Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.relation
% (110491)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1129
% Passive clauses: 10835
% Generated clauses: 32997
% Final active clauses: 1006
% Final passive clauses: 9223
% Input formulas: 69
% Initial clauses: 239
% 
% Duplicate literals: 420
% Trivial inequalities: 166
% Fw subsumption resolutions: 3500
% Fw demodulations: 3473
% Bw demodulations: 111
% 
% Simple tautologies: 289
% Equational tautologies: 42
% Forward subsumptions: 13389
% Bw demodulations to eq. taut.: 10
% 
% Binary resolution: 10296
% Factoring: 969
% Forward superposition: 3876
% Backward superposition: 9489
% Self superposition: 34
% Equality factoring: 26
% Equality resolution: 14
% 
% Split clauses: 869
% Split components: 1739
% SAT solver clauses: 1258
% SAT solver unit clauses: 12
% SAT solver binary clauses: 161
% 
% Sat splits: 869
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 1256
% TWLsolver calls for satisfiability: 525
% 
% Memory used [KB]: 15863
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 144 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.relation
% (110524)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 467
% Passive clauses: 3008
% Generated clauses: 12879
% Final active clauses: 452
% Final passive clauses: 2515
% Input formulas: 69
% Initial clauses: 239
% 
% Duplicate literals: 203
% Trivial inequalities: 80
% Fw subsumption resolutions: 856
% Fw demodulations: 954
% Bw demodulations: 39
% Condensations: 277
% Global subsumptions: 18
% 
% Simple tautologies: 300
% Equational tautologies: 17
% Forward subsumptions: 3687
% 
% Binary resolution: 3480
% Factoring: 124
% Forward superposition: 1714
% Backward superposition: 1807
% Self superposition: 29
% Equality factoring: 26
% Equality resolution: 8
% 
% Split clauses: 32
% Split components: 64
% SAT solver clauses: 3673
% SAT solver unit clauses: 234
% SAT solver binary clauses: 75
% SAT solver learnt clauses: 36
% SAT solver learnt literals: 10
% 
% Sat splits: 32
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 3601
% TWLsolver calls for satisfiability: 2161
% 
% Memory used [KB]: 14967
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.relation
% (110547)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 2718
% Passive clauses: 11489
% Generated clauses: 110711
% Final active clauses: 2162
% Final passive clauses: 3240
% Input formulas: 69
% Initial clauses: 239
% 
% Duplicate literals: 1519
% Trivial inequalities: 107
% Fw subsumption resolutions: 3601
% Fw demodulations: 19305
% Bw demodulations: 4446
% Condensations: 3719
% Global subsumptions: 954
% 
% Simple tautologies: 190
% Equational tautologies: 352
% Forward subsumptions: 64500
% Bw demodulations to eq. taut.: 26
% 
% Unit resulting resolution: 27868
% Factoring: 1
% Forward superposition: 21120
% Backward superposition: 23449
% Self superposition: 67
% Equality factoring: 46
% Equality resolution: 7
% 
% Split clauses: 356
% Split components: 717
% SAT solver clauses: 14806
% SAT solver unit clauses: 2778
% SAT solver binary clauses: 4306
% 
% Sat splits: 356
% Sat splitting refutations: 1648
% 
% Memory used [KB]: 139187
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.relation
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.relation
% SZS output start Proof for thm.relation
fof(f60181,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f592,f720,f2403,f2405,f2556,f2546,f2504,f2610,f2611,f2612,f2613,f3332,f3333,f3345,f3346,f3355,f3356,f3524,f3520,f3519,f3515,f3513,f3509,f3464,f3463,f3462,f3592,f3593,f3594,f3595,f3596,f3597,f3612,f3613,f3767,f3882,f4222,f4223,f4224,f4360,f4452,f4453,f4454,f4625,f4749,f4752,f4763,f4766,f4767,f4901,f4902,f4943,f4947,f4977,f4981,f4991,f4995,f5005,f5016,f5018,f5009,f5019,f5030,f5032,f5043,f5045,f5059,f5049,f5062,f5074,f5077,f5078,f5088,f4764,f5081,f5282,f5283,f5284,f5615,f5614,f5685,f5684,f5702,f5761,f5760,f5833,f5832,f5838,f3361,f6569,f6570,f6571,f6760,f4465,f7505,f7575,f7578,f7576,f12589,f12808,f12809,f12810,f13736,f18758,f18844,f18863,f18845,f18846,f18938,f18939,f19183,f19197,f19370,f19371,f19633,f19652,f20036,f20061,f20082,f18866,f20108,f18865,f20126,f20114,f20044,f20257,f20259,f20260,f20261,f20258,f20262,f20264,f20263,f20265,f20268,f20267,f20269,f20270,f20266,f20271,f20274,f7154,f12461,f20793,f20794,f20803,f20804,f20954,f20958,f20987,f20991,f21002,f21006,f21022,f21025,f21039,f21040,f21051,f21052,f21053,f21057,f21105,f21109,f21118,f21122,f21131,f21142,f21143,f21135,f21161,f21147,f21148,f21149,f21151,f21152,f21153,f21164,f21165,f21166,f21179,f21180,f21181,f21182,f21197,f21186,f21198,f21199,f21200,f21212,f21214,f21215,f21216,f21220,f21231,f21235,f21247,f21248,f21259,f21260,f21261,f21262,f21273,f21274,f21275,f21276,f21289,f21290,f21326,f21306,f21307,f21308,f21309,f21311,f21312,f21313,f21314,f21316,f21317,f21318,f21319,f21334,f21335,f21340,f21341,f21343,f21345,f21346,f21348,f21359,f21360,f21361,f21373,f21376,f21377,f21378,f21434,f21433,f21476,f21479,f21480,f21539,f21538,f21548,f21565,f21566,f21571,f21641,f21640,f21678,f21683,f21684,f21687,f21688,f21689,f21690,f21705,f21729,f21763,f21762,f21768,f21769,f22201,f22415,f22420,f22421,f22424,f22425,f22426,f22427,f22442,f22466,f22483,f22484,f22514,f22513,f22520,f22684,f22683,f23042,f23045,f23046,f23560,f23677,f23691,f23692,f23697,f23708,f24324,f24282,f24388,f24389,f24390,f24391,f24402,f24403,f24412,f24413,f24424,f24428,f24455,f24458,f24467,f24468,f24477,f24480,f24489,f24490,f24491,f24495,f24512,f24515,f24524,f24525,f24556,f24560,f24579,f24590,f24591,f24583,f24609,f24595,f24596,f24597,f24599,f24600,f24601,f24610,f24611,f24612,f24623,f24624,f24625,f24626,f24639,f24630,f24640,f24641,f24642,f24653,f24655,f24656,f24657,f24661,f24708,f24707,f24741,f21491,f24744,f24745,f24796,f24795,f24798,f24813,f21578,f24814,f21579,f24819,f3619,f24828,f24832,f24841,f24842,f24846,f24855,f24856,f24860,f24869,f24884,f24883,f24887,f24898,f24900,f24891,f24899,f24918,f24904,f24905,f24920,f24906,f24921,f24908,f24919,f24909,f24910,f24922,f24923,f24924,f24935,f24936,f24937,f24938,f24939,f24952,f24953,f24943,f24954,f24955,f24966,f24956,f24968,f24969,f24985,f24984,f24988,f25001,f24989,f25002,f24990,f24457,f24994,f3621,f25005,f25009,f25018,f25019,f25023,f25032,f25033,f25037,f25046,f25047,f25058,f25060,f25051,f25059,f25423,f25636,f25622,f25623,f25638,f25624,f25639,f25626,f25637,f25627,f25628,f25640,f25641,f25642,f25653,f25654,f25655,f25656,f25657,f25670,f25671,f25661,f25672,f25673,f25684,f25674,f25686,f25687,f25689,f25702,f25690,f25703,f25691,f24479,f25695,f4371,f25745,f22640,f25870,f25874,f25883,f25884,f25888,f25897,f25925,f25929,f25938,f25939,f25950,f25952,f25943,f25951,f25970,f25956,f25957,f25972,f25958,f25973,f25960,f25971,f25961,f25962,f25974,f25975,f25976,f25987,f25988,f25989,f25990,f25991,f26004,f26005,f25995,f26032,f26033,f26044,f26034,f26046,f26047,f26049,f26062,f26050,f26063,f26051,f24514,f26055,f7516,f26136,f26139,f26140,f26175,f26208,f26209,f26214,f26852,f26867,f26855,f26856,f26897,f26909,f26933,f26946,f26934,f26947,f26939,f26866,f27222,f27223,f27237,f27224,f27239,f27225,f27241,f27226,f27243,f27227,f27245,f27249,f27250,f714,f712,f27251,f27366,f27364,f27543,f27656,f27657,f27666,f27667,f27688,f27692,f4956,f27701,f27705,f4990,f27714,f27718,f5004,f27727,f27738,f27739,f27731,f5017,f27758,f27743,f27744,f21162,f27745,f21163,f27747,f5031,f27748,f27749,f27750,f27760,f27761,f27762,f27763,f27774,f5044,f27775,f27776,f27777,f27790,f5060,f27781,f27791,f27792,f21210,f27793,f27804,f5075,f27806,f21227,f27807,f21228,f27808,f27812,f5092,f28064,f28068,f21246,f28077,f28078,f28088,f28089,f21258,f28090,f28091,f28101,f28102,f21272,f28103,f28104,f21286,f28116,f28117,f21300,f28153,f28132,f28133,f21329,f28134,f21330,f28135,f21327,f28137,f28138,f28139,f28140,f28141,f28143,f3360,f28144,f28145,f3349,f28146,f28157,f21357,f21242,f28158,f28163,f28164,f28165,f28167,f28169,f28170,f28172,f28181,f3338,f28182,f21371,f28183,f28194,f21374,f28196,f21331,f28197,f21388,f28198,f28286,f28285,f28291,f21700,f28296,f28298,f28301,f3362,f28302,f28303,f3350,f28304,f28315,f21728,f838,f28339,f3340,f21741,f28354,f21702,f28355,f21780,f28392,f28391,f28449,f22437,f28454,f28456,f28459,f4466,f28460,f28461,f4469,f28462,f28484,f28483,f28487,f22465,f28511,f4468,f22478,f28526,f22439,f28527,f22495,f28807,f28806,f4765,f28862,f28878,f28865,f28866,f28867,f28880,f29468,f29480,f29483,f29496,f29484,f29497,f29489,f28877,f29512,f29511,f29526,f29525,f29566,f29565,f29610,f29632,f29615,f29617,f29620,f3618,f29621,f29622,f29633,f29623,f29676,f29675,f21243,f29679,f27238,f21560,f29693,f29634,f29694,f29705,f30452,f23561,f30658,f30680,f30663,f30665,f30668,f23054,f30669,f30670,f30681,f30671,f30710,f27234,f30724,f30682,f30725,f30736,f30856,f30861,f30863,f30866,f30867,f30868,f30869,f30912,f30927,f30928,f33683,f33684,f35275,f20272,f36231,f36232,f36416,f36419,f36464,f36465,f36742,f36746,f36757,f36758,f36762,f36773,f36774,f36778,f36789,f36790,f36803,f36805,f36794,f36804,f36826,f36809,f36810,f36828,f36811,f36829,f36813,f36827,f36814,f36815,f36816,f36831,f36832,f36833,f36834,f36847,f36848,f36942,f36943,f36944,f36959,f36960,f36948,f36992,f36993,f37006,f36994,f37008,f37009,f37489,f20275,f37661,f37676,f37662,f37677,f37663,f37667,f13749,f38630,f38666,f38788,f38795,f38849,f39029,f39556,f40018,f40062,f40108,f40161,f40895,f41078,f41621,f41630,f41667,f41676,f41714,f41723,f41960,f41959,f42148,f42147,f42338,f42337,f42442,f42441,f36418,f47572,f30332,f29708,f23440,f22818,f841,f839,f36476,f3880,f840,f47964,f47965,f47974,f47975,f47998,f48002,f20967,f48003,f48004,f48022,f48023,f48030,f48031,f21024,f48032,f48036,f21000,f48037,f48046,f48050,f21015,f48051,f48060,f48073,f48075,f48064,f48074,f48065,f48066,f48096,f48079,f48080,f48098,f48081,f23056,f48083,f48097,f48084,f48085,f48086,f27235,f48087,f48088,f48099,f48100,f21177,f48101,f48114,f48115,f48105,f48106,f48116,f21193,f48117,f21194,f48118,f48131,f48132,f48122,f48133,f48134,f23689,f48135,f48146,f48147,f48149,f23704,f48150,f23705,f48151,f48155,f23055,f48164,f48168,f48178,f48169,f48179,f48180,f48190,f48191,f48189,f48184,f48192,f48193,f48202,f48203,f48201,f48204,f48205,f48214,f48209,f48217,f48218,f48227,f48222,f48253,f48233,f48234,f48258,f48235,f48259,f48236,f48254,f48238,f48239,f48240,f48241,f48242,f48244,f48262,f48245,f48246,f48260,f48247,f48248,f48266,f48289,f48174,f48267,f48272,f48273,f48274,f48276,f48278,f48279,f48281,f48282,f48290,f48257,f48291,f48301,f48292,f48303,f48304,f48296,f48306,f48263,f48307,f48317,f48308,f48312,f48337,f48336,f48369,f48368,f48403,f48402,f48432,f48431,f4751,f48446,f48461,f48449,f48450,f48451,f48463,f49138,f49149,f49139,f49150,f49144,f49175,f48460,f49185,f49214,f49213,f49274,f49273,f49303,f49302,f49338,f49337,f48175,f49843,f49844,f49950,f49954,f49955,f50002,f50003,f50015,f50016,f50007,f50017,f50018,f50029,f50030,f50031,f50032,f50036,f50046,f50047,f50051,f50086,f50064,f50065,f50066,f50067,f50069,f50070,f50071,f50072,f50073,f50075,f50076,f50077,f50078,f50079,f50123,f50124,f50129,f50130,f50131,f50133,f50135,f50136,f50138,f50139,f50175,f50176,f50192,f50196,f49964,f50197,f50206,f50207,f50218,f50219,f50014,f50211,f50220,f50221,f50231,f50232,f50028,f50263,f50264,f50043,f50268,f50277,f50278,f50058,f50282,f50316,f50294,f50295,f30880,f50296,f50089,f50297,f50087,f50299,f50300,f50301,f50302,f50303,f50305,f30883,f50306,f50307,f30881,f50308,f50309,f50320,f50148,f22642,f50321,f50326,f50327,f50328,f27247,f50330,f50332,f21358,f50333,f50335,f21176,f50336,f50345,f30879,f50346,f50347,f50359,f50360,f50351,f50399,f30884,f50400,f30939,f50401,f50405,f22641,f50557,f50573,f50560,f50561,f50562,f50575,f50641,f50651,f50654,f50665,f50655,f50666,f50660,f50572,f52640,f39557,f53130,f53131,f53667,f58811,f58839,f58816,f5061,f58818,f58821,f58842,f58822,f58823,f58840,f58824,f58876,f58838,f58891,f58894,f58843,f58895,f58909,f59042,f59056,f59052,f59049,f38381,f26223,f23823,f23821,f59190,f59156,f59274,f59275,f59276,f59277,f59278,f59438,f59605,f59622,f59608,f59609,f7577,f59610,f30878,f3336,f59656,f30924,f59685,f59700,f59686,f59701,f59692,f30882,f59765,f59883,f59766,f59767,f59768,f59769,f59900,f59884,f59901,f59885,f59902,f59886,f59903,f2404,f59887,f60180])).
fof(f60180,plain,(
  $false | ($spl2 | $spl30)),
  inference(subsumption_resolution,[],[f58200,f4487])).
fof(f4487,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | ($spl2 | $spl30)),
  inference(backward_demodulation,[],[f3880,f714])).
fof(f58200,plain,(
  ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | $spl30),
  inference(subsumption_resolution,[],[f58173,f588])).
fof(f588,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f515])).
fof(f515,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f317])).
fof(f317,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f130])).
fof(f130,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X6] : ! [X8] : (s(X6,X8) = s(X6,X8) <=> p(s(bool,'T_0')))),
  file('thm.relation.EXTEND_RTC_TC','thm.bool.REFL_CLAUSE')).
fof(f58173,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | $spl30),
  inference(superposition,[],[f57610,f1712])).
fof(f1712,plain,(
  s(bool,'T_0') = s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))),
  inference(resolution,[],[f884,f557])).
fof(f557,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f343])).
fof(f343,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f150])).
fof(f150,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X11] : (s(bool,'T_0') = s(bool,X11) <=> p(s(bool,X11)))),
  file('thm.relation.EXTEND_RTC_TC','thm.bool.EQ_CLAUSES|split|1')).
fof(f884,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))),
  inference(resolution,[],[f350,f354])).
fof(f354,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f177])).
fof(f177,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))))),
  inference(ennf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) => p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))))),
  inference(flattening,[],[f75])).
fof(f75,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) => p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))))),
  inference(rectify,[],[f53])).
fof(f53,axiom,(
  ! [X6] : ! [X7,X8,X9] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X8))),s(X6,X9)))) => p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X8),s(X6,X9)))))),
  file('thm.relation.EXTEND_RTC_TC','thm.relation.TC_RULES|split|0')).
fof(f350,plain,(
  p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))))),
  inference(cnf_transformation,[],[f238])).
fof(f238,plain,(
  p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) & p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) & ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13,sK14,sK15,sK16])],[f174])).
fof(f174,plain,(
  ? [X0,X1,X2,X3,X4] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) & ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4)))))),
  inference(flattening,[],[f173])).
fof(f173,plain,(
  ? [X0,X1,X2,X3,X4] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4))))) & ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4)))))),
  inference(ennf_transformation,[],[f72])).
fof(f72,plain,(
  ~! [X0,X1,X2,X3,X4] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4))))) => p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4)))))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ~! [X0] : ! [X1,X2,X3,X4] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4))))) => p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9,X10] : ((p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X8))),s(X6,X9)))) & p(s(bool,'const.relation.RTC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X9),s(X6,X10))))) => p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X8),s(X6,X10)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9,X10] : ((p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X8))),s(X6,X9)))) & p(s(bool,'const.relation.RTC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X9),s(X6,X10))))) => p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X8),s(X6,X10)))))),
  file('thm.relation.EXTEND_RTC_TC','thm.relation.EXTEND_RTC_TC')).
fof(f57610,plain,(
  ( ! [X0] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,X0)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,X0),s(sK12,sK14))))) ) | $spl30),
  inference(forward_demodulation,[],[f722,f3880])).
fof(f722,plain,(
  ( ! [X0] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,X0),s(sK12,sK16)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,X0))))) )),
  inference(resolution,[],[f352,f353])).
fof(f353,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f176])).
fof(f176,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4)))))),
  inference(flattening,[],[f175])).
fof(f175,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4))))) | p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4)))))),
  inference(ennf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0,X1,X2,X3,X4] : ((p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4))))) => p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4)))))),
  inference(flattening,[],[f73])).
fof(f73,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : ((p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4))))) => p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4)))))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X6] : ! [X7,X8,X9,X10] : ((p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X8),s(X6,X9)))) & p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X9),s(X6,X10))))) => p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X8),s(X6,X10)))))),
  file('thm.relation.EXTEND_RTC_TC','thm.relation.TC_RULES|split|1')).
fof(f352,plain,(
  ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))),
  inference(cnf_transformation,[],[f238])).
fof(f59887,plain,(
  ( ! [X15] : (p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK28(happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14)),X15,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X15),s(sK12,sK28(happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14)),X15,sK12)))))) ) | $spl570),
  inference(cnf_transformation,[],[f59887_D])).
fof(f59887_D,plain,(
  ( ! [X15] : (p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK28(happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14)),X15,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X15),s(sK12,sK28(happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14)),X15,sK12)))))) ) <=> ~$spl570),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl570])])).
fof(f2404,plain,(
  ( ! [X3] : (p(s(bool,X3))) ) | $spl8),
  inference(cnf_transformation,[],[f2404_D])).
fof(f2404_D,plain,(
  ( ! [X3] : (p(s(bool,X3))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f59903,plain,(
  ( ! [X14,X15] : (p(s(bool,X14)) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK28(happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14)),X15,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X15),s(sK12,sK28(happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14)),X15,sK12)))))) )),
  inference(resolution,[],[f886,f498])).
fof(f498,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK28(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK28(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) )),
  inference(cnf_transformation,[],[f299])).
fof(f299,plain,(
  ! [X0,X1,X2] : ((! [X8] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X8)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X8))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK26(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK27(X2,X1,X0)))))) & ((! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK28(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK28(X2,X1,X0))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK26,sK27,sK28])],[f298])).
fof(f298,plain,(
  ! [X0,X1,X2] : ((! [X8] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X8)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X8))))) | ? [X7] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X7)))) | ? [X6] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) & ((! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(rectify,[],[f297])).
fof(f297,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) & ((! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(flattening,[],[f296])).
fof(f296,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & ((! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f119])).
fof(f119,plain,(
  ! [X0] : ! [X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X6] : ! [X20,X21] : (? [X8] : (p(s(bool,happ(s(fun(X6,bool),X20),s(X6,X8)))) | p(s(bool,happ(s(fun(X6,bool),X21),s(X6,X8))))) <=> (? [X8] : p(s(bool,happ(s(fun(X6,bool),X20),s(X6,X8)))) | ? [X8] : p(s(bool,happ(s(fun(X6,bool),X21),s(X6,X8))))))),
  file('thm.relation.EXTEND_RTC_TC','thm.bool.EXISTS_OR_THM')).
fof(f886,plain,(
  ( ! [X2] : (p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK35(happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14)),X2,sK12))))) | p(s(bool,X2))) )),
  inference(resolution,[],[f350,f510])).
fof(f510,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK35(X2,X1,X0))))) | p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) )),
  inference(cnf_transformation,[],[f311])).
fof(f311,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK34(X2,X1,X0))))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK35(X2,X1,X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK34,sK35])],[f310])).
fof(f310,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | ? [X5] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(rectify,[],[f309])).
fof(f309,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(flattening,[],[f308])).
fof(f308,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f125])).
fof(f125,plain,(
  ! [X0] : ! [X1,X2] : ((p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X6] : ! [X20,X21] : ((p(s(bool,X20)) | ? [X8] : p(s(bool,happ(s(fun(X6,bool),X21),s(X6,X8))))) <=> ? [X8] : (p(s(bool,X20)) | p(s(bool,happ(s(fun(X6,bool),X21),s(X6,X8))))))),
  file('thm.relation.EXTEND_RTC_TC','thm.bool.RIGHT_OR_EXISTS_THM')).
fof(f59886,plain,(
  ( ! [X13] : (p(s(bool,happ(s(fun(sK12,bool),X13),s(sK12,sK28(X13,happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14)),sK12))))) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK28(X13,happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14)),sK12)))))) ) | $spl568),
  inference(cnf_transformation,[],[f59886_D])).
fof(f59886_D,plain,(
  ( ! [X13] : (p(s(bool,happ(s(fun(sK12,bool),X13),s(sK12,sK28(X13,happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14)),sK12))))) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK28(X13,happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14)),sK12)))))) ) <=> ~$spl568),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl568])])).
fof(f59902,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | p(s(bool,happ(s(fun(sK12,bool),X13),s(sK12,sK28(X13,happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14)),sK12))))) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK28(X13,happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14)),sK12)))))) )),
  inference(resolution,[],[f886,f497])).
fof(f497,plain,(
  ( ! [X2,X0,X5,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK28(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK28(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))) )),
  inference(cnf_transformation,[],[f299])).
fof(f59885,plain,(
  ( ! [X11] : (p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK27(happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14)),X11,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X11),s(sK12,sK26(happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14)),X11,sK12)))))) ) | $spl566),
  inference(cnf_transformation,[],[f59885_D])).
fof(f59885_D,plain,(
  ( ! [X11] : (p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK27(happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14)),X11,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X11),s(sK12,sK26(happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14)),X11,sK12)))))) ) <=> ~$spl566),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl566])])).
fof(f59901,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK27(happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14)),X11,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X11),s(sK12,sK26(happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14)),X11,sK12)))))) )),
  inference(resolution,[],[f886,f496])).
fof(f496,plain,(
  ( ! [X2,X0,X8,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK27(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK26(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X8))))) )),
  inference(cnf_transformation,[],[f299])).
fof(f59884,plain,(
  ( ! [X9] : (p(s(bool,happ(s(fun(sK12,bool),X9),s(sK12,sK27(X9,happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14)),sK12))))) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK26(X9,happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14)),sK12)))))) ) | $spl564),
  inference(cnf_transformation,[],[f59884_D])).
fof(f59884_D,plain,(
  ( ! [X9] : (p(s(bool,happ(s(fun(sK12,bool),X9),s(sK12,sK27(X9,happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14)),sK12))))) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK26(X9,happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14)),sK12)))))) ) <=> ~$spl564),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl564])])).
fof(f59900,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | p(s(bool,happ(s(fun(sK12,bool),X9),s(sK12,sK27(X9,happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14)),sK12))))) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK26(X9,happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14)),sK12)))))) )),
  inference(resolution,[],[f886,f495])).
fof(f495,plain,(
  ( ! [X2,X0,X8,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK27(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK26(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X8))))) )),
  inference(cnf_transformation,[],[f299])).
fof(f59769,plain,(
  ( ! [X14,X15] : (p(s(bool,X14)) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK28(happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14)),X15,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X15),s(sK12,sK28(happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14)),X15,sK12)))))) )),
  inference(resolution,[],[f885,f498])).
fof(f885,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK34(happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14)),X1,sK12))))) | p(s(bool,X1))) )),
  inference(resolution,[],[f350,f508])).
fof(f508,plain,(
  ( ! [X6,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK34(X2,X1,X0))))) | p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) )),
  inference(cnf_transformation,[],[f311])).
fof(f59768,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | p(s(bool,happ(s(fun(sK12,bool),X13),s(sK12,sK28(X13,happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14)),sK12))))) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK28(X13,happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14)),sK12)))))) )),
  inference(resolution,[],[f885,f497])).
fof(f59767,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK27(happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14)),X11,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X11),s(sK12,sK26(happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14)),X11,sK12)))))) )),
  inference(resolution,[],[f885,f496])).
fof(f59766,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | p(s(bool,happ(s(fun(sK12,bool),X9),s(sK12,sK27(X9,happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14)),sK12))))) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK26(X9,happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14)),sK12)))))) )),
  inference(resolution,[],[f885,f495])).
fof(f59883,plain,(
  ( ! [X7] : (p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK35(happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14)),X7,sK12))))) | p(s(bool,X7))) ) | $spl562),
  inference(cnf_transformation,[],[f59883_D])).
fof(f59883_D,plain,(
  ( ! [X7] : (p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK35(happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14)),X7,sK12))))) | p(s(bool,X7))) ) <=> ~$spl562),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl562])])).
fof(f59765,plain,(
  ( ! [X6,X7] : (p(s(bool,X6)) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK35(happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14)),X7,sK12))))) | p(s(bool,X7))) )),
  inference(resolution,[],[f885,f510])).
fof(f30882,plain,(
  s(bool,'F_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | $spl409),
  inference(cnf_transformation,[],[f30882_D])).
fof(f30882_D,plain,(
  s(bool,'F_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) <=> ~$spl409),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl409])])).
fof(f59692,plain,(
  ( ! [X13] : (p(s(bool,X13)) | s(bool,'F_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | ($spl2 | $spl30)),
  inference(resolution,[],[f4510,f3286])).
fof(f3286,plain,(
  ( ! [X4,X5] : (~sP6(X5,'F_0',X4) | p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5)) )),
  inference(subsumption_resolution,[],[f3280,f431])).
fof(f431,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f271])).
fof(f271,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f270])).
fof(f270,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f269])).
fof(f269,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f231])).
fof(f231,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f3280,plain,(
  ( ! [X4,X5] : (~sP8(X4,X5,'F_0') | ~sP6(X5,'F_0',X4) | p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5)) )),
  inference(resolution,[],[f3257,f443])).
fof(f443,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | s(bool,X0) != s(bool,X1) | ~sP9(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f277])).
fof(f277,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP9(X2,X1,X0)) & (~sP9(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f233])).
fof(f233,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP9(X2,X1,X0))),
  inference(definition_folding,[],[f98,f232,f231,f230,f229])).
fof(f229,plain,(
  ! [X0,X1,X2] : (sP6(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f230,plain,(
  ! [X1,X0,X2] : (sP7(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f232,plain,(
  ! [X2,X1,X0] : (sP9(X2,X1,X0) <=> (sP6(X0,X1,X2) & sP7(X1,X0,X2) & sP8(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f98,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f97])).
fof(f97,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X14,X15,X16] : ((p(s(bool,X16)) <=> s(bool,X14) = s(bool,X15)) <=> ((p(s(bool,X16)) | p(s(bool,X15)) | p(s(bool,X14))) & (p(s(bool,X16)) | ~p(s(bool,X14)) | ~p(s(bool,X15))) & (p(s(bool,X15)) | ~p(s(bool,X14)) | ~p(s(bool,X16))) & (p(s(bool,X14)) | ~p(s(bool,X15)) | ~p(s(bool,X16)))))),
  file('thm.relation.EXTEND_RTC_TC','thm.sat.sat_2|49|_prove_ltb')).
fof(f3257,plain,(
  ( ! [X4,X5] : (sP9(X4,'F_0',X5) | ~sP8(X4,X5,'F_0') | ~sP6(X5,'F_0',X4)) )),
  inference(subsumption_resolution,[],[f3255,f571])).
fof(f571,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f165])).
fof(f165,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.relation.EXTEND_RTC_TC',aHL_FALSITY)).
fof(f3255,plain,(
  ( ! [X4,X5] : (sP9(X4,'F_0',X5) | p(s(bool,'F_0')) | ~sP8(X4,X5,'F_0') | ~sP6(X5,'F_0',X4)) )),
  inference(resolution,[],[f3252,f426])).
fof(f426,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | p(s(bool,X1)) | ~sP8(X0,X2,X1) | ~sP7(X1,X2,X0) | ~sP6(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f268])).
fof(f268,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | (sP6(X2,X1,X0) & sP7(X1,X2,X0) & sP8(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP6(X2,X1,X0) | ~sP7(X1,X2,X0) | ~sP8(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f267])).
fof(f267,plain,(
  ! [X2,X1,X0] : ((~sP9(X2,X1,X0) | (sP6(X0,X1,X2) & sP7(X1,X0,X2) & sP8(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP6(X0,X1,X2) | ~sP7(X1,X0,X2) | ~sP8(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP9(X2,X1,X0)))),
  inference(flattening,[],[f266])).
fof(f266,plain,(
  ! [X2,X1,X0] : ((~sP9(X2,X1,X0) | (sP6(X0,X1,X2) & sP7(X1,X0,X2) & sP8(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP6(X0,X1,X2) | ~sP7(X1,X0,X2) | ~sP8(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP9(X2,X1,X0)))),
  inference(nnf_transformation,[],[f232])).
fof(f3252,plain,(
  ( ! [X4,X5] : (sP7('F_0',X4,X5)) )),
  inference(subsumption_resolution,[],[f3246,f433])).
fof(f433,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f274])).
fof(f274,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f273])).
fof(f273,plain,(
  ! [X1,X0,X2] : ((~sP7(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP7(X1,X0,X2)))),
  inference(flattening,[],[f272])).
fof(f272,plain,(
  ! [X1,X0,X2] : ((~sP7(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP7(X1,X0,X2)))),
  inference(nnf_transformation,[],[f230])).
fof(f3246,plain,(
  ( ! [X4,X5] : (sP7('F_0',X4,X5) | p(s(bool,X5))) )),
  inference(resolution,[],[f3243,f431])).
fof(f3243,plain,(
  ( ! [X6,X7] : (~sP8(X7,X6,'F_0') | sP7('F_0',X6,X7)) )),
  inference(subsumption_resolution,[],[f3239,f434])).
fof(f434,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f274])).
fof(f3239,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | ~sP8(X7,X6,'F_0') | sP7('F_0',X6,X7)) )),
  inference(resolution,[],[f3233,f422])).
fof(f422,plain,(
  ( ! [X2,X0,X1] : (sP7(X1,X2,X0) | ~sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f268])).
fof(f3233,plain,(
  ( ! [X2,X3] : (sP9(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP8(X3,X2,'F_0')) )),
  inference(subsumption_resolution,[],[f3232,f557])).
fof(f3232,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP9(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP8(X3,X2,'F_0')) )),
  inference(subsumption_resolution,[],[f3229,f2687])).
fof(f2687,plain,(
  ( ! [X4,X5] : (sP6(X4,'F_0',X5) | s(bool,'T_0') != s(bool,X4)) )),
  inference(subsumption_resolution,[],[f2682,f437])).
fof(f437,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f276])).
fof(f276,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(flattening,[],[f275])).
fof(f275,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(nnf_transformation,[],[f229])).
fof(f2682,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | p(s(bool,X5)) | sP6(X4,'F_0',X5)) )),
  inference(resolution,[],[f2378,f421])).
fof(f421,plain,(
  ( ! [X2,X0,X1] : (sP6(X2,X1,X0) | ~sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f268])).
fof(f2378,plain,(
  ( ! [X6,X7] : (sP9(X7,'F_0',X6) | s(bool,'T_0') != s(bool,X6) | p(s(bool,X7))) )),
  inference(superposition,[],[f2369,f440])).
fof(f440,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f277])).
fof(f2369,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(subsumption_resolution,[],[f2357,f351])).
fof(f351,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))),
  inference(cnf_transformation,[],[f238])).
fof(f2357,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))),
  inference(superposition,[],[f678,f557])).
fof(f678,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))),
  inference(resolution,[],[f351,f552])).
fof(f552,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f341])).
fof(f341,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f146])).
fof(f146,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X11] : (s(bool,'F_0') = s(bool,X11) <=> ~p(s(bool,X11)))),
  file('thm.relation.EXTEND_RTC_TC','thm.bool.EQ_CLAUSES|split|3')).
fof(f3229,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP9(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP8(X3,X2,'F_0') | ~sP6(X2,'F_0',X3)) )),
  inference(resolution,[],[f2688,f425])).
fof(f425,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~p(s(bool,X2)) | ~sP8(X0,X2,X1) | ~sP7(X1,X2,X0) | ~sP6(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f268])).
fof(f2688,plain,(
  ( ! [X6,X7] : (sP7('F_0',X6,X7) | s(bool,'T_0') != s(bool,X6)) )),
  inference(subsumption_resolution,[],[f2683,f433])).
fof(f2683,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X6) | p(s(bool,X7)) | sP7('F_0',X6,X7)) )),
  inference(resolution,[],[f2378,f422])).
fof(f4510,plain,(
  ( ! [X121,X120] : (sP6('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X120,X121)) ) | ($spl2 | $spl30)),
  inference(backward_demodulation,[],[f3880,f1168])).
fof(f1168,plain,(
  ( ! [X121,X120] : (sP6('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X120,X121)) ) | $spl2),
  inference(resolution,[],[f714,f439])).
fof(f439,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f276])).
fof(f59701,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | $spl560),
  inference(cnf_transformation,[],[f59701_D])).
fof(f59701_D,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) <=> ~$spl560),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl560])])).
fof(f59686,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | ($spl2 | $spl10 | $spl30)),
  inference(resolution,[],[f4510,f19474])).
fof(f19474,plain,(
  ( ! [X4,X5,X3] : (~sP6(X4,X3,X5) | p(s(bool,X3)) | p(s(bool,X5)) | s(bool,X3) != s(bool,X4)) ) | $spl10),
  inference(duplicate_literal_removal,[],[f19455])).
fof(f19455,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~sP6(X4,X3,X5) | p(s(bool,X5)) | p(s(bool,X5)) | s(bool,X3) != s(bool,X4)) ) | $spl10),
  inference(resolution,[],[f19390,f443])).
fof(f19390,plain,(
  ( ! [X6,X8,X7] : (sP9(X7,X6,X8) | p(s(bool,X6)) | ~sP6(X8,X6,X7) | p(s(bool,X7))) ) | $spl10),
  inference(resolution,[],[f19385,f431])).
fof(f19385,plain,(
  ( ! [X6,X8,X7] : (~sP8(X7,X8,X6) | p(s(bool,X6)) | sP9(X7,X6,X8) | ~sP6(X8,X6,X7)) ) | $spl10),
  inference(subsumption_resolution,[],[f19380,f553])).
fof(f553,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f341])).
fof(f19380,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') != s(bool,X6) | sP9(X7,X6,X8) | p(s(bool,X6)) | ~sP8(X7,X8,X6) | ~sP6(X8,X6,X7)) ) | $spl10),
  inference(resolution,[],[f19373,f426])).
fof(f19373,plain,(
  ( ! [X10,X11,X9] : (sP7(X10,X9,X11) | s(bool,'F_0') != s(bool,X10)) ) | $spl10),
  inference(subsumption_resolution,[],[f19372,f434])).
fof(f19372,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | s(bool,'F_0') != s(bool,X10) | sP7(X10,X9,X11)) ) | $spl10),
  inference(subsumption_resolution,[],[f19350,f433])).
fof(f19350,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | s(bool,'F_0') != s(bool,X10) | p(s(bool,X11)) | sP7(X10,X9,X11)) ) | $spl10),
  inference(resolution,[],[f19200,f422])).
fof(f19200,plain,(
  ( ! [X6,X8,X7] : (sP9(X6,X7,X8) | ~p(s(bool,X8)) | s(bool,'F_0') != s(bool,X7) | p(s(bool,X6))) ) | $spl10),
  inference(resolution,[],[f19168,f431])).
fof(f19168,plain,(
  ( ! [X4,X5,X3] : (~sP8(X5,X4,X3) | sP9(X5,X3,X4) | ~p(s(bool,X4)) | s(bool,'F_0') != s(bool,X3)) ) | $spl10),
  inference(subsumption_resolution,[],[f19167,f557])).
fof(f19167,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X3) | s(bool,'T_0') != s(bool,X4) | sP9(X5,X3,X4) | ~p(s(bool,X4)) | ~sP8(X5,X4,X3)) ) | $spl10),
  inference(subsumption_resolution,[],[f19161,f19142])).
fof(f19142,plain,(
  ( ! [X6,X8,X7] : (sP6(X6,X7,X8) | s(bool,'F_0') != s(bool,X7) | s(bool,'T_0') != s(bool,X6)) ) | $spl10),
  inference(subsumption_resolution,[],[f19122,f437])).
fof(f19122,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X6) | s(bool,'F_0') != s(bool,X7) | p(s(bool,X8)) | sP6(X6,X7,X8)) ) | $spl10),
  inference(resolution,[],[f2647,f421])).
fof(f2647,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X1,X0) | s(bool,'T_0') != s(bool,X0) | s(bool,'F_0') != s(bool,X1) | p(s(bool,X2))) ) | $spl10),
  inference(superposition,[],[f2405,f440])).
fof(f19161,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X3) | s(bool,'T_0') != s(bool,X4) | sP9(X5,X3,X4) | ~p(s(bool,X4)) | ~sP8(X5,X4,X3) | ~sP6(X4,X3,X5)) ) | $spl10),
  inference(resolution,[],[f19143,f425])).
fof(f19143,plain,(
  ( ! [X10,X11,X9] : (sP7(X10,X9,X11) | s(bool,'F_0') != s(bool,X10) | s(bool,'T_0') != s(bool,X9)) ) | $spl10),
  inference(subsumption_resolution,[],[f19123,f433])).
fof(f19123,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') != s(bool,X9) | s(bool,'F_0') != s(bool,X10) | p(s(bool,X11)) | sP7(X10,X9,X11)) ) | $spl10),
  inference(resolution,[],[f2647,f422])).
fof(f59700,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,X1) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | $spl558),
  inference(cnf_transformation,[],[f59700_D])).
fof(f59700_D,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,X1) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) <=> ~$spl558),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl558])])).
fof(f59685,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'F_0') != s(bool,X1) | s(bool,X1) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | ($spl2 | $spl10 | $spl30)),
  inference(resolution,[],[f4510,f19603])).
fof(f19603,plain,(
  ( ! [X4,X5,X3] : (~sP6(X5,X3,X4) | p(s(bool,X4)) | s(bool,'F_0') != s(bool,X3) | s(bool,X3) != s(bool,X5)) ) | $spl10),
  inference(duplicate_literal_removal,[],[f19584])).
fof(f19584,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X3) | p(s(bool,X4)) | ~sP6(X5,X3,X4) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5)) ) | $spl10),
  inference(resolution,[],[f19386,f443])).
fof(f19386,plain,(
  ( ! [X10,X11,X9] : (sP9(X10,X9,X11) | s(bool,'F_0') != s(bool,X9) | p(s(bool,X10)) | ~sP6(X11,X9,X10)) ) | $spl10),
  inference(subsumption_resolution,[],[f19381,f431])).
fof(f19381,plain,(
  ( ! [X10,X11,X9] : (s(bool,'F_0') != s(bool,X9) | sP9(X10,X9,X11) | p(s(bool,X10)) | ~sP8(X10,X11,X9) | ~sP6(X11,X9,X10)) ) | $spl10),
  inference(resolution,[],[f19373,f427])).
fof(f427,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | p(s(bool,X0)) | ~sP8(X0,X2,X1) | ~sP7(X1,X2,X0) | ~sP6(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f268])).
fof(f30924,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | $spl412),
  inference(cnf_transformation,[],[f30924_D])).
fof(f30924_D,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) <=> ~$spl412),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl412])])).
fof(f59656,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'F_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | s(bool,X1) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | ($spl2 | $spl10 | $spl30)),
  inference(resolution,[],[f4509,f19603])).
fof(f4509,plain,(
  ( ! [X118,X119] : (sP6(X118,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X119)) ) | ($spl2 | $spl30)),
  inference(backward_demodulation,[],[f3880,f1167])).
fof(f1167,plain,(
  ( ! [X118,X119] : (sP6(X118,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X119)) ) | $spl2),
  inference(resolution,[],[f714,f438])).
fof(f438,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f276])).
fof(f3336,plain,(
  ( ! [X8] : (~p(s(bool,X8))) ) | $spl12),
  inference(cnf_transformation,[],[f3336_D])).
fof(f3336_D,plain,(
  ( ! [X8] : (~p(s(bool,X8))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f30878,plain,(
  s(bool,'T_0') = s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | $spl402),
  inference(cnf_transformation,[],[f30878_D])).
fof(f30878_D,plain,(
  s(bool,'T_0') = s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) <=> ~$spl402),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl402])])).
fof(f59610,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | s(bool,'T_0') = s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | ~p(s(bool,X21))) ) | ($spl2 | $spl10 | $spl30 | $spl96)),
  inference(resolution,[],[f4506,f27175])).
fof(f27175,plain,(
  ( ! [X14,X12,X13] : (~sP8(X14,X13,X12) | ~p(s(bool,X13)) | s(bool,'T_0') = s(bool,X12) | ~p(s(bool,X14))) ) | ($spl10 | $spl96)),
  inference(subsumption_resolution,[],[f27163,f7687])).
fof(f7687,plain,(
  ( ! [X52,X53] : (s(bool,X52) != s(bool,X53) | ~p(s(bool,X53)) | s(bool,'T_0') = s(bool,X52)) ) | $spl10),
  inference(subsumption_resolution,[],[f7654,f520])).
fof(f520,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f320])).
fof(f320,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f319])).
fof(f319,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f206])).
fof(f206,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f134])).
fof(f134,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X11] : ((p(s(bool,X11)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X11)))),
  file('thm.relation.EXTEND_RTC_TC','thm.bool.IMP_CLAUSES|split|4')).
fof(f7654,plain,(
  ( ! [X52,X53] : (s(bool,X52) != s(bool,X53) | p(s(bool,'F_0')) | ~p(s(bool,X53)) | s(bool,'T_0') = s(bool,X52)) ) | $spl10),
  inference(superposition,[],[f7571,f584])).
fof(f584,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.relation.EXTEND_RTC_TC',aHL_BOOL_CASES)).
fof(f7571,plain,(
  ( ! [X35,X34] : (s(bool,X34) != s(bool,X35) | p(s(bool,X35)) | ~p(s(bool,X34))) ) | $spl10),
  inference(trivial_inequality_removal,[],[f7546])).
fof(f7546,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X34) != s(bool,X35) | p(s(bool,X35)) | ~p(s(bool,X34))) ) | $spl10),
  inference(superposition,[],[f2627,f559])).
fof(f559,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f344])).
fof(f344,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f151])).
fof(f151,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X11] : (s(bool,'T_0') = s(bool,X11) <=> p(s(bool,X11)))),
  file('thm.relation.EXTEND_RTC_TC','thm.bool.EQ_CLAUSES|split|0')).
fof(f2627,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X5) | s(bool,X4) != s(bool,X5) | p(s(bool,X4))) ) | $spl10),
  inference(superposition,[],[f2405,f555])).
fof(f555,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f342])).
fof(f342,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f149])).
fof(f149,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f148])).
fof(f148,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X11] : (s(bool,'F_0') = s(bool,X11) <=> ~p(s(bool,X11)))),
  file('thm.relation.EXTEND_RTC_TC','thm.bool.EQ_CLAUSES|split|2')).
fof(f27163,plain,(
  ( ! [X14,X12,X13] : (s(bool,'T_0') = s(bool,X12) | ~p(s(bool,X13)) | ~sP8(X14,X13,X12) | s(bool,X12) = s(bool,X13) | ~p(s(bool,X14))) ) | $spl96),
  inference(resolution,[],[f27157,f442])).
fof(f442,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X2)) | ~sP9(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f277])).
fof(f27157,plain,(
  ( ! [X10,X11,X9] : (sP9(X10,X9,X11) | s(bool,'T_0') = s(bool,X9) | ~p(s(bool,X11)) | ~sP8(X10,X11,X9)) ) | $spl96),
  inference(subsumption_resolution,[],[f27154,f439])).
fof(f27154,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') = s(bool,X9) | sP9(X10,X9,X11) | ~p(s(bool,X11)) | ~sP8(X10,X11,X9) | ~sP6(X11,X9,X10)) ) | $spl96),
  inference(resolution,[],[f27150,f425])).
fof(f27150,plain,(
  ( ! [X4,X5,X3] : (sP7(X3,X5,X4) | s(bool,'T_0') = s(bool,X3)) ) | $spl96),
  inference(subsumption_resolution,[],[f27149,f433])).
fof(f27149,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') = s(bool,X3) | p(s(bool,X4)) | sP7(X3,X5,X4)) ) | $spl96),
  inference(subsumption_resolution,[],[f27133,f434])).
fof(f27133,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') = s(bool,X3) | p(s(bool,X4)) | ~p(s(bool,X5)) | sP7(X3,X5,X4)) ) | $spl96),
  inference(resolution,[],[f20263,f422])).
fof(f4506,plain,(
  ( ! [X109,X108] : (sP8(X108,X109,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | ($spl2 | $spl30)),
  inference(backward_demodulation,[],[f3880,f1162])).
fof(f1162,plain,(
  ( ! [X109,X108] : (sP8(X108,X109,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) | $spl2),
  inference(resolution,[],[f714,f429])).
fof(f429,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f271])).
fof(f7577,plain,(
  ( ! [X23] : (s(bool,'T_0') != s(bool,X23)) ) | $spl72),
  inference(cnf_transformation,[],[f7577_D])).
fof(f7577_D,plain,(
  ( ! [X23] : (s(bool,'T_0') != s(bool,X23)) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f59609,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | s(bool,'T_0') != s(bool,X19)) ) | ($spl2 | $spl10 | $spl30)),
  inference(resolution,[],[f4506,f19799])).
fof(f19799,plain,(
  ( ! [X4,X5,X3] : (~sP8(X4,X3,X5) | ~p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5) | s(bool,'T_0') != s(bool,X3)) ) | $spl10),
  inference(subsumption_resolution,[],[f19798,f552])).
fof(f19798,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5) | p(s(bool,X5)) | ~sP8(X4,X3,X5)) ) | $spl10),
  inference(subsumption_resolution,[],[f19797,f19142])).
fof(f19797,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5) | p(s(bool,X5)) | ~sP8(X4,X3,X5) | ~sP6(X3,X5,X4)) ) | $spl10),
  inference(subsumption_resolution,[],[f19758,f19143])).
fof(f19758,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5) | p(s(bool,X5)) | ~sP8(X4,X3,X5) | ~sP7(X5,X3,X4) | ~sP6(X3,X5,X4)) ) | $spl10),
  inference(resolution,[],[f2649,f426])).
fof(f2649,plain,(
  ( ! [X6,X8,X7] : (~sP9(X8,X7,X6) | s(bool,'T_0') != s(bool,X6) | ~p(s(bool,X8)) | s(bool,'F_0') != s(bool,X7)) ) | $spl10),
  inference(superposition,[],[f2405,f442])).
fof(f59608,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | ~p(s(bool,X16)) | ~p(s(bool,X17))) ) | ($spl2 | $spl10 | $spl30)),
  inference(resolution,[],[f4506,f19796])).
fof(f19796,plain,(
  ( ! [X2,X0,X1] : (~sP8(X1,X0,X2) | s(bool,'F_0') != s(bool,X2) | ~p(s(bool,X0)) | ~p(s(bool,X1))) ) | $spl10),
  inference(subsumption_resolution,[],[f19795,f557])).
fof(f19795,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X1)) | s(bool,'F_0') != s(bool,X2) | ~p(s(bool,X0)) | ~sP8(X1,X0,X2)) ) | $spl10),
  inference(subsumption_resolution,[],[f19794,f19142])).
fof(f19794,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X1)) | s(bool,'F_0') != s(bool,X2) | ~p(s(bool,X0)) | ~sP8(X1,X0,X2) | ~sP6(X0,X2,X1)) ) | $spl10),
  inference(subsumption_resolution,[],[f19757,f19143])).
fof(f19757,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X1)) | s(bool,'F_0') != s(bool,X2) | ~p(s(bool,X0)) | ~sP8(X1,X0,X2) | ~sP7(X2,X0,X1) | ~sP6(X0,X2,X1)) ) | $spl10),
  inference(resolution,[],[f2649,f425])).
fof(f59622,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X11) | ~p(s(bool,X11))) ) | $spl556),
  inference(cnf_transformation,[],[f59622_D])).
fof(f59622_D,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X11) | ~p(s(bool,X11))) ) <=> ~$spl556),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl556])])).
fof(f59605,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X11) | ~p(s(bool,X11)) | s(bool,'F_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | ($spl2 | $spl10 | $spl30)),
  inference(resolution,[],[f4506,f19168])).
fof(f59438,plain,(
  $false | ($spl30 | $spl256 | $spl369)),
  inference(subsumption_resolution,[],[f23821,f47721])).
fof(f47721,plain,(
  s(sK12,sK14) != s(sK12,sK15) | ($spl30 | $spl369)),
  inference(backward_demodulation,[],[f3880,f27363])).
fof(f27363,plain,(
  s(sK12,sK15) != s(sK12,sK16) | $spl369),
  inference(cnf_transformation,[],[f27363_D])).
fof(f27363_D,plain,(
  s(sK12,sK15) != s(sK12,sK16) <=> ~$spl369),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl369])])).
fof(f59278,plain,(
  $false | ($spl30 | $spl368)),
  inference(subsumption_resolution,[],[f59210,f27541])).
fof(f27541,plain,(
  ( ! [X127] : (p(s(bool,X127))) ) | $spl368),
  inference(subsumption_resolution,[],[f27416,f1699])).
fof(f1699,plain,(
  ( ! [X134] : (sP11(X134,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(resolution,[],[f884,f475])).
fof(f475,plain,(
  ( ! [X0,X1] : (sP11(X0,X1) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f291])).
fof(f291,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP11(X0,X1)))),
  inference(flattening,[],[f290])).
fof(f290,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP11(X0,X1)))),
  inference(nnf_transformation,[],[f236])).
fof(f236,plain,(
  ! [X0,X1] : (sP11(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f27416,plain,(
  ( ! [X127] : (~sP11(X127,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | p(s(bool,X127))) ) | $spl368),
  inference(backward_demodulation,[],[f27364,f842])).
fof(f842,plain,(
  ( ! [X127] : (~sP11(X127,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | p(s(bool,X127))) )),
  inference(subsumption_resolution,[],[f790,f538])).
fof(f538,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f332])).
fof(f332,plain,(
  ! [X0] : (((~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | p(s(bool,X0))) & (~p(s(bool,X0)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f331])).
fof(f331,plain,(
  ! [X0] : (((~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f141])).
fof(f141,plain,(
  ! [X0] : ((p(s(bool,X0)) | p(s(bool,'F_0'))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f65])).
fof(f65,axiom,(
  ! [X11] : ((p(s(bool,X11)) | p(s(bool,'F_0'))) <=> p(s(bool,X11)))),
  file('thm.relation.EXTEND_RTC_TC','thm.bool.OR_CLAUSES|split|3')).
fof(f790,plain,(
  ( ! [X127] : (p(s(bool,'F_0')) | p(s(bool,X127)) | ~sP11(X127,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) )),
  inference(resolution,[],[f352,f474])).
fof(f474,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP11(X0,X1)) )),
  inference(cnf_transformation,[],[f291])).
fof(f59210,plain,(
  ( ! [X438] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,X438),s(sK12,sK14))))) ) | ($spl30 | $spl368)),
  inference(resolution,[],[f27541,f57610])).
fof(f59277,plain,(
  $false | $spl368),
  inference(subsumption_resolution,[],[f59158,f27541])).
fof(f59158,plain,(
  ~p(s(bool,'T_0')) | $spl368),
  inference(resolution,[],[f27541,f575])).
fof(f575,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f346])).
fof(f346,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f167])).
fof(f167,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f29])).
fof(f29,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.relation.EXTEND_RTC_TC','thm.bool.NOT_CLAUSES|split|2')).
fof(f59276,plain,(
  $false | $spl368),
  inference(subsumption_resolution,[],[f59157,f27541])).
fof(f59157,plain,(
  ~p(s(bool,'T_0')) | $spl368),
  inference(resolution,[],[f27541,f573])).
fof(f573,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f345])).
fof(f345,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f166])).
fof(f166,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f30])).
fof(f30,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.relation.EXTEND_RTC_TC','thm.bool.NOT_CLAUSES|split|1')).
fof(f59275,plain,(
  $false | $spl368),
  inference(subsumption_resolution,[],[f59154,f27541])).
fof(f59154,plain,(
  ( ! [X224] : (~p(s(bool,X224))) ) | $spl368),
  inference(resolution,[],[f27541,f567])).
fof(f567,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f216])).
fof(f216,plain,(
  ! [X0] : ((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f159])).
fof(f159,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) => ~p(s(bool,X0)))),
  inference(flattening,[],[f158])).
fof(f158,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) => ~p(s(bool,X0)))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X11] : ((p(s(bool,X11)) => p(s(bool,'F_0'))) => ~p(s(bool,X11)))),
  file('thm.relation.EXTEND_RTC_TC','thm.bool.IMP_F')).
fof(f59274,plain,(
  $false | $spl368),
  inference(subsumption_resolution,[],[f59150,f27541])).
fof(f59150,plain,(
  ( ! [X220] : (~p(s(bool,X220))) ) | $spl368),
  inference(resolution,[],[f27541,f520])).
fof(f59156,plain,(
  $false | $spl368),
  inference(resolution,[],[f27541,f571])).
fof(f59190,plain,(
  $false | ($spl30 | $spl368)),
  inference(resolution,[],[f27541,f47639])).
fof(f47639,plain,(
  ~p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK14)))) | $spl30),
  inference(backward_demodulation,[],[f3880,f721])).
fof(f721,plain,(
  ~p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK16))))),
  inference(resolution,[],[f352,f354])).
fof(f23821,plain,(
  s(sK12,sK14) = s(sK12,sK15) | $spl256),
  inference(cnf_transformation,[],[f23821_D])).
fof(f23821_D,plain,(
  s(sK12,sK14) = s(sK12,sK15) <=> ~$spl256),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl256])])).
fof(f23823,plain,(
  p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15))),s(sK12,sK14)))) | $spl258),
  inference(cnf_transformation,[],[f23823_D])).
fof(f23823_D,plain,(
  p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15))),s(sK12,sK14)))) <=> ~$spl258),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl258])])).
fof(f26223,plain,(
  p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15))),s(sK12,sK14)))) | s(sK12,sK14) = s(sK12,sK15) | $spl240),
  inference(resolution,[],[f22640,f578])).
fof(f578,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | s(X0,X2) = s(X0,X3) | ~p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f348])).
fof(f348,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | s(X0,X2) = s(X0,X3) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))))) & ((s(X0,X2) != s(X0,X3) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))))) | p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))))),
  inference(flattening,[],[f347])).
fof(f347,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | (s(X0,X2) = s(X0,X3) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))))) & ((s(X0,X2) != s(X0,X3) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))))) | p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))))),
  inference(nnf_transformation,[],[f169])).
fof(f169,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) <=> (s(X0,X2) = s(X0,X3) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))))))),
  inference(flattening,[],[f168])).
fof(f168,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) <=> (s(X0,X2) = s(X0,X3) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))))))),
  inference(rectify,[],[f66])).
fof(f66,axiom,(
  ! [X6] : ! [X7,X8,X9] : (p(s(bool,'const.relation.RC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X8),s(X6,X9)))) <=> (s(X6,X8) = s(X6,X9) | p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X8))),s(X6,X9))))))),
  file('thm.relation.EXTEND_RTC_TC','thm.relation.RC_DEF')).
fof(f38381,plain,(
  p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15))),s(sK12,sK14)))) | s(sK12,sK14) = s(sK12,sK15) | $spl240),
  inference(subsumption_resolution,[],[f38046,f588])).
fof(f38046,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15))),s(sK12,sK14)))) | s(sK12,sK14) = s(sK12,sK15) | $spl240),
  inference(superposition,[],[f578,f23789])).
fof(f23789,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | $spl240),
  inference(resolution,[],[f22640,f557])).
fof(f59049,plain,(
  s(sK12,sK14) = s(sK12,sK15) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15))),s(sK12,sK14)))) | ($spl0 | $spl30)),
  inference(forward_demodulation,[],[f59048,f3880])).
fof(f59048,plain,(
  p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15))),s(sK12,sK14)))) | s(sK12,sK15) = s(sK12,sK16) | ($spl0 | $spl30)),
  inference(forward_demodulation,[],[f32812,f3880])).
fof(f32812,plain,(
  p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15))),s(sK12,sK16)))) | s(sK12,sK15) = s(sK12,sK16) | $spl0),
  inference(subsumption_resolution,[],[f32477,f588])).
fof(f32477,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15))),s(sK12,sK16)))) | s(sK12,sK15) = s(sK12,sK16) | $spl0),
  inference(superposition,[],[f578,f27332])).
fof(f27332,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) | $spl0),
  inference(resolution,[],[f712,f557])).
fof(f59052,plain,(
  s(sK12,sK14) = s(sK12,sK15) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15))),s(sK12,sK14)))) | ($spl0 | $spl30)),
  inference(forward_demodulation,[],[f59051,f3880])).
fof(f59051,plain,(
  p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15))),s(sK12,sK14)))) | s(sK12,sK15) = s(sK12,sK16) | ($spl0 | $spl30)),
  inference(forward_demodulation,[],[f27821,f3880])).
fof(f27821,plain,(
  p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15))),s(sK12,sK16)))) | s(sK12,sK15) = s(sK12,sK16) | $spl0),
  inference(resolution,[],[f712,f578])).
fof(f59056,plain,(
  $false | ($spl30 | $spl257 | $spl368)),
  inference(subsumption_resolution,[],[f59055,f23820])).
fof(f23820,plain,(
  s(sK12,sK14) != s(sK12,sK15) | $spl257),
  inference(cnf_transformation,[],[f23820_D])).
fof(f23820_D,plain,(
  s(sK12,sK14) != s(sK12,sK15) <=> ~$spl257),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl257])])).
fof(f59055,plain,(
  s(sK12,sK14) = s(sK12,sK15) | ($spl30 | $spl368)),
  inference(forward_demodulation,[],[f27364,f3880])).
fof(f59042,plain,(
  $false | ($spl30 | $spl243 | $spl370)),
  inference(subsumption_resolution,[],[f59041,f571])).
fof(f59041,plain,(
  p(s(bool,'F_0')) | ($spl30 | $spl243 | $spl370)),
  inference(forward_demodulation,[],[f58920,f50935])).
fof(f50935,plain,(
  s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | $spl243),
  inference(resolution,[],[f22641,f553])).
fof(f58920,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | ($spl30 | $spl370)),
  inference(resolution,[],[f58698,f354])).
fof(f58698,plain,(
  p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15))),s(sK12,sK14)))) | ($spl30 | $spl370)),
  inference(forward_demodulation,[],[f27366,f3880])).
fof(f58909,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK14)))) ) | $spl554),
  inference(cnf_transformation,[],[f58909_D])).
fof(f58909_D,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK14)))) ) <=> ~$spl554),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl554])])).
fof(f58895,plain,(
  ( ! [X2,X3] : (~p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK14)))) ) | ($spl10 | $spl30)),
  inference(resolution,[],[f53026,f19313])).
fof(f19313,plain,(
  ( ! [X4,X5,X3] : (~sP7(X4,X3,X5) | ~p(s(bool,X3)) | p(s(bool,X5)) | s(bool,X3) != s(bool,X4)) ) | $spl10),
  inference(subsumption_resolution,[],[f19289,f7571])).
fof(f19289,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | ~sP7(X4,X3,X5) | ~p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X3) != s(bool,X4)) ) | $spl10),
  inference(resolution,[],[f19264,f443])).
fof(f19264,plain,(
  ( ! [X2,X0,X1] : (sP9(X1,X2,X0) | ~p(s(bool,X0)) | ~sP7(X2,X0,X1) | ~p(s(bool,X2))) ) | $spl10),
  inference(resolution,[],[f19262,f429])).
fof(f19262,plain,(
  ( ! [X4,X5,X3] : (~sP8(X4,X3,X5) | ~p(s(bool,X3)) | sP9(X4,X5,X3) | ~sP7(X5,X3,X4)) ) | $spl10),
  inference(subsumption_resolution,[],[f19255,f557])).
fof(f19255,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,X3) | sP9(X4,X5,X3) | ~p(s(bool,X3)) | ~sP8(X4,X3,X5) | ~sP7(X5,X3,X4)) ) | $spl10),
  inference(resolution,[],[f19248,f425])).
fof(f19248,plain,(
  ( ! [X6,X8,X7] : (sP6(X7,X6,X8) | s(bool,'T_0') != s(bool,X7)) ) | $spl10),
  inference(subsumption_resolution,[],[f19247,f437])).
fof(f19247,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X7) | p(s(bool,X8)) | sP6(X7,X6,X8)) ) | $spl10),
  inference(subsumption_resolution,[],[f19227,f438])).
fof(f19227,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | s(bool,'T_0') != s(bool,X7) | p(s(bool,X8)) | sP6(X7,X6,X8)) ) | $spl10),
  inference(resolution,[],[f19175,f421])).
fof(f19175,plain,(
  ( ! [X6,X8,X7] : (sP9(X6,X7,X8) | p(s(bool,X7)) | s(bool,'T_0') != s(bool,X8) | p(s(bool,X6))) ) | $spl10),
  inference(resolution,[],[f19159,f431])).
fof(f19159,plain,(
  ( ! [X6,X8,X7] : (~sP8(X8,X7,X6) | sP9(X8,X6,X7) | p(s(bool,X6)) | s(bool,'T_0') != s(bool,X7)) ) | $spl10),
  inference(subsumption_resolution,[],[f19158,f435])).
fof(f435,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f274])).
fof(f19158,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X7) | sP9(X8,X6,X7) | p(s(bool,X6)) | ~sP8(X8,X7,X6) | ~sP7(X6,X7,X8)) ) | $spl10),
  inference(subsumption_resolution,[],[f19150,f553])).
fof(f19150,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,'T_0') != s(bool,X7) | sP9(X8,X6,X7) | p(s(bool,X6)) | ~sP8(X8,X7,X6) | ~sP7(X6,X7,X8)) ) | $spl10),
  inference(resolution,[],[f19142,f426])).
fof(f53026,plain,(
  ( ! [X123,X122] : (sP7(happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK14)),X122,X123)) ) | $spl30),
  inference(resolution,[],[f47639,f435])).
fof(f58843,plain,(
  ( ! [X27] : (s(bool,'T_0') != s(bool,X27) | s(bool,X27) != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK14)))) ) | $spl550),
  inference(cnf_transformation,[],[f58843_D])).
fof(f58843_D,plain,(
  ( ! [X27] : (s(bool,'T_0') != s(bool,X27) | s(bool,X27) != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK14)))) ) <=> ~$spl550),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl550])])).
fof(f58894,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X1) | s(bool,X1) != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK14)))) ) | ($spl10 | $spl30)),
  inference(resolution,[],[f53026,f19575])).
fof(f19575,plain,(
  ( ! [X4,X5,X3] : (~sP7(X5,X3,X4) | p(s(bool,X4)) | s(bool,'T_0') != s(bool,X3) | s(bool,X3) != s(bool,X5)) ) | $spl10),
  inference(duplicate_literal_removal,[],[f19556])).
fof(f19556,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,X3) | p(s(bool,X4)) | ~sP7(X5,X3,X4) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5)) ) | $spl10),
  inference(resolution,[],[f19263,f443])).
fof(f19263,plain,(
  ( ! [X10,X11,X9] : (sP9(X10,X11,X9) | s(bool,'T_0') != s(bool,X9) | p(s(bool,X10)) | ~sP7(X11,X9,X10)) ) | $spl10),
  inference(subsumption_resolution,[],[f19257,f431])).
fof(f19257,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') != s(bool,X9) | sP9(X10,X11,X9) | p(s(bool,X10)) | ~sP8(X10,X9,X11) | ~sP7(X11,X9,X10)) ) | $spl10),
  inference(resolution,[],[f19248,f427])).
fof(f58891,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK14)))) ) | $spl552),
  inference(cnf_transformation,[],[f58891_D])).
fof(f58891_D,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK14)))) ) <=> ~$spl552),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl552])])).
fof(f58838,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK14))) | $spl543),
  inference(cnf_transformation,[],[f58838_D])).
fof(f58838_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK14))) <=> ~$spl543),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl543])])).
fof(f58876,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK14))) | s(bool,X1) != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK14)))) ) | ($spl10 | $spl30)),
  inference(resolution,[],[f53025,f19575])).
fof(f53025,plain,(
  ( ! [X121,X120] : (sP7(X120,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK14)),X121)) ) | $spl30),
  inference(resolution,[],[f47639,f434])).
fof(f58824,plain,(
  ( ! [X30,X29] : (~p(s(bool,X29)) | s(bool,X30) != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK14))) | s(bool,'T_0') != s(bool,X30)) ) | ($spl10 | $spl30)),
  inference(resolution,[],[f53022,f19111])).
fof(f19111,plain,(
  ( ! [X12,X10,X11] : (~sP8(X11,X12,'F_0') | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,'T_0') != s(bool,X10)) ) | $spl10),
  inference(subsumption_resolution,[],[f19096,f2627])).
fof(f19096,plain,(
  ( ! [X12,X10,X11] : (s(bool,'T_0') != s(bool,X10) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | ~p(s(bool,X12)) | ~sP8(X11,X12,'F_0')) ) | $spl10),
  inference(resolution,[],[f2632,f3233])).
fof(f2632,plain,(
  ( ! [X19,X17,X18] : (~sP9(X19,'F_0',X17) | s(bool,'T_0') != s(bool,X18) | ~p(s(bool,X19)) | s(bool,X17) != s(bool,X18)) ) | $spl10),
  inference(superposition,[],[f2405,f442])).
fof(f53022,plain,(
  ( ! [X114,X115] : (sP8(X114,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK14)),X115)) ) | $spl30),
  inference(resolution,[],[f47639,f430])).
fof(f430,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f271])).
fof(f58840,plain,(
  ( ! [X22] : (sP9(X22,'F_0',happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK14)))) ) | $spl546),
  inference(cnf_transformation,[],[f58840_D])).
fof(f58840_D,plain,(
  ( ! [X22] : (sP9(X22,'F_0',happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK14)))) ) <=> ~$spl546),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl546])])).
fof(f58823,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') != s(bool,X27) | sP9(X28,'F_0',happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK14))) | s(bool,X27) != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK14)))) ) | ($spl10 | $spl30)),
  inference(resolution,[],[f53022,f19068])).
fof(f19068,plain,(
  ( ! [X6,X8,X7] : (~sP8(X8,X7,'F_0') | s(bool,'T_0') != s(bool,X6) | sP9(X8,'F_0',X7) | s(bool,X6) != s(bool,X7)) ) | $spl10),
  inference(subsumption_resolution,[],[f19067,f3252])).
fof(f19067,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,X7) | s(bool,'T_0') != s(bool,X6) | sP9(X8,'F_0',X7) | ~sP8(X8,X7,'F_0') | ~sP7('F_0',X7,X8)) ) | $spl10),
  inference(subsumption_resolution,[],[f19061,f2627])).
fof(f19061,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,X7) | s(bool,'T_0') != s(bool,X6) | sP9(X8,'F_0',X7) | ~p(s(bool,X7)) | ~sP8(X8,X7,'F_0') | ~sP7('F_0',X7,X8)) ) | $spl10),
  inference(resolution,[],[f19054,f425])).
fof(f19054,plain,(
  ( ! [X14,X12,X13] : (sP6(X13,'F_0',X14) | s(bool,X12) != s(bool,X13) | s(bool,'T_0') != s(bool,X12)) ) | $spl10),
  inference(subsumption_resolution,[],[f19042,f437])).
fof(f19042,plain,(
  ( ! [X14,X12,X13] : (s(bool,'T_0') != s(bool,X12) | s(bool,X12) != s(bool,X13) | p(s(bool,X14)) | sP6(X13,'F_0',X14)) ) | $spl10),
  inference(resolution,[],[f2630,f421])).
fof(f2630,plain,(
  ( ! [X12,X13,X11] : (sP9(X13,'F_0',X11) | s(bool,'T_0') != s(bool,X12) | s(bool,X11) != s(bool,X12) | p(s(bool,X13))) ) | $spl10),
  inference(superposition,[],[f2405,f440])).
fof(f58822,plain,(
  ( ! [X26] : (s(bool,'T_0') != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK14))) | ~p(s(bool,X26))) ) | $spl30),
  inference(resolution,[],[f53022,f3327])).
fof(f3327,plain,(
  ( ! [X14,X15] : (~sP8(X14,X15,'F_0') | s(bool,'T_0') != s(bool,X15) | ~p(s(bool,X14))) )),
  inference(subsumption_resolution,[],[f3326,f437])).
fof(f3326,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | ~sP8(X14,X15,'F_0') | ~sP6(X15,'F_0',X14)) )),
  inference(subsumption_resolution,[],[f3325,f520])).
fof(f3325,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | p(s(bool,'F_0')) | ~sP8(X14,X15,'F_0') | ~sP6(X15,'F_0',X14)) )),
  inference(subsumption_resolution,[],[f3318,f435])).
fof(f3318,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | p(s(bool,'F_0')) | ~sP8(X14,X15,'F_0') | ~sP7('F_0',X15,X14) | ~sP6(X15,'F_0',X14)) )),
  inference(resolution,[],[f2380,f426])).
fof(f2380,plain,(
  ( ! [X10,X11] : (~sP9(X11,'F_0',X10) | ~p(s(bool,X11)) | s(bool,'T_0') != s(bool,X10)) )),
  inference(superposition,[],[f2369,f442])).
fof(f58842,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK14))) | $spl548),
  inference(cnf_transformation,[],[f58842_D])).
fof(f58842_D,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK14))) <=> ~$spl548),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl548])])).
fof(f58821,plain,(
  ( ! [X25] : (s(bool,'F_0') = s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK14))) | ~p(s(bool,X25))) ) | $spl30),
  inference(resolution,[],[f53022,f3285])).
fof(f3285,plain,(
  ( ! [X2,X3] : (~sP8(X2,X3,'F_0') | s(bool,'F_0') = s(bool,X3) | ~p(s(bool,X2))) )),
  inference(subsumption_resolution,[],[f3279,f437])).
fof(f3279,plain,(
  ( ! [X2,X3] : (~sP8(X2,X3,'F_0') | ~sP6(X3,'F_0',X2) | s(bool,'F_0') = s(bool,X3) | ~p(s(bool,X2))) )),
  inference(resolution,[],[f3257,f442])).
fof(f58818,plain,(
  ( ! [X22] : (sP9(X22,'F_0',happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK14))) | s(bool,'T_0') != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK14)))) ) | $spl30),
  inference(resolution,[],[f53022,f3227])).
fof(f3227,plain,(
  ( ! [X4,X5] : (~sP8(X5,X4,'F_0') | sP9(X5,'F_0',X4) | s(bool,'T_0') != s(bool,X4)) )),
  inference(subsumption_resolution,[],[f3226,f571])).
fof(f3226,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | sP9(X5,'F_0',X4) | p(s(bool,'F_0')) | ~sP8(X5,X4,'F_0')) )),
  inference(subsumption_resolution,[],[f3223,f435])).
fof(f3223,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | sP9(X5,'F_0',X4) | p(s(bool,'F_0')) | ~sP8(X5,X4,'F_0') | ~sP7('F_0',X4,X5)) )),
  inference(resolution,[],[f2687,f426])).
fof(f5061,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,X8)) ) | $spl60),
  inference(cnf_transformation,[],[f5061_D])).
fof(f5061_D,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,X8)) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f58816,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,X19) | s(bool,'T_0') != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK14)))) ) | ($spl10 | $spl30)),
  inference(resolution,[],[f53022,f19799])).
fof(f58839,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK14))) | p(s(bool,X9))) ) | $spl544),
  inference(cnf_transformation,[],[f58839_D])).
fof(f58839_D,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK14))) | p(s(bool,X9))) ) <=> ~$spl544),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl544])])).
fof(f58811,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK14))) | p(s(bool,X9)) | s(bool,'T_0') != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK14)))) ) | ($spl10 | $spl30)),
  inference(resolution,[],[f53022,f19159])).
fof(f53667,plain,(
  $false | ($spl4 | $spl30 | $spl39)),
  inference(subsumption_resolution,[],[f53666,f48545])).
fof(f48545,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | ($spl4 | $spl30)),
  inference(resolution,[],[f4731,f557])).
fof(f4731,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | ($spl4 | $spl30)),
  inference(forward_demodulation,[],[f3741,f3880])).
fof(f3741,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | $spl4),
  inference(subsumption_resolution,[],[f3622,f352])).
fof(f3622,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | $spl4),
  inference(resolution,[],[f838,f355])).
fof(f355,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f179])).
fof(f179,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))))),
  inference(flattening,[],[f178])).
fof(f178,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | (p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))))),
  inference(ennf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) => (p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))))),
  inference(flattening,[],[f77])).
fof(f77,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) => (p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : ! [X7,X8,X9] : (p(s(bool,'const.relation.RTC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X8),s(X6,X9)))) => (p(s(bool,'const.relation.RC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X8),s(X6,X9)))) | p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X8),s(X6,X9))))))),
  file('thm.relation.EXTEND_RTC_TC','thm.relation.RTC_TC_RC')).
fof(f53666,plain,(
  s(bool,'T_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | ($spl30 | $spl39)),
  inference(forward_demodulation,[],[f4468,f3880])).
fof(f53131,plain,(
  $false | ($spl17 | $spl30)),
  inference(subsumption_resolution,[],[f53124,f369])).
fof(f369,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2))))) )),
  inference(cnf_transformation,[],[f90])).
fof(f90,plain,(
  ! [X0,X1,X2] : p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2))))),
  inference(flattening,[],[f89])).
fof(f89,plain,(
  ! [X0] : ! [X1,X2] : p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2))))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X6] : ! [X7,X8] : p(s(bool,'const.relation.RTC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X8),s(X6,X8))))),
  file('thm.relation.EXTEND_RTC_TC','thm.relation.RTC_RULES|split|0')).
fof(f53124,plain,(
  ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | ($spl17 | $spl30)),
  inference(trivial_inequality_removal,[],[f53115])).
fof(f53115,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | ($spl17 | $spl30)),
  inference(superposition,[],[f52915,f559])).
fof(f52915,plain,(
  s(bool,'T_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | ($spl17 | $spl30)),
  inference(forward_demodulation,[],[f3340,f3880])).
fof(f53130,plain,(
  $false | ($spl17 | $spl30)),
  inference(subsumption_resolution,[],[f53125,f369])).
fof(f53125,plain,(
  ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | ($spl17 | $spl30)),
  inference(trivial_inequality_removal,[],[f53114])).
fof(f53114,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | ($spl17 | $spl30)),
  inference(superposition,[],[f52915,f557])).
fof(f39557,plain,(
  ( ! [X325,X326] : (~p(s(bool,X325)) | p(s(bool,X326)) | sP9(X326,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X325)) ) | $spl450),
  inference(cnf_transformation,[],[f39557_D])).
fof(f39557_D,plain,(
  ( ! [X325,X326] : (~p(s(bool,X325)) | p(s(bool,X326)) | sP9(X326,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X325)) ) <=> ~$spl450),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl450])])).
fof(f52640,plain,(
  ( ! [X325,X324,X326] : (p(s(bool,X324)) | ~p(s(bool,X325)) | sP9(X326,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X325) | p(s(bool,X326))) ) | ($spl10 | $spl243)),
  inference(subsumption_resolution,[],[f52386,f553])).
fof(f52386,plain,(
  ( ! [X325,X324,X326] : (s(bool,'F_0') != s(bool,X324) | p(s(bool,X324)) | ~p(s(bool,X325)) | sP9(X326,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X325) | p(s(bool,X326))) ) | ($spl10 | $spl243)),
  inference(superposition,[],[f7606,f50935])).
fof(f7606,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X1) != s(bool,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | sP9(X3,X1,X0) | p(s(bool,X3))) ) | $spl10),
  inference(superposition,[],[f7571,f440])).
fof(f50572,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | $spl531),
  inference(cnf_transformation,[],[f50572_D])).
fof(f50572_D,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) <=> ~$spl531),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl531])])).
fof(f50660,plain,(
  ( ! [X12] : (p(s(bool,X12)) | s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | ($spl6 | $spl30)),
  inference(resolution,[],[f48440,f3286])).
fof(f48440,plain,(
  ( ! [X112,X113] : (sP6('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X112,X113)) ) | ($spl6 | $spl30)),
  inference(forward_demodulation,[],[f3824,f3880])).
fof(f3824,plain,(
  ( ! [X112,X113] : (sP6('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X112,X113)) ) | $spl6),
  inference(resolution,[],[f841,f439])).
fof(f50666,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | $spl540),
  inference(cnf_transformation,[],[f50666_D])).
fof(f50666_D,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) <=> ~$spl540),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl540])])).
fof(f50655,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | ($spl6 | $spl10 | $spl30)),
  inference(resolution,[],[f48440,f19474])).
fof(f50665,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | $spl538),
  inference(cnf_transformation,[],[f50665_D])).
fof(f50665_D,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) <=> ~$spl538),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl538])])).
fof(f50654,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'F_0') != s(bool,X1) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | ($spl6 | $spl10 | $spl30)),
  inference(resolution,[],[f48440,f19603])).
fof(f50651,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | $spl536),
  inference(cnf_transformation,[],[f50651_D])).
fof(f50651_D,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) <=> ~$spl536),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl536])])).
fof(f50641,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | ($spl6 | $spl10 | $spl30)),
  inference(resolution,[],[f48439,f19603])).
fof(f48439,plain,(
  ( ! [X111,X110] : (sP6(X110,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X111)) ) | ($spl6 | $spl30)),
  inference(forward_demodulation,[],[f3823,f3880])).
fof(f3823,plain,(
  ( ! [X111,X110] : (sP6(X110,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X111)) ) | $spl6),
  inference(resolution,[],[f841,f438])).
fof(f50575,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | $spl534),
  inference(cnf_transformation,[],[f50575_D])).
fof(f50575_D,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) <=> ~$spl534),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl534])])).
fof(f50562,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | s(bool,'T_0') = s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | ~p(s(bool,X21))) ) | ($spl6 | $spl10 | $spl30 | $spl96)),
  inference(resolution,[],[f48412,f27175])).
fof(f48412,plain,(
  ( ! [X101,X100] : (sP8(X100,X101,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | ($spl6 | $spl30)),
  inference(forward_demodulation,[],[f3818,f3880])).
fof(f3818,plain,(
  ( ! [X101,X100] : (sP8(X100,X101,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | $spl6),
  inference(resolution,[],[f841,f429])).
fof(f50561,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | s(bool,'T_0') != s(bool,X19)) ) | ($spl6 | $spl10 | $spl30)),
  inference(resolution,[],[f48412,f19799])).
fof(f50560,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | ~p(s(bool,X16)) | ~p(s(bool,X17))) ) | ($spl6 | $spl10 | $spl30)),
  inference(resolution,[],[f48412,f19796])).
fof(f50573,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X11) | ~p(s(bool,X11))) ) | $spl532),
  inference(cnf_transformation,[],[f50573_D])).
fof(f50573_D,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X11) | ~p(s(bool,X11))) ) <=> ~$spl532),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl532])])).
fof(f50557,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X11) | ~p(s(bool,X11)) | s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | ($spl6 | $spl10 | $spl30)),
  inference(resolution,[],[f48412,f19168])).
fof(f22641,plain,(
  ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | $spl243),
  inference(cnf_transformation,[],[f22641_D])).
fof(f22641_D,plain,(
  ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) <=> ~$spl243),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl243])])).
fof(f50405,plain,(
  ( ! [X12] : (p(s(bool,X12)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47742,f943])).
fof(f943,plain,(
  ( ! [X111,X112] : (~sP7(X111,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X112) | p(s(bool,X112)) | ~p(s(bool,X111))) )),
  inference(resolution,[],[f350,f432])).
fof(f432,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f274])).
fof(f47742,plain,(
  ( ! [X107,X108] : (sP7('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X107,X108)) ) | ($spl3 | $spl30)),
  inference(backward_demodulation,[],[f3880,f27994])).
fof(f27994,plain,(
  ( ! [X107,X108] : (sP7('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X107,X108)) ) | $spl3),
  inference(resolution,[],[f713,f435])).
fof(f713,plain,(
  ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | $spl3),
  inference(cnf_transformation,[],[f713_D])).
fof(f713_D,plain,(
  ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f50401,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | ~p(s(bool,X4)) | p(s(bool,X5))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47742,f432])).
fof(f30939,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | $spl414),
  inference(cnf_transformation,[],[f30939_D])).
fof(f30939_D,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) <=> ~$spl414),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl414])])).
fof(f50400,plain,(
  ( ! [X2,X3] : (~p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | ($spl3 | $spl10 | $spl30)),
  inference(resolution,[],[f47742,f19313])).
fof(f30884,plain,(
  ( ! [X27] : (s(bool,'T_0') != s(bool,X27) | s(bool,X27) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | $spl410),
  inference(cnf_transformation,[],[f30884_D])).
fof(f30884_D,plain,(
  ( ! [X27] : (s(bool,'T_0') != s(bool,X27) | s(bool,X27) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) <=> ~$spl410),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl410])])).
fof(f50399,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X1) | s(bool,X1) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | ($spl3 | $spl10 | $spl30)),
  inference(resolution,[],[f47742,f19575])).
fof(f50351,plain,(
  ( ! [X12] : (p(s(bool,X12)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47741,f942])).
fof(f942,plain,(
  ( ! [X109,X110] : (~sP7(happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X109,X110) | p(s(bool,X110)) | ~p(s(bool,X109))) )),
  inference(resolution,[],[f350,f432])).
fof(f47741,plain,(
  ( ! [X105,X106] : (sP7(X105,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X106)) ) | ($spl3 | $spl30)),
  inference(backward_demodulation,[],[f3880,f27993])).
fof(f27993,plain,(
  ( ! [X105,X106] : (sP7(X105,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X106)) ) | $spl3),
  inference(resolution,[],[f713,f434])).
fof(f50360,plain,(
  ( ! [X6,X7] : (sP9(X6,X7,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | $spl528),
  inference(cnf_transformation,[],[f50360_D])).
fof(f50360_D,plain,(
  ( ! [X6,X7] : (sP9(X6,X7,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) <=> ~$spl528),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl528])])).
fof(f50359,plain,(
  ( ! [X6,X7] : (sP9(X6,X7,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | ($spl3 | $spl30)),
  inference(subsumption_resolution,[],[f50358,f47739])).
fof(f47739,plain,(
  ( ! [X99,X100] : (sP8(X99,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X100)) ) | ($spl3 | $spl30)),
  inference(backward_demodulation,[],[f3880,f27990])).
fof(f27990,plain,(
  ( ! [X99,X100] : (sP8(X99,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X100)) ) | $spl3),
  inference(resolution,[],[f713,f430])).
fof(f50358,plain,(
  ( ! [X6,X7] : (sP9(X6,X7,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | ~sP8(X6,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X7)) ) | ($spl3 | $spl30)),
  inference(subsumption_resolution,[],[f50348,f439])).
fof(f50348,plain,(
  ( ! [X6,X7] : (sP9(X6,X7,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | ~sP8(X6,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X7) | ~sP6('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X7,X6)) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47741,f425])).
fof(f50347,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | p(s(bool,X5))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47741,f432])).
fof(f50346,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | p(s(bool,X2)) | s(bool,X3) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | ($spl3 | $spl10 | $spl30)),
  inference(resolution,[],[f47741,f19313])).
fof(f30879,plain,(
  s(bool,'T_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | $spl403),
  inference(cnf_transformation,[],[f30879_D])).
fof(f30879_D,plain,(
  s(bool,'T_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) <=> ~$spl403),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl403])])).
fof(f50345,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | s(bool,X1) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | ($spl3 | $spl10 | $spl30)),
  inference(resolution,[],[f47741,f19575])).
fof(f50336,plain,(
  ( ! [X31] : (p(s(bool,X31)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47740,f940])).
fof(f940,plain,(
  ( ! [X105,X106] : (~sP8(X105,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X106) | p(s(bool,X106)) | ~p(s(bool,X105))) )),
  inference(resolution,[],[f350,f428])).
fof(f428,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f271])).
fof(f47740,plain,(
  ( ! [X101,X102] : (sP8('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X101,X102)) ) | ($spl3 | $spl30)),
  inference(backward_demodulation,[],[f3880,f27991])).
fof(f27991,plain,(
  ( ! [X101,X102] : (sP8('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X101,X102)) ) | $spl3),
  inference(resolution,[],[f713,f431])).
fof(f21176,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X0) | s(bool,X0) != s(bool,X1)) ) | $spl118),
  inference(cnf_transformation,[],[f21176_D])).
fof(f21176_D,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X0) | s(bool,X0) != s(bool,X1)) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f50335,plain,(
  ( ! [X30,X29] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | s(bool,X29) != s(bool,X30) | s(bool,'T_0') != s(bool,X29)) ) | ($spl3 | $spl10 | $spl30)),
  inference(resolution,[],[f47740,f19111])).
fof(f50333,plain,(
  ( ! [X26] : (s(bool,'T_0') != s(bool,X26) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47740,f3327])).
fof(f21358,plain,(
  ( ! [X23] : (s(bool,'F_0') = s(bool,X23)) ) | $spl154),
  inference(cnf_transformation,[],[f21358_D])).
fof(f21358_D,plain,(
  ( ! [X23] : (s(bool,'F_0') = s(bool,X23)) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f50332,plain,(
  ( ! [X25] : (s(bool,'F_0') = s(bool,X25) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47740,f3285])).
fof(f50330,plain,(
  ( ! [X23] : (~p(s(bool,X23)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47740,f3242])).
fof(f3242,plain,(
  ( ! [X0,X1] : (~sP8(X1,X0,'F_0') | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(subsumption_resolution,[],[f3236,f552])).
fof(f3236,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~sP8(X1,X0,'F_0') | s(bool,'F_0') = s(bool,X0) | ~p(s(bool,X1))) )),
  inference(resolution,[],[f3233,f442])).
fof(f27247,plain,(
  ( ! [X30] : (s(bool,'T_0') = s(bool,X30)) ) | $spl366),
  inference(cnf_transformation,[],[f27247_D])).
fof(f27247_D,plain,(
  ( ! [X30] : (s(bool,'T_0') = s(bool,X30)) ) <=> ~$spl366),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl366])])).
fof(f50328,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | s(bool,'T_0') = s(bool,X21) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | ($spl3 | $spl10 | $spl30 | $spl96)),
  inference(resolution,[],[f47740,f27175])).
fof(f50327,plain,(
  ( ! [X19,X18] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | s(bool,'F_0') != s(bool,X18) | s(bool,'T_0') != s(bool,X19)) ) | ($spl3 | $spl10 | $spl30)),
  inference(resolution,[],[f47740,f19799])).
fof(f50326,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,X16) | ~p(s(bool,X17)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | ($spl3 | $spl10 | $spl30)),
  inference(resolution,[],[f47740,f19796])).
fof(f50321,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | ~p(s(bool,X6)) | p(s(bool,X7))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47740,f428])).
fof(f22642,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | $spl242),
  inference(cnf_transformation,[],[f22642_D])).
fof(f22642_D,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) <=> ~$spl242),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl242])])).
fof(f50148,plain,(
  ( ! [X4,X5] : (sP9('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X4,X5) | ~sP6(X5,X4,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | ~sP7(X4,X5,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | $spl526),
  inference(cnf_transformation,[],[f50148_D])).
fof(f50148_D,plain,(
  ( ! [X4,X5] : (sP9('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X4,X5) | ~sP6(X5,X4,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | ~sP7(X4,X5,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) <=> ~$spl526),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl526])])).
fof(f50320,plain,(
  ( ! [X4,X5] : (sP9('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X4,X5) | p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | ~sP7(X4,X5,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | ~sP6(X5,X4,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47740,f427])).
fof(f50309,plain,(
  ( ! [X31] : (p(s(bool,X31)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47739,f939])).
fof(f939,plain,(
  ( ! [X103,X104] : (~sP8(happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X103,X104) | p(s(bool,X104)) | ~p(s(bool,X103))) )),
  inference(resolution,[],[f350,f428])).
fof(f50308,plain,(
  ( ! [X30,X29] : (~p(s(bool,X29)) | s(bool,X30) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | s(bool,'T_0') != s(bool,X30)) ) | ($spl3 | $spl10 | $spl30)),
  inference(resolution,[],[f47739,f19111])).
fof(f30881,plain,(
  ( ! [X22] : (sP9(X22,'F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | $spl406),
  inference(cnf_transformation,[],[f30881_D])).
fof(f30881_D,plain,(
  ( ! [X22] : (sP9(X22,'F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) <=> ~$spl406),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl406])])).
fof(f50307,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') != s(bool,X27) | sP9(X28,'F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | s(bool,X27) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | ($spl3 | $spl10 | $spl30)),
  inference(resolution,[],[f47739,f19068])).
fof(f50306,plain,(
  ( ! [X26] : (s(bool,'T_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | ~p(s(bool,X26))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47739,f3327])).
fof(f30883,plain,(
  s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | $spl408),
  inference(cnf_transformation,[],[f30883_D])).
fof(f30883_D,plain,(
  s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) <=> ~$spl408),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl408])])).
fof(f50305,plain,(
  ( ! [X25] : (s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | ~p(s(bool,X25))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47739,f3285])).
fof(f50303,plain,(
  ( ! [X23] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | ~p(s(bool,X23))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47739,f3242])).
fof(f50302,plain,(
  ( ! [X22] : (sP9(X22,'F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | s(bool,'T_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47739,f3227])).
fof(f50301,plain,(
  ( ! [X21,X20] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | s(bool,'T_0') = s(bool,X20) | ~p(s(bool,X21))) ) | ($spl3 | $spl10 | $spl30 | $spl96)),
  inference(resolution,[],[f47739,f27175])).
fof(f50300,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,X19) | s(bool,'T_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | ($spl3 | $spl10 | $spl30)),
  inference(resolution,[],[f47739,f19799])).
fof(f50299,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,X16) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | ~p(s(bool,X17))) ) | ($spl3 | $spl10 | $spl30)),
  inference(resolution,[],[f47739,f19796])).
fof(f50087,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | ~sP7(X1,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X0)) ) | $spl522),
  inference(cnf_transformation,[],[f50087_D])).
fof(f50087_D,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | ~sP7(X1,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X0)) ) <=> ~$spl522),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl522])])).
fof(f50297,plain,(
  ( ! [X12,X13] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | sP9(X12,X13,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | ~sP7(X13,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X12)) ) | ($spl3 | $spl10 | $spl30)),
  inference(resolution,[],[f47739,f19262])).
fof(f50089,plain,(
  ( ! [X10,X11] : (sP9(X10,X11,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | s(bool,'F_0') != s(bool,X11)) ) | $spl524),
  inference(cnf_transformation,[],[f50089_D])).
fof(f50089_D,plain,(
  ( ! [X10,X11] : (sP9(X10,X11,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | s(bool,'F_0') != s(bool,X11)) ) <=> ~$spl524),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl524])])).
fof(f50296,plain,(
  ( ! [X10,X11] : (sP9(X10,X11,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | s(bool,'F_0') != s(bool,X11)) ) | ($spl3 | $spl10 | $spl30)),
  inference(resolution,[],[f47739,f19168])).
fof(f30880,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | p(s(bool,X9))) ) | $spl404),
  inference(cnf_transformation,[],[f30880_D])).
fof(f30880_D,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | p(s(bool,X9))) ) <=> ~$spl404),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl404])])).
fof(f50295,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | p(s(bool,X9)) | s(bool,'T_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | ($spl3 | $spl10 | $spl30)),
  inference(resolution,[],[f47739,f19159])).
fof(f50294,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | p(s(bool,X7))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47739,f428])).
fof(f50316,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | ~sP7(X1,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X0)) ) | ($spl3 | $spl30)),
  inference(subsumption_resolution,[],[f50291,f439])).
fof(f50291,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | ~sP7(X1,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X0) | ~sP6('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X1,X0)) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47739,f425])).
fof(f50282,plain,(
  ( ! [X10] : (p(s(bool,X10)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47734,f927])).
fof(f927,plain,(
  ( ! [X80,X79] : (~sP4(X79,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X80) | p(s(bool,X80)) | ~p(s(bool,X79))) )),
  inference(resolution,[],[f350,f411])).
fof(f411,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f263])).
fof(f263,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(flattening,[],[f262])).
fof(f262,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(nnf_transformation,[],[f226])).
fof(f226,plain,(
  ! [X0,X1,X2] : (sP4(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f47734,plain,(
  ( ! [X76,X75] : (sP4('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X75,X76)) ) | ($spl3 | $spl30)),
  inference(backward_demodulation,[],[f3880,f27978])).
fof(f27978,plain,(
  ( ! [X76,X75] : (sP4('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X75,X76)) ) | $spl3),
  inference(resolution,[],[f713,f414])).
fof(f414,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f263])).
fof(f50058,plain,(
  ( ! [X2,X3] : (sP5(X2,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X3) | ~p(s(bool,X3))) ) | $spl520),
  inference(cnf_transformation,[],[f50058_D])).
fof(f50058_D,plain,(
  ( ! [X2,X3] : (sP5(X2,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X3) | ~p(s(bool,X3))) ) <=> ~$spl520),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl520])])).
fof(f50278,plain,(
  ( ! [X2,X3] : (sP5(X2,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X3) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | ~p(s(bool,X3))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47734,f407])).
fof(f407,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP4(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f261])).
fof(f261,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | (sP4(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP4(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f260])).
fof(f260,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | (sP4(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP4(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP5(X2,X0,X1)))),
  inference(flattening,[],[f259])).
fof(f259,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | (sP4(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP4(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP5(X2,X0,X1)))),
  inference(nnf_transformation,[],[f227])).
fof(f227,plain,(
  ! [X2,X0,X1] : (sP5(X2,X0,X1) <=> (sP4(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f50277,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | ~p(s(bool,X0)) | p(s(bool,X1))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47734,f411])).
fof(f50268,plain,(
  ( ! [X10] : (p(s(bool,X10)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47733,f926])).
fof(f926,plain,(
  ( ! [X78,X77] : (~sP4(happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X77,X78) | p(s(bool,X78)) | ~p(s(bool,X77))) )),
  inference(resolution,[],[f350,f411])).
fof(f47733,plain,(
  ( ! [X74,X73] : (sP4(X73,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X74)) ) | ($spl3 | $spl30)),
  inference(backward_demodulation,[],[f3880,f27977])).
fof(f27977,plain,(
  ( ! [X74,X73] : (sP4(X73,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X74)) ) | $spl3),
  inference(resolution,[],[f713,f413])).
fof(f413,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f263])).
fof(f50043,plain,(
  ( ! [X2,X3] : (sP5(X2,X3,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | ~p(s(bool,X3))) ) | $spl518),
  inference(cnf_transformation,[],[f50043_D])).
fof(f50043_D,plain,(
  ( ! [X2,X3] : (sP5(X2,X3,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | ~p(s(bool,X3))) ) <=> ~$spl518),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl518])])).
fof(f50264,plain,(
  ( ! [X2,X3] : (sP5(X2,X3,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | ~p(s(bool,X3)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47733,f407])).
fof(f50263,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | p(s(bool,X1))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47733,f411])).
fof(f50028,plain,(
  ( ! [X2,X3] : (sP3('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X2,X3)) ) | $spl516),
  inference(cnf_transformation,[],[f50028_D])).
fof(f50028_D,plain,(
  ( ! [X2,X3] : (sP3('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X2,X3)) ) <=> ~$spl516),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl516])])).
fof(f50232,plain,(
  ( ! [X6,X7] : (sP3('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X6,X7) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | ($spl3 | $spl30)),
  inference(subsumption_resolution,[],[f50223,f400])).
fof(f400,plain,(
  ( ! [X2,X0,X1] : (sP3(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f258])).
fof(f258,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP3(X2,X0,X1)) & (~sP3(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f257])).
fof(f257,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP3(X2,X0,X1)) & (~sP3(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f225])).
fof(f225,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP3(X2,X0,X1))),
  inference(definition_folding,[],[f94,f224,f223])).
fof(f223,plain,(
  ! [X2,X0,X1] : (sP2(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f224,plain,(
  ! [X2,X0,X1] : (sP3(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP2(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f94,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f93])).
fof(f93,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X14,X15,X16] : ((p(s(bool,X16)) <=> (p(s(bool,X15)) | p(s(bool,X14)))) <=> ((p(s(bool,X16)) | ~p(s(bool,X15))) & (p(s(bool,X16)) | ~p(s(bool,X14))) & (p(s(bool,X15)) | p(s(bool,X14)) | ~p(s(bool,X16)))))),
  file('thm.relation.EXTEND_RTC_TC','thm.sat.sat_25_prove_ltb')).
fof(f50223,plain,(
  ( ! [X6,X7] : (sP3('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X6,X7) | p(s(bool,X6)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47730,f392])).
fof(f392,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~sP2(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f253])).
fof(f253,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP2(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP2(X0,X1,X2) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f252])).
fof(f252,plain,(
  ! [X2,X0,X1] : ((~sP3(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP2(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP2(X2,X0,X1) | sP3(X2,X0,X1)))),
  inference(flattening,[],[f251])).
fof(f251,plain,(
  ! [X2,X0,X1] : ((~sP3(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP2(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP2(X2,X0,X1)) | sP3(X2,X0,X1)))),
  inference(nnf_transformation,[],[f224])).
fof(f47730,plain,(
  ( ! [X47,X48] : (sP2('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X47,X48)) ) | ($spl3 | $spl30)),
  inference(backward_demodulation,[],[f3880,f27964])).
fof(f27964,plain,(
  ( ! [X47,X48] : (sP2('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X47,X48)) ) | $spl3),
  inference(resolution,[],[f713,f397])).
fof(f397,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f256])).
fof(f256,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(rectify,[],[f255])).
fof(f255,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(flattening,[],[f254])).
fof(f254,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(nnf_transformation,[],[f223])).
fof(f50231,plain,(
  ( ! [X4,X5] : (sP3('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X4,X5) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | ($spl3 | $spl30)),
  inference(subsumption_resolution,[],[f50222,f399])).
fof(f399,plain,(
  ( ! [X2,X0,X1] : (sP3(X2,X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f258])).
fof(f50222,plain,(
  ( ! [X4,X5] : (sP3('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X4,X5) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | p(s(bool,X5))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47730,f391])).
fof(f391,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~sP2(X0,X1,X2) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f253])).
fof(f50221,plain,(
  ( ! [X2,X3] : (sP3('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X2,X3) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47730,f586])).
fof(f586,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~sP2(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(duplicate_literal_removal,[],[f390])).
fof(f390,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~sP2(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f253])).
fof(f50220,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | p(s(bool,X0)) | p(s(bool,X1))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47730,f394])).
fof(f394,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f256])).
fof(f50211,plain,(
  ( ! [X10] : (p(s(bool,X10)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47726,f897])).
fof(f897,plain,(
  ( ! [X19,X20] : (~sP0(X19,X20,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))) | p(s(bool,X20)) | ~p(s(bool,X19))) )),
  inference(resolution,[],[f350,f377])).
fof(f377,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP0(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f248])).
fof(f248,plain,(
  ! [X0,X1,X2] : ((~sP0(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP0(X0,X1,X2)))),
  inference(rectify,[],[f247])).
fof(f247,plain,(
  ! [X2,X0,X1] : ((~sP0(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP0(X2,X0,X1)))),
  inference(flattening,[],[f246])).
fof(f246,plain,(
  ! [X2,X0,X1] : ((~sP0(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP0(X2,X0,X1)))),
  inference(nnf_transformation,[],[f220])).
fof(f220,plain,(
  ! [X2,X0,X1] : (sP0(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f47726,plain,(
  ( ! [X15,X16] : (sP0('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X15,X16)) ) | ($spl3 | $spl30)),
  inference(backward_demodulation,[],[f3880,f27948])).
fof(f27948,plain,(
  ( ! [X15,X16] : (sP0('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X15,X16)) ) | $spl3),
  inference(resolution,[],[f713,f380])).
fof(f380,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f248])).
fof(f50014,plain,(
  ( ! [X2,X3] : (sP1('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X2,X3)) ) | $spl514),
  inference(cnf_transformation,[],[f50014_D])).
fof(f50014_D,plain,(
  ( ! [X2,X3] : (sP1('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X2,X3)) ) <=> ~$spl514),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl514])])).
fof(f50219,plain,(
  ( ! [X6,X7] : (sP1('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X6,X7) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | ($spl3 | $spl30)),
  inference(subsumption_resolution,[],[f50209,f383])).
fof(f383,plain,(
  ( ! [X2,X0,X1] : (sP1(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f250])).
fof(f250,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP1(X2,X0,X1)) & (~sP1(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f249])).
fof(f249,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP1(X2,X0,X1)) & (~sP1(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f222])).
fof(f222,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP1(X2,X0,X1))),
  inference(definition_folding,[],[f188,f221,f220])).
fof(f221,plain,(
  ! [X2,X0,X1] : (sP1(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP0(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f188,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f92])).
fof(f92,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f91])).
fof(f91,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X14,X15,X16] : ((p(s(bool,X16)) <=> (p(s(bool,X15)) => p(s(bool,X14)))) <=> ((p(s(bool,X16)) | p(s(bool,X15))) & (p(s(bool,X16)) | ~p(s(bool,X14))) & (~p(s(bool,X15)) | p(s(bool,X14)) | ~p(s(bool,X16)))))),
  file('thm.relation.EXTEND_RTC_TC','thm.sat.sat_27_prove_ltb')).
fof(f50209,plain,(
  ( ! [X6,X7] : (sP1('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X6,X7) | p(s(bool,X6)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47726,f375])).
fof(f375,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~sP0(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f245])).
fof(f245,plain,(
  ! [X0,X1,X2] : ((~sP1(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP0(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP0(X0,X1,X2) | sP1(X0,X1,X2)))),
  inference(rectify,[],[f244])).
fof(f244,plain,(
  ! [X2,X0,X1] : ((~sP1(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP0(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP0(X2,X0,X1) | sP1(X2,X0,X1)))),
  inference(flattening,[],[f243])).
fof(f243,plain,(
  ! [X2,X0,X1] : ((~sP1(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP0(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP0(X2,X0,X1)) | sP1(X2,X0,X1)))),
  inference(nnf_transformation,[],[f221])).
fof(f50218,plain,(
  ( ! [X4,X5] : (sP1('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X4,X5) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | ($spl3 | $spl30)),
  inference(subsumption_resolution,[],[f50208,f382])).
fof(f382,plain,(
  ( ! [X2,X0,X1] : (sP1(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f250])).
fof(f50208,plain,(
  ( ! [X4,X5] : (sP1('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X4,X5) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | ~p(s(bool,X5))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47726,f374])).
fof(f374,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~sP0(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f245])).
fof(f50207,plain,(
  ( ! [X2,X3] : (sP1('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X2,X3) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47726,f585])).
fof(f585,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~sP0(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(duplicate_literal_removal,[],[f373])).
fof(f373,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~sP0(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f245])).
fof(f50206,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | p(s(bool,X0)) | ~p(s(bool,X1))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47726,f377])).
fof(f50197,plain,(
  ( ! [X10] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | p(s(bool,X10))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47725,f896])).
fof(f896,plain,(
  ( ! [X17,X18] : (~sP0(happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X17,X18) | ~p(s(bool,X18)) | p(s(bool,X17))) )),
  inference(resolution,[],[f350,f377])).
fof(f47725,plain,(
  ( ! [X14,X13] : (sP0(X13,X14,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | ($spl3 | $spl30)),
  inference(backward_demodulation,[],[f3880,f27947])).
fof(f27947,plain,(
  ( ! [X14,X13] : (sP0(X13,X14,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) | $spl3),
  inference(resolution,[],[f713,f378])).
fof(f378,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f248])).
fof(f49964,plain,(
  ( ! [X8,X9] : (sP1(X8,X9,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | p(s(bool,X9))) ) | $spl512),
  inference(cnf_transformation,[],[f49964_D])).
fof(f49964_D,plain,(
  ( ! [X8,X9] : (sP1(X8,X9,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | p(s(bool,X9))) ) <=> ~$spl512),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl512])])).
fof(f50196,plain,(
  ( ! [X8,X9] : (sP1(X8,X9,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | p(s(bool,X9)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47725,f376])).
fof(f376,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~sP0(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f245])).
fof(f50192,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47725,f377])).
fof(f50176,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | ($spl2 | $spl30)),
  inference(subsumption_resolution,[],[f50168,f538])).
fof(f50168,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | ($spl2 | $spl30)),
  inference(resolution,[],[f4511,f474])).
fof(f4511,plain,(
  ( ! [X134] : (sP11(X134,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | ($spl2 | $spl30)),
  inference(backward_demodulation,[],[f3880,f1179])).
fof(f1179,plain,(
  ( ! [X134] : (sP11(X134,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) | $spl2),
  inference(resolution,[],[f714,f475])).
fof(f50175,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | ($spl2 | $spl30)),
  inference(subsumption_resolution,[],[f50167,f538])).
fof(f50167,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | ($spl2 | $spl30)),
  inference(resolution,[],[f4511,f478])).
fof(f478,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP11(X0,X1)) )),
  inference(cnf_transformation,[],[f293])).
fof(f293,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | (p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP11(X0,X1)))),
  inference(flattening,[],[f292])).
fof(f292,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0')))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP11(X0,X1)))),
  inference(nnf_transformation,[],[f237])).
fof(f237,plain,(
  ! [X0,X1] : (sP11(X0,X1) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(definition_folding,[],[f194,f236])).
fof(f194,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f193])).
fof(f193,plain,(
  ! [X0,X1] : (((p(s(bool,X1)) | p(s(bool,X0))) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | (p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(ennf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(flattening,[],[f108])).
fof(f108,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X17,X6] : ((~(p(s(bool,X6)) | p(s(bool,X17))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X6)) => p(s(bool,'F_0'))) => (~p(s(bool,X17)) => p(s(bool,'F_0')))))),
  file('thm.relation.EXTEND_RTC_TC','thm.sat.sat_9_prove_ltb')).
fof(f50139,plain,(
  ( ! [X31] : (p(s(bool,X31)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47740,f940])).
fof(f50138,plain,(
  ( ! [X30,X29] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | s(bool,X29) != s(bool,X30) | s(bool,'T_0') != s(bool,X29)) ) | ($spl3 | $spl10 | $spl30)),
  inference(resolution,[],[f47740,f19111])).
fof(f50136,plain,(
  ( ! [X26] : (s(bool,'T_0') != s(bool,X26) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47740,f3327])).
fof(f50135,plain,(
  ( ! [X25] : (s(bool,'F_0') = s(bool,X25) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47740,f3285])).
fof(f50133,plain,(
  ( ! [X23] : (~p(s(bool,X23)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47740,f3242])).
fof(f50131,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | s(bool,'T_0') = s(bool,X21) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | ($spl3 | $spl10 | $spl30 | $spl96)),
  inference(resolution,[],[f47740,f27175])).
fof(f50130,plain,(
  ( ! [X19,X18] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | s(bool,'F_0') != s(bool,X18) | s(bool,'T_0') != s(bool,X19)) ) | ($spl3 | $spl10 | $spl30)),
  inference(resolution,[],[f47740,f19799])).
fof(f50129,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,X16) | ~p(s(bool,X17)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | ($spl3 | $spl10 | $spl30)),
  inference(resolution,[],[f47740,f19796])).
fof(f50124,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | ~p(s(bool,X6)) | p(s(bool,X7))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47740,f428])).
fof(f50123,plain,(
  ( ! [X4,X5] : (sP9('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X4,X5) | p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | ~sP7(X4,X5,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | ~sP6(X5,X4,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47740,f427])).
fof(f50079,plain,(
  ( ! [X31] : (p(s(bool,X31)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47739,f939])).
fof(f50078,plain,(
  ( ! [X30,X29] : (~p(s(bool,X29)) | s(bool,X30) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | s(bool,'T_0') != s(bool,X30)) ) | ($spl3 | $spl10 | $spl30)),
  inference(resolution,[],[f47739,f19111])).
fof(f50077,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') != s(bool,X27) | sP9(X28,'F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | s(bool,X27) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | ($spl3 | $spl10 | $spl30)),
  inference(resolution,[],[f47739,f19068])).
fof(f50076,plain,(
  ( ! [X26] : (s(bool,'T_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | ~p(s(bool,X26))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47739,f3327])).
fof(f50075,plain,(
  ( ! [X25] : (s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | ~p(s(bool,X25))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47739,f3285])).
fof(f50073,plain,(
  ( ! [X23] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | ~p(s(bool,X23))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47739,f3242])).
fof(f50072,plain,(
  ( ! [X22] : (sP9(X22,'F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | s(bool,'T_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47739,f3227])).
fof(f50071,plain,(
  ( ! [X21,X20] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | s(bool,'T_0') = s(bool,X20) | ~p(s(bool,X21))) ) | ($spl3 | $spl10 | $spl30 | $spl96)),
  inference(resolution,[],[f47739,f27175])).
fof(f50070,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,X19) | s(bool,'T_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | ($spl3 | $spl10 | $spl30)),
  inference(resolution,[],[f47739,f19799])).
fof(f50069,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,X16) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | ~p(s(bool,X17))) ) | ($spl3 | $spl10 | $spl30)),
  inference(resolution,[],[f47739,f19796])).
fof(f50067,plain,(
  ( ! [X12,X13] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | sP9(X12,X13,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | ~sP7(X13,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X12)) ) | ($spl3 | $spl10 | $spl30)),
  inference(resolution,[],[f47739,f19262])).
fof(f50066,plain,(
  ( ! [X10,X11] : (sP9(X10,X11,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | s(bool,'F_0') != s(bool,X11)) ) | ($spl3 | $spl10 | $spl30)),
  inference(resolution,[],[f47739,f19168])).
fof(f50065,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | p(s(bool,X9)) | s(bool,'T_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | ($spl3 | $spl10 | $spl30)),
  inference(resolution,[],[f47739,f19159])).
fof(f50064,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | p(s(bool,X7))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47739,f428])).
fof(f50086,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | ~sP7(X1,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X0)) ) | ($spl3 | $spl30)),
  inference(subsumption_resolution,[],[f50061,f439])).
fof(f50061,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | ~sP7(X1,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X0) | ~sP6('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X1,X0)) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47739,f425])).
fof(f50051,plain,(
  ( ! [X10] : (p(s(bool,X10)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47734,f927])).
fof(f50047,plain,(
  ( ! [X2,X3] : (sP5(X2,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X3) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | ~p(s(bool,X3))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47734,f407])).
fof(f50046,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | ~p(s(bool,X0)) | p(s(bool,X1))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47734,f411])).
fof(f50036,plain,(
  ( ! [X10] : (p(s(bool,X10)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47733,f926])).
fof(f50032,plain,(
  ( ! [X2,X3] : (sP5(X2,X3,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | ~p(s(bool,X3)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47733,f407])).
fof(f50031,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | p(s(bool,X1))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47733,f411])).
fof(f50030,plain,(
  ( ! [X6,X7] : (sP3('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X6,X7) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | ($spl3 | $spl30)),
  inference(subsumption_resolution,[],[f50020,f400])).
fof(f50020,plain,(
  ( ! [X6,X7] : (sP3('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X6,X7) | p(s(bool,X6)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47730,f392])).
fof(f50029,plain,(
  ( ! [X4,X5] : (sP3('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X4,X5) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | ($spl3 | $spl30)),
  inference(subsumption_resolution,[],[f50019,f399])).
fof(f50019,plain,(
  ( ! [X4,X5] : (sP3('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X4,X5) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | p(s(bool,X5))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47730,f391])).
fof(f50018,plain,(
  ( ! [X2,X3] : (sP3('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X2,X3) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47730,f586])).
fof(f50017,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | p(s(bool,X0)) | p(s(bool,X1))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47730,f394])).
fof(f50007,plain,(
  ( ! [X10] : (p(s(bool,X10)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47726,f897])).
fof(f50016,plain,(
  ( ! [X6,X7] : (sP1('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X6,X7) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | ($spl3 | $spl30)),
  inference(subsumption_resolution,[],[f50005,f383])).
fof(f50005,plain,(
  ( ! [X6,X7] : (sP1('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X6,X7) | p(s(bool,X6)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47726,f375])).
fof(f50015,plain,(
  ( ! [X4,X5] : (sP1('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X4,X5) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | ($spl3 | $spl30)),
  inference(subsumption_resolution,[],[f50004,f382])).
fof(f50004,plain,(
  ( ! [X4,X5] : (sP1('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X4,X5) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | ~p(s(bool,X5))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47726,f374])).
fof(f50003,plain,(
  ( ! [X2,X3] : (sP1('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X2,X3) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47726,f585])).
fof(f50002,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | p(s(bool,X0)) | ~p(s(bool,X1))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47726,f377])).
fof(f49955,plain,(
  ( ! [X10] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | p(s(bool,X10))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47725,f896])).
fof(f49954,plain,(
  ( ! [X8,X9] : (sP1(X8,X9,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | p(s(bool,X9)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47725,f376])).
fof(f49950,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | ($spl3 | $spl30)),
  inference(resolution,[],[f47725,f377])).
fof(f49844,plain,(
  ( ! [X325,X326] : (~p(s(bool,X325)) | p(s(bool,X326)) | sP9(X326,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X325)) ) | $spl510),
  inference(cnf_transformation,[],[f49844_D])).
fof(f49844_D,plain,(
  ( ! [X325,X326] : (~p(s(bool,X325)) | p(s(bool,X326)) | sP9(X326,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X325)) ) <=> ~$spl510),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl510])])).
fof(f49843,plain,(
  ( ! [X325,X324,X326] : (p(s(bool,X324)) | ~p(s(bool,X325)) | sP9(X326,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X325) | p(s(bool,X326))) ) | ($spl10 | $spl30)),
  inference(subsumption_resolution,[],[f49590,f553])).
fof(f49590,plain,(
  ( ! [X325,X324,X326] : (s(bool,'F_0') != s(bool,X324) | p(s(bool,X324)) | ~p(s(bool,X325)) | sP9(X326,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X325) | p(s(bool,X326))) ) | ($spl10 | $spl30)),
  inference(superposition,[],[f7606,f47686])).
fof(f47686,plain,(
  s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | $spl30),
  inference(backward_demodulation,[],[f3880,f806])).
fof(f806,plain,(
  s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))),
  inference(resolution,[],[f352,f553])).
fof(f48175,plain,(
  ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | $spl465),
  inference(cnf_transformation,[],[f48175_D])).
fof(f48175_D,plain,(
  ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) <=> ~$spl465),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl465])])).
fof(f49337,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | $spl30),
  inference(duplicate_literal_removal,[],[f49325])).
fof(f49325,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | ~p(s(bool,X4))) ) | $spl30),
  inference(resolution,[],[f47685,f459])).
fof(f459,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | ~sP10(X0,X1)) )),
  inference(cnf_transformation,[],[f286])).
fof(f286,plain,(
  ! [X0,X1] : ((~sP10(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP10(X0,X1)))),
  inference(rectify,[],[f285])).
fof(f285,plain,(
  ! [X1,X0] : ((~sP10(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP10(X1,X0)))),
  inference(flattening,[],[f284])).
fof(f284,plain,(
  ! [X1,X0] : ((~sP10(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP10(X1,X0)))),
  inference(nnf_transformation,[],[f234])).
fof(f234,plain,(
  ! [X1,X0] : (sP10(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f47685,plain,(
  ( ! [X122] : (sP10('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X122) | ~p(s(bool,X122))) ) | $spl30),
  inference(backward_demodulation,[],[f3880,f785])).
fof(f785,plain,(
  ( ! [X122] : (sP10('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X122) | ~p(s(bool,X122))) )),
  inference(resolution,[],[f352,f463])).
fof(f463,plain,(
  ( ! [X0,X1] : (sP10(X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f286])).
fof(f49338,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | $spl30),
  inference(duplicate_literal_removal,[],[f49322])).
fof(f49322,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | ~p(s(bool,X1)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | $spl30),
  inference(resolution,[],[f47685,f466])).
fof(f466,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | ~sP10(X1,X0)) )),
  inference(cnf_transformation,[],[f287])).
fof(f287,plain,(
  ! [X0,X1] : ((((p(s(bool,X1)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)))) | sP10(X1,X0)) & (~sP10(X1,X0) | ((~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X0)) | p(s(bool,X1))))))),
  inference(nnf_transformation,[],[f235])).
fof(f235,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> sP10(X1,X0))),
  inference(definition_folding,[],[f105,f234])).
fof(f105,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  inference(flattening,[],[f104])).
fof(f104,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X15,X16] : ((p(s(bool,X16)) <=> ~p(s(bool,X15))) <=> ((p(s(bool,X16)) | p(s(bool,X15))) & (~p(s(bool,X15)) | ~p(s(bool,X16)))))),
  file('thm.relation.EXTEND_RTC_TC','thm.sat.sat_29_prove_ltb')).
fof(f49302,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | $spl30),
  inference(duplicate_literal_removal,[],[f49290])).
fof(f49290,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | ~p(s(bool,X4)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | $spl30),
  inference(resolution,[],[f47684,f459])).
fof(f47684,plain,(
  ( ! [X121] : (sP10(X121,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | ~p(s(bool,X121))) ) | $spl30),
  inference(backward_demodulation,[],[f3880,f784])).
fof(f784,plain,(
  ( ! [X121] : (sP10(X121,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | ~p(s(bool,X121))) )),
  inference(resolution,[],[f352,f460])).
fof(f460,plain,(
  ( ! [X0,X1] : (sP10(X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f286])).
fof(f49303,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | $spl30),
  inference(duplicate_literal_removal,[],[f49287])).
fof(f49287,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | ~p(s(bool,X1))) ) | $spl30),
  inference(resolution,[],[f47684,f466])).
fof(f49273,plain,(
  ( ! [X2] : (p(s(bool,X2)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | $spl30),
  inference(duplicate_literal_removal,[],[f49261])).
fof(f49261,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,X2)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | $spl30),
  inference(resolution,[],[f47683,f460])).
fof(f47683,plain,(
  ( ! [X120] : (~sP10('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X120) | p(s(bool,X120))) ) | $spl30),
  inference(backward_demodulation,[],[f3880,f783])).
fof(f783,plain,(
  ( ! [X120] : (~sP10('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X120) | p(s(bool,X120))) )),
  inference(resolution,[],[f352,f458])).
fof(f458,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | ~sP10(X0,X1)) )),
  inference(cnf_transformation,[],[f286])).
fof(f49274,plain,(
  ( ! [X1] : (p(s(bool,X1)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | $spl30),
  inference(duplicate_literal_removal,[],[f49260])).
fof(f49260,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,X1)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | $spl30),
  inference(resolution,[],[f47683,f465])).
fof(f465,plain,(
  ( ! [X0,X1] : (sP10(X1,X0) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f287])).
fof(f49213,plain,(
  ( ! [X3] : (p(s(bool,X3)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | $spl30),
  inference(duplicate_literal_removal,[],[f49202])).
fof(f49202,plain,(
  ( ! [X3] : (p(s(bool,X3)) | p(s(bool,X3)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | $spl30),
  inference(resolution,[],[f47682,f463])).
fof(f47682,plain,(
  ( ! [X119] : (~sP10(X119,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | p(s(bool,X119))) ) | $spl30),
  inference(backward_demodulation,[],[f3880,f782])).
fof(f782,plain,(
  ( ! [X119] : (~sP10(X119,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | p(s(bool,X119))) )),
  inference(resolution,[],[f352,f458])).
fof(f49214,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | $spl30),
  inference(duplicate_literal_removal,[],[f49199])).
fof(f49199,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | p(s(bool,X0))) ) | $spl30),
  inference(resolution,[],[f47682,f464])).
fof(f464,plain,(
  ( ! [X0,X1] : (sP10(X1,X0) | ~p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f287])).
fof(f49185,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | $spl508),
  inference(cnf_transformation,[],[f49185_D])).
fof(f49185_D,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) <=> ~$spl508),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl508])])).
fof(f48460,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | $spl499),
  inference(cnf_transformation,[],[f48460_D])).
fof(f48460_D,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) <=> ~$spl499),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl499])])).
fof(f49175,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'F_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | s(bool,X1) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | ($spl4 | $spl10 | $spl30)),
  inference(resolution,[],[f4735,f19603])).
fof(f4735,plain,(
  ( ! [X118,X117] : (sP6(X117,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X118)) ) | ($spl4 | $spl30)),
  inference(forward_demodulation,[],[f3685,f3880])).
fof(f3685,plain,(
  ( ! [X118,X117] : (sP6(X117,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X118)) ) | $spl4),
  inference(resolution,[],[f838,f438])).
fof(f49144,plain,(
  ( ! [X12] : (p(s(bool,X12)) | s(bool,'F_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | ($spl4 | $spl30)),
  inference(resolution,[],[f4732,f3286])).
fof(f4732,plain,(
  ( ! [X120,X119] : (sP6('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X119,X120)) ) | ($spl4 | $spl30)),
  inference(forward_demodulation,[],[f3686,f3880])).
fof(f3686,plain,(
  ( ! [X120,X119] : (sP6('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X119,X120)) ) | $spl4),
  inference(resolution,[],[f838,f439])).
fof(f49150,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | $spl506),
  inference(cnf_transformation,[],[f49150_D])).
fof(f49150_D,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) <=> ~$spl506),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl506])])).
fof(f49139,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | ($spl4 | $spl10 | $spl30)),
  inference(resolution,[],[f4732,f19474])).
fof(f49149,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,X1) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | $spl504),
  inference(cnf_transformation,[],[f49149_D])).
fof(f49149_D,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,X1) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) <=> ~$spl504),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl504])])).
fof(f49138,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'F_0') != s(bool,X1) | s(bool,X1) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | ($spl4 | $spl10 | $spl30)),
  inference(resolution,[],[f4732,f19603])).
fof(f48463,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | $spl502),
  inference(cnf_transformation,[],[f48463_D])).
fof(f48463_D,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) <=> ~$spl502),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl502])])).
fof(f48451,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | s(bool,'T_0') = s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | ~p(s(bool,X21))) ) | ($spl4 | $spl10 | $spl30 | $spl96)),
  inference(resolution,[],[f4730,f27175])).
fof(f4730,plain,(
  ( ! [X107,X108] : (sP8(X107,X108,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | ($spl4 | $spl30)),
  inference(forward_demodulation,[],[f3680,f3880])).
fof(f3680,plain,(
  ( ! [X107,X108] : (sP8(X107,X108,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | $spl4),
  inference(resolution,[],[f838,f429])).
fof(f48450,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | s(bool,'T_0') != s(bool,X19)) ) | ($spl4 | $spl10 | $spl30)),
  inference(resolution,[],[f4730,f19799])).
fof(f48449,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | ~p(s(bool,X16)) | ~p(s(bool,X17))) ) | ($spl4 | $spl10 | $spl30)),
  inference(resolution,[],[f4730,f19796])).
fof(f48461,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X11) | ~p(s(bool,X11))) ) | $spl500),
  inference(cnf_transformation,[],[f48461_D])).
fof(f48461_D,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X11) | ~p(s(bool,X11))) ) <=> ~$spl500),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl500])])).
fof(f48446,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X11) | ~p(s(bool,X11)) | s(bool,'F_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | ($spl4 | $spl10 | $spl30)),
  inference(resolution,[],[f4730,f19168])).
fof(f4751,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | $spl42),
  inference(cnf_transformation,[],[f4751_D])).
fof(f4751_D,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f48431,plain,(
  ( ! [X3] : (p(s(bool,X3)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | $spl30),
  inference(duplicate_literal_removal,[],[f48417])).
fof(f48417,plain,(
  ( ! [X3] : (p(s(bool,X3)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | p(s(bool,X3))) ) | $spl30),
  inference(resolution,[],[f47632,f458])).
fof(f47632,plain,(
  ( ! [X128] : (sP10(X128,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | p(s(bool,X128))) ) | $spl30),
  inference(backward_demodulation,[],[f3880,f662])).
fof(f662,plain,(
  ( ! [X128] : (sP10(X128,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) | p(s(bool,X128))) )),
  inference(resolution,[],[f351,f463])).
fof(f48432,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | $spl30),
  inference(duplicate_literal_removal,[],[f48416])).
fof(f48416,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,X2)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | $spl30),
  inference(resolution,[],[f47632,f467])).
fof(f467,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | ~sP10(X1,X0)) )),
  inference(cnf_transformation,[],[f287])).
fof(f48402,plain,(
  ( ! [X3] : (p(s(bool,X3)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | $spl30),
  inference(duplicate_literal_removal,[],[f48388])).
fof(f48388,plain,(
  ( ! [X3] : (p(s(bool,X3)) | p(s(bool,X3)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | $spl30),
  inference(resolution,[],[f47631,f458])).
fof(f47631,plain,(
  ( ! [X127] : (sP10('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X127) | p(s(bool,X127))) ) | $spl30),
  inference(backward_demodulation,[],[f3880,f661])).
fof(f661,plain,(
  ( ! [X127] : (sP10('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X127) | p(s(bool,X127))) )),
  inference(resolution,[],[f351,f460])).
fof(f48403,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | $spl30),
  inference(duplicate_literal_removal,[],[f48387])).
fof(f48387,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | p(s(bool,X2))) ) | $spl30),
  inference(resolution,[],[f47631,f467])).
fof(f48368,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | $spl30),
  inference(duplicate_literal_removal,[],[f48355])).
fof(f48355,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | ~p(s(bool,X2))) ) | $spl30),
  inference(resolution,[],[f47630,f460])).
fof(f47630,plain,(
  ( ! [X126] : (~sP10(X126,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | ~p(s(bool,X126))) ) | $spl30),
  inference(backward_demodulation,[],[f3880,f660])).
fof(f660,plain,(
  ( ! [X126] : (~sP10(X126,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) | ~p(s(bool,X126))) )),
  inference(resolution,[],[f351,f459])).
fof(f48369,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | $spl30),
  inference(duplicate_literal_removal,[],[f48354])).
fof(f48354,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | ~p(s(bool,X1))) ) | $spl30),
  inference(resolution,[],[f47630,f465])).
fof(f48336,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | $spl30),
  inference(duplicate_literal_removal,[],[f48324])).
fof(f48324,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | ~p(s(bool,X3))) ) | $spl30),
  inference(resolution,[],[f47629,f463])).
fof(f47629,plain,(
  ( ! [X125] : (~sP10('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X125) | ~p(s(bool,X125))) ) | $spl30),
  inference(backward_demodulation,[],[f3880,f659])).
fof(f659,plain,(
  ( ! [X125] : (~sP10('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X125) | ~p(s(bool,X125))) )),
  inference(resolution,[],[f351,f459])).
fof(f48337,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | $spl30),
  inference(duplicate_literal_removal,[],[f48321])).
fof(f48321,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | $spl30),
  inference(resolution,[],[f47629,f464])).
fof(f48312,plain,(
  ( ! [X12] : (p(s(bool,X12)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | $spl30),
  inference(resolution,[],[f47676,f943])).
fof(f47676,plain,(
  ( ! [X107,X108] : (sP7('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X107,X108)) ) | $spl30),
  inference(backward_demodulation,[],[f3880,f776])).
fof(f776,plain,(
  ( ! [X107,X108] : (sP7('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X107,X108)) )),
  inference(resolution,[],[f352,f435])).
fof(f48308,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | ~p(s(bool,X4)) | p(s(bool,X5))) ) | $spl30),
  inference(resolution,[],[f47676,f432])).
fof(f48317,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | $spl496),
  inference(cnf_transformation,[],[f48317_D])).
fof(f48317_D,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) <=> ~$spl496),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl496])])).
fof(f48307,plain,(
  ( ! [X2,X3] : (~p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | ($spl10 | $spl30)),
  inference(resolution,[],[f47676,f19313])).
fof(f48263,plain,(
  ( ! [X27] : (s(bool,'T_0') != s(bool,X27) | s(bool,X27) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | $spl488),
  inference(cnf_transformation,[],[f48263_D])).
fof(f48263_D,plain,(
  ( ! [X27] : (s(bool,'T_0') != s(bool,X27) | s(bool,X27) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) <=> ~$spl488),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl488])])).
fof(f48306,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X1) | s(bool,X1) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | ($spl10 | $spl30)),
  inference(resolution,[],[f47676,f19575])).
fof(f48296,plain,(
  ( ! [X12] : (p(s(bool,X12)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | $spl30),
  inference(resolution,[],[f47675,f942])).
fof(f47675,plain,(
  ( ! [X105,X106] : (sP7(X105,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X106)) ) | $spl30),
  inference(backward_demodulation,[],[f3880,f775])).
fof(f775,plain,(
  ( ! [X105,X106] : (sP7(X105,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X106)) )),
  inference(resolution,[],[f352,f434])).
fof(f48304,plain,(
  ( ! [X6,X7] : (sP9(X6,X7,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | $spl494),
  inference(cnf_transformation,[],[f48304_D])).
fof(f48304_D,plain,(
  ( ! [X6,X7] : (sP9(X6,X7,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) <=> ~$spl494),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl494])])).
fof(f48303,plain,(
  ( ! [X6,X7] : (sP9(X6,X7,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | $spl30),
  inference(subsumption_resolution,[],[f48302,f47672])).
fof(f47672,plain,(
  ( ! [X99,X100] : (sP8(X99,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X100)) ) | $spl30),
  inference(backward_demodulation,[],[f3880,f772])).
fof(f772,plain,(
  ( ! [X99,X100] : (sP8(X99,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X100)) )),
  inference(resolution,[],[f352,f430])).
fof(f48302,plain,(
  ( ! [X6,X7] : (sP9(X6,X7,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | ~sP8(X6,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X7)) ) | $spl30),
  inference(subsumption_resolution,[],[f48293,f439])).
fof(f48293,plain,(
  ( ! [X6,X7] : (sP9(X6,X7,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | ~sP8(X6,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X7) | ~sP6('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X7,X6)) ) | $spl30),
  inference(resolution,[],[f47675,f425])).
fof(f48292,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | p(s(bool,X5))) ) | $spl30),
  inference(resolution,[],[f47675,f432])).
fof(f48301,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | $spl492),
  inference(cnf_transformation,[],[f48301_D])).
fof(f48301_D,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) <=> ~$spl492),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl492])])).
fof(f48291,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | p(s(bool,X2)) | s(bool,X3) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | ($spl10 | $spl30)),
  inference(resolution,[],[f47675,f19313])).
fof(f48257,plain,(
  s(bool,'T_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | $spl479),
  inference(cnf_transformation,[],[f48257_D])).
fof(f48257_D,plain,(
  s(bool,'T_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) <=> ~$spl479),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl479])])).
fof(f48290,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | s(bool,X1) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | ($spl10 | $spl30)),
  inference(resolution,[],[f47675,f19575])).
fof(f48282,plain,(
  ( ! [X31] : (p(s(bool,X31)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | $spl30),
  inference(resolution,[],[f47673,f940])).
fof(f47673,plain,(
  ( ! [X101,X102] : (sP8('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X101,X102)) ) | $spl30),
  inference(backward_demodulation,[],[f3880,f773])).
fof(f773,plain,(
  ( ! [X101,X102] : (sP8('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X101,X102)) )),
  inference(resolution,[],[f352,f431])).
fof(f48281,plain,(
  ( ! [X30,X29] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | s(bool,X29) != s(bool,X30) | s(bool,'T_0') != s(bool,X29)) ) | ($spl10 | $spl30)),
  inference(resolution,[],[f47673,f19111])).
fof(f48279,plain,(
  ( ! [X26] : (s(bool,'T_0') != s(bool,X26) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | $spl30),
  inference(resolution,[],[f47673,f3327])).
fof(f48278,plain,(
  ( ! [X25] : (s(bool,'F_0') = s(bool,X25) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | $spl30),
  inference(resolution,[],[f47673,f3285])).
fof(f48276,plain,(
  ( ! [X23] : (~p(s(bool,X23)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | $spl30),
  inference(resolution,[],[f47673,f3242])).
fof(f48274,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | s(bool,'T_0') = s(bool,X21) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | ($spl10 | $spl30 | $spl96)),
  inference(resolution,[],[f47673,f27175])).
fof(f48273,plain,(
  ( ! [X19,X18] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | s(bool,'F_0') != s(bool,X18) | s(bool,'T_0') != s(bool,X19)) ) | ($spl10 | $spl30)),
  inference(resolution,[],[f47673,f19799])).
fof(f48272,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,X16) | ~p(s(bool,X17)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | ($spl10 | $spl30)),
  inference(resolution,[],[f47673,f19796])).
fof(f48267,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | ~p(s(bool,X6)) | p(s(bool,X7))) ) | $spl30),
  inference(resolution,[],[f47673,f428])).
fof(f48174,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | $spl464),
  inference(cnf_transformation,[],[f48174_D])).
fof(f48174_D,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) <=> ~$spl464),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl464])])).
fof(f48289,plain,(
  ( ! [X4,X5] : (sP9('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X4,X5) | ~sP6(X5,X4,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | ~sP7(X4,X5,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | $spl490),
  inference(cnf_transformation,[],[f48289_D])).
fof(f48289_D,plain,(
  ( ! [X4,X5] : (sP9('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X4,X5) | ~sP6(X5,X4,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | ~sP7(X4,X5,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) <=> ~$spl490),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl490])])).
fof(f48266,plain,(
  ( ! [X4,X5] : (sP9('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X4,X5) | p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | ~sP7(X4,X5,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | ~sP6(X5,X4,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | $spl30),
  inference(resolution,[],[f47673,f427])).
fof(f48248,plain,(
  ( ! [X31] : (p(s(bool,X31)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | $spl30),
  inference(resolution,[],[f47672,f939])).
fof(f48247,plain,(
  ( ! [X30,X29] : (~p(s(bool,X29)) | s(bool,X30) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | s(bool,'T_0') != s(bool,X30)) ) | ($spl10 | $spl30)),
  inference(resolution,[],[f47672,f19111])).
fof(f48260,plain,(
  ( ! [X22] : (sP9(X22,'F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | $spl484),
  inference(cnf_transformation,[],[f48260_D])).
fof(f48260_D,plain,(
  ( ! [X22] : (sP9(X22,'F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) <=> ~$spl484),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl484])])).
fof(f48246,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') != s(bool,X27) | sP9(X28,'F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | s(bool,X27) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | ($spl10 | $spl30)),
  inference(resolution,[],[f47672,f19068])).
fof(f48245,plain,(
  ( ! [X26] : (s(bool,'T_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | ~p(s(bool,X26))) ) | $spl30),
  inference(resolution,[],[f47672,f3327])).
fof(f48262,plain,(
  s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | $spl486),
  inference(cnf_transformation,[],[f48262_D])).
fof(f48262_D,plain,(
  s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) <=> ~$spl486),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl486])])).
fof(f48244,plain,(
  ( ! [X25] : (s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | ~p(s(bool,X25))) ) | $spl30),
  inference(resolution,[],[f47672,f3285])).
fof(f48242,plain,(
  ( ! [X23] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | ~p(s(bool,X23))) ) | $spl30),
  inference(resolution,[],[f47672,f3242])).
fof(f48241,plain,(
  ( ! [X22] : (sP9(X22,'F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | s(bool,'T_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | $spl30),
  inference(resolution,[],[f47672,f3227])).
fof(f48240,plain,(
  ( ! [X21,X20] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | s(bool,'T_0') = s(bool,X20) | ~p(s(bool,X21))) ) | ($spl10 | $spl30 | $spl96)),
  inference(resolution,[],[f47672,f27175])).
fof(f48239,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,X19) | s(bool,'T_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | ($spl10 | $spl30)),
  inference(resolution,[],[f47672,f19799])).
fof(f48238,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,X16) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | ~p(s(bool,X17))) ) | ($spl10 | $spl30)),
  inference(resolution,[],[f47672,f19796])).
fof(f48254,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | ~sP7(X1,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X0)) ) | $spl476),
  inference(cnf_transformation,[],[f48254_D])).
fof(f48254_D,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | ~sP7(X1,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X0)) ) <=> ~$spl476),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl476])])).
fof(f48236,plain,(
  ( ! [X12,X13] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | sP9(X12,X13,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | ~sP7(X13,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X12)) ) | ($spl10 | $spl30)),
  inference(resolution,[],[f47672,f19262])).
fof(f48259,plain,(
  ( ! [X10,X11] : (sP9(X10,X11,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | s(bool,'F_0') != s(bool,X11)) ) | $spl482),
  inference(cnf_transformation,[],[f48259_D])).
fof(f48259_D,plain,(
  ( ! [X10,X11] : (sP9(X10,X11,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | s(bool,'F_0') != s(bool,X11)) ) <=> ~$spl482),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl482])])).
fof(f48235,plain,(
  ( ! [X10,X11] : (sP9(X10,X11,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | s(bool,'F_0') != s(bool,X11)) ) | ($spl10 | $spl30)),
  inference(resolution,[],[f47672,f19168])).
fof(f48258,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | p(s(bool,X9))) ) | $spl480),
  inference(cnf_transformation,[],[f48258_D])).
fof(f48258_D,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | p(s(bool,X9))) ) <=> ~$spl480),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl480])])).
fof(f48234,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | p(s(bool,X9)) | s(bool,'T_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | ($spl10 | $spl30)),
  inference(resolution,[],[f47672,f19159])).
fof(f48233,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | p(s(bool,X7))) ) | $spl30),
  inference(resolution,[],[f47672,f428])).
fof(f48253,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | ~sP7(X1,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X0)) ) | $spl30),
  inference(subsumption_resolution,[],[f48230,f439])).
fof(f48230,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | ~sP7(X1,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X0) | ~sP6('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X1,X0)) ) | $spl30),
  inference(resolution,[],[f47672,f425])).
fof(f48222,plain,(
  ( ! [X10] : (p(s(bool,X10)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | $spl30),
  inference(resolution,[],[f47664,f927])).
fof(f47664,plain,(
  ( ! [X76,X75] : (sP4('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X75,X76)) ) | $spl30),
  inference(backward_demodulation,[],[f3880,f760])).
fof(f760,plain,(
  ( ! [X76,X75] : (sP4('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X75,X76)) )),
  inference(resolution,[],[f352,f414])).
fof(f48227,plain,(
  ( ! [X2,X3] : (sP5(X2,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X3) | ~p(s(bool,X3))) ) | $spl474),
  inference(cnf_transformation,[],[f48227_D])).
fof(f48227_D,plain,(
  ( ! [X2,X3] : (sP5(X2,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X3) | ~p(s(bool,X3))) ) <=> ~$spl474),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl474])])).
fof(f48218,plain,(
  ( ! [X2,X3] : (sP5(X2,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X3) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | ~p(s(bool,X3))) ) | $spl30),
  inference(resolution,[],[f47664,f407])).
fof(f48217,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | ~p(s(bool,X0)) | p(s(bool,X1))) ) | $spl30),
  inference(resolution,[],[f47664,f411])).
fof(f48209,plain,(
  ( ! [X10] : (p(s(bool,X10)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | $spl30),
  inference(resolution,[],[f47663,f926])).
fof(f47663,plain,(
  ( ! [X74,X73] : (sP4(X73,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X74)) ) | $spl30),
  inference(backward_demodulation,[],[f3880,f759])).
fof(f759,plain,(
  ( ! [X74,X73] : (sP4(X73,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X74)) )),
  inference(resolution,[],[f352,f413])).
fof(f48214,plain,(
  ( ! [X2,X3] : (sP5(X2,X3,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | ~p(s(bool,X3))) ) | $spl472),
  inference(cnf_transformation,[],[f48214_D])).
fof(f48214_D,plain,(
  ( ! [X2,X3] : (sP5(X2,X3,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | ~p(s(bool,X3))) ) <=> ~$spl472),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl472])])).
fof(f48205,plain,(
  ( ! [X2,X3] : (sP5(X2,X3,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | ~p(s(bool,X3)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | $spl30),
  inference(resolution,[],[f47663,f407])).
fof(f48204,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | p(s(bool,X1))) ) | $spl30),
  inference(resolution,[],[f47663,f411])).
fof(f48201,plain,(
  ( ! [X2,X3] : (sP3('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X2,X3)) ) | $spl470),
  inference(cnf_transformation,[],[f48201_D])).
fof(f48201_D,plain,(
  ( ! [X2,X3] : (sP3('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X2,X3)) ) <=> ~$spl470),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl470])])).
fof(f48203,plain,(
  ( ! [X6,X7] : (sP3('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X6,X7) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | $spl30),
  inference(subsumption_resolution,[],[f48195,f400])).
fof(f48195,plain,(
  ( ! [X6,X7] : (sP3('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X6,X7) | p(s(bool,X6)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | $spl30),
  inference(resolution,[],[f47654,f392])).
fof(f47654,plain,(
  ( ! [X47,X48] : (sP2('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X47,X48)) ) | $spl30),
  inference(backward_demodulation,[],[f3880,f746])).
fof(f746,plain,(
  ( ! [X47,X48] : (sP2('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X47,X48)) )),
  inference(resolution,[],[f352,f397])).
fof(f48202,plain,(
  ( ! [X4,X5] : (sP3('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X4,X5) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | $spl30),
  inference(subsumption_resolution,[],[f48194,f399])).
fof(f48194,plain,(
  ( ! [X4,X5] : (sP3('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X4,X5) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | p(s(bool,X5))) ) | $spl30),
  inference(resolution,[],[f47654,f391])).
fof(f48193,plain,(
  ( ! [X2,X3] : (sP3('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X2,X3) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | $spl30),
  inference(resolution,[],[f47654,f586])).
fof(f48192,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | p(s(bool,X0)) | p(s(bool,X1))) ) | $spl30),
  inference(resolution,[],[f47654,f394])).
fof(f48184,plain,(
  ( ! [X10] : (p(s(bool,X10)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | $spl30),
  inference(resolution,[],[f47645,f897])).
fof(f47645,plain,(
  ( ! [X15,X16] : (sP0('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X15,X16)) ) | $spl30),
  inference(backward_demodulation,[],[f3880,f730])).
fof(f730,plain,(
  ( ! [X15,X16] : (sP0('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X15,X16)) )),
  inference(resolution,[],[f352,f380])).
fof(f48189,plain,(
  ( ! [X2,X3] : (sP1('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X2,X3)) ) | $spl468),
  inference(cnf_transformation,[],[f48189_D])).
fof(f48189_D,plain,(
  ( ! [X2,X3] : (sP1('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X2,X3)) ) <=> ~$spl468),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl468])])).
fof(f48191,plain,(
  ( ! [X6,X7] : (sP1('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X6,X7) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | $spl30),
  inference(subsumption_resolution,[],[f48182,f383])).
fof(f48182,plain,(
  ( ! [X6,X7] : (sP1('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X6,X7) | p(s(bool,X6)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | $spl30),
  inference(resolution,[],[f47645,f375])).
fof(f48190,plain,(
  ( ! [X4,X5] : (sP1('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X4,X5) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | $spl30),
  inference(subsumption_resolution,[],[f48181,f382])).
fof(f48181,plain,(
  ( ! [X4,X5] : (sP1('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X4,X5) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | ~p(s(bool,X5))) ) | $spl30),
  inference(resolution,[],[f47645,f374])).
fof(f48180,plain,(
  ( ! [X2,X3] : (sP1('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X2,X3) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | $spl30),
  inference(resolution,[],[f47645,f585])).
fof(f48179,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | p(s(bool,X0)) | ~p(s(bool,X1))) ) | $spl30),
  inference(resolution,[],[f47645,f377])).
fof(f48169,plain,(
  ( ! [X10] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | p(s(bool,X10))) ) | $spl30),
  inference(resolution,[],[f47644,f896])).
fof(f47644,plain,(
  ( ! [X14,X13] : (sP0(X13,X14,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | $spl30),
  inference(backward_demodulation,[],[f3880,f729])).
fof(f729,plain,(
  ( ! [X14,X13] : (sP0(X13,X14,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) )),
  inference(resolution,[],[f352,f378])).
fof(f48178,plain,(
  ( ! [X8,X9] : (sP1(X8,X9,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | p(s(bool,X9))) ) | $spl466),
  inference(cnf_transformation,[],[f48178_D])).
fof(f48178_D,plain,(
  ( ! [X8,X9] : (sP1(X8,X9,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | p(s(bool,X9))) ) <=> ~$spl466),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl466])])).
fof(f48168,plain,(
  ( ! [X8,X9] : (sP1(X8,X9,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))) | p(s(bool,X9)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | $spl30),
  inference(resolution,[],[f47644,f376])).
fof(f48164,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | $spl30),
  inference(resolution,[],[f47644,f377])).
fof(f23055,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | $spl245),
  inference(cnf_transformation,[],[f23055_D])).
fof(f23055_D,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) <=> ~$spl245),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl245])])).
fof(f48155,plain,(
  ( ! [X12] : (p(s(bool,X12)) | s(bool,'F_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | $spl30),
  inference(resolution,[],[f47626,f3286])).
fof(f47626,plain,(
  ( ! [X120,X119] : (sP6('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X119,X120)) ) | $spl30),
  inference(backward_demodulation,[],[f3880,f656])).
fof(f656,plain,(
  ( ! [X120,X119] : (sP6('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X119,X120)) )),
  inference(resolution,[],[f351,f439])).
fof(f48151,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | p(s(bool,X4)) | p(s(bool,X5))) ) | $spl30),
  inference(resolution,[],[f47626,f436])).
fof(f436,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f276])).
fof(f23705,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | $spl254),
  inference(cnf_transformation,[],[f23705_D])).
fof(f23705_D,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) <=> ~$spl254),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl254])])).
fof(f48150,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | ($spl10 | $spl30)),
  inference(resolution,[],[f47626,f19474])).
fof(f23704,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,X1) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | $spl252),
  inference(cnf_transformation,[],[f23704_D])).
fof(f23704_D,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,X1) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) <=> ~$spl252),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl252])])).
fof(f48149,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'F_0') != s(bool,X1) | s(bool,X1) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | ($spl10 | $spl30)),
  inference(resolution,[],[f47626,f19603])).
fof(f48147,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X9)) ) | $spl462),
  inference(cnf_transformation,[],[f48147_D])).
fof(f48147_D,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X9)) ) <=> ~$spl462),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl462])])).
fof(f48146,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X9) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | $spl30),
  inference(subsumption_resolution,[],[f48145,f47620])).
fof(f47620,plain,(
  ( ! [X107,X108] : (sP8(X107,X108,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | $spl30),
  inference(backward_demodulation,[],[f3880,f650])).
fof(f650,plain,(
  ( ! [X107,X108] : (sP8(X107,X108,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) )),
  inference(resolution,[],[f351,f429])).
fof(f48145,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X9) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | ~sP8(X8,X9,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | $spl30),
  inference(subsumption_resolution,[],[f48137,f435])).
fof(f48137,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X9) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | ~sP8(X8,X9,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | ~sP7('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X9,X8)) ) | $spl30),
  inference(resolution,[],[f47625,f426])).
fof(f47625,plain,(
  ( ! [X118,X117] : (sP6(X117,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X118)) ) | $spl30),
  inference(backward_demodulation,[],[f3880,f655])).
fof(f655,plain,(
  ( ! [X118,X117] : (sP6(X117,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X118)) )),
  inference(resolution,[],[f351,f438])).
fof(f48135,plain,(
  ( ! [X4,X5] : (p(s(bool,X4)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | p(s(bool,X5))) ) | $spl30),
  inference(resolution,[],[f47625,f436])).
fof(f23689,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | $spl250),
  inference(cnf_transformation,[],[f23689_D])).
fof(f23689_D,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) <=> ~$spl250),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl250])])).
fof(f48134,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | p(s(bool,X2)) | s(bool,X3) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | ($spl10 | $spl30)),
  inference(resolution,[],[f47625,f19474])).
fof(f48133,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'F_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | s(bool,X1) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | ($spl10 | $spl30)),
  inference(resolution,[],[f47625,f19603])).
fof(f48122,plain,(
  ( ! [X12] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | s(bool,'F_0') != s(bool,X12)) ) | $spl30),
  inference(resolution,[],[f47624,f3286])).
fof(f47624,plain,(
  ( ! [X116,X115] : (sP6(X115,X116,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | $spl30),
  inference(backward_demodulation,[],[f3880,f654])).
fof(f654,plain,(
  ( ! [X116,X115] : (sP6(X115,X116,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) )),
  inference(resolution,[],[f351,f437])).
fof(f48132,plain,(
  ( ! [X10,X11] : (sP9('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X10,X11)) ) | $spl460),
  inference(cnf_transformation,[],[f48132_D])).
fof(f48132_D,plain,(
  ( ! [X10,X11] : (sP9('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X10,X11)) ) <=> ~$spl460),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl460])])).
fof(f48131,plain,(
  ( ! [X10,X11] : (sP9('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X10,X11) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | $spl30),
  inference(subsumption_resolution,[],[f48130,f47623])).
fof(f47623,plain,(
  ( ! [X114,X113] : (sP7(X113,X114,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | $spl30),
  inference(backward_demodulation,[],[f3880,f653])).
fof(f653,plain,(
  ( ! [X114,X113] : (sP7(X113,X114,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) )),
  inference(resolution,[],[f351,f433])).
fof(f48130,plain,(
  ( ! [X10,X11] : (sP9('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X10,X11) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | ~sP7(X10,X11,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | $spl30),
  inference(subsumption_resolution,[],[f48121,f431])).
fof(f48121,plain,(
  ( ! [X10,X11] : (sP9('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X10,X11) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | ~sP8('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X11,X10) | ~sP7(X10,X11,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | $spl30),
  inference(resolution,[],[f47624,f427])).
fof(f48118,plain,(
  ( ! [X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | $spl30),
  inference(resolution,[],[f47624,f436])).
fof(f21194,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | s(bool,X2) != s(bool,X3)) ) | $spl124),
  inference(cnf_transformation,[],[f21194_D])).
fof(f21194_D,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | s(bool,X2) != s(bool,X3)) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f48117,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | s(bool,X2) != s(bool,X3)) ) | ($spl10 | $spl30)),
  inference(resolution,[],[f47624,f19474])).
fof(f21193,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,X0) | s(bool,X0) != s(bool,X1)) ) | $spl122),
  inference(cnf_transformation,[],[f21193_D])).
fof(f21193_D,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,X0) | s(bool,X0) != s(bool,X1)) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f48116,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | s(bool,'F_0') != s(bool,X0) | s(bool,X0) != s(bool,X1)) ) | ($spl10 | $spl30)),
  inference(resolution,[],[f47624,f19603])).
fof(f48106,plain,(
  ( ! [X13] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | ~p(s(bool,X13))) ) | $spl30),
  inference(resolution,[],[f47623,f942])).
fof(f48105,plain,(
  ( ! [X12] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | ~p(s(bool,X12))) ) | $spl30),
  inference(resolution,[],[f47623,f943])).
fof(f48115,plain,(
  ( ! [X10,X11] : (sP9('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X10,X11) | ~sP6(X11,X10,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | $spl458),
  inference(cnf_transformation,[],[f48115_D])).
fof(f48115_D,plain,(
  ( ! [X10,X11] : (sP9('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X10,X11) | ~sP6(X11,X10,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) <=> ~$spl458),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl458])])).
fof(f48114,plain,(
  ( ! [X10,X11] : (sP9('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X10,X11) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | ~sP6(X11,X10,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | $spl30),
  inference(subsumption_resolution,[],[f48104,f431])).
fof(f48104,plain,(
  ( ! [X10,X11] : (sP9('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X10,X11) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | ~sP8('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X11,X10) | ~sP6(X11,X10,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | $spl30),
  inference(resolution,[],[f47623,f427])).
fof(f48101,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | ~p(s(bool,X5)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | $spl30),
  inference(resolution,[],[f47623,f432])).
fof(f21177,plain,(
  ( ! [X2,X3] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,X3)) ) | $spl120),
  inference(cnf_transformation,[],[f21177_D])).
fof(f21177_D,plain,(
  ( ! [X2,X3] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,X3)) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f48100,plain,(
  ( ! [X2,X3] : (~p(s(bool,X2)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | s(bool,X2) != s(bool,X3)) ) | ($spl10 | $spl30)),
  inference(resolution,[],[f47623,f19313])).
fof(f48099,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | s(bool,'T_0') != s(bool,X0) | s(bool,X0) != s(bool,X1)) ) | ($spl10 | $spl30)),
  inference(resolution,[],[f47623,f19575])).
fof(f48088,plain,(
  ( ! [X23] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | ~p(s(bool,X23))) ) | $spl30),
  inference(resolution,[],[f47620,f939])).
fof(f48087,plain,(
  ( ! [X22] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | ~p(s(bool,X22))) ) | $spl30),
  inference(resolution,[],[f47620,f940])).
fof(f27235,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | $spl354),
  inference(cnf_transformation,[],[f27235_D])).
fof(f27235_D,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) <=> ~$spl354),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl354])])).
fof(f48086,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | s(bool,'T_0') = s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | ~p(s(bool,X21))) ) | ($spl10 | $spl30 | $spl96)),
  inference(resolution,[],[f47620,f27175])).
fof(f48085,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | s(bool,'T_0') != s(bool,X19)) ) | ($spl10 | $spl30)),
  inference(resolution,[],[f47620,f19799])).
fof(f48084,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | ~p(s(bool,X16)) | ~p(s(bool,X17))) ) | ($spl10 | $spl30)),
  inference(resolution,[],[f47620,f19796])).
fof(f48097,plain,(
  ( ! [X2,X3] : (sP9(X2,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X3) | ~sP6(X3,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X2)) ) | $spl454),
  inference(cnf_transformation,[],[f48097_D])).
fof(f48097_D,plain,(
  ( ! [X2,X3] : (sP9(X2,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X3) | ~sP6(X3,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X2)) ) <=> ~$spl454),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl454])])).
fof(f48083,plain,(
  ( ! [X14,X15] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | sP9(X14,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X15) | ~sP6(X15,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X14)) ) | ($spl10 | $spl30)),
  inference(resolution,[],[f47620,f19385])).
fof(f23056,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X11) | ~p(s(bool,X11))) ) | $spl246),
  inference(cnf_transformation,[],[f23056_D])).
fof(f23056_D,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X11) | ~p(s(bool,X11))) ) <=> ~$spl246),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl246])])).
fof(f48081,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X11) | ~p(s(bool,X11)) | s(bool,'F_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | ($spl10 | $spl30)),
  inference(resolution,[],[f47620,f19168])).
fof(f48098,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X9) | s(bool,'T_0') != s(bool,X9)) ) | $spl456),
  inference(cnf_transformation,[],[f48098_D])).
fof(f48098_D,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X9) | s(bool,'T_0') != s(bool,X9)) ) <=> ~$spl456),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl456])])).
fof(f48080,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X9) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | s(bool,'T_0') != s(bool,X9)) ) | ($spl10 | $spl30)),
  inference(resolution,[],[f47620,f19159])).
fof(f48079,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | $spl30),
  inference(resolution,[],[f47620,f428])).
fof(f48096,plain,(
  ( ! [X2,X3] : (sP9(X2,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X3) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | ~sP6(X3,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X2)) ) | $spl30),
  inference(subsumption_resolution,[],[f48077,f435])).
fof(f48077,plain,(
  ( ! [X2,X3] : (sP9(X2,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X3) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | ~sP7('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X3,X2) | ~sP6(X3,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X2)) ) | $spl30),
  inference(resolution,[],[f47620,f426])).
fof(f48066,plain,(
  ( ! [X11] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | ~p(s(bool,X11))) ) | $spl30),
  inference(resolution,[],[f47610,f926])).
fof(f47610,plain,(
  ( ! [X81,X82] : (sP4(X81,X82,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | $spl30),
  inference(backward_demodulation,[],[f3880,f637])).
fof(f637,plain,(
  ( ! [X81,X82] : (sP4(X81,X82,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) )),
  inference(resolution,[],[f351,f412])).
fof(f412,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f263])).
fof(f48065,plain,(
  ( ! [X10] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | ~p(s(bool,X10))) ) | $spl30),
  inference(resolution,[],[f47610,f927])).
fof(f48074,plain,(
  ( ! [X4,X5] : (sP5('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X4,X5)) ) | $spl452),
  inference(cnf_transformation,[],[f48074_D])).
fof(f48074_D,plain,(
  ( ! [X4,X5] : (sP5('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X4,X5)) ) <=> ~$spl452),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl452])])).
fof(f48064,plain,(
  ( ! [X8,X9] : (sP5('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X8,X9) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | $spl30),
  inference(resolution,[],[f47610,f587])).
fof(f587,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X0)) | ~sP4(X1,X2,X0)) )),
  inference(duplicate_literal_removal,[],[f410])).
fof(f410,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X0)) | ~sP4(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f261])).
fof(f48075,plain,(
  ( ! [X6,X7] : (sP5('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X6,X7) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | $spl30),
  inference(subsumption_resolution,[],[f48063,f415])).
fof(f415,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f265])).
fof(f265,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP5(X2,X0,X1)) & (~sP5(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f264])).
fof(f264,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP5(X2,X0,X1)) & (~sP5(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f228])).
fof(f228,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP5(X2,X0,X1))),
  inference(definition_folding,[],[f96,f227,f226])).
fof(f96,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f95])).
fof(f95,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X14,X15,X16] : ((p(s(bool,X16)) <=> (p(s(bool,X15)) & p(s(bool,X14)))) <=> ((p(s(bool,X16)) | ~p(s(bool,X15)) | ~p(s(bool,X14))) & (p(s(bool,X15)) | ~p(s(bool,X16))) & (p(s(bool,X14)) | ~p(s(bool,X16)))))),
  file('thm.relation.EXTEND_RTC_TC','thm.sat.sat_23_prove_ltb')).
fof(f48063,plain,(
  ( ! [X6,X7] : (sP5('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X6,X7) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | ~p(s(bool,X7))) ) | $spl30),
  inference(resolution,[],[f47610,f409])).
fof(f409,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X2)) | ~sP4(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f261])).
fof(f48073,plain,(
  ( ! [X4,X5] : (sP5('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X4,X5) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | $spl30),
  inference(subsumption_resolution,[],[f48062,f416])).
fof(f416,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f265])).
fof(f48062,plain,(
  ( ! [X4,X5] : (sP5('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X4,X5) | ~p(s(bool,X4)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | $spl30),
  inference(resolution,[],[f47610,f408])).
fof(f408,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~sP4(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f261])).
fof(f48060,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | $spl30),
  inference(resolution,[],[f47610,f411])).
fof(f48051,plain,(
  ( ! [X10] : (p(s(bool,X10)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | $spl30),
  inference(resolution,[],[f47600,f910])).
fof(f910,plain,(
  ( ! [X45,X46] : (~sP2(happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X45,X46) | p(s(bool,X46)) | p(s(bool,X45))) )),
  inference(resolution,[],[f350,f394])).
fof(f47600,plain,(
  ( ! [X50,X49] : (sP2(X49,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X50)) ) | $spl30),
  inference(backward_demodulation,[],[f3880,f621])).
fof(f621,plain,(
  ( ! [X50,X49] : (sP2(X49,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X50)) )),
  inference(resolution,[],[f351,f396])).
fof(f396,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f256])).
fof(f21015,plain,(
  ( ! [X8,X9] : (sP3(X8,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X9) | p(s(bool,X9))) ) | $spl110),
  inference(cnf_transformation,[],[f21015_D])).
fof(f21015_D,plain,(
  ( ! [X8,X9] : (sP3(X8,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X9) | p(s(bool,X9))) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f48050,plain,(
  ( ! [X8,X9] : (sP3(X8,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X9) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | p(s(bool,X9))) ) | $spl30),
  inference(resolution,[],[f47600,f393])).
fof(f393,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f253])).
fof(f48046,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | p(s(bool,X1))) ) | $spl30),
  inference(resolution,[],[f47600,f394])).
fof(f48037,plain,(
  ( ! [X10] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | p(s(bool,X10))) ) | $spl30),
  inference(resolution,[],[f47599,f910])).
fof(f47599,plain,(
  ( ! [X47,X48] : (sP2(X47,X48,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | $spl30),
  inference(backward_demodulation,[],[f3880,f620])).
fof(f620,plain,(
  ( ! [X47,X48] : (sP2(X47,X48,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) )),
  inference(resolution,[],[f351,f395])).
fof(f395,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f256])).
fof(f21000,plain,(
  ( ! [X8,X9] : (sP3(X8,X9,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | p(s(bool,X9))) ) | $spl108),
  inference(cnf_transformation,[],[f21000_D])).
fof(f21000_D,plain,(
  ( ! [X8,X9] : (sP3(X8,X9,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | p(s(bool,X9))) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f48036,plain,(
  ( ! [X8,X9] : (sP3(X8,X9,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | p(s(bool,X9)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | $spl30),
  inference(resolution,[],[f47599,f393])).
fof(f48032,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | $spl30),
  inference(resolution,[],[f47599,f394])).
fof(f21024,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | $spl112),
  inference(cnf_transformation,[],[f21024_D])).
fof(f21024_D,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f48031,plain,(
  ( ! [X1] : (p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | p(s(bool,X1))) ) | ($spl6 | $spl30)),
  inference(subsumption_resolution,[],[f48025,f538])).
fof(f48025,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | p(s(bool,X1))) ) | ($spl6 | $spl30)),
  inference(resolution,[],[f20933,f474])).
fof(f20933,plain,(
  ( ! [X127] : (sP11('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)),X127)) ) | ($spl6 | $spl30)),
  inference(forward_demodulation,[],[f3836,f3880])).
fof(f3836,plain,(
  ( ! [X127] : (sP11('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X127)) ) | $spl6),
  inference(resolution,[],[f841,f476])).
fof(f476,plain,(
  ( ! [X0,X1] : (sP11(X0,X1) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f291])).
fof(f48030,plain,(
  ( ! [X0] : (p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | p(s(bool,X0))) ) | ($spl6 | $spl30)),
  inference(subsumption_resolution,[],[f48024,f538])).
fof(f48024,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | p(s(bool,X0))) ) | ($spl6 | $spl30)),
  inference(resolution,[],[f20933,f478])).
fof(f48023,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | ($spl6 | $spl30)),
  inference(subsumption_resolution,[],[f48017,f538])).
fof(f48017,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | ($spl6 | $spl30)),
  inference(resolution,[],[f20932,f474])).
fof(f20932,plain,(
  ( ! [X126] : (sP11(X126,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) ) | ($spl6 | $spl30)),
  inference(forward_demodulation,[],[f3835,f3880])).
fof(f3835,plain,(
  ( ! [X126] : (sP11(X126,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | $spl6),
  inference(resolution,[],[f841,f475])).
fof(f48022,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | ($spl6 | $spl30)),
  inference(subsumption_resolution,[],[f48016,f538])).
fof(f48016,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | ($spl6 | $spl30)),
  inference(resolution,[],[f20932,f478])).
fof(f48004,plain,(
  ( ! [X11] : (~p(s(bool,X11)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | $spl30),
  inference(resolution,[],[f47589,f896])).
fof(f47589,plain,(
  ( ! [X21,X22] : (sP0(X21,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X22)) ) | $spl30),
  inference(backward_demodulation,[],[f3880,f607])).
fof(f607,plain,(
  ( ! [X21,X22] : (sP0(X21,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X22)) )),
  inference(resolution,[],[f351,f379])).
fof(f379,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f248])).
fof(f48003,plain,(
  ( ! [X10] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | ~p(s(bool,X10))) ) | $spl30),
  inference(resolution,[],[f47589,f897])).
fof(f20967,plain,(
  ( ! [X8,X9] : (sP1(X8,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X9) | ~p(s(bool,X9))) ) | $spl106),
  inference(cnf_transformation,[],[f20967_D])).
fof(f20967_D,plain,(
  ( ! [X8,X9] : (sP1(X8,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X9) | ~p(s(bool,X9))) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f48002,plain,(
  ( ! [X8,X9] : (sP1(X8,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X9) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | ~p(s(bool,X9))) ) | $spl30),
  inference(resolution,[],[f47589,f376])).
fof(f47998,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | ~p(s(bool,X1))) ) | $spl30),
  inference(resolution,[],[f47589,f377])).
fof(f47975,plain,(
  ( ! [X1] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | p(s(bool,X1))) ) | $spl30),
  inference(subsumption_resolution,[],[f47967,f538])).
fof(f47967,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | p(s(bool,X1))) ) | $spl30),
  inference(resolution,[],[f47634,f474])).
fof(f47634,plain,(
  ( ! [X134] : (sP11('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X134)) ) | $spl30),
  inference(backward_demodulation,[],[f3880,f668])).
fof(f668,plain,(
  ( ! [X134] : (sP11('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X134)) )),
  inference(resolution,[],[f351,f476])).
fof(f47974,plain,(
  ( ! [X0] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | p(s(bool,X0))) ) | $spl30),
  inference(subsumption_resolution,[],[f47966,f538])).
fof(f47966,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | p(s(bool,X0))) ) | $spl30),
  inference(resolution,[],[f47634,f478])).
fof(f47965,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | $spl30),
  inference(subsumption_resolution,[],[f47957,f538])).
fof(f47957,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | $spl30),
  inference(resolution,[],[f47633,f474])).
fof(f47633,plain,(
  ( ! [X133] : (sP11(X133,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | $spl30),
  inference(backward_demodulation,[],[f3880,f667])).
fof(f667,plain,(
  ( ! [X133] : (sP11(X133,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) )),
  inference(resolution,[],[f351,f475])).
fof(f47964,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | $spl30),
  inference(subsumption_resolution,[],[f47956,f538])).
fof(f47956,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | $spl30),
  inference(resolution,[],[f47633,f478])).
fof(f840,plain,(
  ~p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | $spl7),
  inference(cnf_transformation,[],[f840_D])).
fof(f840_D,plain,(
  ~p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f3880,plain,(
  s(sK12,sK14) = s(sK12,sK16) | $spl30),
  inference(cnf_transformation,[],[f3880_D])).
fof(f3880_D,plain,(
  s(sK12,sK14) = s(sK12,sK16) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f36476,plain,(
  s(sK12,sK14) = s(sK12,sK16) | ~p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))),
  inference(resolution,[],[f721,f578])).
fof(f839,plain,(
  ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | $spl5),
  inference(cnf_transformation,[],[f839_D])).
fof(f839_D,plain,(
  ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f841,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | $spl6),
  inference(cnf_transformation,[],[f841_D])).
fof(f841_D,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f22818,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | $spl133),
  inference(resolution,[],[f21243,f355])).
fof(f23440,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))),
  inference(subsumption_resolution,[],[f23095,f538])).
fof(f23095,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))),
  inference(superposition,[],[f355,f806])).
fof(f29708,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))),
  inference(resolution,[],[f29568,f355])).
fof(f29568,plain,(
  ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))),
  inference(subsumption_resolution,[],[f29567,f588])).
fof(f29567,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))),
  inference(forward_demodulation,[],[f29557,f680])).
fof(f680,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))),
  inference(resolution,[],[f351,f557])).
fof(f29557,plain,(
  ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))),
  inference(resolution,[],[f784,f659])).
fof(f30332,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))),
  inference(subsumption_resolution,[],[f29987,f538])).
fof(f29987,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))),
  inference(superposition,[],[f355,f806])).
fof(f47572,plain,(
  $false | $spl7),
  inference(subsumption_resolution,[],[f47571,f571])).
fof(f47571,plain,(
  p(s(bool,'F_0')) | $spl7),
  inference(forward_demodulation,[],[f47570,f22776])).
fof(f22776,plain,(
  s(bool,'F_0') = s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | $spl7),
  inference(resolution,[],[f12719,f553])).
fof(f12719,plain,(
  ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | $spl7),
  inference(subsumption_resolution,[],[f12718,f567])).
fof(f12718,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | $spl7),
  inference(forward_demodulation,[],[f12602,f806])).
fof(f12602,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | $spl7),
  inference(resolution,[],[f840,f355])).
fof(f47570,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))),
  inference(subsumption_resolution,[],[f47541,f588])).
fof(f47541,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))),
  inference(superposition,[],[f595,f2238])).
fof(f2238,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))),
  inference(resolution,[],[f1625,f557])).
fof(f1625,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))),
  inference(resolution,[],[f884,f356])).
fof(f356,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f180])).
fof(f180,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))))),
  inference(ennf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) => p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))))),
  inference(flattening,[],[f79])).
fof(f79,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) => p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6] : ! [X7,X8,X9] : (p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X8),s(X6,X9)))) => p(s(bool,'const.relation.RTC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X8),s(X6,X9)))))),
  file('thm.relation.EXTEND_RTC_TC','thm.relation.TC_RTC')).
fof(f595,plain,(
  ( ! [X2] : (~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,X2),s(sK12,sK15)))) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,X2),s(sK12,sK16))))) )),
  inference(resolution,[],[f351,f364])).
fof(f364,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f185])).
fof(f185,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | ! [X4] : (~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4))))))),
  inference(ennf_transformation,[],[f86])).
fof(f86,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) => ! [X4] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) => p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4))))))),
  inference(flattening,[],[f85])).
fof(f85,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) => ! [X4] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) => p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4))))))),
  inference(rectify,[],[f50])).
fof(f50,axiom,(
  ! [X6] : ! [X7,X8,X9] : (p(s(bool,'const.relation.RTC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X8),s(X6,X9)))) => ! [X10] : (p(s(bool,'const.relation.RTC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X9),s(X6,X10)))) => p(s(bool,'const.relation.RTC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X8),s(X6,X10))))))),
  file('thm.relation.EXTEND_RTC_TC','thm.relation.RTC_RTC')).
fof(f36418,plain,(
  p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) | $spl420),
  inference(cnf_transformation,[],[f36418_D])).
fof(f36418_D,plain,(
  p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) <=> ~$spl420),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl420])])).
fof(f42441,plain,(
  ( ! [X3] : (p(s(bool,X3)) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))))) )),
  inference(duplicate_literal_removal,[],[f42418])).
fof(f42418,plain,(
  ( ! [X3] : (p(s(bool,X3)) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) | p(s(bool,X3))) )),
  inference(resolution,[],[f953,f458])).
fof(f953,plain,(
  ( ! [X128] : (sP10(X128,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))) | p(s(bool,X128))) )),
  inference(resolution,[],[f350,f463])).
fof(f42442,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))))) )),
  inference(duplicate_literal_removal,[],[f42417])).
fof(f42417,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,X2)) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))))) )),
  inference(resolution,[],[f953,f467])).
fof(f42337,plain,(
  ( ! [X3] : (p(s(bool,X3)) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))))) )),
  inference(duplicate_literal_removal,[],[f42314])).
fof(f42314,plain,(
  ( ! [X3] : (p(s(bool,X3)) | p(s(bool,X3)) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))))) )),
  inference(resolution,[],[f952,f458])).
fof(f952,plain,(
  ( ! [X127] : (sP10(happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X127) | p(s(bool,X127))) )),
  inference(resolution,[],[f350,f460])).
fof(f42338,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))))) )),
  inference(duplicate_literal_removal,[],[f42313])).
fof(f42313,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) | p(s(bool,X2))) )),
  inference(resolution,[],[f952,f467])).
fof(f42147,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))))) )),
  inference(duplicate_literal_removal,[],[f42125])).
fof(f42125,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) | ~p(s(bool,X2))) )),
  inference(resolution,[],[f951,f460])).
fof(f951,plain,(
  ( ! [X126] : (~sP10(X126,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))) | ~p(s(bool,X126))) )),
  inference(resolution,[],[f350,f459])).
fof(f42148,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))))) )),
  inference(duplicate_literal_removal,[],[f42124])).
fof(f42124,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) | ~p(s(bool,X1))) )),
  inference(resolution,[],[f951,f465])).
fof(f41959,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))))) )),
  inference(duplicate_literal_removal,[],[f41938])).
fof(f41938,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) | ~p(s(bool,X3))) )),
  inference(resolution,[],[f950,f463])).
fof(f950,plain,(
  ( ! [X125] : (~sP10(happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X125) | ~p(s(bool,X125))) )),
  inference(resolution,[],[f350,f459])).
fof(f41960,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))))) )),
  inference(duplicate_literal_removal,[],[f41935])).
fof(f41935,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))))) )),
  inference(resolution,[],[f950,f464])).
fof(f41723,plain,(
  ( ! [X24] : (p(s(bool,X24)) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))))) )),
  inference(resolution,[],[f779,f946])).
fof(f946,plain,(
  ( ! [X118,X117] : (sP6(X117,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X118)) )),
  inference(resolution,[],[f350,f438])).
fof(f779,plain,(
  ( ! [X114,X113] : (~sP6('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X113,X114) | p(s(bool,X114)) | p(s(bool,X113))) )),
  inference(resolution,[],[f352,f436])).
fof(f41714,plain,(
  ( ! [X14] : (p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) | p(s(bool,X14))) )),
  inference(resolution,[],[f779,f945])).
fof(f945,plain,(
  ( ! [X116,X115] : (sP6(X115,X116,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) )),
  inference(resolution,[],[f350,f437])).
fof(f41676,plain,(
  ( ! [X24] : (p(s(bool,X24)) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))))) )),
  inference(resolution,[],[f778,f947])).
fof(f947,plain,(
  ( ! [X120,X119] : (sP6(happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X119,X120)) )),
  inference(resolution,[],[f350,f439])).
fof(f778,plain,(
  ( ! [X111,X112] : (~sP6(X111,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X112) | p(s(bool,X112)) | p(s(bool,X111))) )),
  inference(resolution,[],[f352,f436])).
fof(f41667,plain,(
  ( ! [X14] : (p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) | p(s(bool,X14))) )),
  inference(resolution,[],[f778,f945])).
fof(f41630,plain,(
  ( ! [X27] : (p(s(bool,X27)) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))))) )),
  inference(resolution,[],[f777,f947])).
fof(f777,plain,(
  ( ! [X109,X110] : (~sP6(X109,X110,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | p(s(bool,X110)) | p(s(bool,X109))) )),
  inference(resolution,[],[f352,f436])).
fof(f41621,plain,(
  ( ! [X17] : (p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) | p(s(bool,X17))) )),
  inference(resolution,[],[f777,f946])).
fof(f41078,plain,(
  ( ! [X8] : (p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) | ~p(s(bool,X8))) )),
  inference(resolution,[],[f745,f911])).
fof(f911,plain,(
  ( ! [X47,X48] : (sP2(X47,X48,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) )),
  inference(resolution,[],[f350,f395])).
fof(f745,plain,(
  ( ! [X45,X46] : (~sP2(X45,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X46) | p(s(bool,X46)) | ~p(s(bool,X45))) )),
  inference(resolution,[],[f352,f394])).
fof(f40895,plain,(
  ( ! [X8] : (p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) | ~p(s(bool,X8))) )),
  inference(resolution,[],[f744,f912])).
fof(f912,plain,(
  ( ! [X50,X49] : (sP2(X49,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X50)) )),
  inference(resolution,[],[f350,f396])).
fof(f744,plain,(
  ( ! [X43,X44] : (~sP2(X43,X44,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | p(s(bool,X44)) | ~p(s(bool,X43))) )),
  inference(resolution,[],[f352,f394])).
fof(f40161,plain,(
  ( ! [X16] : (p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) | ~p(s(bool,X16))) )),
  inference(resolution,[],[f652,f944])).
fof(f944,plain,(
  ( ! [X114,X113] : (sP7(X113,X114,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) )),
  inference(resolution,[],[f350,f433])).
fof(f652,plain,(
  ( ! [X111,X112] : (~sP7(X111,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X112) | p(s(bool,X112)) | ~p(s(bool,X111))) )),
  inference(resolution,[],[f351,f432])).
fof(f40108,plain,(
  ( ! [X16] : (p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) | ~p(s(bool,X16))) )),
  inference(resolution,[],[f651,f944])).
fof(f651,plain,(
  ( ! [X109,X110] : (~sP7('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X109,X110) | p(s(bool,X110)) | ~p(s(bool,X109))) )),
  inference(resolution,[],[f351,f432])).
fof(f40062,plain,(
  ( ! [X8] : (p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) | ~p(s(bool,X8))) )),
  inference(resolution,[],[f649,f941])).
fof(f941,plain,(
  ( ! [X107,X108] : (sP8(X107,X108,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) )),
  inference(resolution,[],[f350,f429])).
fof(f649,plain,(
  ( ! [X105,X106] : (~sP8(X105,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X106) | p(s(bool,X106)) | ~p(s(bool,X105))) )),
  inference(resolution,[],[f351,f428])).
fof(f40018,plain,(
  ( ! [X8] : (p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) | ~p(s(bool,X8))) )),
  inference(resolution,[],[f648,f941])).
fof(f648,plain,(
  ( ! [X103,X104] : (~sP8('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X103,X104) | p(s(bool,X104)) | ~p(s(bool,X103))) )),
  inference(resolution,[],[f351,f428])).
fof(f39556,plain,(
  ( ! [X325,X324,X326] : (p(s(bool,X324)) | ~p(s(bool,X325)) | sP9(X326,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X325) | p(s(bool,X326))) ) | ($spl10 | $spl43)),
  inference(subsumption_resolution,[],[f39302,f553])).
fof(f39302,plain,(
  ( ! [X325,X324,X326] : (s(bool,'F_0') != s(bool,X324) | p(s(bool,X324)) | ~p(s(bool,X325)) | sP9(X326,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X325) | p(s(bool,X326))) ) | ($spl10 | $spl43)),
  inference(superposition,[],[f7606,f29922])).
fof(f29922,plain,(
  s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | $spl43),
  inference(resolution,[],[f28542,f553])).
fof(f28542,plain,(
  ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | $spl43),
  inference(resolution,[],[f4750,f356])).
fof(f4750,plain,(
  ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | $spl43),
  inference(cnf_transformation,[],[f4750_D])).
fof(f4750_D,plain,(
  ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) <=> ~$spl43),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl43])])).
fof(f39029,plain,(
  ( ! [X8] : (p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) | ~p(s(bool,X8))) )),
  inference(resolution,[],[f636,f928])).
fof(f928,plain,(
  ( ! [X81,X82] : (sP4(X81,X82,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) )),
  inference(resolution,[],[f350,f412])).
fof(f636,plain,(
  ( ! [X80,X79] : (~sP4(X79,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X80) | p(s(bool,X80)) | ~p(s(bool,X79))) )),
  inference(resolution,[],[f351,f411])).
fof(f38849,plain,(
  ( ! [X8] : (p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) | ~p(s(bool,X8))) )),
  inference(resolution,[],[f635,f928])).
fof(f635,plain,(
  ( ! [X78,X77] : (~sP4('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X77,X78) | p(s(bool,X78)) | ~p(s(bool,X77))) )),
  inference(resolution,[],[f351,f411])).
fof(f38795,plain,(
  ( ! [X15] : (p(s(bool,X15)) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))))) )),
  inference(resolution,[],[f619,f912])).
fof(f619,plain,(
  ( ! [X45,X46] : (~sP2('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X45,X46) | p(s(bool,X46)) | p(s(bool,X45))) )),
  inference(resolution,[],[f351,f394])).
fof(f38788,plain,(
  ( ! [X8] : (p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) | p(s(bool,X8))) )),
  inference(resolution,[],[f619,f911])).
fof(f38666,plain,(
  ( ! [X8] : (p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) | ~p(s(bool,X8))) )),
  inference(resolution,[],[f606,f898])).
fof(f898,plain,(
  ( ! [X21,X22] : (sP0(X21,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X22)) )),
  inference(resolution,[],[f350,f379])).
fof(f606,plain,(
  ( ! [X19,X20] : (~sP0(X19,X20,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) | p(s(bool,X20)) | ~p(s(bool,X19))) )),
  inference(resolution,[],[f351,f377])).
fof(f38630,plain,(
  ( ! [X14] : (~p(s(bool,X14)) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))))) )),
  inference(resolution,[],[f605,f898])).
fof(f605,plain,(
  ( ! [X17,X18] : (~sP0('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X17,X18) | ~p(s(bool,X18)) | p(s(bool,X17))) )),
  inference(resolution,[],[f351,f377])).
fof(f13749,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))) | $spl75),
  inference(cnf_transformation,[],[f13749_D])).
fof(f13749_D,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))) <=> ~$spl75),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl75])])).
fof(f37667,plain,(
  ( ! [X12] : (p(s(bool,X12)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) )),
  inference(resolution,[],[f947,f3286])).
fof(f37663,plain,(
  ( ! [X4,X5] : (p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) | p(s(bool,X4)) | p(s(bool,X5))) )),
  inference(resolution,[],[f947,f436])).
fof(f37677,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) ) | $spl448),
  inference(cnf_transformation,[],[f37677_D])).
fof(f37677_D,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) ) <=> ~$spl448),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl448])])).
fof(f37662,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) ) | $spl10),
  inference(resolution,[],[f947,f19474])).
fof(f37676,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,X1) != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) ) | $spl446),
  inference(cnf_transformation,[],[f37676_D])).
fof(f37676_D,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,X1) != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) ) <=> ~$spl446),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl446])])).
fof(f37661,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'F_0') != s(bool,X1) | s(bool,X1) != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) ) | $spl10),
  inference(resolution,[],[f947,f19603])).
fof(f20275,plain,(
  ( ! [X134,X133] : (~p(s(bool,X133)) | p(s(bool,X134)) | sP9(X134,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X133)) ) | $spl104),
  inference(cnf_transformation,[],[f20275_D])).
fof(f20275_D,plain,(
  ( ! [X134,X133] : (~p(s(bool,X133)) | p(s(bool,X134)) | sP9(X134,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X133)) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f37489,plain,(
  ( ! [X315,X316,X317] : (p(s(bool,X315)) | ~p(s(bool,X316)) | sP9(X317,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X316) | p(s(bool,X317))) ) | ($spl10 | $spl31)),
  inference(subsumption_resolution,[],[f37240,f553])).
fof(f37240,plain,(
  ( ! [X315,X316,X317] : (s(bool,'F_0') != s(bool,X315) | p(s(bool,X315)) | ~p(s(bool,X316)) | sP9(X317,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X316) | p(s(bool,X317))) ) | ($spl10 | $spl31)),
  inference(superposition,[],[f7606,f36697])).
fof(f36697,plain,(
  s(bool,'F_0') = s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | $spl31),
  inference(resolution,[],[f36602,f553])).
fof(f36602,plain,(
  ~p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | $spl31),
  inference(subsumption_resolution,[],[f36476,f3879])).
fof(f3879,plain,(
  s(sK12,sK14) != s(sK12,sK16) | $spl31),
  inference(cnf_transformation,[],[f3879_D])).
fof(f3879_D,plain,(
  s(sK12,sK14) != s(sK12,sK16) <=> ~$spl31),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl31])])).
fof(f37009,plain,(
  ( ! [X8,X9] : (sP9(X8,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X9)) ) | $spl444),
  inference(cnf_transformation,[],[f37009_D])).
fof(f37009_D,plain,(
  ( ! [X8,X9] : (sP9(X8,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X9)) ) <=> ~$spl444),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl444])])).
fof(f37008,plain,(
  ( ! [X8,X9] : (sP9(X8,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X9) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))))) )),
  inference(subsumption_resolution,[],[f37007,f941])).
fof(f37007,plain,(
  ( ! [X8,X9] : (sP9(X8,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X9) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) | ~sP8(X8,X9,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) )),
  inference(subsumption_resolution,[],[f36996,f435])).
fof(f36996,plain,(
  ( ! [X8,X9] : (sP9(X8,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X9) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) | ~sP8(X8,X9,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))) | ~sP7(happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X9,X8)) )),
  inference(resolution,[],[f946,f426])).
fof(f36994,plain,(
  ( ! [X4,X5] : (p(s(bool,X4)) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) | p(s(bool,X5))) )),
  inference(resolution,[],[f946,f436])).
fof(f37006,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) ) | $spl442),
  inference(cnf_transformation,[],[f37006_D])).
fof(f37006_D,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) ) <=> ~$spl442),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl442])])).
fof(f36993,plain,(
  ( ! [X2,X3] : (p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) | p(s(bool,X2)) | s(bool,X3) != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) ) | $spl10),
  inference(resolution,[],[f946,f19474])).
fof(f36992,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))) | s(bool,X1) != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) ) | $spl10),
  inference(resolution,[],[f946,f19603])).
fof(f36948,plain,(
  ( ! [X12] : (p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) | s(bool,'F_0') != s(bool,X12)) )),
  inference(resolution,[],[f945,f3286])).
fof(f36960,plain,(
  ( ! [X10,X11] : (sP9(happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X10,X11)) ) | $spl440),
  inference(cnf_transformation,[],[f36960_D])).
fof(f36960_D,plain,(
  ( ! [X10,X11] : (sP9(happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X10,X11)) ) <=> ~$spl440),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl440])])).
fof(f36959,plain,(
  ( ! [X10,X11] : (sP9(happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X10,X11) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))))) )),
  inference(subsumption_resolution,[],[f36958,f944])).
fof(f36958,plain,(
  ( ! [X10,X11] : (sP9(happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X10,X11) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) | ~sP7(X10,X11,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) )),
  inference(subsumption_resolution,[],[f36947,f431])).
fof(f36947,plain,(
  ( ! [X10,X11] : (sP9(happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X10,X11) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) | ~sP8(happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X11,X10) | ~sP7(X10,X11,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) )),
  inference(resolution,[],[f945,f427])).
fof(f36944,plain,(
  ( ! [X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))))) )),
  inference(resolution,[],[f945,f436])).
fof(f36943,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) | s(bool,X2) != s(bool,X3)) ) | $spl10),
  inference(resolution,[],[f945,f19474])).
fof(f36942,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) | s(bool,'F_0') != s(bool,X0) | s(bool,X0) != s(bool,X1)) ) | $spl10),
  inference(resolution,[],[f945,f19603])).
fof(f36848,plain,(
  ( ! [X10,X11] : (sP9(happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X10,X11) | ~sP6(X11,X10,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) ) | $spl438),
  inference(cnf_transformation,[],[f36848_D])).
fof(f36848_D,plain,(
  ( ! [X10,X11] : (sP9(happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X10,X11) | ~sP6(X11,X10,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) ) <=> ~$spl438),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl438])])).
fof(f36847,plain,(
  ( ! [X10,X11] : (sP9(happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X10,X11) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) | ~sP6(X11,X10,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) )),
  inference(subsumption_resolution,[],[f36837,f431])).
fof(f36837,plain,(
  ( ! [X10,X11] : (sP9(happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X10,X11) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) | ~sP8(happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X11,X10) | ~sP6(X11,X10,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) )),
  inference(resolution,[],[f944,f427])).
fof(f36834,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | ~p(s(bool,X5)) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))))) )),
  inference(resolution,[],[f944,f432])).
fof(f36833,plain,(
  ( ! [X2,X3] : (~p(s(bool,X2)) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) | s(bool,X2) != s(bool,X3)) ) | $spl10),
  inference(resolution,[],[f944,f19313])).
fof(f36832,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) | s(bool,'T_0') != s(bool,X0) | s(bool,X0) != s(bool,X1)) ) | $spl10),
  inference(resolution,[],[f944,f19575])).
fof(f36831,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))) | $spl436),
  inference(cnf_transformation,[],[f36831_D])).
fof(f36831_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))) <=> ~$spl436),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl436])])).
fof(f36816,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))) | ~p(s(bool,X21))) ) | ($spl10 | $spl96)),
  inference(resolution,[],[f941,f27175])).
fof(f36815,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))) | s(bool,'T_0') != s(bool,X19)) ) | $spl10),
  inference(resolution,[],[f941,f19799])).
fof(f36814,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))) | ~p(s(bool,X16)) | ~p(s(bool,X17))) ) | $spl10),
  inference(resolution,[],[f941,f19796])).
fof(f36827,plain,(
  ( ! [X2,X3] : (sP9(X2,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X3) | ~sP6(X3,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X2)) ) | $spl430),
  inference(cnf_transformation,[],[f36827_D])).
fof(f36827_D,plain,(
  ( ! [X2,X3] : (sP9(X2,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X3) | ~sP6(X3,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X2)) ) <=> ~$spl430),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl430])])).
fof(f36813,plain,(
  ( ! [X14,X15] : (p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) | sP9(X14,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X15) | ~sP6(X15,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X14)) ) | $spl10),
  inference(resolution,[],[f941,f19385])).
fof(f36829,plain,(
  ( ! [X10,X11] : (sP9(X10,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X11) | ~p(s(bool,X11))) ) | $spl434),
  inference(cnf_transformation,[],[f36829_D])).
fof(f36829_D,plain,(
  ( ! [X10,X11] : (sP9(X10,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X11) | ~p(s(bool,X11))) ) <=> ~$spl434),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl434])])).
fof(f36811,plain,(
  ( ! [X10,X11] : (sP9(X10,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X11) | ~p(s(bool,X11)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) ) | $spl10),
  inference(resolution,[],[f941,f19168])).
fof(f36828,plain,(
  ( ! [X8,X9] : (sP9(X8,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X9) | s(bool,'T_0') != s(bool,X9)) ) | $spl432),
  inference(cnf_transformation,[],[f36828_D])).
fof(f36828_D,plain,(
  ( ! [X8,X9] : (sP9(X8,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X9) | s(bool,'T_0') != s(bool,X9)) ) <=> ~$spl432),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl432])])).
fof(f36810,plain,(
  ( ! [X8,X9] : (sP9(X8,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X9) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) | s(bool,'T_0') != s(bool,X9)) ) | $spl10),
  inference(resolution,[],[f941,f19159])).
fof(f36809,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))))) )),
  inference(resolution,[],[f941,f428])).
fof(f36826,plain,(
  ( ! [X2,X3] : (sP9(X2,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X3) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) | ~sP6(X3,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X2)) )),
  inference(subsumption_resolution,[],[f36807,f435])).
fof(f36807,plain,(
  ( ! [X2,X3] : (sP9(X2,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X3) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) | ~sP7(happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X3,X2) | ~sP6(X3,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X2)) )),
  inference(resolution,[],[f941,f426])).
fof(f36804,plain,(
  ( ! [X4,X5] : (sP5(happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X4,X5)) ) | $spl428),
  inference(cnf_transformation,[],[f36804_D])).
fof(f36804_D,plain,(
  ( ! [X4,X5] : (sP5(happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X4,X5)) ) <=> ~$spl428),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl428])])).
fof(f36794,plain,(
  ( ! [X8,X9] : (sP5(happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X8,X9) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))))) )),
  inference(resolution,[],[f928,f587])).
fof(f36805,plain,(
  ( ! [X6,X7] : (sP5(happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X6,X7) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))))) )),
  inference(subsumption_resolution,[],[f36793,f415])).
fof(f36793,plain,(
  ( ! [X6,X7] : (sP5(happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X6,X7) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) | ~p(s(bool,X7))) )),
  inference(resolution,[],[f928,f409])).
fof(f36803,plain,(
  ( ! [X4,X5] : (sP5(happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X4,X5) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))))) )),
  inference(subsumption_resolution,[],[f36792,f416])).
fof(f36792,plain,(
  ( ! [X4,X5] : (sP5(happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X4,X5) | ~p(s(bool,X4)) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))))) )),
  inference(resolution,[],[f928,f408])).
fof(f36790,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))))) )),
  inference(resolution,[],[f928,f411])).
fof(f36789,plain,(
  ( ! [X8,X9] : (sP3(X8,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X9) | p(s(bool,X9))) ) | $spl426),
  inference(cnf_transformation,[],[f36789_D])).
fof(f36789_D,plain,(
  ( ! [X8,X9] : (sP3(X8,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X9) | p(s(bool,X9))) ) <=> ~$spl426),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl426])])).
fof(f36778,plain,(
  ( ! [X8,X9] : (sP3(X8,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X9) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) | p(s(bool,X9))) )),
  inference(resolution,[],[f912,f393])).
fof(f36774,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) | p(s(bool,X1))) )),
  inference(resolution,[],[f912,f394])).
fof(f36773,plain,(
  ( ! [X8,X9] : (sP3(X8,X9,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))) | p(s(bool,X9))) ) | $spl424),
  inference(cnf_transformation,[],[f36773_D])).
fof(f36773_D,plain,(
  ( ! [X8,X9] : (sP3(X8,X9,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))) | p(s(bool,X9))) ) <=> ~$spl424),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl424])])).
fof(f36762,plain,(
  ( ! [X8,X9] : (sP3(X8,X9,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))) | p(s(bool,X9)) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))))) )),
  inference(resolution,[],[f911,f393])).
fof(f36758,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))))) )),
  inference(resolution,[],[f911,f394])).
fof(f36757,plain,(
  ( ! [X8,X9] : (sP1(X8,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X9) | ~p(s(bool,X9))) ) | $spl422),
  inference(cnf_transformation,[],[f36757_D])).
fof(f36757_D,plain,(
  ( ! [X8,X9] : (sP1(X8,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X9) | ~p(s(bool,X9))) ) <=> ~$spl422),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl422])])).
fof(f36746,plain,(
  ( ! [X8,X9] : (sP1(X8,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X9) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) | ~p(s(bool,X9))) )),
  inference(resolution,[],[f898,f376])).
fof(f36742,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) | ~p(s(bool,X1))) )),
  inference(resolution,[],[f898,f377])).
fof(f36465,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) | p(s(bool,X1))) )),
  inference(subsumption_resolution,[],[f36451,f538])).
fof(f36451,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) | p(s(bool,X1))) )),
  inference(resolution,[],[f959,f474])).
fof(f959,plain,(
  ( ! [X134] : (sP11(happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)),X134)) )),
  inference(resolution,[],[f350,f476])).
fof(f36464,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f36450,f538])).
fof(f36450,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) | p(s(bool,X0))) )),
  inference(resolution,[],[f959,f478])).
fof(f36419,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))))) )),
  inference(subsumption_resolution,[],[f36404,f538])).
fof(f36404,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))))) )),
  inference(resolution,[],[f958,f474])).
fof(f958,plain,(
  ( ! [X133] : (sP11(X133,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) )),
  inference(resolution,[],[f350,f475])).
fof(f36416,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))))) )),
  inference(subsumption_resolution,[],[f36403,f538])).
fof(f36403,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15))))) )),
  inference(resolution,[],[f958,f478])).
fof(f36232,plain,(
  ( ! [X325,X326] : (~p(s(bool,X325)) | p(s(bool,X326)) | sP9(X326,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X325)) ) | $spl418),
  inference(cnf_transformation,[],[f36232_D])).
fof(f36232_D,plain,(
  ( ! [X325,X326] : (~p(s(bool,X325)) | p(s(bool,X326)) | sP9(X326,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X325)) ) <=> ~$spl418),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl418])])).
fof(f36231,plain,(
  ( ! [X325,X324,X326] : (p(s(bool,X324)) | ~p(s(bool,X325)) | sP9(X326,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X325) | p(s(bool,X326))) ) | ($spl10 | $spl43)),
  inference(subsumption_resolution,[],[f35983,f553])).
fof(f35983,plain,(
  ( ! [X325,X324,X326] : (s(bool,'F_0') != s(bool,X324) | p(s(bool,X324)) | ~p(s(bool,X325)) | sP9(X326,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X325) | p(s(bool,X326))) ) | ($spl10 | $spl43)),
  inference(superposition,[],[f7606,f28626])).
fof(f28626,plain,(
  s(bool,'F_0') = s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | $spl43),
  inference(resolution,[],[f4750,f553])).
fof(f20272,plain,(
  ( ! [X116,X115] : (~p(s(bool,X115)) | p(s(bool,X116)) | sP9(X116,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X115)) ) | $spl102),
  inference(cnf_transformation,[],[f20272_D])).
fof(f20272_D,plain,(
  ( ! [X116,X115] : (~p(s(bool,X115)) | p(s(bool,X116)) | sP9(X116,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X115)) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f35275,plain,(
  ( ! [X325,X324,X326] : (p(s(bool,X324)) | ~p(s(bool,X325)) | sP9(X326,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X325) | p(s(bool,X326))) ) | ($spl7 | $spl10)),
  inference(subsumption_resolution,[],[f35026,f553])).
fof(f35026,plain,(
  ( ! [X325,X324,X326] : (s(bool,'F_0') != s(bool,X324) | p(s(bool,X324)) | ~p(s(bool,X325)) | sP9(X326,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X325) | p(s(bool,X326))) ) | ($spl7 | $spl10)),
  inference(superposition,[],[f7606,f22776])).
fof(f33684,plain,(
  ( ! [X325,X326] : (~p(s(bool,X325)) | p(s(bool,X326)) | sP9(X326,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X325)) ) | $spl416),
  inference(cnf_transformation,[],[f33684_D])).
fof(f33684_D,plain,(
  ( ! [X325,X326] : (~p(s(bool,X325)) | p(s(bool,X326)) | sP9(X326,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X325)) ) <=> ~$spl416),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl416])])).
fof(f33683,plain,(
  ( ! [X325,X324,X326] : (p(s(bool,X324)) | ~p(s(bool,X325)) | sP9(X326,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X325) | p(s(bool,X326))) ) | ($spl3 | $spl10)),
  inference(subsumption_resolution,[],[f33429,f553])).
fof(f33429,plain,(
  ( ! [X325,X324,X326] : (s(bool,'F_0') != s(bool,X324) | p(s(bool,X324)) | ~p(s(bool,X325)) | sP9(X326,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X325) | p(s(bool,X326))) ) | ($spl3 | $spl10)),
  inference(superposition,[],[f7606,f28024])).
fof(f28024,plain,(
  s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) | $spl3),
  inference(resolution,[],[f713,f553])).
fof(f30928,plain,(
  ( ! [X2,X3] : (~p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | ($spl10 | $spl43)),
  inference(resolution,[],[f29892,f19313])).
fof(f29892,plain,(
  ( ! [X107,X108] : (sP7('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X107,X108)) ) | $spl43),
  inference(resolution,[],[f28542,f435])).
fof(f30927,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X1) | s(bool,X1) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | ($spl10 | $spl43)),
  inference(resolution,[],[f29892,f19575])).
fof(f30912,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | s(bool,X1) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | ($spl10 | $spl43)),
  inference(resolution,[],[f29891,f19575])).
fof(f29891,plain,(
  ( ! [X105,X106] : (sP7(X105,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X106)) ) | $spl43),
  inference(resolution,[],[f28542,f434])).
fof(f30869,plain,(
  ( ! [X30,X29] : (~p(s(bool,X29)) | s(bool,X30) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | s(bool,'T_0') != s(bool,X30)) ) | ($spl10 | $spl43)),
  inference(resolution,[],[f29888,f19111])).
fof(f29888,plain,(
  ( ! [X99,X100] : (sP8(X99,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X100)) ) | $spl43),
  inference(resolution,[],[f28542,f430])).
fof(f30868,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') != s(bool,X27) | sP9(X28,'F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | s(bool,X27) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | ($spl10 | $spl43)),
  inference(resolution,[],[f29888,f19068])).
fof(f30867,plain,(
  ( ! [X26] : (s(bool,'T_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | ~p(s(bool,X26))) ) | $spl43),
  inference(resolution,[],[f29888,f3327])).
fof(f30866,plain,(
  ( ! [X25] : (s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | ~p(s(bool,X25))) ) | $spl43),
  inference(resolution,[],[f29888,f3285])).
fof(f30863,plain,(
  ( ! [X22] : (sP9(X22,'F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | s(bool,'T_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | $spl43),
  inference(resolution,[],[f29888,f3227])).
fof(f30861,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,X19) | s(bool,'T_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | ($spl10 | $spl43)),
  inference(resolution,[],[f29888,f19799])).
fof(f30856,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | p(s(bool,X9)) | s(bool,'T_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | ($spl10 | $spl43)),
  inference(resolution,[],[f29888,f19159])).
fof(f30736,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | $spl400),
  inference(cnf_transformation,[],[f30736_D])).
fof(f30736_D,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) <=> ~$spl400),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl400])])).
fof(f30725,plain,(
  ( ! [X2,X3] : (~p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | ($spl10 | $spl43)),
  inference(resolution,[],[f28596,f19313])).
fof(f28596,plain,(
  ( ! [X109,X108] : (sP7('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X108,X109)) ) | $spl43),
  inference(resolution,[],[f4750,f435])).
fof(f30682,plain,(
  ( ! [X27] : (s(bool,'T_0') != s(bool,X27) | s(bool,X27) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | $spl398),
  inference(cnf_transformation,[],[f30682_D])).
fof(f30682_D,plain,(
  ( ! [X27] : (s(bool,'T_0') != s(bool,X27) | s(bool,X27) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) <=> ~$spl398),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl398])])).
fof(f30724,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X1) | s(bool,X1) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | ($spl10 | $spl43)),
  inference(resolution,[],[f28596,f19575])).
fof(f27234,plain,(
  s(bool,'T_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | $spl355),
  inference(cnf_transformation,[],[f27234_D])).
fof(f27234_D,plain,(
  s(bool,'T_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) <=> ~$spl355),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl355])])).
fof(f30710,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | s(bool,X1) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | ($spl10 | $spl43)),
  inference(resolution,[],[f28595,f19575])).
fof(f28595,plain,(
  ( ! [X107,X106] : (sP7(X106,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X107)) ) | $spl43),
  inference(resolution,[],[f4750,f434])).
fof(f30671,plain,(
  ( ! [X30,X29] : (~p(s(bool,X29)) | s(bool,X30) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | s(bool,'T_0') != s(bool,X30)) ) | ($spl10 | $spl43)),
  inference(resolution,[],[f28592,f19111])).
fof(f28592,plain,(
  ( ! [X101,X100] : (sP8(X100,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X101)) ) | $spl43),
  inference(resolution,[],[f4750,f430])).
fof(f30681,plain,(
  ( ! [X22] : (sP9(X22,'F_0','const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | $spl396),
  inference(cnf_transformation,[],[f30681_D])).
fof(f30681_D,plain,(
  ( ! [X22] : (sP9(X22,'F_0','const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) <=> ~$spl396),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl396])])).
fof(f30670,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') != s(bool,X27) | sP9(X28,'F_0','const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | s(bool,X27) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | ($spl10 | $spl43)),
  inference(resolution,[],[f28592,f19068])).
fof(f30669,plain,(
  ( ! [X26] : (s(bool,'T_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | ~p(s(bool,X26))) ) | $spl43),
  inference(resolution,[],[f28592,f3327])).
fof(f23054,plain,(
  s(bool,'F_0') = s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | $spl244),
  inference(cnf_transformation,[],[f23054_D])).
fof(f23054_D,plain,(
  s(bool,'F_0') = s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) <=> ~$spl244),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl244])])).
fof(f30668,plain,(
  ( ! [X25] : (s(bool,'F_0') = s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | ~p(s(bool,X25))) ) | $spl43),
  inference(resolution,[],[f28592,f3285])).
fof(f30665,plain,(
  ( ! [X22] : (sP9(X22,'F_0','const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | s(bool,'T_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | $spl43),
  inference(resolution,[],[f28592,f3227])).
fof(f30663,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,X19) | s(bool,'T_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | ($spl10 | $spl43)),
  inference(resolution,[],[f28592,f19799])).
fof(f30680,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | p(s(bool,X9))) ) | $spl394),
  inference(cnf_transformation,[],[f30680_D])).
fof(f30680_D,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | p(s(bool,X9))) ) <=> ~$spl394),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl394])])).
fof(f30658,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | p(s(bool,X9)) | s(bool,'T_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | ($spl10 | $spl43)),
  inference(resolution,[],[f28592,f19159])).
fof(f23561,plain,(
  ( ! [X325,X326] : (~p(s(bool,X325)) | p(s(bool,X326)) | sP9(X326,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X325)) ) | $spl248),
  inference(cnf_transformation,[],[f23561_D])).
fof(f23561_D,plain,(
  ( ! [X325,X326] : (~p(s(bool,X325)) | p(s(bool,X326)) | sP9(X326,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X325)) ) <=> ~$spl248),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl248])])).
fof(f30452,plain,(
  ( ! [X325,X324,X326] : (p(s(bool,X324)) | ~p(s(bool,X325)) | sP9(X326,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X325) | p(s(bool,X326))) ) | $spl10),
  inference(subsumption_resolution,[],[f30199,f553])).
fof(f30199,plain,(
  ( ! [X325,X324,X326] : (s(bool,'F_0') != s(bool,X324) | p(s(bool,X324)) | ~p(s(bool,X325)) | sP9(X326,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X325) | p(s(bool,X326))) ) | $spl10),
  inference(superposition,[],[f7606,f806])).
fof(f29705,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) | $spl392),
  inference(cnf_transformation,[],[f29705_D])).
fof(f29705_D,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) <=> ~$spl392),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl392])])).
fof(f29694,plain,(
  ( ! [X2,X3] : (~p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) | ($spl3 | $spl10)),
  inference(resolution,[],[f27994,f19313])).
fof(f29634,plain,(
  ( ! [X27] : (s(bool,'T_0') != s(bool,X27) | s(bool,X27) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) | $spl390),
  inference(cnf_transformation,[],[f29634_D])).
fof(f29634_D,plain,(
  ( ! [X27] : (s(bool,'T_0') != s(bool,X27) | s(bool,X27) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) <=> ~$spl390),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl390])])).
fof(f29693,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X1) | s(bool,X1) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) | ($spl3 | $spl10)),
  inference(resolution,[],[f27994,f19575])).
fof(f21560,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) | $spl180),
  inference(cnf_transformation,[],[f21560_D])).
fof(f21560_D,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f27238,plain,(
  s(bool,'T_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) | $spl359),
  inference(cnf_transformation,[],[f27238_D])).
fof(f27238_D,plain,(
  s(bool,'T_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) <=> ~$spl359),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl359])])).
fof(f29679,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) | s(bool,X1) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) | ($spl3 | $spl10)),
  inference(resolution,[],[f27993,f19575])).
fof(f21243,plain,(
  ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | $spl133),
  inference(cnf_transformation,[],[f21243_D])).
fof(f21243_D,plain,(
  ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) <=> ~$spl133),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl133])])).
fof(f29675,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))) )),
  inference(duplicate_literal_removal,[],[f29666])).
fof(f29666,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | ~p(s(bool,X4))) )),
  inference(resolution,[],[f785,f459])).
fof(f29676,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))) )),
  inference(duplicate_literal_removal,[],[f29663])).
fof(f29663,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | ~p(s(bool,X1)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))) )),
  inference(resolution,[],[f785,f466])).
fof(f29623,plain,(
  ( ! [X30,X29] : (~p(s(bool,X29)) | s(bool,X30) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) | s(bool,'T_0') != s(bool,X30)) ) | ($spl3 | $spl10)),
  inference(resolution,[],[f27990,f19111])).
fof(f29633,plain,(
  ( ! [X22] : (sP9(X22,'F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) | $spl388),
  inference(cnf_transformation,[],[f29633_D])).
fof(f29633_D,plain,(
  ( ! [X22] : (sP9(X22,'F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) <=> ~$spl388),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl388])])).
fof(f29622,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') != s(bool,X27) | sP9(X28,'F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) | s(bool,X27) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) | ($spl3 | $spl10)),
  inference(resolution,[],[f27990,f19068])).
fof(f29621,plain,(
  ( ! [X26] : (s(bool,'T_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) | ~p(s(bool,X26))) ) | $spl3),
  inference(resolution,[],[f27990,f3327])).
fof(f3618,plain,(
  s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) | $spl26),
  inference(cnf_transformation,[],[f3618_D])).
fof(f3618_D,plain,(
  s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f29620,plain,(
  ( ! [X25] : (s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) | ~p(s(bool,X25))) ) | $spl3),
  inference(resolution,[],[f27990,f3285])).
fof(f29617,plain,(
  ( ! [X22] : (sP9(X22,'F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) | s(bool,'T_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) | $spl3),
  inference(resolution,[],[f27990,f3227])).
fof(f29615,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,X19) | s(bool,'T_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) | ($spl3 | $spl10)),
  inference(resolution,[],[f27990,f19799])).
fof(f29632,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) | p(s(bool,X9))) ) | $spl386),
  inference(cnf_transformation,[],[f29632_D])).
fof(f29632_D,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) | p(s(bool,X9))) ) <=> ~$spl386),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl386])])).
fof(f29610,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) | p(s(bool,X9)) | s(bool,'T_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) | ($spl3 | $spl10)),
  inference(resolution,[],[f27990,f19159])).
fof(f29565,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))) )),
  inference(duplicate_literal_removal,[],[f29556])).
fof(f29556,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | ~p(s(bool,X4)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))) )),
  inference(resolution,[],[f784,f459])).
fof(f29566,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))) )),
  inference(duplicate_literal_removal,[],[f29553])).
fof(f29553,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | ~p(s(bool,X1))) )),
  inference(resolution,[],[f784,f466])).
fof(f29525,plain,(
  ( ! [X2] : (p(s(bool,X2)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))) )),
  inference(duplicate_literal_removal,[],[f29516])).
fof(f29516,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,X2)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))) )),
  inference(resolution,[],[f783,f460])).
fof(f29526,plain,(
  ( ! [X1] : (p(s(bool,X1)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))) )),
  inference(duplicate_literal_removal,[],[f29515])).
fof(f29515,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,X1)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))) )),
  inference(resolution,[],[f783,f465])).
fof(f29511,plain,(
  ( ! [X3] : (p(s(bool,X3)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))) )),
  inference(duplicate_literal_removal,[],[f29503])).
fof(f29503,plain,(
  ( ! [X3] : (p(s(bool,X3)) | p(s(bool,X3)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))) )),
  inference(resolution,[],[f782,f463])).
fof(f29512,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))) )),
  inference(duplicate_literal_removal,[],[f29500])).
fof(f29500,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | p(s(bool,X0))) )),
  inference(resolution,[],[f782,f464])).
fof(f28877,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) | $spl375),
  inference(cnf_transformation,[],[f28877_D])).
fof(f28877_D,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) <=> ~$spl375),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl375])])).
fof(f29489,plain,(
  ( ! [X12] : (p(s(bool,X12)) | s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) | $spl0),
  inference(resolution,[],[f27308,f3286])).
fof(f27308,plain,(
  ( ! [X112,X113] : (sP6('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X112,X113)) ) | $spl0),
  inference(resolution,[],[f712,f439])).
fof(f29497,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) | $spl384),
  inference(cnf_transformation,[],[f29497_D])).
fof(f29497_D,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) <=> ~$spl384),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl384])])).
fof(f29484,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) | ($spl0 | $spl10)),
  inference(resolution,[],[f27308,f19474])).
fof(f29496,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) | $spl382),
  inference(cnf_transformation,[],[f29496_D])).
fof(f29496_D,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) <=> ~$spl382),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl382])])).
fof(f29483,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'F_0') != s(bool,X1) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) | ($spl0 | $spl10)),
  inference(resolution,[],[f27308,f19603])).
fof(f29480,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) | $spl380),
  inference(cnf_transformation,[],[f29480_D])).
fof(f29480_D,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) <=> ~$spl380),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl380])])).
fof(f29468,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) | ($spl0 | $spl10)),
  inference(resolution,[],[f27307,f19603])).
fof(f27307,plain,(
  ( ! [X111,X110] : (sP6(X110,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X111)) ) | $spl0),
  inference(resolution,[],[f712,f438])).
fof(f28880,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) | $spl378),
  inference(cnf_transformation,[],[f28880_D])).
fof(f28880_D,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) <=> ~$spl378),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl378])])).
fof(f28867,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | s(bool,'T_0') = s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) | ~p(s(bool,X21))) ) | ($spl0 | $spl10 | $spl96)),
  inference(resolution,[],[f27302,f27175])).
fof(f27302,plain,(
  ( ! [X101,X100] : (sP8(X100,X101,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) | $spl0),
  inference(resolution,[],[f712,f429])).
fof(f28866,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) | s(bool,'T_0') != s(bool,X19)) ) | ($spl0 | $spl10)),
  inference(resolution,[],[f27302,f19799])).
fof(f28865,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) | ~p(s(bool,X16)) | ~p(s(bool,X17))) ) | ($spl0 | $spl10)),
  inference(resolution,[],[f27302,f19796])).
fof(f28878,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X11) | ~p(s(bool,X11))) ) | $spl376),
  inference(cnf_transformation,[],[f28878_D])).
fof(f28878_D,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X11) | ~p(s(bool,X11))) ) <=> ~$spl376),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl376])])).
fof(f28862,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X11) | ~p(s(bool,X11)) | s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) | ($spl0 | $spl10)),
  inference(resolution,[],[f27302,f19168])).
fof(f4765,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | $spl44),
  inference(cnf_transformation,[],[f4765_D])).
fof(f4765_D,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f28806,plain,(
  ( ! [X3] : (p(s(bool,X3)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(duplicate_literal_removal,[],[f28797])).
fof(f28797,plain,(
  ( ! [X3] : (p(s(bool,X3)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | p(s(bool,X3))) )),
  inference(resolution,[],[f662,f458])).
fof(f28807,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(duplicate_literal_removal,[],[f28796])).
fof(f28796,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,X2)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(resolution,[],[f662,f467])).
fof(f22495,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | $spl238),
  inference(cnf_transformation,[],[f22495_D])).
fof(f22495_D,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) <=> ~$spl238),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl238])])).
fof(f28527,plain,(
  ( ! [X2,X3] : (~p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | ($spl7 | $spl10)),
  inference(resolution,[],[f4070,f19313])).
fof(f4070,plain,(
  ( ! [X107,X106] : (sP7('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X106,X107)) ) | $spl7),
  inference(resolution,[],[f840,f435])).
fof(f22439,plain,(
  ( ! [X25] : (s(bool,'T_0') != s(bool,X25) | s(bool,X25) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | $spl230),
  inference(cnf_transformation,[],[f22439_D])).
fof(f22439_D,plain,(
  ( ! [X25] : (s(bool,'T_0') != s(bool,X25) | s(bool,X25) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) <=> ~$spl230),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl230])])).
fof(f28526,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X1) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | ($spl7 | $spl10)),
  inference(resolution,[],[f4070,f19575])).
fof(f22478,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | $spl234),
  inference(cnf_transformation,[],[f22478_D])).
fof(f22478_D,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) <=> ~$spl234),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl234])])).
fof(f4468,plain,(
  s(bool,'T_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | $spl39),
  inference(cnf_transformation,[],[f4468_D])).
fof(f4468_D,plain,(
  s(bool,'T_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) <=> ~$spl39),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl39])])).
fof(f28511,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | ($spl7 | $spl10)),
  inference(resolution,[],[f4069,f19575])).
fof(f4069,plain,(
  ( ! [X105,X104] : (sP7(X104,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X105)) ) | $spl7),
  inference(resolution,[],[f840,f434])).
fof(f22465,plain,(
  ( ! [X4,X5] : (sP9('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X4,X5) | ~sP6(X5,X4,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | ~sP7(X4,X5,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | $spl232),
  inference(cnf_transformation,[],[f22465_D])).
fof(f22465_D,plain,(
  ( ! [X4,X5] : (sP9('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X4,X5) | ~sP6(X5,X4,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | ~sP7(X4,X5,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) <=> ~$spl232),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl232])])).
fof(f28487,plain,(
  ( ! [X4,X5] : (sP9('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X4,X5) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | ~sP7(X4,X5,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | ~sP6(X5,X4,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | $spl7),
  inference(resolution,[],[f4067,f427])).
fof(f4067,plain,(
  ( ! [X101,X100] : (sP8('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X100,X101)) ) | $spl7),
  inference(resolution,[],[f840,f431])).
fof(f28483,plain,(
  ( ! [X3] : (p(s(bool,X3)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(duplicate_literal_removal,[],[f28474])).
fof(f28474,plain,(
  ( ! [X3] : (p(s(bool,X3)) | p(s(bool,X3)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(resolution,[],[f661,f458])).
fof(f28484,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(duplicate_literal_removal,[],[f28473])).
fof(f28473,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | p(s(bool,X2))) )),
  inference(resolution,[],[f661,f467])).
fof(f28462,plain,(
  ( ! [X30,X29] : (~p(s(bool,X29)) | s(bool,X30) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | s(bool,'T_0') != s(bool,X30)) ) | ($spl7 | $spl10)),
  inference(resolution,[],[f4066,f19111])).
fof(f4066,plain,(
  ( ! [X99,X98] : (sP8(X98,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X99)) ) | $spl7),
  inference(resolution,[],[f840,f430])).
fof(f4469,plain,(
  ( ! [X9] : (sP9(X9,'F_0','const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | $spl40),
  inference(cnf_transformation,[],[f4469_D])).
fof(f4469_D,plain,(
  ( ! [X9] : (sP9(X9,'F_0','const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f28461,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') != s(bool,X27) | sP9(X28,'F_0','const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | s(bool,X27) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | ($spl7 | $spl10)),
  inference(resolution,[],[f4066,f19068])).
fof(f28460,plain,(
  ( ! [X26] : (s(bool,'T_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | ~p(s(bool,X26))) ) | $spl7),
  inference(resolution,[],[f4066,f3327])).
fof(f4466,plain,(
  s(bool,'F_0') = s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | $spl36),
  inference(cnf_transformation,[],[f4466_D])).
fof(f4466_D,plain,(
  s(bool,'F_0') = s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f28459,plain,(
  ( ! [X25] : (s(bool,'F_0') = s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | ~p(s(bool,X25))) ) | $spl7),
  inference(resolution,[],[f4066,f3285])).
fof(f28456,plain,(
  ( ! [X22] : (sP9(X22,'F_0','const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | s(bool,'T_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | $spl7),
  inference(resolution,[],[f4066,f3227])).
fof(f28454,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,X19) | s(bool,'T_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | ($spl7 | $spl10)),
  inference(resolution,[],[f4066,f19799])).
fof(f22437,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | p(s(bool,X9))) ) | $spl226),
  inference(cnf_transformation,[],[f22437_D])).
fof(f22437_D,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | p(s(bool,X9))) ) <=> ~$spl226),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl226])])).
fof(f28449,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | p(s(bool,X9)) | s(bool,'T_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | ($spl7 | $spl10)),
  inference(resolution,[],[f4066,f19159])).
fof(f28391,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(duplicate_literal_removal,[],[f28383])).
fof(f28383,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | ~p(s(bool,X2))) )),
  inference(resolution,[],[f660,f460])).
fof(f28392,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(duplicate_literal_removal,[],[f28382])).
fof(f28382,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | ~p(s(bool,X1))) )),
  inference(resolution,[],[f660,f465])).
fof(f21780,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | $spl212),
  inference(cnf_transformation,[],[f21780_D])).
fof(f21780_D,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f28355,plain,(
  ( ! [X2,X3] : (~p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | ($spl5 | $spl10)),
  inference(resolution,[],[f1416,f19313])).
fof(f1416,plain,(
  ( ! [X109,X108] : (sP7('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X108,X109)) ) | $spl5),
  inference(resolution,[],[f839,f435])).
fof(f21702,plain,(
  ( ! [X25] : (s(bool,'T_0') != s(bool,X25) | s(bool,X25) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | $spl204),
  inference(cnf_transformation,[],[f21702_D])).
fof(f21702_D,plain,(
  ( ! [X25] : (s(bool,'T_0') != s(bool,X25) | s(bool,X25) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f28354,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X1) | s(bool,X1) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | ($spl5 | $spl10)),
  inference(resolution,[],[f1416,f19575])).
fof(f21741,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | $spl208),
  inference(cnf_transformation,[],[f21741_D])).
fof(f21741_D,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f3340,plain,(
  s(bool,'T_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | $spl17),
  inference(cnf_transformation,[],[f3340_D])).
fof(f3340_D,plain,(
  s(bool,'T_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f28339,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | s(bool,X1) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | ($spl5 | $spl10)),
  inference(resolution,[],[f1415,f19575])).
fof(f1415,plain,(
  ( ! [X107,X106] : (sP7(X106,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X107)) ) | $spl5),
  inference(resolution,[],[f839,f434])).
fof(f838,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | $spl4),
  inference(cnf_transformation,[],[f838_D])).
fof(f838_D,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f21728,plain,(
  ( ! [X4,X5] : (sP9('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X4,X5) | ~sP6(X5,X4,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | ~sP7(X4,X5,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | $spl206),
  inference(cnf_transformation,[],[f21728_D])).
fof(f21728_D,plain,(
  ( ! [X4,X5] : (sP9('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X4,X5) | ~sP6(X5,X4,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | ~sP7(X4,X5,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f28315,plain,(
  ( ! [X4,X5] : (sP9('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X4,X5) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | ~sP7(X4,X5,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | ~sP6(X5,X4,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | $spl5),
  inference(resolution,[],[f1413,f427])).
fof(f1413,plain,(
  ( ! [X103,X102] : (sP8('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X102,X103)) ) | $spl5),
  inference(resolution,[],[f839,f431])).
fof(f28304,plain,(
  ( ! [X30,X29] : (~p(s(bool,X29)) | s(bool,X30) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | s(bool,'T_0') != s(bool,X30)) ) | ($spl5 | $spl10)),
  inference(resolution,[],[f1412,f19111])).
fof(f1412,plain,(
  ( ! [X101,X100] : (sP8(X100,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X101)) ) | $spl5),
  inference(resolution,[],[f839,f430])).
fof(f3350,plain,(
  ( ! [X9] : (sP9(X9,'F_0','const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | $spl20),
  inference(cnf_transformation,[],[f3350_D])).
fof(f3350_D,plain,(
  ( ! [X9] : (sP9(X9,'F_0','const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f28303,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') != s(bool,X27) | sP9(X28,'F_0','const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | s(bool,X27) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | ($spl5 | $spl10)),
  inference(resolution,[],[f1412,f19068])).
fof(f28302,plain,(
  ( ! [X26] : (s(bool,'T_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | ~p(s(bool,X26))) ) | $spl5),
  inference(resolution,[],[f1412,f3327])).
fof(f3362,plain,(
  s(bool,'F_0') = s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | $spl24),
  inference(cnf_transformation,[],[f3362_D])).
fof(f3362_D,plain,(
  s(bool,'F_0') = s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f28301,plain,(
  ( ! [X25] : (s(bool,'F_0') = s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | ~p(s(bool,X25))) ) | $spl5),
  inference(resolution,[],[f1412,f3285])).
fof(f28298,plain,(
  ( ! [X22] : (sP9(X22,'F_0','const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | s(bool,'T_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | $spl5),
  inference(resolution,[],[f1412,f3227])).
fof(f28296,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,X19) | s(bool,'T_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | ($spl5 | $spl10)),
  inference(resolution,[],[f1412,f19799])).
fof(f21700,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | p(s(bool,X9))) ) | $spl200),
  inference(cnf_transformation,[],[f21700_D])).
fof(f21700_D,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | p(s(bool,X9))) ) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f28291,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | p(s(bool,X9)) | s(bool,'T_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | ($spl5 | $spl10)),
  inference(resolution,[],[f1412,f19159])).
fof(f28285,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(duplicate_literal_removal,[],[f28278])).
fof(f28278,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | ~p(s(bool,X3))) )),
  inference(resolution,[],[f659,f463])).
fof(f28286,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(duplicate_literal_removal,[],[f28275])).
fof(f28275,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(resolution,[],[f659,f464])).
fof(f28198,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | ~p(s(bool,X4)) | p(s(bool,X5))) )),
  inference(resolution,[],[f776,f432])).
fof(f21388,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | $spl160),
  inference(cnf_transformation,[],[f21388_D])).
fof(f21388_D,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f28197,plain,(
  ( ! [X2,X3] : (~p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | $spl10),
  inference(resolution,[],[f776,f19313])).
fof(f21331,plain,(
  ( ! [X25] : (s(bool,'T_0') != s(bool,X25) | s(bool,X25) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | $spl150),
  inference(cnf_transformation,[],[f21331_D])).
fof(f21331_D,plain,(
  ( ! [X25] : (s(bool,'T_0') != s(bool,X25) | s(bool,X25) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f28196,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X1) | s(bool,X1) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | $spl10),
  inference(resolution,[],[f776,f19575])).
fof(f21374,plain,(
  ( ! [X6,X7] : (sP9(X6,X7,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | $spl158),
  inference(cnf_transformation,[],[f21374_D])).
fof(f21374_D,plain,(
  ( ! [X6,X7] : (sP9(X6,X7,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f28194,plain,(
  ( ! [X6,X7] : (sP9(X6,X7,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))) )),
  inference(subsumption_resolution,[],[f28193,f772])).
fof(f28193,plain,(
  ( ! [X6,X7] : (sP9(X6,X7,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | ~sP8(X6,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X7)) )),
  inference(subsumption_resolution,[],[f28184,f439])).
fof(f28184,plain,(
  ( ! [X6,X7] : (sP9(X6,X7,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | ~sP8(X6,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X7) | ~sP6('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X7,X6)) )),
  inference(resolution,[],[f775,f425])).
fof(f28183,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | p(s(bool,X5))) )),
  inference(resolution,[],[f775,f432])).
fof(f21371,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | $spl156),
  inference(cnf_transformation,[],[f21371_D])).
fof(f21371_D,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f28182,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | p(s(bool,X2)) | s(bool,X3) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | $spl10),
  inference(resolution,[],[f775,f19313])).
fof(f3338,plain,(
  s(bool,'T_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | $spl15),
  inference(cnf_transformation,[],[f3338_D])).
fof(f3338_D,plain,(
  s(bool,'T_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f28181,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | s(bool,X1) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | $spl10),
  inference(resolution,[],[f775,f19575])).
fof(f28172,plain,(
  ( ! [X30,X29] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | s(bool,X29) != s(bool,X30) | s(bool,'T_0') != s(bool,X29)) ) | $spl10),
  inference(resolution,[],[f773,f19111])).
fof(f28170,plain,(
  ( ! [X26] : (s(bool,'T_0') != s(bool,X26) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))) )),
  inference(resolution,[],[f773,f3327])).
fof(f28169,plain,(
  ( ! [X25] : (s(bool,'F_0') = s(bool,X25) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))) )),
  inference(resolution,[],[f773,f3285])).
fof(f28167,plain,(
  ( ! [X23] : (~p(s(bool,X23)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))) )),
  inference(resolution,[],[f773,f3242])).
fof(f28165,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | s(bool,'T_0') = s(bool,X21) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))) ) | ($spl10 | $spl96)),
  inference(resolution,[],[f773,f27175])).
fof(f28164,plain,(
  ( ! [X19,X18] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | s(bool,'F_0') != s(bool,X18) | s(bool,'T_0') != s(bool,X19)) ) | $spl10),
  inference(resolution,[],[f773,f19799])).
fof(f28163,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,X16) | ~p(s(bool,X17)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))) ) | $spl10),
  inference(resolution,[],[f773,f19796])).
fof(f28158,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | ~p(s(bool,X6)) | p(s(bool,X7))) )),
  inference(resolution,[],[f773,f428])).
fof(f21242,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | $spl132),
  inference(cnf_transformation,[],[f21242_D])).
fof(f21242_D,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f21357,plain,(
  ( ! [X4,X5] : (sP9('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X4,X5) | ~sP6(X5,X4,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | ~sP7(X4,X5,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | $spl152),
  inference(cnf_transformation,[],[f21357_D])).
fof(f21357_D,plain,(
  ( ! [X4,X5] : (sP9('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X4,X5) | ~sP6(X5,X4,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | ~sP7(X4,X5,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f28157,plain,(
  ( ! [X4,X5] : (sP9('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X4,X5) | p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | ~sP7(X4,X5,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | ~sP6(X5,X4,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) )),
  inference(resolution,[],[f773,f427])).
fof(f28146,plain,(
  ( ! [X30,X29] : (~p(s(bool,X29)) | s(bool,X30) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | s(bool,'T_0') != s(bool,X30)) ) | $spl10),
  inference(resolution,[],[f772,f19111])).
fof(f3349,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | $spl18),
  inference(cnf_transformation,[],[f3349_D])).
fof(f3349_D,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f28145,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') != s(bool,X27) | sP9(X28,'F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | s(bool,X27) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | $spl10),
  inference(resolution,[],[f772,f19068])).
fof(f28144,plain,(
  ( ! [X26] : (s(bool,'T_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | ~p(s(bool,X26))) )),
  inference(resolution,[],[f772,f3327])).
fof(f3360,plain,(
  s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | $spl22),
  inference(cnf_transformation,[],[f3360_D])).
fof(f3360_D,plain,(
  s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f28143,plain,(
  ( ! [X25] : (s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | ~p(s(bool,X25))) )),
  inference(resolution,[],[f772,f3285])).
fof(f28141,plain,(
  ( ! [X23] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | ~p(s(bool,X23))) )),
  inference(resolution,[],[f772,f3242])).
fof(f28140,plain,(
  ( ! [X22] : (sP9(X22,'F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | s(bool,'T_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) )),
  inference(resolution,[],[f772,f3227])).
fof(f28139,plain,(
  ( ! [X21,X20] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | s(bool,'T_0') = s(bool,X20) | ~p(s(bool,X21))) ) | ($spl10 | $spl96)),
  inference(resolution,[],[f772,f27175])).
fof(f28138,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,X19) | s(bool,'T_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | $spl10),
  inference(resolution,[],[f772,f19799])).
fof(f28137,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,X16) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | ~p(s(bool,X17))) ) | $spl10),
  inference(resolution,[],[f772,f19796])).
fof(f21327,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | ~sP7(X1,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X0)) ) | $spl144),
  inference(cnf_transformation,[],[f21327_D])).
fof(f21327_D,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | ~sP7(X1,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X0)) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f28135,plain,(
  ( ! [X12,X13] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | sP9(X12,X13,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | ~sP7(X13,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X12)) ) | $spl10),
  inference(resolution,[],[f772,f19262])).
fof(f21330,plain,(
  ( ! [X10,X11] : (sP9(X10,X11,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | s(bool,'F_0') != s(bool,X11)) ) | $spl148),
  inference(cnf_transformation,[],[f21330_D])).
fof(f21330_D,plain,(
  ( ! [X10,X11] : (sP9(X10,X11,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | s(bool,'F_0') != s(bool,X11)) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f28134,plain,(
  ( ! [X10,X11] : (sP9(X10,X11,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | s(bool,'F_0') != s(bool,X11)) ) | $spl10),
  inference(resolution,[],[f772,f19168])).
fof(f21329,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | p(s(bool,X9))) ) | $spl146),
  inference(cnf_transformation,[],[f21329_D])).
fof(f21329_D,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | p(s(bool,X9))) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f28133,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | p(s(bool,X9)) | s(bool,'T_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | $spl10),
  inference(resolution,[],[f772,f19159])).
fof(f28132,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | p(s(bool,X7))) )),
  inference(resolution,[],[f772,f428])).
fof(f28153,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | ~sP7(X1,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X0)) )),
  inference(subsumption_resolution,[],[f28129,f439])).
fof(f28129,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | ~sP7(X1,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X0) | ~sP6('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X1,X0)) )),
  inference(resolution,[],[f772,f425])).
fof(f21300,plain,(
  ( ! [X2,X3] : (sP5(X2,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X3) | ~p(s(bool,X3))) ) | $spl142),
  inference(cnf_transformation,[],[f21300_D])).
fof(f21300_D,plain,(
  ( ! [X2,X3] : (sP5(X2,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X3) | ~p(s(bool,X3))) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f28117,plain,(
  ( ! [X2,X3] : (sP5(X2,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X3) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | ~p(s(bool,X3))) )),
  inference(resolution,[],[f760,f407])).
fof(f28116,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | ~p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(resolution,[],[f760,f411])).
fof(f21286,plain,(
  ( ! [X2,X3] : (sP5(X2,X3,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | ~p(s(bool,X3))) ) | $spl140),
  inference(cnf_transformation,[],[f21286_D])).
fof(f21286_D,plain,(
  ( ! [X2,X3] : (sP5(X2,X3,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | ~p(s(bool,X3))) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f28104,plain,(
  ( ! [X2,X3] : (sP5(X2,X3,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | ~p(s(bool,X3)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))) )),
  inference(resolution,[],[f759,f407])).
fof(f28103,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | p(s(bool,X1))) )),
  inference(resolution,[],[f759,f411])).
fof(f21272,plain,(
  ( ! [X2,X3] : (sP3('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X2,X3)) ) | $spl138),
  inference(cnf_transformation,[],[f21272_D])).
fof(f21272_D,plain,(
  ( ! [X2,X3] : (sP3('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X2,X3)) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f28102,plain,(
  ( ! [X6,X7] : (sP3('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X6,X7) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))) )),
  inference(subsumption_resolution,[],[f28093,f400])).
fof(f28093,plain,(
  ( ! [X6,X7] : (sP3('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X6,X7) | p(s(bool,X6)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))) )),
  inference(resolution,[],[f746,f392])).
fof(f28101,plain,(
  ( ! [X4,X5] : (sP3('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X4,X5) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))) )),
  inference(subsumption_resolution,[],[f28092,f399])).
fof(f28092,plain,(
  ( ! [X4,X5] : (sP3('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X4,X5) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | p(s(bool,X5))) )),
  inference(resolution,[],[f746,f391])).
fof(f28091,plain,(
  ( ! [X2,X3] : (sP3('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X2,X3) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))) )),
  inference(resolution,[],[f746,f586])).
fof(f28090,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(resolution,[],[f746,f394])).
fof(f21258,plain,(
  ( ! [X2,X3] : (sP1('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X2,X3)) ) | $spl136),
  inference(cnf_transformation,[],[f21258_D])).
fof(f21258_D,plain,(
  ( ! [X2,X3] : (sP1('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X2,X3)) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f28089,plain,(
  ( ! [X6,X7] : (sP1('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X6,X7) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))) )),
  inference(subsumption_resolution,[],[f28080,f383])).
fof(f28080,plain,(
  ( ! [X6,X7] : (sP1('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X6,X7) | p(s(bool,X6)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))) )),
  inference(resolution,[],[f730,f375])).
fof(f28088,plain,(
  ( ! [X4,X5] : (sP1('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X4,X5) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))) )),
  inference(subsumption_resolution,[],[f28079,f382])).
fof(f28079,plain,(
  ( ! [X4,X5] : (sP1('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X4,X5) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | ~p(s(bool,X5))) )),
  inference(resolution,[],[f730,f374])).
fof(f28078,plain,(
  ( ! [X2,X3] : (sP1('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X2,X3) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))) )),
  inference(resolution,[],[f730,f585])).
fof(f28077,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(resolution,[],[f730,f377])).
fof(f21246,plain,(
  ( ! [X8,X9] : (sP1(X8,X9,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | p(s(bool,X9))) ) | $spl134),
  inference(cnf_transformation,[],[f21246_D])).
fof(f21246_D,plain,(
  ( ! [X8,X9] : (sP1(X8,X9,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | p(s(bool,X9))) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f28068,plain,(
  ( ! [X8,X9] : (sP1(X8,X9,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | p(s(bool,X9)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))) )),
  inference(resolution,[],[f729,f376])).
fof(f28064,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))) )),
  inference(resolution,[],[f729,f377])).
fof(f5092,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) | $spl67),
  inference(cnf_transformation,[],[f5092_D])).
fof(f5092_D,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) <=> ~$spl67),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl67])])).
fof(f27812,plain,(
  ( ! [X12] : (p(s(bool,X12)) | s(bool,'F_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) )),
  inference(resolution,[],[f656,f3286])).
fof(f27808,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | p(s(bool,X4)) | p(s(bool,X5))) )),
  inference(resolution,[],[f656,f436])).
fof(f21228,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) | $spl130),
  inference(cnf_transformation,[],[f21228_D])).
fof(f21228_D,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f27807,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) | $spl10),
  inference(resolution,[],[f656,f19474])).
fof(f21227,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,X1) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) | $spl128),
  inference(cnf_transformation,[],[f21227_D])).
fof(f21227_D,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,X1) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f27806,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'F_0') != s(bool,X1) | s(bool,X1) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) | $spl10),
  inference(resolution,[],[f656,f19603])).
fof(f5075,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X5)) ) | $spl62),
  inference(cnf_transformation,[],[f5075_D])).
fof(f5075_D,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X5)) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f27804,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X9) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(subsumption_resolution,[],[f27803,f650])).
fof(f27803,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X9) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | ~sP8(X8,X9,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) )),
  inference(subsumption_resolution,[],[f27795,f435])).
fof(f27795,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X9) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | ~sP8(X8,X9,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) | ~sP7('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X9,X8)) )),
  inference(resolution,[],[f655,f426])).
fof(f27793,plain,(
  ( ! [X4,X5] : (p(s(bool,X4)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | p(s(bool,X5))) )),
  inference(resolution,[],[f655,f436])).
fof(f21210,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) | $spl126),
  inference(cnf_transformation,[],[f21210_D])).
fof(f21210_D,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f27792,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | p(s(bool,X2)) | s(bool,X3) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) | $spl10),
  inference(resolution,[],[f655,f19474])).
fof(f27791,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'F_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) | s(bool,X1) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) | $spl10),
  inference(resolution,[],[f655,f19603])).
fof(f27781,plain,(
  ( ! [X12] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | s(bool,'F_0') != s(bool,X12)) )),
  inference(resolution,[],[f654,f3286])).
fof(f5060,plain,(
  ( ! [X6,X7] : (sP9('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X6,X7)) ) | $spl58),
  inference(cnf_transformation,[],[f5060_D])).
fof(f5060_D,plain,(
  ( ! [X6,X7] : (sP9('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X6,X7)) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f27790,plain,(
  ( ! [X10,X11] : (sP9('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X10,X11) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(subsumption_resolution,[],[f27789,f653])).
fof(f27789,plain,(
  ( ! [X10,X11] : (sP9('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X10,X11) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | ~sP7(X10,X11,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) )),
  inference(subsumption_resolution,[],[f27780,f431])).
fof(f27780,plain,(
  ( ! [X10,X11] : (sP9('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X10,X11) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | ~sP8('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X11,X10) | ~sP7(X10,X11,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) )),
  inference(resolution,[],[f654,f427])).
fof(f27777,plain,(
  ( ! [X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(resolution,[],[f654,f436])).
fof(f27776,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | s(bool,X2) != s(bool,X3)) ) | $spl10),
  inference(resolution,[],[f654,f19474])).
fof(f27775,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | s(bool,'F_0') != s(bool,X0) | s(bool,X0) != s(bool,X1)) ) | $spl10),
  inference(resolution,[],[f654,f19603])).
fof(f5044,plain,(
  ( ! [X6,X7] : (sP9('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X6,X7) | ~sP6(X7,X6,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) | $spl56),
  inference(cnf_transformation,[],[f5044_D])).
fof(f5044_D,plain,(
  ( ! [X6,X7] : (sP9('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X6,X7) | ~sP6(X7,X6,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f27774,plain,(
  ( ! [X10,X11] : (sP9('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X10,X11) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | ~sP6(X11,X10,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) )),
  inference(subsumption_resolution,[],[f27766,f431])).
fof(f27766,plain,(
  ( ! [X10,X11] : (sP9('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X10,X11) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | ~sP8('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X11,X10) | ~sP6(X11,X10,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) )),
  inference(resolution,[],[f653,f427])).
fof(f27763,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | ~p(s(bool,X5)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(resolution,[],[f653,f432])).
fof(f27762,plain,(
  ( ! [X2,X3] : (~p(s(bool,X2)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | s(bool,X2) != s(bool,X3)) ) | $spl10),
  inference(resolution,[],[f653,f19313])).
fof(f27761,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | s(bool,'T_0') != s(bool,X0) | s(bool,X0) != s(bool,X1)) ) | $spl10),
  inference(resolution,[],[f653,f19575])).
fof(f27760,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) | $spl372),
  inference(cnf_transformation,[],[f27760_D])).
fof(f27760_D,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) <=> ~$spl372),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl372])])).
fof(f27750,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | s(bool,'T_0') = s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) | ~p(s(bool,X21))) ) | ($spl10 | $spl96)),
  inference(resolution,[],[f650,f27175])).
fof(f27749,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) | s(bool,'T_0') != s(bool,X19)) ) | $spl10),
  inference(resolution,[],[f650,f19799])).
fof(f27748,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) | ~p(s(bool,X16)) | ~p(s(bool,X17))) ) | $spl10),
  inference(resolution,[],[f650,f19796])).
fof(f5031,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X5) | ~sP6(X5,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X4)) ) | $spl54),
  inference(cnf_transformation,[],[f5031_D])).
fof(f5031_D,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X5) | ~sP6(X5,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X4)) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f27747,plain,(
  ( ! [X14,X15] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | sP9(X14,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X15) | ~sP6(X15,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X14)) ) | $spl10),
  inference(resolution,[],[f650,f19385])).
fof(f21163,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X11) | ~p(s(bool,X11))) ) | $spl116),
  inference(cnf_transformation,[],[f21163_D])).
fof(f21163_D,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X11) | ~p(s(bool,X11))) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f27745,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X11) | ~p(s(bool,X11)) | s(bool,'F_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) | $spl10),
  inference(resolution,[],[f650,f19168])).
fof(f21162,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X9) | s(bool,'T_0') != s(bool,X9)) ) | $spl114),
  inference(cnf_transformation,[],[f21162_D])).
fof(f21162_D,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X9) | s(bool,'T_0') != s(bool,X9)) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f27744,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X9) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | s(bool,'T_0') != s(bool,X9)) ) | $spl10),
  inference(resolution,[],[f650,f19159])).
fof(f27743,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(resolution,[],[f650,f428])).
fof(f27758,plain,(
  ( ! [X2,X3] : (sP9(X2,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X3) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | ~sP6(X3,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X2)) )),
  inference(subsumption_resolution,[],[f27741,f435])).
fof(f27741,plain,(
  ( ! [X2,X3] : (sP9(X2,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X3) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | ~sP7('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X3,X2) | ~sP6(X3,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X2)) )),
  inference(resolution,[],[f650,f426])).
fof(f5017,plain,(
  ( ! [X4,X5] : (sP5('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X4,X5)) ) | $spl52),
  inference(cnf_transformation,[],[f5017_D])).
fof(f5017_D,plain,(
  ( ! [X4,X5] : (sP5('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X4,X5)) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f27731,plain,(
  ( ! [X8,X9] : (sP5('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X8,X9) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(resolution,[],[f637,f587])).
fof(f27739,plain,(
  ( ! [X6,X7] : (sP5('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X6,X7) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(subsumption_resolution,[],[f27730,f415])).
fof(f27730,plain,(
  ( ! [X6,X7] : (sP5('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X6,X7) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | ~p(s(bool,X7))) )),
  inference(resolution,[],[f637,f409])).
fof(f27738,plain,(
  ( ! [X4,X5] : (sP5('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X4,X5) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(subsumption_resolution,[],[f27729,f416])).
fof(f27729,plain,(
  ( ! [X4,X5] : (sP5('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X4,X5) | ~p(s(bool,X4)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(resolution,[],[f637,f408])).
fof(f27727,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(resolution,[],[f637,f411])).
fof(f5004,plain,(
  ( ! [X8,X9] : (sP3(X8,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X9) | p(s(bool,X9))) ) | $spl50),
  inference(cnf_transformation,[],[f5004_D])).
fof(f5004_D,plain,(
  ( ! [X8,X9] : (sP3(X8,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X9) | p(s(bool,X9))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f27718,plain,(
  ( ! [X8,X9] : (sP3(X8,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X9) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | p(s(bool,X9))) )),
  inference(resolution,[],[f621,f393])).
fof(f27714,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | p(s(bool,X1))) )),
  inference(resolution,[],[f621,f394])).
fof(f4990,plain,(
  ( ! [X8,X9] : (sP3(X8,X9,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) | p(s(bool,X9))) ) | $spl48),
  inference(cnf_transformation,[],[f4990_D])).
fof(f4990_D,plain,(
  ( ! [X8,X9] : (sP3(X8,X9,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) | p(s(bool,X9))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f27705,plain,(
  ( ! [X8,X9] : (sP3(X8,X9,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) | p(s(bool,X9)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(resolution,[],[f620,f393])).
fof(f27701,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(resolution,[],[f620,f394])).
fof(f4956,plain,(
  ( ! [X8,X9] : (sP1(X8,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X9) | ~p(s(bool,X9))) ) | $spl46),
  inference(cnf_transformation,[],[f4956_D])).
fof(f4956_D,plain,(
  ( ! [X8,X9] : (sP1(X8,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X9) | ~p(s(bool,X9))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f27692,plain,(
  ( ! [X8,X9] : (sP1(X8,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X9) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | ~p(s(bool,X9))) )),
  inference(resolution,[],[f607,f376])).
fof(f27688,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | ~p(s(bool,X1))) )),
  inference(resolution,[],[f607,f377])).
fof(f27667,plain,(
  ( ! [X1] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | p(s(bool,X1))) )),
  inference(subsumption_resolution,[],[f27659,f538])).
fof(f27659,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | p(s(bool,X1))) )),
  inference(resolution,[],[f668,f474])).
fof(f27666,plain,(
  ( ! [X0] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f27658,f538])).
fof(f27658,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | p(s(bool,X0))) )),
  inference(resolution,[],[f668,f478])).
fof(f27657,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(subsumption_resolution,[],[f27649,f538])).
fof(f27649,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(resolution,[],[f667,f474])).
fof(f27656,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(subsumption_resolution,[],[f27648,f538])).
fof(f27648,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(resolution,[],[f667,f478])).
fof(f27543,plain,(
  $false | ($spl7 | $spl264 | $spl368)),
  inference(subsumption_resolution,[],[f27436,f24514])).
fof(f27436,plain,(
  ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ($spl7 | $spl368)),
  inference(backward_demodulation,[],[f27364,f12719])).
fof(f27364,plain,(
  s(sK12,sK15) = s(sK12,sK16) | $spl368),
  inference(cnf_transformation,[],[f27364_D])).
fof(f27364_D,plain,(
  s(sK12,sK15) = s(sK12,sK16) <=> ~$spl368),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl368])])).
fof(f27366,plain,(
  p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15))),s(sK12,sK16)))) | $spl370),
  inference(cnf_transformation,[],[f27366_D])).
fof(f27366_D,plain,(
  p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15))),s(sK12,sK16)))) <=> ~$spl370),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl370])])).
fof(f27251,plain,(
  p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15))),s(sK12,sK16)))) | s(sK12,sK15) = s(sK12,sK16) | $spl0),
  inference(resolution,[],[f712,f578])).
fof(f712,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | $spl0),
  inference(cnf_transformation,[],[f712_D])).
fof(f712_D,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f714,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | $spl2),
  inference(cnf_transformation,[],[f714_D])).
fof(f714_D,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f27250,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))),
  inference(subsumption_resolution,[],[f2711,f588])).
fof(f2711,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))),
  inference(superposition,[],[f355,f680])).
fof(f27249,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))),
  inference(subsumption_resolution,[],[f6031,f588])).
fof(f6031,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))),
  inference(superposition,[],[f355,f680])).
fof(f27245,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | $spl364),
  inference(cnf_transformation,[],[f27245_D])).
fof(f27245_D,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) <=> ~$spl364),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl364])])).
fof(f27227,plain,(
  ( ! [X24,X25] : (~p(s(bool,X24)) | s(bool,'T_0') = s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | ~p(s(bool,X25))) ) | ($spl10 | $spl96 | $spl240)),
  inference(resolution,[],[f27175,f23759])).
fof(f23759,plain,(
  ( ! [X101,X100] : (sP8(X100,X101,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | $spl240),
  inference(resolution,[],[f22640,f429])).
fof(f27243,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | $spl362),
  inference(cnf_transformation,[],[f27243_D])).
fof(f27243_D,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) <=> ~$spl362),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl362])])).
fof(f27226,plain,(
  ( ! [X23,X22] : (~p(s(bool,X22)) | s(bool,'T_0') = s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~p(s(bool,X23))) ) | ($spl10 | $spl96)),
  inference(resolution,[],[f27175,f1537])).
fof(f1537,plain,(
  ( ! [X101,X100] : (sP8(X100,X101,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(resolution,[],[f882,f429])).
fof(f882,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))),
  inference(resolution,[],[f350,f580])).
fof(f580,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f348])).
fof(f27241,plain,(
  s(bool,'T_0') = s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | $spl360),
  inference(cnf_transformation,[],[f27241_D])).
fof(f27241_D,plain,(
  s(bool,'T_0') = s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) <=> ~$spl360),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl360])])).
fof(f27225,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | s(bool,'T_0') = s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~p(s(bool,X21))) ) | ($spl10 | $spl96)),
  inference(resolution,[],[f27175,f1682])).
fof(f1682,plain,(
  ( ! [X109,X108] : (sP8(X108,X109,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(resolution,[],[f884,f429])).
fof(f27239,plain,(
  s(bool,'T_0') = s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) | $spl358),
  inference(cnf_transformation,[],[f27239_D])).
fof(f27239_D,plain,(
  s(bool,'T_0') = s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) <=> ~$spl358),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl358])])).
fof(f27224,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'T_0') = s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) | ~p(s(bool,X19))) ) | ($spl2 | $spl10 | $spl96)),
  inference(resolution,[],[f27175,f1162])).
fof(f27237,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | $spl356),
  inference(cnf_transformation,[],[f27237_D])).
fof(f27237_D,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) <=> ~$spl356),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl356])])).
fof(f27223,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | s(bool,'T_0') = s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~p(s(bool,X17))) ) | ($spl10 | $spl96)),
  inference(resolution,[],[f27175,f2208])).
fof(f2208,plain,(
  ( ! [X107,X108] : (sP8(X107,X108,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(resolution,[],[f1625,f429])).
fof(f27222,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') = s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | ~p(s(bool,X15))) ) | ($spl10 | $spl42 | $spl96)),
  inference(resolution,[],[f27175,f22578])).
fof(f22578,plain,(
  ( ! [X107,X108] : (sP8(X107,X108,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | $spl42),
  inference(resolution,[],[f4751,f429])).
fof(f26866,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | $spl345),
  inference(cnf_transformation,[],[f26866_D])).
fof(f26866_D,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) <=> ~$spl345),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl345])])).
fof(f26939,plain,(
  ( ! [X12] : (p(s(bool,X12)) | s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | $spl240),
  inference(resolution,[],[f23765,f3286])).
fof(f23765,plain,(
  ( ! [X112,X113] : (sP6('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X112,X113)) ) | $spl240),
  inference(resolution,[],[f22640,f439])).
fof(f26947,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | $spl352),
  inference(cnf_transformation,[],[f26947_D])).
fof(f26947_D,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) <=> ~$spl352),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl352])])).
fof(f26934,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | ($spl10 | $spl240)),
  inference(resolution,[],[f23765,f19474])).
fof(f26946,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | $spl350),
  inference(cnf_transformation,[],[f26946_D])).
fof(f26946_D,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) <=> ~$spl350),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl350])])).
fof(f26933,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'F_0') != s(bool,X1) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | ($spl10 | $spl240)),
  inference(resolution,[],[f23765,f19603])).
fof(f26909,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | $spl348),
  inference(cnf_transformation,[],[f26909_D])).
fof(f26909_D,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) <=> ~$spl348),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl348])])).
fof(f26897,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | ($spl10 | $spl240)),
  inference(resolution,[],[f23764,f19603])).
fof(f23764,plain,(
  ( ! [X111,X110] : (sP6(X110,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X111)) ) | $spl240),
  inference(resolution,[],[f22640,f438])).
fof(f26856,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | s(bool,'T_0') != s(bool,X19)) ) | ($spl10 | $spl240)),
  inference(resolution,[],[f23759,f19799])).
fof(f26855,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | ~p(s(bool,X16)) | ~p(s(bool,X17))) ) | ($spl10 | $spl240)),
  inference(resolution,[],[f23759,f19796])).
fof(f26867,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X11) | ~p(s(bool,X11))) ) | $spl346),
  inference(cnf_transformation,[],[f26867_D])).
fof(f26867_D,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X11) | ~p(s(bool,X11))) ) <=> ~$spl346),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl346])])).
fof(f26852,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X11) | ~p(s(bool,X11)) | s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | ($spl10 | $spl240)),
  inference(resolution,[],[f23759,f19168])).
fof(f26214,plain,(
  ( ! [X12] : (p(s(bool,X12)) | s(bool,'F_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | $spl42),
  inference(resolution,[],[f22584,f3286])).
fof(f22584,plain,(
  ( ! [X120,X119] : (sP6('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X119,X120)) ) | $spl42),
  inference(resolution,[],[f4751,f439])).
fof(f26209,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | ($spl10 | $spl42)),
  inference(resolution,[],[f22584,f19474])).
fof(f26208,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'F_0') != s(bool,X1) | s(bool,X1) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | ($spl10 | $spl42)),
  inference(resolution,[],[f22584,f19603])).
fof(f26175,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'F_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | s(bool,X1) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | ($spl10 | $spl42)),
  inference(resolution,[],[f22583,f19603])).
fof(f22583,plain,(
  ( ! [X118,X117] : (sP6(X117,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X118)) ) | $spl42),
  inference(resolution,[],[f4751,f438])).
fof(f26140,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | s(bool,'T_0') != s(bool,X19)) ) | ($spl10 | $spl42)),
  inference(resolution,[],[f22578,f19799])).
fof(f26139,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | ~p(s(bool,X16)) | ~p(s(bool,X17))) ) | ($spl10 | $spl42)),
  inference(resolution,[],[f22578,f19796])).
fof(f26136,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X11) | ~p(s(bool,X11)) | s(bool,'F_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | ($spl10 | $spl42)),
  inference(resolution,[],[f22578,f19168])).
fof(f7516,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | $spl69),
  inference(cnf_transformation,[],[f7516_D])).
fof(f7516_D,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) <=> ~$spl69),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl69])])).
fof(f26055,plain,(
  ( ! [X12] : (p(s(bool,X12)) | s(bool,'F_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(resolution,[],[f2214,f3286])).
fof(f2214,plain,(
  ( ! [X120,X119] : (sP6('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X119,X120)) )),
  inference(resolution,[],[f1625,f439])).
fof(f24514,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | $spl264),
  inference(cnf_transformation,[],[f24514_D])).
fof(f24514_D,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) <=> ~$spl264),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl264])])).
fof(f26051,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X4)) | p(s(bool,X5))) )),
  inference(resolution,[],[f2214,f436])).
fof(f26063,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl342),
  inference(cnf_transformation,[],[f26063_D])).
fof(f26063_D,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) <=> ~$spl342),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl342])])).
fof(f26050,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl10),
  inference(resolution,[],[f2214,f19474])).
fof(f26062,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,X1) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl340),
  inference(cnf_transformation,[],[f26062_D])).
fof(f26062_D,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,X1) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) <=> ~$spl340),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl340])])).
fof(f26049,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'F_0') != s(bool,X1) | s(bool,X1) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl10),
  inference(resolution,[],[f2214,f19603])).
fof(f26047,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9)) ) | $spl338),
  inference(cnf_transformation,[],[f26047_D])).
fof(f26047_D,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9)) ) <=> ~$spl338),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl338])])).
fof(f26046,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(subsumption_resolution,[],[f26045,f2208])).
fof(f26045,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP8(X8,X9,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(subsumption_resolution,[],[f26036,f435])).
fof(f26036,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP8(X8,X9,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~sP7('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9,X8)) )),
  inference(resolution,[],[f2213,f426])).
fof(f2213,plain,(
  ( ! [X118,X117] : (sP6(X117,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X118)) )),
  inference(resolution,[],[f1625,f438])).
fof(f26034,plain,(
  ( ! [X4,X5] : (p(s(bool,X4)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X5))) )),
  inference(resolution,[],[f2213,f436])).
fof(f26044,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl336),
  inference(cnf_transformation,[],[f26044_D])).
fof(f26044_D,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) <=> ~$spl336),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl336])])).
fof(f26033,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X2)) | s(bool,X3) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl10),
  inference(resolution,[],[f2213,f19474])).
fof(f26032,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'F_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | s(bool,X1) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl10),
  inference(resolution,[],[f2213,f19603])).
fof(f25995,plain,(
  ( ! [X12] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | s(bool,'F_0') != s(bool,X12)) )),
  inference(resolution,[],[f2212,f3286])).
fof(f2212,plain,(
  ( ! [X116,X115] : (sP6(X115,X116,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(resolution,[],[f1625,f437])).
fof(f26005,plain,(
  ( ! [X10,X11] : (sP9('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X10,X11)) ) | $spl334),
  inference(cnf_transformation,[],[f26005_D])).
fof(f26005_D,plain,(
  ( ! [X10,X11] : (sP9('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X10,X11)) ) <=> ~$spl334),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl334])])).
fof(f26004,plain,(
  ( ! [X10,X11] : (sP9('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X10,X11) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(subsumption_resolution,[],[f26003,f2211])).
fof(f2211,plain,(
  ( ! [X114,X113] : (sP7(X113,X114,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(resolution,[],[f1625,f433])).
fof(f26003,plain,(
  ( ! [X10,X11] : (sP9('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X10,X11) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP7(X10,X11,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(subsumption_resolution,[],[f25994,f431])).
fof(f25994,plain,(
  ( ! [X10,X11] : (sP9('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X10,X11) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP8('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X11,X10) | ~sP7(X10,X11,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(resolution,[],[f2212,f427])).
fof(f25991,plain,(
  ( ! [X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f2212,f436])).
fof(f25990,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | s(bool,X2) != s(bool,X3)) ) | $spl10),
  inference(resolution,[],[f2212,f19474])).
fof(f25989,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | s(bool,'F_0') != s(bool,X0) | s(bool,X0) != s(bool,X1)) ) | $spl10),
  inference(resolution,[],[f2212,f19603])).
fof(f25988,plain,(
  ( ! [X10,X11] : (sP9('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X10,X11) | ~sP6(X11,X10,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl332),
  inference(cnf_transformation,[],[f25988_D])).
fof(f25988_D,plain,(
  ( ! [X10,X11] : (sP9('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X10,X11) | ~sP6(X11,X10,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) <=> ~$spl332),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl332])])).
fof(f25987,plain,(
  ( ! [X10,X11] : (sP9('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X10,X11) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP6(X11,X10,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(subsumption_resolution,[],[f25979,f431])).
fof(f25979,plain,(
  ( ! [X10,X11] : (sP9('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X10,X11) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP8('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X11,X10) | ~sP6(X11,X10,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(resolution,[],[f2211,f427])).
fof(f25976,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | ~p(s(bool,X5)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f2211,f432])).
fof(f25975,plain,(
  ( ! [X2,X3] : (~p(s(bool,X2)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | s(bool,X2) != s(bool,X3)) ) | $spl10),
  inference(resolution,[],[f2211,f19313])).
fof(f25974,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | s(bool,'T_0') != s(bool,X0) | s(bool,X0) != s(bool,X1)) ) | $spl10),
  inference(resolution,[],[f2211,f19575])).
fof(f25962,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | s(bool,'T_0') != s(bool,X19)) ) | $spl10),
  inference(resolution,[],[f2208,f19799])).
fof(f25961,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~p(s(bool,X16)) | ~p(s(bool,X17))) ) | $spl10),
  inference(resolution,[],[f2208,f19796])).
fof(f25971,plain,(
  ( ! [X2,X3] : (sP9(X2,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X3) | ~sP6(X3,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X2)) ) | $spl326),
  inference(cnf_transformation,[],[f25971_D])).
fof(f25971_D,plain,(
  ( ! [X2,X3] : (sP9(X2,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X3) | ~sP6(X3,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X2)) ) <=> ~$spl326),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl326])])).
fof(f25960,plain,(
  ( ! [X14,X15] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | sP9(X14,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X15) | ~sP6(X15,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X14)) ) | $spl10),
  inference(resolution,[],[f2208,f19385])).
fof(f25973,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X11) | ~p(s(bool,X11))) ) | $spl330),
  inference(cnf_transformation,[],[f25973_D])).
fof(f25973_D,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X11) | ~p(s(bool,X11))) ) <=> ~$spl330),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl330])])).
fof(f25958,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X11) | ~p(s(bool,X11)) | s(bool,'F_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl10),
  inference(resolution,[],[f2208,f19168])).
fof(f25972,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9) | s(bool,'T_0') != s(bool,X9)) ) | $spl328),
  inference(cnf_transformation,[],[f25972_D])).
fof(f25972_D,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9) | s(bool,'T_0') != s(bool,X9)) ) <=> ~$spl328),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl328])])).
fof(f25957,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | s(bool,'T_0') != s(bool,X9)) ) | $spl10),
  inference(resolution,[],[f2208,f19159])).
fof(f25956,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f2208,f428])).
fof(f25970,plain,(
  ( ! [X2,X3] : (sP9(X2,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X3) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP6(X3,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X2)) )),
  inference(subsumption_resolution,[],[f25954,f435])).
fof(f25954,plain,(
  ( ! [X2,X3] : (sP9(X2,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X3) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP7('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X3,X2) | ~sP6(X3,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X2)) )),
  inference(resolution,[],[f2208,f426])).
fof(f25951,plain,(
  ( ! [X4,X5] : (sP5('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X4,X5)) ) | $spl324),
  inference(cnf_transformation,[],[f25951_D])).
fof(f25951_D,plain,(
  ( ! [X4,X5] : (sP5('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X4,X5)) ) <=> ~$spl324),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl324])])).
fof(f25943,plain,(
  ( ! [X8,X9] : (sP5('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X8,X9) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f2195,f587])).
fof(f2195,plain,(
  ( ! [X81,X82] : (sP4(X81,X82,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(resolution,[],[f1625,f412])).
fof(f25952,plain,(
  ( ! [X6,X7] : (sP5('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X6,X7) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(subsumption_resolution,[],[f25942,f415])).
fof(f25942,plain,(
  ( ! [X6,X7] : (sP5('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X6,X7) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~p(s(bool,X7))) )),
  inference(resolution,[],[f2195,f409])).
fof(f25950,plain,(
  ( ! [X4,X5] : (sP5('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X4,X5) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(subsumption_resolution,[],[f25941,f416])).
fof(f25941,plain,(
  ( ! [X4,X5] : (sP5('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X4,X5) | ~p(s(bool,X4)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f2195,f408])).
fof(f25939,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f2195,f411])).
fof(f25938,plain,(
  ( ! [X8,X9] : (sP3(X8,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9) | p(s(bool,X9))) ) | $spl322),
  inference(cnf_transformation,[],[f25938_D])).
fof(f25938_D,plain,(
  ( ! [X8,X9] : (sP3(X8,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9) | p(s(bool,X9))) ) <=> ~$spl322),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl322])])).
fof(f25929,plain,(
  ( ! [X8,X9] : (sP3(X8,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X9))) )),
  inference(resolution,[],[f2179,f393])).
fof(f2179,plain,(
  ( ! [X50,X49] : (sP2(X49,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X50)) )),
  inference(resolution,[],[f1625,f396])).
fof(f25925,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X1))) )),
  inference(resolution,[],[f2179,f394])).
fof(f25897,plain,(
  ( ! [X8,X9] : (sP3(X8,X9,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | p(s(bool,X9))) ) | $spl320),
  inference(cnf_transformation,[],[f25897_D])).
fof(f25897_D,plain,(
  ( ! [X8,X9] : (sP3(X8,X9,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | p(s(bool,X9))) ) <=> ~$spl320),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl320])])).
fof(f25888,plain,(
  ( ! [X8,X9] : (sP3(X8,X9,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | p(s(bool,X9)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f2178,f393])).
fof(f2178,plain,(
  ( ! [X47,X48] : (sP2(X47,X48,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(resolution,[],[f1625,f395])).
fof(f25884,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f2178,f394])).
fof(f25883,plain,(
  ( ! [X8,X9] : (sP1(X8,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9) | ~p(s(bool,X9))) ) | $spl318),
  inference(cnf_transformation,[],[f25883_D])).
fof(f25883_D,plain,(
  ( ! [X8,X9] : (sP1(X8,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9) | ~p(s(bool,X9))) ) <=> ~$spl318),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl318])])).
fof(f25874,plain,(
  ( ! [X8,X9] : (sP1(X8,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~p(s(bool,X9))) )),
  inference(resolution,[],[f2165,f376])).
fof(f2165,plain,(
  ( ! [X21,X22] : (sP0(X21,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X22)) )),
  inference(resolution,[],[f1625,f379])).
fof(f25870,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~p(s(bool,X1))) )),
  inference(resolution,[],[f2165,f377])).
fof(f22640,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | $spl240),
  inference(cnf_transformation,[],[f22640_D])).
fof(f22640_D,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) <=> ~$spl240),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl240])])).
fof(f25745,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | $spl42),
  inference(resolution,[],[f4751,f355])).
fof(f4371,plain,(
  s(bool,'F_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | $spl35),
  inference(cnf_transformation,[],[f4371_D])).
fof(f4371_D,plain,(
  s(bool,'F_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) <=> ~$spl35),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl35])])).
fof(f25695,plain,(
  ( ! [X12] : (p(s(bool,X12)) | s(bool,'F_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(resolution,[],[f1688,f3286])).
fof(f1688,plain,(
  ( ! [X121,X120] : (sP6('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X120,X121)) )),
  inference(resolution,[],[f884,f439])).
fof(f24479,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | $spl262),
  inference(cnf_transformation,[],[f24479_D])).
fof(f24479_D,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) <=> ~$spl262),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl262])])).
fof(f25691,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X4)) | p(s(bool,X5))) )),
  inference(resolution,[],[f1688,f436])).
fof(f25703,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl316),
  inference(cnf_transformation,[],[f25703_D])).
fof(f25703_D,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) <=> ~$spl316),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl316])])).
fof(f25690,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl10),
  inference(resolution,[],[f1688,f19474])).
fof(f25702,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,X1) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl314),
  inference(cnf_transformation,[],[f25702_D])).
fof(f25702_D,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,X1) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) <=> ~$spl314),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl314])])).
fof(f25689,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'F_0') != s(bool,X1) | s(bool,X1) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl10),
  inference(resolution,[],[f1688,f19603])).
fof(f25687,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9)) ) | $spl312),
  inference(cnf_transformation,[],[f25687_D])).
fof(f25687_D,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9)) ) <=> ~$spl312),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl312])])).
fof(f25686,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9) | p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(subsumption_resolution,[],[f25685,f1682])).
fof(f25685,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9) | p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP8(X8,X9,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(subsumption_resolution,[],[f25676,f435])).
fof(f25676,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9) | p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP8(X8,X9,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~sP7('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9,X8)) )),
  inference(resolution,[],[f1687,f426])).
fof(f1687,plain,(
  ( ! [X118,X119] : (sP6(X118,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X119)) )),
  inference(resolution,[],[f884,f438])).
fof(f25674,plain,(
  ( ! [X4,X5] : (p(s(bool,X4)) | p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X5))) )),
  inference(resolution,[],[f1687,f436])).
fof(f25684,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl310),
  inference(cnf_transformation,[],[f25684_D])).
fof(f25684_D,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) <=> ~$spl310),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl310])])).
fof(f25673,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X2)) | s(bool,X3) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl10),
  inference(resolution,[],[f1687,f19474])).
fof(f25672,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'F_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | s(bool,X1) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl10),
  inference(resolution,[],[f1687,f19603])).
fof(f25661,plain,(
  ( ! [X12] : (p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | s(bool,'F_0') != s(bool,X12)) )),
  inference(resolution,[],[f1686,f3286])).
fof(f1686,plain,(
  ( ! [X116,X117] : (sP6(X116,X117,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(resolution,[],[f884,f437])).
fof(f25671,plain,(
  ( ! [X10,X11] : (sP9('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X10,X11)) ) | $spl308),
  inference(cnf_transformation,[],[f25671_D])).
fof(f25671_D,plain,(
  ( ! [X10,X11] : (sP9('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X10,X11)) ) <=> ~$spl308),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl308])])).
fof(f25670,plain,(
  ( ! [X10,X11] : (sP9('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X10,X11) | p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(subsumption_resolution,[],[f25669,f1685])).
fof(f1685,plain,(
  ( ! [X114,X115] : (sP7(X114,X115,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(resolution,[],[f884,f433])).
fof(f25669,plain,(
  ( ! [X10,X11] : (sP9('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X10,X11) | p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP7(X10,X11,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(subsumption_resolution,[],[f25660,f431])).
fof(f25660,plain,(
  ( ! [X10,X11] : (sP9('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X10,X11) | p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP8('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X11,X10) | ~sP7(X10,X11,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(resolution,[],[f1686,f427])).
fof(f25657,plain,(
  ( ! [X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f1686,f436])).
fof(f25656,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | s(bool,X2) != s(bool,X3)) ) | $spl10),
  inference(resolution,[],[f1686,f19474])).
fof(f25655,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | s(bool,'F_0') != s(bool,X0) | s(bool,X0) != s(bool,X1)) ) | $spl10),
  inference(resolution,[],[f1686,f19603])).
fof(f25654,plain,(
  ( ! [X10,X11] : (sP9('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X10,X11) | ~sP6(X11,X10,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl306),
  inference(cnf_transformation,[],[f25654_D])).
fof(f25654_D,plain,(
  ( ! [X10,X11] : (sP9('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X10,X11) | ~sP6(X11,X10,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) <=> ~$spl306),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl306])])).
fof(f25653,plain,(
  ( ! [X10,X11] : (sP9('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X10,X11) | p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP6(X11,X10,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(subsumption_resolution,[],[f25645,f431])).
fof(f25645,plain,(
  ( ! [X10,X11] : (sP9('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X10,X11) | p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP8('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X11,X10) | ~sP6(X11,X10,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(resolution,[],[f1685,f427])).
fof(f25642,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | ~p(s(bool,X5)) | p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f1685,f432])).
fof(f25641,plain,(
  ( ! [X2,X3] : (~p(s(bool,X2)) | p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | s(bool,X2) != s(bool,X3)) ) | $spl10),
  inference(resolution,[],[f1685,f19313])).
fof(f25640,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | s(bool,'T_0') != s(bool,X0) | s(bool,X0) != s(bool,X1)) ) | $spl10),
  inference(resolution,[],[f1685,f19575])).
fof(f25628,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | s(bool,'T_0') != s(bool,X19)) ) | $spl10),
  inference(resolution,[],[f1682,f19799])).
fof(f25627,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~p(s(bool,X16)) | ~p(s(bool,X17))) ) | $spl10),
  inference(resolution,[],[f1682,f19796])).
fof(f25637,plain,(
  ( ! [X2,X3] : (sP9(X2,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X3) | ~sP6(X3,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X2)) ) | $spl300),
  inference(cnf_transformation,[],[f25637_D])).
fof(f25637_D,plain,(
  ( ! [X2,X3] : (sP9(X2,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X3) | ~sP6(X3,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X2)) ) <=> ~$spl300),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl300])])).
fof(f25626,plain,(
  ( ! [X14,X15] : (p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | sP9(X14,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X15) | ~sP6(X15,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X14)) ) | $spl10),
  inference(resolution,[],[f1682,f19385])).
fof(f25639,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X11) | ~p(s(bool,X11))) ) | $spl304),
  inference(cnf_transformation,[],[f25639_D])).
fof(f25639_D,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X11) | ~p(s(bool,X11))) ) <=> ~$spl304),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl304])])).
fof(f25624,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X11) | ~p(s(bool,X11)) | s(bool,'F_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl10),
  inference(resolution,[],[f1682,f19168])).
fof(f25638,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9) | s(bool,'T_0') != s(bool,X9)) ) | $spl302),
  inference(cnf_transformation,[],[f25638_D])).
fof(f25638_D,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9) | s(bool,'T_0') != s(bool,X9)) ) <=> ~$spl302),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl302])])).
fof(f25623,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9) | p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | s(bool,'T_0') != s(bool,X9)) ) | $spl10),
  inference(resolution,[],[f1682,f19159])).
fof(f25622,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f1682,f428])).
fof(f25636,plain,(
  ( ! [X2,X3] : (sP9(X2,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X3) | p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP6(X3,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X2)) )),
  inference(subsumption_resolution,[],[f25620,f435])).
fof(f25620,plain,(
  ( ! [X2,X3] : (sP9(X2,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X3) | p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP7('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X3,X2) | ~sP6(X3,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X2)) )),
  inference(resolution,[],[f1682,f426])).
fof(f25423,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))),
  inference(subsumption_resolution,[],[f25081,f588])).
fof(f25081,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))),
  inference(superposition,[],[f355,f680])).
fof(f25059,plain,(
  ( ! [X4,X5] : (sP5('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X4,X5)) ) | $spl298),
  inference(cnf_transformation,[],[f25059_D])).
fof(f25059_D,plain,(
  ( ! [X4,X5] : (sP5('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X4,X5)) ) <=> ~$spl298),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl298])])).
fof(f25051,plain,(
  ( ! [X8,X9] : (sP5('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X8,X9) | p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f1669,f587])).
fof(f1669,plain,(
  ( ! [X83,X82] : (sP4(X82,X83,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(resolution,[],[f884,f412])).
fof(f25060,plain,(
  ( ! [X6,X7] : (sP5('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X6,X7) | p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(subsumption_resolution,[],[f25050,f415])).
fof(f25050,plain,(
  ( ! [X6,X7] : (sP5('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X6,X7) | p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~p(s(bool,X7))) )),
  inference(resolution,[],[f1669,f409])).
fof(f25058,plain,(
  ( ! [X4,X5] : (sP5('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X4,X5) | p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(subsumption_resolution,[],[f25049,f416])).
fof(f25049,plain,(
  ( ! [X4,X5] : (sP5('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X4,X5) | ~p(s(bool,X4)) | p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f1669,f408])).
fof(f25047,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f1669,f411])).
fof(f25046,plain,(
  ( ! [X8,X9] : (sP3(X8,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9) | p(s(bool,X9))) ) | $spl296),
  inference(cnf_transformation,[],[f25046_D])).
fof(f25046_D,plain,(
  ( ! [X8,X9] : (sP3(X8,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9) | p(s(bool,X9))) ) <=> ~$spl296),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl296])])).
fof(f25037,plain,(
  ( ! [X8,X9] : (sP3(X8,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9) | p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X9))) )),
  inference(resolution,[],[f1653,f393])).
fof(f1653,plain,(
  ( ! [X50,X51] : (sP2(X50,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X51)) )),
  inference(resolution,[],[f884,f396])).
fof(f25033,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X1))) )),
  inference(resolution,[],[f1653,f394])).
fof(f25032,plain,(
  ( ! [X8,X9] : (sP3(X8,X9,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | p(s(bool,X9))) ) | $spl294),
  inference(cnf_transformation,[],[f25032_D])).
fof(f25032_D,plain,(
  ( ! [X8,X9] : (sP3(X8,X9,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | p(s(bool,X9))) ) <=> ~$spl294),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl294])])).
fof(f25023,plain,(
  ( ! [X8,X9] : (sP3(X8,X9,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | p(s(bool,X9)) | p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f1652,f393])).
fof(f1652,plain,(
  ( ! [X48,X49] : (sP2(X48,X49,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(resolution,[],[f884,f395])).
fof(f25019,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f1652,f394])).
fof(f25018,plain,(
  ( ! [X8,X9] : (sP1(X8,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9) | ~p(s(bool,X9))) ) | $spl292),
  inference(cnf_transformation,[],[f25018_D])).
fof(f25018_D,plain,(
  ( ! [X8,X9] : (sP1(X8,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9) | ~p(s(bool,X9))) ) <=> ~$spl292),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl292])])).
fof(f25009,plain,(
  ( ! [X8,X9] : (sP1(X8,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9) | p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~p(s(bool,X9))) )),
  inference(resolution,[],[f1639,f376])).
fof(f1639,plain,(
  ( ! [X23,X22] : (sP0(X22,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X23)) )),
  inference(resolution,[],[f884,f379])).
fof(f25005,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~p(s(bool,X1))) )),
  inference(resolution,[],[f1639,f377])).
fof(f3621,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | $spl29),
  inference(cnf_transformation,[],[f3621_D])).
fof(f3621_D,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f24994,plain,(
  ( ! [X12] : (p(s(bool,X12)) | s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(resolution,[],[f1543,f3286])).
fof(f1543,plain,(
  ( ! [X112,X113] : (sP6('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X112,X113)) )),
  inference(resolution,[],[f882,f439])).
fof(f24457,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | $spl260),
  inference(cnf_transformation,[],[f24457_D])).
fof(f24457_D,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) <=> ~$spl260),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl260])])).
fof(f24990,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X4)) | p(s(bool,X5))) )),
  inference(resolution,[],[f1543,f436])).
fof(f25002,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl290),
  inference(cnf_transformation,[],[f25002_D])).
fof(f25002_D,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) <=> ~$spl290),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl290])])).
fof(f24989,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl10),
  inference(resolution,[],[f1543,f19474])).
fof(f25001,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl288),
  inference(cnf_transformation,[],[f25001_D])).
fof(f25001_D,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) <=> ~$spl288),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl288])])).
fof(f24988,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'F_0') != s(bool,X1) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl10),
  inference(resolution,[],[f1543,f19603])).
fof(f24984,plain,(
  ( ! [X3] : (p(s(bool,X3)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(duplicate_literal_removal,[],[f24974])).
fof(f24974,plain,(
  ( ! [X3] : (p(s(bool,X3)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | p(s(bool,X3))) )),
  inference(resolution,[],[f662,f458])).
fof(f24985,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(duplicate_literal_removal,[],[f24973])).
fof(f24973,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,X2)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(resolution,[],[f662,f467])).
fof(f24969,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9)) ) | $spl286),
  inference(cnf_transformation,[],[f24969_D])).
fof(f24969_D,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9)) ) <=> ~$spl286),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl286])])).
fof(f24968,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(subsumption_resolution,[],[f24967,f1537])).
fof(f24967,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP8(X8,X9,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(subsumption_resolution,[],[f24958,f435])).
fof(f24958,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP8(X8,X9,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~sP7('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9,X8)) )),
  inference(resolution,[],[f1542,f426])).
fof(f1542,plain,(
  ( ! [X111,X110] : (sP6(X110,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X111)) )),
  inference(resolution,[],[f882,f438])).
fof(f24956,plain,(
  ( ! [X4,X5] : (p(s(bool,X4)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X5))) )),
  inference(resolution,[],[f1542,f436])).
fof(f24966,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl284),
  inference(cnf_transformation,[],[f24966_D])).
fof(f24966_D,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) <=> ~$spl284),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl284])])).
fof(f24955,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X2)) | s(bool,X3) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl10),
  inference(resolution,[],[f1542,f19474])).
fof(f24954,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl10),
  inference(resolution,[],[f1542,f19603])).
fof(f24943,plain,(
  ( ! [X12] : (p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | s(bool,'F_0') != s(bool,X12)) )),
  inference(resolution,[],[f1541,f3286])).
fof(f1541,plain,(
  ( ! [X109,X108] : (sP6(X108,X109,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(resolution,[],[f882,f437])).
fof(f24953,plain,(
  ( ! [X10,X11] : (sP9('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X10,X11)) ) | $spl282),
  inference(cnf_transformation,[],[f24953_D])).
fof(f24953_D,plain,(
  ( ! [X10,X11] : (sP9('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X10,X11)) ) <=> ~$spl282),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl282])])).
fof(f24952,plain,(
  ( ! [X10,X11] : (sP9('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X10,X11) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(subsumption_resolution,[],[f24951,f1540])).
fof(f1540,plain,(
  ( ! [X107,X106] : (sP7(X106,X107,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(resolution,[],[f882,f433])).
fof(f24951,plain,(
  ( ! [X10,X11] : (sP9('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X10,X11) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP7(X10,X11,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(subsumption_resolution,[],[f24942,f431])).
fof(f24942,plain,(
  ( ! [X10,X11] : (sP9('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X10,X11) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP8('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X11,X10) | ~sP7(X10,X11,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(resolution,[],[f1541,f427])).
fof(f24939,plain,(
  ( ! [X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f1541,f436])).
fof(f24938,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | s(bool,X2) != s(bool,X3)) ) | $spl10),
  inference(resolution,[],[f1541,f19474])).
fof(f24937,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | s(bool,'F_0') != s(bool,X0) | s(bool,X0) != s(bool,X1)) ) | $spl10),
  inference(resolution,[],[f1541,f19603])).
fof(f24936,plain,(
  ( ! [X10,X11] : (sP9('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X10,X11) | ~sP6(X11,X10,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl280),
  inference(cnf_transformation,[],[f24936_D])).
fof(f24936_D,plain,(
  ( ! [X10,X11] : (sP9('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X10,X11) | ~sP6(X11,X10,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) <=> ~$spl280),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl280])])).
fof(f24935,plain,(
  ( ! [X10,X11] : (sP9('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X10,X11) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP6(X11,X10,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(subsumption_resolution,[],[f24927,f431])).
fof(f24927,plain,(
  ( ! [X10,X11] : (sP9('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X10,X11) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP8('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X11,X10) | ~sP6(X11,X10,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(resolution,[],[f1540,f427])).
fof(f24924,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | ~p(s(bool,X5)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f1540,f432])).
fof(f24923,plain,(
  ( ! [X2,X3] : (~p(s(bool,X2)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | s(bool,X2) != s(bool,X3)) ) | $spl10),
  inference(resolution,[],[f1540,f19313])).
fof(f24922,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | s(bool,'T_0') != s(bool,X0) | s(bool,X0) != s(bool,X1)) ) | $spl10),
  inference(resolution,[],[f1540,f19575])).
fof(f24910,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | s(bool,'T_0') != s(bool,X19)) ) | $spl10),
  inference(resolution,[],[f1537,f19799])).
fof(f24909,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | ~p(s(bool,X16)) | ~p(s(bool,X17))) ) | $spl10),
  inference(resolution,[],[f1537,f19796])).
fof(f24919,plain,(
  ( ! [X2,X3] : (sP9(X2,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X3) | ~sP6(X3,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X2)) ) | $spl274),
  inference(cnf_transformation,[],[f24919_D])).
fof(f24919_D,plain,(
  ( ! [X2,X3] : (sP9(X2,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X3) | ~sP6(X3,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X2)) ) <=> ~$spl274),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl274])])).
fof(f24908,plain,(
  ( ! [X14,X15] : (p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | sP9(X14,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X15) | ~sP6(X15,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X14)) ) | $spl10),
  inference(resolution,[],[f1537,f19385])).
fof(f24921,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X11) | ~p(s(bool,X11))) ) | $spl278),
  inference(cnf_transformation,[],[f24921_D])).
fof(f24921_D,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X11) | ~p(s(bool,X11))) ) <=> ~$spl278),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl278])])).
fof(f24906,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X11) | ~p(s(bool,X11)) | s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) ) | $spl10),
  inference(resolution,[],[f1537,f19168])).
fof(f24920,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9) | s(bool,'T_0') != s(bool,X9)) ) | $spl276),
  inference(cnf_transformation,[],[f24920_D])).
fof(f24920_D,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9) | s(bool,'T_0') != s(bool,X9)) ) <=> ~$spl276),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl276])])).
fof(f24905,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | s(bool,'T_0') != s(bool,X9)) ) | $spl10),
  inference(resolution,[],[f1537,f19159])).
fof(f24904,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f1537,f428])).
fof(f24918,plain,(
  ( ! [X2,X3] : (sP9(X2,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X3) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP6(X3,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X2)) )),
  inference(subsumption_resolution,[],[f24902,f435])).
fof(f24902,plain,(
  ( ! [X2,X3] : (sP9(X2,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X3) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~sP7('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X3,X2) | ~sP6(X3,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X2)) )),
  inference(resolution,[],[f1537,f426])).
fof(f24899,plain,(
  ( ! [X4,X5] : (sP5('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X4,X5)) ) | $spl272),
  inference(cnf_transformation,[],[f24899_D])).
fof(f24899_D,plain,(
  ( ! [X4,X5] : (sP5('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X4,X5)) ) <=> ~$spl272),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl272])])).
fof(f24891,plain,(
  ( ! [X8,X9] : (sP5('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X8,X9) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f1524,f587])).
fof(f1524,plain,(
  ( ! [X74,X75] : (sP4(X74,X75,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(resolution,[],[f882,f412])).
fof(f24900,plain,(
  ( ! [X6,X7] : (sP5('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X6,X7) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(subsumption_resolution,[],[f24890,f415])).
fof(f24890,plain,(
  ( ! [X6,X7] : (sP5('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X6,X7) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~p(s(bool,X7))) )),
  inference(resolution,[],[f1524,f409])).
fof(f24898,plain,(
  ( ! [X4,X5] : (sP5('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X4,X5) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(subsumption_resolution,[],[f24889,f416])).
fof(f24889,plain,(
  ( ! [X4,X5] : (sP5('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X4,X5) | ~p(s(bool,X4)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f1524,f408])).
fof(f24887,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f1524,f411])).
fof(f24883,plain,(
  ( ! [X3] : (p(s(bool,X3)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(duplicate_literal_removal,[],[f24873])).
fof(f24873,plain,(
  ( ! [X3] : (p(s(bool,X3)) | p(s(bool,X3)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(resolution,[],[f661,f458])).
fof(f24884,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(duplicate_literal_removal,[],[f24872])).
fof(f24872,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | p(s(bool,X2))) )),
  inference(resolution,[],[f661,f467])).
fof(f24869,plain,(
  ( ! [X8,X9] : (sP3(X8,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9) | p(s(bool,X9))) ) | $spl270),
  inference(cnf_transformation,[],[f24869_D])).
fof(f24869_D,plain,(
  ( ! [X8,X9] : (sP3(X8,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9) | p(s(bool,X9))) ) <=> ~$spl270),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl270])])).
fof(f24860,plain,(
  ( ! [X8,X9] : (sP3(X8,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X9))) )),
  inference(resolution,[],[f1508,f393])).
fof(f1508,plain,(
  ( ! [X43,X42] : (sP2(X42,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X43)) )),
  inference(resolution,[],[f882,f396])).
fof(f24856,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X1))) )),
  inference(resolution,[],[f1508,f394])).
fof(f24855,plain,(
  ( ! [X8,X9] : (sP3(X8,X9,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | p(s(bool,X9))) ) | $spl268),
  inference(cnf_transformation,[],[f24855_D])).
fof(f24855_D,plain,(
  ( ! [X8,X9] : (sP3(X8,X9,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | p(s(bool,X9))) ) <=> ~$spl268),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl268])])).
fof(f24846,plain,(
  ( ! [X8,X9] : (sP3(X8,X9,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))) | p(s(bool,X9)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f1507,f393])).
fof(f1507,plain,(
  ( ! [X41,X40] : (sP2(X40,X41,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(resolution,[],[f882,f395])).
fof(f24842,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f1507,f394])).
fof(f24841,plain,(
  ( ! [X8,X9] : (sP1(X8,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9) | ~p(s(bool,X9))) ) | $spl266),
  inference(cnf_transformation,[],[f24841_D])).
fof(f24841_D,plain,(
  ( ! [X8,X9] : (sP1(X8,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9) | ~p(s(bool,X9))) ) <=> ~$spl266),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl266])])).
fof(f24832,plain,(
  ( ! [X8,X9] : (sP1(X8,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X9) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~p(s(bool,X9))) )),
  inference(resolution,[],[f1494,f376])).
fof(f1494,plain,(
  ( ! [X14,X15] : (sP0(X14,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X15)) )),
  inference(resolution,[],[f882,f379])).
fof(f24828,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | ~p(s(bool,X1))) )),
  inference(resolution,[],[f1494,f377])).
fof(f3619,plain,(
  s(bool,'F_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) | $spl27),
  inference(cnf_transformation,[],[f3619_D])).
fof(f3619_D,plain,(
  s(bool,'F_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f24819,plain,(
  ( ! [X12] : (p(s(bool,X12)) | s(bool,'F_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) | $spl2),
  inference(resolution,[],[f1168,f3286])).
fof(f21579,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) | $spl186),
  inference(cnf_transformation,[],[f21579_D])).
fof(f21579_D,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f24814,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) | ($spl2 | $spl10)),
  inference(resolution,[],[f1168,f19474])).
fof(f21578,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,X1) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) | $spl184),
  inference(cnf_transformation,[],[f21578_D])).
fof(f21578_D,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,X1) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f24813,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'F_0') != s(bool,X1) | s(bool,X1) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) | ($spl2 | $spl10)),
  inference(resolution,[],[f1168,f19603])).
fof(f24798,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'F_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) | s(bool,X1) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) | ($spl2 | $spl10)),
  inference(resolution,[],[f1167,f19603])).
fof(f24795,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(duplicate_literal_removal,[],[f24786])).
fof(f24786,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | ~p(s(bool,X2))) )),
  inference(resolution,[],[f660,f460])).
fof(f24796,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(duplicate_literal_removal,[],[f24785])).
fof(f24785,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | ~p(s(bool,X1))) )),
  inference(resolution,[],[f660,f465])).
fof(f24745,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) | s(bool,'T_0') != s(bool,X19)) ) | ($spl2 | $spl10)),
  inference(resolution,[],[f1162,f19799])).
fof(f24744,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) | ~p(s(bool,X16)) | ~p(s(bool,X17))) ) | ($spl2 | $spl10)),
  inference(resolution,[],[f1162,f19796])).
fof(f21491,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X11) | ~p(s(bool,X11))) ) | $spl174),
  inference(cnf_transformation,[],[f21491_D])).
fof(f21491_D,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X11) | ~p(s(bool,X11))) ) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f24741,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X11) | ~p(s(bool,X11)) | s(bool,'F_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) | ($spl2 | $spl10)),
  inference(resolution,[],[f1162,f19168])).
fof(f24707,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(duplicate_literal_removal,[],[f24699])).
fof(f24699,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | ~p(s(bool,X3))) )),
  inference(resolution,[],[f659,f463])).
fof(f24708,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(duplicate_literal_removal,[],[f24696])).
fof(f24696,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(resolution,[],[f659,f464])).
fof(f24661,plain,(
  ( ! [X12] : (p(s(bool,X12)) | s(bool,'F_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) )),
  inference(resolution,[],[f656,f3286])).
fof(f24657,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | p(s(bool,X4)) | p(s(bool,X5))) )),
  inference(resolution,[],[f656,f436])).
fof(f24656,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) | $spl10),
  inference(resolution,[],[f656,f19474])).
fof(f24655,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'F_0') != s(bool,X1) | s(bool,X1) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) | $spl10),
  inference(resolution,[],[f656,f19603])).
fof(f24653,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X9) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(subsumption_resolution,[],[f24652,f650])).
fof(f24652,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X9) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | ~sP8(X8,X9,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) )),
  inference(subsumption_resolution,[],[f24644,f435])).
fof(f24644,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X9) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | ~sP8(X8,X9,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) | ~sP7('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X9,X8)) )),
  inference(resolution,[],[f655,f426])).
fof(f24642,plain,(
  ( ! [X4,X5] : (p(s(bool,X4)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | p(s(bool,X5))) )),
  inference(resolution,[],[f655,f436])).
fof(f24641,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | p(s(bool,X2)) | s(bool,X3) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) | $spl10),
  inference(resolution,[],[f655,f19474])).
fof(f24640,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'F_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) | s(bool,X1) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) | $spl10),
  inference(resolution,[],[f655,f19603])).
fof(f24630,plain,(
  ( ! [X12] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | s(bool,'F_0') != s(bool,X12)) )),
  inference(resolution,[],[f654,f3286])).
fof(f24639,plain,(
  ( ! [X10,X11] : (sP9('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X10,X11) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(subsumption_resolution,[],[f24638,f653])).
fof(f24638,plain,(
  ( ! [X10,X11] : (sP9('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X10,X11) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | ~sP7(X10,X11,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) )),
  inference(subsumption_resolution,[],[f24629,f431])).
fof(f24629,plain,(
  ( ! [X10,X11] : (sP9('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X10,X11) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | ~sP8('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X11,X10) | ~sP7(X10,X11,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) )),
  inference(resolution,[],[f654,f427])).
fof(f24626,plain,(
  ( ! [X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(resolution,[],[f654,f436])).
fof(f24625,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | s(bool,X2) != s(bool,X3)) ) | $spl10),
  inference(resolution,[],[f654,f19474])).
fof(f24624,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | s(bool,'F_0') != s(bool,X0) | s(bool,X0) != s(bool,X1)) ) | $spl10),
  inference(resolution,[],[f654,f19603])).
fof(f24623,plain,(
  ( ! [X10,X11] : (sP9('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X10,X11) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | ~sP6(X11,X10,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) )),
  inference(subsumption_resolution,[],[f24615,f431])).
fof(f24615,plain,(
  ( ! [X10,X11] : (sP9('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X10,X11) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | ~sP8('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X11,X10) | ~sP6(X11,X10,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) )),
  inference(resolution,[],[f653,f427])).
fof(f24612,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | ~p(s(bool,X5)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(resolution,[],[f653,f432])).
fof(f24611,plain,(
  ( ! [X2,X3] : (~p(s(bool,X2)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | s(bool,X2) != s(bool,X3)) ) | $spl10),
  inference(resolution,[],[f653,f19313])).
fof(f24610,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | s(bool,'T_0') != s(bool,X0) | s(bool,X0) != s(bool,X1)) ) | $spl10),
  inference(resolution,[],[f653,f19575])).
fof(f24601,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) | s(bool,'T_0') != s(bool,X19)) ) | $spl10),
  inference(resolution,[],[f650,f19799])).
fof(f24600,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) | ~p(s(bool,X16)) | ~p(s(bool,X17))) ) | $spl10),
  inference(resolution,[],[f650,f19796])).
fof(f24599,plain,(
  ( ! [X14,X15] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | sP9(X14,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X15) | ~sP6(X15,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X14)) ) | $spl10),
  inference(resolution,[],[f650,f19385])).
fof(f24597,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X11) | ~p(s(bool,X11)) | s(bool,'F_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) | $spl10),
  inference(resolution,[],[f650,f19168])).
fof(f24596,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X9) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | s(bool,'T_0') != s(bool,X9)) ) | $spl10),
  inference(resolution,[],[f650,f19159])).
fof(f24595,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(resolution,[],[f650,f428])).
fof(f24609,plain,(
  ( ! [X2,X3] : (sP9(X2,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X3) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | ~sP6(X3,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X2)) )),
  inference(subsumption_resolution,[],[f24593,f435])).
fof(f24593,plain,(
  ( ! [X2,X3] : (sP9(X2,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X3) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | ~sP7('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X3,X2) | ~sP6(X3,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X2)) )),
  inference(resolution,[],[f650,f426])).
fof(f24583,plain,(
  ( ! [X8,X9] : (sP5('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X8,X9) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(resolution,[],[f637,f587])).
fof(f24591,plain,(
  ( ! [X6,X7] : (sP5('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X6,X7) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(subsumption_resolution,[],[f24582,f415])).
fof(f24582,plain,(
  ( ! [X6,X7] : (sP5('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X6,X7) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | ~p(s(bool,X7))) )),
  inference(resolution,[],[f637,f409])).
fof(f24590,plain,(
  ( ! [X4,X5] : (sP5('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X4,X5) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(subsumption_resolution,[],[f24581,f416])).
fof(f24581,plain,(
  ( ! [X4,X5] : (sP5('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X4,X5) | ~p(s(bool,X4)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(resolution,[],[f637,f408])).
fof(f24579,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(resolution,[],[f637,f411])).
fof(f24560,plain,(
  ( ! [X8,X9] : (sP3(X8,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X9) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | p(s(bool,X9))) )),
  inference(resolution,[],[f621,f393])).
fof(f24556,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | p(s(bool,X1))) )),
  inference(resolution,[],[f621,f394])).
fof(f24525,plain,(
  ( ! [X1] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X1))) )),
  inference(subsumption_resolution,[],[f24517,f538])).
fof(f24517,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X1))) )),
  inference(resolution,[],[f2226,f474])).
fof(f2226,plain,(
  ( ! [X134] : (sP11('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X134)) )),
  inference(resolution,[],[f1625,f476])).
fof(f24524,plain,(
  ( ! [X0] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f24516,f538])).
fof(f24516,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X0))) )),
  inference(resolution,[],[f2226,f478])).
fof(f24515,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(subsumption_resolution,[],[f24505,f538])).
fof(f24505,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f2225,f474])).
fof(f2225,plain,(
  ( ! [X133] : (sP11(X133,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(resolution,[],[f1625,f475])).
fof(f24512,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(subsumption_resolution,[],[f24504,f538])).
fof(f24504,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f2225,f478])).
fof(f24495,plain,(
  ( ! [X8,X9] : (sP3(X8,X9,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) | p(s(bool,X9)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(resolution,[],[f620,f393])).
fof(f24491,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(resolution,[],[f620,f394])).
fof(f24490,plain,(
  ( ! [X1] : (p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X1))) )),
  inference(subsumption_resolution,[],[f24482,f538])).
fof(f24482,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X1))) )),
  inference(resolution,[],[f1700,f474])).
fof(f1700,plain,(
  ( ! [X135] : (sP11('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X135)) )),
  inference(resolution,[],[f884,f476])).
fof(f24489,plain,(
  ( ! [X0] : (p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f24481,f538])).
fof(f24481,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X0))) )),
  inference(resolution,[],[f1700,f478])).
fof(f24480,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(subsumption_resolution,[],[f24470,f538])).
fof(f24470,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f1699,f474])).
fof(f24477,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(subsumption_resolution,[],[f24469,f538])).
fof(f24469,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f1699,f478])).
fof(f24468,plain,(
  ( ! [X1] : (p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X1))) )),
  inference(subsumption_resolution,[],[f24460,f538])).
fof(f24460,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X1))) )),
  inference(resolution,[],[f1555,f474])).
fof(f1555,plain,(
  ( ! [X127] : (sP11('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)),X127)) )),
  inference(resolution,[],[f882,f476])).
fof(f24467,plain,(
  ( ! [X0] : (p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f24459,f538])).
fof(f24459,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) | p(s(bool,X0))) )),
  inference(resolution,[],[f1555,f478])).
fof(f24458,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(subsumption_resolution,[],[f24448,f538])).
fof(f24448,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f1554,f474])).
fof(f1554,plain,(
  ( ! [X126] : (sP11(X126,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(resolution,[],[f882,f475])).
fof(f24455,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(subsumption_resolution,[],[f24447,f538])).
fof(f24447,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15))))) )),
  inference(resolution,[],[f1554,f478])).
fof(f24428,plain,(
  ( ! [X8,X9] : (sP1(X8,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X9) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | ~p(s(bool,X9))) )),
  inference(resolution,[],[f607,f376])).
fof(f24424,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | ~p(s(bool,X1))) )),
  inference(resolution,[],[f607,f377])).
fof(f24413,plain,(
  ( ! [X1] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | p(s(bool,X1))) )),
  inference(subsumption_resolution,[],[f24405,f538])).
fof(f24405,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | p(s(bool,X1))) )),
  inference(resolution,[],[f668,f474])).
fof(f24412,plain,(
  ( ! [X0] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f24404,f538])).
fof(f24404,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | p(s(bool,X0))) )),
  inference(resolution,[],[f668,f478])).
fof(f24403,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(subsumption_resolution,[],[f24395,f538])).
fof(f24395,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(resolution,[],[f667,f474])).
fof(f24402,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(subsumption_resolution,[],[f24394,f538])).
fof(f24394,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(resolution,[],[f667,f478])).
fof(f24391,plain,(
  $false | ($spl7 | $spl256)),
  inference(subsumption_resolution,[],[f24284,f24183])).
fof(f24183,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl7 | $spl256)),
  inference(subsumption_resolution,[],[f24182,f12719])).
fof(f24182,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))) ) | $spl256),
  inference(subsumption_resolution,[],[f24174,f538])).
fof(f24174,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))) ) | $spl256),
  inference(resolution,[],[f23842,f478])).
fof(f23842,plain,(
  ( ! [X133] : (sP11(X133,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | $spl256),
  inference(backward_demodulation,[],[f23821,f667])).
fof(f24284,plain,(
  ~p(s(bool,'T_0')) | ($spl7 | $spl256)),
  inference(resolution,[],[f24183,f575])).
fof(f24390,plain,(
  $false | ($spl7 | $spl256)),
  inference(subsumption_resolution,[],[f24283,f24183])).
fof(f24283,plain,(
  ~p(s(bool,'T_0')) | ($spl7 | $spl256)),
  inference(resolution,[],[f24183,f573])).
fof(f24389,plain,(
  $false | ($spl7 | $spl256)),
  inference(subsumption_resolution,[],[f24280,f24183])).
fof(f24280,plain,(
  ( ! [X224] : (~p(s(bool,X224))) ) | ($spl7 | $spl256)),
  inference(resolution,[],[f24183,f567])).
fof(f24388,plain,(
  $false | ($spl7 | $spl256)),
  inference(subsumption_resolution,[],[f24276,f24183])).
fof(f24276,plain,(
  ( ! [X220] : (~p(s(bool,X220))) ) | ($spl7 | $spl256)),
  inference(resolution,[],[f24183,f520])).
fof(f24282,plain,(
  $false | ($spl7 | $spl256)),
  inference(resolution,[],[f24183,f571])).
fof(f24324,plain,(
  $false | ($spl7 | $spl256)),
  inference(resolution,[],[f24183,f12719])).
fof(f23708,plain,(
  p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15))),s(sK12,sK14)))) | s(sK12,sK14) = s(sK12,sK15) | $spl240),
  inference(resolution,[],[f22640,f578])).
fof(f23697,plain,(
  ( ! [X12] : (p(s(bool,X12)) | s(bool,'F_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | $spl42),
  inference(resolution,[],[f22584,f3286])).
fof(f23692,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | ($spl10 | $spl42)),
  inference(resolution,[],[f22584,f19474])).
fof(f23691,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'F_0') != s(bool,X1) | s(bool,X1) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | ($spl10 | $spl42)),
  inference(resolution,[],[f22584,f19603])).
fof(f23677,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'F_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | s(bool,X1) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | ($spl10 | $spl42)),
  inference(resolution,[],[f22583,f19603])).
fof(f23560,plain,(
  ( ! [X325,X324,X326] : (p(s(bool,X324)) | ~p(s(bool,X325)) | sP9(X326,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X325) | p(s(bool,X326))) ) | $spl10),
  inference(subsumption_resolution,[],[f23307,f553])).
fof(f23307,plain,(
  ( ! [X325,X324,X326] : (s(bool,'F_0') != s(bool,X324) | p(s(bool,X324)) | ~p(s(bool,X325)) | sP9(X326,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X325) | p(s(bool,X326))) ) | $spl10),
  inference(superposition,[],[f7606,f806])).
fof(f23046,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | s(bool,'T_0') != s(bool,X19)) ) | ($spl10 | $spl42)),
  inference(resolution,[],[f22578,f19799])).
fof(f23045,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | ~p(s(bool,X16)) | ~p(s(bool,X17))) ) | ($spl10 | $spl42)),
  inference(resolution,[],[f22578,f19796])).
fof(f23042,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X11) | ~p(s(bool,X11)) | s(bool,'F_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | ($spl10 | $spl42)),
  inference(resolution,[],[f22578,f19168])).
fof(f22683,plain,(
  ( ! [X2] : (p(s(bool,X2)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))) )),
  inference(duplicate_literal_removal,[],[f22670])).
fof(f22670,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,X2)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))) )),
  inference(resolution,[],[f783,f460])).
fof(f22684,plain,(
  ( ! [X1] : (p(s(bool,X1)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))) )),
  inference(duplicate_literal_removal,[],[f22669])).
fof(f22669,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,X1)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))) )),
  inference(resolution,[],[f783,f465])).
fof(f22520,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | $spl42),
  inference(resolution,[],[f4751,f355])).
fof(f22513,plain,(
  ( ! [X3] : (p(s(bool,X3)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))) )),
  inference(duplicate_literal_removal,[],[f22501])).
fof(f22501,plain,(
  ( ! [X3] : (p(s(bool,X3)) | p(s(bool,X3)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))) )),
  inference(resolution,[],[f782,f463])).
fof(f22514,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))) )),
  inference(duplicate_literal_removal,[],[f22498])).
fof(f22498,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | p(s(bool,X0))) )),
  inference(resolution,[],[f782,f464])).
fof(f22484,plain,(
  ( ! [X2,X3] : (~p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | ($spl7 | $spl10)),
  inference(resolution,[],[f4070,f19313])).
fof(f22483,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X1) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | ($spl7 | $spl10)),
  inference(resolution,[],[f4070,f19575])).
fof(f22466,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | ($spl7 | $spl10)),
  inference(resolution,[],[f4069,f19575])).
fof(f22442,plain,(
  ( ! [X4,X5] : (sP9('const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X4,X5) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | ~sP7(X4,X5,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | ~sP6(X5,X4,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | $spl7),
  inference(resolution,[],[f4067,f427])).
fof(f22427,plain,(
  ( ! [X28,X27] : (~p(s(bool,X27)) | s(bool,X28) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | s(bool,'T_0') != s(bool,X28)) ) | ($spl7 | $spl10)),
  inference(resolution,[],[f4066,f19111])).
fof(f22426,plain,(
  ( ! [X26,X25] : (s(bool,'T_0') != s(bool,X25) | sP9(X26,'F_0','const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | s(bool,X25) != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | ($spl7 | $spl10)),
  inference(resolution,[],[f4066,f19068])).
fof(f22425,plain,(
  ( ! [X24] : (s(bool,'T_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | ~p(s(bool,X24))) ) | $spl7),
  inference(resolution,[],[f4066,f3327])).
fof(f22424,plain,(
  ( ! [X23] : (s(bool,'F_0') = s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | ~p(s(bool,X23))) ) | $spl7),
  inference(resolution,[],[f4066,f3285])).
fof(f22421,plain,(
  ( ! [X20] : (sP9(X20,'F_0','const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | s(bool,'T_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | $spl7),
  inference(resolution,[],[f4066,f3227])).
fof(f22420,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,X19) | s(bool,'T_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | ($spl7 | $spl10)),
  inference(resolution,[],[f4066,f19799])).
fof(f22415,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | p(s(bool,X9)) | s(bool,'T_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | ($spl7 | $spl10)),
  inference(resolution,[],[f4066,f19159])).
fof(f22201,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))),
  inference(subsumption_resolution,[],[f21859,f588])).
fof(f21859,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))),
  inference(superposition,[],[f355,f680])).
fof(f21769,plain,(
  ( ! [X2,X3] : (~p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | ($spl5 | $spl10)),
  inference(resolution,[],[f1416,f19313])).
fof(f21768,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X1) | s(bool,X1) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | ($spl5 | $spl10)),
  inference(resolution,[],[f1416,f19575])).
fof(f21762,plain,(
  ( ! [X3] : (p(s(bool,X3)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(duplicate_literal_removal,[],[f21749])).
fof(f21749,plain,(
  ( ! [X3] : (p(s(bool,X3)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | p(s(bool,X3))) )),
  inference(resolution,[],[f662,f458])).
fof(f21763,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(duplicate_literal_removal,[],[f21748])).
fof(f21748,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,X2)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(resolution,[],[f662,f467])).
fof(f21729,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | s(bool,X1) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | ($spl5 | $spl10)),
  inference(resolution,[],[f1415,f19575])).
fof(f21705,plain,(
  ( ! [X4,X5] : (sP9('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X4,X5) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | ~sP7(X4,X5,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | ~sP6(X5,X4,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | $spl5),
  inference(resolution,[],[f1413,f427])).
fof(f21690,plain,(
  ( ! [X28,X27] : (~p(s(bool,X27)) | s(bool,X28) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | s(bool,'T_0') != s(bool,X28)) ) | ($spl5 | $spl10)),
  inference(resolution,[],[f1412,f19111])).
fof(f21689,plain,(
  ( ! [X26,X25] : (s(bool,'T_0') != s(bool,X25) | sP9(X26,'F_0','const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | s(bool,X25) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | ($spl5 | $spl10)),
  inference(resolution,[],[f1412,f19068])).
fof(f21688,plain,(
  ( ! [X24] : (s(bool,'T_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | ~p(s(bool,X24))) ) | $spl5),
  inference(resolution,[],[f1412,f3327])).
fof(f21687,plain,(
  ( ! [X23] : (s(bool,'F_0') = s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | ~p(s(bool,X23))) ) | $spl5),
  inference(resolution,[],[f1412,f3285])).
fof(f21684,plain,(
  ( ! [X20] : (sP9(X20,'F_0','const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | s(bool,'T_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | $spl5),
  inference(resolution,[],[f1412,f3227])).
fof(f21683,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,X19) | s(bool,'T_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | ($spl5 | $spl10)),
  inference(resolution,[],[f1412,f19799])).
fof(f21678,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | p(s(bool,X9)) | s(bool,'T_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | ($spl5 | $spl10)),
  inference(resolution,[],[f1412,f19159])).
fof(f21640,plain,(
  ( ! [X3] : (p(s(bool,X3)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(duplicate_literal_removal,[],[f21627])).
fof(f21627,plain,(
  ( ! [X3] : (p(s(bool,X3)) | p(s(bool,X3)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(resolution,[],[f661,f458])).
fof(f21641,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(duplicate_literal_removal,[],[f21626])).
fof(f21626,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | p(s(bool,X2))) )),
  inference(resolution,[],[f661,f467])).
fof(f21571,plain,(
  ( ! [X12] : (p(s(bool,X12)) | s(bool,'F_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) | $spl2),
  inference(resolution,[],[f1168,f3286])).
fof(f21566,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) | ($spl2 | $spl10)),
  inference(resolution,[],[f1168,f19474])).
fof(f21565,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'F_0') != s(bool,X1) | s(bool,X1) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) | ($spl2 | $spl10)),
  inference(resolution,[],[f1168,f19603])).
fof(f21548,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'F_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) | s(bool,X1) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) | ($spl2 | $spl10)),
  inference(resolution,[],[f1167,f19603])).
fof(f21538,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(duplicate_literal_removal,[],[f21526])).
fof(f21526,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | ~p(s(bool,X2))) )),
  inference(resolution,[],[f660,f460])).
fof(f21539,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(duplicate_literal_removal,[],[f21525])).
fof(f21525,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | ~p(s(bool,X1))) )),
  inference(resolution,[],[f660,f465])).
fof(f21480,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) | s(bool,'T_0') != s(bool,X19)) ) | ($spl2 | $spl10)),
  inference(resolution,[],[f1162,f19799])).
fof(f21479,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) | ~p(s(bool,X16)) | ~p(s(bool,X17))) ) | ($spl2 | $spl10)),
  inference(resolution,[],[f1162,f19796])).
fof(f21476,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X11) | ~p(s(bool,X11)) | s(bool,'F_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) | ($spl2 | $spl10)),
  inference(resolution,[],[f1162,f19168])).
fof(f21433,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(duplicate_literal_removal,[],[f21422])).
fof(f21422,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | ~p(s(bool,X3))) )),
  inference(resolution,[],[f659,f463])).
fof(f21434,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(duplicate_literal_removal,[],[f21419])).
fof(f21419,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(resolution,[],[f659,f464])).
fof(f21378,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | ~p(s(bool,X4)) | p(s(bool,X5))) )),
  inference(resolution,[],[f776,f432])).
fof(f21377,plain,(
  ( ! [X2,X3] : (~p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | $spl10),
  inference(resolution,[],[f776,f19313])).
fof(f21376,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X1) | s(bool,X1) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | $spl10),
  inference(resolution,[],[f776,f19575])).
fof(f21373,plain,(
  ( ! [X6,X7] : (sP9(X6,X7,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))) )),
  inference(subsumption_resolution,[],[f21372,f772])).
fof(f21372,plain,(
  ( ! [X6,X7] : (sP9(X6,X7,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | ~sP8(X6,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X7)) )),
  inference(subsumption_resolution,[],[f21362,f439])).
fof(f21362,plain,(
  ( ! [X6,X7] : (sP9(X6,X7,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | ~sP8(X6,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X7) | ~sP6('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X7,X6)) )),
  inference(resolution,[],[f775,f425])).
fof(f21361,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | p(s(bool,X5))) )),
  inference(resolution,[],[f775,f432])).
fof(f21360,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | p(s(bool,X2)) | s(bool,X3) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | $spl10),
  inference(resolution,[],[f775,f19313])).
fof(f21359,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | s(bool,X1) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | $spl10),
  inference(resolution,[],[f775,f19575])).
fof(f21348,plain,(
  ( ! [X28,X27] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | s(bool,X27) != s(bool,X28) | s(bool,'T_0') != s(bool,X27)) ) | $spl10),
  inference(resolution,[],[f773,f19111])).
fof(f21346,plain,(
  ( ! [X24] : (s(bool,'T_0') != s(bool,X24) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))) )),
  inference(resolution,[],[f773,f3327])).
fof(f21345,plain,(
  ( ! [X23] : (s(bool,'F_0') = s(bool,X23) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))) )),
  inference(resolution,[],[f773,f3285])).
fof(f21343,plain,(
  ( ! [X21] : (~p(s(bool,X21)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))) )),
  inference(resolution,[],[f773,f3242])).
fof(f21341,plain,(
  ( ! [X19,X18] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | s(bool,'F_0') != s(bool,X18) | s(bool,'T_0') != s(bool,X19)) ) | $spl10),
  inference(resolution,[],[f773,f19799])).
fof(f21340,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,X16) | ~p(s(bool,X17)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))) ) | $spl10),
  inference(resolution,[],[f773,f19796])).
fof(f21335,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | ~p(s(bool,X6)) | p(s(bool,X7))) )),
  inference(resolution,[],[f773,f428])).
fof(f21334,plain,(
  ( ! [X4,X5] : (sP9('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X4,X5) | p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | ~sP7(X4,X5,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | ~sP6(X5,X4,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) )),
  inference(resolution,[],[f773,f427])).
fof(f21319,plain,(
  ( ! [X28,X27] : (~p(s(bool,X27)) | s(bool,X28) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | s(bool,'T_0') != s(bool,X28)) ) | $spl10),
  inference(resolution,[],[f772,f19111])).
fof(f21318,plain,(
  ( ! [X26,X25] : (s(bool,'T_0') != s(bool,X25) | sP9(X26,'F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | s(bool,X25) != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | $spl10),
  inference(resolution,[],[f772,f19068])).
fof(f21317,plain,(
  ( ! [X24] : (s(bool,'T_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | ~p(s(bool,X24))) )),
  inference(resolution,[],[f772,f3327])).
fof(f21316,plain,(
  ( ! [X23] : (s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | ~p(s(bool,X23))) )),
  inference(resolution,[],[f772,f3285])).
fof(f21314,plain,(
  ( ! [X21] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | ~p(s(bool,X21))) )),
  inference(resolution,[],[f772,f3242])).
fof(f21313,plain,(
  ( ! [X20] : (sP9(X20,'F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | s(bool,'T_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) )),
  inference(resolution,[],[f772,f3227])).
fof(f21312,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,X19) | s(bool,'T_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | $spl10),
  inference(resolution,[],[f772,f19799])).
fof(f21311,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,X16) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | ~p(s(bool,X17))) ) | $spl10),
  inference(resolution,[],[f772,f19796])).
fof(f21309,plain,(
  ( ! [X12,X13] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | sP9(X12,X13,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | ~sP7(X13,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X12)) ) | $spl10),
  inference(resolution,[],[f772,f19262])).
fof(f21308,plain,(
  ( ! [X10,X11] : (sP9(X10,X11,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | s(bool,'F_0') != s(bool,X11)) ) | $spl10),
  inference(resolution,[],[f772,f19168])).
fof(f21307,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | p(s(bool,X9)) | s(bool,'T_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | $spl10),
  inference(resolution,[],[f772,f19159])).
fof(f21306,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | p(s(bool,X7))) )),
  inference(resolution,[],[f772,f428])).
fof(f21326,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | ~sP7(X1,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X0)) )),
  inference(subsumption_resolution,[],[f21303,f439])).
fof(f21303,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | ~sP7(X1,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X0) | ~sP6('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X1,X0)) )),
  inference(resolution,[],[f772,f425])).
fof(f21290,plain,(
  ( ! [X2,X3] : (sP5(X2,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X3) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | ~p(s(bool,X3))) )),
  inference(resolution,[],[f760,f407])).
fof(f21289,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | ~p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(resolution,[],[f760,f411])).
fof(f21276,plain,(
  ( ! [X2,X3] : (sP5(X2,X3,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | ~p(s(bool,X3)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))) )),
  inference(resolution,[],[f759,f407])).
fof(f21275,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | p(s(bool,X1))) )),
  inference(resolution,[],[f759,f411])).
fof(f21274,plain,(
  ( ! [X6,X7] : (sP3('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X6,X7) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))) )),
  inference(subsumption_resolution,[],[f21264,f400])).
fof(f21264,plain,(
  ( ! [X6,X7] : (sP3('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X6,X7) | p(s(bool,X6)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))) )),
  inference(resolution,[],[f746,f392])).
fof(f21273,plain,(
  ( ! [X4,X5] : (sP3('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X4,X5) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))) )),
  inference(subsumption_resolution,[],[f21263,f399])).
fof(f21263,plain,(
  ( ! [X4,X5] : (sP3('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X4,X5) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | p(s(bool,X5))) )),
  inference(resolution,[],[f746,f391])).
fof(f21262,plain,(
  ( ! [X2,X3] : (sP3('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X2,X3) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))) )),
  inference(resolution,[],[f746,f586])).
fof(f21261,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(resolution,[],[f746,f394])).
fof(f21260,plain,(
  ( ! [X6,X7] : (sP1('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X6,X7) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))) )),
  inference(subsumption_resolution,[],[f21250,f383])).
fof(f21250,plain,(
  ( ! [X6,X7] : (sP1('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X6,X7) | p(s(bool,X6)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))) )),
  inference(resolution,[],[f730,f375])).
fof(f21259,plain,(
  ( ! [X4,X5] : (sP1('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X4,X5) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))) )),
  inference(subsumption_resolution,[],[f21249,f382])).
fof(f21249,plain,(
  ( ! [X4,X5] : (sP1('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X4,X5) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | ~p(s(bool,X5))) )),
  inference(resolution,[],[f730,f374])).
fof(f21248,plain,(
  ( ! [X2,X3] : (sP1('const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X2,X3) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))) )),
  inference(resolution,[],[f730,f585])).
fof(f21247,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(resolution,[],[f730,f377])).
fof(f21235,plain,(
  ( ! [X8,X9] : (sP1(X8,X9,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | p(s(bool,X9)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))) )),
  inference(resolution,[],[f729,f376])).
fof(f21231,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))) )),
  inference(resolution,[],[f729,f377])).
fof(f21220,plain,(
  ( ! [X12] : (p(s(bool,X12)) | s(bool,'F_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) )),
  inference(resolution,[],[f656,f3286])).
fof(f21216,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | p(s(bool,X4)) | p(s(bool,X5))) )),
  inference(resolution,[],[f656,f436])).
fof(f21215,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) | $spl10),
  inference(resolution,[],[f656,f19474])).
fof(f21214,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'F_0') != s(bool,X1) | s(bool,X1) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) | $spl10),
  inference(resolution,[],[f656,f19603])).
fof(f21212,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X9) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(subsumption_resolution,[],[f21211,f650])).
fof(f21211,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X9) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | ~sP8(X8,X9,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) )),
  inference(subsumption_resolution,[],[f21202,f435])).
fof(f21202,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X9) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | ~sP8(X8,X9,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) | ~sP7('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X9,X8)) )),
  inference(resolution,[],[f655,f426])).
fof(f21200,plain,(
  ( ! [X4,X5] : (p(s(bool,X4)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | p(s(bool,X5))) )),
  inference(resolution,[],[f655,f436])).
fof(f21199,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | p(s(bool,X2)) | s(bool,X3) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) | $spl10),
  inference(resolution,[],[f655,f19474])).
fof(f21198,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'F_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) | s(bool,X1) != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) | $spl10),
  inference(resolution,[],[f655,f19603])).
fof(f21186,plain,(
  ( ! [X12] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | s(bool,'F_0') != s(bool,X12)) )),
  inference(resolution,[],[f654,f3286])).
fof(f21197,plain,(
  ( ! [X10,X11] : (sP9('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X10,X11) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(subsumption_resolution,[],[f21196,f653])).
fof(f21196,plain,(
  ( ! [X10,X11] : (sP9('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X10,X11) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | ~sP7(X10,X11,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) )),
  inference(subsumption_resolution,[],[f21185,f431])).
fof(f21185,plain,(
  ( ! [X10,X11] : (sP9('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X10,X11) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | ~sP8('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X11,X10) | ~sP7(X10,X11,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) )),
  inference(resolution,[],[f654,f427])).
fof(f21182,plain,(
  ( ! [X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(resolution,[],[f654,f436])).
fof(f21181,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | s(bool,X2) != s(bool,X3)) ) | $spl10),
  inference(resolution,[],[f654,f19474])).
fof(f21180,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | s(bool,'F_0') != s(bool,X0) | s(bool,X0) != s(bool,X1)) ) | $spl10),
  inference(resolution,[],[f654,f19603])).
fof(f21179,plain,(
  ( ! [X10,X11] : (sP9('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X10,X11) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | ~sP6(X11,X10,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) )),
  inference(subsumption_resolution,[],[f21169,f431])).
fof(f21169,plain,(
  ( ! [X10,X11] : (sP9('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X10,X11) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | ~sP8('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X11,X10) | ~sP6(X11,X10,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) )),
  inference(resolution,[],[f653,f427])).
fof(f21166,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | ~p(s(bool,X5)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(resolution,[],[f653,f432])).
fof(f21165,plain,(
  ( ! [X2,X3] : (~p(s(bool,X2)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | s(bool,X2) != s(bool,X3)) ) | $spl10),
  inference(resolution,[],[f653,f19313])).
fof(f21164,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | s(bool,'T_0') != s(bool,X0) | s(bool,X0) != s(bool,X1)) ) | $spl10),
  inference(resolution,[],[f653,f19575])).
fof(f21153,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) | s(bool,'T_0') != s(bool,X19)) ) | $spl10),
  inference(resolution,[],[f650,f19799])).
fof(f21152,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) | ~p(s(bool,X16)) | ~p(s(bool,X17))) ) | $spl10),
  inference(resolution,[],[f650,f19796])).
fof(f21151,plain,(
  ( ! [X14,X15] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | sP9(X14,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X15) | ~sP6(X15,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X14)) ) | $spl10),
  inference(resolution,[],[f650,f19385])).
fof(f21149,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X11) | ~p(s(bool,X11)) | s(bool,'F_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) | $spl10),
  inference(resolution,[],[f650,f19168])).
fof(f21148,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X9) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | s(bool,'T_0') != s(bool,X9)) ) | $spl10),
  inference(resolution,[],[f650,f19159])).
fof(f21147,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(resolution,[],[f650,f428])).
fof(f21161,plain,(
  ( ! [X2,X3] : (sP9(X2,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X3) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | ~sP6(X3,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X2)) )),
  inference(subsumption_resolution,[],[f21145,f435])).
fof(f21145,plain,(
  ( ! [X2,X3] : (sP9(X2,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X3) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | ~sP7('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X3,X2) | ~sP6(X3,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X2)) )),
  inference(resolution,[],[f650,f426])).
fof(f21135,plain,(
  ( ! [X8,X9] : (sP5('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X8,X9) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(resolution,[],[f637,f587])).
fof(f21143,plain,(
  ( ! [X6,X7] : (sP5('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X6,X7) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(subsumption_resolution,[],[f21134,f415])).
fof(f21134,plain,(
  ( ! [X6,X7] : (sP5('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X6,X7) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | ~p(s(bool,X7))) )),
  inference(resolution,[],[f637,f409])).
fof(f21142,plain,(
  ( ! [X4,X5] : (sP5('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X4,X5) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(subsumption_resolution,[],[f21133,f416])).
fof(f21133,plain,(
  ( ! [X4,X5] : (sP5('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X4,X5) | ~p(s(bool,X4)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(resolution,[],[f637,f408])).
fof(f21131,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(resolution,[],[f637,f411])).
fof(f21122,plain,(
  ( ! [X8,X9] : (sP3(X8,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X9) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | p(s(bool,X9))) )),
  inference(resolution,[],[f621,f393])).
fof(f21118,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | p(s(bool,X1))) )),
  inference(resolution,[],[f621,f394])).
fof(f21109,plain,(
  ( ! [X8,X9] : (sP3(X8,X9,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) | p(s(bool,X9)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(resolution,[],[f620,f393])).
fof(f21105,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(resolution,[],[f620,f394])).
fof(f21057,plain,(
  ( ! [X8,X9] : (sP1(X8,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X9) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | ~p(s(bool,X9))) )),
  inference(resolution,[],[f607,f376])).
fof(f21053,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | ~p(s(bool,X1))) )),
  inference(resolution,[],[f607,f377])).
fof(f21052,plain,(
  ( ! [X1] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | p(s(bool,X1))) )),
  inference(subsumption_resolution,[],[f21044,f538])).
fof(f21044,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | p(s(bool,X1))) )),
  inference(resolution,[],[f668,f474])).
fof(f21051,plain,(
  ( ! [X0] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f21043,f538])).
fof(f21043,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | p(s(bool,X0))) )),
  inference(resolution,[],[f668,f478])).
fof(f21040,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(subsumption_resolution,[],[f21032,f538])).
fof(f21032,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(resolution,[],[f667,f474])).
fof(f21039,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(subsumption_resolution,[],[f21031,f538])).
fof(f21031,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(resolution,[],[f667,f478])).
fof(f21025,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | ($spl6 | $spl30)),
  inference(subsumption_resolution,[],[f21017,f538])).
fof(f21017,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | ($spl6 | $spl30)),
  inference(resolution,[],[f20932,f474])).
fof(f21022,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | ($spl6 | $spl30)),
  inference(subsumption_resolution,[],[f21016,f538])).
fof(f21016,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14))))) ) | ($spl6 | $spl30)),
  inference(resolution,[],[f20932,f478])).
fof(f21006,plain,(
  ( ! [X8,X9] : (sP3(X8,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X9) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | p(s(bool,X9))) ) | $spl30),
  inference(resolution,[],[f20429,f393])).
fof(f20429,plain,(
  ( ! [X50,X49] : (sP2(X49,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X50)) ) | $spl30),
  inference(backward_demodulation,[],[f3880,f621])).
fof(f21002,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | p(s(bool,X1))) ) | $spl30),
  inference(resolution,[],[f20429,f394])).
fof(f20991,plain,(
  ( ! [X8,X9] : (sP3(X8,X9,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))) | p(s(bool,X9)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | $spl30),
  inference(resolution,[],[f20428,f393])).
fof(f20428,plain,(
  ( ! [X47,X48] : (sP2(X47,X48,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | $spl30),
  inference(backward_demodulation,[],[f3880,f620])).
fof(f20987,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | $spl30),
  inference(resolution,[],[f20428,f394])).
fof(f20958,plain,(
  ( ! [X8,X9] : (sP1(X8,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X9) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | ~p(s(bool,X9))) ) | $spl30),
  inference(resolution,[],[f20418,f376])).
fof(f20418,plain,(
  ( ! [X21,X22] : (sP0(X21,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X22)) ) | $spl30),
  inference(backward_demodulation,[],[f3880,f607])).
fof(f20954,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | ~p(s(bool,X1))) ) | $spl30),
  inference(resolution,[],[f20418,f377])).
fof(f20804,plain,(
  ( ! [X1] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | p(s(bool,X1))) ) | $spl30),
  inference(subsumption_resolution,[],[f20796,f538])).
fof(f20796,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | p(s(bool,X1))) ) | $spl30),
  inference(resolution,[],[f20463,f474])).
fof(f20463,plain,(
  ( ! [X134] : (sP11('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)),X134)) ) | $spl30),
  inference(backward_demodulation,[],[f3880,f668])).
fof(f20803,plain,(
  ( ! [X0] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | p(s(bool,X0))) ) | $spl30),
  inference(subsumption_resolution,[],[f20795,f538])).
fof(f20795,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) | p(s(bool,X0))) ) | $spl30),
  inference(resolution,[],[f20463,f478])).
fof(f20794,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | $spl30),
  inference(subsumption_resolution,[],[f20786,f538])).
fof(f20786,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | $spl30),
  inference(resolution,[],[f20462,f474])).
fof(f20462,plain,(
  ( ! [X133] : (sP11(X133,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | $spl30),
  inference(backward_demodulation,[],[f3880,f667])).
fof(f20793,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | $spl30),
  inference(subsumption_resolution,[],[f20785,f538])).
fof(f20785,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | $spl30),
  inference(resolution,[],[f20462,f478])).
fof(f12461,plain,(
  s(sK12,sK14) = s(sK12,sK16) | ~p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))),
  inference(resolution,[],[f721,f578])).
fof(f7154,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))),
  inference(subsumption_resolution,[],[f6827,f538])).
fof(f6827,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))),
  inference(superposition,[],[f355,f806])).
fof(f20274,plain,(
  ( ! [X134,X132,X133] : (p(s(bool,X132)) | ~p(s(bool,X133)) | sP9(X134,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X133) | p(s(bool,X134))) ) | ($spl7 | $spl10)),
  inference(subsumption_resolution,[],[f20214,f553])).
fof(f20214,plain,(
  ( ! [X134,X132,X133] : (s(bool,'F_0') != s(bool,X132) | p(s(bool,X132)) | ~p(s(bool,X133)) | sP9(X134,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X133) | p(s(bool,X134))) ) | ($spl7 | $spl10)),
  inference(superposition,[],[f7606,f12686])).
fof(f12686,plain,(
  s(bool,'F_0') = s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | $spl7),
  inference(resolution,[],[f840,f553])).
fof(f20271,plain,(
  ( ! [X116,X114,X115] : (p(s(bool,X114)) | ~p(s(bool,X115)) | sP9(X116,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X115) | p(s(bool,X116))) ) | ($spl5 | $spl10)),
  inference(subsumption_resolution,[],[f20208,f553])).
fof(f20208,plain,(
  ( ! [X116,X114,X115] : (s(bool,'F_0') != s(bool,X114) | p(s(bool,X114)) | ~p(s(bool,X115)) | sP9(X116,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)),X115) | p(s(bool,X116))) ) | ($spl5 | $spl10)),
  inference(superposition,[],[f7606,f5934])).
fof(f5934,plain,(
  s(bool,'F_0') = s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | $spl5),
  inference(resolution,[],[f839,f553])).
fof(f20266,plain,(
  ( ! [X94,X95] : (sP9(X95,'F_0',X94) | p(s(bool,X95)) | ~p(s(bool,X94))) ) | $spl98),
  inference(cnf_transformation,[],[f20266_D])).
fof(f20266_D,plain,(
  ( ! [X94,X95] : (sP9(X95,'F_0',X94) | p(s(bool,X95)) | ~p(s(bool,X94))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f20270,plain,(
  ( ! [X107,X105,X106,X104] : (p(s(bool,X105)) | ~p(s(bool,X106)) | sP9(X107,'F_0',X106) | p(s(bool,X107)) | p(s(bool,X104))) ) | $spl10),
  inference(subsumption_resolution,[],[f20205,f491])).
fof(f491,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f204])).
fof(f204,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f203])).
fof(f203,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X12,X13] : ((p(s(bool,X12)) => p(s(bool,X13))) => ((p(s(bool,X13)) => p(s(bool,X12))) => s(bool,X12) = s(bool,X13)))),
  file('thm.relation.EXTEND_RTC_TC','thm.bool.IMP_ANTISYM_AX')).
fof(f20205,plain,(
  ( ! [X107,X105,X106,X104] : (s(bool,X104) != s(bool,X105) | p(s(bool,X105)) | ~p(s(bool,X106)) | sP9(X107,'F_0',X106) | p(s(bool,X107)) | p(s(bool,X104))) ) | $spl10),
  inference(superposition,[],[f7606,f553])).
fof(f20269,plain,(
  ( ! [X103,X101,X102,X100] : (p(s(bool,X101)) | ~p(s(bool,X102)) | sP9(X103,'F_0',X102) | p(s(bool,X103)) | p(s(bool,X100))) ) | $spl10),
  inference(subsumption_resolution,[],[f20204,f491])).
fof(f20204,plain,(
  ( ! [X103,X101,X102,X100] : (s(bool,X100) != s(bool,X101) | p(s(bool,X101)) | ~p(s(bool,X102)) | sP9(X103,'F_0',X102) | p(s(bool,X103)) | p(s(bool,X100))) ) | $spl10),
  inference(superposition,[],[f7606,f555])).
fof(f20267,plain,(
  ( ! [X92,X93] : (s(bool,X92) != s(bool,X93) | s(bool,'T_0') = s(bool,X92)) ) | $spl100),
  inference(cnf_transformation,[],[f20267_D])).
fof(f20267_D,plain,(
  ( ! [X92,X93] : (s(bool,X92) != s(bool,X93) | s(bool,'T_0') = s(bool,X92)) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f20268,plain,(
  ( ! [X99,X97,X98,X96] : (s(bool,X96) != s(bool,X97) | ~p(s(bool,X98)) | sP9(X99,'F_0',X98) | p(s(bool,X99)) | s(bool,'T_0') = s(bool,X96)) ) | $spl10),
  inference(subsumption_resolution,[],[f20203,f7687])).
fof(f20203,plain,(
  ( ! [X99,X97,X98,X96] : (s(bool,X96) != s(bool,X97) | p(s(bool,X97)) | ~p(s(bool,X98)) | sP9(X99,'F_0',X98) | p(s(bool,X99)) | s(bool,'T_0') = s(bool,X96)) ) | $spl10),
  inference(superposition,[],[f7606,f583])).
fof(f583,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f172])).
fof(f172,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  inference(rectify,[],[f62])).
fof(f62,axiom,(
  ! [X11] : (s(bool,'T_0') = s(bool,X11) | s(bool,'F_0') = s(bool,X11))),
  file('thm.relation.EXTEND_RTC_TC','thm.bool.BOOL_CASES_AX')).
fof(f20265,plain,(
  ( ! [X94,X92,X95,X93] : (s(bool,X92) != s(bool,X93) | ~p(s(bool,X94)) | sP9(X95,'F_0',X94) | p(s(bool,X95)) | s(bool,'T_0') = s(bool,X92)) ) | $spl10),
  inference(subsumption_resolution,[],[f20202,f7687])).
fof(f20202,plain,(
  ( ! [X94,X92,X95,X93] : (s(bool,X92) != s(bool,X93) | p(s(bool,X93)) | ~p(s(bool,X94)) | sP9(X95,'F_0',X94) | p(s(bool,X95)) | s(bool,'T_0') = s(bool,X92)) ) | $spl10),
  inference(superposition,[],[f7606,f584])).
fof(f20263,plain,(
  ( ! [X66,X64,X67] : (sP9(X67,X64,X66) | s(bool,'T_0') = s(bool,X64) | p(s(bool,X67)) | ~p(s(bool,X66))) ) | $spl96),
  inference(cnf_transformation,[],[f20263_D])).
fof(f20263_D,plain,(
  ( ! [X66,X64,X67] : (sP9(X67,X64,X66) | s(bool,'T_0') = s(bool,X64) | p(s(bool,X67)) | ~p(s(bool,X66))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f20264,plain,(
  ( ! [X74,X72,X75,X73] : (p(s(bool,X73)) | ~p(s(bool,X74)) | sP9(X75,X72,X74) | p(s(bool,X75)) | s(bool,'T_0') = s(bool,X72)) ) | $spl10),
  inference(subsumption_resolution,[],[f20197,f553])).
fof(f20197,plain,(
  ( ! [X74,X72,X75,X73] : (s(bool,'F_0') != s(bool,X73) | p(s(bool,X73)) | ~p(s(bool,X74)) | sP9(X75,X72,X74) | p(s(bool,X75)) | s(bool,'T_0') = s(bool,X72)) ) | $spl10),
  inference(superposition,[],[f7606,f584])).
fof(f20262,plain,(
  ( ! [X66,X64,X67,X65] : (p(s(bool,X65)) | ~p(s(bool,X66)) | sP9(X67,X64,X66) | p(s(bool,X67)) | s(bool,'T_0') = s(bool,X64)) ) | $spl10),
  inference(subsumption_resolution,[],[f20195,f553])).
fof(f20195,plain,(
  ( ! [X66,X64,X67,X65] : (s(bool,'F_0') != s(bool,X65) | p(s(bool,X65)) | ~p(s(bool,X66)) | sP9(X67,X64,X66) | p(s(bool,X67)) | s(bool,'T_0') = s(bool,X64)) ) | $spl10),
  inference(superposition,[],[f7606,f583])).
fof(f20258,plain,(
  ( ! [X28,X24,X27] : (~p(s(bool,X27)) | p(s(bool,X24)) | p(s(bool,X28)) | sP9(X28,X24,X27)) ) | $spl94),
  inference(cnf_transformation,[],[f20258_D])).
fof(f20258_D,plain,(
  ( ! [X28,X24,X27] : (~p(s(bool,X27)) | p(s(bool,X24)) | p(s(bool,X28)) | sP9(X28,X24,X27)) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f20261,plain,(
  ( ! [X50,X48,X51,X49] : (p(s(bool,X49)) | ~p(s(bool,X50)) | sP9(X51,X48,X50) | p(s(bool,X51)) | p(s(bool,X48))) ) | $spl10),
  inference(subsumption_resolution,[],[f20191,f553])).
fof(f20191,plain,(
  ( ! [X50,X48,X51,X49] : (s(bool,'F_0') != s(bool,X49) | p(s(bool,X49)) | ~p(s(bool,X50)) | sP9(X51,X48,X50) | p(s(bool,X51)) | p(s(bool,X48))) ) | $spl10),
  inference(superposition,[],[f7606,f555])).
fof(f20260,plain,(
  ( ! [X47,X45,X46,X44] : (p(s(bool,X45)) | ~p(s(bool,X46)) | sP9(X47,X44,X46) | p(s(bool,X47)) | p(s(bool,X44))) ) | $spl10),
  inference(subsumption_resolution,[],[f20190,f553])).
fof(f20190,plain,(
  ( ! [X47,X45,X46,X44] : (s(bool,'F_0') != s(bool,X45) | p(s(bool,X45)) | ~p(s(bool,X46)) | sP9(X47,X44,X46) | p(s(bool,X47)) | p(s(bool,X44))) ) | $spl10),
  inference(superposition,[],[f7606,f553])).
fof(f20259,plain,(
  ( ! [X30,X33,X31,X29,X32] : (p(s(bool,X31)) | ~p(s(bool,X32)) | sP9(X33,X29,X32) | p(s(bool,X33)) | p(s(bool,X29)) | p(s(bool,X30))) ) | $spl10),
  inference(subsumption_resolution,[],[f20187,f491])).
fof(f20187,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s(bool,X30) != s(bool,X31) | p(s(bool,X31)) | ~p(s(bool,X32)) | sP9(X33,X29,X32) | p(s(bool,X33)) | p(s(bool,X29)) | p(s(bool,X30))) ) | $spl10),
  inference(superposition,[],[f7606,f491])).
fof(f20257,plain,(
  ( ! [X28,X26,X24,X27,X25] : (p(s(bool,X26)) | ~p(s(bool,X27)) | sP9(X28,X24,X27) | p(s(bool,X28)) | p(s(bool,X25)) | p(s(bool,X24))) ) | $spl10),
  inference(subsumption_resolution,[],[f20186,f491])).
fof(f20186,plain,(
  ( ! [X28,X26,X24,X27,X25] : (s(bool,X25) != s(bool,X26) | p(s(bool,X26)) | ~p(s(bool,X27)) | sP9(X28,X24,X27) | p(s(bool,X28)) | p(s(bool,X25)) | p(s(bool,X24))) ) | $spl10),
  inference(superposition,[],[f7606,f491])).
fof(f20044,plain,(
  ( ! [X13] : (s(bool,X13) != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK16))) | s(bool,'T_0') != s(bool,X13)) ) | $spl90),
  inference(cnf_transformation,[],[f20044_D])).
fof(f20044_D,plain,(
  ( ! [X13] : (s(bool,X13) != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK16))) | s(bool,'T_0') != s(bool,X13)) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f20114,plain,(
  ( ! [X50,X51] : (p(s(bool,X50)) | s(bool,'T_0') != s(bool,X51) | s(bool,X51) != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK16)))) ) | $spl10),
  inference(resolution,[],[f19575,f12523])).
fof(f12523,plain,(
  ( ! [X123,X122] : (sP7(happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK16)),X122,X123)) )),
  inference(resolution,[],[f721,f435])).
fof(f20126,plain,(
  ( ! [X39] : (s(bool,X39) != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK16)))) ) | $spl92),
  inference(cnf_transformation,[],[f20126_D])).
fof(f20126_D,plain,(
  ( ! [X39] : (s(bool,X39) != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK16)))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f18865,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK16))) | $spl79),
  inference(cnf_transformation,[],[f18865_D])).
fof(f18865_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK16))) <=> ~$spl79),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl79])])).
fof(f20108,plain,(
  ( ! [X39,X38] : (p(s(bool,X38)) | s(bool,'T_0') != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK16))) | s(bool,X39) != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK16)))) ) | $spl10),
  inference(resolution,[],[f19575,f12522])).
fof(f12522,plain,(
  ( ! [X121,X120] : (sP7(X120,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK16)),X121)) )),
  inference(resolution,[],[f721,f434])).
fof(f18866,plain,(
  ( ! [X9] : (sP9(X9,'F_0',happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK16)))) ) | $spl80),
  inference(cnf_transformation,[],[f18866_D])).
fof(f18866_D,plain,(
  ( ! [X9] : (sP9(X9,'F_0',happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK16)))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f20082,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') != s(bool,X12) | sP9(X13,'F_0',happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK16))) | s(bool,X12) != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK16)))) ) | $spl10),
  inference(resolution,[],[f19068,f12519])).
fof(f12519,plain,(
  ( ! [X114,X115] : (sP8(X114,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK16)),X115)) )),
  inference(resolution,[],[f721,f430])).
fof(f20061,plain,(
  ( ! [X24,X25] : (~p(s(bool,X24)) | s(bool,'F_0') != s(bool,X25) | s(bool,'T_0') != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK16)))) ) | $spl10),
  inference(resolution,[],[f19799,f12519])).
fof(f20036,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | s(bool,X13) != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK16))) | s(bool,'T_0') != s(bool,X13)) ) | $spl10),
  inference(resolution,[],[f19111,f12519])).
fof(f19652,plain,(
  ( ! [X47] : (~p(s(bool,X47)) | s(bool,X47) != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK16)))) ) | $spl88),
  inference(cnf_transformation,[],[f19652_D])).
fof(f19652_D,plain,(
  ( ! [X47] : (~p(s(bool,X47)) | s(bool,X47) != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK16)))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f19633,plain,(
  ( ! [X47,X48] : (~p(s(bool,X47)) | p(s(bool,X48)) | s(bool,X47) != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK16)))) ) | $spl10),
  inference(resolution,[],[f19313,f12523])).
fof(f19371,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') != s(bool,X4) | s(bool,X3) != s(bool,X4) | ~p(s(bool,X3))) ) | $spl86),
  inference(cnf_transformation,[],[f19371_D])).
fof(f19371_D,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') != s(bool,X4) | s(bool,X3) != s(bool,X4) | ~p(s(bool,X3))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f19370,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | s(bool,'F_0') != s(bool,X4) | p(s(bool,X5)) | s(bool,X3) != s(bool,X4)) ) | $spl10),
  inference(duplicate_literal_removal,[],[f19348])).
fof(f19348,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | s(bool,'F_0') != s(bool,X4) | p(s(bool,X5)) | p(s(bool,X5)) | s(bool,X3) != s(bool,X4)) ) | $spl10),
  inference(resolution,[],[f19200,f443])).
fof(f19197,plain,(
  ( ! [X24,X25] : (sP9(X24,X25,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK16))) | p(s(bool,X25))) ) | $spl84),
  inference(cnf_transformation,[],[f19197_D])).
fof(f19197_D,plain,(
  ( ! [X24,X25] : (sP9(X24,X25,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK16))) | p(s(bool,X25))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f19183,plain,(
  ( ! [X24,X25] : (sP9(X24,X25,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK16))) | p(s(bool,X25)) | s(bool,'T_0') != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK16)))) ) | $spl10),
  inference(resolution,[],[f19159,f12519])).
fof(f18939,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X7) | s(bool,'T_0') != s(bool,X6) | s(bool,X6) != s(bool,X7)) ) | $spl82),
  inference(cnf_transformation,[],[f18939_D])).
fof(f18939_D,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X7) | s(bool,'T_0') != s(bool,X6) | s(bool,X6) != s(bool,X7)) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f18938,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X6) | s(bool,X6) != s(bool,X7) | p(s(bool,X8)) | s(bool,'F_0') != s(bool,X7)) ) | $spl10),
  inference(duplicate_literal_removal,[],[f18925])).
fof(f18925,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X6) | s(bool,X6) != s(bool,X7) | p(s(bool,X8)) | p(s(bool,X8)) | s(bool,'F_0') != s(bool,X7)) ) | $spl10),
  inference(resolution,[],[f2629,f443])).
fof(f2629,plain,(
  ( ! [X10,X8,X9] : (sP9(X10,X8,'F_0') | s(bool,'T_0') != s(bool,X9) | s(bool,X8) != s(bool,X9) | p(s(bool,X10))) ) | $spl10),
  inference(superposition,[],[f2405,f440])).
fof(f18846,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK16))) | ~p(s(bool,X10))) )),
  inference(resolution,[],[f12519,f3327])).
fof(f18845,plain,(
  ( ! [X9] : (sP9(X9,'F_0',happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK16))) | s(bool,'T_0') != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK16)))) )),
  inference(resolution,[],[f12519,f3227])).
fof(f18863,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK16))) | $spl76),
  inference(cnf_transformation,[],[f18863_D])).
fof(f18863_D,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK16))) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f18844,plain,(
  ( ! [X8] : (s(bool,'F_0') = s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK16))) | ~p(s(bool,X8))) )),
  inference(resolution,[],[f12519,f3285])).
fof(f18758,plain,(
  $false | $spl32),
  inference(subsumption_resolution,[],[f3882,f721])).
fof(f13736,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) )),
  inference(resolution,[],[f947,f3286])).
fof(f12810,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | ~p(s(bool,X10))) ) | $spl7),
  inference(resolution,[],[f4066,f3327])).
fof(f12809,plain,(
  ( ! [X9] : (sP9(X9,'F_0','const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | s(bool,'T_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | $spl7),
  inference(resolution,[],[f4066,f3227])).
fof(f12808,plain,(
  ( ! [X8] : (s(bool,'F_0') = s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | ~p(s(bool,X8))) ) | $spl7),
  inference(resolution,[],[f4066,f3285])).
fof(f12589,plain,(
  $false | ($spl6 | $spl31)),
  inference(subsumption_resolution,[],[f12588,f577])).
fof(f577,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f40])).
fof(f40,axiom,(
  p(s(bool,'T_0'))),
  file('thm.relation.EXTEND_RTC_TC','thm.bool.TRUTH')).
fof(f12588,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl31)),
  inference(forward_demodulation,[],[f12587,f5544])).
fof(f5544,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | $spl6),
  inference(resolution,[],[f841,f557])).
fof(f12587,plain,(
  ~p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | $spl31),
  inference(subsumption_resolution,[],[f12461,f3879])).
fof(f7576,plain,(
  ( ! [X24,X22] : (s(bool,X22) != s(bool,X24) | ~p(s(bool,X22)) | p(s(bool,X24))) ) | $spl70),
  inference(cnf_transformation,[],[f7576_D])).
fof(f7576_D,plain,(
  ( ! [X24,X22] : (s(bool,X22) != s(bool,X24) | ~p(s(bool,X22)) | p(s(bool,X24))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f7578,plain,(
  ( ! [X26,X27,X25] : (s(bool,'T_0') != s(bool,X26) | s(bool,X25) != s(bool,X27) | p(s(bool,X27)) | ~p(s(bool,X25))) ) | $spl10),
  inference(subsumption_resolution,[],[f7542,f556])).
fof(f556,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f343])).
fof(f7542,plain,(
  ( ! [X26,X27,X25] : (s(bool,'T_0') != s(bool,X26) | s(bool,X25) != s(bool,X27) | p(s(bool,X27)) | ~p(s(bool,X26)) | ~p(s(bool,X25))) ) | $spl10),
  inference(superposition,[],[f2627,f494])).
fof(f494,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f204])).
fof(f7575,plain,(
  ( ! [X24,X23,X22] : (s(bool,'T_0') != s(bool,X23) | s(bool,X22) != s(bool,X24) | p(s(bool,X24)) | ~p(s(bool,X22))) ) | $spl10),
  inference(subsumption_resolution,[],[f7541,f556])).
fof(f7541,plain,(
  ( ! [X24,X23,X22] : (s(bool,'T_0') != s(bool,X23) | s(bool,X22) != s(bool,X24) | p(s(bool,X24)) | ~p(s(bool,X22)) | ~p(s(bool,X23))) ) | $spl10),
  inference(superposition,[],[f2627,f494])).
fof(f7505,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,'F_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(resolution,[],[f2214,f3286])).
fof(f4465,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | $spl37),
  inference(cnf_transformation,[],[f4465_D])).
fof(f4465_D,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) <=> ~$spl37),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl37])])).
fof(f6760,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | $spl6),
  inference(resolution,[],[f3824,f3286])).
fof(f6571,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | ~p(s(bool,X10))) ) | $spl5),
  inference(resolution,[],[f1412,f3327])).
fof(f6570,plain,(
  ( ! [X9] : (sP9(X9,'F_0','const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | s(bool,'T_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | $spl5),
  inference(resolution,[],[f1412,f3227])).
fof(f6569,plain,(
  ( ! [X8] : (s(bool,'F_0') = s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | ~p(s(bool,X8))) ) | $spl5),
  inference(resolution,[],[f1412,f3285])).
fof(f3361,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | $spl25),
  inference(cnf_transformation,[],[f3361_D])).
fof(f3361_D,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f5838,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,'F_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | $spl4),
  inference(resolution,[],[f3686,f3286])).
fof(f5832,plain,(
  ( ! [X3] : (p(s(bool,X3)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(duplicate_literal_removal,[],[f5823])).
fof(f5823,plain,(
  ( ! [X3] : (p(s(bool,X3)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | p(s(bool,X3))) )),
  inference(resolution,[],[f662,f458])).
fof(f5833,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(duplicate_literal_removal,[],[f5822])).
fof(f5822,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,X2)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(resolution,[],[f662,f467])).
fof(f5760,plain,(
  ( ! [X3] : (p(s(bool,X3)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(duplicate_literal_removal,[],[f5751])).
fof(f5751,plain,(
  ( ! [X3] : (p(s(bool,X3)) | p(s(bool,X3)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(resolution,[],[f661,f458])).
fof(f5761,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(duplicate_literal_removal,[],[f5750])).
fof(f5750,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | p(s(bool,X2))) )),
  inference(resolution,[],[f661,f467])).
fof(f5702,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,'F_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) | $spl2),
  inference(resolution,[],[f1168,f3286])).
fof(f5684,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(duplicate_literal_removal,[],[f5676])).
fof(f5676,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | ~p(s(bool,X2))) )),
  inference(resolution,[],[f660,f460])).
fof(f5685,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(duplicate_literal_removal,[],[f5675])).
fof(f5675,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | ~p(s(bool,X1))) )),
  inference(resolution,[],[f660,f465])).
fof(f5614,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(duplicate_literal_removal,[],[f5607])).
fof(f5607,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | ~p(s(bool,X3))) )),
  inference(resolution,[],[f659,f463])).
fof(f5615,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(duplicate_literal_removal,[],[f5604])).
fof(f5604,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(resolution,[],[f659,f464])).
fof(f5284,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | ~p(s(bool,X10))) )),
  inference(resolution,[],[f772,f3327])).
fof(f5283,plain,(
  ( ! [X9] : (sP9(X9,'F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | s(bool,'T_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) )),
  inference(resolution,[],[f772,f3227])).
fof(f5282,plain,(
  ( ! [X8] : (s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | ~p(s(bool,X8))) )),
  inference(resolution,[],[f772,f3285])).
fof(f5081,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,'F_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) )),
  inference(resolution,[],[f656,f3286])).
fof(f4764,plain,(
  ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | $spl45),
  inference(cnf_transformation,[],[f4764_D])).
fof(f4764_D,plain,(
  ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) <=> ~$spl45),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl45])])).
fof(f5088,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) | ~sP7(X3,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X2) | ~sP8(X2,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X3)) ) | $spl64),
  inference(cnf_transformation,[],[f5088_D])).
fof(f5088_D,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) | ~sP7(X3,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X2) | ~sP8(X2,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X3)) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f5078,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | ~sP8(X2,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X3) | ~sP7(X3,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X2)) )),
  inference(resolution,[],[f656,f425])).
fof(f5077,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(resolution,[],[f656,f436])).
fof(f5074,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X5) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(subsumption_resolution,[],[f5073,f650])).
fof(f5073,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X5) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | ~sP8(X4,X5,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) )),
  inference(subsumption_resolution,[],[f5064,f435])).
fof(f5064,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X5) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | ~sP8(X4,X5,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) | ~sP7('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X5,X4)) )),
  inference(resolution,[],[f655,f426])).
fof(f5062,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | p(s(bool,X1))) )),
  inference(resolution,[],[f655,f436])).
fof(f5049,plain,(
  ( ! [X8] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | s(bool,'F_0') != s(bool,X8)) )),
  inference(resolution,[],[f654,f3286])).
fof(f5059,plain,(
  ( ! [X6,X7] : (sP9('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X6,X7) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(subsumption_resolution,[],[f5058,f653])).
fof(f5058,plain,(
  ( ! [X6,X7] : (sP9('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X6,X7) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | ~sP7(X6,X7,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) )),
  inference(subsumption_resolution,[],[f5048,f431])).
fof(f5048,plain,(
  ( ! [X6,X7] : (sP9('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X6,X7) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | ~sP8('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X7,X6) | ~sP7(X6,X7,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) )),
  inference(resolution,[],[f654,f427])).
fof(f5045,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(resolution,[],[f654,f436])).
fof(f5043,plain,(
  ( ! [X6,X7] : (sP9('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X6,X7) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | ~sP6(X7,X6,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) )),
  inference(subsumption_resolution,[],[f5035,f431])).
fof(f5035,plain,(
  ( ! [X6,X7] : (sP9('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X6,X7) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | ~sP8('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X7,X6) | ~sP6(X7,X6,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) )),
  inference(resolution,[],[f653,f427])).
fof(f5032,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(resolution,[],[f653,f432])).
fof(f5030,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X5) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | ~sP6(X5,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X4)) )),
  inference(subsumption_resolution,[],[f5021,f435])).
fof(f5021,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X5) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | ~sP7('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X5,X4) | ~sP6(X5,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X4)) )),
  inference(resolution,[],[f650,f426])).
fof(f5019,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(resolution,[],[f650,f428])).
fof(f5009,plain,(
  ( ! [X8,X9] : (sP5('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X8,X9) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(resolution,[],[f637,f587])).
fof(f5018,plain,(
  ( ! [X6,X7] : (sP5('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X6,X7) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(subsumption_resolution,[],[f5008,f415])).
fof(f5008,plain,(
  ( ! [X6,X7] : (sP5('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X6,X7) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | ~p(s(bool,X7))) )),
  inference(resolution,[],[f637,f409])).
fof(f5016,plain,(
  ( ! [X4,X5] : (sP5('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X4,X5) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(subsumption_resolution,[],[f5007,f416])).
fof(f5007,plain,(
  ( ! [X4,X5] : (sP5('const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X4,X5) | ~p(s(bool,X4)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(resolution,[],[f637,f408])).
fof(f5005,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(resolution,[],[f637,f411])).
fof(f4995,plain,(
  ( ! [X8,X9] : (sP3(X8,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X9) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | p(s(bool,X9))) )),
  inference(resolution,[],[f621,f393])).
fof(f4991,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | p(s(bool,X1))) )),
  inference(resolution,[],[f621,f394])).
fof(f4981,plain,(
  ( ! [X8,X9] : (sP3(X8,X9,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))) | p(s(bool,X9)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(resolution,[],[f620,f393])).
fof(f4977,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(resolution,[],[f620,f394])).
fof(f4947,plain,(
  ( ! [X8,X9] : (sP1(X8,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)),X9) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | ~p(s(bool,X9))) )),
  inference(resolution,[],[f607,f376])).
fof(f4943,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | ~p(s(bool,X1))) )),
  inference(resolution,[],[f607,f377])).
fof(f4902,plain,(
  ( ! [X1] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | p(s(bool,X1))) )),
  inference(subsumption_resolution,[],[f4894,f538])).
fof(f4894,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | p(s(bool,X1))) )),
  inference(resolution,[],[f668,f474])).
fof(f4901,plain,(
  ( ! [X0] : (p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f4893,f538])).
fof(f4893,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | p(s(bool,X0))) )),
  inference(resolution,[],[f668,f478])).
fof(f4767,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))),
  inference(resolution,[],[f352,f355])).
fof(f4766,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(subsumption_resolution,[],[f4756,f538])).
fof(f4756,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(resolution,[],[f667,f474])).
fof(f4763,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(subsumption_resolution,[],[f4755,f538])).
fof(f4755,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))) )),
  inference(resolution,[],[f667,f478])).
fof(f4752,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | $spl30),
  inference(subsumption_resolution,[],[f4742,f538])).
fof(f4742,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | $spl30),
  inference(resolution,[],[f4484,f474])).
fof(f4484,plain,(
  ( ! [X133] : (sP11(X133,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14)))) ) | $spl30),
  inference(backward_demodulation,[],[f3880,f667])).
fof(f4749,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | $spl30),
  inference(subsumption_resolution,[],[f4741,f538])).
fof(f4741,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK14))))) ) | $spl30),
  inference(resolution,[],[f4484,f478])).
fof(f4625,plain,(
  $false | ($spl5 | $spl30)),
  inference(subsumption_resolution,[],[f4501,f369])).
fof(f4501,plain,(
  ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK14)))) | ($spl5 | $spl30)),
  inference(backward_demodulation,[],[f3880,f839])).
fof(f4454,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | ~p(s(bool,X10))) ) | $spl7),
  inference(resolution,[],[f4066,f3327])).
fof(f4453,plain,(
  ( ! [X9] : (sP9(X9,'F_0','const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | s(bool,'T_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | $spl7),
  inference(resolution,[],[f4066,f3227])).
fof(f4452,plain,(
  ( ! [X8] : (s(bool,'F_0') = s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | ~p(s(bool,X8))) ) | $spl7),
  inference(resolution,[],[f4066,f3285])).
fof(f4360,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,'F_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(resolution,[],[f1688,f3286])).
fof(f4224,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | ~p(s(bool,X10))) ) | $spl5),
  inference(resolution,[],[f1412,f3327])).
fof(f4223,plain,(
  ( ! [X9] : (sP9(X9,'F_0','const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | s(bool,'T_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | $spl5),
  inference(resolution,[],[f1412,f3227])).
fof(f4222,plain,(
  ( ! [X8] : (s(bool,'F_0') = s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | ~p(s(bool,X8))) ) | $spl5),
  inference(resolution,[],[f1412,f3285])).
fof(f3882,plain,(
  p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK16)))) | $spl32),
  inference(cnf_transformation,[],[f3882_D])).
fof(f3882_D,plain,(
  p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK16)))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f3767,plain,(
  p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK16)))) | s(sK12,sK14) = s(sK12,sK16) | $spl6),
  inference(resolution,[],[f841,f578])).
fof(f3613,plain,(
  ( ! [X16] : (p(s(bool,X16)) | s(bool,'F_0') != s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK15)))) )),
  inference(resolution,[],[f3286,f1543])).
fof(f3612,plain,(
  ( ! [X15] : (p(s(bool,X15)) | s(bool,'F_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) ) | $spl2),
  inference(resolution,[],[f3286,f1168])).
fof(f3597,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f3461,f3336])).
fof(f3461,plain,(
  p(s(bool,'F_0')) | $spl12),
  inference(resolution,[],[f3336,f574])).
fof(f574,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f346])).
fof(f3596,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f3460,f3336])).
fof(f3460,plain,(
  p(s(bool,'F_0')) | $spl12),
  inference(resolution,[],[f3336,f572])).
fof(f572,plain,(
  p(s(bool,'F_0')) | p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f345])).
fof(f3595,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f3456,f3336])).
fof(f3456,plain,(
  ( ! [X231] : (p(s(bool,X231))) ) | $spl12),
  inference(resolution,[],[f3336,f531])).
fof(f531,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f328])).
fof(f328,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))))),
  inference(flattening,[],[f327])).
fof(f327,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f210])).
fof(f210,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f139])).
fof(f139,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X11] : ((p(s(bool,X11)) => p(s(bool,X11))) <=> p(s(bool,'T_0')))),
  file('thm.relation.EXTEND_RTC_TC','thm.bool.IMP_CLAUSES|split|3')).
fof(f3594,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f3454,f3336])).
fof(f3454,plain,(
  p(s(bool,'F_0')) | $spl12),
  inference(resolution,[],[f3336,f528])).
fof(f528,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f326])).
fof(f326,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f325])).
fof(f325,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f209])).
fof(f209,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f138])).
fof(f138,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X11] : ((p(s(bool,'F_0')) => p(s(bool,X11))) <=> p(s(bool,'T_0')))),
  file('thm.relation.EXTEND_RTC_TC','thm.bool.IMP_CLAUSES|split|2')).
fof(f3593,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f3453,f3336])).
fof(f3453,plain,(
  ( ! [X229] : (p(s(bool,X229))) ) | $spl12),
  inference(resolution,[],[f3336,f525])).
fof(f525,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f324])).
fof(f324,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,'T_0'))))),
  inference(flattening,[],[f323])).
fof(f323,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f208])).
fof(f208,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f137])).
fof(f137,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X11] : ((p(s(bool,X11)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  file('thm.relation.EXTEND_RTC_TC','thm.bool.IMP_CLAUSES|split|1')).
fof(f3592,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f3452,f3336])).
fof(f3452,plain,(
  ( ! [X228] : (p(s(bool,X228))) ) | $spl12),
  inference(resolution,[],[f3336,f522])).
fof(f522,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f322])).
fof(f322,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f321])).
fof(f321,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (~p(s(bool,'T_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f207])).
fof(f207,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(ennf_transformation,[],[f136])).
fof(f136,plain,(
  ! [X0] : ((p(s(bool,'T_0')) => p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X11] : ((p(s(bool,'T_0')) => p(s(bool,X11))) <=> p(s(bool,X11)))),
  file('thm.relation.EXTEND_RTC_TC','thm.bool.IMP_CLAUSES|split|0')).
fof(f3462,plain,(
  $false | $spl12),
  inference(resolution,[],[f3336,f576])).
fof(f576,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.relation.EXTEND_RTC_TC',aHL_TRUTH)).
fof(f3463,plain,(
  $false | $spl12),
  inference(resolution,[],[f3336,f577])).
fof(f3464,plain,(
  $false | $spl12),
  inference(resolution,[],[f3336,f588])).
fof(f3509,plain,(
  $false | $spl12),
  inference(resolution,[],[f3336,f350])).
fof(f3513,plain,(
  $false | $spl12),
  inference(resolution,[],[f3336,f369])).
fof(f3515,plain,(
  $false | $spl12),
  inference(resolution,[],[f3336,f1625])).
fof(f3519,plain,(
  $false | ($spl2 | $spl12)),
  inference(resolution,[],[f3336,f714])).
fof(f3520,plain,(
  $false | $spl12),
  inference(resolution,[],[f3336,f884])).
fof(f3524,plain,(
  $false | $spl12),
  inference(resolution,[],[f3336,f882])).
fof(f3356,plain,(
  ( ! [X9] : (s(bool,'F_0') = s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | ~p(s(bool,X9))) ) | $spl5),
  inference(resolution,[],[f3285,f1412])).
fof(f3355,plain,(
  ( ! [X8] : (s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | ~p(s(bool,X8))) )),
  inference(resolution,[],[f3285,f772])).
fof(f3346,plain,(
  ( ! [X9] : (sP9(X9,'F_0','const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | s(bool,'T_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) ) | $spl5),
  inference(resolution,[],[f3227,f1412])).
fof(f3345,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | s(bool,'T_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) )),
  inference(resolution,[],[f3227,f772])).
fof(f3333,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | ~p(s(bool,X9))) ) | $spl5),
  inference(resolution,[],[f3327,f1412])).
fof(f3332,plain,(
  ( ! [X8] : (s(bool,'T_0') != s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))) | ~p(s(bool,X8))) )),
  inference(resolution,[],[f3327,f772])).
fof(f2613,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f2506,f2404])).
fof(f2506,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(resolution,[],[f2404,f575])).
fof(f2612,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f2505,f2404])).
fof(f2505,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(resolution,[],[f2404,f573])).
fof(f2611,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f2502,f2404])).
fof(f2502,plain,(
  ( ! [X224] : (~p(s(bool,X224))) ) | $spl8),
  inference(resolution,[],[f2404,f567])).
fof(f2610,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f2498,f2404])).
fof(f2498,plain,(
  ( ! [X220] : (~p(s(bool,X220))) ) | $spl8),
  inference(resolution,[],[f2404,f520])).
fof(f2504,plain,(
  $false | $spl8),
  inference(resolution,[],[f2404,f571])).
fof(f2546,plain,(
  $false | ($spl5 | $spl8)),
  inference(resolution,[],[f2404,f839])).
fof(f2556,plain,(
  $false | $spl8),
  inference(resolution,[],[f2404,f352])).
fof(f2405,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | s(bool,'T_0') != s(bool,X2)) ) | $spl10),
  inference(cnf_transformation,[],[f2405_D])).
fof(f2405_D,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | s(bool,'T_0') != s(bool,X2)) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f2403,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) )),
  inference(duplicate_literal_removal,[],[f2398])).
fof(f2398,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) )),
  inference(resolution,[],[f2377,f443])).
fof(f2377,plain,(
  ( ! [X4,X5] : (sP9(X5,X4,'F_0') | s(bool,'T_0') != s(bool,X4) | p(s(bool,X5))) )),
  inference(superposition,[],[f2369,f440])).
fof(f720,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14),s(sK12,sK16))))),
  inference(resolution,[],[f352,f355])).
fof(f592,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16)))) | p(s(bool,'const.relation.RC_3'(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15),s(sK12,sK16))))),
  inference(resolution,[],[f351,f355])).
% SZS output end Proof for thm.relation
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 1499
% Passive clauses: 17502
% Generated clauses: 60554
% Final active clauses: 604
% Final passive clauses: 8712
% Input formulas: 69
% Initial clauses: 235
% 
% Selected by SInE selection: 65
% SInE iterations: 1
% 
% Duplicate literals: 1887
% Trivial inequalities: 282
% Fw subsumption resolutions: 7201
% Fw demodulations: 9866
% Bw demodulations: 605
% 
% Simple tautologies: 4092
% Equational tautologies: 3
% Forward subsumptions: 20127
% 
% Binary resolution: 18675
% Forward superposition: 10048
% Backward superposition: 10711
% Equality resolution: 11
% 
% Split clauses: 1152
% Split components: 2639
% SAT solver clauses: 1307
% SAT solver unit clauses: 27
% SAT solver binary clauses: 350
% 
% Sat splits: 1152
% Sat splitting refutations: 47
% 
% Memory used [KB]: 14967
% Time elapsed: 1.206 s
% ------------------------------
% ------------------------------
% Success in time 12.941 s
