Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.arithmetic
% (82556)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 817
% Passive clauses: 6617
% Generated clauses: 19009
% Final active clauses: 763
% Final passive clauses: 5259
% Input formulas: 36
% Initial clauses: 157
% 
% Duplicate literals: 603
% Trivial inequalities: 130
% Fw subsumption resolutions: 2149
% Fw demodulations: 533
% Bw demodulations: 47
% 
% Simple tautologies: 225
% Equational tautologies: 74
% Forward subsumptions: 8116
% 
% Binary resolution: 8415
% Factoring: 1395
% Forward superposition: 3115
% Backward superposition: 2110
% Self superposition: 26
% Equality factoring: 28
% Equality resolution: 34
% 
% Split clauses: 544
% Split components: 1095
% SAT solver clauses: 992
% SAT solver unit clauses: 29
% SAT solver binary clauses: 136
% 
% Sat splits: 544
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 990
% TWLsolver calls for satisfiability: 493
% 
% Memory used [KB]: 7291
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.arithmetic
% (82580)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 467
% Passive clauses: 2310
% Generated clauses: 13236
% Final active clauses: 362
% Final passive clauses: 1765
% Input formulas: 36
% Initial clauses: 157
% 
% Duplicate literals: 278
% Trivial inequalities: 38
% Fw subsumption resolutions: 1264
% Fw demodulations: 832
% Bw demodulations: 111
% Condensations: 261
% Global subsumptions: 13
% 
% Simple tautologies: 253
% Equational tautologies: 54
% Forward subsumptions: 4617
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 4017
% Factoring: 82
% Forward superposition: 1139
% Backward superposition: 2074
% Self superposition: 32
% Equality factoring: 36
% Equality resolution: 9
% 
% Split clauses: 158
% Split components: 321
% SAT solver clauses: 3044
% SAT solver unit clauses: 171
% SAT solver binary clauses: 79
% SAT solver learnt clauses: 28
% SAT solver learnt literals: 15
% 
% Sat splits: 158
% Sat splitting refutations: 37
% 
% TWLsolver clauses: 2990
% TWLsolver calls for satisfiability: 1706
% 
% Memory used [KB]: 11129
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.arithmetic
% (82601)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1688
% Passive clauses: 5609
% Generated clauses: 109655
% Final active clauses: 465
% Final passive clauses: 574
% Input formulas: 36
% Initial clauses: 157
% 
% Duplicate literals: 6877
% Trivial inequalities: 392
% Fw subsumption resolutions: 7335
% Fw demodulations: 9607
% Bw demodulations: 2564
% Condensations: 11920
% Global subsumptions: 2048
% 
% Simple tautologies: 421
% Equational tautologies: 1584
% Forward subsumptions: 59388
% Bw demodulations to eq. taut.: 134
% 
% Unit resulting resolution: 15898
% Factoring: 1
% Forward superposition: 23434
% Backward superposition: 15477
% Self superposition: 499
% Equality factoring: 424
% Equality resolution: 50
% 
% Split clauses: 679
% Split components: 1391
% SAT solver clauses: 11672
% SAT solver unit clauses: 490
% SAT solver binary clauses: 1554
% 
% Sat splits: 679
% Sat splitting refutations: 2654
% 
% Memory used [KB]: 52962
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.arithmetic
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.arithmetic
% SZS output start Proof for thm.arithmetic
fof(f41923,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f359,f362,f365,f368,f371,f374,f381,f382,f383,f384,f385,f386,f387,f388,f389,f390,f421,f422,f505,f508,f511,f514,f517,f520,f567,f569,f698,f877,f880,f883,f886,f914,f915,f917,f1031,f1069,f1086,f1093,f1179,f1182,f1185,f551,f1188,f1191,f1225,f1227,f1238,f1239,f1255,f1276,f1369,f1372,f1375,f1378,f1381,f554,f1425,f1447,f1469,f1492,f1513,f1521,f1524,f1527,f1530,f1533,f909,f1555,f1558,f1557,f1575,f1596,f1744,f1747,f1750,f1753,f1756,f1782,f1784,f1556,f1785,f1864,f1867,f1870,f1778,f1873,f1876,f2096,f2566,f2593,f2572,f2596,f2599,f2598,f2677,f2697,f2846,f2093,f2092,f1861,f1776,f1518,f1366,f1283,f544,f392,f2094,f543,f539,f1282,f542,f540,f2944,f2947,f2953,f2721,f2731,f2730,f2957,f2717,f2707,f2959,f2703,f2425,f1849,f2961,f1845,f1835,f2963,f1831,f1817,f2966,f1816,f2968,f1801,f2971,f1786,f2974,f2972,f1727,f2976,f2969,f1713,f2461,f2477,f1814,f2977,f2496,f1800,f2978,f2516,f2637,f2657,f3068,f3122,f3123,f3124,f2536,f2557,f2615,f3256,f3259,f3265,f3272,f3273,f3274,f3275,f3276,f3277,f3278,f3279,f3280,f3281,f3296,f3429,f3297,f3431,f398,f3825,f3705,f3832,f541,f3805,f3834,f3885,f3886,f3945,f3947,f3953,f3897,f3957,f3541,f3544,f3965,f3966,f3969,f3985,f1072,f4003,f1089,f4008,f393,f391,f4011,f395,f4014,f399,f4017,f401,f4020,f4023,f4030,f4031,f411,f4032,f4033,f412,f4034,f4035,f410,f4036,f4037,f409,f4038,f416,f418,f4039,f414,f420,f4134,f1052,f4275,f397,f4251,f402,f4281,f2950,f4399,f4376,f1051,f4420,f1260,f4440,f1281,f4599,f1431,f4620,f1453,f4641,f1475,f4663,f1496,f4683,f1517,f4687,f3430,f2951,f4850,f1579,f4870,f1601,f5100,f5353,f5513,f5540,f5548,f5897,f5956,f6014,f6122,f6680,f6681,f6723,f6748,f6824,f6809,f6810,f6831,f6832,f6841,f6850,f6856,f6865,f6891,f6875,f6895,f7019,f7050,f7034,f7053,f7271,f7310,f7403,f3833,f7453,f7482,f7532,f7514,f7540,f7548,f7577,f7606,f7635,f7666,f7714,f6892,f7696,f7720,f7729,f7757,f7816,f3289,f7817,f8006,f3283,f8027,f8028,f8041,f8029,f8030,f8042,f8031,f8032,f8040,f8033,f8034,f8039,f8035,f8036,f8533,f3282,f8560,f8561,f8628,f8045,f908,f8629,f1780,f8945,f3285,f396,f8974,f8975,f10299,f10341,f10302,f3288,f10342,f10305,f10306,f10346,f10349,f10347,f10391,f10392,f10389,f10387,f7815,f7814,f7813,f7812,f7811,f7810,f3287,f550,f7809,f7808,f3286,f549,f10579,f10581,f1554,f10588,f10784,f10589,f10785,f10590,f10786,f10591,f10787,f10592,f10788,f10593,f10789,f10594,f10790,f10595,f10791,f10596,f10792,f10597,f7474,f10598,f10793,f10599,f10794,f10600,f10795,f10601,f10796,f10602,f10797,f10603,f7503,f10604,f10798,f10605,f10799,f10606,f10800,f10607,f10801,f10608,f10802,f10609,f7569,f10610,f7598,f10611,f7626,f10614,f10803,f10617,f10804,f10620,f10805,f10623,f10806,f10649,f3291,f10650,f10651,f10652,f10653,f10654,f10834,f10657,f10835,f10658,f10836,f10659,f10837,f10660,f10838,f10661,f10839,f10662,f10840,f10663,f10841,f10664,f7657,f10665,f10842,f10666,f10843,f10667,f10844,f10668,f10845,f10669,f10846,f10670,f10847,f10671,f10848,f10672,f10849,f10673,f10850,f10674,f10851,f10675,f10852,f10676,f10831,f10677,f10830,f10678,f10832,f10679,f10833,f10680,f10681,f10682,f10683,f10685,f10857,f10687,f10858,f10688,f10859,f10689,f10860,f10690,f10861,f10691,f10862,f10692,f10863,f10693,f7685,f10694,f7748,f10695,f10864,f10696,f10865,f10697,f10866,f10698,f10867,f10699,f10868,f10700,f10869,f10701,f10870,f10702,f10871,f10703,f10872,f10704,f10873,f10705,f10874,f10706,f10853,f10707,f10854,f10708,f10855,f10709,f10856,f10710,f10711,f10712,f10713,f10714,f10879,f10715,f10880,f10716,f10881,f10718,f10882,f10720,f10883,f10721,f10884,f10722,f10885,f10723,f7780,f10724,f10886,f10725,f10887,f10726,f10888,f10727,f10889,f10728,f10890,f10729,f10891,f10730,f10892,f10731,f10893,f10732,f10894,f10733,f10875,f10734,f10876,f10735,f10877,f10736,f10878,f10737,f10895,f10738,f10896,f10739,f10897,f10740,f10741,f10742,f10743,f10744,f10745,f10746,f10747,f10748,f10904,f10749,f10905,f10750,f10902,f10751,f10903,f10752,f10901,f10753,f10900,f10754,f10899,f10755,f10898,f10906,f10908,f10758,f10759,f10910,f10907,f10911,f10909,f10912,f6742,f10763,f6767,f10771,f10913,f10772,f10914,f10773,f10915,f6835,f6821,f6820,f6834,f6833,f11019,f11016,f11007,f11006,f11086,f11087,f11088,f11089,f11090,f11091,f11630,f11629,f11633,f1777,f6825,f11618,f11619,f6830,f6829,f11638,f11639,f11640,f11641,f11642,f1553,f11681,f11682,f6828,f11749,f11750,f11789,f11791,f11790,f11833,f11834,f11835,f11836,f11851,f11854,f11857,f11858,f11893,f11895,f11894,f11903,f11906,f11909,f11910,f11943,f11945,f11944,f12139,f12142,f12145,f12146,f12175,f12177,f12176,f12336,f12339,f12342,f12343,f12370,f12372,f12371,f12379,f12380,f10343,f12464,f12465,f11788,f12637,f12638,f12827,f12828,f13038,f13039,f12172,f13125,f13126,f12369,f13843,f13844,f13883,f13885,f13884,f13919,f10344,f13920,f10345,f13959,f13961,f13960,f14110,f14113,f14116,f14117,f14152,f14154,f14153,f14302,f14305,f14308,f11891,f14309,f11892,f14344,f14346,f14345,f14501,f14504,f14507,f14508,f14537,f14539,f14538,f14608,f14611,f8982,f14614,f12173,f14615,f12174,f14644,f14646,f10348,f14645,f14653,f14654,f13882,f14696,f12415,f14697,f12416,f13958,f14766,f14767,f14828,f12664,f14829,f12665,f14889,f14890,f14536,f14947,f13061,f14948,f13062,f14643,f16883,f8044,f1781,f17186,f17189,f17188,f14944,f14642,f14941,f14641,f14886,f14535,f14883,f14534,f13122,f12368,f13119,f12367,f13035,f12171,f13032,f407,f12170,f17222,f1779,f8983,f17223,f907,f17799,f17800,f17797,f17795,f17808,f17810,f553,f19747,f19548,f11890,f19547,f11942,f19546,f14151,f19545,f14343,f19884,f19882,f19857,f19833,f19831,f19793,f19791,f19786,f19784,f19922,f14342,f19923,f14341,f19924,f14150,f19925,f14149,f19926,f11941,f19927,f11940,f19928,f11889,f19929,f8046,f11888,f19931,f21663,f21676,f20825,f20822,f20794,f20791,f20763,f20760,f20729,f20726,f20314,f20311,f20282,f20279,f20250,f20247,f20218,f20215,f20186,f20183,f20151,f20148,f2595,f20099,f20096,f19979,f19976,f19947,f19944,f8038,f19523,f19520,f2594,f19097,f19094,f21697,f15682,f15679,f21698,f15656,f15653,f21699,f15629,f15626,f21700,f15603,f15600,f21701,f15577,f15574,f21702,f15549,f15546,f21703,f15523,f15520,f21704,f15496,f15493,f21705,f15471,f15468,f21706,f14819,f14816,f21707,f14757,f14754,f21708,f14711,f14708,f21709,f14668,f14665,f21710,f14317,f14314,f21711,f14234,f14231,f21712,f14206,f14203,f21713,f14177,f14174,f21714,f14125,f14122,f21715,f14098,f14095,f21716,f14070,f14067,f21717,f14040,f14037,f21718,f14012,f14009,f21719,f13983,f13980,f21720,f13934,f13931,f21721,f13907,f13904,f21722,f13858,f13855,f21723,f13826,f13823,f21724,f13797,f13794,f21725,f13768,f13765,f21726,f13736,f13733,f21727,f13706,f13703,f21728,f13676,f13673,f21729,f13646,f13643,f21730,f13616,f13613,f21731,f13586,f13583,f21732,f13244,f13241,f21734,f12818,f12815,f21735,f12700,f12697,f21683,f12680,f12677,f21684,f12628,f12625,f21736,f12598,f12595,f21685,f12578,f12575,f21686,f12556,f12553,f21687,f12535,f12532,f21688,f12514,f12511,f21689,f12479,f12476,f21737,f12451,f12448,f21690,f12431,f12428,f21691,f12394,f12391,f21738,f11918,f11915,f21739,f11866,f11863,f21740,f11764,f11761,f21741,f11734,f11731,f21742,f11704,f11701,f21743,f11655,f11652,f21745,f11605,f11602,f21746,f11173,f11170,f21748,f10367,f10364,f21750,f10320,f10317,f21751,f7708,f7705,f21752,f7526,f7523,f21753,f7285,f7282,f21693,f7046,f7043,f21754,f6887,f6884,f21755,f6862,f6859,f21694,f6847,f6844,f21695,f6134,f6131,f8985,f5519,f408,f5516,f21692,f4833,f405,f403,f23239,f8043,f2958,f23253,f23254,f1783,f552,f23300,f23301,f23302,f23303,f23304,f23305,f23306,f24211,f1226,f24230,f41922])).
fof(f41922,plain,(
  $false | ($spl2 | $spl30 | $spl42)),
  inference(subsumption_resolution,[],[f5779,f543])).
fof(f5779,plain,(
  s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',sK14))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',sK14))) | ($spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5778,f393])).
fof(f5778,plain,(
  s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',sK14))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',sK14))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) | $spl30),
  inference(subsumption_resolution,[],[f5579,f328])).
fof(f328,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f188])).
fof(f188,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f187])).
fof(f187,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f74])).
fof(f74,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X13] : ((p(s(bool,X13)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X13)))),
  file('thm.arithmetic.MODEQ_SYM','thm.bool.IMP_CLAUSES|split|4')).
fof(f5579,plain,(
  p(s(bool,'F_0')) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',sK14))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',sK14))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) | $spl30),
  inference(superposition,[],[f242,f413])).
fof(f413,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK13))) | $spl30),
  inference(cnf_transformation,[],[f413_D])).
fof(f413_D,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK13))) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f242,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num',X2),s('type.num.num',X1)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X1),s('type.num.num',X0))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X2),s('type.num.num',X0))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f151])).
fof(f151,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num',X2),s('type.num.num',X1)))) | (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num',X1) = s('type.num.num',X2)) | (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) & s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X1),s('type.num.num',X0))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X2),s('type.num.num',X0))))) & (((s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num',X1) != s('type.num.num',X2)) & (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X1),s('type.num.num',X0))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X2),s('type.num.num',X0))))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num',X2),s('type.num.num',X1))))))),
  inference(flattening,[],[f150])).
fof(f150,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num',X2),s('type.num.num',X1)))) | ((s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num',X1) = s('type.num.num',X2)) | (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) & s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X1),s('type.num.num',X0))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X2),s('type.num.num',X0)))))) & (((s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num',X1) != s('type.num.num',X2)) & (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X1),s('type.num.num',X0))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X2),s('type.num.num',X0))))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num',X2),s('type.num.num',X1))))))),
  inference(nnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num',X2),s('type.num.num',X1)))) <=> ((s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num',X1) = s('type.num.num',X2)) | (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) & s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X1),s('type.num.num',X0))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X2),s('type.num.num',X0))))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X8,X9,X10] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X8),s('type.num.num',X10),s('type.num.num',X9)))) <=> ((s('type.num.num','const.num.|48|_0') = s('type.num.num',X8) & s('type.num.num',X9) = s('type.num.num',X10)) | (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X8)))) & s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X9),s('type.num.num',X8))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X10),s('type.num.num',X8))))))),
  file('thm.arithmetic.MODEQ_SYM','thm.arithmetic.MODEQ_THM')).
fof(f24230,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK12))) | $spl661),
  inference(cnf_transformation,[],[f24230_D])).
fof(f24230_D,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK12))) <=> ~$spl661),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl661])])).
fof(f1226,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl62),
  inference(cnf_transformation,[],[f1226_D])).
fof(f1226_D,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f24211,plain,(
  ( ! [X21] : (p(s(bool,X21)) | s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK12)))) ) | ($spl24 | $spl30 | $spl98)),
  inference(resolution,[],[f17016,f5439])).
fof(f5439,plain,(
  ( ! [X23,X22] : (sP0('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK12)),X22,X23)) ) | $spl98),
  inference(resolution,[],[f2597,f218])).
fof(f218,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f144])).
fof(f144,plain,(
  ! [X0,X1,X2] : ((~sP0(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP0(X0,X1,X2)))),
  inference(flattening,[],[f143])).
fof(f143,plain,(
  ! [X0,X1,X2] : ((~sP0(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP0(X0,X1,X2)))),
  inference(nnf_transformation,[],[f115])).
fof(f115,plain,(
  ! [X0,X1,X2] : (sP0(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f2597,plain,(
  p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK12)))) | $spl98),
  inference(cnf_transformation,[],[f2597_D])).
fof(f2597_D,plain,(
  p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK12)))) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f17016,plain,(
  ( ! [X2,X3] : (~sP0(X3,'F_0',X2) | p(s(bool,X2)) | s(bool,'F_0') != s(bool,X3)) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f17007,f210])).
fof(f210,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f139])).
fof(f139,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(rectify,[],[f138])).
fof(f138,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(flattening,[],[f137])).
fof(f137,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(nnf_transformation,[],[f117])).
fof(f117,plain,(
  ! [X2,X0,X1] : (sP2(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f17007,plain,(
  ( ! [X2,X3] : (~sP2(X2,X3,'F_0') | ~sP0(X3,'F_0',X2) | p(s(bool,X2)) | s(bool,'F_0') != s(bool,X3)) ) | ($spl24 | $spl30)),
  inference(resolution,[],[f16990,f222])).
fof(f222,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | s(bool,X0) != s(bool,X1) | ~sP3(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP3(X2,X1,X0)) & (~sP3(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP3(X2,X1,X0))),
  inference(definition_folding,[],[f40,f118,f117,f116,f115])).
fof(f116,plain,(
  ! [X1,X0,X2] : (sP1(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f118,plain,(
  ! [X2,X1,X0] : (sP3(X2,X1,X0) <=> (sP0(X0,X1,X2) & sP1(X1,X0,X2) & sP2(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f40,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X16,X17,X18] : ((p(s(bool,X18)) <=> s(bool,X16) = s(bool,X17)) <=> ((p(s(bool,X18)) | p(s(bool,X17)) | p(s(bool,X16))) & (p(s(bool,X18)) | ~p(s(bool,X16)) | ~p(s(bool,X17))) & (p(s(bool,X17)) | ~p(s(bool,X16)) | ~p(s(bool,X18))) & (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)))))),
  file('thm.arithmetic.MODEQ_SYM','thm.sat.sat_2|49|_prove_ltb')).
fof(f16990,plain,(
  ( ! [X4,X5] : (sP3(X4,'F_0',X5) | ~sP2(X4,X5,'F_0') | ~sP0(X5,'F_0',X4)) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f16988,f348])).
fof(f348,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f86])).
fof(f86,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.arithmetic.MODEQ_SYM',aHL_FALSITY)).
fof(f16988,plain,(
  ( ! [X4,X5] : (sP3(X4,'F_0',X5) | p(s(bool,'F_0')) | ~sP2(X4,X5,'F_0') | ~sP0(X5,'F_0',X4)) ) | ($spl24 | $spl30)),
  inference(resolution,[],[f16985,f205])).
fof(f205,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X1)) | ~sP2(X0,X2,X1) | ~sP1(X1,X2,X0) | ~sP0(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f136])).
fof(f136,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | (sP0(X2,X1,X0) & sP1(X1,X2,X0) & sP2(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP0(X2,X1,X0) | ~sP1(X1,X2,X0) | ~sP2(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f135])).
fof(f135,plain,(
  ! [X2,X1,X0] : ((~sP3(X2,X1,X0) | (sP0(X0,X1,X2) & sP1(X1,X0,X2) & sP2(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP0(X0,X1,X2) | ~sP1(X1,X0,X2) | ~sP2(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP3(X2,X1,X0)))),
  inference(flattening,[],[f134])).
fof(f134,plain,(
  ! [X2,X1,X0] : ((~sP3(X2,X1,X0) | (sP0(X0,X1,X2) & sP1(X1,X0,X2) & sP2(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP0(X0,X1,X2) | ~sP1(X1,X0,X2) | ~sP2(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP3(X2,X1,X0)))),
  inference(nnf_transformation,[],[f118])).
fof(f16985,plain,(
  ( ! [X4,X5] : (sP1('F_0',X4,X5)) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f16965,f212])).
fof(f212,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f142])).
fof(f142,plain,(
  ! [X0,X1,X2] : ((~sP1(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1,X2)))),
  inference(rectify,[],[f141])).
fof(f141,plain,(
  ! [X1,X0,X2] : ((~sP1(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0,X2)))),
  inference(flattening,[],[f140])).
fof(f140,plain,(
  ! [X1,X0,X2] : ((~sP1(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0,X2)))),
  inference(nnf_transformation,[],[f116])).
fof(f16965,plain,(
  ( ! [X4,X5] : (sP1('F_0',X4,X5) | p(s(bool,X5))) ) | ($spl24 | $spl30)),
  inference(resolution,[],[f16960,f210])).
fof(f16960,plain,(
  ( ! [X6,X7] : (~sP2(X7,X6,'F_0') | sP1('F_0',X6,X7)) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f16953,f213])).
fof(f213,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f142])).
fof(f16953,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | ~sP2(X7,X6,'F_0') | sP1('F_0',X6,X7)) ) | ($spl24 | $spl30)),
  inference(resolution,[],[f16947,f201])).
fof(f201,plain,(
  ( ! [X2,X0,X1] : (sP1(X1,X2,X0) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f136])).
fof(f16947,plain,(
  ( ! [X2,X3] : (sP3(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP2(X3,X2,'F_0')) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f16946,f233])).
fof(f233,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f148])).
fof(f148,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f48])).
fof(f48,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X13] : (s(bool,'T_0') = s(bool,X13) <=> p(s(bool,X13)))),
  file('thm.arithmetic.MODEQ_SYM','thm.bool.EQ_CLAUSES|split|0')).
fof(f16946,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP3(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP2(X3,X2,'F_0')) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f16943,f16933])).
fof(f16933,plain,(
  ( ! [X4,X5] : (sP0(X4,'F_0',X5) | s(bool,'T_0') != s(bool,X4)) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f16927,f216])).
fof(f216,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f144])).
fof(f16927,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | p(s(bool,X5)) | sP0(X4,'F_0',X5)) ) | ($spl24 | $spl30)),
  inference(resolution,[],[f5938,f200])).
fof(f200,plain,(
  ( ! [X2,X0,X1] : (sP0(X2,X1,X0) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f136])).
fof(f5938,plain,(
  ( ! [X6,X5] : (sP3(X6,'F_0',X5) | s(bool,'T_0') != s(bool,X5) | p(s(bool,X6))) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f5931,f219])).
fof(f219,plain,(
  ( ! [X2,X0,X1] : (sP3(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f145])).
fof(f5931,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f5922,f410])).
fof(f5922,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',sK12)))) | $spl30),
  inference(superposition,[],[f5544,f233])).
fof(f5544,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',sK12))) | $spl30),
  inference(backward_demodulation,[],[f413,f199])).
fof(f199,plain,(
  s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',sK12)))),
  inference(cnf_transformation,[],[f133])).
fof(f133,plain,(
  s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',sK12)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13,sK14])],[f89])).
fof(f89,plain,(
  ? [X0,X1,X2] : s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X2),s('type.num.num',X0),s('type.num.num',X1))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X2),s('type.num.num',X1),s('type.num.num',X0)))),
  inference(ennf_transformation,[],[f38])).
fof(f38,plain,(
  ~! [X0,X1,X2] : s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X2),s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X2),s('type.num.num',X1),s('type.num.num',X0)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X8),s('type.num.num',X6),s('type.num.num',X7))) = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X8),s('type.num.num',X7),s('type.num.num',X6)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X8),s('type.num.num',X6),s('type.num.num',X7))) = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X8),s('type.num.num',X7),s('type.num.num',X6)))),
  file('thm.arithmetic.MODEQ_SYM','thm.arithmetic.MODEQ_SYM')).
fof(f16943,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP3(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP2(X3,X2,'F_0') | ~sP0(X2,'F_0',X3)) ) | ($spl24 | $spl30)),
  inference(resolution,[],[f16934,f204])).
fof(f204,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2)) | ~sP2(X0,X2,X1) | ~sP1(X1,X2,X0) | ~sP0(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f136])).
fof(f16934,plain,(
  ( ! [X6,X7] : (sP1('F_0',X6,X7) | s(bool,'T_0') != s(bool,X6)) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f16928,f212])).
fof(f16928,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X6) | p(s(bool,X7)) | sP1('F_0',X6,X7)) ) | ($spl24 | $spl30)),
  inference(resolution,[],[f5938,f201])).
fof(f23306,plain,(
  $false | ($spl90 | $spl92)),
  inference(trivial_inequality_removal,[],[f23290])).
fof(f23290,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl90 | $spl92)),
  inference(superposition,[],[f1783,f1780])).
fof(f23305,plain,(
  $false | ($spl36 | $spl92)),
  inference(trivial_inequality_removal,[],[f23291])).
fof(f23291,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl36 | $spl92)),
  inference(superposition,[],[f1783,f420])).
fof(f23304,plain,(
  $false | ($spl92 | $spl132)),
  inference(trivial_inequality_removal,[],[f23293])).
fof(f23293,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl92 | $spl132)),
  inference(superposition,[],[f1783,f3291])).
fof(f23303,plain,(
  $false | ($spl56 | $spl92)),
  inference(trivial_inequality_removal,[],[f23296])).
fof(f23296,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl56 | $spl92)),
  inference(superposition,[],[f1783,f10440])).
fof(f10440,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | $spl56),
  inference(resolution,[],[f1052,f233])).
fof(f23302,plain,(
  $false | ($spl18 | $spl92)),
  inference(trivial_inequality_removal,[],[f23297])).
fof(f23297,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl18 | $spl92)),
  inference(superposition,[],[f1783,f21795])).
fof(f21795,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) | $spl18),
  inference(resolution,[],[f405,f233])).
fof(f23301,plain,(
  $false | ($spl2 | $spl92)),
  inference(trivial_inequality_removal,[],[f23299])).
fof(f23299,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl2 | $spl92)),
  inference(superposition,[],[f1783,f4059])).
fof(f4059,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14))) | $spl2),
  inference(resolution,[],[f393,f233])).
fof(f23300,plain,(
  $false | $spl92),
  inference(equality_resolution,[],[f1783])).
fof(f552,plain,(
  ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | $spl47),
  inference(cnf_transformation,[],[f552_D])).
fof(f552_D,plain,(
  ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) <=> ~$spl47),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl47])])).
fof(f1783,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,X10)) ) | $spl92),
  inference(cnf_transformation,[],[f1783_D])).
fof(f1783_D,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,X10)) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f23254,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,X10) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))))) ) | ($spl13 | $spl90)),
  inference(subsumption_resolution,[],[f20160,f232])).
fof(f232,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f148])).
fof(f20160,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,X10) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,X10))) ) | ($spl13 | $spl90)),
  inference(superposition,[],[f19932,f226])).
fof(f226,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f91])).
fof(f91,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f90])).
fof(f90,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f41])).
fof(f41,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X11,X12] : ((p(s(bool,X11)) => p(s(bool,X12))) => ((p(s(bool,X12)) => p(s(bool,X11))) => s(bool,X11) = s(bool,X12)))),
  file('thm.arithmetic.MODEQ_SYM','thm.bool.IMP_ANTISYM_AX')).
fof(f19932,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) | ($spl13 | $spl90)),
  inference(forward_demodulation,[],[f401,f1780])).
fof(f23253,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,X11) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))))) ) | ($spl13 | $spl90)),
  inference(subsumption_resolution,[],[f20161,f232])).
fof(f20161,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,X11) | ~p(s(bool,X11)) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))))) ) | ($spl13 | $spl90)),
  inference(superposition,[],[f19932,f226])).
fof(f2958,plain,(
  ( ! [X1] : (~p(s(bool,X1))) ) | $spl104),
  inference(cnf_transformation,[],[f2958_D])).
fof(f2958_D,plain,(
  ( ! [X1] : (~p(s(bool,X1))) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f8043,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) | $spl196),
  inference(cnf_transformation,[],[f8043_D])).
fof(f8043_D,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f23239,plain,(
  ( ! [X13] : (s(bool,'F_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) | ~p(s(bool,X13))) ) | ($spl24 | $spl30 | $spl47)),
  inference(resolution,[],[f16958,f7190])).
fof(f7190,plain,(
  ( ! [X8,X9] : (sP2(X8,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)),X9)) ) | $spl47),
  inference(resolution,[],[f552,f209])).
fof(f209,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f139])).
fof(f16958,plain,(
  ( ! [X0,X1] : (~sP2(X1,X0,'F_0') | s(bool,'F_0') = s(bool,X0) | ~p(s(bool,X1))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f16950,f229])).
fof(f229,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f45])).
fof(f45,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f44])).
fof(f44,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X13] : (s(bool,'F_0') = s(bool,X13) <=> ~p(s(bool,X13)))),
  file('thm.arithmetic.MODEQ_SYM','thm.bool.EQ_CLAUSES|split|3')).
fof(f16950,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~sP2(X1,X0,'F_0') | s(bool,'F_0') = s(bool,X0) | ~p(s(bool,X1))) ) | ($spl24 | $spl30)),
  inference(resolution,[],[f16947,f221])).
fof(f221,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X2)) | ~sP3(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f145])).
fof(f403,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X0),s('type.num.num',X1))))) ) | $spl16),
  inference(cnf_transformation,[],[f403_D])).
fof(f403_D,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X0),s('type.num.num',X1))))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f405,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | $spl18),
  inference(cnf_transformation,[],[f405_D])).
fof(f405_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f4833,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X0),s('type.num.num',X1))))) ) | $spl101),
  inference(trivial_inequality_removal,[],[f4826])).
fof(f4826,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.|48|_0') | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X0),s('type.num.num',X1))))) ) | $spl101),
  inference(superposition,[],[f2951,f237])).
fof(f237,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num',X2),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f151])).
fof(f21692,plain,(
  p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | $spl544),
  inference(cnf_transformation,[],[f21692_D])).
fof(f21692_D,plain,(
  p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) <=> ~$spl544),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl544])])).
fof(f5516,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X0),s('type.num.num',X1))))) ) | $spl98),
  inference(superposition,[],[f2597,f237])).
fof(f408,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X7),s('type.num.num',X6)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK12)))) ) | $spl22),
  inference(cnf_transformation,[],[f408_D])).
fof(f408_D,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X7),s('type.num.num',X6)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK12)))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f5519,plain,(
  ( ! [X6,X7] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X7),s('type.num.num',X6))))) ) | $spl98),
  inference(superposition,[],[f2597,f239])).
fof(f239,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X1),s('type.num.num',X0))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X2),s('type.num.num',X0))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num',X2),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f151])).
fof(f8985,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | $spl202),
  inference(cnf_transformation,[],[f8985_D])).
fof(f8985_D,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f6131,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X0),s('type.num.num',X1))))) ) | $spl36),
  inference(superposition,[],[f420,f237])).
fof(f6134,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X7),s('type.num.num',X6))))) ) | $spl36),
  inference(superposition,[],[f420,f239])).
fof(f21695,plain,(
  ( ! [X9] : (sP11(X9,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl548),
  inference(cnf_transformation,[],[f21695_D])).
fof(f21695_D,plain,(
  ( ! [X9] : (sP11(X9,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl548),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl548])])).
fof(f6844,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X1),s('type.num.num',X2))))) ) | $spl98),
  inference(superposition,[],[f5501,f237])).
fof(f5501,plain,(
  ( ! [X128] : (sP11(X128,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK12)))) ) | $spl98),
  inference(resolution,[],[f2597,f311])).
fof(f311,plain,(
  ( ! [X0,X1] : (sP11(X0,X1) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f182])).
fof(f182,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP11(X0,X1)))),
  inference(flattening,[],[f181])).
fof(f181,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP11(X0,X1)))),
  inference(nnf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0,X1] : (sP11(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f6847,plain,(
  ( ! [X10,X11,X9] : (sP11(X9,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X10),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X11),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X11),s('type.num.num',X10))))) ) | $spl98),
  inference(superposition,[],[f5501,f239])).
fof(f21694,plain,(
  ( ! [X9] : (sP11('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X9)) ) | $spl546),
  inference(cnf_transformation,[],[f21694_D])).
fof(f21694_D,plain,(
  ( ! [X9] : (sP11('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X9)) ) <=> ~$spl546),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl546])])).
fof(f6859,plain,(
  ( ! [X2,X0,X1] : (sP11('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X0) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X1),s('type.num.num',X2))))) ) | $spl98),
  inference(superposition,[],[f5502,f237])).
fof(f5502,plain,(
  ( ! [X129] : (sP11('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK12)),X129)) ) | $spl98),
  inference(resolution,[],[f2597,f312])).
fof(f312,plain,(
  ( ! [X0,X1] : (sP11(X0,X1) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f182])).
fof(f6862,plain,(
  ( ! [X10,X11,X9] : (sP11('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X9) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X10),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X11),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X11),s('type.num.num',X10))))) ) | $spl98),
  inference(superposition,[],[f5502,f239])).
fof(f21755,plain,(
  ( ! [X12,X13] : (sP3(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X12)) | s(bool,'F_0') != s(bool,X13)) ) | $spl658),
  inference(cnf_transformation,[],[f21755_D])).
fof(f21755_D,plain,(
  ( ! [X12,X13] : (sP3(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X12)) | s(bool,'F_0') != s(bool,X13)) ) <=> ~$spl658),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl658])])).
fof(f6884,plain,(
  ( ! [X2,X0,X3,X1] : (sP3(X0,X1,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') != s(bool,X1) | p(s(bool,X0)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl30),
  inference(superposition,[],[f5912,f237])).
fof(f5912,plain,(
  ( ! [X0,X1] : (sP3(X1,X0,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',sK12))) | s(bool,'F_0') != s(bool,X0) | p(s(bool,X1))) ) | $spl30),
  inference(superposition,[],[f5544,f219])).
fof(f6887,plain,(
  ( ! [X14,X12,X15,X13] : (sP3(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') != s(bool,X13) | p(s(bool,X12)) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl30),
  inference(superposition,[],[f5912,f239])).
fof(f21754,plain,(
  ( ! [X12,X13] : (sP3(X12,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')),X13) | p(s(bool,X12)) | s(bool,'F_0') != s(bool,X13)) ) | $spl656),
  inference(cnf_transformation,[],[f21754_D])).
fof(f21754_D,plain,(
  ( ! [X12,X13] : (sP3(X12,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')),X13) | p(s(bool,X12)) | s(bool,'F_0') != s(bool,X13)) ) <=> ~$spl656),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl656])])).
fof(f7043,plain,(
  ( ! [X2,X0,X3,X1] : (sP3(X0,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')),X1) | s(bool,'F_0') != s(bool,X1) | p(s(bool,X0)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl30),
  inference(superposition,[],[f5913,f237])).
fof(f5913,plain,(
  ( ! [X2,X3] : (sP3(X3,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',sK12)),X2) | s(bool,'F_0') != s(bool,X2) | p(s(bool,X3))) ) | $spl30),
  inference(superposition,[],[f5544,f219])).
fof(f7046,plain,(
  ( ! [X14,X12,X15,X13] : (sP3(X12,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')),X13) | s(bool,'F_0') != s(bool,X13) | p(s(bool,X12)) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl30),
  inference(superposition,[],[f5913,f239])).
fof(f21693,plain,(
  ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | $spl545),
  inference(cnf_transformation,[],[f21693_D])).
fof(f21693_D,plain,(
  ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) <=> ~$spl545),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl545])])).
fof(f7282,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X0),s('type.num.num',X1))))) ) | $spl47),
  inference(superposition,[],[f552,f237])).
fof(f7285,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X7),s('type.num.num',X6))))) ) | $spl47),
  inference(superposition,[],[f552,f239])).
fof(f21753,plain,(
  ( ! [X12,X13] : (~sP3(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') != s(bool,X13) | ~p(s(bool,X12))) ) | $spl654),
  inference(cnf_transformation,[],[f21753_D])).
fof(f21753_D,plain,(
  ( ! [X12,X13] : (~sP3(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') != s(bool,X13) | ~p(s(bool,X12))) ) <=> ~$spl654),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl654])])).
fof(f7523,plain,(
  ( ! [X2,X0,X3,X1] : (~sP3(X0,X1,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X1) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl30),
  inference(superposition,[],[f5914,f237])).
fof(f5914,plain,(
  ( ! [X4,X5] : (~sP3(X5,X4,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',sK12))) | ~p(s(bool,X5)) | s(bool,'F_0') != s(bool,X4)) ) | $spl30),
  inference(superposition,[],[f5544,f221])).
fof(f7526,plain,(
  ( ! [X14,X12,X15,X13] : (~sP3(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X12)) | s(bool,'F_0') != s(bool,X13) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl30),
  inference(superposition,[],[f5914,f239])).
fof(f21752,plain,(
  ( ! [X12,X13] : (~sP3(X12,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')),X13) | s(bool,'F_0') != s(bool,X13) | ~p(s(bool,X12))) ) | $spl652),
  inference(cnf_transformation,[],[f21752_D])).
fof(f21752_D,plain,(
  ( ! [X12,X13] : (~sP3(X12,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')),X13) | s(bool,'F_0') != s(bool,X13) | ~p(s(bool,X12))) ) <=> ~$spl652),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl652])])).
fof(f7705,plain,(
  ( ! [X2,X0,X3,X1] : (~sP3(X0,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')),X1) | ~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X1) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl30),
  inference(superposition,[],[f5915,f237])).
fof(f5915,plain,(
  ( ! [X6,X7] : (~sP3(X7,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',sK12)),X6) | ~p(s(bool,X7)) | s(bool,'F_0') != s(bool,X6)) ) | $spl30),
  inference(superposition,[],[f5544,f221])).
fof(f7708,plain,(
  ( ! [X14,X12,X15,X13] : (~sP3(X12,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')),X13) | ~p(s(bool,X12)) | s(bool,'F_0') != s(bool,X13) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl30),
  inference(superposition,[],[f5915,f239])).
fof(f21751,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',sK14),s('type.num.num',X12))))) ) | $spl650),
  inference(cnf_transformation,[],[f21751_D])).
fof(f21751_D,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',sK14),s('type.num.num',X12))))) ) <=> ~$spl650),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl650])])).
fof(f10317,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X1),s('type.num.num',sK14),s('type.num.num',X0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl30),
  inference(superposition,[],[f5895,f237])).
fof(f5895,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X2),s('type.num.num',sK13),s('type.num.num',sK12))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X3)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X3),s('type.num.num',sK14),s('type.num.num',X2))))) ) | $spl30),
  inference(superposition,[],[f5544,f238])).
fof(f238,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | s('type.num.num',X1) = s('type.num.num',X2) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num',X2),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f151])).
fof(f10320,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',sK14),s('type.num.num',X12)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl30),
  inference(superposition,[],[f5895,f239])).
fof(f21750,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | $spl649),
  inference(cnf_transformation,[],[f21750_D])).
fof(f21750_D,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) <=> ~$spl649),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl649])])).
fof(f10364,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X0),s('type.num.num',X1))))) ) | $spl131),
  inference(superposition,[],[f3288,f237])).
fof(f10367,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X7),s('type.num.num',X6))))) ) | $spl131),
  inference(superposition,[],[f3288,f239])).
fof(f21748,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | $spl646),
  inference(cnf_transformation,[],[f21748_D])).
fof(f21748_D,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) <=> ~$spl646),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl646])])).
fof(f11170,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X0),s('type.num.num',X1))))) ) | $spl132),
  inference(superposition,[],[f3291,f237])).
fof(f11173,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X7),s('type.num.num',X6))))) ) | $spl132),
  inference(superposition,[],[f3291,f239])).
fof(f21746,plain,(
  ( ! [X9] : (s(bool,X9) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) | ~p(s(bool,X9))) ) | $spl644),
  inference(cnf_transformation,[],[f21746_D])).
fof(f21746_D,plain,(
  ( ! [X9] : (s(bool,X9) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) | ~p(s(bool,X9))) ) <=> ~$spl644),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl644])])).
fof(f11602,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) | ~p(s(bool,X0)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X1),s('type.num.num',X2))))) ) | $spl128),
  inference(superposition,[],[f3287,f237])).
fof(f11605,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) | ~p(s(bool,X9)) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X10),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X11),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X11),s('type.num.num',X10))))) ) | $spl128),
  inference(superposition,[],[f3287,f239])).
fof(f21745,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) | $spl643),
  inference(cnf_transformation,[],[f21745_D])).
fof(f21745_D,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) <=> ~$spl643),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl643])])).
fof(f11652,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X0),s('type.num.num',X1))))) ) | ($spl5 | $spl132)),
  inference(superposition,[],[f11643,f237])).
fof(f11643,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) | ($spl5 | $spl132)),
  inference(forward_demodulation,[],[f395,f3291])).
fof(f11655,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X7),s('type.num.num',X6))))) ) | ($spl5 | $spl132)),
  inference(superposition,[],[f11643,f239])).
fof(f21743,plain,(
  ( ! [X9] : (sP11(X9,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')))) ) | $spl640),
  inference(cnf_transformation,[],[f21743_D])).
fof(f21743_D,plain,(
  ( ! [X9] : (sP11(X9,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl640),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl640])])).
fof(f11701,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X1),s('type.num.num',X2))))) ) | $spl132),
  inference(superposition,[],[f11512,f237])).
fof(f11512,plain,(
  ( ! [X279] : (sP11(X279,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl132),
  inference(subsumption_resolution,[],[f11361,f353])).
fof(f353,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f8])).
fof(f8,axiom,(
  p(s(bool,'T_0'))),
  file('thm.arithmetic.MODEQ_SYM','thm.bool.TRUTH')).
fof(f11361,plain,(
  ( ! [X279] : (~p(s(bool,'T_0')) | sP11(X279,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl132),
  inference(superposition,[],[f311,f3291])).
fof(f11704,plain,(
  ( ! [X10,X11,X9] : (sP11(X9,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X10),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X11),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X11),s('type.num.num',X10))))) ) | $spl132),
  inference(superposition,[],[f11512,f239])).
fof(f21742,plain,(
  ( ! [X9] : (sP11('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')),X9)) ) | $spl638),
  inference(cnf_transformation,[],[f21742_D])).
fof(f21742_D,plain,(
  ( ! [X9] : (sP11('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')),X9)) ) <=> ~$spl638),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl638])])).
fof(f11731,plain,(
  ( ! [X2,X0,X1] : (sP11('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')),X0) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X1),s('type.num.num',X2))))) ) | $spl132),
  inference(superposition,[],[f11513,f237])).
fof(f11513,plain,(
  ( ! [X280] : (sP11('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X280)) ) | $spl132),
  inference(subsumption_resolution,[],[f11362,f353])).
fof(f11362,plain,(
  ( ! [X280] : (~p(s(bool,'T_0')) | sP11('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X280)) ) | $spl132),
  inference(superposition,[],[f312,f3291])).
fof(f11734,plain,(
  ( ! [X10,X11,X9] : (sP11('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')),X9) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X10),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X11),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X11),s('type.num.num',X10))))) ) | $spl132),
  inference(superposition,[],[f11513,f239])).
fof(f21741,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',X12),s('type.num.num',sK14))))) ) | $spl636),
  inference(cnf_transformation,[],[f21741_D])).
fof(f21741_D,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',X12),s('type.num.num',sK14))))) ) <=> ~$spl636),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl636])])).
fof(f11761,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X1),s('type.num.num',X0),s('type.num.num',sK14)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl30),
  inference(superposition,[],[f5896,f237])).
fof(f5896,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X4),s('type.num.num',sK13),s('type.num.num',sK12))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X5)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X5),s('type.num.num',X4),s('type.num.num',sK14))))) ) | $spl30),
  inference(superposition,[],[f5544,f238])).
fof(f11764,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',X12),s('type.num.num',sK14)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl30),
  inference(superposition,[],[f5896,f239])).
fof(f21740,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X12),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',sK13),s('type.num.num',X12))))) ) | $spl634),
  inference(cnf_transformation,[],[f21740_D])).
fof(f21740_D,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X12),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',sK13),s('type.num.num',X12))))) ) <=> ~$spl634),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl634])])).
fof(f11863,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X1),s('type.num.num',sK13),s('type.num.num',X0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl30),
  inference(superposition,[],[f5901,f237])).
fof(f5901,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X2),s('type.num.num',sK12))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X3)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X3),s('type.num.num',sK13),s('type.num.num',X2))))) ) | $spl30),
  inference(superposition,[],[f5544,f238])).
fof(f11866,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X12),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',sK13),s('type.num.num',X12)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl30),
  inference(superposition,[],[f5901,f239])).
fof(f21739,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X12),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',X12),s('type.num.num',sK13))))) ) | $spl632),
  inference(cnf_transformation,[],[f21739_D])).
fof(f21739_D,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X12),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',X12),s('type.num.num',sK13))))) ) <=> ~$spl632),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl632])])).
fof(f11915,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X1),s('type.num.num',X0),s('type.num.num',sK13)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl30),
  inference(superposition,[],[f5902,f237])).
fof(f5902,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X4),s('type.num.num',sK12))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X5)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X5),s('type.num.num',X4),s('type.num.num',sK13))))) ) | $spl30),
  inference(superposition,[],[f5544,f238])).
fof(f11918,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X12),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',X12),s('type.num.num',sK13)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl30),
  inference(superposition,[],[f5902,f239])).
fof(f21738,plain,(
  ( ! [X12,X13] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',sK14),s('type.num.num',X12))))) ) | $spl630),
  inference(cnf_transformation,[],[f21738_D])).
fof(f21738_D,plain,(
  ( ! [X12,X13] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',sK14),s('type.num.num',X12))))) ) <=> ~$spl630),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl630])])).
fof(f12391,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X1),s('type.num.num',sK14),s('type.num.num',X0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) )),
  inference(superposition,[],[f360,f237])).
fof(f360,plain,(
  ( ! [X2,X3] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X2),s('type.num.num',sK12),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X2),s('type.num.num',sK13),s('type.num.num',sK12))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X3)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X3),s('type.num.num',sK14),s('type.num.num',X2))))) )),
  inference(superposition,[],[f199,f238])).
fof(f12394,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',sK14),s('type.num.num',X12)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) )),
  inference(superposition,[],[f360,f239])).
fof(f21691,plain,(
  ( ! [X12,X13] : (sP2(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl542),
  inference(cnf_transformation,[],[f21691_D])).
fof(f21691_D,plain,(
  ( ! [X12,X13] : (sP2(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl542),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl542])])).
fof(f12428,plain,(
  ( ! [X2,X0,X3,X1] : (sP2(X0,X1,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl98),
  inference(superposition,[],[f5433,f237])).
fof(f5433,plain,(
  ( ! [X10,X11] : (sP2(X10,X11,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK12)))) ) | $spl98),
  inference(resolution,[],[f2597,f208])).
fof(f208,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f139])).
fof(f12431,plain,(
  ( ! [X14,X12,X15,X13] : (sP2(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl98),
  inference(superposition,[],[f5433,f239])).
fof(f21690,plain,(
  ( ! [X12,X13] : (sP1(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl540),
  inference(cnf_transformation,[],[f21690_D])).
fof(f21690_D,plain,(
  ( ! [X12,X13] : (sP1(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl540),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl540])])).
fof(f12448,plain,(
  ( ! [X2,X0,X3,X1] : (sP1(X0,X1,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl98),
  inference(superposition,[],[f5436,f237])).
fof(f5436,plain,(
  ( ! [X17,X16] : (sP1(X16,X17,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK12)))) ) | $spl98),
  inference(resolution,[],[f2597,f212])).
fof(f12451,plain,(
  ( ! [X14,X12,X15,X13] : (sP1(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl98),
  inference(superposition,[],[f5436,f239])).
fof(f21737,plain,(
  ( ! [X12,X13] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',X12),s('type.num.num',sK14))))) ) | $spl628),
  inference(cnf_transformation,[],[f21737_D])).
fof(f21737_D,plain,(
  ( ! [X12,X13] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',X12),s('type.num.num',sK14))))) ) <=> ~$spl628),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl628])])).
fof(f12476,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X1),s('type.num.num',X0),s('type.num.num',sK14)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) )),
  inference(superposition,[],[f361,f237])).
fof(f361,plain,(
  ( ! [X4,X5] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X4),s('type.num.num',sK12),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X4),s('type.num.num',sK13),s('type.num.num',sK12))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X5)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X5),s('type.num.num',X4),s('type.num.num',sK14))))) )),
  inference(superposition,[],[f199,f238])).
fof(f12479,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',X12),s('type.num.num',sK14)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) )),
  inference(superposition,[],[f361,f239])).
fof(f21689,plain,(
  ( ! [X12,X13] : (sP0(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl538),
  inference(cnf_transformation,[],[f21689_D])).
fof(f21689_D,plain,(
  ( ! [X12,X13] : (sP0(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl538),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl538])])).
fof(f12511,plain,(
  ( ! [X2,X0,X3,X1] : (sP0(X0,X1,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl98),
  inference(superposition,[],[f5437,f237])).
fof(f5437,plain,(
  ( ! [X19,X18] : (sP0(X18,X19,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK12)))) ) | $spl98),
  inference(resolution,[],[f2597,f216])).
fof(f12514,plain,(
  ( ! [X14,X12,X15,X13] : (sP0(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl98),
  inference(superposition,[],[f5437,f239])).
fof(f21688,plain,(
  ( ! [X12,X13] : (sP0(X12,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X13)) ) | $spl536),
  inference(cnf_transformation,[],[f21688_D])).
fof(f21688_D,plain,(
  ( ! [X12,X13] : (sP0(X12,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X13)) ) <=> ~$spl536),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl536])])).
fof(f12532,plain,(
  ( ! [X2,X0,X3,X1] : (sP0(X0,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X1) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl98),
  inference(superposition,[],[f5438,f237])).
fof(f5438,plain,(
  ( ! [X21,X20] : (sP0(X20,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK12)),X21)) ) | $spl98),
  inference(resolution,[],[f2597,f217])).
fof(f217,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f144])).
fof(f12535,plain,(
  ( ! [X14,X12,X15,X13] : (sP0(X12,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X13) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl98),
  inference(superposition,[],[f5438,f239])).
fof(f21687,plain,(
  ( ! [X12,X13] : (sP0('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X12,X13)) ) | $spl534),
  inference(cnf_transformation,[],[f21687_D])).
fof(f21687_D,plain,(
  ( ! [X12,X13] : (sP0('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X12,X13)) ) <=> ~$spl534),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl534])])).
fof(f12553,plain,(
  ( ! [X2,X0,X3,X1] : (sP0('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X0,X1) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl98),
  inference(superposition,[],[f5439,f237])).
fof(f12556,plain,(
  ( ! [X14,X12,X15,X13] : (sP0('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X12,X13) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl98),
  inference(superposition,[],[f5439,f239])).
fof(f21686,plain,(
  ( ! [X12,X13] : (sP4(X12,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X13)) ) | $spl532),
  inference(cnf_transformation,[],[f21686_D])).
fof(f21686_D,plain,(
  ( ! [X12,X13] : (sP4(X12,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X13)) ) <=> ~$spl532),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl532])])).
fof(f12575,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X0,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X1) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl98),
  inference(superposition,[],[f5455,f237])).
fof(f5455,plain,(
  ( ! [X45,X44] : (sP4(X44,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK12)),X45)) ) | $spl98),
  inference(resolution,[],[f2597,f252])).
fof(f252,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f157])).
fof(f157,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f156])).
fof(f156,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(flattening,[],[f155])).
fof(f155,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(nnf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X2,X0,X1] : (sP4(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f12578,plain,(
  ( ! [X14,X12,X15,X13] : (sP4(X12,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X13) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl98),
  inference(superposition,[],[f5455,f239])).
fof(f21685,plain,(
  ( ! [X12,X13] : (sP6(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl530),
  inference(cnf_transformation,[],[f21685_D])).
fof(f21685_D,plain,(
  ( ! [X12,X13] : (sP6(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl530),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl530])])).
fof(f12595,plain,(
  ( ! [X2,X0,X3,X1] : (sP6(X0,X1,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl98),
  inference(superposition,[],[f5468,f237])).
fof(f5468,plain,(
  ( ! [X70,X71] : (sP6(X70,X71,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK12)))) ) | $spl98),
  inference(resolution,[],[f2597,f268])).
fof(f268,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f165])).
fof(f165,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(rectify,[],[f164])).
fof(f164,plain,(
  ! [X2,X0,X1] : ((~sP6(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP6(X2,X0,X1)))),
  inference(flattening,[],[f163])).
fof(f163,plain,(
  ! [X2,X0,X1] : ((~sP6(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP6(X2,X0,X1)))),
  inference(nnf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X2,X0,X1] : (sP6(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f12598,plain,(
  ( ! [X14,X12,X15,X13] : (sP6(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl98),
  inference(superposition,[],[f5468,f239])).
fof(f21736,plain,(
  ( ! [X12,X13] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X12),s('type.num.num','const.num.|48|_0'))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',X12))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',sK13),s('type.num.num',X12))))) ) | $spl626),
  inference(cnf_transformation,[],[f21736_D])).
fof(f21736_D,plain,(
  ( ! [X12,X13] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X12),s('type.num.num','const.num.|48|_0'))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',X12))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',sK13),s('type.num.num',X12))))) ) <=> ~$spl626),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl626])])).
fof(f12625,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X1),s('type.num.num',sK13),s('type.num.num',X0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) )),
  inference(superposition,[],[f366,f237])).
fof(f366,plain,(
  ( ! [X2,X3] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X2),s('type.num.num',sK12))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',X2))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X3)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X3),s('type.num.num',sK13),s('type.num.num',X2))))) )),
  inference(superposition,[],[f199,f238])).
fof(f12628,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X12),s('type.num.num','const.num.|48|_0'))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',X12))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',sK13),s('type.num.num',X12)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) )),
  inference(superposition,[],[f366,f239])).
fof(f21684,plain,(
  ( ! [X12,X13] : (sP6(X12,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X13)) ) | $spl528),
  inference(cnf_transformation,[],[f21684_D])).
fof(f21684_D,plain,(
  ( ! [X12,X13] : (sP6(X12,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X13)) ) <=> ~$spl528),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl528])])).
fof(f12677,plain,(
  ( ! [X2,X0,X3,X1] : (sP6(X0,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X1) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl98),
  inference(superposition,[],[f5469,f237])).
fof(f5469,plain,(
  ( ! [X72,X73] : (sP6(X72,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK12)),X73)) ) | $spl98),
  inference(resolution,[],[f2597,f269])).
fof(f269,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f165])).
fof(f12680,plain,(
  ( ! [X14,X12,X15,X13] : (sP6(X12,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X13) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl98),
  inference(superposition,[],[f5469,f239])).
fof(f21683,plain,(
  ( ! [X12,X13] : (sP8(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl526),
  inference(cnf_transformation,[],[f21683_D])).
fof(f21683_D,plain,(
  ( ! [X12,X13] : (sP8(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl526),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl526])])).
fof(f12697,plain,(
  ( ! [X2,X0,X3,X1] : (sP8(X0,X1,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl98),
  inference(superposition,[],[f5485,f237])).
fof(f5485,plain,(
  ( ! [X105,X104] : (sP8(X104,X105,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK12)))) ) | $spl98),
  inference(resolution,[],[f2597,f285])).
fof(f285,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f172])).
fof(f172,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(flattening,[],[f171])).
fof(f171,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(nnf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0,X1,X2] : (sP8(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f12700,plain,(
  ( ! [X14,X12,X15,X13] : (sP8(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl98),
  inference(superposition,[],[f5485,f239])).
fof(f21735,plain,(
  ( ! [X12,X13] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X12),s('type.num.num','const.num.|48|_0'))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',X12))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',X12),s('type.num.num',sK13))))) ) | $spl624),
  inference(cnf_transformation,[],[f21735_D])).
fof(f21735_D,plain,(
  ( ! [X12,X13] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X12),s('type.num.num','const.num.|48|_0'))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',X12))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',X12),s('type.num.num',sK13))))) ) <=> ~$spl624),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl624])])).
fof(f12815,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X1),s('type.num.num',X0),s('type.num.num',sK13)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) )),
  inference(superposition,[],[f367,f237])).
fof(f367,plain,(
  ( ! [X4,X5] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X4),s('type.num.num',sK12))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',X4))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X5)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X5),s('type.num.num',X4),s('type.num.num',sK13))))) )),
  inference(superposition,[],[f199,f238])).
fof(f12818,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X12),s('type.num.num','const.num.|48|_0'))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',X12))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',X12),s('type.num.num',sK13)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) )),
  inference(superposition,[],[f367,f239])).
fof(f21734,plain,(
  ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)))) | $spl623),
  inference(cnf_transformation,[],[f21734_D])).
fof(f21734_D,plain,(
  ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)))) <=> ~$spl623),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl623])])).
fof(f13241,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X0),s('type.num.num',X1))))) ) | $spl128),
  inference(superposition,[],[f6819,f237])).
fof(f6819,plain,(
  ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | $spl128),
  inference(equality_resolution,[],[f3287])).
fof(f13244,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X7),s('type.num.num',X6))))) ) | $spl128),
  inference(superposition,[],[f6819,f239])).
fof(f21732,plain,(
  ( ! [X12,X13] : (sP2(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')))) ) | $spl620),
  inference(cnf_transformation,[],[f21732_D])).
fof(f21732_D,plain,(
  ( ! [X12,X13] : (sP2(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl620),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl620])])).
fof(f13583,plain,(
  ( ! [X2,X0,X3,X1] : (sP2(X0,X1,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl132),
  inference(superposition,[],[f11404,f237])).
fof(f11404,plain,(
  ( ! [X19,X20] : (sP2(X19,X20,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl132),
  inference(subsumption_resolution,[],[f11203,f353])).
fof(f11203,plain,(
  ( ! [X19,X20] : (~p(s(bool,'T_0')) | sP2(X19,X20,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl132),
  inference(superposition,[],[f208,f3291])).
fof(f13586,plain,(
  ( ! [X14,X12,X15,X13] : (sP2(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl132),
  inference(superposition,[],[f11404,f239])).
fof(f21731,plain,(
  ( ! [X12,X13] : (sP1(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')))) ) | $spl618),
  inference(cnf_transformation,[],[f21731_D])).
fof(f21731_D,plain,(
  ( ! [X12,X13] : (sP1(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl618),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl618])])).
fof(f13613,plain,(
  ( ! [X2,X0,X3,X1] : (sP1(X0,X1,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl132),
  inference(superposition,[],[f11409,f237])).
fof(f11409,plain,(
  ( ! [X31,X32] : (sP1(X31,X32,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl132),
  inference(subsumption_resolution,[],[f11209,f353])).
fof(f11209,plain,(
  ( ! [X31,X32] : (~p(s(bool,'T_0')) | sP1(X31,X32,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl132),
  inference(superposition,[],[f212,f3291])).
fof(f13616,plain,(
  ( ! [X14,X12,X15,X13] : (sP1(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl132),
  inference(superposition,[],[f11409,f239])).
fof(f21730,plain,(
  ( ! [X12,X13] : (sP0(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')))) ) | $spl616),
  inference(cnf_transformation,[],[f21730_D])).
fof(f21730_D,plain,(
  ( ! [X12,X13] : (sP0(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl616),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl616])])).
fof(f13643,plain,(
  ( ! [X2,X0,X3,X1] : (sP0(X0,X1,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl132),
  inference(superposition,[],[f11410,f237])).
fof(f11410,plain,(
  ( ! [X43,X44] : (sP0(X43,X44,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl132),
  inference(subsumption_resolution,[],[f11215,f353])).
fof(f11215,plain,(
  ( ! [X43,X44] : (~p(s(bool,'T_0')) | sP0(X43,X44,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl132),
  inference(superposition,[],[f216,f3291])).
fof(f13646,plain,(
  ( ! [X14,X12,X15,X13] : (sP0(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl132),
  inference(superposition,[],[f11410,f239])).
fof(f21729,plain,(
  ( ! [X12,X13] : (sP0(X12,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')),X13)) ) | $spl614),
  inference(cnf_transformation,[],[f21729_D])).
fof(f21729_D,plain,(
  ( ! [X12,X13] : (sP0(X12,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')),X13)) ) <=> ~$spl614),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl614])])).
fof(f13673,plain,(
  ( ! [X2,X0,X3,X1] : (sP0(X0,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')),X1) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl132),
  inference(superposition,[],[f11411,f237])).
fof(f11411,plain,(
  ( ! [X45,X46] : (sP0(X45,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X46)) ) | $spl132),
  inference(subsumption_resolution,[],[f11216,f353])).
fof(f11216,plain,(
  ( ! [X45,X46] : (~p(s(bool,'T_0')) | sP0(X45,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X46)) ) | $spl132),
  inference(superposition,[],[f217,f3291])).
fof(f13676,plain,(
  ( ! [X14,X12,X15,X13] : (sP0(X12,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')),X13) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl132),
  inference(superposition,[],[f11411,f239])).
fof(f21728,plain,(
  ( ! [X12,X13] : (sP0('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')),X12,X13)) ) | $spl612),
  inference(cnf_transformation,[],[f21728_D])).
fof(f21728_D,plain,(
  ( ! [X12,X13] : (sP0('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')),X12,X13)) ) <=> ~$spl612),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl612])])).
fof(f13703,plain,(
  ( ! [X2,X0,X3,X1] : (sP0('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')),X0,X1) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl132),
  inference(superposition,[],[f11412,f237])).
fof(f11412,plain,(
  ( ! [X47,X48] : (sP0('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X47,X48)) ) | $spl132),
  inference(subsumption_resolution,[],[f11217,f353])).
fof(f11217,plain,(
  ( ! [X47,X48] : (~p(s(bool,'T_0')) | sP0('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X47,X48)) ) | $spl132),
  inference(superposition,[],[f218,f3291])).
fof(f13706,plain,(
  ( ! [X14,X12,X15,X13] : (sP0('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')),X12,X13) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl132),
  inference(superposition,[],[f11412,f239])).
fof(f21727,plain,(
  ( ! [X12,X13] : (sP4(X12,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')),X13)) ) | $spl610),
  inference(cnf_transformation,[],[f21727_D])).
fof(f21727_D,plain,(
  ( ! [X12,X13] : (sP4(X12,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')),X13)) ) <=> ~$spl610),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl610])])).
fof(f13733,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X0,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')),X1) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl132),
  inference(superposition,[],[f11444,f237])).
fof(f11444,plain,(
  ( ! [X109,X110] : (sP4(X109,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X110)) ) | $spl132),
  inference(subsumption_resolution,[],[f11267,f353])).
fof(f11267,plain,(
  ( ! [X109,X110] : (~p(s(bool,'T_0')) | sP4(X109,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X110)) ) | $spl132),
  inference(superposition,[],[f252,f3291])).
fof(f13736,plain,(
  ( ! [X14,X12,X15,X13] : (sP4(X12,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')),X13) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl132),
  inference(superposition,[],[f11444,f239])).
fof(f21726,plain,(
  ( ! [X12,X13] : (sP6(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')))) ) | $spl608),
  inference(cnf_transformation,[],[f21726_D])).
fof(f21726_D,plain,(
  ( ! [X12,X13] : (sP6(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl608),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl608])])).
fof(f13765,plain,(
  ( ! [X2,X0,X3,X1] : (sP6(X0,X1,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl132),
  inference(superposition,[],[f11465,f237])).
fof(f11465,plain,(
  ( ! [X167,X168] : (sP6(X167,X168,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl132),
  inference(subsumption_resolution,[],[f11296,f353])).
fof(f11296,plain,(
  ( ! [X167,X168] : (~p(s(bool,'T_0')) | sP6(X167,X168,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl132),
  inference(superposition,[],[f268,f3291])).
fof(f13768,plain,(
  ( ! [X14,X12,X15,X13] : (sP6(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl132),
  inference(superposition,[],[f11465,f239])).
fof(f21725,plain,(
  ( ! [X12,X13] : (sP6(X12,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')),X13)) ) | $spl606),
  inference(cnf_transformation,[],[f21725_D])).
fof(f21725_D,plain,(
  ( ! [X12,X13] : (sP6(X12,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')),X13)) ) <=> ~$spl606),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl606])])).
fof(f13794,plain,(
  ( ! [X2,X0,X3,X1] : (sP6(X0,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')),X1) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl132),
  inference(superposition,[],[f11466,f237])).
fof(f11466,plain,(
  ( ! [X169,X170] : (sP6(X169,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X170)) ) | $spl132),
  inference(subsumption_resolution,[],[f11297,f353])).
fof(f11297,plain,(
  ( ! [X169,X170] : (~p(s(bool,'T_0')) | sP6(X169,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X170)) ) | $spl132),
  inference(superposition,[],[f269,f3291])).
fof(f13797,plain,(
  ( ! [X14,X12,X15,X13] : (sP6(X12,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')),X13) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl132),
  inference(superposition,[],[f11466,f239])).
fof(f21724,plain,(
  ( ! [X12,X13] : (sP8(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')))) ) | $spl604),
  inference(cnf_transformation,[],[f21724_D])).
fof(f21724_D,plain,(
  ( ! [X12,X13] : (sP8(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl604),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl604])])).
fof(f13823,plain,(
  ( ! [X2,X0,X3,X1] : (sP8(X0,X1,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl132),
  inference(superposition,[],[f11488,f237])).
fof(f11488,plain,(
  ( ! [X227,X228] : (sP8(X227,X228,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl132),
  inference(subsumption_resolution,[],[f11326,f353])).
fof(f11326,plain,(
  ( ! [X227,X228] : (~p(s(bool,'T_0')) | sP8(X227,X228,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl132),
  inference(superposition,[],[f285,f3291])).
fof(f13826,plain,(
  ( ! [X14,X12,X15,X13] : (sP8(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl132),
  inference(superposition,[],[f11488,f239])).
fof(f21723,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',sK14),s('type.num.num',X12)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK14),s('type.num.num',X13))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X12),s('type.num.num',X13)))) ) | $spl602),
  inference(cnf_transformation,[],[f21723_D])).
fof(f21723_D,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',sK14),s('type.num.num',X12)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK14),s('type.num.num',X13))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X12),s('type.num.num',X13)))) ) <=> ~$spl602),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl602])])).
fof(f13855,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK14),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X0),s('type.num.num',X1))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X1),s('type.num.num',sK14),s('type.num.num',X0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl30),
  inference(superposition,[],[f5898,f237])).
fof(f5898,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X8),s('type.num.num',sK13),s('type.num.num',sK12))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK14),s('type.num.num',X9))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X8),s('type.num.num',X9))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X9),s('type.num.num',sK14),s('type.num.num',X8))))) ) | $spl30),
  inference(superposition,[],[f5544,f240])).
fof(f240,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X1),s('type.num.num',X0))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X2),s('type.num.num',X0))) | s('type.num.num',X1) = s('type.num.num',X2) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num',X2),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f151])).
fof(f13858,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK14),s('type.num.num',X13))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X12),s('type.num.num',X13))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',sK14),s('type.num.num',X12)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl30),
  inference(superposition,[],[f5898,f239])).
fof(f21722,plain,(
  ( ! [X12,X13] : (sP2(X12,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X13)) ) | $spl600),
  inference(cnf_transformation,[],[f21722_D])).
fof(f21722_D,plain,(
  ( ! [X12,X13] : (sP2(X12,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X13)) ) <=> ~$spl600),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl600])])).
fof(f13904,plain,(
  ( ! [X2,X0,X3,X1] : (sP2(X0,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X1) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl47),
  inference(superposition,[],[f7190,f237])).
fof(f13907,plain,(
  ( ! [X14,X12,X15,X13] : (sP2(X12,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X13) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl47),
  inference(superposition,[],[f7190,f239])).
fof(f21721,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',X12),s('type.num.num',sK14)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK14),s('type.num.num',X13))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X12),s('type.num.num',X13)))) ) | $spl598),
  inference(cnf_transformation,[],[f21721_D])).
fof(f21721_D,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',X12),s('type.num.num',sK14)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK14),s('type.num.num',X13))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X12),s('type.num.num',X13)))) ) <=> ~$spl598),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl598])])).
fof(f13931,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK14),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X0),s('type.num.num',X1))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X1),s('type.num.num',X0),s('type.num.num',sK14)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl30),
  inference(superposition,[],[f5899,f237])).
fof(f5899,plain,(
  ( ! [X10,X11] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X10),s('type.num.num',sK13),s('type.num.num',sK12))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK14),s('type.num.num',X11))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X10),s('type.num.num',X11))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X11),s('type.num.num',X10),s('type.num.num',sK14))))) ) | $spl30),
  inference(superposition,[],[f5544,f240])).
fof(f13934,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK14),s('type.num.num',X13))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X12),s('type.num.num',X13))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',X12),s('type.num.num',sK14)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl30),
  inference(superposition,[],[f5899,f239])).
fof(f21720,plain,(
  ( ! [X12,X13] : (sP2('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X12,X13)) ) | $spl596),
  inference(cnf_transformation,[],[f21720_D])).
fof(f21720_D,plain,(
  ( ! [X12,X13] : (sP2('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X12,X13)) ) <=> ~$spl596),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl596])])).
fof(f13980,plain,(
  ( ! [X2,X0,X3,X1] : (sP2('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X0,X1) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl47),
  inference(superposition,[],[f7191,f237])).
fof(f7191,plain,(
  ( ! [X10,X11] : (sP2('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)),X10,X11)) ) | $spl47),
  inference(resolution,[],[f552,f210])).
fof(f13983,plain,(
  ( ! [X14,X12,X15,X13] : (sP2('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X12,X13) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl47),
  inference(superposition,[],[f7191,f239])).
fof(f21719,plain,(
  ( ! [X12,X13] : (sP1(X12,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X13)) ) | $spl594),
  inference(cnf_transformation,[],[f21719_D])).
fof(f21719_D,plain,(
  ( ! [X12,X13] : (sP1(X12,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X13)) ) <=> ~$spl594),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl594])])).
fof(f14009,plain,(
  ( ! [X2,X0,X3,X1] : (sP1(X0,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X1) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl47),
  inference(superposition,[],[f7193,f237])).
fof(f7193,plain,(
  ( ! [X14,X15] : (sP1(X14,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)),X15)) ) | $spl47),
  inference(resolution,[],[f552,f213])).
fof(f14012,plain,(
  ( ! [X14,X12,X15,X13] : (sP1(X12,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X13) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl47),
  inference(superposition,[],[f7193,f239])).
fof(f21718,plain,(
  ( ! [X12,X13] : (sP1('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X12,X13)) ) | $spl592),
  inference(cnf_transformation,[],[f21718_D])).
fof(f21718_D,plain,(
  ( ! [X12,X13] : (sP1('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X12,X13)) ) <=> ~$spl592),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl592])])).
fof(f14037,plain,(
  ( ! [X2,X0,X3,X1] : (sP1('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X0,X1) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl47),
  inference(superposition,[],[f7194,f237])).
fof(f7194,plain,(
  ( ! [X17,X16] : (sP1('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)),X16,X17)) ) | $spl47),
  inference(resolution,[],[f552,f214])).
fof(f214,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f142])).
fof(f14040,plain,(
  ( ! [X14,X12,X15,X13] : (sP1('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X12,X13) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl47),
  inference(superposition,[],[f7194,f239])).
fof(f21717,plain,(
  ( ! [X12,X13] : (sP4(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl590),
  inference(cnf_transformation,[],[f21717_D])).
fof(f21717_D,plain,(
  ( ! [X12,X13] : (sP4(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl590),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl590])])).
fof(f14067,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X0,X1,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl47),
  inference(superposition,[],[f7212,f237])).
fof(f7212,plain,(
  ( ! [X43,X42] : (sP4(X42,X43,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) ) | $spl47),
  inference(resolution,[],[f552,f251])).
fof(f251,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f157])).
fof(f14070,plain,(
  ( ! [X14,X12,X15,X13] : (sP4(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl47),
  inference(superposition,[],[f7212,f239])).
fof(f21716,plain,(
  ( ! [X12,X13] : (sP4('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X12,X13)) ) | $spl588),
  inference(cnf_transformation,[],[f21716_D])).
fof(f21716_D,plain,(
  ( ! [X12,X13] : (sP4('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X12,X13)) ) <=> ~$spl588),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl588])])).
fof(f14095,plain,(
  ( ! [X2,X0,X3,X1] : (sP4('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X0,X1) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl47),
  inference(superposition,[],[f7213,f237])).
fof(f7213,plain,(
  ( ! [X45,X44] : (sP4('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)),X44,X45)) ) | $spl47),
  inference(resolution,[],[f552,f253])).
fof(f253,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f157])).
fof(f14098,plain,(
  ( ! [X14,X12,X15,X13] : (sP4('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X12,X13) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl47),
  inference(superposition,[],[f7213,f239])).
fof(f21715,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X12),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',sK13),s('type.num.num',X12)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X13))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X12),s('type.num.num',X13)))) ) | $spl586),
  inference(cnf_transformation,[],[f21715_D])).
fof(f21715_D,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X12),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',sK13),s('type.num.num',X12)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X13))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X12),s('type.num.num',X13)))) ) <=> ~$spl586),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl586])])).
fof(f14122,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X0),s('type.num.num',X1))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X1),s('type.num.num',sK13),s('type.num.num',X0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl30),
  inference(superposition,[],[f5904,f237])).
fof(f5904,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X8),s('type.num.num',sK12))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X9))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X8),s('type.num.num',X9))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X9),s('type.num.num',sK13),s('type.num.num',X8))))) ) | $spl30),
  inference(superposition,[],[f5544,f240])).
fof(f14125,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X12),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X13))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X12),s('type.num.num',X13))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',sK13),s('type.num.num',X12)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl30),
  inference(superposition,[],[f5904,f239])).
fof(f21714,plain,(
  ( ! [X12,X13] : (sP6('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X12,X13)) ) | $spl584),
  inference(cnf_transformation,[],[f21714_D])).
fof(f21714_D,plain,(
  ( ! [X12,X13] : (sP6('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X12,X13)) ) <=> ~$spl584),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl584])])).
fof(f14174,plain,(
  ( ! [X2,X0,X3,X1] : (sP6('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X0,X1) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl47),
  inference(superposition,[],[f7229,f237])).
fof(f7229,plain,(
  ( ! [X76,X77] : (sP6('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)),X76,X77)) ) | $spl47),
  inference(resolution,[],[f552,f270])).
fof(f270,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f165])).
fof(f14177,plain,(
  ( ! [X14,X12,X15,X13] : (sP6('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X12,X13) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl47),
  inference(superposition,[],[f7229,f239])).
fof(f21713,plain,(
  ( ! [X12,X13] : (sP8(X12,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X13)) ) | $spl582),
  inference(cnf_transformation,[],[f21713_D])).
fof(f21713_D,plain,(
  ( ! [X12,X13] : (sP8(X12,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X13)) ) <=> ~$spl582),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl582])])).
fof(f14203,plain,(
  ( ! [X2,X0,X3,X1] : (sP8(X0,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X1) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl47),
  inference(superposition,[],[f7242,f237])).
fof(f7242,plain,(
  ( ! [X103,X102] : (sP8(X102,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)),X103)) ) | $spl47),
  inference(resolution,[],[f552,f286])).
fof(f286,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f172])).
fof(f14206,plain,(
  ( ! [X14,X12,X15,X13] : (sP8(X12,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X13) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl47),
  inference(superposition,[],[f7242,f239])).
fof(f21712,plain,(
  ( ! [X12,X13] : (sP8('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X12,X13)) ) | $spl580),
  inference(cnf_transformation,[],[f21712_D])).
fof(f21712_D,plain,(
  ( ! [X12,X13] : (sP8('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X12,X13)) ) <=> ~$spl580),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl580])])).
fof(f14231,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X0,X1) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl47),
  inference(superposition,[],[f7243,f237])).
fof(f7243,plain,(
  ( ! [X105,X104] : (sP8('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)),X104,X105)) ) | $spl47),
  inference(resolution,[],[f552,f287])).
fof(f287,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f172])).
fof(f14234,plain,(
  ( ! [X14,X12,X15,X13] : (sP8('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X12,X13) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl47),
  inference(superposition,[],[f7243,f239])).
fof(f21711,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X12),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',X12),s('type.num.num',sK13)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X13))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X12),s('type.num.num',X13)))) ) | $spl578),
  inference(cnf_transformation,[],[f21711_D])).
fof(f21711_D,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X12),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',X12),s('type.num.num',sK13)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X13))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X12),s('type.num.num',X13)))) ) <=> ~$spl578),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl578])])).
fof(f14314,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X0),s('type.num.num',X1))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X1),s('type.num.num',X0),s('type.num.num',sK13)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl30),
  inference(superposition,[],[f5905,f237])).
fof(f5905,plain,(
  ( ! [X10,X11] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X10),s('type.num.num',sK12))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X11))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X10),s('type.num.num',X11))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X11),s('type.num.num',X10),s('type.num.num',sK13))))) ) | $spl30),
  inference(superposition,[],[f5544,f240])).
fof(f14317,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X12),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X13))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X12),s('type.num.num',X13))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',X12),s('type.num.num',sK13)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl30),
  inference(superposition,[],[f5905,f239])).
fof(f21710,plain,(
  ( ! [X12,X13] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',sK14),s('type.num.num',X12)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK14),s('type.num.num',X13))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X12),s('type.num.num',X13)))) ) | $spl576),
  inference(cnf_transformation,[],[f21710_D])).
fof(f21710_D,plain,(
  ( ! [X12,X13] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',sK14),s('type.num.num',X12)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK14),s('type.num.num',X13))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X12),s('type.num.num',X13)))) ) <=> ~$spl576),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl576])])).
fof(f14665,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK14),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X0),s('type.num.num',X1))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X1),s('type.num.num',sK14),s('type.num.num',X0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) )),
  inference(superposition,[],[f363,f237])).
fof(f363,plain,(
  ( ! [X8,X9] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X8),s('type.num.num',sK12),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X8),s('type.num.num',sK13),s('type.num.num',sK12))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK14),s('type.num.num',X9))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X8),s('type.num.num',X9))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X9),s('type.num.num',sK14),s('type.num.num',X8))))) )),
  inference(superposition,[],[f199,f240])).
fof(f14668,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK14),s('type.num.num',X13))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X12),s('type.num.num',X13))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',sK14),s('type.num.num',X12)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) )),
  inference(superposition,[],[f363,f239])).
fof(f21709,plain,(
  ( ! [X12,X13] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',X12),s('type.num.num',sK14)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK14),s('type.num.num',X13))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X12),s('type.num.num',X13)))) ) | $spl574),
  inference(cnf_transformation,[],[f21709_D])).
fof(f21709_D,plain,(
  ( ! [X12,X13] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',X12),s('type.num.num',sK14)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK14),s('type.num.num',X13))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X12),s('type.num.num',X13)))) ) <=> ~$spl574),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl574])])).
fof(f14708,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK14),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X0),s('type.num.num',X1))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X1),s('type.num.num',X0),s('type.num.num',sK14)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) )),
  inference(superposition,[],[f364,f237])).
fof(f364,plain,(
  ( ! [X10,X11] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X10),s('type.num.num',sK12),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X10),s('type.num.num',sK13),s('type.num.num',sK12))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK14),s('type.num.num',X11))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X10),s('type.num.num',X11))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X11),s('type.num.num',X10),s('type.num.num',sK14))))) )),
  inference(superposition,[],[f199,f240])).
fof(f14711,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK14),s('type.num.num',X13))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X12),s('type.num.num',X13))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',X12),s('type.num.num',sK14)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) )),
  inference(superposition,[],[f364,f239])).
fof(f21708,plain,(
  ( ! [X12,X13] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X12),s('type.num.num','const.num.|48|_0'))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',X12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',sK13),s('type.num.num',X12)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X13))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X12),s('type.num.num',X13)))) ) | $spl572),
  inference(cnf_transformation,[],[f21708_D])).
fof(f21708_D,plain,(
  ( ! [X12,X13] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X12),s('type.num.num','const.num.|48|_0'))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',X12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',sK13),s('type.num.num',X12)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X13))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X12),s('type.num.num',X13)))) ) <=> ~$spl572),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl572])])).
fof(f14754,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X0),s('type.num.num',X1))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X1),s('type.num.num',sK13),s('type.num.num',X0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) )),
  inference(superposition,[],[f369,f237])).
fof(f369,plain,(
  ( ! [X8,X9] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X8),s('type.num.num',sK12))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',X8))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X9))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X8),s('type.num.num',X9))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X9),s('type.num.num',sK13),s('type.num.num',X8))))) )),
  inference(superposition,[],[f199,f240])).
fof(f14757,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X12),s('type.num.num','const.num.|48|_0'))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',X12))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X13))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X12),s('type.num.num',X13))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',sK13),s('type.num.num',X12)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) )),
  inference(superposition,[],[f369,f239])).
fof(f21707,plain,(
  ( ! [X12,X13] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X12),s('type.num.num','const.num.|48|_0'))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',X12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',X12),s('type.num.num',sK13)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X13))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X12),s('type.num.num',X13)))) ) | $spl570),
  inference(cnf_transformation,[],[f21707_D])).
fof(f21707_D,plain,(
  ( ! [X12,X13] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X12),s('type.num.num','const.num.|48|_0'))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',X12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',X12),s('type.num.num',sK13)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X13))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X12),s('type.num.num',X13)))) ) <=> ~$spl570),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl570])])).
fof(f14816,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X0),s('type.num.num',X1))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X1),s('type.num.num',X0),s('type.num.num',sK13)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) )),
  inference(superposition,[],[f370,f237])).
fof(f370,plain,(
  ( ! [X10,X11] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X10),s('type.num.num',sK12))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',X10))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X11))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X10),s('type.num.num',X11))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X11),s('type.num.num',X10),s('type.num.num',sK13))))) )),
  inference(superposition,[],[f199,f240])).
fof(f14819,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X12),s('type.num.num','const.num.|48|_0'))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',X12))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X13))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X12),s('type.num.num',X13))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',X12),s('type.num.num',sK13)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) )),
  inference(superposition,[],[f370,f239])).
fof(f21706,plain,(
  ( ! [X12,X13] : (sP2(X12,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)),X13)) ) | $spl568),
  inference(cnf_transformation,[],[f21706_D])).
fof(f21706_D,plain,(
  ( ! [X12,X13] : (sP2(X12,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)),X13)) ) <=> ~$spl568),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl568])])).
fof(f15468,plain,(
  ( ! [X2,X0,X3,X1] : (sP2(X0,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)),X1) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl128),
  inference(superposition,[],[f13155,f237])).
fof(f13155,plain,(
  ( ! [X8,X9] : (sP2(X8,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)),X9)) ) | $spl128),
  inference(resolution,[],[f6819,f209])).
fof(f15471,plain,(
  ( ! [X14,X12,X15,X13] : (sP2(X12,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)),X13) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl128),
  inference(superposition,[],[f13155,f239])).
fof(f21705,plain,(
  ( ! [X12,X13] : (sP2('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)),X12,X13)) ) | $spl566),
  inference(cnf_transformation,[],[f21705_D])).
fof(f21705_D,plain,(
  ( ! [X12,X13] : (sP2('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)),X12,X13)) ) <=> ~$spl566),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl566])])).
fof(f15493,plain,(
  ( ! [X2,X0,X3,X1] : (sP2('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)),X0,X1) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl128),
  inference(superposition,[],[f13156,f237])).
fof(f13156,plain,(
  ( ! [X10,X11] : (sP2('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)),X10,X11)) ) | $spl128),
  inference(resolution,[],[f6819,f210])).
fof(f15496,plain,(
  ( ! [X14,X12,X15,X13] : (sP2('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)),X12,X13) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl128),
  inference(superposition,[],[f13156,f239])).
fof(f21704,plain,(
  ( ! [X12,X13] : (sP1(X12,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)),X13)) ) | $spl564),
  inference(cnf_transformation,[],[f21704_D])).
fof(f21704_D,plain,(
  ( ! [X12,X13] : (sP1(X12,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)),X13)) ) <=> ~$spl564),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl564])])).
fof(f15520,plain,(
  ( ! [X2,X0,X3,X1] : (sP1(X0,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)),X1) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl128),
  inference(superposition,[],[f13158,f237])).
fof(f13158,plain,(
  ( ! [X14,X15] : (sP1(X14,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)),X15)) ) | $spl128),
  inference(resolution,[],[f6819,f213])).
fof(f15523,plain,(
  ( ! [X14,X12,X15,X13] : (sP1(X12,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)),X13) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl128),
  inference(superposition,[],[f13158,f239])).
fof(f21703,plain,(
  ( ! [X12,X13] : (sP1('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)),X12,X13)) ) | $spl562),
  inference(cnf_transformation,[],[f21703_D])).
fof(f21703_D,plain,(
  ( ! [X12,X13] : (sP1('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)),X12,X13)) ) <=> ~$spl562),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl562])])).
fof(f15546,plain,(
  ( ! [X2,X0,X3,X1] : (sP1('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)),X0,X1) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl128),
  inference(superposition,[],[f13159,f237])).
fof(f13159,plain,(
  ( ! [X17,X16] : (sP1('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)),X16,X17)) ) | $spl128),
  inference(resolution,[],[f6819,f214])).
fof(f15549,plain,(
  ( ! [X14,X12,X15,X13] : (sP1('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)),X12,X13) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl128),
  inference(superposition,[],[f13159,f239])).
fof(f21702,plain,(
  ( ! [X12,X13] : (sP4(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)))) ) | $spl560),
  inference(cnf_transformation,[],[f21702_D])).
fof(f21702_D,plain,(
  ( ! [X12,X13] : (sP4(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)))) ) <=> ~$spl560),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl560])])).
fof(f15574,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X0,X1,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl128),
  inference(superposition,[],[f13177,f237])).
fof(f13177,plain,(
  ( ! [X43,X42] : (sP4(X42,X43,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) ) | $spl128),
  inference(resolution,[],[f6819,f251])).
fof(f15577,plain,(
  ( ! [X14,X12,X15,X13] : (sP4(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl128),
  inference(superposition,[],[f13177,f239])).
fof(f21701,plain,(
  ( ! [X12,X13] : (sP4('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)),X12,X13)) ) | $spl558),
  inference(cnf_transformation,[],[f21701_D])).
fof(f21701_D,plain,(
  ( ! [X12,X13] : (sP4('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)),X12,X13)) ) <=> ~$spl558),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl558])])).
fof(f15600,plain,(
  ( ! [X2,X0,X3,X1] : (sP4('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)),X0,X1) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl128),
  inference(superposition,[],[f13178,f237])).
fof(f13178,plain,(
  ( ! [X45,X44] : (sP4('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)),X44,X45)) ) | $spl128),
  inference(resolution,[],[f6819,f253])).
fof(f15603,plain,(
  ( ! [X14,X12,X15,X13] : (sP4('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)),X12,X13) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl128),
  inference(superposition,[],[f13178,f239])).
fof(f21700,plain,(
  ( ! [X12,X13] : (sP6('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)),X12,X13)) ) | $spl556),
  inference(cnf_transformation,[],[f21700_D])).
fof(f21700_D,plain,(
  ( ! [X12,X13] : (sP6('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)),X12,X13)) ) <=> ~$spl556),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl556])])).
fof(f15626,plain,(
  ( ! [X2,X0,X3,X1] : (sP6('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)),X0,X1) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl128),
  inference(superposition,[],[f13194,f237])).
fof(f13194,plain,(
  ( ! [X76,X77] : (sP6('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)),X76,X77)) ) | $spl128),
  inference(resolution,[],[f6819,f270])).
fof(f15629,plain,(
  ( ! [X14,X12,X15,X13] : (sP6('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)),X12,X13) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl128),
  inference(superposition,[],[f13194,f239])).
fof(f21699,plain,(
  ( ! [X12,X13] : (sP8(X12,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)),X13)) ) | $spl554),
  inference(cnf_transformation,[],[f21699_D])).
fof(f21699_D,plain,(
  ( ! [X12,X13] : (sP8(X12,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)),X13)) ) <=> ~$spl554),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl554])])).
fof(f15653,plain,(
  ( ! [X2,X0,X3,X1] : (sP8(X0,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)),X1) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl128),
  inference(superposition,[],[f13207,f237])).
fof(f13207,plain,(
  ( ! [X103,X102] : (sP8(X102,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)),X103)) ) | $spl128),
  inference(resolution,[],[f6819,f286])).
fof(f15656,plain,(
  ( ! [X14,X12,X15,X13] : (sP8(X12,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)),X13) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl128),
  inference(superposition,[],[f13207,f239])).
fof(f21698,plain,(
  ( ! [X12,X13] : (sP8('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)),X12,X13)) ) | $spl552),
  inference(cnf_transformation,[],[f21698_D])).
fof(f21698_D,plain,(
  ( ! [X12,X13] : (sP8('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)),X12,X13)) ) <=> ~$spl552),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl552])])).
fof(f15679,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)),X0,X1) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl128),
  inference(superposition,[],[f13208,f237])).
fof(f13208,plain,(
  ( ! [X105,X104] : (sP8('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)),X104,X105)) ) | $spl128),
  inference(resolution,[],[f6819,f287])).
fof(f15682,plain,(
  ( ! [X14,X12,X15,X13] : (sP8('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)),X12,X13) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl128),
  inference(superposition,[],[f13208,f239])).
fof(f21697,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) | $spl550),
  inference(cnf_transformation,[],[f21697_D])).
fof(f21697_D,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) <=> ~$spl550),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl550])])).
fof(f19094,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X0),s('type.num.num',X1))))) ) | $spl78),
  inference(superposition,[],[f1553,f237])).
fof(f19097,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X7),s('type.num.num',X6))))) ) | $spl78),
  inference(superposition,[],[f1553,f239])).
fof(f2594,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | $spl96),
  inference(cnf_transformation,[],[f2594_D])).
fof(f2594_D,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f19520,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X0),s('type.num.num',X1))))) ) | $spl90),
  inference(superposition,[],[f1780,f237])).
fof(f19523,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X7),s('type.num.num',X6))))) ) | $spl90),
  inference(superposition,[],[f1780,f239])).
fof(f8038,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X0),s('type.num.num',X1))))) ) | $spl13),
  inference(trivial_inequality_removal,[],[f8009])).
fof(f8009,plain,(
  ( ! [X0,X1] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X0),s('type.num.num',X1))))) ) | $spl13),
  inference(superposition,[],[f401,f237])).
fof(f19944,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X1),s('type.num.num',X2))))) ) | $spl90),
  inference(superposition,[],[f19909,f237])).
fof(f19909,plain,(
  ( ! [X282] : (sP11(X282,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')))) ) | $spl90),
  inference(subsumption_resolution,[],[f19723,f353])).
fof(f19723,plain,(
  ( ! [X282] : (~p(s(bool,'T_0')) | sP11(X282,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')))) ) | $spl90),
  inference(superposition,[],[f311,f1780])).
fof(f19947,plain,(
  ( ! [X10,X11,X9] : (sP11(X9,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X10),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X11),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X11),s('type.num.num',X10))))) ) | $spl90),
  inference(superposition,[],[f19909,f239])).
fof(f19976,plain,(
  ( ! [X2,X0,X1] : (sP11('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X0) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X1),s('type.num.num',X2))))) ) | $spl90),
  inference(superposition,[],[f19910,f237])).
fof(f19910,plain,(
  ( ! [X283] : (sP11('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')),X283)) ) | $spl90),
  inference(subsumption_resolution,[],[f19724,f353])).
fof(f19724,plain,(
  ( ! [X283] : (~p(s(bool,'T_0')) | sP11('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')),X283)) ) | $spl90),
  inference(superposition,[],[f312,f1780])).
fof(f19979,plain,(
  ( ! [X10,X11,X9] : (sP11('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X9) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X10),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X11),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X11),s('type.num.num',X10))))) ) | $spl90),
  inference(superposition,[],[f19910,f239])).
fof(f20096,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X0),s('type.num.num',X1))))) ) | $spl47),
  inference(superposition,[],[f552,f237])).
fof(f20099,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X7),s('type.num.num',X6))))) ) | $spl47),
  inference(superposition,[],[f552,f239])).
fof(f2595,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | $spl97),
  inference(cnf_transformation,[],[f2595_D])).
fof(f2595_D,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) <=> ~$spl97),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl97])])).
fof(f20148,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X0),s('type.num.num',X1))))) ) | ($spl13 | $spl90)),
  inference(superposition,[],[f19932,f237])).
fof(f20151,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X7),s('type.num.num',X6))))) ) | ($spl13 | $spl90)),
  inference(superposition,[],[f19932,f239])).
fof(f20183,plain,(
  ( ! [X2,X0,X3,X1] : (sP2(X0,X1,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl90),
  inference(superposition,[],[f19787,f237])).
fof(f19787,plain,(
  ( ! [X23,X22] : (sP2(X22,X23,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')))) ) | $spl90),
  inference(subsumption_resolution,[],[f19565,f353])).
fof(f19565,plain,(
  ( ! [X23,X22] : (~p(s(bool,'T_0')) | sP2(X22,X23,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')))) ) | $spl90),
  inference(superposition,[],[f208,f1780])).
fof(f20186,plain,(
  ( ! [X14,X12,X15,X13] : (sP2(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl90),
  inference(superposition,[],[f19787,f239])).
fof(f20215,plain,(
  ( ! [X2,X0,X3,X1] : (sP1(X0,X1,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl90),
  inference(superposition,[],[f19794,f237])).
fof(f19794,plain,(
  ( ! [X35,X34] : (sP1(X34,X35,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')))) ) | $spl90),
  inference(subsumption_resolution,[],[f19571,f353])).
fof(f19571,plain,(
  ( ! [X35,X34] : (~p(s(bool,'T_0')) | sP1(X34,X35,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')))) ) | $spl90),
  inference(superposition,[],[f212,f1780])).
fof(f20218,plain,(
  ( ! [X14,X12,X15,X13] : (sP1(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl90),
  inference(superposition,[],[f19794,f239])).
fof(f20247,plain,(
  ( ! [X2,X0,X3,X1] : (sP0(X0,X1,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl90),
  inference(superposition,[],[f19795,f237])).
fof(f19795,plain,(
  ( ! [X47,X46] : (sP0(X46,X47,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')))) ) | $spl90),
  inference(subsumption_resolution,[],[f19577,f353])).
fof(f19577,plain,(
  ( ! [X47,X46] : (~p(s(bool,'T_0')) | sP0(X46,X47,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')))) ) | $spl90),
  inference(superposition,[],[f216,f1780])).
fof(f20250,plain,(
  ( ! [X14,X12,X15,X13] : (sP0(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl90),
  inference(superposition,[],[f19795,f239])).
fof(f20279,plain,(
  ( ! [X2,X0,X3,X1] : (sP0(X0,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X1) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl90),
  inference(superposition,[],[f19796,f237])).
fof(f19796,plain,(
  ( ! [X48,X49] : (sP0(X48,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')),X49)) ) | $spl90),
  inference(subsumption_resolution,[],[f19578,f353])).
fof(f19578,plain,(
  ( ! [X48,X49] : (~p(s(bool,'T_0')) | sP0(X48,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')),X49)) ) | $spl90),
  inference(superposition,[],[f217,f1780])).
fof(f20282,plain,(
  ( ! [X14,X12,X15,X13] : (sP0(X12,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X13) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl90),
  inference(superposition,[],[f19796,f239])).
fof(f20311,plain,(
  ( ! [X2,X0,X3,X1] : (sP0('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X0,X1) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl90),
  inference(superposition,[],[f19797,f237])).
fof(f19797,plain,(
  ( ! [X50,X51] : (sP0('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')),X50,X51)) ) | $spl90),
  inference(subsumption_resolution,[],[f19579,f353])).
fof(f19579,plain,(
  ( ! [X50,X51] : (~p(s(bool,'T_0')) | sP0('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')),X50,X51)) ) | $spl90),
  inference(superposition,[],[f218,f1780])).
fof(f20314,plain,(
  ( ! [X14,X12,X15,X13] : (sP0('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X12,X13) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl90),
  inference(superposition,[],[f19797,f239])).
fof(f20726,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X0,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X1) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl90),
  inference(superposition,[],[f19834,f237])).
fof(f19834,plain,(
  ( ! [X112,X113] : (sP4(X112,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')),X113)) ) | $spl90),
  inference(subsumption_resolution,[],[f19629,f353])).
fof(f19629,plain,(
  ( ! [X112,X113] : (~p(s(bool,'T_0')) | sP4(X112,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')),X113)) ) | $spl90),
  inference(superposition,[],[f252,f1780])).
fof(f20729,plain,(
  ( ! [X14,X12,X15,X13] : (sP4(X12,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X13) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl90),
  inference(superposition,[],[f19834,f239])).
fof(f20760,plain,(
  ( ! [X2,X0,X3,X1] : (sP6(X0,X1,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl90),
  inference(superposition,[],[f19858,f237])).
fof(f19858,plain,(
  ( ! [X171,X170] : (sP6(X170,X171,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')))) ) | $spl90),
  inference(subsumption_resolution,[],[f19658,f353])).
fof(f19658,plain,(
  ( ! [X171,X170] : (~p(s(bool,'T_0')) | sP6(X170,X171,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')))) ) | $spl90),
  inference(superposition,[],[f268,f1780])).
fof(f20763,plain,(
  ( ! [X14,X12,X15,X13] : (sP6(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl90),
  inference(superposition,[],[f19858,f239])).
fof(f20791,plain,(
  ( ! [X2,X0,X3,X1] : (sP6(X0,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X1) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl90),
  inference(superposition,[],[f19859,f237])).
fof(f19859,plain,(
  ( ! [X173,X172] : (sP6(X172,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')),X173)) ) | $spl90),
  inference(subsumption_resolution,[],[f19659,f353])).
fof(f19659,plain,(
  ( ! [X173,X172] : (~p(s(bool,'T_0')) | sP6(X172,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')),X173)) ) | $spl90),
  inference(superposition,[],[f269,f1780])).
fof(f20794,plain,(
  ( ! [X14,X12,X15,X13] : (sP6(X12,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X13) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl90),
  inference(superposition,[],[f19859,f239])).
fof(f20822,plain,(
  ( ! [X2,X0,X3,X1] : (sP8(X0,X1,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl90),
  inference(superposition,[],[f19885,f237])).
fof(f19885,plain,(
  ( ! [X231,X230] : (sP8(X230,X231,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')))) ) | $spl90),
  inference(subsumption_resolution,[],[f19688,f353])).
fof(f19688,plain,(
  ( ! [X231,X230] : (~p(s(bool,'T_0')) | sP8(X230,X231,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')))) ) | $spl90),
  inference(superposition,[],[f285,f1780])).
fof(f20825,plain,(
  ( ! [X14,X12,X15,X13] : (sP8(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl90),
  inference(superposition,[],[f19885,f239])).
fof(f21676,plain,(
  ( ! [X14,X13] : (s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X13),s('type.num.num',X14)))) ) | $spl524),
  inference(cnf_transformation,[],[f21676_D])).
fof(f21676_D,plain,(
  ( ! [X14,X13] : (s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X13),s('type.num.num',X14)))) ) <=> ~$spl524),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl524])])).
fof(f21663,plain,(
  ( ! [X14,X15,X13] : (s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X13),s('type.num.num',X14))) | ~p(s(bool,X15))) ) | ($spl16 | $spl24 | $spl30)),
  inference(resolution,[],[f17046,f4693])).
fof(f4693,plain,(
  ( ! [X23,X21,X22,X20] : (sP2(X20,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X21),s('type.num.num',X22)),X23)) ) | $spl16),
  inference(resolution,[],[f403,f209])).
fof(f17046,plain,(
  ( ! [X12,X13] : (~sP2(X12,X13,'F_0') | s(bool,'T_0') != s(bool,X13) | ~p(s(bool,X12))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f17045,f232])).
fof(f17045,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | s(bool,'T_0') != s(bool,X13) | ~p(s(bool,X13)) | ~sP2(X12,X13,'F_0')) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f17044,f16933])).
fof(f17044,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | s(bool,'T_0') != s(bool,X13) | ~p(s(bool,X13)) | ~sP2(X12,X13,'F_0') | ~sP0(X13,'F_0',X12)) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f17034,f16934])).
fof(f17034,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | s(bool,'T_0') != s(bool,X13) | ~p(s(bool,X13)) | ~sP2(X12,X13,'F_0') | ~sP1('F_0',X13,X12) | ~sP0(X13,'F_0',X12)) ) | ($spl24 | $spl30)),
  inference(resolution,[],[f5940,f204])).
fof(f5940,plain,(
  ( ! [X10,X9] : (~sP3(X10,'F_0',X9) | ~p(s(bool,X10)) | s(bool,'T_0') != s(bool,X9)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f5931,f221])).
fof(f19931,plain,(
  $false | ($spl90 | $spl191)),
  inference(subsumption_resolution,[],[f19930,f354])).
fof(f354,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.arithmetic.MODEQ_SYM',aHL_TRUTH)).
fof(f19930,plain,(
  ~p(s(bool,'T_0')) | ($spl90 | $spl191)),
  inference(forward_demodulation,[],[f8039,f1780])).
fof(f11888,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num',X2),s('type.num.num',X3)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X1),s('type.num.num',sK13),s('type.num.num',X0))))) ) | $spl430),
  inference(cnf_transformation,[],[f11888_D])).
fof(f11888_D,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num',X2),s('type.num.num',X3)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X1),s('type.num.num',sK13),s('type.num.num',X0))))) ) <=> ~$spl430),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl430])])).
fof(f8046,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) | $spl199),
  inference(cnf_transformation,[],[f8046_D])).
fof(f8046_D,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) <=> ~$spl199),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl199])])).
fof(f19929,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X1),s('type.num.num',sK13),s('type.num.num',X0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl90),
  inference(forward_demodulation,[],[f12631,f1780])).
fof(f12631,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X1),s('type.num.num',sK13),s('type.num.num',X0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num',X2),s('type.num.num',X3))))) )),
  inference(superposition,[],[f366,f237])).
fof(f11889,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',X15),s('type.num.num',X14)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',X12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',X12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',sK13),s('type.num.num',X12))))) ) | $spl432),
  inference(cnf_transformation,[],[f11889_D])).
fof(f11889_D,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',X15),s('type.num.num',X14)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',X12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',X12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',sK13),s('type.num.num',X12))))) ) <=> ~$spl432),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl432])])).
fof(f19928,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',sK13),s('type.num.num',X12)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',X12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',X12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl90),
  inference(forward_demodulation,[],[f12634,f1780])).
fof(f12634,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',sK13),s('type.num.num',X12)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',X12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',X12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',X15),s('type.num.num',X14))))) )),
  inference(superposition,[],[f366,f239])).
fof(f11940,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num',X2),s('type.num.num',X3)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X1),s('type.num.num',X0),s('type.num.num',sK13))))) ) | $spl442),
  inference(cnf_transformation,[],[f11940_D])).
fof(f11940_D,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num',X2),s('type.num.num',X3)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X1),s('type.num.num',X0),s('type.num.num',sK13))))) ) <=> ~$spl442),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl442])])).
fof(f19927,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X1),s('type.num.num',X0),s('type.num.num',sK13)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl90),
  inference(forward_demodulation,[],[f12821,f1780])).
fof(f12821,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X1),s('type.num.num',X0),s('type.num.num',sK13)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num',X2),s('type.num.num',X3))))) )),
  inference(superposition,[],[f367,f237])).
fof(f11941,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',X15),s('type.num.num',X14)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',X12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',X12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',X12),s('type.num.num',sK13))))) ) | $spl444),
  inference(cnf_transformation,[],[f11941_D])).
fof(f11941_D,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',X15),s('type.num.num',X14)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',X12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',X12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',X12),s('type.num.num',sK13))))) ) <=> ~$spl444),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl444])])).
fof(f19926,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',X12),s('type.num.num',sK13)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',X12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',X12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl90),
  inference(forward_demodulation,[],[f12824,f1780])).
fof(f12824,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',X12),s('type.num.num',sK13)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',X12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',X12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',X15),s('type.num.num',X14))))) )),
  inference(superposition,[],[f367,f239])).
fof(f14149,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X0),s('type.num.num',X1))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num',X2),s('type.num.num',X3)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X1),s('type.num.num',sK13),s('type.num.num',X0))))) ) | $spl490),
  inference(cnf_transformation,[],[f14149_D])).
fof(f14149_D,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X0),s('type.num.num',X1))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num',X2),s('type.num.num',X3)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X1),s('type.num.num',sK13),s('type.num.num',X0))))) ) <=> ~$spl490),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl490])])).
fof(f19925,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X0),s('type.num.num',X1))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X1),s('type.num.num',sK13),s('type.num.num',X0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl90),
  inference(forward_demodulation,[],[f14760,f1780])).
fof(f14760,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X0),s('type.num.num',X1))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X1),s('type.num.num',sK13),s('type.num.num',X0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num',X2),s('type.num.num',X3))))) )),
  inference(superposition,[],[f369,f237])).
fof(f14150,plain,(
  ( ! [X14,X12,X15,X13] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X13))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X12),s('type.num.num',X13))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',X15),s('type.num.num',X14)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',X12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',X12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',sK13),s('type.num.num',X12))))) ) | $spl492),
  inference(cnf_transformation,[],[f14150_D])).
fof(f14150_D,plain,(
  ( ! [X14,X12,X15,X13] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X13))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X12),s('type.num.num',X13))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',X15),s('type.num.num',X14)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',X12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',X12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',sK13),s('type.num.num',X12))))) ) <=> ~$spl492),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl492])])).
fof(f19924,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X13))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X12),s('type.num.num',X13))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',sK13),s('type.num.num',X12)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',X12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',X12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl90),
  inference(forward_demodulation,[],[f14763,f1780])).
fof(f14763,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X13))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X12),s('type.num.num',X13))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',sK13),s('type.num.num',X12)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',X12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',X12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',X15),s('type.num.num',X14))))) )),
  inference(superposition,[],[f369,f239])).
fof(f14341,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X0),s('type.num.num',X1))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num',X2),s('type.num.num',X3)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X1),s('type.num.num',X0),s('type.num.num',sK13))))) ) | $spl498),
  inference(cnf_transformation,[],[f14341_D])).
fof(f14341_D,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X0),s('type.num.num',X1))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num',X2),s('type.num.num',X3)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X1),s('type.num.num',X0),s('type.num.num',sK13))))) ) <=> ~$spl498),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl498])])).
fof(f19923,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X0),s('type.num.num',X1))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X1),s('type.num.num',X0),s('type.num.num',sK13)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl90),
  inference(forward_demodulation,[],[f14822,f1780])).
fof(f14822,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X0),s('type.num.num',X1))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X1),s('type.num.num',X0),s('type.num.num',sK13)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num',X2),s('type.num.num',X3))))) )),
  inference(superposition,[],[f370,f237])).
fof(f14342,plain,(
  ( ! [X14,X12,X15,X13] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X13))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X12),s('type.num.num',X13))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',X15),s('type.num.num',X14)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',X12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',X12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',X12),s('type.num.num',sK13))))) ) | $spl500),
  inference(cnf_transformation,[],[f14342_D])).
fof(f14342_D,plain,(
  ( ! [X14,X12,X15,X13] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X13))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X12),s('type.num.num',X13))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',X15),s('type.num.num',X14)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',X12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',X12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',X12),s('type.num.num',sK13))))) ) <=> ~$spl500),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl500])])).
fof(f19922,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X13))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X12),s('type.num.num',X13))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',X12),s('type.num.num',sK13)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',X12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',X12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl90),
  inference(forward_demodulation,[],[f14825,f1780])).
fof(f14825,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X13))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X12),s('type.num.num',X13))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',X12),s('type.num.num',sK13)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',X12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',X12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',X15),s('type.num.num',X14))))) )),
  inference(superposition,[],[f370,f239])).
fof(f19784,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | p(s(bool,X19))) ) | ($spl12 | $spl47 | $spl90)),
  inference(subsumption_resolution,[],[f19783,f8767])).
fof(f8767,plain,(
  ( ! [X23,X22] : (sP2('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')),X22,X23)) ) | ($spl12 | $spl47)),
  inference(subsumption_resolution,[],[f8579,f552])).
fof(f8579,plain,(
  ( ! [X23,X22] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | sP2('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')),X22,X23)) ) | $spl12),
  inference(superposition,[],[f210,f400])).
fof(f400,plain,(
  s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0'))) | $spl12),
  inference(cnf_transformation,[],[f400_D])).
fof(f400_D,plain,(
  s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0'))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f19783,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | p(s(bool,X19)) | ~sP2('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')),X18,X19)) ) | $spl90),
  inference(subsumption_resolution,[],[f19563,f337])).
fof(f337,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f194])).
fof(f194,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f193])).
fof(f193,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f108])).
fof(f108,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X13] : ((p(s(bool,'F_0')) => p(s(bool,X13))) <=> p(s(bool,'T_0')))),
  file('thm.arithmetic.MODEQ_SYM','thm.bool.IMP_CLAUSES|split|2')).
fof(f19563,plain,(
  ( ! [X19,X18] : (~p(s(bool,'T_0')) | ~p(s(bool,X18)) | p(s(bool,X19)) | ~sP2('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')),X18,X19)) ) | $spl90),
  inference(superposition,[],[f207,f1780])).
fof(f207,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f139])).
fof(f19786,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | p(s(bool,X21))) ) | ($spl12 | $spl47 | $spl90)),
  inference(subsumption_resolution,[],[f19785,f8766])).
fof(f8766,plain,(
  ( ! [X21,X20] : (sP2(X20,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')),X21)) ) | ($spl12 | $spl47)),
  inference(subsumption_resolution,[],[f8578,f552])).
fof(f8578,plain,(
  ( ! [X21,X20] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | sP2(X20,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')),X21)) ) | $spl12),
  inference(superposition,[],[f209,f400])).
fof(f19785,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | p(s(bool,X21)) | ~sP2(X20,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')),X21)) ) | $spl90),
  inference(subsumption_resolution,[],[f19564,f337])).
fof(f19564,plain,(
  ( ! [X21,X20] : (~p(s(bool,'T_0')) | ~p(s(bool,X20)) | p(s(bool,X21)) | ~sP2(X20,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')),X21)) ) | $spl90),
  inference(superposition,[],[f207,f1780])).
fof(f19791,plain,(
  ( ! [X30,X31] : (~p(s(bool,X30)) | p(s(bool,X31))) ) | ($spl12 | $spl47 | $spl90)),
  inference(subsumption_resolution,[],[f19790,f8770])).
fof(f8770,plain,(
  ( ! [X35,X34] : (sP1('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')),X34,X35)) ) | ($spl12 | $spl47)),
  inference(subsumption_resolution,[],[f8585,f552])).
fof(f8585,plain,(
  ( ! [X35,X34] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | sP1('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')),X34,X35)) ) | $spl12),
  inference(superposition,[],[f214,f400])).
fof(f19790,plain,(
  ( ! [X30,X31] : (~p(s(bool,X30)) | p(s(bool,X31)) | ~sP1('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')),X30,X31)) ) | $spl90),
  inference(subsumption_resolution,[],[f19569,f337])).
fof(f19569,plain,(
  ( ! [X30,X31] : (~p(s(bool,'T_0')) | ~p(s(bool,X30)) | p(s(bool,X31)) | ~sP1('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')),X30,X31)) ) | $spl90),
  inference(superposition,[],[f211,f1780])).
fof(f211,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP1(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f142])).
fof(f19793,plain,(
  ( ! [X33,X32] : (~p(s(bool,X32)) | p(s(bool,X33))) ) | ($spl12 | $spl47 | $spl90)),
  inference(subsumption_resolution,[],[f19792,f8769])).
fof(f8769,plain,(
  ( ! [X33,X32] : (sP1(X32,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')),X33)) ) | ($spl12 | $spl47)),
  inference(subsumption_resolution,[],[f8584,f552])).
fof(f8584,plain,(
  ( ! [X33,X32] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | sP1(X32,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')),X33)) ) | $spl12),
  inference(superposition,[],[f213,f400])).
fof(f19792,plain,(
  ( ! [X33,X32] : (~p(s(bool,X32)) | p(s(bool,X33)) | ~sP1(X32,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')),X33)) ) | $spl90),
  inference(subsumption_resolution,[],[f19570,f337])).
fof(f19570,plain,(
  ( ! [X33,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X32)) | p(s(bool,X33)) | ~sP1(X32,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')),X33)) ) | $spl90),
  inference(superposition,[],[f211,f1780])).
fof(f19831,plain,(
  ( ! [X107,X106] : (p(s(bool,X106)) | ~p(s(bool,X107))) ) | ($spl12 | $spl47 | $spl90)),
  inference(subsumption_resolution,[],[f19830,f8802])).
fof(f8802,plain,(
  ( ! [X111,X110] : (sP4('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')),X110,X111)) ) | ($spl12 | $spl47)),
  inference(subsumption_resolution,[],[f8645,f552])).
fof(f8645,plain,(
  ( ! [X111,X110] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | sP4('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')),X110,X111)) ) | $spl12),
  inference(superposition,[],[f253,f400])).
fof(f19830,plain,(
  ( ! [X107,X106] : (p(s(bool,X106)) | ~p(s(bool,X107)) | ~sP4('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')),X106,X107)) ) | $spl90),
  inference(subsumption_resolution,[],[f19626,f337])).
fof(f19626,plain,(
  ( ! [X107,X106] : (~p(s(bool,'T_0')) | p(s(bool,X106)) | ~p(s(bool,X107)) | ~sP4('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')),X106,X107)) ) | $spl90),
  inference(superposition,[],[f250,f1780])).
fof(f250,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f157])).
fof(f19833,plain,(
  ( ! [X109,X108] : (~p(s(bool,X108)) | p(s(bool,X109))) ) | ($spl12 | $spl47 | $spl90)),
  inference(subsumption_resolution,[],[f19832,f8801])).
fof(f8801,plain,(
  ( ! [X107,X106] : (sP4(X106,X107,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')))) ) | ($spl12 | $spl47)),
  inference(subsumption_resolution,[],[f8643,f552])).
fof(f8643,plain,(
  ( ! [X107,X106] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | sP4(X106,X107,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')))) ) | $spl12),
  inference(superposition,[],[f251,f400])).
fof(f19832,plain,(
  ( ! [X109,X108] : (~p(s(bool,X108)) | p(s(bool,X109)) | ~sP4(X108,X109,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')))) ) | $spl90),
  inference(subsumption_resolution,[],[f19627,f337])).
fof(f19627,plain,(
  ( ! [X109,X108] : (~p(s(bool,'T_0')) | ~p(s(bool,X108)) | p(s(bool,X109)) | ~sP4(X108,X109,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')))) ) | $spl90),
  inference(superposition,[],[f250,f1780])).
fof(f19857,plain,(
  ( ! [X169,X168] : (p(s(bool,X168)) | p(s(bool,X169))) ) | ($spl12 | $spl47 | $spl90)),
  inference(subsumption_resolution,[],[f19856,f8816])).
fof(f8816,plain,(
  ( ! [X171,X170] : (sP6('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')),X170,X171)) ) | ($spl12 | $spl47)),
  inference(subsumption_resolution,[],[f8675,f552])).
fof(f8675,plain,(
  ( ! [X171,X170] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | sP6('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')),X170,X171)) ) | $spl12),
  inference(superposition,[],[f270,f400])).
fof(f19856,plain,(
  ( ! [X169,X168] : (p(s(bool,X168)) | p(s(bool,X169)) | ~sP6('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')),X168,X169)) ) | $spl90),
  inference(subsumption_resolution,[],[f19657,f330])).
fof(f330,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f190])).
fof(f190,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f189])).
fof(f189,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (~p(s(bool,'T_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f106])).
fof(f106,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(ennf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0] : ((p(s(bool,'T_0')) => p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X13] : ((p(s(bool,'T_0')) => p(s(bool,X13))) <=> p(s(bool,X13)))),
  file('thm.arithmetic.MODEQ_SYM','thm.bool.IMP_CLAUSES|split|0')).
fof(f19657,plain,(
  ( ! [X169,X168] : (~p(s(bool,'T_0')) | p(s(bool,X168)) | p(s(bool,X169)) | ~sP6('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')),X168,X169)) ) | $spl90),
  inference(superposition,[],[f267,f1780])).
fof(f267,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f165])).
fof(f19882,plain,(
  ( ! [X227,X226] : (~p(s(bool,X226)) | p(s(bool,X227))) ) | ($spl12 | $spl47 | $spl90)),
  inference(subsumption_resolution,[],[f19881,f8830])).
fof(f8830,plain,(
  ( ! [X231,X230] : (sP8('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')),X230,X231)) ) | ($spl12 | $spl47)),
  inference(subsumption_resolution,[],[f8705,f552])).
fof(f8705,plain,(
  ( ! [X231,X230] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | sP8('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')),X230,X231)) ) | $spl12),
  inference(superposition,[],[f287,f400])).
fof(f19881,plain,(
  ( ! [X227,X226] : (~p(s(bool,X226)) | p(s(bool,X227)) | ~sP8('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')),X226,X227)) ) | $spl90),
  inference(subsumption_resolution,[],[f19686,f337])).
fof(f19686,plain,(
  ( ! [X227,X226] : (~p(s(bool,'T_0')) | ~p(s(bool,X226)) | p(s(bool,X227)) | ~sP8('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')),X226,X227)) ) | $spl90),
  inference(superposition,[],[f284,f1780])).
fof(f284,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f172])).
fof(f19884,plain,(
  ( ! [X229,X228] : (~p(s(bool,X228)) | p(s(bool,X229))) ) | ($spl12 | $spl47 | $spl90)),
  inference(subsumption_resolution,[],[f19883,f8829])).
fof(f8829,plain,(
  ( ! [X229,X228] : (sP8(X228,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')),X229)) ) | ($spl12 | $spl47)),
  inference(subsumption_resolution,[],[f8704,f552])).
fof(f8704,plain,(
  ( ! [X229,X228] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | sP8(X228,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')),X229)) ) | $spl12),
  inference(superposition,[],[f286,f400])).
fof(f19883,plain,(
  ( ! [X229,X228] : (~p(s(bool,X228)) | p(s(bool,X229)) | ~sP8(X228,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')),X229)) ) | $spl90),
  inference(subsumption_resolution,[],[f19687,f337])).
fof(f19687,plain,(
  ( ! [X229,X228] : (~p(s(bool,'T_0')) | ~p(s(bool,X228)) | p(s(bool,X229)) | ~sP8(X228,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')),X229)) ) | $spl90),
  inference(superposition,[],[f284,f1780])).
fof(f14343,plain,(
  ( ! [X25] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X25))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X25))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X25),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))))) ) | $spl502),
  inference(cnf_transformation,[],[f14343_D])).
fof(f14343_D,plain,(
  ( ! [X25] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X25))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X25))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X25),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))))) ) <=> ~$spl502),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl502])])).
fof(f19545,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X0))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))))) ) | $spl90),
  inference(superposition,[],[f370,f1780])).
fof(f14151,plain,(
  ( ! [X25] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X25))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X25))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X25),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))))) ) | $spl494),
  inference(cnf_transformation,[],[f14151_D])).
fof(f14151_D,plain,(
  ( ! [X25] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X25))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X25))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X25),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))))) ) <=> ~$spl494),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl494])])).
fof(f19546,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X1))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X1),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))))) ) | $spl90),
  inference(superposition,[],[f369,f1780])).
fof(f11942,plain,(
  ( ! [X25] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X25)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X25),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))))) ) | $spl446),
  inference(cnf_transformation,[],[f11942_D])).
fof(f11942_D,plain,(
  ( ! [X25] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X25)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X25),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))))) ) <=> ~$spl446),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl446])])).
fof(f19547,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X2)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X2),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))))) ) | $spl90),
  inference(superposition,[],[f367,f1780])).
fof(f11890,plain,(
  ( ! [X25] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X25)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X25),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))))) ) | $spl434),
  inference(cnf_transformation,[],[f11890_D])).
fof(f11890_D,plain,(
  ( ! [X25] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X25)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X25),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))))) ) <=> ~$spl434),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl434])])).
fof(f19548,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X3)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X3),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))))) ) | $spl90),
  inference(superposition,[],[f366,f1780])).
fof(f19747,plain,(
  $false | ($spl24 | $spl30 | $spl54 | $spl90)),
  inference(subsumption_resolution,[],[f19513,f5931])).
fof(f19513,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl54 | $spl90)),
  inference(backward_demodulation,[],[f1780,f908])).
fof(f553,plain,(
  p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')))) | $spl48),
  inference(cnf_transformation,[],[f553_D])).
fof(f553_D,plain,(
  p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')))) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f17810,plain,(
  ( ! [X8] : (p(s(bool,X8)) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))))) ) | ($spl21 | $spl52)),
  inference(subsumption_resolution,[],[f17809,f229])).
fof(f17809,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,X8) | p(s(bool,X8)) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))))) ) | ($spl21 | $spl52)),
  inference(forward_demodulation,[],[f17214,f906])).
fof(f906,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) | $spl52),
  inference(cnf_transformation,[],[f906_D])).
fof(f906_D,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f17214,plain,(
  ( ! [X8] : (s(bool,X8) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) | p(s(bool,X8)) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))))) ) | $spl21),
  inference(superposition,[],[f407,f223])).
fof(f223,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f91])).
fof(f17808,plain,(
  ( ! [X9] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X9))) ) | ($spl21 | $spl52)),
  inference(subsumption_resolution,[],[f17807,f229])).
fof(f17807,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,X9) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X9))) ) | ($spl21 | $spl52)),
  inference(forward_demodulation,[],[f17215,f906])).
fof(f17215,plain,(
  ( ! [X9] : (s(bool,X9) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X9))) ) | $spl21),
  inference(superposition,[],[f407,f223])).
fof(f17795,plain,(
  ( ! [X8] : (p(s(bool,X8)) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))))) ) | ($spl21 | $spl52)),
  inference(subsumption_resolution,[],[f17780,f229])).
fof(f17780,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,X8) | p(s(bool,X8)) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))))) ) | ($spl21 | $spl52)),
  inference(superposition,[],[f17232,f223])).
fof(f17232,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | ($spl21 | $spl52)),
  inference(backward_demodulation,[],[f906,f407])).
fof(f17797,plain,(
  ( ! [X9] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X9))) ) | ($spl21 | $spl52)),
  inference(subsumption_resolution,[],[f17781,f229])).
fof(f17781,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,X9) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X9))) ) | ($spl21 | $spl52)),
  inference(superposition,[],[f17232,f223])).
fof(f17800,plain,(
  $false | ($spl21 | $spl49 | $spl52)),
  inference(subsumption_resolution,[],[f17791,f554])).
fof(f17791,plain,(
  p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')))) | ($spl21 | $spl52)),
  inference(trivial_inequality_removal,[],[f17785])).
fof(f17785,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')))) | ($spl21 | $spl52)),
  inference(superposition,[],[f17232,f231])).
fof(f231,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f46])).
fof(f46,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X13] : (s(bool,'F_0') = s(bool,X13) <=> ~p(s(bool,X13)))),
  file('thm.arithmetic.MODEQ_SYM','thm.bool.EQ_CLAUSES|split|2')).
fof(f17799,plain,(
  $false | ($spl21 | $spl49 | $spl52)),
  inference(subsumption_resolution,[],[f17792,f554])).
fof(f17792,plain,(
  p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')))) | ($spl21 | $spl52)),
  inference(trivial_inequality_removal,[],[f17784])).
fof(f17784,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')))) | ($spl21 | $spl52)),
  inference(superposition,[],[f17232,f229])).
fof(f907,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) | $spl53),
  inference(cnf_transformation,[],[f907_D])).
fof(f907_D,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) <=> ~$spl53),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl53])])).
fof(f17223,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) | s(bool,'T_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | $spl21),
  inference(superposition,[],[f407,f236])).
fof(f236,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.arithmetic.MODEQ_SYM',aHL_BOOL_CASES)).
fof(f8983,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | $spl200),
  inference(cnf_transformation,[],[f8983_D])).
fof(f8983_D,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f1779,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) | $spl89),
  inference(cnf_transformation,[],[f1779_D])).
fof(f1779_D,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) <=> ~$spl89),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl89])])).
fof(f17222,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) | s(bool,'F_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | $spl21),
  inference(superposition,[],[f407,f236])).
fof(f12170,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num',X2),s('type.num.num',X3)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X1),s('type.num.num',sK12),s('type.num.num',X0))))) ) | $spl450),
  inference(cnf_transformation,[],[f12170_D])).
fof(f12170_D,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num',X2),s('type.num.num',X3)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X1),s('type.num.num',sK12),s('type.num.num',X0))))) ) <=> ~$spl450),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl450])])).
fof(f407,plain,(
  s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | $spl21),
  inference(cnf_transformation,[],[f407_D])).
fof(f407_D,plain,(
  s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f13032,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X1),s('type.num.num',sK12),s('type.num.num',X0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num',X2),s('type.num.num',X3))))) )),
  inference(superposition,[],[f372,f237])).
fof(f372,plain,(
  ( ! [X2,X3] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X2),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',X2))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X3)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X3),s('type.num.num',sK12),s('type.num.num',X2))))) )),
  inference(superposition,[],[f199,f238])).
fof(f12171,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',X15),s('type.num.num',X14)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',X12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',X12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',sK12),s('type.num.num',X12))))) ) | $spl452),
  inference(cnf_transformation,[],[f12171_D])).
fof(f12171_D,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',X15),s('type.num.num',X14)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',X12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',X12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',sK12),s('type.num.num',X12))))) ) <=> ~$spl452),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl452])])).
fof(f13035,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',sK12),s('type.num.num',X12)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',X12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',X12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',X15),s('type.num.num',X14))))) )),
  inference(superposition,[],[f372,f239])).
fof(f12367,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num',X2),s('type.num.num',X3)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X1),s('type.num.num',X0),s('type.num.num',sK12))))) ) | $spl462),
  inference(cnf_transformation,[],[f12367_D])).
fof(f12367_D,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num',X2),s('type.num.num',X3)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X1),s('type.num.num',X0),s('type.num.num',sK12))))) ) <=> ~$spl462),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl462])])).
fof(f13119,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X1),s('type.num.num',X0),s('type.num.num',sK12)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num',X2),s('type.num.num',X3))))) )),
  inference(superposition,[],[f373,f237])).
fof(f373,plain,(
  ( ! [X4,X5] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X4),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',X4))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X5)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X5),s('type.num.num',X4),s('type.num.num',sK12))))) )),
  inference(superposition,[],[f199,f238])).
fof(f12368,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',X15),s('type.num.num',X14)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',X12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',X12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',X12),s('type.num.num',sK12))))) ) | $spl464),
  inference(cnf_transformation,[],[f12368_D])).
fof(f12368_D,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',X15),s('type.num.num',X14)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',X12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',X12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',X12),s('type.num.num',sK12))))) ) <=> ~$spl464),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl464])])).
fof(f13122,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',X12),s('type.num.num',sK12)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',X12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',X12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',X15),s('type.num.num',X14))))) )),
  inference(superposition,[],[f373,f239])).
fof(f14534,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X0),s('type.num.num',X1))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num',X2),s('type.num.num',X3)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X1),s('type.num.num',sK12),s('type.num.num',X0))))) ) | $spl506),
  inference(cnf_transformation,[],[f14534_D])).
fof(f14534_D,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X0),s('type.num.num',X1))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num',X2),s('type.num.num',X3)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X1),s('type.num.num',sK12),s('type.num.num',X0))))) ) <=> ~$spl506),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl506])])).
fof(f14883,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X0),s('type.num.num',X1))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X1),s('type.num.num',sK12),s('type.num.num',X0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num',X2),s('type.num.num',X3))))) )),
  inference(superposition,[],[f375,f237])).
fof(f375,plain,(
  ( ! [X8,X9] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X8),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',X8))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',X9))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X8),s('type.num.num',X9))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X9),s('type.num.num',sK12),s('type.num.num',X8))))) )),
  inference(superposition,[],[f199,f240])).
fof(f14535,plain,(
  ( ! [X14,X12,X15,X13] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',X13))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X12),s('type.num.num',X13))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',X15),s('type.num.num',X14)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',X12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',X12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',sK12),s('type.num.num',X12))))) ) | $spl508),
  inference(cnf_transformation,[],[f14535_D])).
fof(f14535_D,plain,(
  ( ! [X14,X12,X15,X13] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',X13))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X12),s('type.num.num',X13))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',X15),s('type.num.num',X14)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',X12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',X12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',sK12),s('type.num.num',X12))))) ) <=> ~$spl508),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl508])])).
fof(f14886,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',X13))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X12),s('type.num.num',X13))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',sK12),s('type.num.num',X12)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',X12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',X12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',X15),s('type.num.num',X14))))) )),
  inference(superposition,[],[f375,f239])).
fof(f14641,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X0),s('type.num.num',X1))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num',X2),s('type.num.num',X3)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X1),s('type.num.num',X0),s('type.num.num',sK12))))) ) | $spl514),
  inference(cnf_transformation,[],[f14641_D])).
fof(f14641_D,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X0),s('type.num.num',X1))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num',X2),s('type.num.num',X3)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X1),s('type.num.num',X0),s('type.num.num',sK12))))) ) <=> ~$spl514),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl514])])).
fof(f14941,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X0),s('type.num.num',X1))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X1),s('type.num.num',X0),s('type.num.num',sK12)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num',X2),s('type.num.num',X3))))) )),
  inference(superposition,[],[f376,f237])).
fof(f376,plain,(
  ( ! [X10,X11] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X10),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',X10))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',X11))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X10),s('type.num.num',X11))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X11),s('type.num.num',X10),s('type.num.num',sK12))))) )),
  inference(superposition,[],[f199,f240])).
fof(f14642,plain,(
  ( ! [X14,X12,X15,X13] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',X13))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X12),s('type.num.num',X13))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',X15),s('type.num.num',X14)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',X12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',X12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',X12),s('type.num.num',sK12))))) ) | $spl516),
  inference(cnf_transformation,[],[f14642_D])).
fof(f14642_D,plain,(
  ( ! [X14,X12,X15,X13] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',X13))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X12),s('type.num.num',X13))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',X15),s('type.num.num',X14)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',X12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',X12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',X12),s('type.num.num',sK12))))) ) <=> ~$spl516),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl516])])).
fof(f14944,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',X13))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X12),s('type.num.num',X13))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',X12),s('type.num.num',sK12)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',X12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',X12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',X15),s('type.num.num',X14))))) )),
  inference(superposition,[],[f376,f239])).
fof(f17188,plain,(
  p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)))) | $spl522),
  inference(cnf_transformation,[],[f17188_D])).
fof(f17188_D,plain,(
  p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)))) <=> ~$spl522),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl522])])).
fof(f17189,plain,(
  ( ! [X9] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)))) | p(s(bool,X9))) ) | $spl53),
  inference(subsumption_resolution,[],[f17172,f229])).
fof(f17172,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,X9) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)))) | p(s(bool,X9))) ) | $spl53),
  inference(superposition,[],[f907,f223])).
fof(f17186,plain,(
  ( ! [X8] : (p(s(bool,X8)) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))))) ) | $spl53),
  inference(subsumption_resolution,[],[f17171,f229])).
fof(f17171,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,X8) | p(s(bool,X8)) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))))) ) | $spl53),
  inference(superposition,[],[f907,f223])).
fof(f1781,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0'))) | $spl91),
  inference(cnf_transformation,[],[f1781_D])).
fof(f1781_D,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0'))) <=> ~$spl91),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl91])])).
fof(f8044,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) | $spl197),
  inference(cnf_transformation,[],[f8044_D])).
fof(f8044_D,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) <=> ~$spl197),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl197])])).
fof(f16883,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) | s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0'))) | ($spl12 | $spl160)),
  inference(superposition,[],[f6892,f400])).
fof(f14643,plain,(
  ( ! [X25] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X25))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',X25))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X25),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))))) ) | $spl518),
  inference(cnf_transformation,[],[f14643_D])).
fof(f14643_D,plain,(
  ( ! [X25] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X25))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',X25))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X25),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))))) ) <=> ~$spl518),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl518])])).
fof(f13062,plain,(
  ( ! [X30,X28,X31] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X28),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',X28))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X28),s('type.num.num',X31),s('type.num.num',X30)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X30),s('type.num.num',X28))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X31),s('type.num.num',X28)))) ) | $spl480),
  inference(cnf_transformation,[],[f13062_D])).
fof(f13062_D,plain,(
  ( ! [X30,X28,X31] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X28),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',X28))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X28),s('type.num.num',X31),s('type.num.num',X30)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X30),s('type.num.num',X28))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X31),s('type.num.num',X28)))) ) <=> ~$spl480),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl480])])).
fof(f14948,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X28),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',X28))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X29))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',X29))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X29),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X30),s('type.num.num',X28))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X31),s('type.num.num',X28))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X28),s('type.num.num',X31),s('type.num.num',X30))))) )),
  inference(superposition,[],[f376,f239])).
fof(f13061,plain,(
  ( ! [X26,X24,X27] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X24),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',X24))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X24),s('type.num.num',X26),s('type.num.num',X27)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X24))))) ) | $spl478),
  inference(cnf_transformation,[],[f13061_D])).
fof(f13061_D,plain,(
  ( ! [X26,X24,X27] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X24),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',X24))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X24),s('type.num.num',X26),s('type.num.num',X27)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X24))))) ) <=> ~$spl478),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl478])])).
fof(f14947,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X24),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',X24))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X25))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',X25))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X25),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X24)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X24),s('type.num.num',X26),s('type.num.num',X27))))) )),
  inference(superposition,[],[f376,f237])).
fof(f14536,plain,(
  ( ! [X25] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X25))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',X25))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X25),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0'))))) ) | $spl510),
  inference(cnf_transformation,[],[f14536_D])).
fof(f14536_D,plain,(
  ( ! [X25] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X25))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',X25))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X25),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0'))))) ) <=> ~$spl510),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl510])])).
fof(f14890,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X28),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',X28))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X29))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',X29))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X29),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X30),s('type.num.num',X28))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X31),s('type.num.num',X28))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X28),s('type.num.num',X31),s('type.num.num',X30))))) )),
  inference(superposition,[],[f375,f239])).
fof(f14889,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X24),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',X24))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X25))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',X25))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X25),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X24)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X24),s('type.num.num',X26),s('type.num.num',X27))))) )),
  inference(superposition,[],[f375,f237])).
fof(f12665,plain,(
  ( ! [X30,X28,X31] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X28),s('type.num.num',sK12))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',X28))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X28),s('type.num.num',X31),s('type.num.num',X30)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X30),s('type.num.num',X28))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X31),s('type.num.num',X28)))) ) | $spl476),
  inference(cnf_transformation,[],[f12665_D])).
fof(f12665_D,plain,(
  ( ! [X30,X28,X31] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X28),s('type.num.num',sK12))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',X28))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X28),s('type.num.num',X31),s('type.num.num',X30)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X30),s('type.num.num',X28))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X31),s('type.num.num',X28)))) ) <=> ~$spl476),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl476])])).
fof(f14829,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X28),s('type.num.num',sK12))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',X28))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X29))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X29))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X29),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X30),s('type.num.num',X28))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X31),s('type.num.num',X28))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X28),s('type.num.num',X31),s('type.num.num',X30))))) )),
  inference(superposition,[],[f370,f239])).
fof(f12664,plain,(
  ( ! [X26,X24,X27] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X24),s('type.num.num',sK12))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',X24))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X24),s('type.num.num',X26),s('type.num.num',X27)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X24))))) ) | $spl474),
  inference(cnf_transformation,[],[f12664_D])).
fof(f12664_D,plain,(
  ( ! [X26,X24,X27] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X24),s('type.num.num',sK12))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',X24))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X24),s('type.num.num',X26),s('type.num.num',X27)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X24))))) ) <=> ~$spl474),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl474])])).
fof(f14828,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X24),s('type.num.num',sK12))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',X24))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X25))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X25))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X25),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X24)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X24),s('type.num.num',X26),s('type.num.num',X27))))) )),
  inference(superposition,[],[f370,f237])).
fof(f14767,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X28),s('type.num.num',sK12))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',X28))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X29))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X29))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X29),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X30),s('type.num.num',X28))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X31),s('type.num.num',X28))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X28),s('type.num.num',X31),s('type.num.num',X30))))) )),
  inference(superposition,[],[f369,f239])).
fof(f14766,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X24),s('type.num.num',sK12))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',X24))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X25))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X25))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X25),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X24)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X24),s('type.num.num',X26),s('type.num.num',X27))))) )),
  inference(superposition,[],[f369,f237])).
fof(f13958,plain,(
  ( ! [X25] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X25))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK14),s('type.num.num',X25))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X25),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14))))) ) | $spl486),
  inference(cnf_transformation,[],[f13958_D])).
fof(f13958_D,plain,(
  ( ! [X25] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X25))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK14),s('type.num.num',X25))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X25),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14))))) ) <=> ~$spl486),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl486])])).
fof(f12416,plain,(
  ( ! [X30,X28,X31] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X28),s('type.num.num',sK12),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X28),s('type.num.num',sK13),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X28),s('type.num.num',X31),s('type.num.num',X30)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X30),s('type.num.num',X28))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X31),s('type.num.num',X28)))) ) | $spl472),
  inference(cnf_transformation,[],[f12416_D])).
fof(f12416_D,plain,(
  ( ! [X30,X28,X31] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X28),s('type.num.num',sK12),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X28),s('type.num.num',sK13),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X28),s('type.num.num',X31),s('type.num.num',X30)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X30),s('type.num.num',X28))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X31),s('type.num.num',X28)))) ) <=> ~$spl472),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl472])])).
fof(f14697,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X28),s('type.num.num',sK12),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X28),s('type.num.num',sK13),s('type.num.num',sK12))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X29))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK14),s('type.num.num',X29))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X29),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X30),s('type.num.num',X28))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X31),s('type.num.num',X28))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X28),s('type.num.num',X31),s('type.num.num',X30))))) )),
  inference(superposition,[],[f364,f239])).
fof(f12415,plain,(
  ( ! [X26,X24,X27] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X24),s('type.num.num',sK12),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X24),s('type.num.num',sK13),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X24),s('type.num.num',X26),s('type.num.num',X27)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X24))))) ) | $spl470),
  inference(cnf_transformation,[],[f12415_D])).
fof(f12415_D,plain,(
  ( ! [X26,X24,X27] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X24),s('type.num.num',sK12),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X24),s('type.num.num',sK13),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X24),s('type.num.num',X26),s('type.num.num',X27)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X24))))) ) <=> ~$spl470),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl470])])).
fof(f14696,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X24),s('type.num.num',sK12),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X24),s('type.num.num',sK13),s('type.num.num',sK12))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X25))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK14),s('type.num.num',X25))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X25),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X24)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X24),s('type.num.num',X26),s('type.num.num',X27))))) )),
  inference(superposition,[],[f364,f237])).
fof(f13882,plain,(
  ( ! [X25] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X25))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK14),s('type.num.num',X25))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X25),s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | $spl482),
  inference(cnf_transformation,[],[f13882_D])).
fof(f13882_D,plain,(
  ( ! [X25] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X25))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK14),s('type.num.num',X25))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X25),s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) <=> ~$spl482),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl482])])).
fof(f14654,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X28),s('type.num.num',sK12),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X28),s('type.num.num',sK13),s('type.num.num',sK12))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X29))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK14),s('type.num.num',X29))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X29),s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X30),s('type.num.num',X28))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X31),s('type.num.num',X28))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X28),s('type.num.num',X31),s('type.num.num',X30))))) )),
  inference(superposition,[],[f363,f239])).
fof(f14653,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X24),s('type.num.num',sK12),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X24),s('type.num.num',sK13),s('type.num.num',sK12))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X25))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK14),s('type.num.num',X25))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X25),s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X24)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X24),s('type.num.num',X26),s('type.num.num',X27))))) )),
  inference(superposition,[],[f363,f237])).
fof(f14645,plain,(
  ( ! [X19,X17] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',X19))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X17),s('type.num.num',X19))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',X17)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X19),s('type.num.num',X17),s('type.num.num',sK12))))) ) | $spl520),
  inference(cnf_transformation,[],[f14645_D])).
fof(f14645_D,plain,(
  ( ! [X19,X17] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',X19))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X17),s('type.num.num',X19))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',X17)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X19),s('type.num.num',X17),s('type.num.num',sK12))))) ) <=> ~$spl520),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl520])])).
fof(f10348,plain,(
  ( ! [X18] : (s(bool,'F_0') != s(bool,X18)) ) | $spl216),
  inference(cnf_transformation,[],[f10348_D])).
fof(f10348_D,plain,(
  ( ! [X18] : (s(bool,'F_0') != s(bool,X18)) ) <=> ~$spl216),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl216])])).
fof(f14646,plain,(
  ( ! [X21,X22,X20] : (s(bool,'F_0') != s(bool,X21) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',X22))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X20),s('type.num.num',X22))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X22),s('type.num.num',X20),s('type.num.num',sK12)))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',X20))))) ) | $spl30),
  inference(subsumption_resolution,[],[f14626,f228])).
fof(f228,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f146])).
fof(f14626,plain,(
  ( ! [X21,X22,X20] : (s(bool,'F_0') != s(bool,X21) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',X22))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X20),s('type.num.num',X22))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X22),s('type.num.num',X20),s('type.num.num',sK12)))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',X20)))) | p(s(bool,X21))) ) | $spl30),
  inference(superposition,[],[f5911,f223])).
fof(f5911,plain,(
  ( ! [X10,X11] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',X10))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',X11))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X10),s('type.num.num',X11))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X11),s('type.num.num',X10),s('type.num.num',sK12))))) ) | $spl30),
  inference(superposition,[],[f5544,f240])).
fof(f14644,plain,(
  ( ! [X19,X17,X18] : (s(bool,'F_0') != s(bool,X18) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',X19))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X17),s('type.num.num',X19))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X19),s('type.num.num',X17),s('type.num.num',sK12)))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',X17))))) ) | $spl30),
  inference(subsumption_resolution,[],[f14625,f228])).
fof(f14625,plain,(
  ( ! [X19,X17,X18] : (s(bool,'F_0') != s(bool,X18) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',X19))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X17),s('type.num.num',X19))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X19),s('type.num.num',X17),s('type.num.num',sK12)))) | p(s(bool,X18)) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',X17))))) ) | $spl30),
  inference(superposition,[],[f5911,f223])).
fof(f12174,plain,(
  ( ! [X30,X28,X31] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',X28))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X28),s('type.num.num',X31),s('type.num.num',X30)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X30),s('type.num.num',X28))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X31),s('type.num.num',X28)))) ) | $spl458),
  inference(cnf_transformation,[],[f12174_D])).
fof(f12174_D,plain,(
  ( ! [X30,X28,X31] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',X28))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X28),s('type.num.num',X31),s('type.num.num',X30)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X30),s('type.num.num',X28))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X31),s('type.num.num',X28)))) ) <=> ~$spl458),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl458])])).
fof(f14615,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',X28))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X29))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',X29))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X29),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X30),s('type.num.num',X28))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X31),s('type.num.num',X28))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X28),s('type.num.num',X31),s('type.num.num',X30))))) ) | $spl30),
  inference(superposition,[],[f5911,f239])).
fof(f12173,plain,(
  ( ! [X26,X24,X27] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',X24))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X24),s('type.num.num',X26),s('type.num.num',X27)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X24))))) ) | $spl456),
  inference(cnf_transformation,[],[f12173_D])).
fof(f12173_D,plain,(
  ( ! [X26,X24,X27] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',X24))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X24),s('type.num.num',X26),s('type.num.num',X27)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X24))))) ) <=> ~$spl456),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl456])])).
fof(f14614,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',X24))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X25))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',X25))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X25),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X24)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X24),s('type.num.num',X26),s('type.num.num',X27))))) ) | $spl30),
  inference(superposition,[],[f5911,f237])).
fof(f8982,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | $spl201),
  inference(cnf_transformation,[],[f8982_D])).
fof(f8982_D,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) <=> ~$spl201),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl201])])).
fof(f14611,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',X13))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X12),s('type.num.num',X13))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',X12),s('type.num.num',sK12)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',X12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',X12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl30),
  inference(superposition,[],[f5911,f239])).
fof(f14608,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X0),s('type.num.num',X1))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X1),s('type.num.num',X0),s('type.num.num',sK12)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl30),
  inference(superposition,[],[f5911,f237])).
fof(f14538,plain,(
  ( ! [X19,X17] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',X19))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X17),s('type.num.num',X19))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',X17)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X19),s('type.num.num',sK12),s('type.num.num',X17))))) ) | $spl512),
  inference(cnf_transformation,[],[f14538_D])).
fof(f14538_D,plain,(
  ( ! [X19,X17] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',X19))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X17),s('type.num.num',X19))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',X17)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X19),s('type.num.num',sK12),s('type.num.num',X17))))) ) <=> ~$spl512),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl512])])).
fof(f14539,plain,(
  ( ! [X21,X22,X20] : (s(bool,'F_0') != s(bool,X21) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',X22))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X20),s('type.num.num',X22))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X22),s('type.num.num',sK12),s('type.num.num',X20)))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',X20))))) ) | $spl30),
  inference(subsumption_resolution,[],[f14519,f228])).
fof(f14519,plain,(
  ( ! [X21,X22,X20] : (s(bool,'F_0') != s(bool,X21) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',X22))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X20),s('type.num.num',X22))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X22),s('type.num.num',sK12),s('type.num.num',X20)))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',X20)))) | p(s(bool,X21))) ) | $spl30),
  inference(superposition,[],[f5910,f223])).
fof(f5910,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',X8))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',X9))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X8),s('type.num.num',X9))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X9),s('type.num.num',sK12),s('type.num.num',X8))))) ) | $spl30),
  inference(superposition,[],[f5544,f240])).
fof(f14537,plain,(
  ( ! [X19,X17,X18] : (s(bool,'F_0') != s(bool,X18) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',X19))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X17),s('type.num.num',X19))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X19),s('type.num.num',sK12),s('type.num.num',X17)))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',X17))))) ) | $spl30),
  inference(subsumption_resolution,[],[f14518,f228])).
fof(f14518,plain,(
  ( ! [X19,X17,X18] : (s(bool,'F_0') != s(bool,X18) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',X19))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X17),s('type.num.num',X19))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X19),s('type.num.num',sK12),s('type.num.num',X17)))) | p(s(bool,X18)) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',X17))))) ) | $spl30),
  inference(superposition,[],[f5910,f223])).
fof(f14508,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',X28))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X29))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',X29))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X29),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X30),s('type.num.num',X28))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X31),s('type.num.num',X28))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X28),s('type.num.num',X31),s('type.num.num',X30))))) ) | $spl30),
  inference(superposition,[],[f5910,f239])).
fof(f14507,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',X24))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X25))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',X25))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X25),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X24)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X24),s('type.num.num',X26),s('type.num.num',X27))))) ) | $spl30),
  inference(superposition,[],[f5910,f237])).
fof(f14504,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',X13))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X12),s('type.num.num',X13))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',sK12),s('type.num.num',X12)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',X12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',X12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl30),
  inference(superposition,[],[f5910,f239])).
fof(f14501,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X0),s('type.num.num',X1))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X1),s('type.num.num',sK12),s('type.num.num',X0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl30),
  inference(superposition,[],[f5910,f237])).
fof(f14345,plain,(
  ( ! [X19,X17] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X19))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X17),s('type.num.num',X19))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X17),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X19),s('type.num.num',X17),s('type.num.num',sK13))))) ) | $spl504),
  inference(cnf_transformation,[],[f14345_D])).
fof(f14345_D,plain,(
  ( ! [X19,X17] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X19))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X17),s('type.num.num',X19))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X17),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X19),s('type.num.num',X17),s('type.num.num',sK13))))) ) <=> ~$spl504),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl504])])).
fof(f14346,plain,(
  ( ! [X21,X22,X20] : (s(bool,'F_0') != s(bool,X21) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X22))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X20),s('type.num.num',X22))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X22),s('type.num.num',X20),s('type.num.num',sK13)))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X20),s('type.num.num',sK12))))) ) | $spl30),
  inference(subsumption_resolution,[],[f14326,f228])).
fof(f14326,plain,(
  ( ! [X21,X22,X20] : (s(bool,'F_0') != s(bool,X21) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X22))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X20),s('type.num.num',X22))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X22),s('type.num.num',X20),s('type.num.num',sK13)))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X20),s('type.num.num',sK12)))) | p(s(bool,X21))) ) | $spl30),
  inference(superposition,[],[f5905,f223])).
fof(f14344,plain,(
  ( ! [X19,X17,X18] : (s(bool,'F_0') != s(bool,X18) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X19))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X17),s('type.num.num',X19))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X19),s('type.num.num',X17),s('type.num.num',sK13)))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X17),s('type.num.num',sK12))))) ) | $spl30),
  inference(subsumption_resolution,[],[f14325,f228])).
fof(f14325,plain,(
  ( ! [X19,X17,X18] : (s(bool,'F_0') != s(bool,X18) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X19))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X17),s('type.num.num',X19))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X19),s('type.num.num',X17),s('type.num.num',sK13)))) | p(s(bool,X18)) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X17),s('type.num.num',sK12))))) ) | $spl30),
  inference(superposition,[],[f5905,f223])).
fof(f11892,plain,(
  ( ! [X30,X28,X31] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X28),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X28),s('type.num.num',X31),s('type.num.num',X30)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X30),s('type.num.num',X28))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X31),s('type.num.num',X28)))) ) | $spl438),
  inference(cnf_transformation,[],[f11892_D])).
fof(f11892_D,plain,(
  ( ! [X30,X28,X31] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X28),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X28),s('type.num.num',X31),s('type.num.num',X30)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X30),s('type.num.num',X28))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X31),s('type.num.num',X28)))) ) <=> ~$spl438),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl438])])).
fof(f14309,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X28),s('type.num.num',sK12))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X29))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X29))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X29),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X30),s('type.num.num',X28))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X31),s('type.num.num',X28))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X28),s('type.num.num',X31),s('type.num.num',X30))))) ) | $spl30),
  inference(superposition,[],[f5905,f239])).
fof(f11891,plain,(
  ( ! [X26,X24,X27] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X24),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X24),s('type.num.num',X26),s('type.num.num',X27)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X24))))) ) | $spl436),
  inference(cnf_transformation,[],[f11891_D])).
fof(f11891_D,plain,(
  ( ! [X26,X24,X27] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X24),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X24),s('type.num.num',X26),s('type.num.num',X27)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X24))))) ) <=> ~$spl436),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl436])])).
fof(f14308,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X24),s('type.num.num',sK12))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X25))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X25))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X25),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X24)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X24),s('type.num.num',X26),s('type.num.num',X27))))) ) | $spl30),
  inference(superposition,[],[f5905,f237])).
fof(f14305,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X13))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X12),s('type.num.num',X13))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',X12),s('type.num.num',sK13)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',X12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',X12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl30),
  inference(superposition,[],[f5905,f239])).
fof(f14302,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X0),s('type.num.num',X1))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X1),s('type.num.num',X0),s('type.num.num',sK13)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl30),
  inference(superposition,[],[f5905,f237])).
fof(f14153,plain,(
  ( ! [X19,X17] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X19))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X17),s('type.num.num',X19))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X17),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X19),s('type.num.num',sK13),s('type.num.num',X17))))) ) | $spl496),
  inference(cnf_transformation,[],[f14153_D])).
fof(f14153_D,plain,(
  ( ! [X19,X17] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X19))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X17),s('type.num.num',X19))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X17),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X19),s('type.num.num',sK13),s('type.num.num',X17))))) ) <=> ~$spl496),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl496])])).
fof(f14154,plain,(
  ( ! [X21,X22,X20] : (s(bool,'F_0') != s(bool,X21) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X22))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X20),s('type.num.num',X22))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X22),s('type.num.num',sK13),s('type.num.num',X20)))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X20),s('type.num.num',sK12))))) ) | $spl30),
  inference(subsumption_resolution,[],[f14134,f228])).
fof(f14134,plain,(
  ( ! [X21,X22,X20] : (s(bool,'F_0') != s(bool,X21) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X22))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X20),s('type.num.num',X22))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X22),s('type.num.num',sK13),s('type.num.num',X20)))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X20),s('type.num.num',sK12)))) | p(s(bool,X21))) ) | $spl30),
  inference(superposition,[],[f5904,f223])).
fof(f14152,plain,(
  ( ! [X19,X17,X18] : (s(bool,'F_0') != s(bool,X18) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X19))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X17),s('type.num.num',X19))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X19),s('type.num.num',sK13),s('type.num.num',X17)))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X17),s('type.num.num',sK12))))) ) | $spl30),
  inference(subsumption_resolution,[],[f14133,f228])).
fof(f14133,plain,(
  ( ! [X19,X17,X18] : (s(bool,'F_0') != s(bool,X18) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X19))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X17),s('type.num.num',X19))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X19),s('type.num.num',sK13),s('type.num.num',X17)))) | p(s(bool,X18)) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X17),s('type.num.num',sK12))))) ) | $spl30),
  inference(superposition,[],[f5904,f223])).
fof(f14117,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X28),s('type.num.num',sK12))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X29))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X29))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X29),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X30),s('type.num.num',X28))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X31),s('type.num.num',X28))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X28),s('type.num.num',X31),s('type.num.num',X30))))) ) | $spl30),
  inference(superposition,[],[f5904,f239])).
fof(f14116,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X24),s('type.num.num',sK12))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X25))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X25))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X25),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X24)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X24),s('type.num.num',X26),s('type.num.num',X27))))) ) | $spl30),
  inference(superposition,[],[f5904,f237])).
fof(f14113,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X13))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X12),s('type.num.num',X13))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',sK13),s('type.num.num',X12)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',X12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',X12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl30),
  inference(superposition,[],[f5904,f239])).
fof(f14110,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X0),s('type.num.num',X1))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X1),s('type.num.num',sK13),s('type.num.num',X0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl30),
  inference(superposition,[],[f5904,f237])).
fof(f13960,plain,(
  ( ! [X20,X18] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK14),s('type.num.num',X20))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X18),s('type.num.num',X20))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X18),s('type.num.num',sK13),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X20),s('type.num.num',X18),s('type.num.num',sK14))))) ) | $spl488),
  inference(cnf_transformation,[],[f13960_D])).
fof(f13960_D,plain,(
  ( ! [X20,X18] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK14),s('type.num.num',X20))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X18),s('type.num.num',X20))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X18),s('type.num.num',sK13),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X20),s('type.num.num',X18),s('type.num.num',sK14))))) ) <=> ~$spl488),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl488])])).
fof(f13961,plain,(
  ( ! [X23,X21,X22] : (s(bool,'F_0') != s(bool,X22) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK14),s('type.num.num',X23))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X21),s('type.num.num',X23))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X23),s('type.num.num',X21),s('type.num.num',sK14)))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X21),s('type.num.num',sK13),s('type.num.num',sK12))))) ) | $spl30),
  inference(subsumption_resolution,[],[f13944,f228])).
fof(f13944,plain,(
  ( ! [X23,X21,X22] : (s(bool,'F_0') != s(bool,X22) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK14),s('type.num.num',X23))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X21),s('type.num.num',X23))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X23),s('type.num.num',X21),s('type.num.num',sK14)))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X21),s('type.num.num',sK13),s('type.num.num',sK12)))) | p(s(bool,X22))) ) | $spl30),
  inference(superposition,[],[f5899,f223])).
fof(f13959,plain,(
  ( ! [X19,X20,X18] : (s(bool,'F_0') != s(bool,X19) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK14),s('type.num.num',X20))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X18),s('type.num.num',X20))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X20),s('type.num.num',X18),s('type.num.num',sK14)))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X18),s('type.num.num',sK13),s('type.num.num',sK12))))) ) | $spl30),
  inference(subsumption_resolution,[],[f13943,f228])).
fof(f13943,plain,(
  ( ! [X19,X20,X18] : (s(bool,'F_0') != s(bool,X19) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK14),s('type.num.num',X20))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X18),s('type.num.num',X20))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X20),s('type.num.num',X18),s('type.num.num',sK14)))) | p(s(bool,X19)) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X18),s('type.num.num',sK13),s('type.num.num',sK12))))) ) | $spl30),
  inference(superposition,[],[f5899,f223])).
fof(f10345,plain,(
  ( ! [X30,X28,X31] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X28),s('type.num.num',sK13),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X28),s('type.num.num',X31),s('type.num.num',X30)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X30),s('type.num.num',X28))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X31),s('type.num.num',X28)))) ) | $spl212),
  inference(cnf_transformation,[],[f10345_D])).
fof(f10345_D,plain,(
  ( ! [X30,X28,X31] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X28),s('type.num.num',sK13),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X28),s('type.num.num',X31),s('type.num.num',X30)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X30),s('type.num.num',X28))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X31),s('type.num.num',X28)))) ) <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f13920,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X28),s('type.num.num',sK13),s('type.num.num',sK12))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X29))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK14),s('type.num.num',X29))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X29),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X30),s('type.num.num',X28))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X31),s('type.num.num',X28))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X28),s('type.num.num',X31),s('type.num.num',X30))))) ) | $spl30),
  inference(superposition,[],[f5899,f239])).
fof(f10344,plain,(
  ( ! [X26,X24,X27] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X24),s('type.num.num',sK13),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X24),s('type.num.num',X26),s('type.num.num',X27)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X24))))) ) | $spl210),
  inference(cnf_transformation,[],[f10344_D])).
fof(f10344_D,plain,(
  ( ! [X26,X24,X27] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X24),s('type.num.num',sK13),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X24),s('type.num.num',X26),s('type.num.num',X27)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X24))))) ) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f13919,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X24),s('type.num.num',sK13),s('type.num.num',sK12))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X25))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK14),s('type.num.num',X25))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X25),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X24)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X24),s('type.num.num',X26),s('type.num.num',X27))))) ) | $spl30),
  inference(superposition,[],[f5899,f237])).
fof(f13884,plain,(
  ( ! [X20,X18] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK14),s('type.num.num',X20))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X18),s('type.num.num',X20))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X18),s('type.num.num',sK13),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X20),s('type.num.num',sK14),s('type.num.num',X18))))) ) | $spl484),
  inference(cnf_transformation,[],[f13884_D])).
fof(f13884_D,plain,(
  ( ! [X20,X18] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK14),s('type.num.num',X20))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X18),s('type.num.num',X20))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X18),s('type.num.num',sK13),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X20),s('type.num.num',sK14),s('type.num.num',X18))))) ) <=> ~$spl484),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl484])])).
fof(f13885,plain,(
  ( ! [X23,X21,X22] : (s(bool,'F_0') != s(bool,X22) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK14),s('type.num.num',X23))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X21),s('type.num.num',X23))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X23),s('type.num.num',sK14),s('type.num.num',X21)))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X21),s('type.num.num',sK13),s('type.num.num',sK12))))) ) | $spl30),
  inference(subsumption_resolution,[],[f13868,f228])).
fof(f13868,plain,(
  ( ! [X23,X21,X22] : (s(bool,'F_0') != s(bool,X22) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK14),s('type.num.num',X23))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X21),s('type.num.num',X23))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X23),s('type.num.num',sK14),s('type.num.num',X21)))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X21),s('type.num.num',sK13),s('type.num.num',sK12)))) | p(s(bool,X22))) ) | $spl30),
  inference(superposition,[],[f5898,f223])).
fof(f13883,plain,(
  ( ! [X19,X20,X18] : (s(bool,'F_0') != s(bool,X19) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK14),s('type.num.num',X20))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X18),s('type.num.num',X20))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X20),s('type.num.num',sK14),s('type.num.num',X18)))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X18),s('type.num.num',sK13),s('type.num.num',sK12))))) ) | $spl30),
  inference(subsumption_resolution,[],[f13867,f228])).
fof(f13867,plain,(
  ( ! [X19,X20,X18] : (s(bool,'F_0') != s(bool,X19) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK14),s('type.num.num',X20))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X18),s('type.num.num',X20))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X20),s('type.num.num',sK14),s('type.num.num',X18)))) | p(s(bool,X19)) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X18),s('type.num.num',sK13),s('type.num.num',sK12))))) ) | $spl30),
  inference(superposition,[],[f5898,f223])).
fof(f13844,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X28),s('type.num.num',sK13),s('type.num.num',sK12))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X29))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK14),s('type.num.num',X29))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X29),s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X30),s('type.num.num',X28))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X31),s('type.num.num',X28))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X28),s('type.num.num',X31),s('type.num.num',X30))))) ) | $spl30),
  inference(superposition,[],[f5898,f239])).
fof(f13843,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X24),s('type.num.num',sK13),s('type.num.num',sK12))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X25))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK14),s('type.num.num',X25))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X25),s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X24)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X24),s('type.num.num',X26),s('type.num.num',X27))))) ) | $spl30),
  inference(superposition,[],[f5898,f237])).
fof(f12369,plain,(
  ( ! [X25] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X25)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X25),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))))) ) | $spl466),
  inference(cnf_transformation,[],[f12369_D])).
fof(f12369_D,plain,(
  ( ! [X25] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X25)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X25),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))))) ) <=> ~$spl466),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl466])])).
fof(f13126,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X28),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',X28))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X29)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X29),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X30),s('type.num.num',X28))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X31),s('type.num.num',X28))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X28),s('type.num.num',X31),s('type.num.num',X30))))) )),
  inference(superposition,[],[f373,f239])).
fof(f13125,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X24),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',X24))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X25)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X25),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X24)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X24),s('type.num.num',X26),s('type.num.num',X27))))) )),
  inference(superposition,[],[f373,f237])).
fof(f12172,plain,(
  ( ! [X25] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X25)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X25),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0'))))) ) | $spl454),
  inference(cnf_transformation,[],[f12172_D])).
fof(f12172_D,plain,(
  ( ! [X25] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X25)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X25),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0'))))) ) <=> ~$spl454),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl454])])).
fof(f13039,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X28),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',X28))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X29)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X29),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X30),s('type.num.num',X28))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X31),s('type.num.num',X28))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X28),s('type.num.num',X31),s('type.num.num',X30))))) )),
  inference(superposition,[],[f372,f239])).
fof(f13038,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X24),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',X24))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X25)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X25),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X24)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X24),s('type.num.num',X26),s('type.num.num',X27))))) )),
  inference(superposition,[],[f372,f237])).
fof(f12828,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X28),s('type.num.num',sK12))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',X28))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X29)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X29),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X30),s('type.num.num',X28))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X31),s('type.num.num',X28))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X28),s('type.num.num',X31),s('type.num.num',X30))))) )),
  inference(superposition,[],[f367,f239])).
fof(f12827,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X24),s('type.num.num',sK12))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',X24))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X25)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X25),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X24)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X24),s('type.num.num',X26),s('type.num.num',X27))))) )),
  inference(superposition,[],[f367,f237])).
fof(f12638,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X28),s('type.num.num',sK12))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',X28))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X29)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X29),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X30),s('type.num.num',X28))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X31),s('type.num.num',X28))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X28),s('type.num.num',X31),s('type.num.num',X30))))) )),
  inference(superposition,[],[f366,f239])).
fof(f12637,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X24),s('type.num.num',sK12))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',X24))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X25)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X25),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X24)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X24),s('type.num.num',X26),s('type.num.num',X27))))) )),
  inference(superposition,[],[f366,f237])).
fof(f11788,plain,(
  ( ! [X25] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X25)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X25),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14))))) ) | $spl426),
  inference(cnf_transformation,[],[f11788_D])).
fof(f11788_D,plain,(
  ( ! [X25] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X25)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X25),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14))))) ) <=> ~$spl426),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl426])])).
fof(f12465,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X28),s('type.num.num',sK12),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X28),s('type.num.num',sK13),s('type.num.num',sK12))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X29)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X29),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X30),s('type.num.num',X28))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X31),s('type.num.num',X28))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X28),s('type.num.num',X31),s('type.num.num',X30))))) )),
  inference(superposition,[],[f361,f239])).
fof(f12464,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X24),s('type.num.num',sK12),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X24),s('type.num.num',sK13),s('type.num.num',sK12))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X25)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X25),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X24)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X24),s('type.num.num',X26),s('type.num.num',X27))))) )),
  inference(superposition,[],[f361,f237])).
fof(f10343,plain,(
  ( ! [X25] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X25)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X25),s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | $spl208),
  inference(cnf_transformation,[],[f10343_D])).
fof(f10343_D,plain,(
  ( ! [X25] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X25)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X25),s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f12380,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X28),s('type.num.num',sK12),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X28),s('type.num.num',sK13),s('type.num.num',sK12))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X29)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X29),s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X30),s('type.num.num',X28))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X31),s('type.num.num',X28))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X28),s('type.num.num',X31),s('type.num.num',X30))))) )),
  inference(superposition,[],[f360,f239])).
fof(f12379,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X24),s('type.num.num',sK12),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X24),s('type.num.num',sK13),s('type.num.num',sK12))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X25)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X25),s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X24)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X24),s('type.num.num',X26),s('type.num.num',X27))))) )),
  inference(superposition,[],[f360,f237])).
fof(f12371,plain,(
  ( ! [X19,X17] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X19)))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',X17)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X19),s('type.num.num',X17),s('type.num.num',sK12))))) ) | $spl468),
  inference(cnf_transformation,[],[f12371_D])).
fof(f12371_D,plain,(
  ( ! [X19,X17] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X19)))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',X17)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X19),s('type.num.num',X17),s('type.num.num',sK12))))) ) <=> ~$spl468),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl468])])).
fof(f12372,plain,(
  ( ! [X21,X22,X20] : (s(bool,'F_0') != s(bool,X21) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X22)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X22),s('type.num.num',X20),s('type.num.num',sK12)))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',X20))))) ) | $spl30),
  inference(subsumption_resolution,[],[f12354,f228])).
fof(f12354,plain,(
  ( ! [X21,X22,X20] : (s(bool,'F_0') != s(bool,X21) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X22)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X22),s('type.num.num',X20),s('type.num.num',sK12)))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',X20)))) | p(s(bool,X21))) ) | $spl30),
  inference(superposition,[],[f5908,f223])).
fof(f5908,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',X4))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X5)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X5),s('type.num.num',X4),s('type.num.num',sK12))))) ) | $spl30),
  inference(superposition,[],[f5544,f238])).
fof(f12370,plain,(
  ( ! [X19,X17,X18] : (s(bool,'F_0') != s(bool,X18) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X19)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X19),s('type.num.num',X17),s('type.num.num',sK12)))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',X17))))) ) | $spl30),
  inference(subsumption_resolution,[],[f12353,f228])).
fof(f12353,plain,(
  ( ! [X19,X17,X18] : (s(bool,'F_0') != s(bool,X18) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X19)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X19),s('type.num.num',X17),s('type.num.num',sK12)))) | p(s(bool,X18)) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',X17))))) ) | $spl30),
  inference(superposition,[],[f5908,f223])).
fof(f12343,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',X28))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X29)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X29),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X30),s('type.num.num',X28))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X31),s('type.num.num',X28))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X28),s('type.num.num',X31),s('type.num.num',X30))))) ) | $spl30),
  inference(superposition,[],[f5908,f239])).
fof(f12342,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',X24))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X25)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X25),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X24)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X24),s('type.num.num',X26),s('type.num.num',X27))))) ) | $spl30),
  inference(superposition,[],[f5908,f237])).
fof(f12339,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',X12),s('type.num.num',sK12)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',X12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',X12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl30),
  inference(superposition,[],[f5908,f239])).
fof(f12336,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X1),s('type.num.num',X0),s('type.num.num',sK12)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl30),
  inference(superposition,[],[f5908,f237])).
fof(f12176,plain,(
  ( ! [X19,X17] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X19)))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',X17)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X19),s('type.num.num',sK12),s('type.num.num',X17))))) ) | $spl460),
  inference(cnf_transformation,[],[f12176_D])).
fof(f12176_D,plain,(
  ( ! [X19,X17] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X19)))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',X17)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X19),s('type.num.num',sK12),s('type.num.num',X17))))) ) <=> ~$spl460),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl460])])).
fof(f12177,plain,(
  ( ! [X21,X22,X20] : (s(bool,'F_0') != s(bool,X21) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X22)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X22),s('type.num.num',sK12),s('type.num.num',X20)))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',X20))))) ) | $spl30),
  inference(subsumption_resolution,[],[f12157,f228])).
fof(f12157,plain,(
  ( ! [X21,X22,X20] : (s(bool,'F_0') != s(bool,X21) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X22)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X22),s('type.num.num',sK12),s('type.num.num',X20)))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',X20)))) | p(s(bool,X21))) ) | $spl30),
  inference(superposition,[],[f5907,f223])).
fof(f5907,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',X2))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X3)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X3),s('type.num.num',sK12),s('type.num.num',X2))))) ) | $spl30),
  inference(superposition,[],[f5544,f238])).
fof(f12175,plain,(
  ( ! [X19,X17,X18] : (s(bool,'F_0') != s(bool,X18) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X19)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X19),s('type.num.num',sK12),s('type.num.num',X17)))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',X17))))) ) | $spl30),
  inference(subsumption_resolution,[],[f12156,f228])).
fof(f12156,plain,(
  ( ! [X19,X17,X18] : (s(bool,'F_0') != s(bool,X18) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X19)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X19),s('type.num.num',sK12),s('type.num.num',X17)))) | p(s(bool,X18)) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',X17))))) ) | $spl30),
  inference(superposition,[],[f5907,f223])).
fof(f12146,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',X28))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X29)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X29),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X30),s('type.num.num',X28))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X31),s('type.num.num',X28))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X28),s('type.num.num',X31),s('type.num.num',X30))))) ) | $spl30),
  inference(superposition,[],[f5907,f239])).
fof(f12145,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',X24))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X25)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X25),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X24)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X24),s('type.num.num',X26),s('type.num.num',X27))))) ) | $spl30),
  inference(superposition,[],[f5907,f237])).
fof(f12142,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',sK12),s('type.num.num',X12)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',X12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',X12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl30),
  inference(superposition,[],[f5907,f239])).
fof(f12139,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X1),s('type.num.num',sK12),s('type.num.num',X0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl30),
  inference(superposition,[],[f5907,f237])).
fof(f11944,plain,(
  ( ! [X19,X17] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X19)))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X17),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X19),s('type.num.num',X17),s('type.num.num',sK13))))) ) | $spl448),
  inference(cnf_transformation,[],[f11944_D])).
fof(f11944_D,plain,(
  ( ! [X19,X17] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X19)))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X17),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X19),s('type.num.num',X17),s('type.num.num',sK13))))) ) <=> ~$spl448),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl448])])).
fof(f11945,plain,(
  ( ! [X21,X22,X20] : (s(bool,'F_0') != s(bool,X21) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X22)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X22),s('type.num.num',X20),s('type.num.num',sK13)))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X20),s('type.num.num',sK12))))) ) | $spl30),
  inference(subsumption_resolution,[],[f11927,f228])).
fof(f11927,plain,(
  ( ! [X21,X22,X20] : (s(bool,'F_0') != s(bool,X21) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X22)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X22),s('type.num.num',X20),s('type.num.num',sK13)))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X20),s('type.num.num',sK12)))) | p(s(bool,X21))) ) | $spl30),
  inference(superposition,[],[f5902,f223])).
fof(f11943,plain,(
  ( ! [X19,X17,X18] : (s(bool,'F_0') != s(bool,X18) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X19)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X19),s('type.num.num',X17),s('type.num.num',sK13)))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X17),s('type.num.num',sK12))))) ) | $spl30),
  inference(subsumption_resolution,[],[f11926,f228])).
fof(f11926,plain,(
  ( ! [X19,X17,X18] : (s(bool,'F_0') != s(bool,X18) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X19)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X19),s('type.num.num',X17),s('type.num.num',sK13)))) | p(s(bool,X18)) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X17),s('type.num.num',sK12))))) ) | $spl30),
  inference(superposition,[],[f5902,f223])).
fof(f11910,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X28),s('type.num.num',sK12))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X29)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X29),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X30),s('type.num.num',X28))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X31),s('type.num.num',X28))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X28),s('type.num.num',X31),s('type.num.num',X30))))) ) | $spl30),
  inference(superposition,[],[f5902,f239])).
fof(f11909,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X24),s('type.num.num',sK12))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X25)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X25),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X24)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X24),s('type.num.num',X26),s('type.num.num',X27))))) ) | $spl30),
  inference(superposition,[],[f5902,f237])).
fof(f11906,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',X12),s('type.num.num',sK13)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',X12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',X12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl30),
  inference(superposition,[],[f5902,f239])).
fof(f11903,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X1),s('type.num.num',X0),s('type.num.num',sK13)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl30),
  inference(superposition,[],[f5902,f237])).
fof(f11894,plain,(
  ( ! [X19,X17] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X19)))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X17),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X19),s('type.num.num',sK13),s('type.num.num',X17))))) ) | $spl440),
  inference(cnf_transformation,[],[f11894_D])).
fof(f11894_D,plain,(
  ( ! [X19,X17] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X19)))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X17),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X19),s('type.num.num',sK13),s('type.num.num',X17))))) ) <=> ~$spl440),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl440])])).
fof(f11895,plain,(
  ( ! [X21,X22,X20] : (s(bool,'F_0') != s(bool,X21) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X22)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X22),s('type.num.num',sK13),s('type.num.num',X20)))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X20),s('type.num.num',sK12))))) ) | $spl30),
  inference(subsumption_resolution,[],[f11875,f228])).
fof(f11875,plain,(
  ( ! [X21,X22,X20] : (s(bool,'F_0') != s(bool,X21) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X22)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X22),s('type.num.num',sK13),s('type.num.num',X20)))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X20),s('type.num.num',sK12)))) | p(s(bool,X21))) ) | $spl30),
  inference(superposition,[],[f5901,f223])).
fof(f11893,plain,(
  ( ! [X19,X17,X18] : (s(bool,'F_0') != s(bool,X18) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X19)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X19),s('type.num.num',sK13),s('type.num.num',X17)))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X17),s('type.num.num',sK12))))) ) | $spl30),
  inference(subsumption_resolution,[],[f11874,f228])).
fof(f11874,plain,(
  ( ! [X19,X17,X18] : (s(bool,'F_0') != s(bool,X18) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X19)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X19),s('type.num.num',sK13),s('type.num.num',X17)))) | p(s(bool,X18)) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X17),s('type.num.num',sK12))))) ) | $spl30),
  inference(superposition,[],[f5901,f223])).
fof(f11858,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X28),s('type.num.num',sK12))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X29)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X29),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X30),s('type.num.num',X28))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X31),s('type.num.num',X28))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X28),s('type.num.num',X31),s('type.num.num',X30))))) ) | $spl30),
  inference(superposition,[],[f5901,f239])).
fof(f11857,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X24),s('type.num.num',sK12))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X25)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X25),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X24)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X24),s('type.num.num',X26),s('type.num.num',X27))))) ) | $spl30),
  inference(superposition,[],[f5901,f237])).
fof(f11854,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',sK13),s('type.num.num',X12)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',X12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',X12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl30),
  inference(superposition,[],[f5901,f239])).
fof(f11851,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X1),s('type.num.num',sK13),s('type.num.num',X0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl30),
  inference(superposition,[],[f5901,f237])).
fof(f11836,plain,(
  $false | ($spl30 | $spl216)),
  inference(trivial_inequality_removal,[],[f11826])).
fof(f11826,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl30 | $spl216)),
  inference(superposition,[],[f10348,f413])).
fof(f11835,plain,(
  $false | ($spl11 | $spl216)),
  inference(trivial_inequality_removal,[],[f11830])).
fof(f11830,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl11 | $spl216)),
  inference(superposition,[],[f10348,f3467])).
fof(f3467,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) | $spl11),
  inference(resolution,[],[f398,f229])).
fof(f11834,plain,(
  $false | ($spl19 | $spl216)),
  inference(trivial_inequality_removal,[],[f11832])).
fof(f11832,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl19 | $spl216)),
  inference(superposition,[],[f10348,f4299])).
fof(f4299,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) | $spl19),
  inference(resolution,[],[f404,f229])).
fof(f404,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | $spl19),
  inference(cnf_transformation,[],[f404_D])).
fof(f404_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f11833,plain,(
  $false | $spl216),
  inference(equality_resolution,[],[f10348])).
fof(f11790,plain,(
  ( ! [X20,X18] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X20)))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X18),s('type.num.num',sK13),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X20),s('type.num.num',X18),s('type.num.num',sK14))))) ) | $spl428),
  inference(cnf_transformation,[],[f11790_D])).
fof(f11790_D,plain,(
  ( ! [X20,X18] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X20)))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X18),s('type.num.num',sK13),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X20),s('type.num.num',X18),s('type.num.num',sK14))))) ) <=> ~$spl428),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl428])])).
fof(f11791,plain,(
  ( ! [X23,X21,X22] : (s(bool,'F_0') != s(bool,X22) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X23)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X23),s('type.num.num',X21),s('type.num.num',sK14)))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X21),s('type.num.num',sK13),s('type.num.num',sK12))))) ) | $spl30),
  inference(subsumption_resolution,[],[f11774,f228])).
fof(f11774,plain,(
  ( ! [X23,X21,X22] : (s(bool,'F_0') != s(bool,X22) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X23)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X23),s('type.num.num',X21),s('type.num.num',sK14)))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X21),s('type.num.num',sK13),s('type.num.num',sK12)))) | p(s(bool,X22))) ) | $spl30),
  inference(superposition,[],[f5896,f223])).
fof(f11789,plain,(
  ( ! [X19,X20,X18] : (s(bool,'F_0') != s(bool,X19) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X20)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X20),s('type.num.num',X18),s('type.num.num',sK14)))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X18),s('type.num.num',sK13),s('type.num.num',sK12))))) ) | $spl30),
  inference(subsumption_resolution,[],[f11773,f228])).
fof(f11773,plain,(
  ( ! [X19,X20,X18] : (s(bool,'F_0') != s(bool,X19) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X20)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X20),s('type.num.num',X18),s('type.num.num',sK14)))) | p(s(bool,X19)) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X18),s('type.num.num',sK13),s('type.num.num',sK12))))) ) | $spl30),
  inference(superposition,[],[f5896,f223])).
fof(f11750,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X28),s('type.num.num',sK13),s('type.num.num',sK12))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X29)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X29),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X30),s('type.num.num',X28))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X31),s('type.num.num',X28))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X28),s('type.num.num',X31),s('type.num.num',X30))))) ) | $spl30),
  inference(superposition,[],[f5896,f239])).
fof(f11749,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X24),s('type.num.num',sK13),s('type.num.num',sK12))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X25)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X25),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X24)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X24),s('type.num.num',X26),s('type.num.num',X27))))) ) | $spl30),
  inference(superposition,[],[f5896,f237])).
fof(f6828,plain,(
  ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | $spl153),
  inference(cnf_transformation,[],[f6828_D])).
fof(f6828_D,plain,(
  ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) <=> ~$spl153),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl153])])).
fof(f11682,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,X11) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))))) ) | ($spl5 | $spl132)),
  inference(subsumption_resolution,[],[f11671,f232])).
fof(f11671,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,X11) | ~p(s(bool,X11)) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))))) ) | ($spl5 | $spl132)),
  inference(superposition,[],[f11643,f226])).
fof(f11681,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,X10) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))))) ) | ($spl5 | $spl132)),
  inference(subsumption_resolution,[],[f11670,f232])).
fof(f11670,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,X10) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | ~p(s(bool,X10))) ) | ($spl5 | $spl132)),
  inference(superposition,[],[f11643,f226])).
fof(f1553,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) | $spl78),
  inference(cnf_transformation,[],[f1553_D])).
fof(f1553_D,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f11642,plain,(
  ( ! [X24] : (~p(s(bool,X24)) | s(bool,'F_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) ) | $spl128),
  inference(subsumption_resolution,[],[f11626,f233])).
fof(f11626,plain,(
  ( ! [X24] : (s(bool,'T_0') != s(bool,X24) | ~p(s(bool,X24)) | s(bool,'F_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) ) | $spl128),
  inference(superposition,[],[f3287,f236])).
fof(f11641,plain,(
  ( ! [X23] : (s(bool,'T_0') != s(bool,X23) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))))) ) | $spl128),
  inference(subsumption_resolution,[],[f11625,f232])).
fof(f11625,plain,(
  ( ! [X23] : (s(bool,'T_0') != s(bool,X23) | ~p(s(bool,X23)) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))))) ) | $spl128),
  inference(superposition,[],[f3287,f235])).
fof(f235,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f149])).
fof(f149,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X13] : (s(bool,'T_0') = s(bool,X13) <=> p(s(bool,X13)))),
  file('thm.arithmetic.MODEQ_SYM','thm.bool.EQ_CLAUSES|split|1')).
fof(f11640,plain,(
  ( ! [X22] : (s(bool,'T_0') != s(bool,X22) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))))) ) | $spl128),
  inference(subsumption_resolution,[],[f11624,f232])).
fof(f11624,plain,(
  ( ! [X22] : (s(bool,'T_0') != s(bool,X22) | ~p(s(bool,X22)) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))))) ) | $spl128),
  inference(superposition,[],[f3287,f233])).
fof(f11639,plain,(
  ( ! [X19,X18] : (~p(s(bool,X19)) | ~p(s(bool,X18)) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))))) ) | $spl128),
  inference(subsumption_resolution,[],[f11621,f226])).
fof(f11621,plain,(
  ( ! [X19,X18] : (s(bool,X18) != s(bool,X19) | ~p(s(bool,X19)) | ~p(s(bool,X18)) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))))) ) | $spl128),
  inference(superposition,[],[f3287,f226])).
fof(f11638,plain,(
  ( ! [X17,X16] : (~p(s(bool,X17)) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | ~p(s(bool,X16))) ) | $spl128),
  inference(subsumption_resolution,[],[f11620,f226])).
fof(f11620,plain,(
  ( ! [X17,X16] : (s(bool,X16) != s(bool,X17) | ~p(s(bool,X17)) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | ~p(s(bool,X16))) ) | $spl128),
  inference(superposition,[],[f3287,f226])).
fof(f6829,plain,(
  p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | $spl152),
  inference(cnf_transformation,[],[f6829_D])).
fof(f6829_D,plain,(
  p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f6830,plain,(
  ( ! [X12,X13] : (s(bool,X12) != s(bool,X13) | p(s(bool,X12)) | ~p(s(bool,X13))) ) | $spl154),
  inference(cnf_transformation,[],[f6830_D])).
fof(f6830_D,plain,(
  ( ! [X12,X13] : (s(bool,X12) != s(bool,X13) | p(s(bool,X12)) | ~p(s(bool,X13))) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f11619,plain,(
  ( ! [X14,X15] : (s(bool,X14) != s(bool,X15) | ~p(s(bool,X15)) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | p(s(bool,X14))) ) | $spl128),
  inference(superposition,[],[f3287,f223])).
fof(f11618,plain,(
  ( ! [X12,X13] : (s(bool,X12) != s(bool,X13) | ~p(s(bool,X13)) | p(s(bool,X12)) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))))) ) | $spl128),
  inference(superposition,[],[f3287,f223])).
fof(f6825,plain,(
  ( ! [X24] : (s(bool,'F_0') = s(bool,X24)) ) | $spl150),
  inference(cnf_transformation,[],[f6825_D])).
fof(f6825_D,plain,(
  ( ! [X24] : (s(bool,'F_0') = s(bool,X24)) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f1777,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) | $spl87),
  inference(cnf_transformation,[],[f1777_D])).
fof(f1777_D,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) <=> ~$spl87),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl87])])).
fof(f11633,plain,(
  ( ! [X24] : (s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) | s(bool,'F_0') = s(bool,X24)) ) | $spl128),
  inference(subsumption_resolution,[],[f11588,f229])).
fof(f11588,plain,(
  ( ! [X24] : (s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) | ~p(s(bool,X24)) | s(bool,'F_0') = s(bool,X24)) ) | $spl128),
  inference(superposition,[],[f3287,f236])).
fof(f11629,plain,(
  ( ! [X23] : (s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) | ~p(s(bool,X23))) ) | $spl128),
  inference(duplicate_literal_removal,[],[f11587])).
fof(f11587,plain,(
  ( ! [X23] : (s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) | ~p(s(bool,X23)) | ~p(s(bool,X23))) ) | $spl128),
  inference(superposition,[],[f3287,f235])).
fof(f11630,plain,(
  ( ! [X22] : (s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) | ~p(s(bool,X22))) ) | $spl128),
  inference(duplicate_literal_removal,[],[f11586])).
fof(f11586,plain,(
  ( ! [X22] : (s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) | ~p(s(bool,X22)) | ~p(s(bool,X22))) ) | $spl128),
  inference(superposition,[],[f3287,f233])).
fof(f11091,plain,(
  $false | $spl104),
  inference(subsumption_resolution,[],[f11005,f2958])).
fof(f11005,plain,(
  p(s(bool,'F_0')) | $spl104),
  inference(resolution,[],[f2958,f351])).
fof(f351,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f198])).
fof(f198,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f88])).
fof(f88,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f17])).
fof(f17,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.arithmetic.MODEQ_SYM','thm.bool.NOT_CLAUSES|split|2')).
fof(f11090,plain,(
  $false | $spl104),
  inference(subsumption_resolution,[],[f11004,f2958])).
fof(f11004,plain,(
  p(s(bool,'F_0')) | $spl104),
  inference(resolution,[],[f2958,f349])).
fof(f349,plain,(
  p(s(bool,'F_0')) | p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f197])).
fof(f197,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f87])).
fof(f87,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f16])).
fof(f16,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.arithmetic.MODEQ_SYM','thm.bool.NOT_CLAUSES|split|1')).
fof(f11089,plain,(
  $false | $spl104),
  inference(subsumption_resolution,[],[f11002,f2958])).
fof(f11002,plain,(
  ( ! [X215] : (p(s(bool,X215))) ) | $spl104),
  inference(resolution,[],[f2958,f339])).
fof(f339,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f196])).
fof(f196,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))))),
  inference(flattening,[],[f195])).
fof(f195,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X13] : ((p(s(bool,X13)) => p(s(bool,X13))) <=> p(s(bool,'T_0')))),
  file('thm.arithmetic.MODEQ_SYM','thm.bool.IMP_CLAUSES|split|3')).
fof(f11088,plain,(
  $false | $spl104),
  inference(subsumption_resolution,[],[f11000,f2958])).
fof(f11000,plain,(
  p(s(bool,'F_0')) | $spl104),
  inference(resolution,[],[f2958,f336])).
fof(f336,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f194])).
fof(f11087,plain,(
  $false | $spl104),
  inference(subsumption_resolution,[],[f10999,f2958])).
fof(f10999,plain,(
  ( ! [X213] : (p(s(bool,X213))) ) | $spl104),
  inference(resolution,[],[f2958,f333])).
fof(f333,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,'T_0'))))),
  inference(flattening,[],[f191])).
fof(f191,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f107])).
fof(f107,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X13] : ((p(s(bool,X13)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  file('thm.arithmetic.MODEQ_SYM','thm.bool.IMP_CLAUSES|split|1')).
fof(f11086,plain,(
  $false | $spl104),
  inference(subsumption_resolution,[],[f10998,f2958])).
fof(f10998,plain,(
  ( ! [X212] : (p(s(bool,X212))) ) | $spl104),
  inference(resolution,[],[f2958,f330])).
fof(f11006,plain,(
  $false | $spl104),
  inference(resolution,[],[f2958,f353])).
fof(f11007,plain,(
  $false | $spl104),
  inference(resolution,[],[f2958,f354])).
fof(f11016,plain,(
  $false | ($spl98 | $spl104)),
  inference(resolution,[],[f2958,f2597])).
fof(f11019,plain,(
  $false | ($spl56 | $spl104)),
  inference(resolution,[],[f2958,f1052])).
fof(f6833,plain,(
  ( ! [X22] : (s(bool,'T_0') != s(bool,X22) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))))) ) | $spl128),
  inference(subsumption_resolution,[],[f6815,f232])).
fof(f6815,plain,(
  ( ! [X22] : (s(bool,'T_0') != s(bool,X22) | ~p(s(bool,X22)) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))))) ) | $spl128),
  inference(superposition,[],[f3287,f233])).
fof(f6834,plain,(
  ( ! [X23] : (s(bool,'T_0') != s(bool,X23) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))))) ) | $spl128),
  inference(subsumption_resolution,[],[f6816,f232])).
fof(f6816,plain,(
  ( ! [X23] : (s(bool,'T_0') != s(bool,X23) | ~p(s(bool,X23)) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))))) ) | $spl128),
  inference(superposition,[],[f3287,f235])).
fof(f6820,plain,(
  ( ! [X23] : (s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) | ~p(s(bool,X23))) ) | $spl128),
  inference(duplicate_literal_removal,[],[f6780])).
fof(f6780,plain,(
  ( ! [X23] : (s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) | ~p(s(bool,X23)) | ~p(s(bool,X23))) ) | $spl128),
  inference(superposition,[],[f3287,f235])).
fof(f6821,plain,(
  ( ! [X22] : (s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) | ~p(s(bool,X22))) ) | $spl128),
  inference(duplicate_literal_removal,[],[f6779])).
fof(f6779,plain,(
  ( ! [X22] : (s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) | ~p(s(bool,X22)) | ~p(s(bool,X22))) ) | $spl128),
  inference(superposition,[],[f3287,f233])).
fof(f6835,plain,(
  ( ! [X24] : (~p(s(bool,X24)) | s(bool,'F_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) ) | $spl128),
  inference(subsumption_resolution,[],[f6817,f233])).
fof(f6817,plain,(
  ( ! [X24] : (s(bool,'T_0') != s(bool,X24) | ~p(s(bool,X24)) | s(bool,'F_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) ) | $spl128),
  inference(superposition,[],[f3287,f236])).
fof(f10915,plain,(
  ( ! [X290,X289] : (sP9('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X289,X290) | ~sP8(X289,X290,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl424),
  inference(cnf_transformation,[],[f10915_D])).
fof(f10915_D,plain,(
  ( ! [X290,X289] : (sP9('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X289,X290) | ~sP8(X289,X290,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) <=> ~$spl424),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl424])])).
fof(f10773,plain,(
  ( ! [X290,X289] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP9('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X289,X290) | ~sP8(X289,X290,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl4),
  inference(superposition,[],[f357,f394])).
fof(f394,plain,(
  s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(cnf_transformation,[],[f394_D])).
fof(f394_D,plain,(
  s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f357,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | p(s(bool,X0)) | ~sP8(X1,X2,X0)) )),
  inference(duplicate_literal_removal,[],[f283])).
fof(f283,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X0)) | ~sP8(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f170])).
fof(f170,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | (sP8(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP8(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f169])).
fof(f169,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | (sP8(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP8(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f168])).
fof(f168,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | (sP8(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP8(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f127])).
fof(f127,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> (sP8(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f10914,plain,(
  ( ! [X288,X287] : (sP7('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X287,X288) | ~sP6('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X287,X288)) ) | $spl422),
  inference(cnf_transformation,[],[f10914_D])).
fof(f10914_D,plain,(
  ( ! [X288,X287] : (sP7('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X287,X288) | ~sP6('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X287,X288)) ) <=> ~$spl422),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl422])])).
fof(f10772,plain,(
  ( ! [X288,X287] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP7('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X287,X288) | ~sP6('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X287,X288)) ) | $spl4),
  inference(superposition,[],[f356,f394])).
fof(f356,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~sP6(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(duplicate_literal_removal,[],[f263])).
fof(f263,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~sP6(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f162])).
fof(f162,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP6(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP6(X0,X1,X2) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f161])).
fof(f161,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP6(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP6(X2,X0,X1) | sP7(X2,X0,X1)))),
  inference(flattening,[],[f160])).
fof(f160,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP6(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP6(X2,X0,X1)) | sP7(X2,X0,X1)))),
  inference(nnf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X2,X0,X1] : (sP7(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP6(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f10913,plain,(
  ( ! [X285,X286] : (sP5('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X285,X286) | ~sP4('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X285,X286)) ) | $spl420),
  inference(cnf_transformation,[],[f10913_D])).
fof(f10913_D,plain,(
  ( ! [X285,X286] : (sP5('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X285,X286) | ~sP4('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X285,X286)) ) <=> ~$spl420),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl420])])).
fof(f10771,plain,(
  ( ! [X285,X286] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP5('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X285,X286) | ~sP4('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X285,X286)) ) | $spl4),
  inference(superposition,[],[f355,f394])).
fof(f355,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | ~sP4(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(duplicate_literal_removal,[],[f246])).
fof(f246,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | ~sP4(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f154])).
fof(f154,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP4(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP4(X0,X1,X2) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f153])).
fof(f153,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP4(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP4(X2,X0,X1) | sP5(X2,X0,X1)))),
  inference(flattening,[],[f152])).
fof(f152,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP4(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP4(X2,X0,X1)) | sP5(X2,X0,X1)))),
  inference(nnf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X2,X0,X1] : (sP5(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP4(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f6767,plain,(
  ( ! [X9] : (sP11('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X9)) ) | $spl148),
  inference(cnf_transformation,[],[f6767_D])).
fof(f6767_D,plain,(
  ( ! [X9] : (sP11('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X9)) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f10763,plain,(
  ( ! [X284] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP11('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X284)) ) | $spl4),
  inference(superposition,[],[f317,f394])).
fof(f317,plain,(
  ( ! [X0,X1] : (sP11(X0,X1) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f184])).
fof(f184,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | (p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP11(X0,X1)))),
  inference(flattening,[],[f183])).
fof(f183,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0')))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP11(X0,X1)))),
  inference(nnf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0,X1] : (sP11(X0,X1) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(definition_folding,[],[f96,f131])).
fof(f96,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f95])).
fof(f95,plain,(
  ! [X0,X1] : (((p(s(bool,X1)) | p(s(bool,X0))) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | (p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(ennf_transformation,[],[f62])).
fof(f62,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(flattening,[],[f61])).
fof(f61,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X15,X14] : ((~(p(s(bool,X14)) | p(s(bool,X15))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X14)) => p(s(bool,'F_0'))) => (~p(s(bool,X15)) => p(s(bool,'F_0')))))),
  file('thm.arithmetic.MODEQ_SYM','thm.sat.sat_9_prove_ltb')).
fof(f6742,plain,(
  ( ! [X9] : (sP11(X9,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl146),
  inference(cnf_transformation,[],[f6742_D])).
fof(f6742_D,plain,(
  ( ! [X9] : (sP11(X9,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f10912,plain,(
  ( ! [X283] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP11(X283,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl4),
  inference(subsumption_resolution,[],[f10762,f328])).
fof(f10762,plain,(
  ( ! [X283] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP11(X283,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | p(s(bool,'F_0'))) ) | $spl4),
  inference(superposition,[],[f316,f394])).
fof(f316,plain,(
  ( ! [X0,X1] : (sP11(X0,X1) | p(s(bool,'F_0')) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f184])).
fof(f10909,plain,(
  ( ! [X278] : (p(s(bool,X278)) | ~sP11('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X278)) ) | $spl418),
  inference(cnf_transformation,[],[f10909_D])).
fof(f10909_D,plain,(
  ( ! [X278] : (p(s(bool,X278)) | ~sP11('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X278)) ) <=> ~$spl418),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl418])])).
fof(f10911,plain,(
  ( ! [X282] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | p(s(bool,X282)) | ~sP11('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X282)) ) | $spl4),
  inference(subsumption_resolution,[],[f10761,f348])).
fof(f10761,plain,(
  ( ! [X282] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | p(s(bool,'F_0')) | p(s(bool,X282)) | ~sP11('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X282)) ) | $spl4),
  inference(superposition,[],[f314,f394])).
fof(f314,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP11(X0,X1)) )),
  inference(cnf_transformation,[],[f184])).
fof(f10907,plain,(
  ( ! [X277] : (p(s(bool,X277)) | ~sP11(X277,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl416),
  inference(cnf_transformation,[],[f10907_D])).
fof(f10907_D,plain,(
  ( ! [X277] : (p(s(bool,X277)) | ~sP11(X277,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) <=> ~$spl416),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl416])])).
fof(f10910,plain,(
  ( ! [X281] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | p(s(bool,X281)) | ~sP11(X281,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl4),
  inference(subsumption_resolution,[],[f10760,f348])).
fof(f10760,plain,(
  ( ! [X281] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | p(s(bool,'F_0')) | p(s(bool,X281)) | ~sP11(X281,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl4),
  inference(superposition,[],[f314,f394])).
fof(f10759,plain,(
  ( ! [X280] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP11('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X280)) ) | $spl4),
  inference(superposition,[],[f312,f394])).
fof(f10758,plain,(
  ( ! [X279] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP11(X279,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl4),
  inference(superposition,[],[f311,f394])).
fof(f10908,plain,(
  ( ! [X278] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | p(s(bool,X278)) | ~sP11('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X278)) ) | $spl4),
  inference(subsumption_resolution,[],[f10757,f348])).
fof(f10757,plain,(
  ( ! [X278] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | p(s(bool,'F_0')) | p(s(bool,X278)) | ~sP11('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X278)) ) | $spl4),
  inference(superposition,[],[f310,f394])).
fof(f310,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP11(X0,X1)) )),
  inference(cnf_transformation,[],[f182])).
fof(f10906,plain,(
  ( ! [X277] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | p(s(bool,X277)) | ~sP11(X277,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl4),
  inference(subsumption_resolution,[],[f10756,f348])).
fof(f10756,plain,(
  ( ! [X277] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | p(s(bool,'F_0')) | p(s(bool,X277)) | ~sP11(X277,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl4),
  inference(superposition,[],[f310,f394])).
fof(f10898,plain,(
  ( ! [X261] : (p(s(bool,X261)) | ~sP10(X261,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl400),
  inference(cnf_transformation,[],[f10898_D])).
fof(f10898_D,plain,(
  ( ! [X261] : (p(s(bool,X261)) | ~sP10(X261,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) <=> ~$spl400),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl400])])).
fof(f10755,plain,(
  ( ! [X276] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | p(s(bool,X276)) | ~sP10(X276,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl4),
  inference(superposition,[],[f303,f394])).
fof(f303,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | ~sP10(X1,X0)) )),
  inference(cnf_transformation,[],[f178])).
fof(f178,plain,(
  ! [X0,X1] : ((((p(s(bool,X1)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)))) | sP10(X1,X0)) & (~sP10(X1,X0) | ((~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X0)) | p(s(bool,X1))))))),
  inference(nnf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> sP10(X1,X0))),
  inference(definition_folding,[],[f58,f129])).
fof(f129,plain,(
  ! [X1,X0] : (sP10(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f58,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  inference(flattening,[],[f57])).
fof(f57,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X17,X18] : ((p(s(bool,X18)) <=> ~p(s(bool,X17))) <=> ((p(s(bool,X18)) | p(s(bool,X17))) & (~p(s(bool,X17)) | ~p(s(bool,X18)))))),
  file('thm.arithmetic.MODEQ_SYM','thm.sat.sat_29_prove_ltb')).
fof(f10899,plain,(
  ( ! [X262] : (p(s(bool,X262)) | ~sP10('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X262)) ) | $spl402),
  inference(cnf_transformation,[],[f10899_D])).
fof(f10899_D,plain,(
  ( ! [X262] : (p(s(bool,X262)) | ~sP10('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X262)) ) <=> ~$spl402),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl402])])).
fof(f10754,plain,(
  ( ! [X275] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | p(s(bool,X275)) | ~sP10('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X275)) ) | $spl4),
  inference(superposition,[],[f303,f394])).
fof(f10900,plain,(
  ( ! [X263] : (~p(s(bool,X263)) | ~sP10('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X263)) ) | $spl404),
  inference(cnf_transformation,[],[f10900_D])).
fof(f10900_D,plain,(
  ( ! [X263] : (~p(s(bool,X263)) | ~sP10('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X263)) ) <=> ~$spl404),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl404])])).
fof(f10753,plain,(
  ( ! [X274] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | ~p(s(bool,X274)) | ~sP10('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X274)) ) | $spl4),
  inference(superposition,[],[f302,f394])).
fof(f302,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | ~sP10(X1,X0)) )),
  inference(cnf_transformation,[],[f178])).
fof(f10901,plain,(
  ( ! [X264] : (~p(s(bool,X264)) | ~sP10(X264,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl406),
  inference(cnf_transformation,[],[f10901_D])).
fof(f10901_D,plain,(
  ( ! [X264] : (~p(s(bool,X264)) | ~sP10(X264,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) <=> ~$spl406),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl406])])).
fof(f10752,plain,(
  ( ! [X273] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | ~p(s(bool,X273)) | ~sP10(X273,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl4),
  inference(superposition,[],[f302,f394])).
fof(f10903,plain,(
  ( ! [X266] : (sP10('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X266) | p(s(bool,X266))) ) | $spl410),
  inference(cnf_transformation,[],[f10903_D])).
fof(f10903_D,plain,(
  ( ! [X266] : (sP10('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X266) | p(s(bool,X266))) ) <=> ~$spl410),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl410])])).
fof(f10751,plain,(
  ( ! [X272] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP10('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X272) | p(s(bool,X272))) ) | $spl4),
  inference(superposition,[],[f301,f394])).
fof(f301,plain,(
  ( ! [X0,X1] : (sP10(X1,X0) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f178])).
fof(f10902,plain,(
  ( ! [X265] : (sP10(X265,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | ~p(s(bool,X265))) ) | $spl408),
  inference(cnf_transformation,[],[f10902_D])).
fof(f10902_D,plain,(
  ( ! [X265] : (sP10(X265,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | ~p(s(bool,X265))) ) <=> ~$spl408),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl408])])).
fof(f10750,plain,(
  ( ! [X271] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP10(X271,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | ~p(s(bool,X271))) ) | $spl4),
  inference(superposition,[],[f301,f394])).
fof(f10905,plain,(
  ( ! [X268] : (sP10(X268,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | p(s(bool,X268))) ) | $spl414),
  inference(cnf_transformation,[],[f10905_D])).
fof(f10905_D,plain,(
  ( ! [X268] : (sP10(X268,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | p(s(bool,X268))) ) <=> ~$spl414),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl414])])).
fof(f10749,plain,(
  ( ! [X270] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP10(X270,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | p(s(bool,X270))) ) | $spl4),
  inference(superposition,[],[f300,f394])).
fof(f300,plain,(
  ( ! [X0,X1] : (sP10(X1,X0) | ~p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f178])).
fof(f10904,plain,(
  ( ! [X267] : (sP10('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X267) | ~p(s(bool,X267))) ) | $spl412),
  inference(cnf_transformation,[],[f10904_D])).
fof(f10904_D,plain,(
  ( ! [X267] : (sP10('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X267) | ~p(s(bool,X267))) ) <=> ~$spl412),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl412])])).
fof(f10748,plain,(
  ( ! [X269] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP10('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X269) | ~p(s(bool,X269))) ) | $spl4),
  inference(superposition,[],[f300,f394])).
fof(f10747,plain,(
  ( ! [X268] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP10(X268,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | p(s(bool,X268))) ) | $spl4),
  inference(superposition,[],[f299,f394])).
fof(f299,plain,(
  ( ! [X0,X1] : (sP10(X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f177])).
fof(f177,plain,(
  ! [X0,X1] : ((~sP10(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP10(X0,X1)))),
  inference(rectify,[],[f176])).
fof(f176,plain,(
  ! [X1,X0] : ((~sP10(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP10(X1,X0)))),
  inference(flattening,[],[f175])).
fof(f175,plain,(
  ! [X1,X0] : ((~sP10(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP10(X1,X0)))),
  inference(nnf_transformation,[],[f129])).
fof(f10746,plain,(
  ( ! [X267] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP10('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X267) | ~p(s(bool,X267))) ) | $spl4),
  inference(superposition,[],[f299,f394])).
fof(f10745,plain,(
  ( ! [X266] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP10('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X266) | p(s(bool,X266))) ) | $spl4),
  inference(superposition,[],[f296,f394])).
fof(f296,plain,(
  ( ! [X0,X1] : (sP10(X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f177])).
fof(f10744,plain,(
  ( ! [X265] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP10(X265,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | ~p(s(bool,X265))) ) | $spl4),
  inference(superposition,[],[f296,f394])).
fof(f10743,plain,(
  ( ! [X264] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | ~p(s(bool,X264)) | ~sP10(X264,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl4),
  inference(superposition,[],[f295,f394])).
fof(f295,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | ~sP10(X0,X1)) )),
  inference(cnf_transformation,[],[f177])).
fof(f10742,plain,(
  ( ! [X263] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | ~p(s(bool,X263)) | ~sP10('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X263)) ) | $spl4),
  inference(superposition,[],[f295,f394])).
fof(f10741,plain,(
  ( ! [X262] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | p(s(bool,X262)) | ~sP10('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X262)) ) | $spl4),
  inference(superposition,[],[f294,f394])).
fof(f294,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | ~sP10(X0,X1)) )),
  inference(cnf_transformation,[],[f177])).
fof(f10740,plain,(
  ( ! [X261] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | p(s(bool,X261)) | ~sP10(X261,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl4),
  inference(superposition,[],[f294,f394])).
fof(f10897,plain,(
  ( ! [X259,X260] : (p(s(bool,X259)) | ~sP9(X259,X260,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | ~p(s(bool,X260))) ) | $spl398),
  inference(cnf_transformation,[],[f10897_D])).
fof(f10897_D,plain,(
  ( ! [X259,X260] : (p(s(bool,X259)) | ~sP9(X259,X260,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | ~p(s(bool,X260))) ) <=> ~$spl398),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl398])])).
fof(f10739,plain,(
  ( ! [X259,X260] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | p(s(bool,X259)) | ~p(s(bool,X260)) | ~sP9(X259,X260,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl4),
  inference(superposition,[],[f293,f394])).
fof(f293,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | ~sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f174])).
fof(f174,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f173])).
fof(f173,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP9(X2,X0,X1))),
  inference(definition_folding,[],[f56,f127,f126])).
fof(f56,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f55])).
fof(f55,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X16,X17,X18] : ((p(s(bool,X18)) <=> (p(s(bool,X17)) & p(s(bool,X16)))) <=> ((p(s(bool,X18)) | ~p(s(bool,X17)) | ~p(s(bool,X16))) & (p(s(bool,X17)) | ~p(s(bool,X18))) & (p(s(bool,X16)) | ~p(s(bool,X18)))))),
  file('thm.arithmetic.MODEQ_SYM','thm.sat.sat_23_prove_ltb')).
fof(f10896,plain,(
  ( ! [X257,X258] : (p(s(bool,X257)) | ~sP9(X257,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X258) | ~p(s(bool,X258))) ) | $spl396),
  inference(cnf_transformation,[],[f10896_D])).
fof(f10896_D,plain,(
  ( ! [X257,X258] : (p(s(bool,X257)) | ~sP9(X257,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X258) | ~p(s(bool,X258))) ) <=> ~$spl396),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl396])])).
fof(f10738,plain,(
  ( ! [X257,X258] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | p(s(bool,X257)) | ~p(s(bool,X258)) | ~sP9(X257,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X258)) ) | $spl4),
  inference(superposition,[],[f293,f394])).
fof(f10895,plain,(
  ( ! [X255,X256] : (~p(s(bool,X255)) | ~sP9('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X255,X256) | ~p(s(bool,X256))) ) | $spl394),
  inference(cnf_transformation,[],[f10895_D])).
fof(f10895_D,plain,(
  ( ! [X255,X256] : (~p(s(bool,X255)) | ~sP9('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X255,X256) | ~p(s(bool,X256))) ) <=> ~$spl394),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl394])])).
fof(f10737,plain,(
  ( ! [X255,X256] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | ~p(s(bool,X255)) | ~p(s(bool,X256)) | ~sP9('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X255,X256)) ) | $spl4),
  inference(superposition,[],[f293,f394])).
fof(f10878,plain,(
  ( ! [X208,X207] : (p(s(bool,X207)) | ~sP9('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X207,X208)) ) | $spl360),
  inference(cnf_transformation,[],[f10878_D])).
fof(f10878_D,plain,(
  ( ! [X208,X207] : (p(s(bool,X207)) | ~sP9('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X207,X208)) ) <=> ~$spl360),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl360])])).
fof(f10736,plain,(
  ( ! [X253,X254] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | p(s(bool,X253)) | ~sP9('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X253,X254)) ) | $spl4),
  inference(superposition,[],[f292,f394])).
fof(f292,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X2)) | ~sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f174])).
fof(f10877,plain,(
  ( ! [X206,X205] : (~p(s(bool,X205)) | ~sP9(X205,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X206)) ) | $spl358),
  inference(cnf_transformation,[],[f10877_D])).
fof(f10877_D,plain,(
  ( ! [X206,X205] : (~p(s(bool,X205)) | ~sP9(X205,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X206)) ) <=> ~$spl358),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl358])])).
fof(f10735,plain,(
  ( ! [X251,X252] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | ~p(s(bool,X251)) | ~sP9(X251,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X252)) ) | $spl4),
  inference(superposition,[],[f292,f394])).
fof(f10876,plain,(
  ( ! [X204,X203] : (p(s(bool,X203)) | ~sP9('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X204,X203)) ) | $spl356),
  inference(cnf_transformation,[],[f10876_D])).
fof(f10876_D,plain,(
  ( ! [X204,X203] : (p(s(bool,X203)) | ~sP9('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X204,X203)) ) <=> ~$spl356),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl356])])).
fof(f10734,plain,(
  ( ! [X249,X250] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | p(s(bool,X249)) | ~sP9('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X250,X249)) ) | $spl4),
  inference(superposition,[],[f291,f394])).
fof(f291,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f174])).
fof(f10875,plain,(
  ( ! [X202,X201] : (~p(s(bool,X201)) | ~sP9(X201,X202,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl354),
  inference(cnf_transformation,[],[f10875_D])).
fof(f10875_D,plain,(
  ( ! [X202,X201] : (~p(s(bool,X201)) | ~sP9(X201,X202,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) <=> ~$spl354),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl354])])).
fof(f10733,plain,(
  ( ! [X248,X247] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | ~p(s(bool,X247)) | ~sP9(X247,X248,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl4),
  inference(superposition,[],[f291,f394])).
fof(f10894,plain,(
  ( ! [X246,X245] : (sP9('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X245,X246) | ~p(s(bool,X246)) | ~p(s(bool,X245))) ) | $spl392),
  inference(cnf_transformation,[],[f10894_D])).
fof(f10894_D,plain,(
  ( ! [X246,X245] : (sP9('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X245,X246) | ~p(s(bool,X246)) | ~p(s(bool,X245))) ) <=> ~$spl392),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl392])])).
fof(f10732,plain,(
  ( ! [X246,X245] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP9('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X245,X246) | ~p(s(bool,X245)) | ~p(s(bool,X246))) ) | $spl4),
  inference(superposition,[],[f290,f394])).
fof(f290,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f174])).
fof(f10893,plain,(
  ( ! [X243,X244] : (sP9(X243,X244,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | ~p(s(bool,X243)) | ~p(s(bool,X244))) ) | $spl390),
  inference(cnf_transformation,[],[f10893_D])).
fof(f10893_D,plain,(
  ( ! [X243,X244] : (sP9(X243,X244,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | ~p(s(bool,X243)) | ~p(s(bool,X244))) ) <=> ~$spl390),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl390])])).
fof(f10731,plain,(
  ( ! [X243,X244] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP9(X243,X244,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | ~p(s(bool,X244)) | ~p(s(bool,X243))) ) | $spl4),
  inference(superposition,[],[f290,f394])).
fof(f10892,plain,(
  ( ! [X241,X242] : (sP9(X241,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X242) | ~p(s(bool,X241)) | ~p(s(bool,X242))) ) | $spl388),
  inference(cnf_transformation,[],[f10892_D])).
fof(f10892_D,plain,(
  ( ! [X241,X242] : (sP9(X241,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X242) | ~p(s(bool,X241)) | ~p(s(bool,X242))) ) <=> ~$spl388),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl388])])).
fof(f10730,plain,(
  ( ! [X241,X242] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP9(X241,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X242) | ~p(s(bool,X242)) | ~p(s(bool,X241))) ) | $spl4),
  inference(superposition,[],[f290,f394])).
fof(f10891,plain,(
  ( ! [X239,X240] : (sP9(X239,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X240) | p(s(bool,X239))) ) | $spl386),
  inference(cnf_transformation,[],[f10891_D])).
fof(f10891_D,plain,(
  ( ! [X239,X240] : (sP9(X239,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X240) | p(s(bool,X239))) ) <=> ~$spl386),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl386])])).
fof(f10729,plain,(
  ( ! [X239,X240] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP9(X239,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X240) | p(s(bool,X239))) ) | $spl4),
  inference(superposition,[],[f289,f394])).
fof(f289,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f174])).
fof(f10890,plain,(
  ( ! [X237,X238] : (sP9('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X237,X238) | p(s(bool,X237))) ) | $spl384),
  inference(cnf_transformation,[],[f10890_D])).
fof(f10890_D,plain,(
  ( ! [X237,X238] : (sP9('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X237,X238) | p(s(bool,X237))) ) <=> ~$spl384),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl384])])).
fof(f10728,plain,(
  ( ! [X237,X238] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP9('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X237,X238) | p(s(bool,X237))) ) | $spl4),
  inference(superposition,[],[f289,f394])).
fof(f10889,plain,(
  ( ! [X235,X236] : (sP9(X235,X236,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | p(s(bool,X235))) ) | $spl382),
  inference(cnf_transformation,[],[f10889_D])).
fof(f10889_D,plain,(
  ( ! [X235,X236] : (sP9(X235,X236,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | p(s(bool,X235))) ) <=> ~$spl382),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl382])])).
fof(f10727,plain,(
  ( ! [X235,X236] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP9(X235,X236,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | p(s(bool,X235))) ) | $spl4),
  inference(superposition,[],[f288,f394])).
fof(f288,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f174])).
fof(f10888,plain,(
  ( ! [X233,X234] : (sP9('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X233,X234) | p(s(bool,X234))) ) | $spl380),
  inference(cnf_transformation,[],[f10888_D])).
fof(f10888_D,plain,(
  ( ! [X233,X234] : (sP9('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X233,X234) | p(s(bool,X234))) ) <=> ~$spl380),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl380])])).
fof(f10726,plain,(
  ( ! [X233,X234] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP9('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X233,X234) | p(s(bool,X234))) ) | $spl4),
  inference(superposition,[],[f288,f394])).
fof(f10887,plain,(
  ( ! [X231,X232] : (sP8('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X231,X232)) ) | $spl378),
  inference(cnf_transformation,[],[f10887_D])).
fof(f10887_D,plain,(
  ( ! [X231,X232] : (sP8('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X231,X232)) ) <=> ~$spl378),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl378])])).
fof(f10725,plain,(
  ( ! [X231,X232] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP8('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X231,X232)) ) | $spl4),
  inference(superposition,[],[f287,f394])).
fof(f10886,plain,(
  ( ! [X229,X230] : (sP8(X229,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X230)) ) | $spl376),
  inference(cnf_transformation,[],[f10886_D])).
fof(f10886_D,plain,(
  ( ! [X229,X230] : (sP8(X229,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X230)) ) <=> ~$spl376),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl376])])).
fof(f10724,plain,(
  ( ! [X229,X230] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP8(X229,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X230)) ) | $spl4),
  inference(superposition,[],[f286,f394])).
fof(f7780,plain,(
  ( ! [X12,X13] : (sP8(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl188),
  inference(cnf_transformation,[],[f7780_D])).
fof(f7780_D,plain,(
  ( ! [X12,X13] : (sP8(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f10723,plain,(
  ( ! [X227,X228] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP8(X227,X228,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl4),
  inference(superposition,[],[f285,f394])).
fof(f10885,plain,(
  ( ! [X225,X226] : (~p(s(bool,X225)) | ~sP8(X225,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X226) | p(s(bool,X226))) ) | $spl374),
  inference(cnf_transformation,[],[f10885_D])).
fof(f10885_D,plain,(
  ( ! [X225,X226] : (~p(s(bool,X225)) | ~sP8(X225,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X226) | p(s(bool,X226))) ) <=> ~$spl374),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl374])])).
fof(f10722,plain,(
  ( ! [X225,X226] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | ~p(s(bool,X225)) | p(s(bool,X226)) | ~sP8(X225,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X226)) ) | $spl4),
  inference(superposition,[],[f284,f394])).
fof(f10884,plain,(
  ( ! [X223,X224] : (~p(s(bool,X223)) | ~sP8('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X223,X224) | p(s(bool,X224))) ) | $spl372),
  inference(cnf_transformation,[],[f10884_D])).
fof(f10884_D,plain,(
  ( ! [X223,X224] : (~p(s(bool,X223)) | ~sP8('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X223,X224) | p(s(bool,X224))) ) <=> ~$spl372),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl372])])).
fof(f10721,plain,(
  ( ! [X223,X224] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | ~p(s(bool,X223)) | p(s(bool,X224)) | ~sP8('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X223,X224)) ) | $spl4),
  inference(superposition,[],[f284,f394])).
fof(f10883,plain,(
  ( ! [X222,X221] : (~p(s(bool,X221)) | ~sP8(X221,X222,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | ~p(s(bool,X222))) ) | $spl370),
  inference(cnf_transformation,[],[f10883_D])).
fof(f10883_D,plain,(
  ( ! [X222,X221] : (~p(s(bool,X221)) | ~sP8(X221,X222,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | ~p(s(bool,X222))) ) <=> ~$spl370),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl370])])).
fof(f10720,plain,(
  ( ! [X222,X221] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | ~p(s(bool,X221)) | ~p(s(bool,X222)) | ~sP8(X221,X222,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl4),
  inference(superposition,[],[f284,f394])).
fof(f10882,plain,(
  ( ! [X218,X217] : (sP9('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X217,X218) | ~sP8(X217,X218,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | ~p(s(bool,X218))) ) | $spl368),
  inference(cnf_transformation,[],[f10882_D])).
fof(f10882_D,plain,(
  ( ! [X218,X217] : (sP9('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X217,X218) | ~sP8(X217,X218,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | ~p(s(bool,X218))) ) <=> ~$spl368),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl368])])).
fof(f10718,plain,(
  ( ! [X218,X217] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP9('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X217,X218) | ~p(s(bool,X218)) | ~sP8(X217,X218,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl4),
  inference(superposition,[],[f282,f394])).
fof(f282,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X2)) | ~sP8(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f170])).
fof(f10881,plain,(
  ( ! [X213,X214] : (sP9('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X213,X214) | ~sP8(X213,X214,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | ~p(s(bool,X213))) ) | $spl366),
  inference(cnf_transformation,[],[f10881_D])).
fof(f10881_D,plain,(
  ( ! [X213,X214] : (sP9('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X213,X214) | ~sP8(X213,X214,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | ~p(s(bool,X213))) ) <=> ~$spl366),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl366])])).
fof(f10716,plain,(
  ( ! [X213,X214] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP9('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X213,X214) | ~p(s(bool,X213)) | ~sP8(X213,X214,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl4),
  inference(superposition,[],[f281,f394])).
fof(f281,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~sP8(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f170])).
fof(f10880,plain,(
  ( ! [X212,X211] : (sP9(X211,X212,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | ~sP8(X212,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X211) | ~p(s(bool,X212))) ) | $spl364),
  inference(cnf_transformation,[],[f10880_D])).
fof(f10880_D,plain,(
  ( ! [X212,X211] : (sP9(X211,X212,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | ~sP8(X212,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X211) | ~p(s(bool,X212))) ) <=> ~$spl364),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl364])])).
fof(f10715,plain,(
  ( ! [X212,X211] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP9(X211,X212,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | ~p(s(bool,X212)) | ~sP8(X212,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X211)) ) | $spl4),
  inference(superposition,[],[f280,f394])).
fof(f280,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP8(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f170])).
fof(f10879,plain,(
  ( ! [X210,X209] : (sP9(X209,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X210) | ~sP8('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X210,X209) | ~p(s(bool,X210))) ) | $spl362),
  inference(cnf_transformation,[],[f10879_D])).
fof(f10879_D,plain,(
  ( ! [X210,X209] : (sP9(X209,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X210) | ~sP8('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X210,X209) | ~p(s(bool,X210))) ) <=> ~$spl362),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl362])])).
fof(f10714,plain,(
  ( ! [X210,X209] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP9(X209,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X210) | ~p(s(bool,X210)) | ~sP8('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X210,X209)) ) | $spl4),
  inference(superposition,[],[f280,f394])).
fof(f10713,plain,(
  ( ! [X208,X207] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | p(s(bool,X207)) | ~sP9('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X207,X208)) ) | $spl4),
  inference(superposition,[],[f279,f394])).
fof(f279,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f170])).
fof(f10712,plain,(
  ( ! [X206,X205] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | ~p(s(bool,X205)) | ~sP9(X205,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X206)) ) | $spl4),
  inference(superposition,[],[f279,f394])).
fof(f10711,plain,(
  ( ! [X204,X203] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | p(s(bool,X203)) | ~sP9('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X204,X203)) ) | $spl4),
  inference(superposition,[],[f278,f394])).
fof(f278,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X2)) | ~sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f170])).
fof(f10710,plain,(
  ( ! [X202,X201] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | ~p(s(bool,X201)) | ~sP9(X201,X202,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl4),
  inference(superposition,[],[f278,f394])).
fof(f10856,plain,(
  ( ! [X147,X148] : (p(s(bool,X147)) | ~sP7(X147,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X148)) ) | $spl316),
  inference(cnf_transformation,[],[f10856_D])).
fof(f10856_D,plain,(
  ( ! [X147,X148] : (p(s(bool,X147)) | ~sP7(X147,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X148)) ) <=> ~$spl316),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl316])])).
fof(f10709,plain,(
  ( ! [X200,X199] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | p(s(bool,X199)) | ~sP7(X199,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X200)) ) | $spl4),
  inference(superposition,[],[f276,f394])).
fof(f276,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~sP7(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f167])).
fof(f167,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f166])).
fof(f166,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP7(X2,X0,X1))),
  inference(definition_folding,[],[f54,f124,f123])).
fof(f54,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f53])).
fof(f53,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X16,X17,X18] : ((p(s(bool,X18)) <=> (p(s(bool,X17)) | p(s(bool,X16)))) <=> ((p(s(bool,X18)) | ~p(s(bool,X17))) & (p(s(bool,X18)) | ~p(s(bool,X16))) & (p(s(bool,X17)) | p(s(bool,X16)) | ~p(s(bool,X18)))))),
  file('thm.arithmetic.MODEQ_SYM','thm.sat.sat_25_prove_ltb')).
fof(f10855,plain,(
  ( ! [X146,X145] : (~p(s(bool,X145)) | ~sP7('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X145,X146)) ) | $spl314),
  inference(cnf_transformation,[],[f10855_D])).
fof(f10855_D,plain,(
  ( ! [X146,X145] : (~p(s(bool,X145)) | ~sP7('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X145,X146)) ) <=> ~$spl314),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl314])])).
fof(f10708,plain,(
  ( ! [X198,X197] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | ~p(s(bool,X197)) | ~sP7('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X197,X198)) ) | $spl4),
  inference(superposition,[],[f276,f394])).
fof(f10854,plain,(
  ( ! [X144,X143] : (p(s(bool,X143)) | ~sP7(X143,X144,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl312),
  inference(cnf_transformation,[],[f10854_D])).
fof(f10854_D,plain,(
  ( ! [X144,X143] : (p(s(bool,X143)) | ~sP7(X143,X144,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) <=> ~$spl312),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl312])])).
fof(f10707,plain,(
  ( ! [X196,X195] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | p(s(bool,X195)) | ~sP7(X195,X196,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl4),
  inference(superposition,[],[f275,f394])).
fof(f275,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X1)) | ~sP7(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f167])).
fof(f10853,plain,(
  ( ! [X142,X141] : (~p(s(bool,X141)) | ~sP7('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X142,X141)) ) | $spl310),
  inference(cnf_transformation,[],[f10853_D])).
fof(f10853_D,plain,(
  ( ! [X142,X141] : (~p(s(bool,X141)) | ~sP7('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X142,X141)) ) <=> ~$spl310),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl310])])).
fof(f10706,plain,(
  ( ! [X194,X193] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | ~p(s(bool,X193)) | ~sP7('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X194,X193)) ) | $spl4),
  inference(superposition,[],[f275,f394])).
fof(f10874,plain,(
  ( ! [X191,X192] : (p(s(bool,X191)) | ~sP7('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X191,X192) | p(s(bool,X192))) ) | $spl352),
  inference(cnf_transformation,[],[f10874_D])).
fof(f10874_D,plain,(
  ( ! [X191,X192] : (p(s(bool,X191)) | ~sP7('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X191,X192) | p(s(bool,X192))) ) <=> ~$spl352),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl352])])).
fof(f10705,plain,(
  ( ! [X191,X192] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | p(s(bool,X191)) | p(s(bool,X192)) | ~sP7('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X191,X192)) ) | $spl4),
  inference(superposition,[],[f274,f394])).
fof(f274,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP7(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f167])).
fof(f10873,plain,(
  ( ! [X189,X190] : (p(s(bool,X189)) | ~p(s(bool,X190)) | ~sP7(X190,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X189)) ) | $spl350),
  inference(cnf_transformation,[],[f10873_D])).
fof(f10873_D,plain,(
  ( ! [X189,X190] : (p(s(bool,X189)) | ~p(s(bool,X190)) | ~sP7(X190,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X189)) ) <=> ~$spl350),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl350])])).
fof(f10704,plain,(
  ( ! [X189,X190] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | p(s(bool,X189)) | ~p(s(bool,X190)) | ~sP7(X190,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X189)) ) | $spl4),
  inference(superposition,[],[f274,f394])).
fof(f10872,plain,(
  ( ! [X187,X188] : (p(s(bool,X187)) | ~p(s(bool,X188)) | ~sP7(X188,X187,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl348),
  inference(cnf_transformation,[],[f10872_D])).
fof(f10872_D,plain,(
  ( ! [X187,X188] : (p(s(bool,X187)) | ~p(s(bool,X188)) | ~sP7(X188,X187,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) <=> ~$spl348),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl348])])).
fof(f10703,plain,(
  ( ! [X187,X188] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | p(s(bool,X187)) | ~p(s(bool,X188)) | ~sP7(X188,X187,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl4),
  inference(superposition,[],[f274,f394])).
fof(f10871,plain,(
  ( ! [X185,X186] : (sP7('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X185,X186) | ~p(s(bool,X185))) ) | $spl346),
  inference(cnf_transformation,[],[f10871_D])).
fof(f10871_D,plain,(
  ( ! [X185,X186] : (sP7('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X185,X186) | ~p(s(bool,X185))) ) <=> ~$spl346),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl346])])).
fof(f10702,plain,(
  ( ! [X185,X186] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP7('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X185,X186) | ~p(s(bool,X185))) ) | $spl4),
  inference(superposition,[],[f273,f394])).
fof(f273,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f167])).
fof(f10870,plain,(
  ( ! [X184,X183] : (sP7(X183,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X184) | ~p(s(bool,X183))) ) | $spl344),
  inference(cnf_transformation,[],[f10870_D])).
fof(f10870_D,plain,(
  ( ! [X184,X183] : (sP7(X183,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X184) | ~p(s(bool,X183))) ) <=> ~$spl344),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl344])])).
fof(f10701,plain,(
  ( ! [X184,X183] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP7(X183,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X184) | ~p(s(bool,X183))) ) | $spl4),
  inference(superposition,[],[f273,f394])).
fof(f10869,plain,(
  ( ! [X182,X181] : (sP7('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X181,X182) | ~p(s(bool,X182))) ) | $spl342),
  inference(cnf_transformation,[],[f10869_D])).
fof(f10869_D,plain,(
  ( ! [X182,X181] : (sP7('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X181,X182) | ~p(s(bool,X182))) ) <=> ~$spl342),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl342])])).
fof(f10700,plain,(
  ( ! [X182,X181] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP7('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X181,X182) | ~p(s(bool,X182))) ) | $spl4),
  inference(superposition,[],[f272,f394])).
fof(f272,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f167])).
fof(f10868,plain,(
  ( ! [X180,X179] : (sP7(X179,X180,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | ~p(s(bool,X179))) ) | $spl340),
  inference(cnf_transformation,[],[f10868_D])).
fof(f10868_D,plain,(
  ( ! [X180,X179] : (sP7(X179,X180,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | ~p(s(bool,X179))) ) <=> ~$spl340),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl340])])).
fof(f10699,plain,(
  ( ! [X180,X179] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP7(X179,X180,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | ~p(s(bool,X179))) ) | $spl4),
  inference(superposition,[],[f272,f394])).
fof(f10867,plain,(
  ( ! [X177,X178] : (sP7(X177,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X178) | p(s(bool,X177)) | p(s(bool,X178))) ) | $spl338),
  inference(cnf_transformation,[],[f10867_D])).
fof(f10867_D,plain,(
  ( ! [X177,X178] : (sP7(X177,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X178) | p(s(bool,X177)) | p(s(bool,X178))) ) <=> ~$spl338),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl338])])).
fof(f10698,plain,(
  ( ! [X177,X178] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP7(X177,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X178) | p(s(bool,X178)) | p(s(bool,X177))) ) | $spl4),
  inference(superposition,[],[f271,f394])).
fof(f271,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f167])).
fof(f10866,plain,(
  ( ! [X175,X176] : (sP7(X175,X176,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | p(s(bool,X175)) | p(s(bool,X176))) ) | $spl336),
  inference(cnf_transformation,[],[f10866_D])).
fof(f10866_D,plain,(
  ( ! [X175,X176] : (sP7(X175,X176,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | p(s(bool,X175)) | p(s(bool,X176))) ) <=> ~$spl336),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl336])])).
fof(f10697,plain,(
  ( ! [X175,X176] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP7(X175,X176,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | p(s(bool,X176)) | p(s(bool,X175))) ) | $spl4),
  inference(superposition,[],[f271,f394])).
fof(f10865,plain,(
  ( ! [X173,X174] : (sP7('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X173,X174) | p(s(bool,X174)) | p(s(bool,X173))) ) | $spl334),
  inference(cnf_transformation,[],[f10865_D])).
fof(f10865_D,plain,(
  ( ! [X173,X174] : (sP7('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X173,X174) | p(s(bool,X174)) | p(s(bool,X173))) ) <=> ~$spl334),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl334])])).
fof(f10696,plain,(
  ( ! [X173,X174] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP7('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X173,X174) | p(s(bool,X173)) | p(s(bool,X174))) ) | $spl4),
  inference(superposition,[],[f271,f394])).
fof(f10864,plain,(
  ( ! [X171,X172] : (sP6('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X171,X172)) ) | $spl332),
  inference(cnf_transformation,[],[f10864_D])).
fof(f10864_D,plain,(
  ( ! [X171,X172] : (sP6('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X171,X172)) ) <=> ~$spl332),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl332])])).
fof(f10695,plain,(
  ( ! [X171,X172] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP6('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X171,X172)) ) | $spl4),
  inference(superposition,[],[f270,f394])).
fof(f7748,plain,(
  ( ! [X12,X13] : (sP6(X12,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X13)) ) | $spl186),
  inference(cnf_transformation,[],[f7748_D])).
fof(f7748_D,plain,(
  ( ! [X12,X13] : (sP6(X12,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X13)) ) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f10694,plain,(
  ( ! [X169,X170] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP6(X169,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X170)) ) | $spl4),
  inference(superposition,[],[f269,f394])).
fof(f7685,plain,(
  ( ! [X12,X13] : (sP6(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl182),
  inference(cnf_transformation,[],[f7685_D])).
fof(f7685_D,plain,(
  ( ! [X12,X13] : (sP6(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f10693,plain,(
  ( ! [X167,X168] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP6(X167,X168,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl4),
  inference(superposition,[],[f268,f394])).
fof(f10863,plain,(
  ( ! [X165,X166] : (p(s(bool,X165)) | ~sP6('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X165,X166) | p(s(bool,X166))) ) | $spl330),
  inference(cnf_transformation,[],[f10863_D])).
fof(f10863_D,plain,(
  ( ! [X165,X166] : (p(s(bool,X165)) | ~sP6('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X165,X166) | p(s(bool,X166))) ) <=> ~$spl330),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl330])])).
fof(f10692,plain,(
  ( ! [X165,X166] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | p(s(bool,X165)) | p(s(bool,X166)) | ~sP6('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X165,X166)) ) | $spl4),
  inference(superposition,[],[f267,f394])).
fof(f10862,plain,(
  ( ! [X163,X164] : (~p(s(bool,X163)) | ~sP6(X163,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X164) | p(s(bool,X164))) ) | $spl328),
  inference(cnf_transformation,[],[f10862_D])).
fof(f10862_D,plain,(
  ( ! [X163,X164] : (~p(s(bool,X163)) | ~sP6(X163,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X164) | p(s(bool,X164))) ) <=> ~$spl328),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl328])])).
fof(f10691,plain,(
  ( ! [X163,X164] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | ~p(s(bool,X163)) | p(s(bool,X164)) | ~sP6(X163,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X164)) ) | $spl4),
  inference(superposition,[],[f267,f394])).
fof(f10861,plain,(
  ( ! [X161,X162] : (~p(s(bool,X161)) | ~sP6(X161,X162,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | p(s(bool,X162))) ) | $spl326),
  inference(cnf_transformation,[],[f10861_D])).
fof(f10861_D,plain,(
  ( ! [X161,X162] : (~p(s(bool,X161)) | ~sP6(X161,X162,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | p(s(bool,X162))) ) <=> ~$spl326),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl326])])).
fof(f10690,plain,(
  ( ! [X161,X162] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | ~p(s(bool,X161)) | p(s(bool,X162)) | ~sP6(X161,X162,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl4),
  inference(superposition,[],[f267,f394])).
fof(f10860,plain,(
  ( ! [X159,X160] : (sP7(X159,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X160) | p(s(bool,X160)) | ~sP6(X159,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X160)) ) | $spl324),
  inference(cnf_transformation,[],[f10860_D])).
fof(f10860_D,plain,(
  ( ! [X159,X160] : (sP7(X159,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X160) | p(s(bool,X160)) | ~sP6(X159,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X160)) ) <=> ~$spl324),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl324])])).
fof(f10689,plain,(
  ( ! [X159,X160] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP7(X159,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X160) | ~sP6(X159,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X160) | p(s(bool,X160))) ) | $spl4),
  inference(superposition,[],[f266,f394])).
fof(f266,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~sP6(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f162])).
fof(f10859,plain,(
  ( ! [X158,X157] : (sP7(X157,X158,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | p(s(bool,X158)) | ~sP6(X157,X158,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl322),
  inference(cnf_transformation,[],[f10859_D])).
fof(f10859_D,plain,(
  ( ! [X158,X157] : (sP7(X157,X158,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | p(s(bool,X158)) | ~sP6(X157,X158,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) <=> ~$spl322),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl322])])).
fof(f10688,plain,(
  ( ! [X158,X157] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP7(X157,X158,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | ~sP6(X157,X158,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | p(s(bool,X158))) ) | $spl4),
  inference(superposition,[],[f266,f394])).
fof(f10858,plain,(
  ( ! [X156,X155] : (sP7('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X155,X156) | p(s(bool,X155)) | ~sP6('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X155,X156)) ) | $spl320),
  inference(cnf_transformation,[],[f10858_D])).
fof(f10858_D,plain,(
  ( ! [X156,X155] : (sP7('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X155,X156) | p(s(bool,X155)) | ~sP6('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X155,X156)) ) <=> ~$spl320),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl320])])).
fof(f10687,plain,(
  ( ! [X156,X155] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP7('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X155,X156) | ~sP6('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X155,X156) | p(s(bool,X155))) ) | $spl4),
  inference(superposition,[],[f265,f394])).
fof(f265,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~sP6(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f162])).
fof(f10857,plain,(
  ( ! [X152,X151] : (sP7('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X151,X152) | p(s(bool,X152)) | ~sP6('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X151,X152)) ) | $spl318),
  inference(cnf_transformation,[],[f10857_D])).
fof(f10857_D,plain,(
  ( ! [X152,X151] : (sP7('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X151,X152) | p(s(bool,X152)) | ~sP6('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X151,X152)) ) <=> ~$spl318),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl318])])).
fof(f10685,plain,(
  ( ! [X152,X151] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP7('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X151,X152) | ~sP6('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X151,X152) | p(s(bool,X152))) ) | $spl4),
  inference(superposition,[],[f264,f394])).
fof(f264,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~sP6(X0,X1,X2) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f162])).
fof(f10683,plain,(
  ( ! [X147,X148] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | p(s(bool,X147)) | ~sP7(X147,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X148)) ) | $spl4),
  inference(superposition,[],[f261,f394])).
fof(f261,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | ~sP7(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f162])).
fof(f10682,plain,(
  ( ! [X146,X145] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | ~p(s(bool,X145)) | ~sP7('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X145,X146)) ) | $spl4),
  inference(superposition,[],[f261,f394])).
fof(f10681,plain,(
  ( ! [X144,X143] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | p(s(bool,X143)) | ~sP7(X143,X144,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl4),
  inference(superposition,[],[f260,f394])).
fof(f260,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X0)) | ~sP7(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f162])).
fof(f10680,plain,(
  ( ! [X142,X141] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | ~p(s(bool,X141)) | ~sP7('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X142,X141)) ) | $spl4),
  inference(superposition,[],[f260,f394])).
fof(f10833,plain,(
  ( ! [X88,X87] : (p(s(bool,X87)) | ~sP5(X87,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X88)) ) | $spl270),
  inference(cnf_transformation,[],[f10833_D])).
fof(f10833_D,plain,(
  ( ! [X88,X87] : (p(s(bool,X87)) | ~sP5(X87,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X88)) ) <=> ~$spl270),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl270])])).
fof(f10679,plain,(
  ( ! [X140,X139] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | p(s(bool,X139)) | ~sP5(X139,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X140)) ) | $spl4),
  inference(superposition,[],[f259,f394])).
fof(f259,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~sP5(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f159])).
fof(f159,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP5(X2,X0,X1)) & (~sP5(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f158])).
fof(f158,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP5(X2,X0,X1)) & (~sP5(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP5(X2,X0,X1))),
  inference(definition_folding,[],[f92,f121,f120])).
fof(f92,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f51])).
fof(f51,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X16,X17,X18] : ((p(s(bool,X18)) <=> (p(s(bool,X17)) => p(s(bool,X16)))) <=> ((p(s(bool,X18)) | p(s(bool,X17))) & (p(s(bool,X18)) | ~p(s(bool,X16))) & (~p(s(bool,X17)) | p(s(bool,X16)) | ~p(s(bool,X18)))))),
  file('thm.arithmetic.MODEQ_SYM','thm.sat.sat_27_prove_ltb')).
fof(f10832,plain,(
  ( ! [X85,X86] : (~p(s(bool,X85)) | ~sP5('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X85,X86)) ) | $spl268),
  inference(cnf_transformation,[],[f10832_D])).
fof(f10832_D,plain,(
  ( ! [X85,X86] : (~p(s(bool,X85)) | ~sP5('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X85,X86)) ) <=> ~$spl268),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl268])])).
fof(f10678,plain,(
  ( ! [X138,X137] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | ~p(s(bool,X137)) | ~sP5('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X137,X138)) ) | $spl4),
  inference(superposition,[],[f259,f394])).
fof(f10830,plain,(
  ( ! [X81,X82] : (p(s(bool,X81)) | ~sP5(X81,X82,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl264),
  inference(cnf_transformation,[],[f10830_D])).
fof(f10830_D,plain,(
  ( ! [X81,X82] : (p(s(bool,X81)) | ~sP5(X81,X82,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) <=> ~$spl264),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl264])])).
fof(f10677,plain,(
  ( ! [X136,X135] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | p(s(bool,X135)) | ~sP5(X135,X136,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl4),
  inference(superposition,[],[f258,f394])).
fof(f258,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | p(s(bool,X1)) | ~sP5(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f159])).
fof(f10831,plain,(
  ( ! [X83,X84] : (p(s(bool,X83)) | ~sP5('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X84,X83)) ) | $spl266),
  inference(cnf_transformation,[],[f10831_D])).
fof(f10831_D,plain,(
  ( ! [X83,X84] : (p(s(bool,X83)) | ~sP5('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X84,X83)) ) <=> ~$spl266),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl266])])).
fof(f10676,plain,(
  ( ! [X134,X133] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | p(s(bool,X133)) | ~sP5('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X134,X133)) ) | $spl4),
  inference(superposition,[],[f258,f394])).
fof(f10852,plain,(
  ( ! [X132,X131] : (p(s(bool,X131)) | ~sP5('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X131,X132) | ~p(s(bool,X132))) ) | $spl308),
  inference(cnf_transformation,[],[f10852_D])).
fof(f10852_D,plain,(
  ( ! [X132,X131] : (p(s(bool,X131)) | ~sP5('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X131,X132) | ~p(s(bool,X132))) ) <=> ~$spl308),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl308])])).
fof(f10675,plain,(
  ( ! [X132,X131] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | p(s(bool,X131)) | ~p(s(bool,X132)) | ~sP5('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X131,X132)) ) | $spl4),
  inference(superposition,[],[f257,f394])).
fof(f257,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP5(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f159])).
fof(f10851,plain,(
  ( ! [X130,X129] : (p(s(bool,X129)) | ~p(s(bool,X130)) | ~sP5(X130,X129,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl306),
  inference(cnf_transformation,[],[f10851_D])).
fof(f10851_D,plain,(
  ( ! [X130,X129] : (p(s(bool,X129)) | ~p(s(bool,X130)) | ~sP5(X130,X129,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) <=> ~$spl306),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl306])])).
fof(f10674,plain,(
  ( ! [X130,X129] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | p(s(bool,X129)) | ~p(s(bool,X130)) | ~sP5(X130,X129,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl4),
  inference(superposition,[],[f257,f394])).
fof(f10850,plain,(
  ( ! [X127,X128] : (~p(s(bool,X127)) | ~p(s(bool,X128)) | ~sP5(X128,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X127)) ) | $spl304),
  inference(cnf_transformation,[],[f10850_D])).
fof(f10850_D,plain,(
  ( ! [X127,X128] : (~p(s(bool,X127)) | ~p(s(bool,X128)) | ~sP5(X128,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X127)) ) <=> ~$spl304),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl304])])).
fof(f10673,plain,(
  ( ! [X127,X128] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | ~p(s(bool,X127)) | ~p(s(bool,X128)) | ~sP5(X128,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X127)) ) | $spl4),
  inference(superposition,[],[f257,f394])).
fof(f10849,plain,(
  ( ! [X125,X126] : (sP5('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X125,X126) | ~p(s(bool,X125))) ) | $spl302),
  inference(cnf_transformation,[],[f10849_D])).
fof(f10849_D,plain,(
  ( ! [X125,X126] : (sP5('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X125,X126) | ~p(s(bool,X125))) ) <=> ~$spl302),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl302])])).
fof(f10672,plain,(
  ( ! [X125,X126] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP5('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X125,X126) | ~p(s(bool,X125))) ) | $spl4),
  inference(superposition,[],[f256,f394])).
fof(f256,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f159])).
fof(f10848,plain,(
  ( ! [X123,X124] : (sP5(X123,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X124) | ~p(s(bool,X123))) ) | $spl300),
  inference(cnf_transformation,[],[f10848_D])).
fof(f10848_D,plain,(
  ( ! [X123,X124] : (sP5(X123,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X124) | ~p(s(bool,X123))) ) <=> ~$spl300),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl300])])).
fof(f10671,plain,(
  ( ! [X123,X124] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP5(X123,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X124) | ~p(s(bool,X123))) ) | $spl4),
  inference(superposition,[],[f256,f394])).
fof(f10847,plain,(
  ( ! [X121,X122] : (sP5('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X121,X122) | p(s(bool,X122))) ) | $spl298),
  inference(cnf_transformation,[],[f10847_D])).
fof(f10847_D,plain,(
  ( ! [X121,X122] : (sP5('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X121,X122) | p(s(bool,X122))) ) <=> ~$spl298),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl298])])).
fof(f10670,plain,(
  ( ! [X121,X122] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP5('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X121,X122) | p(s(bool,X122))) ) | $spl4),
  inference(superposition,[],[f255,f394])).
fof(f255,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f159])).
fof(f10846,plain,(
  ( ! [X120,X119] : (sP5(X119,X120,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | ~p(s(bool,X119))) ) | $spl296),
  inference(cnf_transformation,[],[f10846_D])).
fof(f10846_D,plain,(
  ( ! [X120,X119] : (sP5(X119,X120,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | ~p(s(bool,X119))) ) <=> ~$spl296),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl296])])).
fof(f10669,plain,(
  ( ! [X120,X119] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP5(X119,X120,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | ~p(s(bool,X119))) ) | $spl4),
  inference(superposition,[],[f255,f394])).
fof(f10845,plain,(
  ( ! [X118,X117] : (sP5(X117,X118,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | p(s(bool,X117)) | p(s(bool,X118))) ) | $spl294),
  inference(cnf_transformation,[],[f10845_D])).
fof(f10845_D,plain,(
  ( ! [X118,X117] : (sP5(X117,X118,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | p(s(bool,X117)) | p(s(bool,X118))) ) <=> ~$spl294),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl294])])).
fof(f10668,plain,(
  ( ! [X118,X117] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP5(X117,X118,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | p(s(bool,X118)) | p(s(bool,X117))) ) | $spl4),
  inference(superposition,[],[f254,f394])).
fof(f254,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f159])).
fof(f10844,plain,(
  ( ! [X116,X115] : (sP5(X115,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X116) | p(s(bool,X115)) | ~p(s(bool,X116))) ) | $spl292),
  inference(cnf_transformation,[],[f10844_D])).
fof(f10844_D,plain,(
  ( ! [X116,X115] : (sP5(X115,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X116) | p(s(bool,X115)) | ~p(s(bool,X116))) ) <=> ~$spl292),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl292])])).
fof(f10667,plain,(
  ( ! [X116,X115] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP5(X115,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X116) | ~p(s(bool,X116)) | p(s(bool,X115))) ) | $spl4),
  inference(superposition,[],[f254,f394])).
fof(f10843,plain,(
  ( ! [X114,X113] : (sP5('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X113,X114) | ~p(s(bool,X114)) | p(s(bool,X113))) ) | $spl290),
  inference(cnf_transformation,[],[f10843_D])).
fof(f10843_D,plain,(
  ( ! [X114,X113] : (sP5('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X113,X114) | ~p(s(bool,X114)) | p(s(bool,X113))) ) <=> ~$spl290),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl290])])).
fof(f10666,plain,(
  ( ! [X114,X113] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP5('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X113,X114) | p(s(bool,X113)) | ~p(s(bool,X114))) ) | $spl4),
  inference(superposition,[],[f254,f394])).
fof(f10842,plain,(
  ( ! [X111,X112] : (sP4('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X111,X112)) ) | $spl288),
  inference(cnf_transformation,[],[f10842_D])).
fof(f10842_D,plain,(
  ( ! [X111,X112] : (sP4('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X111,X112)) ) <=> ~$spl288),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl288])])).
fof(f10665,plain,(
  ( ! [X111,X112] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP4('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X111,X112)) ) | $spl4),
  inference(superposition,[],[f253,f394])).
fof(f7657,plain,(
  ( ! [X12,X13] : (sP4(X12,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X13)) ) | $spl180),
  inference(cnf_transformation,[],[f7657_D])).
fof(f7657_D,plain,(
  ( ! [X12,X13] : (sP4(X12,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X13)) ) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f10664,plain,(
  ( ! [X109,X110] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP4(X109,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X110)) ) | $spl4),
  inference(superposition,[],[f252,f394])).
fof(f10841,plain,(
  ( ! [X107,X108] : (sP4(X107,X108,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl286),
  inference(cnf_transformation,[],[f10841_D])).
fof(f10841_D,plain,(
  ( ! [X107,X108] : (sP4(X107,X108,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) <=> ~$spl286),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl286])])).
fof(f10663,plain,(
  ( ! [X107,X108] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP4(X107,X108,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl4),
  inference(superposition,[],[f251,f394])).
fof(f10840,plain,(
  ( ! [X105,X106] : (~p(s(bool,X105)) | ~sP4(X105,X106,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | p(s(bool,X106))) ) | $spl284),
  inference(cnf_transformation,[],[f10840_D])).
fof(f10840_D,plain,(
  ( ! [X105,X106] : (~p(s(bool,X105)) | ~sP4(X105,X106,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | p(s(bool,X106))) ) <=> ~$spl284),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl284])])).
fof(f10662,plain,(
  ( ! [X105,X106] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | ~p(s(bool,X105)) | p(s(bool,X106)) | ~sP4(X105,X106,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl4),
  inference(superposition,[],[f250,f394])).
fof(f10839,plain,(
  ( ! [X103,X104] : (p(s(bool,X103)) | ~sP4('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X103,X104) | ~p(s(bool,X104))) ) | $spl282),
  inference(cnf_transformation,[],[f10839_D])).
fof(f10839_D,plain,(
  ( ! [X103,X104] : (p(s(bool,X103)) | ~sP4('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X103,X104) | ~p(s(bool,X104))) ) <=> ~$spl282),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl282])])).
fof(f10661,plain,(
  ( ! [X103,X104] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | p(s(bool,X103)) | ~p(s(bool,X104)) | ~sP4('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X103,X104)) ) | $spl4),
  inference(superposition,[],[f250,f394])).
fof(f10838,plain,(
  ( ! [X101,X102] : (~p(s(bool,X101)) | ~sP4(X101,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X102) | ~p(s(bool,X102))) ) | $spl280),
  inference(cnf_transformation,[],[f10838_D])).
fof(f10838_D,plain,(
  ( ! [X101,X102] : (~p(s(bool,X101)) | ~sP4(X101,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X102) | ~p(s(bool,X102))) ) <=> ~$spl280),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl280])])).
fof(f10660,plain,(
  ( ! [X101,X102] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | ~p(s(bool,X101)) | ~p(s(bool,X102)) | ~sP4(X101,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X102)) ) | $spl4),
  inference(superposition,[],[f250,f394])).
fof(f10837,plain,(
  ( ! [X99,X100] : (sP5(X99,X100,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | p(s(bool,X100)) | ~sP4(X99,X100,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl278),
  inference(cnf_transformation,[],[f10837_D])).
fof(f10837_D,plain,(
  ( ! [X99,X100] : (sP5(X99,X100,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | p(s(bool,X100)) | ~sP4(X99,X100,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) <=> ~$spl278),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl278])])).
fof(f10659,plain,(
  ( ! [X99,X100] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP5(X99,X100,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | ~sP4(X99,X100,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | p(s(bool,X100))) ) | $spl4),
  inference(superposition,[],[f249,f394])).
fof(f249,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | ~sP4(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f154])).
fof(f10836,plain,(
  ( ! [X97,X98] : (sP5(X97,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X98) | ~p(s(bool,X98)) | ~sP4(X97,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X98)) ) | $spl276),
  inference(cnf_transformation,[],[f10836_D])).
fof(f10836_D,plain,(
  ( ! [X97,X98] : (sP5(X97,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X98) | ~p(s(bool,X98)) | ~sP4(X97,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X98)) ) <=> ~$spl276),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl276])])).
fof(f10658,plain,(
  ( ! [X97,X98] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP5(X97,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X98) | ~sP4(X97,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X98) | ~p(s(bool,X98))) ) | $spl4),
  inference(superposition,[],[f249,f394])).
fof(f10835,plain,(
  ( ! [X95,X96] : (sP5('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X95,X96) | p(s(bool,X95)) | ~sP4('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X95,X96)) ) | $spl274),
  inference(cnf_transformation,[],[f10835_D])).
fof(f10835_D,plain,(
  ( ! [X95,X96] : (sP5('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X95,X96) | p(s(bool,X95)) | ~sP4('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X95,X96)) ) <=> ~$spl274),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl274])])).
fof(f10657,plain,(
  ( ! [X95,X96] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP5('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X95,X96) | ~sP4('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X95,X96) | p(s(bool,X95))) ) | $spl4),
  inference(superposition,[],[f248,f394])).
fof(f248,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | ~sP4(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f154])).
fof(f10834,plain,(
  ( ! [X90,X89] : (sP5('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X89,X90) | ~p(s(bool,X90)) | ~sP4('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X89,X90)) ) | $spl272),
  inference(cnf_transformation,[],[f10834_D])).
fof(f10834_D,plain,(
  ( ! [X90,X89] : (sP5('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X89,X90) | ~p(s(bool,X90)) | ~sP4('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X89,X90)) ) <=> ~$spl272),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl272])])).
fof(f10654,plain,(
  ( ! [X90,X89] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP5('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X89,X90) | ~sP4('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X89,X90) | ~p(s(bool,X90))) ) | $spl4),
  inference(superposition,[],[f247,f394])).
fof(f247,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | ~sP4(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f154])).
fof(f10653,plain,(
  ( ! [X88,X87] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | p(s(bool,X87)) | ~sP5(X87,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X88)) ) | $spl4),
  inference(superposition,[],[f244,f394])).
fof(f244,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | ~sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f154])).
fof(f10652,plain,(
  ( ! [X85,X86] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | ~p(s(bool,X85)) | ~sP5('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X85,X86)) ) | $spl4),
  inference(superposition,[],[f244,f394])).
fof(f10651,plain,(
  ( ! [X83,X84] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | p(s(bool,X83)) | ~sP5('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X84,X83)) ) | $spl4),
  inference(superposition,[],[f243,f394])).
fof(f243,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | p(s(bool,X0)) | ~sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f154])).
fof(f10650,plain,(
  ( ! [X81,X82] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | p(s(bool,X81)) | ~sP5(X81,X82,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl4),
  inference(superposition,[],[f243,f394])).
fof(f3291,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | $spl132),
  inference(cnf_transformation,[],[f3291_D])).
fof(f3291_D,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f10649,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) | s(bool,'T_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(superposition,[],[f236,f394])).
fof(f10806,plain,(
  ( ! [X72,X71] : (s(bool,X71) != s(bool,X72) | ~sP3('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X72,X71)) ) | $spl262),
  inference(cnf_transformation,[],[f10806_D])).
fof(f10806_D,plain,(
  ( ! [X72,X71] : (s(bool,X71) != s(bool,X72) | ~sP3('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X72,X71)) ) <=> ~$spl262),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl262])])).
fof(f10623,plain,(
  ( ! [X72,X71] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | s(bool,X71) != s(bool,X72) | ~sP3('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X72,X71)) ) | $spl4),
  inference(superposition,[],[f222,f394])).
fof(f10805,plain,(
  ( ! [X66,X65] : (s(bool,X65) = s(bool,X66) | ~sP3('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X66,X65)) ) | $spl260),
  inference(cnf_transformation,[],[f10805_D])).
fof(f10805_D,plain,(
  ( ! [X66,X65] : (s(bool,X65) = s(bool,X66) | ~sP3('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X66,X65)) ) <=> ~$spl260),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl260])])).
fof(f10620,plain,(
  ( ! [X66,X65] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | s(bool,X65) = s(bool,X66) | ~sP3('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X66,X65)) ) | $spl4),
  inference(superposition,[],[f221,f394])).
fof(f10804,plain,(
  ( ! [X59,X60] : (sP3('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X59,X60) | s(bool,X59) != s(bool,X60)) ) | $spl258),
  inference(cnf_transformation,[],[f10804_D])).
fof(f10804_D,plain,(
  ( ! [X59,X60] : (sP3('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X59,X60) | s(bool,X59) != s(bool,X60)) ) <=> ~$spl258),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl258])])).
fof(f10617,plain,(
  ( ! [X59,X60] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP3('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X59,X60) | s(bool,X59) != s(bool,X60)) ) | $spl4),
  inference(superposition,[],[f220,f394])).
fof(f220,plain,(
  ( ! [X2,X0,X1] : (sP3(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f145])).
fof(f10803,plain,(
  ( ! [X54,X53] : (sP3('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X53,X54) | s(bool,X53) = s(bool,X54)) ) | $spl256),
  inference(cnf_transformation,[],[f10803_D])).
fof(f10803_D,plain,(
  ( ! [X54,X53] : (sP3('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X53,X54) | s(bool,X53) = s(bool,X54)) ) <=> ~$spl256),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl256])])).
fof(f10614,plain,(
  ( ! [X54,X53] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP3('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X53,X54) | s(bool,X53) = s(bool,X54)) ) | $spl4),
  inference(superposition,[],[f219,f394])).
fof(f7626,plain,(
  ( ! [X12,X13] : (sP0('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X12,X13)) ) | $spl178),
  inference(cnf_transformation,[],[f7626_D])).
fof(f7626_D,plain,(
  ( ! [X12,X13] : (sP0('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X12,X13)) ) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f10611,plain,(
  ( ! [X47,X48] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP0('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X47,X48)) ) | $spl4),
  inference(superposition,[],[f218,f394])).
fof(f7598,plain,(
  ( ! [X12,X13] : (sP0(X12,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X13)) ) | $spl176),
  inference(cnf_transformation,[],[f7598_D])).
fof(f7598_D,plain,(
  ( ! [X12,X13] : (sP0(X12,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X13)) ) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f10610,plain,(
  ( ! [X45,X46] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP0(X45,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X46)) ) | $spl4),
  inference(superposition,[],[f217,f394])).
fof(f7569,plain,(
  ( ! [X12,X13] : (sP0(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl174),
  inference(cnf_transformation,[],[f7569_D])).
fof(f7569_D,plain,(
  ( ! [X12,X13] : (sP0(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f10609,plain,(
  ( ! [X43,X44] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP0(X43,X44,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl4),
  inference(superposition,[],[f216,f394])).
fof(f10802,plain,(
  ( ! [X41,X42] : (p(s(bool,X41)) | ~sP0('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X41,X42) | p(s(bool,X42))) ) | $spl254),
  inference(cnf_transformation,[],[f10802_D])).
fof(f10802_D,plain,(
  ( ! [X41,X42] : (p(s(bool,X41)) | ~sP0('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X41,X42) | p(s(bool,X42))) ) <=> ~$spl254),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl254])])).
fof(f10608,plain,(
  ( ! [X41,X42] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | p(s(bool,X41)) | p(s(bool,X42)) | ~sP0('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X41,X42)) ) | $spl4),
  inference(superposition,[],[f215,f394])).
fof(f215,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP0(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f144])).
fof(f10801,plain,(
  ( ! [X39,X40] : (p(s(bool,X39)) | ~sP0(X39,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X40) | p(s(bool,X40))) ) | $spl252),
  inference(cnf_transformation,[],[f10801_D])).
fof(f10801_D,plain,(
  ( ! [X39,X40] : (p(s(bool,X39)) | ~sP0(X39,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X40) | p(s(bool,X40))) ) <=> ~$spl252),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl252])])).
fof(f10607,plain,(
  ( ! [X39,X40] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | p(s(bool,X39)) | p(s(bool,X40)) | ~sP0(X39,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X40)) ) | $spl4),
  inference(superposition,[],[f215,f394])).
fof(f10800,plain,(
  ( ! [X37,X38] : (p(s(bool,X37)) | ~sP0(X37,X38,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | p(s(bool,X38))) ) | $spl250),
  inference(cnf_transformation,[],[f10800_D])).
fof(f10800_D,plain,(
  ( ! [X37,X38] : (p(s(bool,X37)) | ~sP0(X37,X38,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | p(s(bool,X38))) ) <=> ~$spl250),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl250])])).
fof(f10606,plain,(
  ( ! [X37,X38] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | p(s(bool,X37)) | p(s(bool,X38)) | ~sP0(X37,X38,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl4),
  inference(superposition,[],[f215,f394])).
fof(f10799,plain,(
  ( ! [X35,X36] : (sP1('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X35,X36)) ) | $spl248),
  inference(cnf_transformation,[],[f10799_D])).
fof(f10799_D,plain,(
  ( ! [X35,X36] : (sP1('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X35,X36)) ) <=> ~$spl248),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl248])])).
fof(f10605,plain,(
  ( ! [X35,X36] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP1('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X35,X36)) ) | $spl4),
  inference(superposition,[],[f214,f394])).
fof(f10798,plain,(
  ( ! [X33,X34] : (sP1(X33,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X34)) ) | $spl246),
  inference(cnf_transformation,[],[f10798_D])).
fof(f10798_D,plain,(
  ( ! [X33,X34] : (sP1(X33,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X34)) ) <=> ~$spl246),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl246])])).
fof(f10604,plain,(
  ( ! [X33,X34] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP1(X33,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X34)) ) | $spl4),
  inference(superposition,[],[f213,f394])).
fof(f7503,plain,(
  ( ! [X12,X13] : (sP1(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl170),
  inference(cnf_transformation,[],[f7503_D])).
fof(f7503_D,plain,(
  ( ! [X12,X13] : (sP1(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f10603,plain,(
  ( ! [X31,X32] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP1(X31,X32,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl4),
  inference(superposition,[],[f212,f394])).
fof(f10797,plain,(
  ( ! [X30,X29] : (~p(s(bool,X29)) | ~sP1(X29,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X30) | p(s(bool,X30))) ) | $spl244),
  inference(cnf_transformation,[],[f10797_D])).
fof(f10797_D,plain,(
  ( ! [X30,X29] : (~p(s(bool,X29)) | ~sP1(X29,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X30) | p(s(bool,X30))) ) <=> ~$spl244),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl244])])).
fof(f10602,plain,(
  ( ! [X30,X29] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | ~p(s(bool,X29)) | p(s(bool,X30)) | ~sP1(X29,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X30)) ) | $spl4),
  inference(superposition,[],[f211,f394])).
fof(f10796,plain,(
  ( ! [X28,X27] : (~p(s(bool,X27)) | ~sP1('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X27,X28) | p(s(bool,X28))) ) | $spl242),
  inference(cnf_transformation,[],[f10796_D])).
fof(f10796_D,plain,(
  ( ! [X28,X27] : (~p(s(bool,X27)) | ~sP1('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X27,X28) | p(s(bool,X28))) ) <=> ~$spl242),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl242])])).
fof(f10601,plain,(
  ( ! [X28,X27] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | ~p(s(bool,X27)) | p(s(bool,X28)) | ~sP1('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X27,X28)) ) | $spl4),
  inference(superposition,[],[f211,f394])).
fof(f10795,plain,(
  ( ! [X26,X25] : (~p(s(bool,X25)) | ~sP1(X25,X26,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | ~p(s(bool,X26))) ) | $spl240),
  inference(cnf_transformation,[],[f10795_D])).
fof(f10795_D,plain,(
  ( ! [X26,X25] : (~p(s(bool,X25)) | ~sP1(X25,X26,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | ~p(s(bool,X26))) ) <=> ~$spl240),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl240])])).
fof(f10600,plain,(
  ( ! [X26,X25] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | ~p(s(bool,X25)) | ~p(s(bool,X26)) | ~sP1(X25,X26,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl4),
  inference(superposition,[],[f211,f394])).
fof(f10794,plain,(
  ( ! [X24,X23] : (sP2('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X23,X24)) ) | $spl238),
  inference(cnf_transformation,[],[f10794_D])).
fof(f10794_D,plain,(
  ( ! [X24,X23] : (sP2('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X23,X24)) ) <=> ~$spl238),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl238])])).
fof(f10599,plain,(
  ( ! [X24,X23] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP2('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X23,X24)) ) | $spl4),
  inference(superposition,[],[f210,f394])).
fof(f10793,plain,(
  ( ! [X21,X22] : (sP2(X21,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X22)) ) | $spl236),
  inference(cnf_transformation,[],[f10793_D])).
fof(f10793_D,plain,(
  ( ! [X21,X22] : (sP2(X21,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X22)) ) <=> ~$spl236),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl236])])).
fof(f10598,plain,(
  ( ! [X21,X22] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP2(X21,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X22)) ) | $spl4),
  inference(superposition,[],[f209,f394])).
fof(f7474,plain,(
  ( ! [X12,X13] : (sP2(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl168),
  inference(cnf_transformation,[],[f7474_D])).
fof(f7474_D,plain,(
  ( ! [X12,X13] : (sP2(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f10597,plain,(
  ( ! [X19,X20] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP2(X19,X20,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl4),
  inference(superposition,[],[f208,f394])).
fof(f10792,plain,(
  ( ! [X17,X18] : (~p(s(bool,X17)) | ~sP2(X17,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X18) | p(s(bool,X18))) ) | $spl234),
  inference(cnf_transformation,[],[f10792_D])).
fof(f10792_D,plain,(
  ( ! [X17,X18] : (~p(s(bool,X17)) | ~sP2(X17,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X18) | p(s(bool,X18))) ) <=> ~$spl234),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl234])])).
fof(f10596,plain,(
  ( ! [X17,X18] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | ~p(s(bool,X17)) | p(s(bool,X18)) | ~sP2(X17,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X18)) ) | $spl4),
  inference(superposition,[],[f207,f394])).
fof(f10791,plain,(
  ( ! [X15,X16] : (~p(s(bool,X15)) | ~sP2('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X15,X16) | p(s(bool,X16))) ) | $spl232),
  inference(cnf_transformation,[],[f10791_D])).
fof(f10791_D,plain,(
  ( ! [X15,X16] : (~p(s(bool,X15)) | ~sP2('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X15,X16) | p(s(bool,X16))) ) <=> ~$spl232),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl232])])).
fof(f10595,plain,(
  ( ! [X15,X16] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | ~p(s(bool,X15)) | p(s(bool,X16)) | ~sP2('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X15,X16)) ) | $spl4),
  inference(superposition,[],[f207,f394])).
fof(f10790,plain,(
  ( ! [X14,X13] : (~p(s(bool,X13)) | ~sP2(X13,X14,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | ~p(s(bool,X14))) ) | $spl230),
  inference(cnf_transformation,[],[f10790_D])).
fof(f10790_D,plain,(
  ( ! [X14,X13] : (~p(s(bool,X13)) | ~sP2(X13,X14,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | ~p(s(bool,X14))) ) <=> ~$spl230),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl230])])).
fof(f10594,plain,(
  ( ! [X14,X13] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | ~p(s(bool,X13)) | ~p(s(bool,X14)) | ~sP2(X13,X14,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl4),
  inference(superposition,[],[f207,f394])).
fof(f10789,plain,(
  ( ! [X12,X11] : (sP3('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X11,X12) | ~sP0(X12,X11,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | ~sP1(X11,X12,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | ~sP2('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X12,X11)) ) | $spl228),
  inference(cnf_transformation,[],[f10789_D])).
fof(f10789_D,plain,(
  ( ! [X12,X11] : (sP3('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X11,X12) | ~sP0(X12,X11,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | ~sP1(X11,X12,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | ~sP2('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X12,X11)) ) <=> ~$spl228),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl228])])).
fof(f10593,plain,(
  ( ! [X12,X11] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP3('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X11,X12) | ~sP2('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X12,X11) | ~sP1(X11,X12,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | ~sP0(X12,X11,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl4),
  inference(superposition,[],[f206,f394])).
fof(f206,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X0)) | ~sP2(X0,X2,X1) | ~sP1(X1,X2,X0) | ~sP0(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f136])).
fof(f10788,plain,(
  ( ! [X10,X9] : (sP3(X9,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X10) | ~sP0(X10,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X9) | ~sP1('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X10,X9) | ~sP2(X9,X10,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl226),
  inference(cnf_transformation,[],[f10788_D])).
fof(f10788_D,plain,(
  ( ! [X10,X9] : (sP3(X9,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X10) | ~sP0(X10,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X9) | ~sP1('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X10,X9) | ~sP2(X9,X10,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) <=> ~$spl226),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl226])])).
fof(f10592,plain,(
  ( ! [X10,X9] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP3(X9,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X10) | ~sP2(X9,X10,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | ~sP1('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X10,X9) | ~sP0(X10,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X9)) ) | $spl4),
  inference(superposition,[],[f205,f394])).
fof(f10787,plain,(
  ( ! [X8,X7] : (sP3(X7,X8,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | ~sP0('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X8,X7) | ~sP1(X8,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X7) | ~sP2(X7,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X8)) ) | $spl224),
  inference(cnf_transformation,[],[f10787_D])).
fof(f10787_D,plain,(
  ( ! [X8,X7] : (sP3(X7,X8,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | ~sP0('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X8,X7) | ~sP1(X8,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X7) | ~sP2(X7,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X8)) ) <=> ~$spl224),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl224])])).
fof(f10591,plain,(
  ( ! [X8,X7] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | sP3(X7,X8,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | ~sP2(X7,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X8) | ~sP1(X8,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X7) | ~sP0('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X8,X7)) ) | $spl4),
  inference(superposition,[],[f204,f394])).
fof(f10786,plain,(
  ( ! [X6,X5] : (~p(s(bool,X5)) | ~sP3(X5,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X6) | p(s(bool,X6))) ) | $spl222),
  inference(cnf_transformation,[],[f10786_D])).
fof(f10786_D,plain,(
  ( ! [X6,X5] : (~p(s(bool,X5)) | ~sP3(X5,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X6) | p(s(bool,X6))) ) <=> ~$spl222),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl222])])).
fof(f10590,plain,(
  ( ! [X6,X5] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | ~p(s(bool,X5)) | p(s(bool,X6)) | ~sP3(X5,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X6)) ) | $spl4),
  inference(superposition,[],[f203,f394])).
fof(f203,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f136])).
fof(f10785,plain,(
  ( ! [X4,X3] : (~p(s(bool,X3)) | ~sP3('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X3,X4) | p(s(bool,X4))) ) | $spl220),
  inference(cnf_transformation,[],[f10785_D])).
fof(f10785_D,plain,(
  ( ! [X4,X3] : (~p(s(bool,X3)) | ~sP3('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X3,X4) | p(s(bool,X4))) ) <=> ~$spl220),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl220])])).
fof(f10589,plain,(
  ( ! [X4,X3] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | ~p(s(bool,X3)) | p(s(bool,X4)) | ~sP3('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X3,X4)) ) | $spl4),
  inference(superposition,[],[f203,f394])).
fof(f10784,plain,(
  ( ! [X2,X1] : (~p(s(bool,X1)) | ~sP3(X1,X2,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | ~p(s(bool,X2))) ) | $spl218),
  inference(cnf_transformation,[],[f10784_D])).
fof(f10784_D,plain,(
  ( ! [X2,X1] : (~p(s(bool,X1)) | ~sP3(X1,X2,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | ~p(s(bool,X2))) ) <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f10588,plain,(
  ( ! [X2,X1] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP3(X1,X2,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl4),
  inference(superposition,[],[f203,f394])).
fof(f1554,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) | $spl79),
  inference(cnf_transformation,[],[f1554_D])).
fof(f1554_D,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) <=> ~$spl79),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl79])])).
fof(f10581,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl4 | $spl30)),
  inference(superposition,[],[f5895,f394])).
fof(f10579,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) | s(bool,'T_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(superposition,[],[f394,f236])).
fof(f549,plain,(
  p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) | $spl44),
  inference(cnf_transformation,[],[f549_D])).
fof(f549_D,plain,(
  p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f3286,plain,(
  ( ! [X8] : (s(bool,X8) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) | p(s(bool,X8))) ) | $spl126),
  inference(cnf_transformation,[],[f3286_D])).
fof(f3286_D,plain,(
  ( ! [X8] : (s(bool,X8) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) | p(s(bool,X8))) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f7808,plain,(
  ( ! [X8] : (s(bool,X8) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) | p(s(bool,X8)) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))))) ) | $spl5),
  inference(superposition,[],[f395,f223])).
fof(f7809,plain,(
  ( ! [X9] : (s(bool,X9) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) | p(s(bool,X9))) ) | $spl5),
  inference(superposition,[],[f395,f223])).
fof(f550,plain,(
  ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) | $spl45),
  inference(cnf_transformation,[],[f550_D])).
fof(f550_D,plain,(
  ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) <=> ~$spl45),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl45])])).
fof(f3287,plain,(
  ( ! [X10] : (s(bool,X10) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) | ~p(s(bool,X10))) ) | $spl128),
  inference(cnf_transformation,[],[f3287_D])).
fof(f3287_D,plain,(
  ( ! [X10] : (s(bool,X10) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) | ~p(s(bool,X10))) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f7810,plain,(
  ( ! [X10] : (s(bool,X10) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) | ~p(s(bool,X10))) ) | $spl5),
  inference(superposition,[],[f395,f226])).
fof(f7811,plain,(
  ( ! [X11] : (s(bool,X11) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) | ~p(s(bool,X11)) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))))) ) | $spl5),
  inference(superposition,[],[f395,f226])).
fof(f7812,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) | $spl5),
  inference(superposition,[],[f395,f229])).
fof(f7813,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) | $spl5),
  inference(superposition,[],[f395,f231])).
fof(f7814,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) | $spl5),
  inference(superposition,[],[f395,f233])).
fof(f7815,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) | $spl5),
  inference(superposition,[],[f395,f235])).
fof(f10387,plain,(
  ( ! [X8] : (p(s(bool,X8)) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))))) ) | $spl131),
  inference(subsumption_resolution,[],[f10374,f229])).
fof(f10374,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,X8) | p(s(bool,X8)) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))))) ) | $spl131),
  inference(superposition,[],[f3288,f223])).
fof(f10389,plain,(
  ( ! [X9] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) | p(s(bool,X9))) ) | $spl131),
  inference(subsumption_resolution,[],[f10375,f229])).
fof(f10375,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,X9) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) | p(s(bool,X9))) ) | $spl131),
  inference(superposition,[],[f3288,f223])).
fof(f10392,plain,(
  $false | ($spl45 | $spl131)),
  inference(subsumption_resolution,[],[f10385,f550])).
fof(f10385,plain,(
  p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) | $spl131),
  inference(trivial_inequality_removal,[],[f10379])).
fof(f10379,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) | $spl131),
  inference(superposition,[],[f3288,f231])).
fof(f10391,plain,(
  $false | ($spl45 | $spl131)),
  inference(subsumption_resolution,[],[f10386,f550])).
fof(f10386,plain,(
  p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) | $spl131),
  inference(trivial_inequality_removal,[],[f10378])).
fof(f10378,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) | $spl131),
  inference(superposition,[],[f3288,f229])).
fof(f10347,plain,(
  ( ! [X19,X17] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X19)))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X17),s('type.num.num',sK13),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X19),s('type.num.num',sK14),s('type.num.num',X17))))) ) | $spl214),
  inference(cnf_transformation,[],[f10347_D])).
fof(f10347_D,plain,(
  ( ! [X19,X17] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X19)))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X17),s('type.num.num',sK13),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X19),s('type.num.num',sK14),s('type.num.num',X17))))) ) <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f10349,plain,(
  ( ! [X21,X22,X20] : (s(bool,'F_0') != s(bool,X21) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X22)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X22),s('type.num.num',sK14),s('type.num.num',X20)))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X20),s('type.num.num',sK13),s('type.num.num',sK12))))) ) | $spl30),
  inference(subsumption_resolution,[],[f10329,f228])).
fof(f10329,plain,(
  ( ! [X21,X22,X20] : (s(bool,'F_0') != s(bool,X21) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X22)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X22),s('type.num.num',sK14),s('type.num.num',X20)))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X20),s('type.num.num',sK13),s('type.num.num',sK12)))) | p(s(bool,X21))) ) | $spl30),
  inference(superposition,[],[f5895,f223])).
fof(f10346,plain,(
  ( ! [X19,X17,X18] : (s(bool,'F_0') != s(bool,X18) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X19)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X19),s('type.num.num',sK14),s('type.num.num',X17)))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X17),s('type.num.num',sK13),s('type.num.num',sK12))))) ) | $spl30),
  inference(subsumption_resolution,[],[f10328,f228])).
fof(f10328,plain,(
  ( ! [X19,X17,X18] : (s(bool,'F_0') != s(bool,X18) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X19)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X19),s('type.num.num',sK14),s('type.num.num',X17)))) | p(s(bool,X18)) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X17),s('type.num.num',sK13),s('type.num.num',sK12))))) ) | $spl30),
  inference(superposition,[],[f5895,f223])).
fof(f10306,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X28),s('type.num.num',sK13),s('type.num.num',sK12))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X29)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X29),s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X30),s('type.num.num',X28))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X31),s('type.num.num',X28))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X28),s('type.num.num',X31),s('type.num.num',X30))))) ) | $spl30),
  inference(superposition,[],[f5895,f239])).
fof(f10305,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X24),s('type.num.num',sK13),s('type.num.num',sK12))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X25)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X25),s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X24)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X24),s('type.num.num',X26),s('type.num.num',X27))))) ) | $spl30),
  inference(superposition,[],[f5895,f237])).
fof(f10342,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',X15),s('type.num.num',X14)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',X12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',X12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',sK14),s('type.num.num',X12))))) ) | $spl206),
  inference(cnf_transformation,[],[f10342_D])).
fof(f10342_D,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',X15),s('type.num.num',X14)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',X12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',X12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',sK14),s('type.num.num',X12))))) ) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f3288,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | $spl131),
  inference(cnf_transformation,[],[f3288_D])).
fof(f3288_D,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) <=> ~$spl131),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl131])])).
fof(f10302,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X13),s('type.num.num',sK14),s('type.num.num',X12)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',X12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',X12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl30),
  inference(superposition,[],[f5895,f239])).
fof(f10341,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num',X2),s('type.num.num',X3)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X1),s('type.num.num',sK14),s('type.num.num',X0))))) ) | $spl204),
  inference(cnf_transformation,[],[f10341_D])).
fof(f10341_D,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num',X2),s('type.num.num',X3)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X1),s('type.num.num',sK14),s('type.num.num',X0))))) ) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f10299,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X1),s('type.num.num',sK14),s('type.num.num',X0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl30),
  inference(superposition,[],[f5895,f237])).
fof(f8975,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) | s(bool,'T_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | $spl21),
  inference(superposition,[],[f407,f236])).
fof(f8974,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) | s(bool,'F_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | $spl21),
  inference(superposition,[],[f407,f236])).
fof(f396,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X7),s('type.num.num',X6)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK14)))) ) | $spl6),
  inference(cnf_transformation,[],[f396_D])).
fof(f396_D,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X7),s('type.num.num',X6)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK14)))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f3285,plain,(
  s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | $spl125),
  inference(cnf_transformation,[],[f3285_D])).
fof(f3285_D,plain,(
  s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) <=> ~$spl125),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl125])])).
fof(f8945,plain,(
  ( ! [X6,X7] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X7),s('type.num.num',X6))))) ) | $spl21),
  inference(superposition,[],[f407,f239])).
fof(f1780,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0'))) | $spl90),
  inference(cnf_transformation,[],[f1780_D])).
fof(f1780_D,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0'))) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f8629,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) | s(bool,'T_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0'))) | $spl12),
  inference(superposition,[],[f236,f400])).
fof(f908,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0'))) | $spl54),
  inference(cnf_transformation,[],[f908_D])).
fof(f908_D,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0'))) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f8045,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) | $spl198),
  inference(cnf_transformation,[],[f8045_D])).
fof(f8045_D,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f8628,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) | s(bool,'F_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0'))) | $spl12),
  inference(superposition,[],[f236,f400])).
fof(f8561,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) | s(bool,'T_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0'))) | $spl12),
  inference(superposition,[],[f400,f236])).
fof(f8560,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) | s(bool,'F_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0'))) | $spl12),
  inference(superposition,[],[f400,f236])).
fof(f3282,plain,(
  s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0'))) | $spl122),
  inference(cnf_transformation,[],[f3282_D])).
fof(f3282_D,plain,(
  s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0'))) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f8533,plain,(
  ( ! [X6,X7] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X7),s('type.num.num',X6))))) ) | $spl12),
  inference(superposition,[],[f400,f239])).
fof(f8036,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) | s(bool,'T_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0'))) | $spl13),
  inference(superposition,[],[f401,f236])).
fof(f8035,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) | s(bool,'F_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0'))) | $spl13),
  inference(superposition,[],[f401,f236])).
fof(f8039,plain,(
  ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')))) | $spl191),
  inference(cnf_transformation,[],[f8039_D])).
fof(f8039_D,plain,(
  ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')))) <=> ~$spl191),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl191])])).
fof(f8034,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')))) | $spl13),
  inference(superposition,[],[f401,f235])).
fof(f8033,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')))) | $spl13),
  inference(superposition,[],[f401,f233])).
fof(f8040,plain,(
  p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')))) | $spl190),
  inference(cnf_transformation,[],[f8040_D])).
fof(f8040_D,plain,(
  p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')))) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f8032,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')))) | $spl13),
  inference(superposition,[],[f401,f231])).
fof(f8031,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')))) | $spl13),
  inference(superposition,[],[f401,f229])).
fof(f8042,plain,(
  ( ! [X10] : (s(bool,X10) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) | ~p(s(bool,X10))) ) | $spl194),
  inference(cnf_transformation,[],[f8042_D])).
fof(f8042_D,plain,(
  ( ! [X10] : (s(bool,X10) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) | ~p(s(bool,X10))) ) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f8030,plain,(
  ( ! [X11] : (s(bool,X11) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) | ~p(s(bool,X11)) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0'))))) ) | $spl13),
  inference(superposition,[],[f401,f226])).
fof(f8029,plain,(
  ( ! [X10] : (s(bool,X10) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X10))) ) | $spl13),
  inference(superposition,[],[f401,f226])).
fof(f8041,plain,(
  ( ! [X8] : (s(bool,X8) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) | p(s(bool,X8))) ) | $spl192),
  inference(cnf_transformation,[],[f8041_D])).
fof(f8041_D,plain,(
  ( ! [X8] : (s(bool,X8) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) | p(s(bool,X8))) ) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f8028,plain,(
  ( ! [X9] : (s(bool,X9) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X9))) ) | $spl13),
  inference(superposition,[],[f401,f223])).
fof(f8027,plain,(
  ( ! [X8] : (s(bool,X8) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) | p(s(bool,X8)) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0'))))) ) | $spl13),
  inference(superposition,[],[f401,f223])).
fof(f3283,plain,(
  s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0'))) | $spl123),
  inference(cnf_transformation,[],[f3283_D])).
fof(f3283_D,plain,(
  s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0'))) <=> ~$spl123),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl123])])).
fof(f8006,plain,(
  ( ! [X6,X7] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X7),s('type.num.num',X6))))) ) | $spl13),
  inference(superposition,[],[f401,f239])).
fof(f7817,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) | s(bool,'T_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | $spl5),
  inference(superposition,[],[f395,f236])).
fof(f3289,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | $spl130),
  inference(cnf_transformation,[],[f3289_D])).
fof(f3289_D,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f7816,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) | s(bool,'F_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | $spl5),
  inference(superposition,[],[f395,f236])).
fof(f7757,plain,(
  ( ! [X14,X12,X15,X13] : (sP8(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl24),
  inference(superposition,[],[f1151,f239])).
fof(f1151,plain,(
  ( ! [X105,X104] : (sP8(X104,X105,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl24),
  inference(resolution,[],[f410,f285])).
fof(f7729,plain,(
  ( ! [X14,X12,X15,X13] : (sP6(X12,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X13) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl24),
  inference(superposition,[],[f1135,f239])).
fof(f1135,plain,(
  ( ! [X72,X73] : (sP6(X72,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',sK12)),X73)) ) | $spl24),
  inference(resolution,[],[f410,f269])).
fof(f7720,plain,(
  ( ! [X12,X13] : (~sP3(X12,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X13) | s(bool,'F_0') != s(bool,X13) | ~p(s(bool,X12))) ) | $spl184),
  inference(cnf_transformation,[],[f7720_D])).
fof(f7720_D,plain,(
  ( ! [X12,X13] : (~sP3(X12,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X13) | s(bool,'F_0') != s(bool,X13) | ~p(s(bool,X12))) ) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f7696,plain,(
  ( ! [X14,X12,X15,X13] : (~sP3(X12,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X13) | ~p(s(bool,X12)) | s(bool,'F_0') != s(bool,X13) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl30),
  inference(superposition,[],[f5915,f239])).
fof(f6892,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | s(bool,'T_0') != s(bool,X2)) ) | $spl160),
  inference(cnf_transformation,[],[f6892_D])).
fof(f6892_D,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | s(bool,'T_0') != s(bool,X2)) ) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f7714,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X5) | ~p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5)) ) | ($spl30 | $spl36)),
  inference(forward_demodulation,[],[f7712,f420])).
fof(f7712,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5) | s(bool,X5) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl30),
  inference(duplicate_literal_removal,[],[f7688])).
fof(f7688,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5) | s(bool,X5) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',sK12))) | ~p(s(bool,X4))) ) | $spl30),
  inference(resolution,[],[f5915,f220])).
fof(f7666,plain,(
  ( ! [X14,X12,X15,X13] : (sP6(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl24),
  inference(superposition,[],[f1134,f239])).
fof(f1134,plain,(
  ( ! [X70,X71] : (sP6(X70,X71,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl24),
  inference(resolution,[],[f410,f268])).
fof(f7635,plain,(
  ( ! [X14,X12,X15,X13] : (sP4(X12,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X13) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl24),
  inference(superposition,[],[f1121,f239])).
fof(f1121,plain,(
  ( ! [X45,X44] : (sP4(X44,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',sK12)),X45)) ) | $spl24),
  inference(resolution,[],[f410,f252])).
fof(f7606,plain,(
  ( ! [X14,X12,X15,X13] : (sP0('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X12,X13) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl24),
  inference(superposition,[],[f1105,f239])).
fof(f1105,plain,(
  ( ! [X23,X22] : (sP0('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',sK12)),X22,X23)) ) | $spl24),
  inference(resolution,[],[f410,f218])).
fof(f7577,plain,(
  ( ! [X14,X12,X15,X13] : (sP0(X12,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X13) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl24),
  inference(superposition,[],[f1104,f239])).
fof(f1104,plain,(
  ( ! [X21,X20] : (sP0(X20,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',sK12)),X21)) ) | $spl24),
  inference(resolution,[],[f410,f217])).
fof(f7548,plain,(
  ( ! [X14,X12,X15,X13] : (sP0(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl24),
  inference(superposition,[],[f1103,f239])).
fof(f1103,plain,(
  ( ! [X19,X18] : (sP0(X18,X19,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl24),
  inference(resolution,[],[f410,f216])).
fof(f7540,plain,(
  ( ! [X12,X13] : (~sP3(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | s(bool,'F_0') != s(bool,X13) | ~p(s(bool,X12))) ) | $spl172),
  inference(cnf_transformation,[],[f7540_D])).
fof(f7540_D,plain,(
  ( ! [X12,X13] : (~sP3(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | s(bool,'F_0') != s(bool,X13) | ~p(s(bool,X12))) ) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f7514,plain,(
  ( ! [X14,X12,X15,X13] : (~sP3(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | ~p(s(bool,X12)) | s(bool,'F_0') != s(bool,X13) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl30),
  inference(superposition,[],[f5914,f239])).
fof(f7532,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X5) | ~p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5)) ) | ($spl30 | $spl36)),
  inference(forward_demodulation,[],[f7530,f420])).
fof(f7530,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5) | s(bool,X5) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl30),
  inference(duplicate_literal_removal,[],[f7506])).
fof(f7506,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5) | s(bool,X5) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',sK12))) | ~p(s(bool,X4))) ) | $spl30),
  inference(resolution,[],[f5914,f220])).
fof(f7482,plain,(
  ( ! [X14,X12,X15,X13] : (sP1(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl24),
  inference(superposition,[],[f1102,f239])).
fof(f1102,plain,(
  ( ! [X17,X16] : (sP1(X16,X17,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl24),
  inference(resolution,[],[f410,f212])).
fof(f7453,plain,(
  ( ! [X14,X12,X15,X13] : (sP2(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl24),
  inference(superposition,[],[f1099,f239])).
fof(f1099,plain,(
  ( ! [X10,X11] : (sP2(X10,X11,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl24),
  inference(resolution,[],[f410,f208])).
fof(f3833,plain,(
  ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')))) | $spl143),
  inference(cnf_transformation,[],[f3833_D])).
fof(f3833_D,plain,(
  ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')))) <=> ~$spl143),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl143])])).
fof(f7403,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X7),s('type.num.num',X6))))) ) | $spl49),
  inference(superposition,[],[f554,f239])).
fof(f7310,plain,(
  ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | $spl167),
  inference(cnf_transformation,[],[f7310_D])).
fof(f7310_D,plain,(
  ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) <=> ~$spl167),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl167])])).
fof(f7271,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X7),s('type.num.num',X6))))) ) | $spl47),
  inference(superposition,[],[f552,f239])).
fof(f7053,plain,(
  ( ! [X12,X13] : (sP3(X12,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X13) | p(s(bool,X12)) | s(bool,'F_0') != s(bool,X13)) ) | $spl164),
  inference(cnf_transformation,[],[f7053_D])).
fof(f7053_D,plain,(
  ( ! [X12,X13] : (sP3(X12,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X13) | p(s(bool,X12)) | s(bool,'F_0') != s(bool,X13)) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f7034,plain,(
  ( ! [X14,X12,X15,X13] : (sP3(X12,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X13) | s(bool,'F_0') != s(bool,X13) | p(s(bool,X12)) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl30),
  inference(superposition,[],[f5913,f239])).
fof(f7050,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | s(bool,'F_0') != s(bool,X2) | p(s(bool,X3))) ) | ($spl30 | $spl36)),
  inference(forward_demodulation,[],[f7049,f420])).
fof(f7049,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,X2) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl30),
  inference(duplicate_literal_removal,[],[f7026])).
fof(f7026,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,X2) | p(s(bool,X3)) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl30),
  inference(resolution,[],[f5913,f222])).
fof(f7019,plain,(
  $false | ($spl41 | $spl44 | $spl57)),
  inference(subsumption_resolution,[],[f7018,f542])).
fof(f7018,plain,(
  s('type.num.num',sK12) = s('type.num.num',sK13) | ($spl44 | $spl57)),
  inference(subsumption_resolution,[],[f6897,f1051])).
fof(f6897,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | s('type.num.num',sK12) = s('type.num.num',sK13) | $spl44),
  inference(resolution,[],[f549,f238])).
fof(f6895,plain,(
  ( ! [X12,X13] : (sP3(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | p(s(bool,X12)) | s(bool,'F_0') != s(bool,X13)) ) | $spl162),
  inference(cnf_transformation,[],[f6895_D])).
fof(f6895_D,plain,(
  ( ! [X12,X13] : (sP3(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | p(s(bool,X12)) | s(bool,'F_0') != s(bool,X13)) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f6875,plain,(
  ( ! [X14,X12,X15,X13] : (sP3(X12,X13,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | s(bool,'F_0') != s(bool,X13) | p(s(bool,X12)) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl30),
  inference(superposition,[],[f5912,f239])).
fof(f6891,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | s(bool,'F_0') != s(bool,X2) | p(s(bool,X3))) ) | ($spl30 | $spl36)),
  inference(forward_demodulation,[],[f6890,f420])).
fof(f6890,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,X2) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl30),
  inference(duplicate_literal_removal,[],[f6867])).
fof(f6867,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,X2) | p(s(bool,X3)) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl30),
  inference(resolution,[],[f5912,f222])).
fof(f6865,plain,(
  ( ! [X9] : (sP11('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK12)),X9)) ) | $spl158),
  inference(cnf_transformation,[],[f6865_D])).
fof(f6865_D,plain,(
  ( ! [X9] : (sP11('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK12)),X9)) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f6856,plain,(
  ( ! [X10,X11,X9] : (sP11('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK12)),X9) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X10),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X11),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X11),s('type.num.num',X10))))) ) | $spl98),
  inference(superposition,[],[f5502,f239])).
fof(f6850,plain,(
  ( ! [X9] : (sP11(X9,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK12)))) ) | $spl156),
  inference(cnf_transformation,[],[f6850_D])).
fof(f6850_D,plain,(
  ( ! [X9] : (sP11(X9,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK12)))) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f6841,plain,(
  ( ! [X10,X11,X9] : (sP11(X9,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK12))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X10),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X11),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X11),s('type.num.num',X10))))) ) | $spl98),
  inference(superposition,[],[f5501,f239])).
fof(f6832,plain,(
  ( ! [X19,X18] : (~p(s(bool,X19)) | ~p(s(bool,X18)) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))))) ) | $spl128),
  inference(subsumption_resolution,[],[f6812,f226])).
fof(f6812,plain,(
  ( ! [X19,X18] : (s(bool,X18) != s(bool,X19) | ~p(s(bool,X19)) | ~p(s(bool,X18)) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))))) ) | $spl128),
  inference(superposition,[],[f3287,f226])).
fof(f6831,plain,(
  ( ! [X17,X16] : (~p(s(bool,X17)) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | ~p(s(bool,X16))) ) | $spl128),
  inference(subsumption_resolution,[],[f6811,f226])).
fof(f6811,plain,(
  ( ! [X17,X16] : (s(bool,X16) != s(bool,X17) | ~p(s(bool,X17)) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | ~p(s(bool,X16))) ) | $spl128),
  inference(superposition,[],[f3287,f226])).
fof(f6810,plain,(
  ( ! [X14,X15] : (s(bool,X14) != s(bool,X15) | ~p(s(bool,X15)) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13)))) | p(s(bool,X14))) ) | $spl128),
  inference(superposition,[],[f3287,f223])).
fof(f6809,plain,(
  ( ! [X12,X13] : (s(bool,X12) != s(bool,X13) | ~p(s(bool,X13)) | p(s(bool,X12)) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))))) ) | $spl128),
  inference(superposition,[],[f3287,f223])).
fof(f6824,plain,(
  ( ! [X24] : (s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) | s(bool,'F_0') = s(bool,X24)) ) | $spl128),
  inference(subsumption_resolution,[],[f6781,f229])).
fof(f6781,plain,(
  ( ! [X24] : (s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) | ~p(s(bool,X24)) | s(bool,'F_0') = s(bool,X24)) ) | $spl128),
  inference(superposition,[],[f3287,f236])).
fof(f6748,plain,(
  ( ! [X10,X11,X9] : (sP11('const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)),X9) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X10),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X11),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X11),s('type.num.num',X10))))) ) | $spl24),
  inference(superposition,[],[f1168,f239])).
fof(f1168,plain,(
  ( ! [X129] : (sP11('const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',sK12)),X129)) ) | $spl24),
  inference(resolution,[],[f410,f312])).
fof(f6723,plain,(
  ( ! [X10,X11,X9] : (sP11(X9,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X10),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X11),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X11),s('type.num.num',X10))))) ) | $spl24),
  inference(superposition,[],[f1167,f239])).
fof(f1167,plain,(
  ( ! [X128] : (sP11(X128,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl24),
  inference(resolution,[],[f410,f311])).
fof(f6681,plain,(
  $false | ($spl26 | $spl30)),
  inference(trivial_inequality_removal,[],[f6678])).
fof(f6678,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl26 | $spl30)),
  inference(superposition,[],[f6643,f413])).
fof(f6643,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,X8)) ) | ($spl26 | $spl30)),
  inference(subsumption_resolution,[],[f6642,f230])).
fof(f230,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f147])).
fof(f6642,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,X8) | p(s(bool,X8))) ) | ($spl26 | $spl30)),
  inference(forward_demodulation,[],[f411,f413])).
fof(f6680,plain,(
  $false | ($spl26 | $spl30)),
  inference(equality_resolution,[],[f6643])).
fof(f6122,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X7),s('type.num.num',X6))))) ) | $spl36),
  inference(superposition,[],[f420,f239])).
fof(f6014,plain,(
  $false | ($spl30 | $spl34)),
  inference(subsumption_resolution,[],[f418,f5544])).
fof(f5956,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X7),s('type.num.num',X6))))) ) | $spl33),
  inference(superposition,[],[f416,f239])).
fof(f5897,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X7),s('type.num.num',X6))))) ) | $spl30),
  inference(superposition,[],[f5544,f239])).
fof(f5548,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X7),s('type.num.num',X6))))) ) | $spl30),
  inference(superposition,[],[f413,f239])).
fof(f5540,plain,(
  p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK12)))) | $spl144),
  inference(cnf_transformation,[],[f5540_D])).
fof(f5540_D,plain,(
  p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK12)))) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f5513,plain,(
  ( ! [X6,X7] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK12)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X7),s('type.num.num',X6))))) ) | $spl98),
  inference(superposition,[],[f2597,f239])).
fof(f5353,plain,(
  $false | ($spl2 | $spl99)),
  inference(subsumption_resolution,[],[f5351,f393])).
fof(f5351,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) | $spl99),
  inference(trivial_inequality_removal,[],[f5240])).
fof(f5240,plain,(
  s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',sK14))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',sK14))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) | $spl99),
  inference(resolution,[],[f2598,f242])).
fof(f5100,plain,(
  ( ! [X6,X7] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X7),s('type.num.num',X6))))) ) | $spl24),
  inference(superposition,[],[f410,f239])).
fof(f1601,plain,(
  ( ! [X12,X13] : (sP8(X12,X13,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl84),
  inference(cnf_transformation,[],[f1601_D])).
fof(f1601_D,plain,(
  ( ! [X12,X13] : (sP8(X12,X13,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f4870,plain,(
  ( ! [X14,X12,X15,X13] : (sP8(X12,X13,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl2),
  inference(superposition,[],[f995,f239])).
fof(f995,plain,(
  ( ! [X107,X106] : (sP8(X106,X107,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) ) | $spl2),
  inference(resolution,[],[f393,f285])).
fof(f1579,plain,(
  ( ! [X12,X13] : (sP6(X12,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X13)) ) | $spl82),
  inference(cnf_transformation,[],[f1579_D])).
fof(f1579_D,plain,(
  ( ! [X12,X13] : (sP6(X12,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X13)) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f4850,plain,(
  ( ! [X14,X12,X15,X13] : (sP6(X12,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X13) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl2),
  inference(superposition,[],[f979,f239])).
fof(f979,plain,(
  ( ! [X74,X75] : (sP6(X74,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),X75)) ) | $spl2),
  inference(resolution,[],[f393,f269])).
fof(f2951,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK12) | $spl101),
  inference(cnf_transformation,[],[f2951_D])).
fof(f2951_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK12) <=> ~$spl101),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl101])])).
fof(f3430,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) != s('type.num.num',X1)) ) | $spl136),
  inference(cnf_transformation,[],[f3430_D])).
fof(f3430_D,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) != s('type.num.num',X1)) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f4687,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) != s('type.num.num',X1) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK12)) ) | $spl16),
  inference(resolution,[],[f403,f241])).
fof(f241,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',X0),s('type.num.num',X2),s('type.num.num',X1)))) | s('type.num.num',X1) != s('type.num.num',X2) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f151])).
fof(f1517,plain,(
  ( ! [X12,X13] : (sP6(X12,X13,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl76),
  inference(cnf_transformation,[],[f1517_D])).
fof(f1517_D,plain,(
  ( ! [X12,X13] : (sP6(X12,X13,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f4683,plain,(
  ( ! [X14,X12,X15,X13] : (sP6(X12,X13,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl2),
  inference(superposition,[],[f978,f239])).
fof(f978,plain,(
  ( ! [X72,X73] : (sP6(X72,X73,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) ) | $spl2),
  inference(resolution,[],[f393,f268])).
fof(f1496,plain,(
  ( ! [X12,X13] : (sP4(X12,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X13)) ) | $spl74),
  inference(cnf_transformation,[],[f1496_D])).
fof(f1496_D,plain,(
  ( ! [X12,X13] : (sP4(X12,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X13)) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f4663,plain,(
  ( ! [X14,X12,X15,X13] : (sP4(X12,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X13) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl2),
  inference(superposition,[],[f965,f239])).
fof(f965,plain,(
  ( ! [X47,X46] : (sP4(X46,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),X47)) ) | $spl2),
  inference(resolution,[],[f393,f252])).
fof(f1475,plain,(
  ( ! [X12,X13] : (sP0('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X12,X13)) ) | $spl72),
  inference(cnf_transformation,[],[f1475_D])).
fof(f1475_D,plain,(
  ( ! [X12,X13] : (sP0('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X12,X13)) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f4641,plain,(
  ( ! [X14,X12,X15,X13] : (sP0('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X12,X13) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl2),
  inference(superposition,[],[f949,f239])).
fof(f949,plain,(
  ( ! [X24,X25] : (sP0('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),X24,X25)) ) | $spl2),
  inference(resolution,[],[f393,f218])).
fof(f1453,plain,(
  ( ! [X12,X13] : (sP0(X12,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X13)) ) | $spl70),
  inference(cnf_transformation,[],[f1453_D])).
fof(f1453_D,plain,(
  ( ! [X12,X13] : (sP0(X12,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X13)) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f4620,plain,(
  ( ! [X14,X12,X15,X13] : (sP0(X12,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X13) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl2),
  inference(superposition,[],[f948,f239])).
fof(f948,plain,(
  ( ! [X23,X22] : (sP0(X22,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),X23)) ) | $spl2),
  inference(resolution,[],[f393,f217])).
fof(f1431,plain,(
  ( ! [X12,X13] : (sP0(X12,X13,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl68),
  inference(cnf_transformation,[],[f1431_D])).
fof(f1431_D,plain,(
  ( ! [X12,X13] : (sP0(X12,X13,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f4599,plain,(
  ( ! [X14,X12,X15,X13] : (sP0(X12,X13,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl2),
  inference(superposition,[],[f947,f239])).
fof(f947,plain,(
  ( ! [X21,X20] : (sP0(X20,X21,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) ) | $spl2),
  inference(resolution,[],[f393,f216])).
fof(f1281,plain,(
  ( ! [X12,X13] : (sP1(X12,X13,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl66),
  inference(cnf_transformation,[],[f1281_D])).
fof(f1281_D,plain,(
  ( ! [X12,X13] : (sP1(X12,X13,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f4440,plain,(
  ( ! [X14,X12,X15,X13] : (sP1(X12,X13,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl2),
  inference(superposition,[],[f946,f239])).
fof(f946,plain,(
  ( ! [X19,X18] : (sP1(X18,X19,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) ) | $spl2),
  inference(resolution,[],[f393,f212])).
fof(f1260,plain,(
  ( ! [X12,X13] : (sP2(X12,X13,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl64),
  inference(cnf_transformation,[],[f1260_D])).
fof(f1260_D,plain,(
  ( ! [X12,X13] : (sP2(X12,X13,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f4420,plain,(
  ( ! [X14,X12,X15,X13] : (sP2(X12,X13,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl2),
  inference(superposition,[],[f943,f239])).
fof(f943,plain,(
  ( ! [X12,X13] : (sP2(X12,X13,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) ) | $spl2),
  inference(resolution,[],[f393,f208])).
fof(f1051,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | $spl57),
  inference(cnf_transformation,[],[f1051_D])).
fof(f1051_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) <=> ~$spl57),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl57])])).
fof(f4376,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X7),s('type.num.num',X6))))) ) | $spl19),
  inference(superposition,[],[f404,f239])).
fof(f4399,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X0),s('type.num.num',X1))))) ) | $spl19),
  inference(subsumption_resolution,[],[f4373,f404])).
fof(f4373,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X0),s('type.num.num',X1))))) ) | $spl19),
  inference(superposition,[],[f404,f237])).
fof(f2950,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK12) | $spl100),
  inference(cnf_transformation,[],[f2950_D])).
fof(f2950_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK12) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f4281,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',sK12) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X0),s('type.num.num',X1))))) ) | $spl19),
  inference(resolution,[],[f404,f237])).
fof(f402,plain,(
  ( ! [X6,X7] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK13))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK13))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK13),s('type.num.num',X7),s('type.num.num',X6))))) ) | $spl14),
  inference(cnf_transformation,[],[f402_D])).
fof(f402_D,plain,(
  ( ! [X6,X7] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK13))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK13))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK13),s('type.num.num',X7),s('type.num.num',X6))))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f4251,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK13))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK13))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK13),s('type.num.num',X7),s('type.num.num',X6))))) ) | $spl11),
  inference(superposition,[],[f398,f239])).
fof(f397,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK13),s('type.num.num',X0),s('type.num.num',X1))))) ) | $spl8),
  inference(cnf_transformation,[],[f397_D])).
fof(f397_D,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK13),s('type.num.num',X0),s('type.num.num',X1))))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f4275,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK13),s('type.num.num',X0),s('type.num.num',X1))))) ) | $spl11),
  inference(subsumption_resolution,[],[f4248,f398])).
fof(f4248,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK13),s('type.num.num',X0),s('type.num.num',X1))))) ) | $spl11),
  inference(superposition,[],[f398,f237])).
fof(f1052,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | $spl56),
  inference(cnf_transformation,[],[f1052_D])).
fof(f1052_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f4134,plain,(
  ( ! [X6,X7] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X7),s('type.num.num',X6))))) ) | $spl2),
  inference(superposition,[],[f393,f239])).
fof(f420,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',sK12))) | $spl36),
  inference(cnf_transformation,[],[f420_D])).
fof(f420_D,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',sK12))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f414,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK13))) | $spl31),
  inference(cnf_transformation,[],[f414_D])).
fof(f414_D,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK13))) <=> ~$spl31),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl31])])).
fof(f4039,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK13))) | s(bool,'T_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',sK12)))),
  inference(superposition,[],[f199,f236])).
fof(f418,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',sK12))) | $spl34),
  inference(cnf_transformation,[],[f418_D])).
fof(f418_D,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',sK12))) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f416,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK13))) | $spl33),
  inference(cnf_transformation,[],[f416_D])).
fof(f416_D,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK13))) <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f4038,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK13))) | s(bool,'F_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',sK12)))),
  inference(superposition,[],[f199,f236])).
fof(f409,plain,(
  ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',sK12)))) | $spl25),
  inference(cnf_transformation,[],[f409_D])).
fof(f409_D,plain,(
  ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',sK12)))) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f4037,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK13))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',sK12))))),
  inference(superposition,[],[f199,f235])).
fof(f4036,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK13))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',sK12))))),
  inference(superposition,[],[f199,f233])).
fof(f410,plain,(
  p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',sK12)))) | $spl24),
  inference(cnf_transformation,[],[f410_D])).
fof(f410_D,plain,(
  p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',sK12)))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f4035,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK13))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',sK12))))),
  inference(superposition,[],[f199,f231])).
fof(f4034,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK13))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',sK12))))),
  inference(superposition,[],[f199,f229])).
fof(f412,plain,(
  ( ! [X10] : (s(bool,X10) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK13))) | ~p(s(bool,X10))) ) | $spl28),
  inference(cnf_transformation,[],[f412_D])).
fof(f412_D,plain,(
  ( ! [X10] : (s(bool,X10) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK13))) | ~p(s(bool,X10))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f4033,plain,(
  ( ! [X11] : (s(bool,X11) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK13))) | ~p(s(bool,X11)) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',sK12))))) )),
  inference(superposition,[],[f199,f226])).
fof(f4032,plain,(
  ( ! [X10] : (s(bool,X10) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK13))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',sK12)))) | ~p(s(bool,X10))) )),
  inference(superposition,[],[f199,f226])).
fof(f411,plain,(
  ( ! [X8] : (s(bool,X8) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK13))) | p(s(bool,X8))) ) | $spl26),
  inference(cnf_transformation,[],[f411_D])).
fof(f411_D,plain,(
  ( ! [X8] : (s(bool,X8) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK13))) | p(s(bool,X8))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f4031,plain,(
  ( ! [X9] : (s(bool,X9) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK13))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',sK12)))) | p(s(bool,X9))) )),
  inference(superposition,[],[f199,f223])).
fof(f4030,plain,(
  ( ! [X8] : (s(bool,X8) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK13))) | p(s(bool,X8)) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',sK12))))) )),
  inference(superposition,[],[f199,f223])).
fof(f4023,plain,(
  ( ! [X6,X7] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X7),s('type.num.num',X6))))) )),
  inference(superposition,[],[f199,f239])).
fof(f4020,plain,(
  ( ! [X0,X1] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(superposition,[],[f199,f237])).
fof(f401,plain,(
  s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0'))) | $spl13),
  inference(cnf_transformation,[],[f401_D])).
fof(f401_D,plain,(
  s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0'))) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f4017,plain,(
  ( ! [X6,X7] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK13))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK13))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK13),s('type.num.num',X7),s('type.num.num',X6))))) )),
  inference(superposition,[],[f199,f239])).
fof(f399,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)))) | $spl10),
  inference(cnf_transformation,[],[f399_D])).
fof(f399_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f4014,plain,(
  ( ! [X0,X1] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK13),s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(superposition,[],[f199,f237])).
fof(f395,plain,(
  s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | $spl5),
  inference(cnf_transformation,[],[f395_D])).
fof(f395_D,plain,(
  s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f4011,plain,(
  ( ! [X6,X7] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X7),s('type.num.num',X6))))) )),
  inference(superposition,[],[f199,f239])).
fof(f391,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X0),s('type.num.num',X1))))) ) | $spl0),
  inference(cnf_transformation,[],[f391_D])).
fof(f391_D,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X0),s('type.num.num',X1))))) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f393,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) | $spl2),
  inference(cnf_transformation,[],[f393_D])).
fof(f393_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f4008,plain,(
  ( ! [X0,X1] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(superposition,[],[f199,f237])).
fof(f1089,plain,(
  ( ! [X9] : (sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X9)) ) | $spl60),
  inference(cnf_transformation,[],[f1089_D])).
fof(f1089_D,plain,(
  ( ! [X9] : (sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X9)) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f4003,plain,(
  ( ! [X10,X11,X9] : (sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X9) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X10),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X11),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X11),s('type.num.num',X10))))) ) | $spl2),
  inference(superposition,[],[f1012,f239])).
fof(f1012,plain,(
  ( ! [X131] : (sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),X131)) ) | $spl2),
  inference(resolution,[],[f393,f312])).
fof(f1072,plain,(
  ( ! [X9] : (sP11(X9,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl58),
  inference(cnf_transformation,[],[f1072_D])).
fof(f1072_D,plain,(
  ( ! [X9] : (sP11(X9,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f3985,plain,(
  ( ! [X10,X11,X9] : (sP11(X9,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X10),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X11),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X11),s('type.num.num',X10))))) ) | $spl2),
  inference(superposition,[],[f1011,f239])).
fof(f1011,plain,(
  ( ! [X130] : (sP11(X130,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) ) | $spl2),
  inference(resolution,[],[f393,f311])).
fof(f3969,plain,(
  ( ! [X6,X7] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',sK12) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK13))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK13))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK13),s('type.num.num',X7),s('type.num.num',X6))))) ) | $spl41),
  inference(superposition,[],[f542,f239])).
fof(f3966,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',sK12) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK13),s('type.num.num',X0),s('type.num.num',X1))))) ) | $spl41),
  inference(superposition,[],[f542,f237])).
fof(f3965,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X0),s('type.num.num',X1))))) ) | $spl39),
  inference(trivial_inequality_removal,[],[f3958])).
fof(f3958,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.|48|_0') | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X0),s('type.num.num',X1))))) ) | $spl39),
  inference(superposition,[],[f540,f237])).
fof(f3544,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK13))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK13))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK13),s('type.num.num',X7),s('type.num.num',X6))))) ) | $spl11),
  inference(superposition,[],[f398,f239])).
fof(f3541,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK13),s('type.num.num',X0),s('type.num.num',X1))))) ) | $spl11),
  inference(superposition,[],[f398,f237])).
fof(f3957,plain,(
  $false | $spl40),
  inference(trivial_inequality_removal,[],[f3956])).
fof(f3956,plain,(
  s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK12))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK12))) | $spl40),
  inference(forward_demodulation,[],[f199,f541])).
fof(f3897,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3))))) ) | ($spl40 | $spl135)),
  inference(forward_demodulation,[],[f3896,f541])).
fof(f3896,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))))) ) | ($spl40 | $spl135)),
  inference(subsumption_resolution,[],[f3847,f3429])).
fof(f3847,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X2),s('type.num.num',X3)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK13)) ) | $spl40),
  inference(superposition,[],[f237,f541])).
fof(f3953,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X6),s('type.num.num',X7)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))))) ) | ($spl40 | $spl135)),
  inference(forward_demodulation,[],[f3952,f541])).
fof(f3952,plain,(
  ( ! [X6,X7] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK13),s('type.num.num',X6),s('type.num.num',X7))))) ) | ($spl40 | $spl135)),
  inference(subsumption_resolution,[],[f3849,f3429])).
fof(f3849,plain,(
  ( ! [X6,X7] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK13) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK13),s('type.num.num',X6),s('type.num.num',X7))))) ) | $spl40),
  inference(superposition,[],[f237,f541])).
fof(f3947,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | ($spl11 | $spl40)),
  inference(forward_demodulation,[],[f3946,f541])).
fof(f3946,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK13),s('type.num.num',X0),s('type.num.num',X1))))) ) | ($spl11 | $spl40)),
  inference(forward_demodulation,[],[f3541,f541])).
fof(f3945,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X7),s('type.num.num',X6)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK12))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | ($spl11 | $spl40)),
  inference(forward_demodulation,[],[f3944,f541])).
fof(f3944,plain,(
  ( ! [X6,X7] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK12))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK13),s('type.num.num',X7),s('type.num.num',X6))))) ) | ($spl11 | $spl40)),
  inference(forward_demodulation,[],[f3544,f541])).
fof(f3886,plain,(
  $false | ($spl11 | $spl18 | $spl40)),
  inference(subsumption_resolution,[],[f3836,f405])).
fof(f3836,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ($spl11 | $spl40)),
  inference(backward_demodulation,[],[f541,f398])).
fof(f3885,plain,(
  $false | ($spl38 | $spl40)),
  inference(trivial_inequality_removal,[],[f3838])).
fof(f3838,plain,(
  s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK12))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK12))) | ($spl38 | $spl40)),
  inference(backward_demodulation,[],[f541,f3242])).
fof(f3242,plain,(
  s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | $spl38),
  inference(forward_demodulation,[],[f199,f539])).
fof(f3834,plain,(
  p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')))) | $spl142),
  inference(cnf_transformation,[],[f3834_D])).
fof(f3834_D,plain,(
  p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')))) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f3805,plain,(
  ( ! [X6,X7] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X7),s('type.num.num',X6))))) ) | ($spl24 | $spl38)),
  inference(superposition,[],[f3614,f239])).
fof(f3614,plain,(
  p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) | ($spl24 | $spl38)),
  inference(forward_demodulation,[],[f410,f539])).
fof(f541,plain,(
  s('type.num.num',sK12) = s('type.num.num',sK13) | $spl40),
  inference(cnf_transformation,[],[f541_D])).
fof(f541_D,plain,(
  s('type.num.num',sK12) = s('type.num.num',sK13) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f3832,plain,(
  s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | $spl140),
  inference(cnf_transformation,[],[f3832_D])).
fof(f3832_D,plain,(
  s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f3705,plain,(
  s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK12) = s('type.num.num',sK13) | ($spl24 | $spl38)),
  inference(resolution,[],[f3614,f240])).
fof(f3825,plain,(
  $false | ($spl3 | $spl24 | $spl38 | $spl41)),
  inference(subsumption_resolution,[],[f3824,f542])).
fof(f3824,plain,(
  s('type.num.num',sK12) = s('type.num.num',sK13) | ($spl3 | $spl24 | $spl38)),
  inference(subsumption_resolution,[],[f3703,f2952])).
fof(f2952,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | ($spl3 | $spl38)),
  inference(forward_demodulation,[],[f392,f539])).
fof(f3703,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | s('type.num.num',sK12) = s('type.num.num',sK13) | ($spl24 | $spl38)),
  inference(resolution,[],[f3614,f238])).
fof(f398,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)))) | $spl11),
  inference(cnf_transformation,[],[f398_D])).
fof(f398_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f3431,plain,(
  ( ! [X2,X3] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X2),s('type.num.num',sK13))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X3),s('type.num.num',sK13)))) ) | $spl138),
  inference(cnf_transformation,[],[f3431_D])).
fof(f3431_D,plain,(
  ( ! [X2,X3] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X2),s('type.num.num',sK13))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X3),s('type.num.num',sK13)))) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f3297,plain,(
  ( ! [X2,X3] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X2),s('type.num.num',sK13))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X3),s('type.num.num',sK13))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))))) ) | $spl8),
  inference(resolution,[],[f397,f242])).
fof(f3429,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK13) | $spl135),
  inference(cnf_transformation,[],[f3429_D])).
fof(f3429_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK13) <=> ~$spl135),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl135])])).
fof(f3296,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) != s('type.num.num',X1) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK13)) ) | $spl8),
  inference(resolution,[],[f397,f241])).
fof(f3281,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) | s(bool,'T_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | $spl38),
  inference(superposition,[],[f3242,f236])).
fof(f3280,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) | s(bool,'F_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | $spl38),
  inference(superposition,[],[f3242,f236])).
fof(f3279,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) | $spl38),
  inference(superposition,[],[f3242,f235])).
fof(f3278,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) | $spl38),
  inference(superposition,[],[f3242,f233])).
fof(f3277,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) | $spl38),
  inference(superposition,[],[f3242,f231])).
fof(f3276,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) | $spl38),
  inference(superposition,[],[f3242,f229])).
fof(f3275,plain,(
  ( ! [X11] : (s(bool,X11) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) | ~p(s(bool,X11)) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))))) ) | $spl38),
  inference(superposition,[],[f3242,f226])).
fof(f3274,plain,(
  ( ! [X10] : (s(bool,X10) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) | ~p(s(bool,X10))) ) | $spl38),
  inference(superposition,[],[f3242,f226])).
fof(f3273,plain,(
  ( ! [X9] : (s(bool,X9) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) | p(s(bool,X9))) ) | $spl38),
  inference(superposition,[],[f3242,f223])).
fof(f3272,plain,(
  ( ! [X8] : (s(bool,X8) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) | p(s(bool,X8)) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))))) ) | $spl38),
  inference(superposition,[],[f3242,f223])).
fof(f3265,plain,(
  ( ! [X6,X7] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X7),s('type.num.num',X6))))) ) | $spl38),
  inference(superposition,[],[f3242,f239])).
fof(f3259,plain,(
  ( ! [X6,X7] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK13))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK13))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK13),s('type.num.num',X7),s('type.num.num',X6))))) ) | $spl38),
  inference(superposition,[],[f3242,f239])).
fof(f3256,plain,(
  ( ! [X0,X1] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK13),s('type.num.num',X0),s('type.num.num',X1))))) ) | $spl38),
  inference(superposition,[],[f3242,f237])).
fof(f2615,plain,(
  ( ! [X14,X12,X15,X13] : (sP0('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X12,X13) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl18),
  inference(superposition,[],[f2343,f239])).
fof(f2343,plain,(
  ( ! [X24,X25] : (sP0('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)),X24,X25)) ) | $spl18),
  inference(resolution,[],[f405,f218])).
fof(f2557,plain,(
  ( ! [X14,X12,X15,X13] : (sP0(X12,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X13) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl18),
  inference(superposition,[],[f2342,f239])).
fof(f2342,plain,(
  ( ! [X23,X22] : (sP0(X22,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)),X23)) ) | $spl18),
  inference(resolution,[],[f405,f217])).
fof(f2536,plain,(
  ( ! [X14,X12,X15,X13] : (sP0(X12,X13,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl18),
  inference(superposition,[],[f2341,f239])).
fof(f2341,plain,(
  ( ! [X21,X20] : (sP0(X20,X21,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) ) | $spl18),
  inference(resolution,[],[f405,f216])).
fof(f3124,plain,(
  $false | $spl62),
  inference(subsumption_resolution,[],[f3070,f1226])).
fof(f3070,plain,(
  ~p(s(bool,'T_0')) | $spl62),
  inference(resolution,[],[f1226,f352])).
fof(f352,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f198])).
fof(f3123,plain,(
  $false | $spl62),
  inference(subsumption_resolution,[],[f3069,f1226])).
fof(f3069,plain,(
  ~p(s(bool,'T_0')) | $spl62),
  inference(resolution,[],[f1226,f350])).
fof(f350,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f197])).
fof(f3122,plain,(
  $false | $spl62),
  inference(subsumption_resolution,[],[f3066,f1226])).
fof(f3066,plain,(
  ( ! [X220] : (~p(s(bool,X220))) ) | $spl62),
  inference(resolution,[],[f1226,f328])).
fof(f3068,plain,(
  $false | $spl62),
  inference(resolution,[],[f1226,f348])).
fof(f2657,plain,(
  ( ! [X14,X12,X15,X13] : (sP6(X12,X13,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl18),
  inference(superposition,[],[f2372,f239])).
fof(f2372,plain,(
  ( ! [X72,X73] : (sP6(X72,X73,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) ) | $spl18),
  inference(resolution,[],[f405,f268])).
fof(f2637,plain,(
  ( ! [X14,X12,X15,X13] : (sP4(X12,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X13) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl18),
  inference(superposition,[],[f2359,f239])).
fof(f2359,plain,(
  ( ! [X47,X46] : (sP4(X46,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)),X47)) ) | $spl18),
  inference(resolution,[],[f405,f252])).
fof(f2516,plain,(
  ( ! [X14,X12,X15,X13] : (sP1(X12,X13,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl18),
  inference(superposition,[],[f2340,f239])).
fof(f2340,plain,(
  ( ! [X19,X18] : (sP1(X18,X19,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) ) | $spl18),
  inference(resolution,[],[f405,f212])).
fof(f2978,plain,(
  ( ! [X6,X7] : (sP3('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X6,X7) | ~sP1(X6,X7,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl120),
  inference(cnf_transformation,[],[f2978_D])).
fof(f2978_D,plain,(
  ( ! [X6,X7] : (sP3('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X6,X7) | ~sP1(X6,X7,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f1800,plain,(
  ( ! [X6,X7] : (sP3('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X6,X7) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | ~sP1(X6,X7,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl68),
  inference(subsumption_resolution,[],[f1789,f210])).
fof(f1789,plain,(
  ( ! [X6,X7] : (sP3('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X6,X7) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | ~sP2('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X7,X6) | ~sP1(X6,X7,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl68),
  inference(resolution,[],[f1431,f206])).
fof(f2496,plain,(
  ( ! [X14,X12,X15,X13] : (sP2(X12,X13,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl18),
  inference(superposition,[],[f2337,f239])).
fof(f2337,plain,(
  ( ! [X12,X13] : (sP2(X12,X13,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) ) | $spl18),
  inference(resolution,[],[f405,f208])).
fof(f2977,plain,(
  ( ! [X4,X5] : (sP3(X4,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X5) | ~sP2(X4,X5,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl118),
  inference(cnf_transformation,[],[f2977_D])).
fof(f2977_D,plain,(
  ( ! [X4,X5] : (sP3(X4,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X5) | ~sP2(X4,X5,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f1814,plain,(
  ( ! [X4,X5] : (sP3(X4,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X5) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | ~sP2(X4,X5,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl70),
  inference(subsumption_resolution,[],[f1803,f214])).
fof(f1803,plain,(
  ( ! [X4,X5] : (sP3(X4,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X5) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | ~sP2(X4,X5,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~sP1('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X5,X4)) ) | $spl70),
  inference(resolution,[],[f1453,f205])).
fof(f2477,plain,(
  ( ! [X10,X11,X9] : (sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X9) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X10),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X11),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X11),s('type.num.num',X10))))) ) | $spl18),
  inference(superposition,[],[f2406,f239])).
fof(f2406,plain,(
  ( ! [X131] : (sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)),X131)) ) | $spl18),
  inference(resolution,[],[f405,f312])).
fof(f2461,plain,(
  ( ! [X10,X11,X9] : (sP11(X9,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X10),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X11),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X11),s('type.num.num',X10))))) ) | $spl18),
  inference(superposition,[],[f2405,f239])).
fof(f2405,plain,(
  ( ! [X130] : (sP11(X130,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) ) | $spl18),
  inference(resolution,[],[f405,f311])).
fof(f1713,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | $spl64),
  inference(resolution,[],[f1260,f207])).
fof(f2969,plain,(
  ( ! [X4,X5] : (sP3(X4,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X5)) ) | $spl114),
  inference(cnf_transformation,[],[f2969_D])).
fof(f2969_D,plain,(
  ( ! [X4,X5] : (sP3(X4,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X5)) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f2976,plain,(
  ( ! [X4,X5] : (sP3(X4,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X5) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | ($spl64 | $spl70)),
  inference(subsumption_resolution,[],[f1726,f1453])).
fof(f1726,plain,(
  ( ! [X4,X5] : (sP3(X4,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X5) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | ~sP0(X5,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X4)) ) | $spl64),
  inference(subsumption_resolution,[],[f1715,f214])).
fof(f1715,plain,(
  ( ! [X4,X5] : (sP3(X4,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X5) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | ~sP1('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X5,X4) | ~sP0(X5,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X4)) ) | $spl64),
  inference(resolution,[],[f1260,f205])).
fof(f1727,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | $spl66),
  inference(resolution,[],[f1281,f211])).
fof(f2972,plain,(
  ( ! [X6,X7] : (sP3('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X6,X7)) ) | $spl116),
  inference(cnf_transformation,[],[f2972_D])).
fof(f2972_D,plain,(
  ( ! [X6,X7] : (sP3('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X6,X7)) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f2974,plain,(
  ( ! [X6,X7] : (sP3('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X6,X7) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | ($spl66 | $spl68)),
  inference(subsumption_resolution,[],[f1740,f1431])).
fof(f1740,plain,(
  ( ! [X6,X7] : (sP3('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X6,X7) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | ~sP0(X7,X6,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl66),
  inference(subsumption_resolution,[],[f1730,f210])).
fof(f1730,plain,(
  ( ! [X6,X7] : (sP3('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X6,X7) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | ~sP2('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X7,X6) | ~sP0(X7,X6,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl66),
  inference(resolution,[],[f1281,f206])).
fof(f1786,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | $spl68),
  inference(resolution,[],[f1431,f215])).
fof(f2971,plain,(
  ( ! [X6,X7] : (sP3('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X6,X7) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | ($spl66 | $spl68)),
  inference(subsumption_resolution,[],[f1800,f1281])).
fof(f1801,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X1))) ) | $spl70),
  inference(resolution,[],[f1453,f215])).
fof(f2968,plain,(
  ( ! [X4,X5] : (sP3(X4,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X5) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | ($spl64 | $spl70)),
  inference(subsumption_resolution,[],[f1814,f1260])).
fof(f1816,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X0)) | p(s(bool,X1))) ) | $spl72),
  inference(resolution,[],[f1475,f215])).
fof(f2966,plain,(
  ( ! [X2,X3] : (sP3(X2,X3,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~sP1(X3,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X2) | ~sP2(X2,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X3)) ) | $spl112),
  inference(cnf_transformation,[],[f2966_D])).
fof(f2966_D,plain,(
  ( ! [X2,X3] : (sP3(X2,X3,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~sP1(X3,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X2) | ~sP2(X2,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X3)) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f1817,plain,(
  ( ! [X2,X3] : (sP3(X2,X3,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | ~sP2(X2,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X3) | ~sP1(X3,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X2)) ) | $spl72),
  inference(resolution,[],[f1475,f204])).
fof(f1831,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X1))) ) | $spl74),
  inference(resolution,[],[f1496,f250])).
fof(f2963,plain,(
  ( ! [X8,X9] : (sP5(X8,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X9) | ~p(s(bool,X9))) ) | $spl110),
  inference(cnf_transformation,[],[f2963_D])).
fof(f2963_D,plain,(
  ( ! [X8,X9] : (sP5(X8,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X9) | ~p(s(bool,X9))) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f1835,plain,(
  ( ! [X8,X9] : (sP5(X8,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X9) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X9))) ) | $spl74),
  inference(resolution,[],[f1496,f249])).
fof(f1845,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | $spl76),
  inference(resolution,[],[f1517,f267])).
fof(f2961,plain,(
  ( ! [X8,X9] : (sP7(X8,X9,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X9))) ) | $spl108),
  inference(cnf_transformation,[],[f2961_D])).
fof(f2961_D,plain,(
  ( ! [X8,X9] : (sP7(X8,X9,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X9))) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f1849,plain,(
  ( ! [X8,X9] : (sP7(X8,X9,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X9)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | $spl76),
  inference(resolution,[],[f1517,f266])).
fof(f2425,plain,(
  ( ! [X6,X7] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X7),s('type.num.num',X6))))) ) | $spl18),
  inference(superposition,[],[f405,f239])).
fof(f2703,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X1))) ) | $spl82),
  inference(resolution,[],[f1579,f267])).
fof(f2959,plain,(
  ( ! [X8,X9] : (sP7(X8,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X9) | p(s(bool,X9))) ) | $spl106),
  inference(cnf_transformation,[],[f2959_D])).
fof(f2959_D,plain,(
  ( ! [X8,X9] : (sP7(X8,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X9) | p(s(bool,X9))) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f2707,plain,(
  ( ! [X8,X9] : (sP7(X8,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X9) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X9))) ) | $spl82),
  inference(resolution,[],[f1579,f266])).
fof(f2717,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | $spl84),
  inference(resolution,[],[f1601,f284])).
fof(f2957,plain,(
  ( ! [X8,X9] : (sP9('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X8,X9)) ) | $spl102),
  inference(cnf_transformation,[],[f2957_D])).
fof(f2957_D,plain,(
  ( ! [X8,X9] : (sP9('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X8,X9)) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f2730,plain,(
  ( ! [X4,X5] : (sP9('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X4,X5) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | $spl84),
  inference(subsumption_resolution,[],[f2719,f289])).
fof(f2719,plain,(
  ( ! [X4,X5] : (sP9('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X4,X5) | ~p(s(bool,X4)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | $spl84),
  inference(resolution,[],[f1601,f281])).
fof(f2731,plain,(
  ( ! [X6,X7] : (sP9('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X6,X7) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | $spl84),
  inference(subsumption_resolution,[],[f2720,f288])).
fof(f2720,plain,(
  ( ! [X6,X7] : (sP9('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X6,X7) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X7))) ) | $spl84),
  inference(resolution,[],[f1601,f282])).
fof(f2721,plain,(
  ( ! [X8,X9] : (sP9('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X8,X9) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | $spl84),
  inference(resolution,[],[f1601,f357])).
fof(f2953,plain,(
  $false | ($spl3 | $spl38 | $spl56)),
  inference(subsumption_resolution,[],[f2952,f1052])).
fof(f2947,plain,(
  ( ! [X6,X7] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',sK12) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK13))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK13))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK13),s('type.num.num',X7),s('type.num.num',X6))))) ) | $spl41),
  inference(superposition,[],[f542,f239])).
fof(f2944,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',sK12) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK13),s('type.num.num',X0),s('type.num.num',X1))))) ) | $spl41),
  inference(superposition,[],[f542,f237])).
fof(f540,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK14) | $spl39),
  inference(cnf_transformation,[],[f540_D])).
fof(f540_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK14) <=> ~$spl39),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl39])])).
fof(f542,plain,(
  s('type.num.num',sK12) != s('type.num.num',sK13) | $spl41),
  inference(cnf_transformation,[],[f542_D])).
fof(f542_D,plain,(
  s('type.num.num',sK12) != s('type.num.num',sK13) <=> ~$spl41),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl41])])).
fof(f1282,plain,(
  s('type.num.num',sK12) != s('type.num.num',sK13) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK14) | $spl25),
  inference(resolution,[],[f409,f241])).
fof(f539,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK14) | $spl38),
  inference(cnf_transformation,[],[f539_D])).
fof(f539_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK14) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f543,plain,(
  s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',sK14))) | $spl42),
  inference(cnf_transformation,[],[f543_D])).
fof(f543_D,plain,(
  s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',sK14))) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f2094,plain,(
  s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',sK14))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK14) | $spl32),
  inference(subsumption_resolution,[],[f1892,f353])).
fof(f1892,plain,(
  ~p(s(bool,'T_0')) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',sK14))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK14) | $spl32),
  inference(superposition,[],[f239,f415])).
fof(f415,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK13))) | $spl32),
  inference(cnf_transformation,[],[f415_D])).
fof(f415_D,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK13))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f392,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) | $spl3),
  inference(cnf_transformation,[],[f392_D])).
fof(f392_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f544,plain,(
  s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',sK14))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',sK14))) | $spl43),
  inference(cnf_transformation,[],[f544_D])).
fof(f544_D,plain,(
  s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',sK14))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',sK14))) <=> ~$spl43),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl43])])).
fof(f1283,plain,(
  s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',sK14))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',sK14))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) | $spl25),
  inference(resolution,[],[f409,f242])).
fof(f1366,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X0),s('type.num.num',X1))))) ) | $spl25),
  inference(superposition,[],[f409,f237])).
fof(f1518,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X0),s('type.num.num',X1))))) ) | $spl31),
  inference(superposition,[],[f414,f237])).
fof(f1776,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) | $spl86),
  inference(cnf_transformation,[],[f1776_D])).
fof(f1776_D,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f1861,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X0),s('type.num.num',X1))))) ) | $spl32),
  inference(superposition,[],[f415,f237])).
fof(f2092,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK14) | $spl32),
  inference(subsumption_resolution,[],[f1890,f330])).
fof(f1890,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK14) | $spl32),
  inference(superposition,[],[f237,f415])).
fof(f2093,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) | s('type.num.num',sK12) = s('type.num.num',sK13) | $spl32),
  inference(subsumption_resolution,[],[f1891,f330])).
fof(f1891,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) | s('type.num.num',sK12) = s('type.num.num',sK13) | $spl32),
  inference(superposition,[],[f238,f415])).
fof(f2846,plain,(
  $false | ($spl2 | $spl25 | $spl40)),
  inference(subsumption_resolution,[],[f2844,f393])).
fof(f2844,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) | ($spl25 | $spl40)),
  inference(trivial_inequality_removal,[],[f2733])).
fof(f2733,plain,(
  s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',sK14))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',sK14))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) | ($spl25 | $spl40)),
  inference(resolution,[],[f2223,f242])).
fof(f2223,plain,(
  ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK12)))) | ($spl25 | $spl40)),
  inference(backward_demodulation,[],[f541,f409])).
fof(f2697,plain,(
  ( ! [X14,X12,X15,X13] : (sP8(X12,X13,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl18),
  inference(superposition,[],[f2389,f239])).
fof(f2389,plain,(
  ( ! [X107,X106] : (sP8(X106,X107,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) ) | $spl18),
  inference(resolution,[],[f405,f285])).
fof(f2677,plain,(
  ( ! [X14,X12,X15,X13] : (sP6(X12,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X13) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl18),
  inference(superposition,[],[f2373,f239])).
fof(f2373,plain,(
  ( ! [X74,X75] : (sP6(X74,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)),X75)) ) | $spl18),
  inference(resolution,[],[f405,f269])).
fof(f2598,plain,(
  ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK12)))) | $spl99),
  inference(cnf_transformation,[],[f2598_D])).
fof(f2598_D,plain,(
  ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK12)))) <=> ~$spl99),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl99])])).
fof(f2599,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,X11) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK12))))) ) | ($spl32 | $spl40)),
  inference(subsumption_resolution,[],[f2582,f232])).
fof(f2582,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,X11) | ~p(s(bool,X11)) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK12))))) ) | ($spl32 | $spl40)),
  inference(superposition,[],[f2480,f226])).
fof(f2480,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK12))) | ($spl32 | $spl40)),
  inference(forward_demodulation,[],[f1860,f541])).
fof(f1860,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',sK12))) | $spl32),
  inference(backward_demodulation,[],[f415,f199])).
fof(f2596,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,X10) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK12))))) ) | ($spl32 | $spl40)),
  inference(subsumption_resolution,[],[f2581,f232])).
fof(f2581,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,X10) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK12)))) | ~p(s(bool,X10))) ) | ($spl32 | $spl40)),
  inference(superposition,[],[f2480,f226])).
fof(f2572,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X7),s('type.num.num',X6))))) ) | ($spl32 | $spl40)),
  inference(superposition,[],[f2480,f239])).
fof(f2593,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK12))) | $spl95),
  inference(cnf_transformation,[],[f2593_D])).
fof(f2593_D,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK12))) <=> ~$spl95),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl95])])).
fof(f2566,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK12))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X7),s('type.num.num',X6))))) ) | ($spl32 | $spl40)),
  inference(superposition,[],[f2480,f239])).
fof(f2096,plain,(
  s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',sK14))) | s('type.num.num',sK12) = s('type.num.num',sK13) | $spl32),
  inference(subsumption_resolution,[],[f1893,f353])).
fof(f1893,plain,(
  ~p(s(bool,'T_0')) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',sK14))) | s('type.num.num',sK12) = s('type.num.num',sK13) | $spl32),
  inference(superposition,[],[f240,f415])).
fof(f1876,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK13))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK13))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK13),s('type.num.num',X7),s('type.num.num',X6))))) ) | $spl32),
  inference(superposition,[],[f415,f239])).
fof(f1873,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK13),s('type.num.num',X0),s('type.num.num',X1))))) ) | $spl32),
  inference(superposition,[],[f415,f237])).
fof(f1778,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) | $spl88),
  inference(cnf_transformation,[],[f1778_D])).
fof(f1778_D,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f1870,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X7),s('type.num.num',X6))))) ) | $spl32),
  inference(superposition,[],[f415,f239])).
fof(f1867,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X0),s('type.num.num',X1))))) ) | $spl32),
  inference(superposition,[],[f415,f237])).
fof(f1864,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X7),s('type.num.num',X6))))) ) | $spl32),
  inference(superposition,[],[f415,f239])).
fof(f1785,plain,(
  $false | ($spl31 | $spl33)),
  inference(subsumption_resolution,[],[f1773,f414])).
fof(f1773,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK13))) | $spl33),
  inference(trivial_inequality_removal,[],[f1771])).
fof(f1771,plain,(
  s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK13))) | $spl33),
  inference(superposition,[],[f416,f236])).
fof(f1556,plain,(
  ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK13)))) | $spl81),
  inference(cnf_transformation,[],[f1556_D])).
fof(f1556_D,plain,(
  ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK13)))) <=> ~$spl81),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl81])])).
fof(f1784,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,X11) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK13))))) ) | $spl33),
  inference(subsumption_resolution,[],[f1766,f232])).
fof(f1766,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,X11) | ~p(s(bool,X11)) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK13))))) ) | $spl33),
  inference(superposition,[],[f416,f226])).
fof(f1782,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,X10) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK13))))) ) | $spl33),
  inference(subsumption_resolution,[],[f1765,f232])).
fof(f1765,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,X10) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK13)))) | ~p(s(bool,X10))) ) | $spl33),
  inference(superposition,[],[f416,f226])).
fof(f1756,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK13))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK13))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK13),s('type.num.num',X7),s('type.num.num',X6))))) ) | $spl33),
  inference(superposition,[],[f416,f239])).
fof(f1753,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK13),s('type.num.num',X0),s('type.num.num',X1))))) ) | $spl33),
  inference(superposition,[],[f416,f237])).
fof(f1750,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X7),s('type.num.num',X6))))) ) | $spl33),
  inference(superposition,[],[f416,f239])).
fof(f1747,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X0),s('type.num.num',X1))))) ) | $spl33),
  inference(superposition,[],[f416,f237])).
fof(f1744,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X7),s('type.num.num',X6))))) ) | $spl33),
  inference(superposition,[],[f416,f239])).
fof(f1596,plain,(
  ( ! [X14,X12,X15,X13] : (sP8(X12,X13,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl2),
  inference(superposition,[],[f995,f239])).
fof(f1575,plain,(
  ( ! [X14,X12,X15,X13] : (sP6(X12,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X13) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl2),
  inference(superposition,[],[f979,f239])).
fof(f1557,plain,(
  p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK13)))) | $spl80),
  inference(cnf_transformation,[],[f1557_D])).
fof(f1557_D,plain,(
  p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK13)))) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f1558,plain,(
  ( ! [X9] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK13)))) | p(s(bool,X9))) ) | $spl31),
  inference(subsumption_resolution,[],[f1541,f229])).
fof(f1541,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,X9) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK13)))) | p(s(bool,X9))) ) | $spl31),
  inference(superposition,[],[f414,f223])).
fof(f1555,plain,(
  ( ! [X8] : (p(s(bool,X8)) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK13))))) ) | $spl31),
  inference(subsumption_resolution,[],[f1540,f229])).
fof(f1540,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,X8) | p(s(bool,X8)) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK13))))) ) | $spl31),
  inference(superposition,[],[f414,f223])).
fof(f909,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0'))) | $spl55),
  inference(cnf_transformation,[],[f909_D])).
fof(f909_D,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0'))) <=> ~$spl55),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl55])])).
fof(f1533,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK13))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK13))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK13),s('type.num.num',X7),s('type.num.num',X6))))) ) | $spl31),
  inference(superposition,[],[f414,f239])).
fof(f1530,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK13),s('type.num.num',X0),s('type.num.num',X1))))) ) | $spl31),
  inference(superposition,[],[f414,f237])).
fof(f1527,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X7),s('type.num.num',X6))))) ) | $spl31),
  inference(superposition,[],[f414,f239])).
fof(f1524,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X0),s('type.num.num',X1))))) ) | $spl31),
  inference(superposition,[],[f414,f237])).
fof(f1521,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X7),s('type.num.num',X6))))) ) | $spl31),
  inference(superposition,[],[f414,f239])).
fof(f1513,plain,(
  ( ! [X14,X12,X15,X13] : (sP6(X12,X13,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl2),
  inference(superposition,[],[f978,f239])).
fof(f1492,plain,(
  ( ! [X14,X12,X15,X13] : (sP4(X12,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X13) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl2),
  inference(superposition,[],[f965,f239])).
fof(f1469,plain,(
  ( ! [X14,X12,X15,X13] : (sP0('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X12,X13) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl2),
  inference(superposition,[],[f949,f239])).
fof(f1447,plain,(
  ( ! [X14,X12,X15,X13] : (sP0(X12,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X13) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl2),
  inference(superposition,[],[f948,f239])).
fof(f1425,plain,(
  ( ! [X14,X12,X15,X13] : (sP0(X12,X13,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl2),
  inference(superposition,[],[f947,f239])).
fof(f554,plain,(
  ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')))) | $spl49),
  inference(cnf_transformation,[],[f554_D])).
fof(f554_D,plain,(
  ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')))) <=> ~$spl49),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl49])])).
fof(f1381,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X7),s('type.num.num',X6))))) ) | $spl25),
  inference(superposition,[],[f409,f239])).
fof(f1378,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X0),s('type.num.num',X1))))) ) | $spl25),
  inference(superposition,[],[f409,f237])).
fof(f1375,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK13))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK13))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK13),s('type.num.num',X7),s('type.num.num',X6))))) ) | $spl25),
  inference(superposition,[],[f409,f239])).
fof(f1372,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK13),s('type.num.num',X0),s('type.num.num',X1))))) ) | $spl25),
  inference(superposition,[],[f409,f237])).
fof(f1369,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X7),s('type.num.num',X6))))) ) | $spl25),
  inference(superposition,[],[f409,f239])).
fof(f1276,plain,(
  ( ! [X14,X12,X15,X13] : (sP1(X12,X13,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl2),
  inference(superposition,[],[f946,f239])).
fof(f1255,plain,(
  ( ! [X14,X12,X15,X13] : (sP2(X12,X13,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X15),s('type.num.num',X14))))) ) | $spl2),
  inference(superposition,[],[f943,f239])).
fof(f1239,plain,(
  ( ! [X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X1))) ) | $spl60),
  inference(subsumption_resolution,[],[f1229,f348])).
fof(f1229,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X1))) ) | $spl60),
  inference(resolution,[],[f1089,f310])).
fof(f1238,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X0))) ) | $spl60),
  inference(subsumption_resolution,[],[f1228,f348])).
fof(f1228,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X0))) ) | $spl60),
  inference(resolution,[],[f1089,f314])).
fof(f1227,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | $spl58),
  inference(subsumption_resolution,[],[f1216,f348])).
fof(f1216,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | $spl58),
  inference(resolution,[],[f1072,f310])).
fof(f1225,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | $spl58),
  inference(subsumption_resolution,[],[f1215,f348])).
fof(f1215,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | $spl58),
  inference(resolution,[],[f1072,f314])).
fof(f1191,plain,(
  ( ! [X6,X7] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X7),s('type.num.num',X6))))) ) | $spl24),
  inference(superposition,[],[f410,f239])).
fof(f1188,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X0),s('type.num.num',X1))))) ) | $spl24),
  inference(superposition,[],[f410,f237])).
fof(f551,plain,(
  p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | $spl46),
  inference(cnf_transformation,[],[f551_D])).
fof(f551_D,plain,(
  p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f1185,plain,(
  ( ! [X6,X7] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK13))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK13))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK13),s('type.num.num',X7),s('type.num.num',X6))))) ) | $spl24),
  inference(superposition,[],[f410,f239])).
fof(f1182,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK13),s('type.num.num',X0),s('type.num.num',X1))))) ) | $spl24),
  inference(superposition,[],[f410,f237])).
fof(f1179,plain,(
  ( ! [X6,X7] : (p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X7),s('type.num.num',X6))))) ) | $spl24),
  inference(superposition,[],[f410,f239])).
fof(f1093,plain,(
  s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',sK14))) | s('type.num.num',sK12) = s('type.num.num',sK13) | $spl24),
  inference(resolution,[],[f410,f240])).
fof(f1086,plain,(
  ( ! [X10,X11,X9] : (sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X9) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X10),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X11),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X11),s('type.num.num',X10))))) ) | $spl2),
  inference(superposition,[],[f1012,f239])).
fof(f1069,plain,(
  ( ! [X10,X11,X9] : (sP11(X9,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X10),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X11),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X11),s('type.num.num',X10))))) ) | $spl2),
  inference(superposition,[],[f1011,f239])).
fof(f1031,plain,(
  ( ! [X6,X7] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X7),s('type.num.num',X6))))) ) | $spl2),
  inference(superposition,[],[f393,f239])).
fof(f917,plain,(
  $false | ($spl0 | $spl24)),
  inference(subsumption_resolution,[],[f410,f391])).
fof(f915,plain,(
  $false | ($spl0 | $spl31)),
  inference(subsumption_resolution,[],[f904,f391])).
fof(f904,plain,(
  p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK13)))) | $spl31),
  inference(trivial_inequality_removal,[],[f898])).
fof(f898,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK13)))) | $spl31),
  inference(superposition,[],[f414,f231])).
fof(f914,plain,(
  $false | ($spl0 | $spl31)),
  inference(subsumption_resolution,[],[f905,f391])).
fof(f905,plain,(
  p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK13)))) | $spl31),
  inference(trivial_inequality_removal,[],[f897])).
fof(f897,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK13)))) | $spl31),
  inference(superposition,[],[f414,f229])).
fof(f886,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK13))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK13))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK13),s('type.num.num',X7),s('type.num.num',X6))))) ) | $spl31),
  inference(superposition,[],[f414,f239])).
fof(f883,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK13),s('type.num.num',X0),s('type.num.num',X1))))) ) | $spl31),
  inference(superposition,[],[f414,f237])).
fof(f880,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X7),s('type.num.num',X6))))) ) | $spl31),
  inference(superposition,[],[f414,f239])).
fof(f877,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X0),s('type.num.num',X1))))) ) | $spl31),
  inference(superposition,[],[f414,f237])).
fof(f698,plain,(
  ( ! [X2,X3] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X2),s('type.num.num',sK14))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X3),s('type.num.num',sK14)))) ) | $spl50),
  inference(cnf_transformation,[],[f698_D])).
fof(f698_D,plain,(
  ( ! [X2,X3] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X2),s('type.num.num',sK14))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X3),s('type.num.num',sK14)))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f569,plain,(
  ( ! [X2,X3] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X2),s('type.num.num',sK14))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X3),s('type.num.num',sK14))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14))))) ) | $spl0),
  inference(resolution,[],[f391,f242])).
fof(f567,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X0),s('type.num.num',X1))))) ) | $spl39),
  inference(trivial_inequality_removal,[],[f560])).
fof(f560,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.|48|_0') | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X0),s('type.num.num',X1))))) ) | $spl39),
  inference(superposition,[],[f540,f237])).
fof(f520,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X7),s('type.num.num',X6))))) ) | $spl25),
  inference(superposition,[],[f409,f239])).
fof(f517,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X0),s('type.num.num',X1))))) ) | $spl25),
  inference(superposition,[],[f409,f237])).
fof(f514,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK13))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK13))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK13),s('type.num.num',X7),s('type.num.num',X6))))) ) | $spl25),
  inference(superposition,[],[f409,f239])).
fof(f511,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK13),s('type.num.num',X0),s('type.num.num',X1))))) ) | $spl25),
  inference(superposition,[],[f409,f237])).
fof(f508,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X7),s('type.num.num',X6))))) ) | $spl25),
  inference(superposition,[],[f409,f239])).
fof(f505,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X0),s('type.num.num',X1))))) ) | $spl25),
  inference(superposition,[],[f409,f237])).
fof(f422,plain,(
  s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK12),s('type.num.num',sK14))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num',sK14))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) | $spl25),
  inference(resolution,[],[f409,f242])).
fof(f421,plain,(
  s('type.num.num',sK12) != s('type.num.num',sK13) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK14) | $spl25),
  inference(resolution,[],[f409,f241])).
fof(f390,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK13))) | s(bool,'T_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',sK12)))),
  inference(superposition,[],[f199,f236])).
fof(f389,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK13))) | s(bool,'F_0') = s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',sK12)))),
  inference(superposition,[],[f199,f236])).
fof(f388,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK13))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',sK12))))),
  inference(superposition,[],[f199,f235])).
fof(f387,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK13))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',sK12))))),
  inference(superposition,[],[f199,f233])).
fof(f386,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK13))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',sK12))))),
  inference(superposition,[],[f199,f231])).
fof(f385,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK13))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',sK12))))),
  inference(superposition,[],[f199,f229])).
fof(f384,plain,(
  ( ! [X11] : (s(bool,X11) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK13))) | ~p(s(bool,X11)) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',sK12))))) )),
  inference(superposition,[],[f199,f226])).
fof(f383,plain,(
  ( ! [X10] : (s(bool,X10) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK13))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',sK12)))) | ~p(s(bool,X10))) )),
  inference(superposition,[],[f199,f226])).
fof(f382,plain,(
  ( ! [X9] : (s(bool,X9) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK13))) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',sK12)))) | p(s(bool,X9))) )),
  inference(superposition,[],[f199,f223])).
fof(f381,plain,(
  ( ! [X8] : (s(bool,X8) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num',sK13))) | p(s(bool,X8)) | p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num',sK12))))) )),
  inference(superposition,[],[f199,f223])).
fof(f374,plain,(
  ( ! [X6,X7] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK12))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X7),s('type.num.num',X6))))) )),
  inference(superposition,[],[f199,f239])).
fof(f371,plain,(
  ( ! [X0,X1] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK13),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK12),s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(superposition,[],[f199,f237])).
fof(f368,plain,(
  ( ! [X6,X7] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK13))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK13))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK13),s('type.num.num',X7),s('type.num.num',X6))))) )),
  inference(superposition,[],[f199,f239])).
fof(f365,plain,(
  ( ! [X0,X1] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',sK12),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK13),s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(superposition,[],[f199,f237])).
fof(f362,plain,(
  ( ! [X6,X7] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X7),s('type.num.num',X6))))) )),
  inference(superposition,[],[f199,f239])).
fof(f359,plain,(
  ( ! [X0,X1] : (s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12),s('type.num.num',sK13))) != s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) | ~p(s(bool,'const.arithmetic.MODEQ_3'(s('type.num.num',sK14),s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(superposition,[],[f199,f237])).
% SZS output end Proof for thm.arithmetic
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 861
% Passive clauses: 13083
% Generated clauses: 42016
% Final active clauses: 673
% Final passive clauses: 7902
% Input formulas: 36
% Initial clauses: 156
% 
% Selected by SInE selection: 35
% SInE iterations: 1
% 
% Duplicate literals: 918
% Trivial inequalities: 233
% Fw subsumption resolutions: 6503
% Fw demodulations: 5761
% Bw demodulations: 32
% 
% Simple tautologies: 1845
% Equational tautologies: 11
% Forward subsumptions: 13287
% 
% Binary resolution: 9725
% Forward superposition: 13898
% Backward superposition: 4095
% Equality resolution: 15
% 
% Split clauses: 727
% Split components: 1581
% SAT solver clauses: 788
% SAT solver unit clauses: 15
% SAT solver binary clauses: 89
% 
% Sat splits: 727
% Sat splitting refutations: 47
% 
% Memory used [KB]: 12409
% Time elapsed: 0.632 s
% ------------------------------
% ------------------------------
% Success in time 12.392 s
