Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.relation
% (117175)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 68
% Passive clauses: 1669
% Generated clauses: 12649
% Final active clauses: 66
% Final passive clauses: 1538
% Input formulas: 19
% Initial clauses: 31
% 
% Duplicate literals: 898
% Trivial inequalities: 171
% Fw subsumption resolutions: 1733
% Fw demodulations: 419
% 
% Simple tautologies: 291
% Equational tautologies: 99
% Forward subsumptions: 5510
% Fw demodulations to eq. taut.: 1
% 
% Binary resolution: 51
% Forward superposition: 5653
% Backward superposition: 2942
% Self superposition: 355
% Equality factoring: 283
% Equality resolution: 14
% 
% Split clauses: 674
% Split components: 1362
% SAT solver clauses: 1844
% SAT solver unit clauses: 252
% SAT solver binary clauses: 513
% 
% Sat splits: 674
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1836
% TWLsolver calls for satisfiability: 71
% 
% Memory used [KB]: 6268
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 294 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.relation
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.relation
% SZS output start Proof for thm.relation
fof(f5356,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f5355,f63,f97,f1286,f1287,f85,f109,f86,f110,f559,f564,f84,f108,f87,f111])).
fof(f111,plain,(
  ( ! [X0,X1] : (s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),X1))))) = s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),X1)))))) ) | $spl16),
  inference(cnf_transformation,[],[f111_D])).
fof(f111_D,plain,(
  ( ! [X0,X1] : (s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),X1))))) = s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),X1)))))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f87,plain,(
  ( ! [X0,X1] : (s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),X1))))) = s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),X1)))))) )),
  inference(cnf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0,X1] : s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),X1))))) = s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),X1)))))),
  inference(flattening,[],[f37])).
fof(f37,plain,(
  ! [X0] : ! [X1] : s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),X1))))) = s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),X1)))))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X6] : ! [X7] : s(fun(X6,fun(X6,bool)),'const.relation.inv_1'(s(fun(X6,fun(X6,bool)),'const.relation.TC_1'(s(fun(X6,fun(X6,bool)),X7))))) = s(fun(X6,fun(X6,bool)),'const.relation.TC_1'(s(fun(X6,fun(X6,bool)),'const.relation.inv_1'(s(fun(X6,fun(X6,bool)),X7)))))),
  file('thm.relation.inv_EQC|split|0','thm.relation.inv_TC')).
fof(f108,plain,(
  ( ! [X0,X1] : (s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1))) = s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))))) ) | $spl10),
  inference(cnf_transformation,[],[f108_D])).
fof(f108_D,plain,(
  ( ! [X0,X1] : (s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1))) = s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f84,plain,(
  ( ! [X0,X1] : (s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1))) = s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))))) )),
  inference(cnf_transformation,[],[f32])).
fof(f32,plain,(
  ! [X0,X1] : s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1))) = s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))))),
  inference(flattening,[],[f31])).
fof(f31,plain,(
  ! [X0] : ! [X1] : s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1))) = s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X6] : ! [X7] : s(fun(X6,fun(X6,bool)),'const.relation.SC_1'(s(fun(X6,fun(X6,bool)),X7))) = s(fun(X6,fun(X6,bool)),'const.relation.inv_1'(s(fun(X6,fun(X6,bool)),'const.relation.SC_1'(s(fun(X6,fun(X6,bool)),X7)))))),
  file('thm.relation.inv_EQC|split|0','thm.relation.inv_SC|split|0')).
fof(f564,plain,(
  ( ! [X0,X1] : (s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1))))) = s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))))))) ) | $spl92),
  inference(cnf_transformation,[],[f564_D])).
fof(f564_D,plain,(
  ( ! [X0,X1] : (s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1))))) = s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))))))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f559,plain,(
  ( ! [X0,X1] : (s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1))))) = s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))))))) ) | ($spl10 | $spl16)),
  inference(superposition,[],[f111,f108])).
fof(f110,plain,(
  ( ! [X0,X1] : (s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),X1))))) = s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),X1)))))) ) | $spl14),
  inference(cnf_transformation,[],[f110_D])).
fof(f110_D,plain,(
  ( ! [X0,X1] : (s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),X1))))) = s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),X1)))))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f86,plain,(
  ( ! [X0,X1] : (s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),X1))))) = s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),X1)))))) )),
  inference(cnf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0,X1] : s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),X1))))) = s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),X1)))))),
  inference(flattening,[],[f35])).
fof(f35,plain,(
  ! [X0] : ! [X1] : s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),X1))))) = s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),X1)))))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X6] : ! [X7] : s(fun(X6,fun(X6,bool)),'const.relation.inv_1'(s(fun(X6,fun(X6,bool)),'const.relation.RC_1'(s(fun(X6,fun(X6,bool)),X7))))) = s(fun(X6,fun(X6,bool)),'const.relation.RC_1'(s(fun(X6,fun(X6,bool)),'const.relation.inv_1'(s(fun(X6,fun(X6,bool)),X7)))))),
  file('thm.relation.inv_EQC|split|0','thm.relation.inv_RC')).
fof(f109,plain,(
  ( ! [X0,X1] : (s(fun(X0,fun(X0,bool)),'const.relation.EQC_1'(s(fun(X0,fun(X0,bool)),X1))) = s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))))))) ) | $spl12),
  inference(cnf_transformation,[],[f109_D])).
fof(f109_D,plain,(
  ( ! [X0,X1] : (s(fun(X0,fun(X0,bool)),'const.relation.EQC_1'(s(fun(X0,fun(X0,bool)),X1))) = s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))))))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f85,plain,(
  ( ! [X0,X1] : (s(fun(X0,fun(X0,bool)),'const.relation.EQC_1'(s(fun(X0,fun(X0,bool)),X1))) = s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))))))) )),
  inference(cnf_transformation,[],[f34])).
fof(f34,plain,(
  ! [X0,X1] : s(fun(X0,fun(X0,bool)),'const.relation.EQC_1'(s(fun(X0,fun(X0,bool)),X1))) = s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))))))),
  inference(flattening,[],[f33])).
fof(f33,plain,(
  ! [X0] : ! [X1] : s(fun(X0,fun(X0,bool)),'const.relation.EQC_1'(s(fun(X0,fun(X0,bool)),X1))) = s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))))))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X6] : ! [X7] : s(fun(X6,fun(X6,bool)),'const.relation.EQC_1'(s(fun(X6,fun(X6,bool)),X7))) = s(fun(X6,fun(X6,bool)),'const.relation.RC_1'(s(fun(X6,fun(X6,bool)),'const.relation.TC_1'(s(fun(X6,fun(X6,bool)),'const.relation.SC_1'(s(fun(X6,fun(X6,bool)),X7)))))))),
  file('thm.relation.inv_EQC|split|0','thm.relation.EQC_DEF')).
fof(f1287,plain,(
  ( ! [X4,X5] : (s(fun(X4,fun(X4,bool)),'const.relation.EQC_1'(s(fun(X4,fun(X4,bool)),X5))) = s(fun(X4,fun(X4,bool)),'const.relation.inv_1'(s(fun(X4,fun(X4,bool)),'const.relation.EQC_1'(s(fun(X4,fun(X4,bool)),X5)))))) ) | $spl162),
  inference(cnf_transformation,[],[f1287_D])).
fof(f1287_D,plain,(
  ( ! [X4,X5] : (s(fun(X4,fun(X4,bool)),'const.relation.EQC_1'(s(fun(X4,fun(X4,bool)),X5))) = s(fun(X4,fun(X4,bool)),'const.relation.inv_1'(s(fun(X4,fun(X4,bool)),'const.relation.EQC_1'(s(fun(X4,fun(X4,bool)),X5)))))) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f1286,plain,(
  ( ! [X4,X5] : (s(fun(X4,fun(X4,bool)),'const.relation.EQC_1'(s(fun(X4,fun(X4,bool)),X5))) = s(fun(X4,fun(X4,bool)),'const.relation.inv_1'(s(fun(X4,fun(X4,bool)),'const.relation.EQC_1'(s(fun(X4,fun(X4,bool)),X5)))))) ) | ($spl12 | $spl14 | $spl92)),
  inference(forward_demodulation,[],[f1274,f109])).
fof(f1274,plain,(
  ( ! [X4,X5] : (s(fun(X4,fun(X4,bool)),'const.relation.RC_1'(s(fun(X4,fun(X4,bool)),'const.relation.TC_1'(s(fun(X4,fun(X4,bool)),'const.relation.SC_1'(s(fun(X4,fun(X4,bool)),X5))))))) = s(fun(X4,fun(X4,bool)),'const.relation.inv_1'(s(fun(X4,fun(X4,bool)),'const.relation.RC_1'(s(fun(X4,fun(X4,bool)),'const.relation.TC_1'(s(fun(X4,fun(X4,bool)),'const.relation.SC_1'(s(fun(X4,fun(X4,bool)),X5)))))))))) ) | ($spl14 | $spl92)),
  inference(superposition,[],[f110,f564])).
fof(f97,plain,(
  s(fun(sK0,fun(sK0,bool)),'const.relation.EQC_1'(s(fun(sK0,fun(sK0,bool)),sK1))) != s(fun(sK0,fun(sK0,bool)),'const.relation.inv_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.EQC_1'(s(fun(sK0,fun(sK0,bool)),sK1))))) | $spl1),
  inference(cnf_transformation,[],[f97_D])).
fof(f97_D,plain,(
  s(fun(sK0,fun(sK0,bool)),'const.relation.EQC_1'(s(fun(sK0,fun(sK0,bool)),sK1))) != s(fun(sK0,fun(sK0,bool)),'const.relation.inv_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.EQC_1'(s(fun(sK0,fun(sK0,bool)),sK1))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f63,plain,(
  s(fun(sK0,fun(sK0,bool)),'const.relation.EQC_1'(s(fun(sK0,fun(sK0,bool)),sK1))) != s(fun(sK0,fun(sK0,bool)),'const.relation.inv_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.EQC_1'(s(fun(sK0,fun(sK0,bool)),sK1)))))),
  inference(cnf_transformation,[],[f48])).
fof(f48,plain,(
  s(fun(sK0,fun(sK0,bool)),'const.relation.EQC_1'(s(fun(sK0,fun(sK0,bool)),sK1))) != s(fun(sK0,fun(sK0,bool)),'const.relation.inv_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.EQC_1'(s(fun(sK0,fun(sK0,bool)),sK1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f46])).
fof(f46,plain,(
  ? [X0,X1] : s(fun(X0,fun(X0,bool)),'const.relation.EQC_1'(s(fun(X0,fun(X0,bool)),X1))) != s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),'const.relation.EQC_1'(s(fun(X0,fun(X0,bool)),X1)))))),
  inference(ennf_transformation,[],[f22])).
fof(f22,plain,(
  ~! [X0,X1] : s(fun(X0,fun(X0,bool)),'const.relation.EQC_1'(s(fun(X0,fun(X0,bool)),X1))) = s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),'const.relation.EQC_1'(s(fun(X0,fun(X0,bool)),X1)))))),
  inference(flattening,[],[f21])).
fof(f21,plain,(
  ~! [X0] : ! [X1] : s(fun(X0,fun(X0,bool)),'const.relation.EQC_1'(s(fun(X0,fun(X0,bool)),X1))) = s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),'const.relation.EQC_1'(s(fun(X0,fun(X0,bool)),X1)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7] : s(fun(X6,fun(X6,bool)),'const.relation.EQC_1'(s(fun(X6,fun(X6,bool)),X7))) = s(fun(X6,fun(X6,bool)),'const.relation.inv_1'(s(fun(X6,fun(X6,bool)),'const.relation.EQC_1'(s(fun(X6,fun(X6,bool)),X7)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7] : s(fun(X6,fun(X6,bool)),'const.relation.EQC_1'(s(fun(X6,fun(X6,bool)),X7))) = s(fun(X6,fun(X6,bool)),'const.relation.inv_1'(s(fun(X6,fun(X6,bool)),'const.relation.EQC_1'(s(fun(X6,fun(X6,bool)),X7)))))),
  file('thm.relation.inv_EQC|split|0','thm.relation.inv_EQC|split|0')).
fof(f5355,plain,(
  $false | ($spl1 | $spl162)),
  inference(trivial_inequality_removal,[],[f5338])).
fof(f5338,plain,(
  s(fun(sK0,fun(sK0,bool)),'const.relation.EQC_1'(s(fun(sK0,fun(sK0,bool)),sK1))) != s(fun(sK0,fun(sK0,bool)),'const.relation.EQC_1'(s(fun(sK0,fun(sK0,bool)),sK1))) | ($spl1 | $spl162)),
  inference(superposition,[],[f97,f1287])).
% SZS output end Proof for thm.relation
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 53
% Passive clauses: 169
% Generated clauses: 5143
% Final active clauses: 47
% Final passive clauses: 98
% Input formulas: 19
% Initial clauses: 31
% 
% Duplicate literals: 816
% Trivial inequalities: 74
% Fw subsumption resolutions: 229
% Fw demodulations: 78
% Bw demodulations: 15
% Condensations: 1168
% Global subsumptions: 4
% 
% Simple tautologies: 134
% Equational tautologies: 39
% Forward subsumptions: 1896
% Fw demodulations to eq. taut.: 1
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 4
% Forward superposition: 1092
% Backward superposition: 1199
% Self superposition: 161
% Equality factoring: 124
% Equality resolution: 3
% 
% Split clauses: 68
% Split components: 143
% SAT solver clauses: 616
% SAT solver unit clauses: 40
% SAT solver binary clauses: 77
% SAT solver learnt clauses: 9
% 
% Sat splits: 68
% Sat splitting refutations: 11
% 
% TWLsolver clauses: 599
% TWLsolver calls for satisfiability: 86
% 
% Memory used [KB]: 4093
% Time elapsed: 0.376 s
% ------------------------------
% ------------------------------
% Success in time 0.989 s
