Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.relation
% (51531)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 210
% Passive clauses: 3722
% Generated clauses: 20410
% Final active clauses: 200
% Final passive clauses: 3265
% Input formulas: 27
% Initial clauses: 71
% 
% Duplicate literals: 993
% Trivial inequalities: 327
% Fw subsumption resolutions: 3195
% Fw demodulations: 2850
% Bw demodulations: 6
% 
% Simple tautologies: 355
% Equational tautologies: 18
% Forward subsumptions: 7578
% 
% Binary resolution: 1228
% Forward superposition: 6996
% Backward superposition: 4362
% Self superposition: 72
% Equality factoring: 32
% Equality resolution: 51
% 
% Split clauses: 861
% Split components: 1730
% SAT solver clauses: 1617
% SAT solver unit clauses: 13
% SAT solver binary clauses: 140
% 
% Sat splits: 861
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 1616
% TWLsolver calls for satisfiability: 328
% 
% Memory used [KB]: 11385
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 294 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.relation
% (51536)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 116
% Passive clauses: 1280
% Generated clauses: 11698
% Final active clauses: 101
% Final passive clauses: 686
% Input formulas: 27
% Initial clauses: 71
% 
% Duplicate literals: 375
% Trivial inequalities: 159
% Fw subsumption resolutions: 773
% Fw demodulations: 1910
% Bw demodulations: 339
% Condensations: 462
% Global subsumptions: 8
% 
% Simple tautologies: 166
% Equational tautologies: 19
% Forward subsumptions: 4503
% Bw demodulations to eq. taut.: 152
% 
% Binary resolution: 389
% Factoring: 2
% Forward superposition: 3525
% Backward superposition: 2317
% Self superposition: 77
% Equality factoring: 48
% Equality resolution: 16
% 
% Split clauses: 130
% Split components: 269
% SAT solver clauses: 2050
% SAT solver unit clauses: 76
% SAT solver binary clauses: 44
% SAT solver learnt clauses: 16
% 
% Sat splits: 130
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 2018
% TWLsolver calls for satisfiability: 343
% 
% Memory used [KB]: 11641
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.relation
% (51546)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 557
% Passive clauses: 7838
% Generated clauses: 170360
% Final active clauses: 64
% Final passive clauses: 887
% Input formulas: 27
% Initial clauses: 71
% 
% Duplicate literals: 5441
% Trivial inequalities: 829
% Fw subsumption resolutions: 13515
% Fw demodulations: 28973
% Bw demodulations: 5334
% Condensations: 11973
% Global subsumptions: 486
% 
% Simple tautologies: 998
% Equational tautologies: 1842
% Forward subsumptions: 95654
% Fw demodulations to eq. taut.: 149
% Bw demodulations to eq. taut.: 319
% 
% Unit resulting resolution: 9407
% Forward superposition: 33041
% Backward superposition: 23266
% Self superposition: 636
% Equality factoring: 325
% Equality resolution: 75
% 
% Split clauses: 539
% Split components: 1093
% SAT solver clauses: 10100
% SAT solver unit clauses: 82
% SAT solver binary clauses: 598
% 
% Sat splits: 539
% Sat splitting refutations: 696
% 
% Memory used [KB]: 77013
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.relation
% (51922)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 697
% Passive clauses: 30080
% Generated clauses: 131793
% Final active clauses: 125
% Final passive clauses: 17609
% Input formulas: 27
% Initial clauses: 71
% 
% Selected by SInE selection: 27
% SInE iterations: 1
% 
% Duplicate literals: 3942
% Trivial inequalities: 1579
% Fw subsumption resolutions: 21171
% Fw demodulations: 14931
% Bw demodulations: 335
% 
% Simple tautologies: 2939
% Equational tautologies: 98
% Forward subsumptions: 48455
% Bw demodulations to eq. taut.: 7
% 
% Binary resolution: 11021
% Factoring: 11
% Forward superposition: 37813
% Backward superposition: 34389
% Self superposition: 499
% Equality factoring: 216
% Equality resolution: 74
% 
% Split clauses: 9242
% Split components: 19675
% SAT solver clauses: 9539
% SAT solver unit clauses: 27
% SAT solver binary clauses: 2707
% 
% Sat splits: 9242
% Sat splitting refutations: 125
% 
% Memory used [KB]: 36715
% Time elapsed: 4.0000 s
% ------------------------------
% ------------------------------
% remaining time: 142 next slice time: 6
lrs+1002_2:3_br=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:sas=minisat:stl=30:sac=on:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none:sp=reverse_arity:urr=on_4 on thm.relation
% (52038)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 66
% Passive clauses: 1696
% Generated clauses: 10774
% Final active clauses: 34
% Final passive clauses: 462
% Input formulas: 27
% Initial clauses: 71
% 
% Duplicate literals: 318
% Trivial inequalities: 251
% Fw subsumption resolutions: 1623
% Fw demodulations: 706
% Bw demodulations: 1162
% Global subsumptions: 109
% 
% Simple tautologies: 373
% Equational tautologies: 62
% Forward subsumptions: 5134
% Bw demodulations to eq. taut.: 4
% 
% Unit resulting resolution: 334
% Forward superposition: 2901
% Backward superposition: 1987
% Self superposition: 74
% Equality factoring: 50
% Equality resolution: 10
% 
% Split clauses: 17
% Split components: 35
% SAT solver clauses: 2040
% SAT solver unit clauses: 157
% SAT solver binary clauses: 133
% 
% Sat splits: 17
% Sat splitting refutations: 216
% 
% Memory used [KB]: 7291
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 135 next slice time: 15
lrs+1011_12_bs=on:bsr=unit_only:cond=on:gs=on:gsaa=from_current:gsssp=full:nwc=1.1:sas=minisat:stl=60:sos=all:sac=on:sdd=off:sser=off:sfr=on:ssfp=100000:ssfq=1.2:ssnc=none:sp=reverse_arity:updr=off_14 on thm.relation
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.relation
% SZS output start Proof for thm.relation
fof(f1121,plain,(
  $false),
  inference(subsumption_resolution,[],[f1120,f187])).
fof(f187,plain,(
  p(s(bool,'T_0'))),
  inference(condensation,[],[f148])).
fof(f148,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f95])).
fof(f95,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))))),
  inference(flattening,[],[f94])).
fof(f94,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f65])).
fof(f65,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f45])).
fof(f45,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X9] : ((p(s(bool,X9)) => p(s(bool,X9))) <=> p(s(bool,'T_0')))),
  file('thm.relation.antisymmetric_RINTER|split|1','thm.bool.IMP_CLAUSES|split|3')).
fof(f1120,plain,(
  ~p(s(bool,'T_0'))),
  inference(forward_demodulation,[],[f1119,f1019])).
fof(f1019,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))),
  inference(resolution,[],[f900,f120])).
fof(f120,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f32])).
fof(f32,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X9] : (s(bool,'T_0') = s(bool,X9) <=> p(s(bool,X9)))),
  file('thm.relation.antisymmetric_RINTER|split|1','thm.bool.EQ_CLAUSES|split|0')).
fof(f900,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))),
  inference(resolution,[],[f281,f168])).
fof(f168,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X4))),s(X1,X5)))) | ~p(s(bool,'const.relation.RINTER_4'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3),s(X0,X4),s(X1,X5))))) )),
  inference(cnf_transformation,[],[f108])).
fof(f108,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((~p(s(bool,'const.relation.RINTER_4'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3),s(X0,X4),s(X1,X5)))) | (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X4))),s(X1,X5)))))) & (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X4))),s(X1,X5)))) | p(s(bool,'const.relation.RINTER_4'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3),s(X0,X4),s(X1,X5))))))),
  inference(flattening,[],[f107])).
fof(f107,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((~p(s(bool,'const.relation.RINTER_4'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3),s(X0,X4),s(X1,X5)))) | (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X4))),s(X1,X5)))))) & ((~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X4))),s(X1,X5))))) | p(s(bool,'const.relation.RINTER_4'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3),s(X0,X4),s(X1,X5))))))),
  inference(nnf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (p(s(bool,'const.relation.RINTER_4'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3),s(X0,X4),s(X1,X5)))) <=> (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X4))),s(X1,X5))))))),
  inference(flattening,[],[f58])).
fof(f58,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5] : (p(s(bool,'const.relation.RINTER_4'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3),s(X0,X4),s(X1,X5)))) <=> (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X4))),s(X1,X5))))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X6,X18] : ! [X8,X7,X10,X11] : (p(s(bool,'const.relation.RINTER_4'(s(fun(X6,fun(X18,bool)),X8),s(fun(X6,fun(X18,bool)),X7),s(X6,X10),s(X18,X11)))) <=> (p(s(bool,happ(s(fun(X18,bool),happ(s(fun(X6,fun(X18,bool)),X8),s(X6,X10))),s(X18,X11)))) & p(s(bool,happ(s(fun(X18,bool),happ(s(fun(X6,fun(X18,bool)),X7),s(X6,X10))),s(X18,X11))))))),
  file('thm.relation.antisymmetric_RINTER|split|1','thm.relation.RINTER')).
fof(f281,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))),
  inference(forward_demodulation,[],[f245,f166])).
fof(f166,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'const.relation.RINTER_4'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3),s(X0,X4),s(X1,X5))) = s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),'const.relation.RINTER_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3))),s(X0,X4))),s(X1,X5)))) )),
  inference(cnf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(bool,'const.relation.RINTER_4'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3),s(X0,X4),s(X1,X5))) = s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),'const.relation.RINTER_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3))),s(X0,X4))),s(X1,X5)))),
  inference(flattening,[],[f56])).
fof(f56,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5] : s(bool,'const.relation.RINTER_4'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3),s(X0,X4),s(X1,X5))) = s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),'const.relation.RINTER_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3))),s(X0,X4))),s(X1,X5)))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X6,X18] : ! [X19,X20,X21,X22] : s(bool,'const.relation.RINTER_4'(s(fun(X6,fun(X18,bool)),X19),s(fun(X6,fun(X18,bool)),X20),s(X6,X21),s(X18,X22))) = s(bool,happ(s(fun(X18,bool),happ(s(fun(X6,fun(X18,bool)),'const.relation.RINTER_2'(s(fun(X6,fun(X18,bool)),X19),s(fun(X6,fun(X18,bool)),X20))),s(X6,X21))),s(X18,X22)))),
  file('thm.relation.antisymmetric_RINTER|split|1','ARITY_2_4_const.relation.RINTER')).
fof(f245,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1))),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))),
  inference(resolution,[],[f114,f158])).
fof(f158,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK3(X1,X0)))),s(X0,sK4(X1,X0)))))) )),
  inference(cnf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X4,X5] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X4)))) | s(X0,X4) = s(X0,X5))) & ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK3(X1,X0)))),s(X0,sK4(X1,X0))))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK4(X1,X0)))),s(X0,sK3(X1,X0))))) & s(X0,sK3(X1,X0)) != s(X0,sK4(X1,X0))) | p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3,sK4])],[f101])).
fof(f101,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X4,X5] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X4)))) | s(X0,X4) = s(X0,X5))) & (? [X2,X3] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2)))) & s(X0,X2) != s(X0,X3)) | p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(rectify,[],[f100])).
fof(f100,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X2,X3] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2)))) | s(X0,X2) = s(X0,X3))) & (? [X2,X3] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2)))) & s(X0,X2) != s(X0,X3)) | p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(nnf_transformation,[],[f69])).
fof(f69,plain,(
  ! [X0,X1] : (p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> ! [X2,X3] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2)))) | s(X0,X2) = s(X0,X3)))),
  inference(flattening,[],[f68])).
fof(f68,plain,(
  ! [X0,X1] : (p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> ! [X2,X3] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2))))) | s(X0,X2) = s(X0,X3)))),
  inference(ennf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0,X1] : (p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> ! [X2,X3] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2))))) => s(X0,X2) = s(X0,X3)))),
  inference(flattening,[],[f49])).
fof(f49,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> ! [X2,X3] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2))))) => s(X0,X2) = s(X0,X3)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X6] : ! [X17] : (p(s(bool,'const.relation.antisymmetric_1'(s(fun(X6,fun(X6,bool)),X17)))) <=> ! [X10,X11] : ((p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X17),s(X6,X10))),s(X6,X11)))) & p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X17),s(X6,X11))),s(X6,X10))))) => s(X6,X10) = s(X6,X11)))),
  file('thm.relation.antisymmetric_RINTER|split|1','thm.relation.antisymmetric_def')).
fof(f114,plain,(
  ~p(s(bool,'const.relation.antisymmetric_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1))))))),
  inference(cnf_transformation,[],[f75])).
fof(f75,plain,(
  p(s(bool,'const.relation.antisymmetric_1'(s(fun(sK0,fun(sK0,bool)),sK1)))) & ~p(s(bool,'const.relation.antisymmetric_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f62])).
fof(f62,plain,(
  ? [X0,X1,X2] : (p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) & ~p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.RINTER_2'(s(fun(X0,fun(X0,bool)),X2),s(fun(X0,fun(X0,bool)),X1)))))))),
  inference(ennf_transformation,[],[f30])).
fof(f30,plain,(
  ~! [X0,X1,X2] : (p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) => p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.RINTER_2'(s(fun(X0,fun(X0,bool)),X2),s(fun(X0,fun(X0,bool)),X1)))))))),
  inference(flattening,[],[f29])).
fof(f29,plain,(
  ~! [X0] : ! [X1,X2] : (p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) => p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.RINTER_2'(s(fun(X0,fun(X0,bool)),X2),s(fun(X0,fun(X0,bool)),X1)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : (p(s(bool,'const.relation.antisymmetric_1'(s(fun(X6,fun(X6,bool)),X7)))) => p(s(bool,'const.relation.antisymmetric_1'(s(fun(X6,fun(X6,bool)),'const.relation.RINTER_2'(s(fun(X6,fun(X6,bool)),X8),s(fun(X6,fun(X6,bool)),X7)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : (p(s(bool,'const.relation.antisymmetric_1'(s(fun(X6,fun(X6,bool)),X7)))) => p(s(bool,'const.relation.antisymmetric_1'(s(fun(X6,fun(X6,bool)),'const.relation.RINTER_2'(s(fun(X6,fun(X6,bool)),X8),s(fun(X6,fun(X6,bool)),X7)))))))),
  file('thm.relation.antisymmetric_RINTER|split|1','thm.relation.antisymmetric_RINTER|split|1')).
fof(f1119,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))),
  inference(subsumption_resolution,[],[f1082,f247])).
fof(f247,plain,(
  s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)) != s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),
  inference(resolution,[],[f114,f160])).
fof(f160,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) | s(X0,sK3(X1,X0)) != s(X0,sK4(X1,X0))) )),
  inference(cnf_transformation,[],[f102])).
fof(f1082,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)) = s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),
  inference(resolution,[],[f282,f700])).
fof(f700,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),X5),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X4),s(sK0,X3)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X3))),s(sK0,X4)))) | s(sK0,X3) = s(sK0,X4)) )),
  inference(resolution,[],[f201,f168])).
fof(f201,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1))),s(sK0,X0)))) | s(sK0,X0) = s(sK0,X1) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,X1))))) )),
  inference(resolution,[],[f113,f157])).
fof(f157,plain,(
  ( ! [X4,X0,X5,X1] : (s(X0,X4) = s(X0,X5) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X5)))) | ~p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(cnf_transformation,[],[f102])).
fof(f113,plain,(
  p(s(bool,'const.relation.antisymmetric_1'(s(fun(sK0,fun(sK0,bool)),sK1))))),
  inference(cnf_transformation,[],[f75])).
fof(f282,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))),
  inference(forward_demodulation,[],[f246,f166])).
fof(f246,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))),
  inference(resolution,[],[f114,f159])).
fof(f159,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK4(X1,X0)))),s(X0,sK3(X1,X0)))))) )),
  inference(cnf_transformation,[],[f102])).
% SZS output end Proof for thm.relation
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 63
% Passive clauses: 182
% Generated clauses: 922
% Final active clauses: 58
% Final passive clauses: 157
% Input formulas: 27
% Initial clauses: 71
% 
% Duplicate literals: 2
% Trivial inequalities: 12
% Fw subsumption resolutions: 60
% Fw demodulations: 196
% Bw demodulations: 6
% Condensations: 17
% 
% Simple tautologies: 40
% Forward subsumptions: 423
% 
% Binary resolution: 165
% Forward superposition: 334
% Backward superposition: 145
% 
% SAT solver clauses: 184
% SAT solver unit clauses: 16
% SAT solver binary clauses: 47
% 
% Memory used [KB]: 2174
% Time elapsed: 0.100 s
% ------------------------------
% ------------------------------
% Success in time 16.55 s
