Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.measure
% (49550)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 232
% Passive clauses: 2799
% Generated clauses: 18453
% Final active clauses: 222
% Final passive clauses: 2336
% Input formulas: 32
% Initial clauses: 85
% 
% Duplicate literals: 945
% Trivial inequalities: 256
% Fw subsumption resolutions: 2987
% Fw demodulations: 2321
% Bw demodulations: 6
% 
% Simple tautologies: 350
% Equational tautologies: 18
% Forward subsumptions: 7283
% 
% Binary resolution: 1715
% Forward superposition: 6702
% Backward superposition: 3072
% Self superposition: 72
% Equality factoring: 32
% Equality resolution: 47
% 
% Split clauses: 832
% Split components: 1672
% SAT solver clauses: 1704
% SAT solver unit clauses: 15
% SAT solver binary clauses: 119
% 
% Sat splits: 832
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 1703
% TWLsolver calls for satisfiability: 313
% 
% Memory used [KB]: 10746
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.measure
% (49600)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 181
% Passive clauses: 714
% Generated clauses: 12632
% Final active clauses: 164
% Final passive clauses: 475
% Input formulas: 32
% Initial clauses: 85
% 
% Duplicate literals: 657
% Trivial inequalities: 125
% Fw subsumption resolutions: 1304
% Fw demodulations: 1535
% Bw demodulations: 71
% Condensations: 976
% Global subsumptions: 8
% 
% Simple tautologies: 257
% Equational tautologies: 23
% Forward subsumptions: 4894
% 
% Binary resolution: 431
% Factoring: 1
% Forward superposition: 3705
% Backward superposition: 2972
% Self superposition: 76
% Equality factoring: 52
% Equality resolution: 30
% 
% Split clauses: 127
% Split components: 259
% SAT solver clauses: 1312
% SAT solver unit clauses: 92
% SAT solver binary clauses: 38
% SAT solver learnt clauses: 16
% 
% Sat splits: 127
% Sat splitting refutations: 7
% 
% TWLsolver clauses: 1280
% TWLsolver calls for satisfiability: 680
% 
% Memory used [KB]: 11897
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.measure
% (49634)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 496
% Passive clauses: 3610
% Generated clauses: 158908
% Final active clauses: 88
% Final passive clauses: 291
% Input formulas: 32
% Initial clauses: 85
% 
% Duplicate literals: 7216
% Trivial inequalities: 1084
% Fw subsumption resolutions: 18282
% Fw demodulations: 13401
% Bw demodulations: 1738
% Condensations: 20393
% Global subsumptions: 978
% 
% Simple tautologies: 1495
% Equational tautologies: 2657
% Forward subsumptions: 88558
% Fw demodulations to eq. taut.: 2
% Bw demodulations to eq. taut.: 94
% 
% Unit resulting resolution: 10893
% Forward superposition: 31984
% Backward superposition: 35801
% Self superposition: 966
% Equality factoring: 557
% Equality resolution: 116
% 
% Split clauses: 419
% Split components: 850
% SAT solver clauses: 6341
% SAT solver unit clauses: 77
% SAT solver binary clauses: 592
% 
% Sat splits: 419
% Sat splitting refutations: 913
% 
% Memory used [KB]: 67418
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.measure
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.measure
% SZS output start Proof for thm.measure
fof(f1854,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f219,f220,f221,f222,f223,f248,f247,f224,f226,f237,f245,f283,f277,f327,f371,f373,f370,f369,f408,f409,f374,f407,f406,f372,f515,f532,f517,f541,f554,f543,f563,f519,f576,f521,f565,f619,f621,f622,f624,f626,f637,f653,f655,f656,f657,f660,f670,f686,f688,f689,f691,f693,f704,f891,f893,f894,f896,f898,f909,f924,f926,f927,f928,f931,f941,f678,f1183,f1182,f1162,f1161,f1202,f1203,f1204,f1205,f1206,f1207,f1213,f1215,f1216,f1217,f1220,f1230,f1259,f1261,f1262,f1264,f1266,f1277,f646,f1294,f1296,f1297,f1299,f1301,f1312,f917,f1340,f1341,f1349,f1372,f1371,f1370,f1374,f1377,f1379,f1380,f1381,f520,f1585,f1584,f1583,f1587,f1590,f1592,f1593,f1596,f1595,f1599,f1601,f1602,f1603,f1606,f1616,f1237,f1630,f1632,f1633,f246,f1634,f250,f1637,f249,f1647,f251,f1654,f1848])).
fof(f1848,plain,(
  $false | ($spl18 | $spl24 | $spl27)),
  inference(subsumption_resolution,[],[f1847,f189])).
fof(f189,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.measure.POW_ALGEBRA',aHL_TRUTH)).
fof(f1847,plain,(
  ~p(s(bool,'T_0')) | ($spl18 | $spl24 | $spl27)),
  inference(forward_demodulation,[],[f1846,f1614])).
fof(f1614,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2))) | $spl18),
  inference(resolution,[],[f517,f184])).
fof(f184,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X8] : (s(bool,'T_0') = s(bool,X8) <=> p(s(bool,X8)))),
  file('thm.measure.POW_ALGEBRA','thm.bool.EQ_CLAUSES|split|0')).
fof(f1846,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) | ($spl24 | $spl27)),
  inference(subsumption_resolution,[],[f1816,f543])).
fof(f1816,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK5('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) | $spl27),
  inference(resolution,[],[f565,f210])).
fof(f210,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X3)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X3)))) | (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X3)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))))) & (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X3)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X3))))))),
  inference(flattening,[],[f127])).
fof(f127,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X3)))) | (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X3)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))))) & ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X3)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X3))))))),
  inference(nnf_transformation,[],[f68])).
fof(f68,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X3)))) <=> (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X3)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))))),
  inference(flattening,[],[f67])).
fof(f67,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X3)))) <=> (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X3)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X6] : ! [X18,X8,X19] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),X18),s(fun(X6,bool),X8))),s(fun(X6,bool),X19)))) <=> (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X18),s(fun(X6,bool),X19)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X19))))))),
  file('thm.measure.POW_ALGEBRA','thm.pred_set.UNION_SUBSET')).
fof(f1654,plain,(
  ( ! [X3] : (p(s(bool,X3)) | s(bool,X3) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK5('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) ) | $spl46),
  inference(cnf_transformation,[],[f1654_D])).
fof(f1654_D,plain,(
  ( ! [X3] : (p(s(bool,X3)) | s(bool,X3) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK5('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f251,plain,(
  ( ! [X15] : (p(s(bool,X15))) ) | $spl12),
  inference(cnf_transformation,[],[f251_D])).
fof(f251_D,plain,(
  ( ! [X15] : (p(s(bool,X15))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f1647,plain,(
  ( ! [X26,X27] : (p(s(bool,X26)) | p(s(bool,X27)) | s(bool,X27) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK5('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) ) | $spl24),
  inference(resolution,[],[f543,f216])).
fof(f216,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f137])).
fof(f137,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f92])).
fof(f92,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f91])).
fof(f91,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f76])).
fof(f76,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X9,X14,X10,X15] : ((s(bool,X9) = s(bool,X14) & (p(s(bool,X14)) => s(bool,X10) = s(bool,X15))) => ((p(s(bool,X9)) => p(s(bool,X10))) <=> (p(s(bool,X14)) => p(s(bool,X15)))))),
  file('thm.measure.POW_ALGEBRA','thm.bool.IMP_CONG')).
fof(f249,plain,(
  ( ! [X6,X8] : (s(bool,X6) != s(bool,X8) | ~p(s(bool,X6))) ) | $spl8),
  inference(cnf_transformation,[],[f249_D])).
fof(f249_D,plain,(
  ( ! [X6,X8] : (s(bool,X6) != s(bool,X8) | ~p(s(bool,X6))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f1637,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X23) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK5('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2))) | s(bool,X24) != s(bool,X25)) ) | $spl24),
  inference(resolution,[],[f543,f140])).
fof(f140,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f92])).
fof(f250,plain,(
  ( ! [X17] : (~p(s(bool,X17))) ) | $spl10),
  inference(cnf_transformation,[],[f250_D])).
fof(f250_D,plain,(
  ( ! [X17] : (~p(s(bool,X17))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f1634,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X16) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK5('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) ) | $spl24),
  inference(resolution,[],[f543,f139])).
fof(f139,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f92])).
fof(f246,plain,(
  ( ! [X2,X0] : (s(bool,X0) != s(bool,X2) | p(s(bool,X0))) ) | $spl2),
  inference(cnf_transformation,[],[f246_D])).
fof(f246_D,plain,(
  ( ! [X2,X0] : (s(bool,X0) != s(bool,X2) | p(s(bool,X0))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f1633,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK5('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) ) | $spl24),
  inference(resolution,[],[f543,f138])).
fof(f138,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f92])).
fof(f1632,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X8) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK5('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2))) | s(bool,X9) != s(bool,X10)) ) | $spl24),
  inference(resolution,[],[f543,f136])).
fof(f136,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f92])).
fof(f1630,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK5('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) ) | $spl24),
  inference(resolution,[],[f543,f134])).
fof(f134,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f92])).
fof(f1237,plain,(
  ( ! [X3] : (p(s(bool,X3)) | s(bool,X3) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) ) | $spl34),
  inference(cnf_transformation,[],[f1237_D])).
fof(f1237_D,plain,(
  ( ! [X3] : (p(s(bool,X3)) | s(bool,X3) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f1616,plain,(
  ( ! [X26,X27] : (p(s(bool,X26)) | p(s(bool,X27)) | s(bool,X27) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) ) | $spl18),
  inference(resolution,[],[f517,f216])).
fof(f1606,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X23) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2))) | s(bool,X24) != s(bool,X25)) ) | $spl18),
  inference(resolution,[],[f517,f140])).
fof(f1603,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X16) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) ) | $spl18),
  inference(resolution,[],[f517,f139])).
fof(f1602,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) ) | $spl18),
  inference(resolution,[],[f517,f138])).
fof(f1601,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X8) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2))) | s(bool,X9) != s(bool,X10)) ) | $spl18),
  inference(resolution,[],[f517,f136])).
fof(f1599,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) ) | $spl18),
  inference(resolution,[],[f517,f134])).
fof(f1595,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK6('const.pred_set.POW_1'(s(fun(sK1,bool),sK2)),sK2,sK1)),s(fun(sK1,bool),sK2)))) | $spl45),
  inference(cnf_transformation,[],[f1595_D])).
fof(f1595_D,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK6('const.pred_set.POW_1'(s(fun(sK1,bool),sK2)),sK2,sK1)),s(fun(sK1,bool),sK2)))) <=> ~$spl45),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl45])])).
fof(f1596,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK6('const.pred_set.POW_1'(s(fun(sK1,bool),sK2)),sK2,sK1)),s(fun(sK1,bool),sK2))))) ) | $spl32),
  inference(subsumption_resolution,[],[f1579,f184])).
fof(f1579,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,X3)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK6('const.pred_set.POW_1'(s(fun(sK1,bool),sK2)),sK2,sK1)),s(fun(sK1,bool),sK2))))) ) | $spl32),
  inference(superposition,[],[f917,f184])).
fof(f1593,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK6('const.pred_set.POW_1'(s(fun(sK1,bool),sK2)),sK2,sK1)),s(fun(sK1,bool),sK2))))) ) | $spl32),
  inference(subsumption_resolution,[],[f1578,f184])).
fof(f1578,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,X2)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK6('const.pred_set.POW_1'(s(fun(sK1,bool),sK2)),sK2,sK1)),s(fun(sK1,bool),sK2))))) ) | $spl32),
  inference(superposition,[],[f917,f186])).
fof(f186,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X8] : (s(bool,'T_0') = s(bool,X8) <=> p(s(bool,X8)))),
  file('thm.measure.POW_ALGEBRA','thm.bool.EQ_CLAUSES|split|1')).
fof(f1592,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK6('const.pred_set.POW_1'(s(fun(sK1,bool),sK2)),sK2,sK1)),s(fun(sK1,bool),sK2))) | $spl42),
  inference(cnf_transformation,[],[f1592_D])).
fof(f1592_D,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK6('const.pred_set.POW_1'(s(fun(sK1,bool),sK2)),sK2,sK1)),s(fun(sK1,bool),sK2))) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f1590,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK6('const.pred_set.POW_1'(s(fun(sK1,bool),sK2)),sK2,sK1)),s(fun(sK1,bool),sK2)))) ) | $spl32),
  inference(subsumption_resolution,[],[f1576,f184])).
fof(f1576,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0)) | s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK6('const.pred_set.POW_1'(s(fun(sK1,bool),sK2)),sK2,sK1)),s(fun(sK1,bool),sK2)))) ) | $spl32),
  inference(superposition,[],[f917,f207])).
fof(f207,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.measure.POW_ALGEBRA',aHL_BOOL_CASES)).
fof(f1587,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK6('const.pred_set.POW_1'(s(fun(sK1,bool),sK2)),sK2,sK1)),s(fun(sK1,bool),sK2))) | $spl41),
  inference(cnf_transformation,[],[f1587_D])).
fof(f1587_D,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK6('const.pred_set.POW_1'(s(fun(sK1,bool),sK2)),sK2,sK1)),s(fun(sK1,bool),sK2))) <=> ~$spl41),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl41])])).
fof(f1583,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK6('const.pred_set.POW_1'(s(fun(sK1,bool),sK2)),sK2,sK1)),s(fun(sK1,bool),sK2))) | ~p(s(bool,X3))) ) | $spl32),
  inference(duplicate_literal_removal,[],[f1571])).
fof(f1571,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK6('const.pred_set.POW_1'(s(fun(sK1,bool),sK2)),sK2,sK1)),s(fun(sK1,bool),sK2))) | ~p(s(bool,X3)) | ~p(s(bool,X3))) ) | $spl32),
  inference(superposition,[],[f917,f184])).
fof(f1584,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK6('const.pred_set.POW_1'(s(fun(sK1,bool),sK2)),sK2,sK1)),s(fun(sK1,bool),sK2))) | ~p(s(bool,X2))) ) | $spl32),
  inference(duplicate_literal_removal,[],[f1570])).
fof(f1570,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK6('const.pred_set.POW_1'(s(fun(sK1,bool),sK2)),sK2,sK1)),s(fun(sK1,bool),sK2))) | ~p(s(bool,X2)) | ~p(s(bool,X2))) ) | $spl32),
  inference(superposition,[],[f917,f186])).
fof(f1585,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK6('const.pred_set.POW_1'(s(fun(sK1,bool),sK2)),sK2,sK1)),s(fun(sK1,bool),sK2))) | ~p(s(bool,X0))) ) | $spl32),
  inference(subsumption_resolution,[],[f1568,f179])).
fof(f179,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f117])).
fof(f117,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f50])).
fof(f50,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X8] : (s(bool,'F_0') = s(bool,X8) <=> ~p(s(bool,X8)))),
  file('thm.measure.POW_ALGEBRA','thm.bool.EQ_CLAUSES|split|2')).
fof(f1568,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK6('const.pred_set.POW_1'(s(fun(sK1,bool),sK2)),sK2,sK1)),s(fun(sK1,bool),sK2))) | ~p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) | $spl32),
  inference(superposition,[],[f917,f207])).
fof(f520,plain,(
  p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))) | $spl22),
  inference(cnf_transformation,[],[f520_D])).
fof(f520_D,plain,(
  p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f1381,plain,(
  ( ! [X5] : (p(s(bool,X5)) | p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))) ) | $spl30),
  inference(subsumption_resolution,[],[f1368,f180])).
fof(f180,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f117])).
fof(f1368,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,X5) | p(s(bool,X5)) | p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))) ) | $spl30),
  inference(superposition,[],[f678,f180])).
fof(f1380,plain,(
  ( ! [X4] : (p(s(bool,X4)) | p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))) ) | $spl30),
  inference(subsumption_resolution,[],[f1367,f180])).
fof(f1367,plain,(
  ( ! [X4] : (s(bool,'F_0') != s(bool,X4) | p(s(bool,X4)) | p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))) ) | $spl30),
  inference(superposition,[],[f678,f182])).
fof(f182,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f52])).
fof(f52,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X8] : (s(bool,'F_0') = s(bool,X8) <=> ~p(s(bool,X8)))),
  file('thm.measure.POW_ALGEBRA','thm.bool.EQ_CLAUSES|split|3')).
fof(f1379,plain,(
  s(bool,'T_0') = s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))) | $spl38),
  inference(cnf_transformation,[],[f1379_D])).
fof(f1379_D,plain,(
  s(bool,'T_0') = s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f1377,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'T_0') = s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))) ) | $spl30),
  inference(subsumption_resolution,[],[f1364,f180])).
fof(f1364,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,X1)) | s(bool,'T_0') = s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))) ) | $spl30),
  inference(superposition,[],[f678,f207])).
fof(f1374,plain,(
  s(bool,'F_0') != s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))) | $spl37),
  inference(cnf_transformation,[],[f1374_D])).
fof(f1374_D,plain,(
  s(bool,'F_0') != s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))) <=> ~$spl37),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl37])])).
fof(f1370,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))) | p(s(bool,X5))) ) | $spl30),
  inference(duplicate_literal_removal,[],[f1360])).
fof(f1360,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))) | p(s(bool,X5)) | p(s(bool,X5))) ) | $spl30),
  inference(superposition,[],[f678,f180])).
fof(f1371,plain,(
  ( ! [X4] : (s(bool,'F_0') != s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))) | p(s(bool,X4))) ) | $spl30),
  inference(duplicate_literal_removal,[],[f1359])).
fof(f1359,plain,(
  ( ! [X4] : (s(bool,'F_0') != s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))) | p(s(bool,X4)) | p(s(bool,X4))) ) | $spl30),
  inference(superposition,[],[f678,f182])).
fof(f1372,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))) | p(s(bool,X1))) ) | $spl30),
  inference(subsumption_resolution,[],[f1356,f183])).
fof(f183,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f119])).
fof(f1356,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))) | p(s(bool,X1)) | s(bool,'T_0') = s(bool,X1)) ) | $spl30),
  inference(superposition,[],[f678,f207])).
fof(f1349,plain,(
  $false | $spl23),
  inference(subsumption_resolution,[],[f1285,f1258])).
fof(f1258,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK6('const.pred_set.POW_1'(s(fun(sK1,bool),sK2)),sK2,sK1)),s(fun(sK1,bool),sK2)))) | $spl23),
  inference(resolution,[],[f521,f213])).
fof(f213,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),X1),s(fun(fun(X0,bool),bool),X2)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),sK6(X2,X1,X0)),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),X1),s(fun(fun(X0,bool),bool),X2)))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X4),s(fun(fun(X0,bool),bool),X2)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X4),s(fun(X0,bool),X1)))))) & ((p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),sK6(X2,X1,X0)),s(fun(fun(X0,bool),bool),X2)))) & ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),sK6(X2,X1,X0)),s(fun(X0,bool),X1))))) | p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),X1),s(fun(fun(X0,bool),bool),X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK6])],[f130])).
fof(f130,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),X1),s(fun(fun(X0,bool),bool),X2)))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X4),s(fun(fun(X0,bool),bool),X2)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X4),s(fun(X0,bool),X1)))))) & (? [X3] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),X2)))) & ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X1))))) | p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),X1),s(fun(fun(X0,bool),bool),X2))))))),
  inference(rectify,[],[f129])).
fof(f129,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),X1),s(fun(fun(X0,bool),bool),X2)))) | ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),X2)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X1)))))) & (? [X3] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),X2)))) & ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X1))))) | p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),X1),s(fun(fun(X0,bool),bool),X2))))))),
  inference(nnf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),X1),s(fun(fun(X0,bool),bool),X2)))) <=> ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),X2)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X1))))))),
  inference(ennf_transformation,[],[f70])).
fof(f70,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),X1),s(fun(fun(X0,bool),bool),X2)))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),X2)))) => p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X1))))))),
  inference(flattening,[],[f69])).
fof(f69,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),X1),s(fun(fun(X0,bool),bool),X2)))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),X2)))) => p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X1))))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X6] : ! [X7,X16] : (p(s(bool,'const.measure.subset_class_2'(s(fun(X6,bool),X7),s(fun(fun(X6,bool),bool),X16)))) <=> ! [X9] : (p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),X9),s(fun(fun(X6,bool),bool),X16)))) => p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X9),s(fun(X6,bool),X7))))))),
  file('thm.measure.POW_ALGEBRA','thm.measure.subset_class_def')).
fof(f1285,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK6('const.pred_set.POW_1'(s(fun(sK1,bool),sK2)),sK2,sK1)),s(fun(sK1,bool),sK2)))) | $spl23),
  inference(forward_demodulation,[],[f1257,f204])).
fof(f204,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.pred_set.POW_1'(s(fun(X0,bool),X1)))))) )),
  inference(cnf_transformation,[],[f64])).
fof(f64,plain,(
  ! [X0,X1,X2] : s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.pred_set.POW_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f63])).
fof(f63,plain,(
  ! [X0] : ! [X1,X2] : s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.pred_set.POW_1'(s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X6] : ! [X20,X21] : s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X21),s(fun(X6,bool),X20))) = s(bool,'const.bool.IN_2'(s(fun(X6,bool),X21),s(fun(fun(X6,bool),bool),'const.pred_set.POW_1'(s(fun(X6,bool),X20)))))),
  file('thm.measure.POW_ALGEBRA','thm.pred_set.IN_POW')).
fof(f1257,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK6('const.pred_set.POW_1'(s(fun(sK1,bool),sK2)),sK2,sK1)),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))) | $spl23),
  inference(resolution,[],[f521,f212])).
fof(f212,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),X1),s(fun(fun(X0,bool),bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),sK6(X2,X1,X0)),s(fun(fun(X0,bool),bool),X2))))) )),
  inference(cnf_transformation,[],[f131])).
fof(f1341,plain,(
  $false | $spl4),
  inference(trivial_inequality_removal,[],[f1331])).
fof(f1331,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl4),
  inference(superposition,[],[f247,f233])).
fof(f233,plain,(
  s(bool,'F_0') = s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))),
  inference(resolution,[],[f132,f180])).
fof(f132,plain,(
  ~p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))))),
  inference(cnf_transformation,[],[f90])).
fof(f90,plain,(
  ~p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1,sK2])],[f75])).
fof(f75,plain,(
  ? [X0,X1] : ~p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),'const.pair.,_2'(s(fun(X0,bool),X1),s(fun(fun(X0,bool),bool),'const.pred_set.POW_1'(s(fun(X0,bool),X1))))))))),
  inference(ennf_transformation,[],[f35])).
fof(f35,plain,(
  ~! [X0,X1] : p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),'const.pair.,_2'(s(fun(X0,bool),X1),s(fun(fun(X0,bool),bool),'const.pred_set.POW_1'(s(fun(X0,bool),X1))))))))),
  inference(flattening,[],[f34])).
fof(f34,plain,(
  ~! [X0] : ! [X1] : p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),'const.pair.,_2'(s(fun(X0,bool),X1),s(fun(fun(X0,bool),bool),'const.pred_set.POW_1'(s(fun(X0,bool),X1))))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7] : p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),'const.pair.,_2'(s(fun(X6,bool),X7),s(fun(fun(X6,bool),bool),'const.pred_set.POW_1'(s(fun(X6,bool),X7))))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7] : p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),'const.pair.,_2'(s(fun(X6,bool),X7),s(fun(fun(X6,bool),bool),'const.pred_set.POW_1'(s(fun(X6,bool),X7))))))))),
  file('thm.measure.POW_ALGEBRA','thm.measure.POW_ALGEBRA')).
fof(f1340,plain,(
  $false | $spl4),
  inference(equality_resolution,[],[f247])).
fof(f917,plain,(
  ( ! [X4] : (s(bool,X4) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK6('const.pred_set.POW_1'(s(fun(sK1,bool),sK2)),sK2,sK1)),s(fun(sK1,bool),sK2))) | ~p(s(bool,X4))) ) | $spl32),
  inference(cnf_transformation,[],[f917_D])).
fof(f917_D,plain,(
  ( ! [X4] : (s(bool,X4) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK6('const.pred_set.POW_1'(s(fun(sK1,bool),sK2)),sK2,sK1)),s(fun(sK1,bool),sK2))) | ~p(s(bool,X4))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f1312,plain,(
  ( ! [X28,X27] : (p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X28) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK6('const.pred_set.POW_1'(s(fun(sK1,bool),sK2)),sK2,sK1)),s(fun(sK1,bool),sK2)))) ) | $spl23),
  inference(resolution,[],[f1258,f216])).
fof(f1301,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK6('const.pred_set.POW_1'(s(fun(sK1,bool),sK2)),sK2,sK1)),s(fun(sK1,bool),sK2))) | s(bool,X24) != s(bool,X26)) ) | $spl23),
  inference(resolution,[],[f1258,f140])).
fof(f1299,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK6('const.pred_set.POW_1'(s(fun(sK1,bool),sK2)),sK2,sK1)),s(fun(sK1,bool),sK2)))) ) | $spl23),
  inference(resolution,[],[f1258,f139])).
fof(f1297,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK6('const.pred_set.POW_1'(s(fun(sK1,bool),sK2)),sK2,sK1)),s(fun(sK1,bool),sK2)))) ) | $spl23),
  inference(resolution,[],[f1258,f138])).
fof(f1296,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK6('const.pred_set.POW_1'(s(fun(sK1,bool),sK2)),sK2,sK1)),s(fun(sK1,bool),sK2))) | s(bool,X9) != s(bool,X11)) ) | $spl23),
  inference(resolution,[],[f1258,f136])).
fof(f1294,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK6('const.pred_set.POW_1'(s(fun(sK1,bool),sK2)),sK2,sK1)),s(fun(sK1,bool),sK2)))) ) | $spl23),
  inference(resolution,[],[f1258,f134])).
fof(f646,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))) ) | $spl28),
  inference(cnf_transformation,[],[f646_D])).
fof(f646_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f1277,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))) ) | $spl23),
  inference(resolution,[],[f521,f216])).
fof(f1266,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))) | s(bool,X21) != s(bool,X23)) ) | $spl23),
  inference(resolution,[],[f521,f140])).
fof(f1264,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))) ) | $spl23),
  inference(resolution,[],[f521,f139])).
fof(f1262,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))) ) | $spl23),
  inference(resolution,[],[f521,f138])).
fof(f1261,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))) | s(bool,X6) != s(bool,X8)) ) | $spl23),
  inference(resolution,[],[f521,f136])).
fof(f1259,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))) ) | $spl23),
  inference(resolution,[],[f521,f134])).
fof(f1230,plain,(
  ( ! [X26,X27] : (p(s(bool,X26)) | p(s(bool,X27)) | s(bool,X27) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) ) | $spl18),
  inference(resolution,[],[f517,f216])).
fof(f1220,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X23) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2))) | s(bool,X24) != s(bool,X25)) ) | $spl18),
  inference(resolution,[],[f517,f140])).
fof(f1217,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X16) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) ) | $spl18),
  inference(resolution,[],[f517,f139])).
fof(f1216,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) ) | $spl18),
  inference(resolution,[],[f517,f138])).
fof(f1215,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X8) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2))) | s(bool,X9) != s(bool,X10)) ) | $spl18),
  inference(resolution,[],[f517,f136])).
fof(f1213,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) ) | $spl18),
  inference(resolution,[],[f517,f134])).
fof(f1207,plain,(
  $false | ($spl22 | $spl28)),
  inference(subsumption_resolution,[],[f1178,f1132])).
fof(f1132,plain,(
  ( ! [X1] : (~p(s(bool,X1))) ) | ($spl22 | $spl28)),
  inference(subsumption_resolution,[],[f1131,f186])).
fof(f1131,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,X1) | ~p(s(bool,X1))) ) | ($spl22 | $spl28)),
  inference(forward_demodulation,[],[f646,f939])).
fof(f939,plain,(
  s(bool,'T_0') = s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))) | $spl22),
  inference(resolution,[],[f520,f184])).
fof(f1178,plain,(
  ( ! [X80,X81,X82] : (p(s(bool,'const.measure.subset_class_2'(s(fun(X80,bool),X81),s(fun(fun(X80,bool),bool),X82))))) ) | ($spl22 | $spl28)),
  inference(resolution,[],[f1132,f212])).
fof(f1206,plain,(
  $false | ($spl22 | $spl28)),
  inference(subsumption_resolution,[],[f1166,f1132])).
fof(f1166,plain,(
  ( ! [X52,X53,X51] : (p(s(bool,'const.bool.IN_2'(s(fun(X51,bool),sK6(X52,X53,X51)),s(fun(fun(X51,bool),bool),X52))))) ) | ($spl22 | $spl28)),
  inference(resolution,[],[f1132,f212])).
fof(f1205,plain,(
  $false | ($spl22 | $spl28)),
  inference(subsumption_resolution,[],[f1158,f1132])).
fof(f1158,plain,(
  ( ! [X44] : (p(s(bool,X44))) ) | ($spl22 | $spl28)),
  inference(resolution,[],[f1132,f161])).
fof(f161,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f106])).
fof(f106,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f105])).
fof(f105,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (~p(s(bool,'T_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(ennf_transformation,[],[f44])).
fof(f44,plain,(
  ! [X0] : ((p(s(bool,'T_0')) => p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X8] : ((p(s(bool,'T_0')) => p(s(bool,X8))) <=> p(s(bool,X8)))),
  file('thm.measure.POW_ALGEBRA','thm.bool.IMP_CLAUSES|split|0')).
fof(f1204,plain,(
  $false | ($spl22 | $spl28)),
  inference(subsumption_resolution,[],[f1157,f1132])).
fof(f1157,plain,(
  ( ! [X43] : (p(s(bool,X43))) ) | ($spl22 | $spl28)),
  inference(resolution,[],[f1132,f158])).
fof(f158,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,'T_0'))))),
  inference(flattening,[],[f103])).
fof(f103,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X8] : ((p(s(bool,X8)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  file('thm.measure.POW_ALGEBRA','thm.bool.IMP_CLAUSES|split|1')).
fof(f1203,plain,(
  $false | ($spl22 | $spl28)),
  inference(subsumption_resolution,[],[f1155,f1132])).
fof(f1155,plain,(
  p(s(bool,'F_0')) | ($spl22 | $spl28)),
  inference(resolution,[],[f1132,f155])).
fof(f155,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f101])).
fof(f101,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X8] : ((p(s(bool,'F_0')) => p(s(bool,X8))) <=> p(s(bool,'T_0')))),
  file('thm.measure.POW_ALGEBRA','thm.bool.IMP_CLAUSES|split|2')).
fof(f1202,plain,(
  $false | ($spl22 | $spl28)),
  inference(subsumption_resolution,[],[f1153,f1132])).
fof(f1153,plain,(
  ( ! [X40] : (p(s(bool,X40))) ) | ($spl22 | $spl28)),
  inference(resolution,[],[f1132,f152])).
fof(f152,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))))),
  inference(flattening,[],[f99])).
fof(f99,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f41])).
fof(f41,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X8] : ((p(s(bool,X8)) => p(s(bool,X8))) <=> p(s(bool,'T_0')))),
  file('thm.measure.POW_ALGEBRA','thm.bool.IMP_CLAUSES|split|3')).
fof(f1161,plain,(
  $false | ($spl22 | $spl28)),
  inference(resolution,[],[f1132,f188])).
fof(f188,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'T_0'))),
  file('thm.measure.POW_ALGEBRA','thm.bool.TRUTH')).
fof(f1162,plain,(
  $false | ($spl22 | $spl28)),
  inference(resolution,[],[f1132,f189])).
fof(f1182,plain,(
  $false | ($spl22 | $spl28)),
  inference(resolution,[],[f1132,f215])).
fof(f215,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0,X1] : p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X0,bool),X1))))),
  inference(flattening,[],[f73])).
fof(f73,plain,(
  ! [X0] : ! [X1] : p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X0,bool),X1))))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X6] : ! [X18] : p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),'const.pred_set.EMPTY_0'),s(fun(X6,bool),X18))))),
  file('thm.measure.POW_ALGEBRA','thm.pred_set.EMPTY_SUBSET')).
fof(f1183,plain,(
  $false | ($spl22 | $spl28)),
  inference(resolution,[],[f1132,f214])).
fof(f214,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0,X1,X2] : p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X1))))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ! [X0] : ! [X1,X2] : p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X1))))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X6] : ! [X18,X8] : p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),'const.pred_set.DIFF_2'(s(fun(X6,bool),X18),s(fun(X6,bool),X8))),s(fun(X6,bool),X18))))),
  file('thm.measure.POW_ALGEBRA','thm.pred_set.DIFF_SUBSET')).
fof(f678,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))) | p(s(bool,X2))) ) | $spl30),
  inference(cnf_transformation,[],[f678_D])).
fof(f678_D,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))) | p(s(bool,X2))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f941,plain,(
  ( ! [X26,X25] : (p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X26) != s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))) ) | $spl22),
  inference(resolution,[],[f520,f216])).
fof(f931,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))) | s(bool,X23) != s(bool,X24)) ) | $spl22),
  inference(resolution,[],[f520,f140])).
fof(f928,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))) ) | $spl22),
  inference(resolution,[],[f520,f139])).
fof(f927,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))) ) | $spl22),
  inference(resolution,[],[f520,f138])).
fof(f926,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))) | s(bool,X8) != s(bool,X9)) ) | $spl22),
  inference(resolution,[],[f520,f136])).
fof(f924,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))) ) | $spl22),
  inference(resolution,[],[f520,f134])).
fof(f909,plain,(
  ( ! [X28,X27] : (p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X28) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK6('const.pred_set.POW_1'(s(fun(sK1,bool),sK2)),sK2,sK1)),s(fun(sK1,bool),sK2)))) ) | $spl23),
  inference(resolution,[],[f685,f216])).
fof(f685,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK6('const.pred_set.POW_1'(s(fun(sK1,bool),sK2)),sK2,sK1)),s(fun(sK1,bool),sK2)))) | $spl23),
  inference(resolution,[],[f521,f213])).
fof(f898,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK6('const.pred_set.POW_1'(s(fun(sK1,bool),sK2)),sK2,sK1)),s(fun(sK1,bool),sK2))) | s(bool,X24) != s(bool,X26)) ) | $spl23),
  inference(resolution,[],[f685,f140])).
fof(f896,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK6('const.pred_set.POW_1'(s(fun(sK1,bool),sK2)),sK2,sK1)),s(fun(sK1,bool),sK2)))) ) | $spl23),
  inference(resolution,[],[f685,f139])).
fof(f894,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK6('const.pred_set.POW_1'(s(fun(sK1,bool),sK2)),sK2,sK1)),s(fun(sK1,bool),sK2)))) ) | $spl23),
  inference(resolution,[],[f685,f138])).
fof(f893,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK6('const.pred_set.POW_1'(s(fun(sK1,bool),sK2)),sK2,sK1)),s(fun(sK1,bool),sK2))) | s(bool,X9) != s(bool,X11)) ) | $spl23),
  inference(resolution,[],[f685,f136])).
fof(f891,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK6('const.pred_set.POW_1'(s(fun(sK1,bool),sK2)),sK2,sK1)),s(fun(sK1,bool),sK2)))) ) | $spl23),
  inference(resolution,[],[f685,f134])).
fof(f704,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))) ) | $spl23),
  inference(resolution,[],[f521,f216])).
fof(f693,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))) | s(bool,X21) != s(bool,X23)) ) | $spl23),
  inference(resolution,[],[f521,f140])).
fof(f691,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))) ) | $spl23),
  inference(resolution,[],[f521,f139])).
fof(f689,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))) ) | $spl23),
  inference(resolution,[],[f521,f138])).
fof(f688,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))) | s(bool,X6) != s(bool,X8)) ) | $spl23),
  inference(resolution,[],[f521,f136])).
fof(f686,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))) ) | $spl23),
  inference(resolution,[],[f521,f134])).
fof(f670,plain,(
  ( ! [X26,X25] : (p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X26) != s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))) ) | $spl22),
  inference(resolution,[],[f520,f216])).
fof(f660,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))) | s(bool,X23) != s(bool,X24)) ) | $spl22),
  inference(resolution,[],[f520,f140])).
fof(f657,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))) ) | $spl22),
  inference(resolution,[],[f520,f139])).
fof(f656,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))) ) | $spl22),
  inference(resolution,[],[f520,f138])).
fof(f655,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))) | s(bool,X8) != s(bool,X9)) ) | $spl22),
  inference(resolution,[],[f520,f136])).
fof(f653,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))) ) | $spl22),
  inference(resolution,[],[f520,f134])).
fof(f637,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))) ) | $spl23),
  inference(resolution,[],[f521,f216])).
fof(f626,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))) | s(bool,X21) != s(bool,X23)) ) | $spl23),
  inference(resolution,[],[f521,f140])).
fof(f624,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))) ) | $spl23),
  inference(resolution,[],[f521,f139])).
fof(f622,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))) ) | $spl23),
  inference(resolution,[],[f521,f138])).
fof(f621,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))) | s(bool,X6) != s(bool,X8)) ) | $spl23),
  inference(resolution,[],[f521,f136])).
fof(f619,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))) ) | $spl23),
  inference(resolution,[],[f521,f134])).
fof(f565,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK5('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)))),s(fun(sK1,bool),sK2)))) | $spl27),
  inference(cnf_transformation,[],[f565_D])).
fof(f565_D,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK5('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)))),s(fun(sK1,bool),sK2)))) <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f521,plain,(
  ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))) | $spl23),
  inference(cnf_transformation,[],[f521_D])).
fof(f521_D,plain,(
  ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f576,plain,(
  ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK5('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)))),s(fun(sK1,bool),sK2))))),
  inference(forward_demodulation,[],[f575,f203])).
fof(f203,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),'const.pair.,_2'(s(fun(X0,bool),X1),s(fun(fun(X0,bool),bool),X2)))))) )),
  inference(cnf_transformation,[],[f62])).
fof(f62,plain,(
  ! [X0,X1,X2] : s(fun(X0,bool),X1) = s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),'const.pair.,_2'(s(fun(X0,bool),X1),s(fun(fun(X0,bool),bool),X2)))))),
  inference(flattening,[],[f61])).
fof(f61,plain,(
  ! [X0] : ! [X1,X2] : s(fun(X0,bool),X1) = s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),'const.pair.,_2'(s(fun(X0,bool),X1),s(fun(fun(X0,bool),bool),X2)))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X6] : ! [X9,X10] : s(fun(X6,bool),X9) = s(fun(X6,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),'const.pair.,_2'(s(fun(X6,bool),X9),s(fun(fun(X6,bool),bool),X10)))))),
  file('thm.measure.POW_ALGEBRA','thm.measure.space_def')).
fof(f575,plain,(
  ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK5('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)))),s(fun(sK1,bool),sK2))))),
  inference(forward_demodulation,[],[f574,f202])).
fof(f202,plain,(
  ( ! [X2,X0,X1] : (s(fun(fun(X0,bool),bool),X2) = s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),'const.pair.,_2'(s(fun(X0,bool),X1),s(fun(fun(X0,bool),bool),X2)))))) )),
  inference(cnf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0,X1,X2] : s(fun(fun(X0,bool),bool),X2) = s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),'const.pair.,_2'(s(fun(X0,bool),X1),s(fun(fun(X0,bool),bool),X2)))))),
  inference(flattening,[],[f59])).
fof(f59,plain,(
  ! [X0] : ! [X1,X2] : s(fun(fun(X0,bool),bool),X2) = s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),'const.pair.,_2'(s(fun(X0,bool),X1),s(fun(fun(X0,bool),bool),X2)))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X6] : ! [X9,X10] : s(fun(fun(X6,bool),bool),X10) = s(fun(fun(X6,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),'const.pair.,_2'(s(fun(X6,bool),X9),s(fun(fun(X6,bool),bool),X10)))))),
  file('thm.measure.POW_ALGEBRA','thm.measure.subsets_def')).
fof(f574,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK5('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)))),s(fun(sK1,bool),sK2)))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))))))),
  inference(subsumption_resolution,[],[f573,f215])).
fof(f573,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(sK1,bool),sK2)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK5('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)))),s(fun(sK1,bool),sK2)))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))))))),
  inference(forward_demodulation,[],[f572,f204])).
fof(f572,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK5('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)))),s(fun(sK1,bool),sK2)))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))))))),
  inference(forward_demodulation,[],[f571,f202])).
fof(f571,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK5('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)))),s(fun(sK1,bool),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))))))),
  inference(subsumption_resolution,[],[f570,f214])).
fof(f570,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK2),s(fun(sK1,bool),sK3('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)))),s(fun(sK1,bool),sK2)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK5('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)))),s(fun(sK1,bool),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))))))),
  inference(forward_demodulation,[],[f569,f203])).
fof(f569,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(sK1,bool),sK3('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)))),s(fun(sK1,bool),sK2)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK5('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)))),s(fun(sK1,bool),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))))))),
  inference(forward_demodulation,[],[f568,f204])).
fof(f568,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(sK1,bool),sK3('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)))),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK5('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)))),s(fun(sK1,bool),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))))))),
  inference(forward_demodulation,[],[f567,f202])).
fof(f567,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK5('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)))),s(fun(sK1,bool),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(sK1,bool),sK3('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))))))),
  inference(forward_demodulation,[],[f566,f204])).
fof(f566,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK5('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)))),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(sK1,bool),sK3('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))))))),
  inference(forward_demodulation,[],[f505,f202])).
fof(f505,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK5('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(sK1,bool),sK3('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))))))),
  inference(resolution,[],[f218,f199])).
fof(f199,plain,(
  ( ! [X0,X1] : (sP0(X0,X1) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),sK4(X1,X0)),s(fun(X0,bool),sK5(X1,X0)))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),sK3(X1,X0)))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) )),
  inference(cnf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ! [X7] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X7),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X7))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) & ! [X5,X6] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X5),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X6),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X5),s(fun(X0,bool),X6))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))))) & (~p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),sK3(X1,X0)),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),sK3(X1,X0)))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) | (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),sK4(X1,X0)),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),sK5(X1,X0)),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),sK4(X1,X0)),s(fun(X0,bool),sK5(X1,X0)))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) | sP0(X0,X1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3,sK4,sK5])],[f123])).
fof(f123,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ! [X7] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X7),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X7))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) & ! [X5,X6] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X5),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X6),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X5),s(fun(X0,bool),X6))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))))) & (~p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ? [X4] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X4),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X4))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) | ? [X2,X3] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) | sP0(X0,X1)))),
  inference(rectify,[],[f122])).
fof(f122,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X4),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X4))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) & ! [X2,X3] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))))) & (~p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ? [X4] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X4),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X4))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) | ? [X2,X3] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) | sP0(X0,X1)))),
  inference(flattening,[],[f121])).
fof(f121,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X4),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X4))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) & ! [X2,X3] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))))) & ((~p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ? [X4] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X4),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X4))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) | ? [X2,X3] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f88])).
fof(f88,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X4),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X4))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) & ! [X2,X3] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f218,plain,(
  ~sP0(sK1,'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))),
  inference(resolution,[],[f132,f201])).
fof(f201,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) | sP0(X0,X1)) & (~sP0(X0,X1) | p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))),
  inference(nnf_transformation,[],[f89])).
fof(f89,plain,(
  ! [X0,X1] : (p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) <=> sP0(X0,X1))),
  inference(definition_folding,[],[f86,f88])).
fof(f86,plain,(
  ! [X0,X1] : (p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) <=> (p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X4),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X4))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) & ! [X2,X3] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))))),
  inference(flattening,[],[f85])).
fof(f85,plain,(
  ! [X0,X1] : (p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) <=> (p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X4),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X4))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) & ! [X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))))),
  inference(ennf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0,X1] : (p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) <=> (p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ! [X4] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X4),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) => p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X4))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) & ! [X2,X3] : ((p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) => p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))))),
  inference(flattening,[],[f57])).
fof(f57,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) <=> (p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ! [X4] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X4),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) => p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X4))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) & ! [X2,X3] : ((p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) => p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X6] : ! [X17] : (p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X17)))) <=> (p(s(bool,'const.measure.subset_class_2'(s(fun(X6,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X17))),s(fun(fun(X6,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X17)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X6,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X17)))))) & ! [X18] : (p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),X18),s(fun(fun(X6,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X17)))))) => p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),'const.pred_set.DIFF_2'(s(fun(X6,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X17))),s(fun(X6,bool),X18))),s(fun(fun(X6,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X17))))))) & ! [X18,X8] : ((p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),X18),s(fun(fun(X6,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X17)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),X8),s(fun(fun(X6,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X17))))))) => p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),X18),s(fun(X6,bool),X8))),s(fun(fun(X6,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X17)))))))))),
  file('thm.measure.POW_ALGEBRA','thm.measure.algebra_def')).
fof(f519,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK3('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) | $spl20),
  inference(cnf_transformation,[],[f519_D])).
fof(f519_D,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK3('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f563,plain,(
  ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK3('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK5('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)))),s(fun(sK1,bool),sK2))))),
  inference(forward_demodulation,[],[f562,f203])).
fof(f562,plain,(
  ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK3('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK5('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)))),s(fun(sK1,bool),sK2))))),
  inference(forward_demodulation,[],[f561,f202])).
fof(f561,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK3('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK5('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)))),s(fun(sK1,bool),sK2)))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))))))),
  inference(subsumption_resolution,[],[f560,f215])).
fof(f560,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(sK1,bool),sK2)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK3('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK5('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)))),s(fun(sK1,bool),sK2)))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))))))),
  inference(forward_demodulation,[],[f559,f204])).
fof(f559,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK3('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK5('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)))),s(fun(sK1,bool),sK2)))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))))))),
  inference(forward_demodulation,[],[f558,f202])).
fof(f558,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK3('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK5('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)))),s(fun(sK1,bool),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))))))),
  inference(forward_demodulation,[],[f557,f204])).
fof(f557,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK3('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK5('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)))),s(fun(sK1,bool),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))))))),
  inference(forward_demodulation,[],[f556,f202])).
fof(f556,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK5('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)))),s(fun(sK1,bool),sK2)))) | p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK3('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))))))),
  inference(forward_demodulation,[],[f555,f204])).
fof(f555,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK5('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)))),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK3('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))))))),
  inference(forward_demodulation,[],[f504,f202])).
fof(f504,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK5('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK3('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))))))),
  inference(resolution,[],[f218,f198])).
fof(f198,plain,(
  ( ! [X0,X1] : (sP0(X0,X1) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),sK4(X1,X0)),s(fun(X0,bool),sK5(X1,X0)))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),sK3(X1,X0)),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) )),
  inference(cnf_transformation,[],[f124])).
fof(f543,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK5('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) | $spl24),
  inference(cnf_transformation,[],[f543_D])).
fof(f543_D,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK5('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f554,plain,(
  ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK5('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2))))),
  inference(forward_demodulation,[],[f553,f203])).
fof(f553,plain,(
  ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK5('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2))))),
  inference(forward_demodulation,[],[f552,f202])).
fof(f552,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK5('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))))))),
  inference(subsumption_resolution,[],[f551,f215])).
fof(f551,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(sK1,bool),sK2)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK5('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))))))),
  inference(forward_demodulation,[],[f550,f204])).
fof(f550,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK5('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))))))),
  inference(forward_demodulation,[],[f549,f202])).
fof(f549,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK5('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))))))),
  inference(subsumption_resolution,[],[f548,f214])).
fof(f548,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK2),s(fun(sK1,bool),sK3('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)))),s(fun(sK1,bool),sK2)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK5('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))))))),
  inference(forward_demodulation,[],[f547,f203])).
fof(f547,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(sK1,bool),sK3('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)))),s(fun(sK1,bool),sK2)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK5('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))))))),
  inference(forward_demodulation,[],[f546,f204])).
fof(f546,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(sK1,bool),sK3('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)))),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK5('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))))))),
  inference(forward_demodulation,[],[f545,f202])).
fof(f545,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK5('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(sK1,bool),sK3('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))))))),
  inference(forward_demodulation,[],[f544,f204])).
fof(f544,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK5('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(sK1,bool),sK3('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))))))),
  inference(forward_demodulation,[],[f503,f202])).
fof(f503,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK5('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(sK1,bool),sK3('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))))))),
  inference(resolution,[],[f218,f197])).
fof(f197,plain,(
  ( ! [X0,X1] : (sP0(X0,X1) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),sK5(X1,X0)),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),sK3(X1,X0)))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) )),
  inference(cnf_transformation,[],[f124])).
fof(f541,plain,(
  ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK3('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK5('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2))))),
  inference(forward_demodulation,[],[f540,f203])).
fof(f540,plain,(
  ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK3('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK5('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2))))),
  inference(forward_demodulation,[],[f539,f202])).
fof(f539,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK3('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK5('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))))))),
  inference(subsumption_resolution,[],[f538,f215])).
fof(f538,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(sK1,bool),sK2)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK3('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK5('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))))))),
  inference(forward_demodulation,[],[f537,f204])).
fof(f537,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK3('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK5('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))))))),
  inference(forward_demodulation,[],[f536,f202])).
fof(f536,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK3('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK5('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))))))),
  inference(forward_demodulation,[],[f535,f204])).
fof(f535,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK3('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK5('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))))))),
  inference(forward_demodulation,[],[f534,f202])).
fof(f534,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK5('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) | p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK3('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))))))),
  inference(forward_demodulation,[],[f533,f204])).
fof(f533,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK5('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK3('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))))))),
  inference(forward_demodulation,[],[f502,f202])).
fof(f502,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK5('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK3('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))))))),
  inference(resolution,[],[f218,f196])).
fof(f196,plain,(
  ( ! [X0,X1] : (sP0(X0,X1) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),sK5(X1,X0)),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),sK3(X1,X0)),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) )),
  inference(cnf_transformation,[],[f124])).
fof(f517,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) | $spl18),
  inference(cnf_transformation,[],[f517_D])).
fof(f517_D,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f532,plain,(
  ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2))))),
  inference(forward_demodulation,[],[f531,f203])).
fof(f531,plain,(
  ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2))))),
  inference(forward_demodulation,[],[f530,f202])).
fof(f530,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))))))),
  inference(subsumption_resolution,[],[f529,f215])).
fof(f529,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(sK1,bool),sK2)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))))))),
  inference(forward_demodulation,[],[f528,f204])).
fof(f528,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))))))),
  inference(forward_demodulation,[],[f527,f202])).
fof(f527,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))))))),
  inference(subsumption_resolution,[],[f526,f214])).
fof(f526,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK2),s(fun(sK1,bool),sK3('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)))),s(fun(sK1,bool),sK2)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))))))),
  inference(forward_demodulation,[],[f525,f203])).
fof(f525,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(sK1,bool),sK3('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)))),s(fun(sK1,bool),sK2)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))))))),
  inference(forward_demodulation,[],[f524,f204])).
fof(f524,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(sK1,bool),sK3('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)))),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))))))),
  inference(forward_demodulation,[],[f523,f202])).
fof(f523,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(sK1,bool),sK3('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))))))),
  inference(forward_demodulation,[],[f522,f204])).
fof(f522,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(sK1,bool),sK3('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))))))),
  inference(forward_demodulation,[],[f501,f202])).
fof(f501,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(sK1,bool),sK3('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))))))),
  inference(resolution,[],[f218,f195])).
fof(f195,plain,(
  ( ! [X0,X1] : (sP0(X0,X1) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),sK4(X1,X0)),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),sK3(X1,X0)))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) )),
  inference(cnf_transformation,[],[f124])).
fof(f515,plain,(
  ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK3('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2))))),
  inference(forward_demodulation,[],[f514,f203])).
fof(f514,plain,(
  ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK3('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2))))),
  inference(forward_demodulation,[],[f513,f202])).
fof(f513,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK3('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))))))),
  inference(subsumption_resolution,[],[f512,f215])).
fof(f512,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(sK1,bool),sK2)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK3('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))))))),
  inference(forward_demodulation,[],[f511,f204])).
fof(f511,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK3('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))))))),
  inference(forward_demodulation,[],[f510,f202])).
fof(f510,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK3('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))))))),
  inference(forward_demodulation,[],[f509,f204])).
fof(f509,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK3('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))))))),
  inference(forward_demodulation,[],[f508,f202])).
fof(f508,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(sK1,bool),sK2)))) | p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK3('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))))))),
  inference(forward_demodulation,[],[f507,f204])).
fof(f507,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK3('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))))))),
  inference(forward_demodulation,[],[f500,f202])).
fof(f500,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK4('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK3('const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))),sK1)),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))))))),
  inference(resolution,[],[f218,f194])).
fof(f194,plain,(
  ( ! [X0,X1] : (sP0(X0,X1) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),sK4(X1,X0)),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),sK3(X1,X0)),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) )),
  inference(cnf_transformation,[],[f124])).
fof(f372,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,X1)) ) | $spl14),
  inference(cnf_transformation,[],[f372_D])).
fof(f372_D,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,X1)) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f406,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X7) | ~p(s(bool,X6))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f382])).
fof(f382,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X7) | ~p(s(bool,X6)) | ~p(s(bool,X6))) ) | $spl8),
  inference(superposition,[],[f249,f184])).
fof(f407,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X5) | ~p(s(bool,X4))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f381])).
fof(f381,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X5) | ~p(s(bool,X4)) | ~p(s(bool,X4))) ) | $spl8),
  inference(superposition,[],[f249,f186])).
fof(f374,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,X3)) ) | $spl16),
  inference(cnf_transformation,[],[f374_D])).
fof(f374_D,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,X3)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f409,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,X3) | ~p(s(bool,X2))) ) | $spl8),
  inference(subsumption_resolution,[],[f380,f249])).
fof(f380,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,X3) | ~p(s(bool,X2)) | s(bool,'T_0') = s(bool,X2)) ) | $spl8),
  inference(superposition,[],[f249,f207])).
fof(f408,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X1) | ~p(s(bool,X0))) ) | $spl8),
  inference(subsumption_resolution,[],[f379,f249])).
fof(f379,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X1) | ~p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) | $spl8),
  inference(superposition,[],[f249,f207])).
fof(f369,plain,(
  ( ! [X10,X11] : (s(bool,'F_0') != s(bool,X11) | p(s(bool,X10))) ) | $spl2),
  inference(duplicate_literal_removal,[],[f347])).
fof(f347,plain,(
  ( ! [X10,X11] : (s(bool,'F_0') != s(bool,X11) | p(s(bool,X10)) | p(s(bool,X10))) ) | $spl2),
  inference(superposition,[],[f246,f180])).
fof(f370,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') != s(bool,X9) | p(s(bool,X8))) ) | $spl2),
  inference(duplicate_literal_removal,[],[f346])).
fof(f346,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') != s(bool,X9) | p(s(bool,X8)) | p(s(bool,X8))) ) | $spl2),
  inference(superposition,[],[f246,f182])).
fof(f373,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,X3) | p(s(bool,X2))) ) | $spl2),
  inference(subsumption_resolution,[],[f343,f246])).
fof(f343,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,X3) | p(s(bool,X2)) | s(bool,'T_0') = s(bool,X2)) ) | $spl2),
  inference(superposition,[],[f246,f207])).
fof(f371,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X1) | p(s(bool,X0))) ) | $spl2),
  inference(subsumption_resolution,[],[f342,f246])).
fof(f342,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X1) | p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) | $spl2),
  inference(superposition,[],[f246,f207])).
fof(f327,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f281,f251])).
fof(f281,plain,(
  ( ! [X42] : (~p(s(bool,X42))) ) | $spl12),
  inference(resolution,[],[f251,f150])).
fof(f150,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f98])).
fof(f98,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f97])).
fof(f97,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f40])).
fof(f40,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X8] : ((p(s(bool,X8)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X8)))),
  file('thm.measure.POW_ALGEBRA','thm.bool.IMP_CLAUSES|split|4')).
fof(f277,plain,(
  $false | $spl12),
  inference(resolution,[],[f251,f187])).
fof(f187,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f56])).
fof(f56,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.measure.POW_ALGEBRA',aHL_FALSITY)).
fof(f283,plain,(
  $false | $spl12),
  inference(resolution,[],[f251,f132])).
fof(f245,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))) ) | $spl0),
  inference(cnf_transformation,[],[f245_D])).
fof(f245_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f237,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))) )),
  inference(resolution,[],[f132,f216])).
fof(f226,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))) | s(bool,X21) != s(bool,X23)) )),
  inference(resolution,[],[f132,f140])).
fof(f224,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))) )),
  inference(resolution,[],[f132,f139])).
fof(f247,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))) ) | $spl4),
  inference(cnf_transformation,[],[f247_D])).
fof(f247_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f248,plain,(
  ( ! [X4,X3] : (~p(s(bool,X3)) | s(bool,X3) != s(bool,X4) | p(s(bool,X4))) ) | $spl6),
  inference(cnf_transformation,[],[f248_D])).
fof(f248_D,plain,(
  ( ! [X4,X3] : (~p(s(bool,X3)) | s(bool,X3) != s(bool,X4) | p(s(bool,X4))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f223,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))) | s(bool,X12) != s(bool,X13)) )),
  inference(resolution,[],[f132,f138])).
fof(f222,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))) )),
  inference(resolution,[],[f132,f138])).
fof(f221,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))) | s(bool,X6) != s(bool,X8)) )),
  inference(resolution,[],[f132,f136])).
fof(f220,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2))))))) | s(bool,X3) != s(bool,X4)) )),
  inference(resolution,[],[f132,f134])).
fof(f219,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),'const.pair.,_2'(s(fun(sK1,bool),sK2),s(fun(fun(sK1,bool),bool),'const.pred_set.POW_1'(s(fun(sK1,bool),sK2)))))))) )),
  inference(resolution,[],[f132,f134])).
% SZS output end Proof for thm.measure
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 88
% Passive clauses: 131
% Generated clauses: 1631
% Final active clauses: 65
% Final passive clauses: 16
% Input formulas: 32
% Initial clauses: 84
% 
% Selected by SInE selection: 31
% SInE iterations: 1
% 
% Duplicate literals: 39
% Trivial inequalities: 27
% Fw subsumption resolutions: 297
% Fw demodulations: 173
% Bw demodulations: 8
% 
% Simple tautologies: 105
% Forward subsumptions: 773
% 
% Binary resolution: 386
% Forward superposition: 346
% Backward superposition: 305
% Equality resolution: 9
% 
% Split clauses: 94
% Split components: 202
% SAT solver clauses: 112
% SAT solver unit clauses: 6
% SAT solver binary clauses: 20
% 
% Sat splits: 94
% Sat splitting refutations: 17
% 
% Memory used [KB]: 1663
% Time elapsed: 0.031 s
% ------------------------------
% ------------------------------
% Success in time 11.751 s
