Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.relation
% (96243)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 126
% Passive clauses: 2472
% Generated clauses: 18989
% Final active clauses: 115
% Final passive clauses: 2203
% Input formulas: 13
% Initial clauses: 17
% 
% Duplicate literals: 905
% Trivial inequalities: 234
% Fw subsumption resolutions: 2207
% Fw demodulations: 3603
% Bw demodulations: 7
% 
% Simple tautologies: 288
% Equational tautologies: 29
% Forward subsumptions: 8782
% Fw demodulations to eq. taut.: 8
% 
% Binary resolution: 10
% Forward superposition: 7035
% Backward superposition: 4872
% Self superposition: 186
% Equality factoring: 62
% Equality resolution: 29
% 
% Split clauses: 336
% Split components: 686
% SAT solver clauses: 601
% SAT solver unit clauses: 10
% SAT solver binary clauses: 71
% 
% Sat splits: 336
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 597
% TWLsolver calls for satisfiability: 78
% 
% Memory used [KB]: 12920
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 294 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.relation
% (96265)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 98
% Passive clauses: 855
% Generated clauses: 14640
% Final active clauses: 89
% Final passive clauses: 742
% Input formulas: 13
% Initial clauses: 17
% 
% Duplicate literals: 594
% Trivial inequalities: 159
% Fw subsumption resolutions: 1120
% Fw demodulations: 2668
% Bw demodulations: 20
% Condensations: 1086
% 
% Simple tautologies: 131
% Equational tautologies: 25
% Forward subsumptions: 6332
% Fw demodulations to eq. taut.: 8
% 
% Binary resolution: 1
% Forward superposition: 4745
% Backward superposition: 3130
% Self superposition: 157
% Equality factoring: 44
% Equality resolution: 20
% 
% Split clauses: 62
% Split components: 131
% SAT solver clauses: 1654
% SAT solver unit clauses: 22
% SAT solver binary clauses: 24
% 
% Sat splits: 62
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 1654
% TWLsolver calls for satisfiability: 129
% 
% Memory used [KB]: 12792
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.relation
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.relation
% SZS output start Proof for thm.relation
fof(f20998,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f46,f66,f64,f45,f67,f63,f83,f390,f418,f416,f1111,f1114,f1113,f1243,f1254,f1251,f1453,f2119,f2120,f2124,f2126,f2130,f2104,f2131,f2088,f2145,f2143,f2125,f2164,f2166,f2352,f2341,f3065,f2142,f5289,f5271,f5312,f5524,f5526,f5528,f5529,f5533,f5534,f5535,f5536,f5541,f5546,f3029,f5548,f5549,f3064,f3062,f3059,f3042,f5550,f1112,f5551,f5538,f5554,f5557,f5560,f5563,f5564,f5628,f5651,f5653,f5657,f1262,f1260,f5661,f5525,f2634,f2637,f86,f5693,f5697,f5698,f5699,f5701,f5702,f5706,f5710,f5712,f5713,f5720,f5721,f5724,f5726,f5727,f5728,f5729,f5730,f5732,f5736,f5752,f5770,f5782,f5795,f5797,f5806,f5808,f5812,f5814,f5819,f5822,f5825,f5828,f5830,f5835,f5836,f5837,f5841,f5843,f5846,f5848,f5854,f5856,f5858,f5860,f5862,f5864,f5866,f5868,f5870,f5871,f5873,f5874,f5875,f5878,f5881,f5884,f5891,f5894,f5906,f5909,f5913,f5920,f5924,f5932,f5935,f5937,f5939,f5942,f5947,f5953,f5957,f5960,f5971,f5975,f5977,f5984,f5988,f5990,f5998,f6002,f6004,f6009,f6013,f6014,f6023,f6025,f6028,f6031,f6032,f6033,f6035,f6036,f6037,f6038,f6040,f6041,f6048,f6049,f6053,f6055,f6057,f6058,f6061,f6063,f6069,f6070,f6073,f6075,f6076,f7509,f7511,f7510,f9980,f84,f9979,f10025,f10295,f10296,f10300,f10301,f15006,f15010,f5539,f15023,f14913,f15008,f15084,f17525,f17530,f17544,f17551,f5547,f17415,f17527,f17647,f20927])).
fof(f20927,plain,(
  $false | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f20926,f68])).
fof(f68,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))) )),
  inference(forward_demodulation,[],[f52,f51])).
fof(f51,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))) )),
  inference(cnf_transformation,[],[f19])).
fof(f19,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))),
  inference(flattening,[],[f18])).
fof(f18,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X9] : ! [X8] : s(X9,X8) = s(X9,'const.combin.I_1'(s(X9,X8)))),
  file('thm.relation.INVOL','thm.combin.I_THM')).
fof(f52,plain,(
  ( ! [X0,X1] : (s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f21])).
fof(f21,plain,(
  ! [X0,X1] : s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))),
  inference(flattening,[],[f20])).
fof(f20,plain,(
  ! [X0] : ! [X1] : s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X9] : ! [X13] : s(X9,'const.combin.I_1'(s(X9,X13))) = s(X9,happ(s(fun(X9,X9),'const.combin.I_0'),s(X9,X13)))),
  file('thm.relation.INVOL','ARITY_0_1_const.combin.I')).
fof(f20926,plain,(
  s(sK0,sK4('const.combin.I_0','const.combin.o_2'(s(fun(sK0,sK0),sK1),s(fun(sK0,sK0),sK1)),sK0,sK0)) != s(sK0,happ(s(fun(sK0,sK0),'const.combin.I_0'),s(sK0,sK4('const.combin.I_0','const.combin.o_2'(s(fun(sK0,sK0),sK1),s(fun(sK0,sK0),sK1)),sK0,sK0)))) | ($spl1 | $spl4)),
  inference(forward_demodulation,[],[f20817,f2366])).
fof(f2366,plain,(
  ( ! [X1] : (s(sK0,X1) = s(sK0,'const.combin.o_3'(s(fun(sK0,sK0),sK1),s(fun(sK0,sK0),sK1),s(sK0,X1)))) ) | $spl4),
  inference(superposition,[],[f67,f60])).
fof(f60,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5)))))) )),
  inference(cnf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5)))))),
  inference(flattening,[],[f30])).
fof(f30,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5)))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X10,X9,X12] : ! [X7,X11,X8] : s(X10,'const.combin.o_3'(s(fun(X9,X10),X7),s(fun(X12,X9),X11),s(X12,X8))) = s(X10,happ(s(fun(X9,X10),X7),s(X9,happ(s(fun(X12,X9),X11),s(X12,X8)))))),
  file('thm.relation.INVOL','thm.combin.o_THM')).
fof(f20817,plain,(
  s(sK0,happ(s(fun(sK0,sK0),'const.combin.I_0'),s(sK0,sK4('const.combin.I_0','const.combin.o_2'(s(fun(sK0,sK0),sK1),s(fun(sK0,sK0),sK1)),sK0,sK0)))) != s(sK0,'const.combin.o_3'(s(fun(sK0,sK0),sK1),s(fun(sK0,sK0),sK1),s(sK0,sK4('const.combin.I_0','const.combin.o_2'(s(fun(sK0,sK0),sK1),s(fun(sK0,sK0),sK1)),sK0,sK0)))) | $spl1),
  inference(unit_resulting_resolution,[],[f2353,f1942])).
fof(f1942,plain,(
  ( ! [X6,X4,X2,X7,X5,X3] : (s(X2,happ(s(fun(X3,X2),X7),s(X3,sK4(X7,'const.combin.o_2'(s(fun(X4,X2),X5),s(fun(X3,X4),X6)),X3,X2)))) != s(X2,'const.combin.o_3'(s(fun(X4,X2),X5),s(fun(X3,X4),X6),s(X3,sK4(X7,'const.combin.o_2'(s(fun(X4,X2),X5),s(fun(X3,X4),X6)),X3,X2)))) | s(fun(X3,X2),X7) = s(fun(X3,X2),'const.combin.o_2'(s(fun(X4,X2),X5),s(fun(X3,X4),X6)))) )),
  inference(superposition,[],[f59,f61])).
fof(f61,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))) )),
  inference(cnf_transformation,[],[f33])).
fof(f33,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(flattening,[],[f32])).
fof(f32,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X10,X12,X9] : ! [X13,X14,X15] : s(X10,'const.combin.o_3'(s(fun(X12,X10),X13),s(fun(X9,X12),X14),s(X9,X15))) = s(X10,happ(s(fun(X9,X10),'const.combin.o_2'(s(fun(X12,X10),X13),s(fun(X9,X12),X14))),s(X9,X15)))),
  file('thm.relation.INVOL','ARITY_2_3_const.combin.o')).
fof(f59,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,sK4(X3,X2,X1,X0)))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,sK4(X3,X2,X1,X0)))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)) )),
  inference(cnf_transformation,[],[f44])).
fof(f44,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X5] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) & (s(X0,happ(s(fun(X1,X0),X2),s(X1,sK4(X3,X2,X1,X0)))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,sK4(X3,X2,X1,X0)))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f43])).
fof(f43,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X5] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) & (? [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(rectify,[],[f42])).
fof(f42,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4)))) & (? [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(nnf_transformation,[],[f29])).
fof(f29,plain,(
  ! [X0,X1,X2,X3] : (s(fun(X1,X0),X2) = s(fun(X1,X0),X3) <=> ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))))),
  inference(flattening,[],[f28])).
fof(f28,plain,(
  ! [X0,X1] : ! [X2,X3] : (s(fun(X1,X0),X2) = s(fun(X1,X0),X3) <=> ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X10,X9] : ! [X7,X11] : (s(fun(X9,X10),X7) = s(fun(X9,X10),X11) <=> ! [X8] : s(X10,happ(s(fun(X9,X10),X7),s(X9,X8))) = s(X10,happ(s(fun(X9,X10),X11),s(X9,X8))))),
  file('thm.relation.INVOL','thm.bool.FUN_EQ_THM')).
fof(f2353,plain,(
  s(fun(sK0,sK0),'const.combin.I_0') != s(fun(sK0,sK0),'const.combin.o_2'(s(fun(sK0,sK0),sK1),s(fun(sK0,sK0),sK1))) | $spl1),
  inference(unit_resulting_resolution,[],[f64,f56])).
fof(f56,plain,(
  ( ! [X0,X1] : (s(fun(X0,X0),'const.combin.I_0') != s(fun(X0,X0),'const.combin.o_2'(s(fun(X0,X0),X1),s(fun(X0,X0),X1))) | p(s(bool,'const.relation.INVOL_1'(s(fun(X0,X0),X1))))) )),
  inference(cnf_transformation,[],[f40])).
fof(f40,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.INVOL_1'(s(fun(X0,X0),X1)))) | s(fun(X0,X0),'const.combin.I_0') = s(fun(X0,X0),'const.combin.o_2'(s(fun(X0,X0),X1),s(fun(X0,X0),X1)))) & (s(fun(X0,X0),'const.combin.I_0') != s(fun(X0,X0),'const.combin.o_2'(s(fun(X0,X0),X1),s(fun(X0,X0),X1))) | p(s(bool,'const.relation.INVOL_1'(s(fun(X0,X0),X1))))))),
  inference(nnf_transformation,[],[f25])).
fof(f25,plain,(
  ! [X0,X1] : (p(s(bool,'const.relation.INVOL_1'(s(fun(X0,X0),X1)))) <=> s(fun(X0,X0),'const.combin.I_0') = s(fun(X0,X0),'const.combin.o_2'(s(fun(X0,X0),X1),s(fun(X0,X0),X1))))),
  inference(flattening,[],[f24])).
fof(f24,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.relation.INVOL_1'(s(fun(X0,X0),X1)))) <=> s(fun(X0,X0),'const.combin.I_0') = s(fun(X0,X0),'const.combin.o_2'(s(fun(X0,X0),X1),s(fun(X0,X0),X1))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6] : ! [X7] : (p(s(bool,'const.relation.INVOL_1'(s(fun(X6,X6),X7)))) <=> s(fun(X6,X6),'const.combin.I_0') = s(fun(X6,X6),'const.combin.o_2'(s(fun(X6,X6),X7),s(fun(X6,X6),X7))))),
  file('thm.relation.INVOL','thm.relation.INVOL_DEF')).
fof(f17647,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.combin.o_3'(s(fun(bool,bool),'const.combin.I_0'),s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0')))) ) | $spl54),
  inference(cnf_transformation,[],[f17647_D])).
fof(f17647_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.combin.o_3'(s(fun(bool,bool),'const.combin.I_0'),s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0')))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f17527,plain,(
  s(bool,'T_0') != s(bool,'const.combin.o_3'(s(fun(bool,bool),'const.combin.I_0'),s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | $spl53),
  inference(cnf_transformation,[],[f17527_D])).
fof(f17527_D,plain,(
  s(bool,'T_0') != s(bool,'const.combin.o_3'(s(fun(bool,bool),'const.combin.I_0'),s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) <=> ~$spl53),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl53])])).
fof(f17415,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.combin.o_3'(s(fun(bool,bool),'const.combin.I_0'),s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,X0) = s(bool,'const.combin.o_3'(s(fun(bool,bool),'const.combin.I_0'),s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0')))) ) | ($spl9 | $spl24 | $spl29 | $spl49)),
  inference(equality_factoring,[],[f15704])).
fof(f15704,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.combin.o_3'(s(fun(bool,bool),'const.combin.I_0'),s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl9 | $spl24 | $spl29 | $spl49)),
  inference(subsumption_resolution,[],[f15608,f86])).
fof(f15608,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X0) = s(bool,'const.combin.o_3'(s(fun(bool,bool),'const.combin.I_0'),s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl24 | $spl29 | $spl49)),
  inference(superposition,[],[f15288,f73])).
fof(f73,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f50,f50])).
fof(f50,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.relation.INVOL',aHL_BOOL_CASES)).
fof(f15288,plain,(
  s(bool,'F_0') = s(bool,'const.combin.o_3'(s(fun(bool,bool),'const.combin.I_0'),s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | ($spl24 | $spl29 | $spl49)),
  inference(forward_demodulation,[],[f15122,f51])).
fof(f15122,plain,(
  s(bool,'const.combin.I_1'(s(bool,'F_0'))) = s(bool,'const.combin.o_3'(s(fun(bool,bool),'const.combin.I_0'),s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | ($spl24 | $spl29 | $spl49)),
  inference(unit_resulting_resolution,[],[f2142,f1305,f15008,f2125])).
fof(f1305,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,X3))) = s(X0,'const.combin.o_3'(s(fun(X0,X0),'const.combin.I_0'),s(fun(X1,X0),X2),s(X1,X3)))) )),
  inference(superposition,[],[f60,f68])).
fof(f5547,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl40),
  inference(cnf_transformation,[],[f5547_D])).
fof(f5547_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f17551,plain,(
  ( ! [X101] : (s(bool,'T_0') != s(bool,'const.combin.o_3'(s(fun(bool,bool),'const.combin.I_0'),s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X101)) ) | ($spl9 | $spl24 | $spl29 | $spl49)),
  inference(subsumption_resolution,[],[f17303,f6306])).
fof(f6306,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl29),
  inference(superposition,[],[f2142,f50])).
fof(f17303,plain,(
  ( ! [X101] : (s(bool,'T_0') != s(bool,'const.combin.o_3'(s(fun(bool,bool),'const.combin.I_0'),s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X101) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X101)))) ) | ($spl9 | $spl24 | $spl29 | $spl49)),
  inference(superposition,[],[f6306,f15704])).
fof(f17544,plain,(
  ( ! [X70] : (s(bool,'T_0') != s(bool,'const.combin.o_3'(s(fun(bool,bool),'const.combin.I_0'),s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X70)) ) | ($spl9 | $spl24 | $spl29 | $spl49)),
  inference(subsumption_resolution,[],[f17287,f15199])).
fof(f15199,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl49),
  inference(superposition,[],[f15008,f50])).
fof(f17287,plain,(
  ( ! [X70] : (s(bool,'T_0') != s(bool,'const.combin.o_3'(s(fun(bool,bool),'const.combin.I_0'),s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X70) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X70)))) ) | ($spl9 | $spl24 | $spl29 | $spl49)),
  inference(superposition,[],[f15199,f15704])).
fof(f17530,plain,(
  ( ! [X49] : (s(bool,'T_0') != s(bool,'const.combin.o_3'(s(fun(bool,bool),'const.combin.I_0'),s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X49)) ) | ($spl9 | $spl16 | $spl24 | $spl29 | $spl49)),
  inference(subsumption_resolution,[],[f17273,f15199])).
fof(f17273,plain,(
  ( ! [X49] : (s(bool,'T_0') != s(bool,'const.combin.o_3'(s(fun(bool,bool),'const.combin.I_0'),s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X49))) | s(bool,'T_0') = s(bool,X49)) ) | ($spl9 | $spl16 | $spl24 | $spl29 | $spl49)),
  inference(superposition,[],[f7671,f15704])).
fof(f7671,plain,(
  ( ! [X30] : (s(bool,'T_0') != s(bool,X30) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X30)))) ) | $spl16),
  inference(forward_demodulation,[],[f7614,f51])).
fof(f7614,plain,(
  ( ! [X30] : (s(bool,X30) != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X30)))) ) | $spl16),
  inference(superposition,[],[f1113,f68])).
fof(f17525,plain,(
  ( ! [X47] : (s(bool,'T_0') != s(bool,'const.combin.o_3'(s(fun(bool,bool),'const.combin.I_0'),s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X47)) ) | ($spl1 | $spl9 | $spl24 | $spl29 | $spl49)),
  inference(subsumption_resolution,[],[f17271,f6306])).
fof(f17271,plain,(
  ( ! [X47] : (s(bool,'T_0') != s(bool,'const.combin.o_3'(s(fun(bool,bool),'const.combin.I_0'),s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X47))) | s(bool,'T_0') = s(bool,X47)) ) | ($spl1 | $spl9 | $spl24 | $spl29 | $spl49)),
  inference(superposition,[],[f6237,f15704])).
fof(f6237,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X0)))) ) | $spl1),
  inference(duplicate_literal_removal,[],[f6236])).
fof(f6236,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X0))) | s(bool,'T_0') != s(bool,X0) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X0)))) ) | $spl1),
  inference(forward_demodulation,[],[f4334,f51])).
fof(f4334,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'const.combin.I_1'(s(bool,X0))))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X0)))) ) | $spl1),
  inference(equality_factoring,[],[f2659])).
fof(f2659,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X1))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X0)))) ) | $spl1),
  inference(superposition,[],[f2434,f2434])).
fof(f2434,plain,(
  ( ! [X14] : (s(bool,X14) = s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X14)))) ) | $spl1),
  inference(superposition,[],[f2359,f51])).
fof(f2359,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,X0)) ) | $spl1),
  inference(subsumption_resolution,[],[f2355,f47])).
fof(f47,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'T_0'))),
  file('thm.relation.INVOL','thm.bool.TRUTH')).
fof(f2355,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,X0)) ) | $spl1),
  inference(superposition,[],[f64,f73])).
fof(f15084,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0')))) ) | $spl50),
  inference(cnf_transformation,[],[f15084_D])).
fof(f15084_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0')))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f15008,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | $spl49),
  inference(cnf_transformation,[],[f15008_D])).
fof(f15008_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) <=> ~$spl49),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl49])])).
fof(f14913,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,X0) = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0')))) ) | ($spl1 | $spl9 | $spl13 | $spl26)),
  inference(equality_factoring,[],[f14564])).
fof(f14564,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl1 | $spl9 | $spl13 | $spl26)),
  inference(condensation,[],[f14563])).
fof(f14563,plain,(
  ( ! [X33,X34] : (s(bool,'T_0') = s(bool,X34) | s(bool,X34) = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X33)) ) | ($spl1 | $spl9 | $spl13 | $spl26)),
  inference(trivial_inequality_removal,[],[f14528])).
fof(f14528,plain,(
  ( ! [X33,X34] : (s(bool,X33) != s(bool,X33) | s(bool,'T_0') = s(bool,X34) | s(bool,X34) = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X33)) ) | ($spl1 | $spl9 | $spl13 | $spl26)),
  inference(superposition,[],[f10305,f50])).
fof(f10305,plain,(
  ( ! [X43,X42] : (s(bool,'F_0') != s(bool,X42) | s(bool,'T_0') = s(bool,X43) | s(bool,X43) = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X42)))) ) | ($spl1 | $spl9 | $spl13 | $spl26)),
  inference(forward_demodulation,[],[f3702,f6130])).
fof(f6130,plain,(
  s(bool,'F_0') = s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | ($spl1 | $spl9 | $spl13)),
  inference(global_subsumption,[],[f3314,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1556,f1534,f1324,f473,f2377,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f86,f418,f1956,f467,f49])).
fof(f49,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f17])).
fof(f17,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.relation.INVOL',aHL_FALSITY)).
fof(f467,plain,(
  ( ! [X21,X19,X20] : (s(bool,'T_0') = s(bool,sK3(X19,X20,X21,bool)) | s(fun(bool,X21),X19) = s(fun(bool,X21),X20) | s(X21,f(s(bool,'F_0'))) != s(X21,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f57,f50])).
fof(f1956,plain,(
  ( ! [X2,X0,X1] : (s(X2,happ(s(fun(bool,X2),X0),s(bool,'F_0'))) != s(X2,happ(s(fun(bool,X2),X1),s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK4(X0,X1,bool,X2))) )),
  inference(superposition,[],[f59,f50])).
fof(f4930,plain,(
  ( ! [X173,X171,X169,X174,X172,X170] : (s(fun(X169,X170),X174) = s(fun(X169,X170),happ(s(fun(X171,fun(X169,X170)),X172),s(X171,X173))) | s(X170,happ(s(fun(X169,X170),X174),s(X169,sK4('const.combin.o_3'(s(fun(fun(X169,X170),fun(X169,X170)),'const.combin.I_0'),s(fun(X171,fun(X169,X170)),X172),s(X171,X173)),X174,X169,X170)))) != s(X170,happ(s(fun(X169,X170),happ(s(fun(X171,fun(X169,X170)),X172),s(X171,X173))),s(X169,sK4('const.combin.o_3'(s(fun(fun(X169,X170),fun(X169,X170)),'const.combin.I_0'),s(fun(X171,fun(X169,X170)),X172),s(X171,X173)),X174,X169,X170))))) )),
  inference(forward_demodulation,[],[f4875,f1305])).
fof(f4875,plain,(
  ( ! [X173,X171,X169,X174,X172,X170] : (s(X170,happ(s(fun(X169,X170),X174),s(X169,sK4('const.combin.o_3'(s(fun(fun(X169,X170),fun(X169,X170)),'const.combin.I_0'),s(fun(X171,fun(X169,X170)),X172),s(X171,X173)),X174,X169,X170)))) != s(X170,happ(s(fun(X169,X170),happ(s(fun(X171,fun(X169,X170)),X172),s(X171,X173))),s(X169,sK4('const.combin.o_3'(s(fun(fun(X169,X170),fun(X169,X170)),'const.combin.I_0'),s(fun(X171,fun(X169,X170)),X172),s(X171,X173)),X174,X169,X170)))) | s(fun(X169,X170),X174) = s(fun(X169,X170),'const.combin.o_3'(s(fun(fun(X169,X170),fun(X169,X170)),'const.combin.I_0'),s(fun(X171,fun(X169,X170)),X172),s(X171,X173)))) )),
  inference(superposition,[],[f59,f1305])).
fof(f4929,plain,(
  ( ! [X167,X165,X163,X166,X164,X168] : (s(fun(X163,X164),X168) = s(fun(X163,X164),happ(s(fun(X165,fun(X163,X164)),X166),s(X165,X167))) | s(X164,happ(s(fun(X163,X164),X168),s(X163,sK4(X168,'const.combin.o_3'(s(fun(fun(X163,X164),fun(X163,X164)),'const.combin.I_0'),s(fun(X165,fun(X163,X164)),X166),s(X165,X167)),X163,X164)))) != s(X164,happ(s(fun(X163,X164),happ(s(fun(X165,fun(X163,X164)),X166),s(X165,X167))),s(X163,sK4(X168,'const.combin.o_3'(s(fun(fun(X163,X164),fun(X163,X164)),'const.combin.I_0'),s(fun(X165,fun(X163,X164)),X166),s(X165,X167)),X163,X164))))) )),
  inference(forward_demodulation,[],[f4874,f1305])).
fof(f4874,plain,(
  ( ! [X167,X165,X163,X166,X164,X168] : (s(X164,happ(s(fun(X163,X164),X168),s(X163,sK4(X168,'const.combin.o_3'(s(fun(fun(X163,X164),fun(X163,X164)),'const.combin.I_0'),s(fun(X165,fun(X163,X164)),X166),s(X165,X167)),X163,X164)))) != s(X164,happ(s(fun(X163,X164),happ(s(fun(X165,fun(X163,X164)),X166),s(X165,X167))),s(X163,sK4(X168,'const.combin.o_3'(s(fun(fun(X163,X164),fun(X163,X164)),'const.combin.I_0'),s(fun(X165,fun(X163,X164)),X166),s(X165,X167)),X163,X164)))) | s(fun(X163,X164),X168) = s(fun(X163,X164),'const.combin.o_3'(s(fun(fun(X163,X164),fun(X163,X164)),'const.combin.I_0'),s(fun(X165,fun(X163,X164)),X166),s(X165,X167)))) )),
  inference(superposition,[],[f59,f1305])).
fof(f4900,plain,(
  ( ! [X59,X57,X58] : (s(bool,'T_0') = s(bool,happ(s(fun(X57,bool),X58),s(X57,X59))) | s(bool,'F_0') = s(bool,'const.combin.o_3'(s(fun(bool,bool),'const.combin.I_0'),s(fun(X57,bool),X58),s(X57,X59)))) )),
  inference(forward_demodulation,[],[f4846,f51])).
fof(f4846,plain,(
  ( ! [X59,X57,X58] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,happ(s(fun(X57,bool),X58),s(X57,X59))))) | s(bool,'F_0') = s(bool,'const.combin.o_3'(s(fun(bool,bool),'const.combin.I_0'),s(fun(X57,bool),X58),s(X57,X59)))) )),
  inference(superposition,[],[f75,f1305])).
fof(f4825,plain,(
  ( ! [X14,X17,X15,X16] : (s(bool,X17) = s(bool,happ(s(fun(X14,bool),X15),s(X14,X16))) | s(bool,'T_0') = s(bool,X17) | s(bool,'T_0') = s(bool,'const.combin.o_3'(s(fun(bool,bool),'const.combin.I_0'),s(fun(X14,bool),X15),s(X14,X16)))) )),
  inference(superposition,[],[f1305,f73])).
fof(f4824,plain,(
  ( ! [X12,X10,X13,X11] : (s(bool,'T_0') = s(bool,happ(s(fun(X10,bool),X11),s(X10,X12))) | s(bool,X13) = s(bool,'const.combin.o_3'(s(fun(bool,bool),'const.combin.I_0'),s(fun(X10,bool),X11),s(X10,X12))) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f1305,f73])).
fof(f3871,plain,(
  ( ! [X144,X142,X140,X145,X143,X141] : (s(fun(X140,X141),X145) = s(fun(X140,X141),happ(s(fun(X142,fun(X140,X141)),X143),s(X142,X144))) | s(X141,happ(s(fun(X140,X141),X145),s(X140,sK4('const.combin.o_3'(s(fun(X142,fun(X140,X141)),X143),s(fun(X142,X142),'const.combin.I_0'),s(X142,X144)),X145,X140,X141)))) != s(X141,happ(s(fun(X140,X141),happ(s(fun(X142,fun(X140,X141)),X143),s(X142,X144))),s(X140,sK4('const.combin.o_3'(s(fun(X142,fun(X140,X141)),X143),s(fun(X142,X142),'const.combin.I_0'),s(X142,X144)),X145,X140,X141))))) )),
  inference(forward_demodulation,[],[f3827,f1293])).
fof(f3827,plain,(
  ( ! [X144,X142,X140,X145,X143,X141] : (s(X141,happ(s(fun(X140,X141),X145),s(X140,sK4('const.combin.o_3'(s(fun(X142,fun(X140,X141)),X143),s(fun(X142,X142),'const.combin.I_0'),s(X142,X144)),X145,X140,X141)))) != s(X141,happ(s(fun(X140,X141),happ(s(fun(X142,fun(X140,X141)),X143),s(X142,X144))),s(X140,sK4('const.combin.o_3'(s(fun(X142,fun(X140,X141)),X143),s(fun(X142,X142),'const.combin.I_0'),s(X142,X144)),X145,X140,X141)))) | s(fun(X140,X141),X145) = s(fun(X140,X141),'const.combin.o_3'(s(fun(X142,fun(X140,X141)),X143),s(fun(X142,X142),'const.combin.I_0'),s(X142,X144)))) )),
  inference(superposition,[],[f59,f1293])).
fof(f3870,plain,(
  ( ! [X134,X138,X136,X135,X139,X137] : (s(fun(X134,X135),X139) = s(fun(X134,X135),happ(s(fun(X136,fun(X134,X135)),X137),s(X136,X138))) | s(X135,happ(s(fun(X134,X135),X139),s(X134,sK4(X139,'const.combin.o_3'(s(fun(X136,fun(X134,X135)),X137),s(fun(X136,X136),'const.combin.I_0'),s(X136,X138)),X134,X135)))) != s(X135,happ(s(fun(X134,X135),happ(s(fun(X136,fun(X134,X135)),X137),s(X136,X138))),s(X134,sK4(X139,'const.combin.o_3'(s(fun(X136,fun(X134,X135)),X137),s(fun(X136,X136),'const.combin.I_0'),s(X136,X138)),X134,X135))))) )),
  inference(forward_demodulation,[],[f3826,f1293])).
fof(f3826,plain,(
  ( ! [X134,X138,X136,X135,X139,X137] : (s(X135,happ(s(fun(X134,X135),X139),s(X134,sK4(X139,'const.combin.o_3'(s(fun(X136,fun(X134,X135)),X137),s(fun(X136,X136),'const.combin.I_0'),s(X136,X138)),X134,X135)))) != s(X135,happ(s(fun(X134,X135),happ(s(fun(X136,fun(X134,X135)),X137),s(X136,X138))),s(X134,sK4(X139,'const.combin.o_3'(s(fun(X136,fun(X134,X135)),X137),s(fun(X136,X136),'const.combin.I_0'),s(X136,X138)),X134,X135)))) | s(fun(X134,X135),X139) = s(fun(X134,X135),'const.combin.o_3'(s(fun(X136,fun(X134,X135)),X137),s(fun(X136,X136),'const.combin.I_0'),s(X136,X138)))) )),
  inference(superposition,[],[f59,f1293])).
fof(f3805,plain,(
  ( ! [X54,X56,X55] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,happ(s(fun(X54,bool),X55),s(X54,X56))))) | s(bool,'F_0') = s(bool,'const.combin.o_3'(s(fun(X54,bool),X55),s(fun(X54,X54),'const.combin.I_0'),s(X54,X56)))) )),
  inference(superposition,[],[f96,f1293])).
fof(f3846,plain,(
  ( ! [X50,X48,X49] : (s(bool,'T_0') = s(bool,happ(s(fun(X48,bool),X49),s(X48,X50))) | s(bool,'F_0') = s(bool,'const.combin.o_3'(s(fun(X48,bool),X49),s(fun(X48,X48),'const.combin.I_0'),s(X48,X50)))) )),
  inference(forward_demodulation,[],[f3803,f51])).
fof(f3803,plain,(
  ( ! [X50,X48,X49] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,happ(s(fun(X48,bool),X49),s(X48,X50))))) | s(bool,'F_0') = s(bool,'const.combin.o_3'(s(fun(X48,bool),X49),s(fun(X48,X48),'const.combin.I_0'),s(X48,X50)))) )),
  inference(superposition,[],[f75,f1293])).
fof(f3786,plain,(
  ( ! [X14,X12,X13,X11] : (s(bool,X14) = s(bool,happ(s(fun(X11,bool),X12),s(X11,X13))) | s(bool,'T_0') = s(bool,X14) | s(bool,'T_0') = s(bool,'const.combin.o_3'(s(fun(X11,bool),X12),s(fun(X11,X11),'const.combin.I_0'),s(X11,X13)))) )),
  inference(superposition,[],[f1293,f73])).
fof(f3785,plain,(
  ( ! [X10,X8,X7,X9] : (s(bool,'T_0') = s(bool,happ(s(fun(X7,bool),X8),s(X7,X9))) | s(bool,X10) = s(bool,'const.combin.o_3'(s(fun(X7,bool),X8),s(fun(X7,X7),'const.combin.I_0'),s(X7,X9))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f1293,f73])).
fof(f3784,plain,(
  ( ! [X6,X4,X5] : (s(bool,'F_0') = s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))) | s(bool,'T_0') = s(bool,'const.combin.o_3'(s(fun(X4,bool),X5),s(fun(X4,X4),'const.combin.I_0'),s(X4,X6)))) )),
  inference(superposition,[],[f1293,f50])).
fof(f3353,plain,(
  ( ! [X30] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),X30),s(bool,sK4('const.combin.I_0',X30,bool,bool)))) | s(fun(bool,bool),'const.combin.I_0') = s(fun(bool,bool),X30) | s(bool,'F_0') = s(bool,sK4('const.combin.I_0',X30,bool,bool))) )),
  inference(superposition,[],[f59,f96])).
fof(f3352,plain,(
  ( ! [X29] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),X29),s(bool,sK4(X29,'const.combin.I_0',bool,bool)))) | s(fun(bool,bool),'const.combin.I_0') = s(fun(bool,bool),X29) | s(bool,'F_0') = s(bool,sK4(X29,'const.combin.I_0',bool,bool))) )),
  inference(superposition,[],[f59,f96])).
fof(f3350,plain,(
  ( ! [X24,X23,X25] : (s(X24,happ(s(fun(bool,X24),X25),s(bool,'T_0'))) = s(X24,'const.combin.o_3'(s(fun(bool,X24),X25),s(fun(bool,bool),'const.combin.I_0'),s(bool,X23))) | s(bool,'F_0') = s(bool,X23)) )),
  inference(superposition,[],[f60,f96])).
fof(f3332,plain,(
  ( ! [X14,X12,X13] : (s(bool,'T_0') = s(bool,'const.combin.o_3'(s(fun(bool,bool),'const.combin.I_0'),s(fun(X12,bool),X13),s(X12,X14))) | s(bool,'F_0') = s(bool,happ(s(fun(X12,bool),X13),s(X12,X14)))) )),
  inference(superposition,[],[f96,f60])).
fof(f3323,plain,(
  ( ! [X30,X33,X31,X34,X32] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'const.combin.o_3'(s(fun(X31,bool),X32),s(fun(X30,X31),X33),s(X30,X34))))) | s(bool,'F_0') = s(bool,happ(s(fun(X30,bool),'const.combin.o_2'(s(fun(X31,bool),X32),s(fun(X30,X31),X33))),s(X30,X34)))) )),
  inference(superposition,[],[f96,f61])).
fof(f3322,plain,(
  ( ! [X28,X26,X29,X27,X25] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'const.combin.o_3'(s(fun(X25,bool),X26),s(fun(X27,X25),X28),s(X27,X29))))) | s(bool,'F_0') = s(bool,happ(s(fun(X25,bool),X26),s(X25,happ(s(fun(X27,X25),X28),s(X27,X29)))))) )),
  inference(superposition,[],[f96,f60])).
fof(f2594,plain,(
  ( ! [X28,X26,X24,X27,X25] : (s(bool,'T_0') = s(bool,'const.combin.o_3'(s(fun(X25,bool),X26),s(fun(X24,X25),X27),s(X24,X28))) | s(bool,'F_0') = s(bool,happ(s(fun(X24,bool),'const.combin.o_2'(s(fun(X25,bool),X26),s(fun(X24,X25),X27))),s(X24,X28)))) )),
  inference(forward_demodulation,[],[f2555,f51])).
fof(f2555,plain,(
  ( ! [X28,X26,X24,X27,X25] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'const.combin.o_3'(s(fun(X25,bool),X26),s(fun(X24,X25),X27),s(X24,X28))))) | s(bool,'F_0') = s(bool,happ(s(fun(X24,bool),'const.combin.o_2'(s(fun(X25,bool),X26),s(fun(X24,X25),X27))),s(X24,X28)))) )),
  inference(superposition,[],[f75,f61])).
fof(f2377,plain,(
  s(bool,'F_0') = s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | ($spl9 | $spl13)),
  inference(unit_resulting_resolution,[],[f86,f418,f73])).
fof(f473,plain,(
  ( ! [X21,X19,X20] : (s(bool,'F_0') != s(bool,f(s(X19,sK3(X20,X21,bool,X19)))) | s(fun(X19,bool),X20) = s(fun(X19,bool),X21) | s(bool,'T_0') = s(bool,g(s(X19,sK3(X20,X21,bool,X19))))) )),
  inference(superposition,[],[f57,f50])).
fof(f1324,plain,(
  ( ! [X47,X45,X48,X46,X49] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'const.combin.o_3'(s(fun(X45,bool),X46),s(fun(X47,X45),X48),s(X47,X49))))) | s(bool,'F_0') = s(bool,happ(s(fun(X45,bool),X46),s(X45,happ(s(fun(X47,X45),X48),s(X47,X49)))))) )),
  inference(superposition,[],[f75,f60])).
fof(f1534,plain,(
  s(bool,'F_0') = s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | ($spl9 | $spl13)),
  inference(unit_resulting_resolution,[],[f86,f418,f73])).
fof(f1556,plain,(
  ( ! [X5] : (s(bool,X5) = s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,X5)) ) | $spl13),
  inference(trivial_inequality_removal,[],[f1551])).
fof(f1551,plain,(
  ( ! [X5] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X5) = s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,X5)) ) | $spl13),
  inference(superposition,[],[f418,f73])).
fof(f1969,plain,(
  ( ! [X17,X15,X18,X16] : (s(bool,X18) != s(bool,happ(s(fun(X15,bool),X17),s(X15,sK4(X16,X17,X15,bool)))) | s(fun(X15,bool),X16) = s(fun(X15,bool),X17) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,happ(s(fun(X15,bool),X16),s(X15,sK4(X16,X17,X15,bool))))) )),
  inference(superposition,[],[f59,f73])).
fof(f1968,plain,(
  ( ! [X14,X12,X13,X11] : (s(bool,'T_0') != s(bool,happ(s(fun(X11,bool),X13),s(X11,sK4(X12,X13,X11,bool)))) | s(fun(X11,bool),X12) = s(fun(X11,bool),X13) | s(bool,X14) = s(bool,happ(s(fun(X11,bool),X12),s(X11,sK4(X12,X13,X11,bool)))) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f59,f73])).
fof(f1967,plain,(
  ( ! [X10,X8,X9] : (s(bool,'F_0') != s(bool,happ(s(fun(X8,bool),X10),s(X8,sK4(X9,X10,X8,bool)))) | s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,'T_0') = s(bool,happ(s(fun(X8,bool),X9),s(X8,sK4(X9,X10,X8,bool))))) )),
  inference(superposition,[],[f59,f50])).
fof(f1966,plain,(
  ( ! [X6,X4,X2,X7,X5,X3] : (s(X2,happ(s(fun(X3,X2),X7),s(X3,sK4('const.combin.o_2'(s(fun(X4,X2),X5),s(fun(X3,X4),X6)),X7,X3,X2)))) != s(X2,'const.combin.o_3'(s(fun(X4,X2),X5),s(fun(X3,X4),X6),s(X3,sK4('const.combin.o_2'(s(fun(X4,X2),X5),s(fun(X3,X4),X6)),X7,X3,X2)))) | s(fun(X3,X2),X7) = s(fun(X3,X2),'const.combin.o_2'(s(fun(X4,X2),X5),s(fun(X3,X4),X6)))) )),
  inference(superposition,[],[f59,f61])).
fof(f1965,plain,(
  ( ! [X0,X1] : (s(X0,sK4('const.combin.I_0',X1,X0,X0)) != s(X0,happ(s(fun(X0,X0),X1),s(X0,sK4('const.combin.I_0',X1,X0,X0)))) | s(fun(X0,X0),'const.combin.I_0') = s(fun(X0,X0),X1)) )),
  inference(superposition,[],[f59,f68])).
fof(f1958,plain,(
  ( ! [X10,X8,X7,X9] : (s(X9,happ(s(fun(bool,X9),X7),s(bool,X10))) != s(X9,happ(s(fun(bool,X9),X8),s(bool,X10))) | s(fun(bool,X9),X7) = s(fun(bool,X9),X8) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,sK4(X7,X8,bool,X9))) )),
  inference(superposition,[],[f59,f73])).
fof(f1957,plain,(
  ( ! [X6,X4,X5,X3] : (s(X5,happ(s(fun(bool,X5),X3),s(bool,'T_0'))) != s(X5,happ(s(fun(bool,X5),X4),s(bool,'T_0'))) | s(fun(bool,X5),X3) = s(fun(bool,X5),X4) | s(bool,X6) = s(bool,sK4(X3,X4,bool,X5)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f59,f73])).
fof(f1986,plain,(
  ( ! [X23,X21,X19,X17,X22,X20,X18,X16] : (s(fun(X16,X17),X23) = s(fun(X16,X17),'const.combin.o_3'(s(fun(X19,fun(X16,X17)),X20),s(fun(X18,X19),X21),s(X18,X22))) | s(X17,happ(s(fun(X16,X17),X23),s(X16,sK4(happ(s(fun(X18,fun(X16,X17)),'const.combin.o_2'(s(fun(X19,fun(X16,X17)),X20),s(fun(X18,X19),X21))),s(X18,X22)),X23,X16,X17)))) != s(X17,happ(s(fun(X16,X17),'const.combin.o_3'(s(fun(X19,fun(X16,X17)),X20),s(fun(X18,X19),X21),s(X18,X22))),s(X16,sK4(happ(s(fun(X18,fun(X16,X17)),'const.combin.o_2'(s(fun(X19,fun(X16,X17)),X20),s(fun(X18,X19),X21))),s(X18,X22)),X23,X16,X17))))) )),
  inference(forward_demodulation,[],[f1955,f61])).
fof(f1955,plain,(
  ( ! [X23,X21,X19,X17,X22,X20,X18,X16] : (s(X17,happ(s(fun(X16,X17),X23),s(X16,sK4(happ(s(fun(X18,fun(X16,X17)),'const.combin.o_2'(s(fun(X19,fun(X16,X17)),X20),s(fun(X18,X19),X21))),s(X18,X22)),X23,X16,X17)))) != s(X17,happ(s(fun(X16,X17),'const.combin.o_3'(s(fun(X19,fun(X16,X17)),X20),s(fun(X18,X19),X21),s(X18,X22))),s(X16,sK4(happ(s(fun(X18,fun(X16,X17)),'const.combin.o_2'(s(fun(X19,fun(X16,X17)),X20),s(fun(X18,X19),X21))),s(X18,X22)),X23,X16,X17)))) | s(fun(X16,X17),X23) = s(fun(X16,X17),happ(s(fun(X18,fun(X16,X17)),'const.combin.o_2'(s(fun(X19,fun(X16,X17)),X20),s(fun(X18,X19),X21))),s(X18,X22)))) )),
  inference(superposition,[],[f59,f61])).
fof(f1985,plain,(
  ( ! [X14,X12,X10,X8,X15,X13,X11,X9] : (s(fun(X8,X9),X15) = s(fun(X8,X9),'const.combin.o_3'(s(fun(X10,fun(X8,X9)),X11),s(fun(X12,X10),X13),s(X12,X14))) | s(X9,happ(s(fun(X8,X9),X15),s(X8,sK4(happ(s(fun(X10,fun(X8,X9)),X11),s(X10,happ(s(fun(X12,X10),X13),s(X12,X14)))),X15,X8,X9)))) != s(X9,happ(s(fun(X8,X9),'const.combin.o_3'(s(fun(X10,fun(X8,X9)),X11),s(fun(X12,X10),X13),s(X12,X14))),s(X8,sK4(happ(s(fun(X10,fun(X8,X9)),X11),s(X10,happ(s(fun(X12,X10),X13),s(X12,X14)))),X15,X8,X9))))) )),
  inference(forward_demodulation,[],[f1954,f60])).
fof(f1954,plain,(
  ( ! [X14,X12,X10,X8,X15,X13,X11,X9] : (s(X9,happ(s(fun(X8,X9),X15),s(X8,sK4(happ(s(fun(X10,fun(X8,X9)),X11),s(X10,happ(s(fun(X12,X10),X13),s(X12,X14)))),X15,X8,X9)))) != s(X9,happ(s(fun(X8,X9),'const.combin.o_3'(s(fun(X10,fun(X8,X9)),X11),s(fun(X12,X10),X13),s(X12,X14))),s(X8,sK4(happ(s(fun(X10,fun(X8,X9)),X11),s(X10,happ(s(fun(X12,X10),X13),s(X12,X14)))),X15,X8,X9)))) | s(fun(X8,X9),X15) = s(fun(X8,X9),happ(s(fun(X10,fun(X8,X9)),X11),s(X10,happ(s(fun(X12,X10),X13),s(X12,X14)))))) )),
  inference(superposition,[],[f59,f60])).
fof(f1984,plain,(
  ( ! [X6,X4,X7,X5] : (s(X5,happ(s(fun(X4,X5),X6),s(X4,sK4(happ(s(fun(fun(X4,X5),fun(X4,X5)),'const.combin.I_0'),s(fun(X4,X5),X6)),X7,X4,X5)))) != s(X5,happ(s(fun(X4,X5),X7),s(X4,sK4(happ(s(fun(fun(X4,X5),fun(X4,X5)),'const.combin.I_0'),s(fun(X4,X5),X6)),X7,X4,X5)))) | s(fun(X4,X5),X6) = s(fun(X4,X5),X7)) )),
  inference(forward_demodulation,[],[f1953,f68])).
fof(f1953,plain,(
  ( ! [X6,X4,X7,X5] : (s(X5,happ(s(fun(X4,X5),X6),s(X4,sK4(happ(s(fun(fun(X4,X5),fun(X4,X5)),'const.combin.I_0'),s(fun(X4,X5),X6)),X7,X4,X5)))) != s(X5,happ(s(fun(X4,X5),X7),s(X4,sK4(happ(s(fun(fun(X4,X5),fun(X4,X5)),'const.combin.I_0'),s(fun(X4,X5),X6)),X7,X4,X5)))) | s(fun(X4,X5),X7) = s(fun(X4,X5),happ(s(fun(fun(X4,X5),fun(X4,X5)),'const.combin.I_0'),s(fun(X4,X5),X6)))) )),
  inference(superposition,[],[f59,f68])).
fof(f1983,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,happ(s(fun(X0,X1),X2),s(X0,sK4('const.combin.I_1'(s(fun(X0,X1),X2)),X3,X0,X1)))) != s(X1,happ(s(fun(X0,X1),X3),s(X0,sK4('const.combin.I_1'(s(fun(X0,X1),X2)),X3,X0,X1)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(forward_demodulation,[],[f1952,f51])).
fof(f1952,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,happ(s(fun(X0,X1),X2),s(X0,sK4('const.combin.I_1'(s(fun(X0,X1),X2)),X3,X0,X1)))) != s(X1,happ(s(fun(X0,X1),X3),s(X0,sK4('const.combin.I_1'(s(fun(X0,X1),X2)),X3,X0,X1)))) | s(fun(X0,X1),X3) = s(fun(X0,X1),'const.combin.I_1'(s(fun(X0,X1),X2)))) )),
  inference(superposition,[],[f59,f51])).
fof(f1945,plain,(
  ( ! [X17,X15,X18,X16] : (s(bool,X18) != s(bool,happ(s(fun(X15,bool),X17),s(X15,sK4(X17,X16,X15,bool)))) | s(fun(X15,bool),X16) = s(fun(X15,bool),X17) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,happ(s(fun(X15,bool),X16),s(X15,sK4(X17,X16,X15,bool))))) )),
  inference(superposition,[],[f59,f73])).
fof(f1944,plain,(
  ( ! [X14,X12,X13,X11] : (s(bool,'T_0') != s(bool,happ(s(fun(X11,bool),X13),s(X11,sK4(X13,X12,X11,bool)))) | s(fun(X11,bool),X12) = s(fun(X11,bool),X13) | s(bool,X14) = s(bool,happ(s(fun(X11,bool),X12),s(X11,sK4(X13,X12,X11,bool)))) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f59,f73])).
fof(f1943,plain,(
  ( ! [X10,X8,X9] : (s(bool,'F_0') != s(bool,happ(s(fun(X8,bool),X10),s(X8,sK4(X10,X9,X8,bool)))) | s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,'T_0') = s(bool,happ(s(fun(X8,bool),X9),s(X8,sK4(X10,X9,X8,bool))))) )),
  inference(superposition,[],[f59,f50])).
fof(f1941,plain,(
  ( ! [X0,X1] : (s(X0,sK4(X1,'const.combin.I_0',X0,X0)) != s(X0,happ(s(fun(X0,X0),X1),s(X0,sK4(X1,'const.combin.I_0',X0,X0)))) | s(fun(X0,X0),'const.combin.I_0') = s(fun(X0,X0),X1)) )),
  inference(superposition,[],[f59,f68])).
fof(f1980,plain,(
  ( ! [X23,X21,X19,X17,X22,X20,X18,X16] : (s(fun(X16,X17),X23) = s(fun(X16,X17),'const.combin.o_3'(s(fun(X19,fun(X16,X17)),X20),s(fun(X18,X19),X21),s(X18,X22))) | s(X17,happ(s(fun(X16,X17),X23),s(X16,sK4(X23,happ(s(fun(X18,fun(X16,X17)),'const.combin.o_2'(s(fun(X19,fun(X16,X17)),X20),s(fun(X18,X19),X21))),s(X18,X22)),X16,X17)))) != s(X17,happ(s(fun(X16,X17),'const.combin.o_3'(s(fun(X19,fun(X16,X17)),X20),s(fun(X18,X19),X21),s(X18,X22))),s(X16,sK4(X23,happ(s(fun(X18,fun(X16,X17)),'const.combin.o_2'(s(fun(X19,fun(X16,X17)),X20),s(fun(X18,X19),X21))),s(X18,X22)),X16,X17))))) )),
  inference(forward_demodulation,[],[f1940,f61])).
fof(f1940,plain,(
  ( ! [X23,X21,X19,X17,X22,X20,X18,X16] : (s(X17,happ(s(fun(X16,X17),X23),s(X16,sK4(X23,happ(s(fun(X18,fun(X16,X17)),'const.combin.o_2'(s(fun(X19,fun(X16,X17)),X20),s(fun(X18,X19),X21))),s(X18,X22)),X16,X17)))) != s(X17,happ(s(fun(X16,X17),'const.combin.o_3'(s(fun(X19,fun(X16,X17)),X20),s(fun(X18,X19),X21),s(X18,X22))),s(X16,sK4(X23,happ(s(fun(X18,fun(X16,X17)),'const.combin.o_2'(s(fun(X19,fun(X16,X17)),X20),s(fun(X18,X19),X21))),s(X18,X22)),X16,X17)))) | s(fun(X16,X17),X23) = s(fun(X16,X17),happ(s(fun(X18,fun(X16,X17)),'const.combin.o_2'(s(fun(X19,fun(X16,X17)),X20),s(fun(X18,X19),X21))),s(X18,X22)))) )),
  inference(superposition,[],[f59,f61])).
fof(f1979,plain,(
  ( ! [X14,X12,X10,X8,X15,X13,X11,X9] : (s(fun(X8,X9),X15) = s(fun(X8,X9),'const.combin.o_3'(s(fun(X10,fun(X8,X9)),X11),s(fun(X12,X10),X13),s(X12,X14))) | s(X9,happ(s(fun(X8,X9),X15),s(X8,sK4(X15,happ(s(fun(X10,fun(X8,X9)),X11),s(X10,happ(s(fun(X12,X10),X13),s(X12,X14)))),X8,X9)))) != s(X9,happ(s(fun(X8,X9),'const.combin.o_3'(s(fun(X10,fun(X8,X9)),X11),s(fun(X12,X10),X13),s(X12,X14))),s(X8,sK4(X15,happ(s(fun(X10,fun(X8,X9)),X11),s(X10,happ(s(fun(X12,X10),X13),s(X12,X14)))),X8,X9))))) )),
  inference(forward_demodulation,[],[f1939,f60])).
fof(f1939,plain,(
  ( ! [X14,X12,X10,X8,X15,X13,X11,X9] : (s(X9,happ(s(fun(X8,X9),X15),s(X8,sK4(X15,happ(s(fun(X10,fun(X8,X9)),X11),s(X10,happ(s(fun(X12,X10),X13),s(X12,X14)))),X8,X9)))) != s(X9,happ(s(fun(X8,X9),'const.combin.o_3'(s(fun(X10,fun(X8,X9)),X11),s(fun(X12,X10),X13),s(X12,X14))),s(X8,sK4(X15,happ(s(fun(X10,fun(X8,X9)),X11),s(X10,happ(s(fun(X12,X10),X13),s(X12,X14)))),X8,X9)))) | s(fun(X8,X9),X15) = s(fun(X8,X9),happ(s(fun(X10,fun(X8,X9)),X11),s(X10,happ(s(fun(X12,X10),X13),s(X12,X14)))))) )),
  inference(superposition,[],[f59,f60])).
fof(f1978,plain,(
  ( ! [X6,X4,X7,X5] : (s(X5,happ(s(fun(X4,X5),X6),s(X4,sK4(X7,happ(s(fun(fun(X4,X5),fun(X4,X5)),'const.combin.I_0'),s(fun(X4,X5),X6)),X4,X5)))) != s(X5,happ(s(fun(X4,X5),X7),s(X4,sK4(X7,happ(s(fun(fun(X4,X5),fun(X4,X5)),'const.combin.I_0'),s(fun(X4,X5),X6)),X4,X5)))) | s(fun(X4,X5),X6) = s(fun(X4,X5),X7)) )),
  inference(forward_demodulation,[],[f1938,f68])).
fof(f1938,plain,(
  ( ! [X6,X4,X7,X5] : (s(X5,happ(s(fun(X4,X5),X6),s(X4,sK4(X7,happ(s(fun(fun(X4,X5),fun(X4,X5)),'const.combin.I_0'),s(fun(X4,X5),X6)),X4,X5)))) != s(X5,happ(s(fun(X4,X5),X7),s(X4,sK4(X7,happ(s(fun(fun(X4,X5),fun(X4,X5)),'const.combin.I_0'),s(fun(X4,X5),X6)),X4,X5)))) | s(fun(X4,X5),X7) = s(fun(X4,X5),happ(s(fun(fun(X4,X5),fun(X4,X5)),'const.combin.I_0'),s(fun(X4,X5),X6)))) )),
  inference(superposition,[],[f59,f68])).
fof(f1977,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,happ(s(fun(X0,X1),X2),s(X0,sK4(X3,'const.combin.I_1'(s(fun(X0,X1),X2)),X0,X1)))) != s(X1,happ(s(fun(X0,X1),X3),s(X0,sK4(X3,'const.combin.I_1'(s(fun(X0,X1),X2)),X0,X1)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(forward_demodulation,[],[f1937,f51])).
fof(f1937,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,happ(s(fun(X0,X1),X2),s(X0,sK4(X3,'const.combin.I_1'(s(fun(X0,X1),X2)),X0,X1)))) != s(X1,happ(s(fun(X0,X1),X3),s(X0,sK4(X3,'const.combin.I_1'(s(fun(X0,X1),X2)),X0,X1)))) | s(fun(X0,X1),X3) = s(fun(X0,X1),'const.combin.I_1'(s(fun(X0,X1),X2)))) )),
  inference(superposition,[],[f59,f51])).
fof(f1766,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X6,happ(s(fun(X0,X6),X7),s(X0,'const.combin.o_3'(s(fun(X2,X0),X3),s(fun(X1,X2),X4),s(X1,X5))))) = s(X6,'const.combin.o_3'(s(fun(X0,X6),X7),s(fun(X1,X0),'const.combin.o_2'(s(fun(X2,X0),X3),s(fun(X1,X2),X4))),s(X1,X5)))) )),
  inference(superposition,[],[f60,f61])).
fof(f1759,plain,(
  ( ! [X30,X28,X26,X29,X27,X25] : (s(bool,X30) = s(bool,'const.combin.o_3'(s(fun(X26,bool),X27),s(fun(X25,X26),X28),s(X25,X29))) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,happ(s(fun(X25,bool),'const.combin.o_2'(s(fun(X26,bool),X27),s(fun(X25,X26),X28))),s(X25,X29)))) )),
  inference(superposition,[],[f61,f73])).
fof(f1758,plain,(
  ( ! [X24,X23,X21,X19,X22,X20] : (s(bool,'T_0') = s(bool,'const.combin.o_3'(s(fun(X20,bool),X21),s(fun(X19,X20),X22),s(X19,X23))) | s(bool,X24) = s(bool,happ(s(fun(X19,bool),'const.combin.o_2'(s(fun(X20,bool),X21),s(fun(X19,X20),X22))),s(X19,X23))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f61,f73])).
fof(f1757,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s(bool,'T_0') = s(bool,happ(s(fun(X14,bool),'const.combin.o_2'(s(fun(X15,bool),X16),s(fun(X14,X15),X17))),s(X14,X18))) | s(bool,'F_0') = s(bool,'const.combin.o_3'(s(fun(X15,bool),X16),s(fun(X14,X15),X17),s(X14,X18)))) )),
  inference(superposition,[],[f61,f50])).
fof(f1756,plain,(
  ( ! [X6,X12,X10,X8,X7,X13,X11,X9] : (s(X6,'const.combin.o_3'(s(fun(X8,X6),X9),s(fun(X7,X8),X10),s(X7,happ(s(fun(X11,X7),X12),s(X11,X13))))) = s(X6,'const.combin.o_3'(s(fun(X7,X6),'const.combin.o_2'(s(fun(X8,X6),X9),s(fun(X7,X8),X10))),s(fun(X11,X7),X12),s(X11,X13)))) )),
  inference(superposition,[],[f61,f60])).
fof(f75,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X2))) | s(bool,'F_0') = s(bool,X2)) )),
  inference(superposition,[],[f50,f51])).
fof(f96,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X2))) | s(bool,'F_0') = s(bool,X2)) )),
  inference(superposition,[],[f68,f50])).
fof(f122,plain,(
  ( ! [X4] : (s(bool,'T_0') = s(bool,X4) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X4)))) )),
  inference(forward_demodulation,[],[f114,f51])).
fof(f114,plain,(
  ( ! [X4] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X4))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X4)))) )),
  inference(superposition,[],[f75,f68])).
fof(f1296,plain,(
  ( ! [X19,X17,X15,X18,X16] : (s(X18,happ(s(fun(bool,X18),X19),s(bool,'F_0'))) = s(X18,'const.combin.o_3'(s(fun(bool,X18),X19),s(fun(X15,bool),X16),s(X15,X17))) | s(bool,'T_0') = s(bool,happ(s(fun(X15,bool),X16),s(X15,X17)))) )),
  inference(superposition,[],[f60,f50])).
fof(f1307,plain,(
  ( ! [X14,X12,X15,X13,X11] : (s(bool,'T_0') = s(bool,happ(s(fun(X11,bool),X12),s(X11,happ(s(fun(X13,X11),X14),s(X13,X15))))) | s(bool,'F_0') = s(bool,'const.combin.o_3'(s(fun(X11,bool),X12),s(fun(X13,X11),X14),s(X13,X15)))) )),
  inference(superposition,[],[f60,f50])).
fof(f471,plain,(
  ( ! [X14,X12,X13,X11] : (s(bool,X14) != s(bool,f(s(X11,sK3(X12,X13,bool,X11)))) | s(fun(X11,bool),X12) = s(fun(X11,bool),X13) | s(bool,'T_0') = s(bool,g(s(X11,sK3(X12,X13,bool,X11)))) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f57,f73])).
fof(f469,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(X3,sK3(X4,X5,bool,X3)))) | s(fun(X3,bool),X4) = s(fun(X3,bool),X5) | s(bool,X6) = s(bool,g(s(X3,sK3(X4,X5,bool,X3)))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f57,f73])).
fof(f348,plain,(
  ( ! [X14,X12,X13] : (s(fun(X12,X12),'const.combin.I_0') = s(fun(X12,X12),'const.combin.o_2'(s(fun(X12,X12),X13),s(fun(X12,X12),X13))) | s(bool,X14) = s(bool,'const.relation.INVOL_1'(s(fun(X12,X12),X13))) | s(bool,'T_0') = s(bool,X14)) )),
  inference(subsumption_resolution,[],[f218,f47])).
fof(f218,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,'T_0')) | s(fun(X12,X12),'const.combin.I_0') = s(fun(X12,X12),'const.combin.o_2'(s(fun(X12,X12),X13),s(fun(X12,X12),X13))) | s(bool,X14) = s(bool,'const.relation.INVOL_1'(s(fun(X12,X12),X13))) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f55,f73])).
fof(f55,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.relation.INVOL_1'(s(fun(X0,X0),X1)))) | s(fun(X0,X0),'const.combin.I_0') = s(fun(X0,X0),'const.combin.o_2'(s(fun(X0,X0),X1),s(fun(X0,X0),X1)))) )),
  inference(cnf_transformation,[],[f40])).
fof(f1298,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (s(X30,happ(s(fun(bool,X30),X31),s(bool,X29))) = s(X30,'const.combin.o_3'(s(fun(bool,X30),X31),s(fun(X26,bool),X27),s(X26,X28))) | s(bool,'T_0') = s(bool,X29) | s(bool,'T_0') = s(bool,happ(s(fun(X26,bool),X27),s(X26,X28)))) )),
  inference(superposition,[],[f60,f73])).
fof(f1297,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (s(X23,happ(s(fun(bool,X23),X24),s(bool,'T_0'))) = s(X23,'const.combin.o_3'(s(fun(bool,X23),X24),s(fun(X20,bool),X21),s(X20,X22))) | s(bool,X25) = s(bool,happ(s(fun(X20,bool),X21),s(X20,X22))) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f60,f73])).
fof(f1309,plain,(
  ( ! [X26,X24,X23,X27,X25,X22] : (s(bool,X27) = s(bool,'const.combin.o_3'(s(fun(X22,bool),X23),s(fun(X24,X22),X25),s(X24,X26))) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,happ(s(fun(X22,bool),X23),s(X22,happ(s(fun(X24,X22),X25),s(X24,X26)))))) )),
  inference(superposition,[],[f60,f73])).
fof(f1308,plain,(
  ( ! [X21,X19,X17,X20,X18,X16] : (s(bool,'T_0') = s(bool,'const.combin.o_3'(s(fun(X16,bool),X17),s(fun(X18,X16),X19),s(X18,X20))) | s(bool,X21) = s(bool,happ(s(fun(X16,bool),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20))))) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f60,f73])).
fof(f1294,plain,(
  ( ! [X6,X4,X10,X8,X7,X5,X11,X9] : (s(X10,happ(s(fun(X4,X10),X11),s(X4,'const.combin.o_3'(s(fun(X5,X4),X6),s(fun(X7,X5),X8),s(X7,X9))))) = s(X10,'const.combin.o_3'(s(fun(X4,X10),X11),s(fun(X5,X4),X6),s(X5,happ(s(fun(X7,X5),X8),s(X7,X9)))))) )),
  inference(superposition,[],[f60,f60])).
fof(f1293,plain,(
  ( ! [X2,X0,X3,X1] : (s(X2,happ(s(fun(X0,X2),X3),s(X0,X1))) = s(X2,'const.combin.o_3'(s(fun(X0,X2),X3),s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))) )),
  inference(superposition,[],[f60,f68])).
fof(f465,plain,(
  ( ! [X14,X12,X13,X11] : (s(X13,f(s(bool,X14))) != s(X13,g(s(bool,X14))) | s(fun(bool,X13),X11) = s(fun(bool,X13),X12) | s(bool,'T_0') = s(bool,sK3(X11,X12,X13,bool)) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f57,f73])).
fof(f463,plain,(
  ( ! [X6,X4,X5,X3] : (s(X5,f(s(bool,'T_0'))) != s(X5,g(s(bool,'T_0'))) | s(fun(bool,X5),X3) = s(fun(bool,X5),X4) | s(bool,X6) = s(bool,sK3(X3,X4,X5,bool)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f57,f73])).
fof(f205,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f47,f73])).
fof(f164,plain,(
  ( ! [X17,X18] : (s(bool,'T_0') = s(bool,X17) | s(bool,X18) = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X17))) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f73,f68])).
fof(f163,plain,(
  ( ! [X15,X16] : (s(bool,'T_0') = s(bool,X15) | s(bool,X16) = s(bool,'const.combin.I_1'(s(bool,X15))) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f73,f51])).
fof(f153,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f73,f73])).
fof(f141,plain,(
  ( ! [X24,X25] : (s(bool,X24) = s(bool,X25) | s(bool,'T_0') = s(bool,X25) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X24)))) )),
  inference(superposition,[],[f73,f68])).
fof(f140,plain,(
  ( ! [X23,X22] : (s(bool,X22) = s(bool,X23) | s(bool,'T_0') = s(bool,X23) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X22)))) )),
  inference(superposition,[],[f73,f51])).
fof(f76,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f49,f50])).
fof(f58,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) )),
  inference(cnf_transformation,[],[f44])).
fof(f57,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK3(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK3(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f41])).
fof(f41,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK3(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK3(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3])],[f35])).
fof(f35,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f27])).
fof(f27,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f26])).
fof(f26,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.relation.INVOL',aHL_EXT)).
fof(f3314,plain,(
  ( ! [X18] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X18))) | s(bool,'F_0') = s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,X18)) ) | $spl1),
  inference(superposition,[],[f96,f2359])).
fof(f3702,plain,(
  ( ! [X43,X42] : (s(bool,X42) != s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,X43) | s(bool,X43) = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X42)))) ) | $spl26),
  inference(superposition,[],[f2131,f68])).
fof(f15023,plain,(
  ( ! [X94] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X94)) ) | ($spl1 | $spl9 | $spl13 | $spl26 | $spl29)),
  inference(subsumption_resolution,[],[f14813,f6306])).
fof(f14813,plain,(
  ( ! [X94] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X94) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X94)))) ) | ($spl1 | $spl9 | $spl13 | $spl26 | $spl29)),
  inference(superposition,[],[f6306,f14564])).
fof(f5539,plain,(
  ( ! [X18] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X18)))) ) | $spl38),
  inference(cnf_transformation,[],[f5539_D])).
fof(f5539_D,plain,(
  ( ! [X18] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X18)))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f15010,plain,(
  ( ! [X49] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X49)))) ) | ($spl1 | $spl9 | $spl13 | $spl16 | $spl26)),
  inference(subsumption_resolution,[],[f14791,f7671])).
fof(f14791,plain,(
  ( ! [X49] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X49))) | s(bool,'T_0') = s(bool,X49)) ) | ($spl1 | $spl9 | $spl13 | $spl16 | $spl26)),
  inference(superposition,[],[f7671,f14564])).
fof(f15006,plain,(
  ( ! [X47] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X47)) ) | ($spl1 | $spl9 | $spl13 | $spl26 | $spl29)),
  inference(subsumption_resolution,[],[f14789,f6306])).
fof(f14789,plain,(
  ( ! [X47] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X47))) | s(bool,'T_0') = s(bool,X47)) ) | ($spl1 | $spl9 | $spl13 | $spl26)),
  inference(superposition,[],[f6237,f14564])).
fof(f10301,plain,(
  $false | ($spl1 | $spl34)),
  inference(subsumption_resolution,[],[f10299,f68])).
fof(f10299,plain,(
  s(sK0,sK4('const.combin.I_0','const.combin.o_2'(s(fun(sK0,sK0),sK1),s(fun(sK0,sK0),sK1)),sK0,sK0)) != s(sK0,happ(s(fun(sK0,sK0),'const.combin.I_0'),s(sK0,sK4('const.combin.I_0','const.combin.o_2'(s(fun(sK0,sK0),sK1),s(fun(sK0,sK0),sK1)),sK0,sK0)))) | ($spl1 | $spl34)),
  inference(backward_demodulation,[],[f10297,f2536])).
fof(f2536,plain,(
  s(sK0,sK4('const.combin.I_0','const.combin.o_2'(s(fun(sK0,sK0),sK1),s(fun(sK0,sK0),sK1)),sK0,sK0)) != s(sK0,happ(s(fun(sK0,sK0),'const.combin.o_2'(s(fun(sK0,sK0),sK1),s(fun(sK0,sK0),sK1))),s(sK0,sK4('const.combin.I_0','const.combin.o_2'(s(fun(sK0,sK0),sK1),s(fun(sK0,sK0),sK1)),sK0,sK0)))) | $spl1),
  inference(forward_demodulation,[],[f2534,f68])).
fof(f2534,plain,(
  s(sK0,happ(s(fun(sK0,sK0),'const.combin.I_0'),s(sK0,sK4('const.combin.I_0','const.combin.o_2'(s(fun(sK0,sK0),sK1),s(fun(sK0,sK0),sK1)),sK0,sK0)))) != s(sK0,happ(s(fun(sK0,sK0),'const.combin.o_2'(s(fun(sK0,sK0),sK1),s(fun(sK0,sK0),sK1))),s(sK0,sK4('const.combin.I_0','const.combin.o_2'(s(fun(sK0,sK0),sK1),s(fun(sK0,sK0),sK1)),sK0,sK0)))) | $spl1),
  inference(unit_resulting_resolution,[],[f2353,f59])).
fof(f10297,plain,(
  ( ! [X0,X1] : (s(fun(X0,X0),'const.combin.I_0') = s(fun(X0,X0),'const.combin.o_2'(s(fun(X0,X0),X1),s(fun(X0,X0),X1)))) ) | $spl34),
  inference(subsumption_resolution,[],[f10230,f47])).
fof(f10230,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s(fun(X0,X0),'const.combin.I_0') = s(fun(X0,X0),'const.combin.o_2'(s(fun(X0,X0),X1),s(fun(X0,X0),X1)))) ) | $spl34),
  inference(backward_demodulation,[],[f9974,f55])).
fof(f9974,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl34),
  inference(condensation,[],[f9826])).
fof(f9826,plain,(
  ( ! [X43,X42] : (s(bool,X42) = s(bool,X43) | s(bool,'T_0') = s(bool,X43)) ) | $spl34),
  inference(superposition,[],[f5535,f68])).
fof(f10300,plain,(
  $false | ($spl1 | $spl34)),
  inference(subsumption_resolution,[],[f10298,f68])).
fof(f10298,plain,(
  s(sK0,sK4('const.combin.o_2'(s(fun(sK0,sK0),sK1),s(fun(sK0,sK0),sK1)),'const.combin.I_0',sK0,sK0)) != s(sK0,happ(s(fun(sK0,sK0),'const.combin.I_0'),s(sK0,sK4('const.combin.o_2'(s(fun(sK0,sK0),sK1),s(fun(sK0,sK0),sK1)),'const.combin.I_0',sK0,sK0)))) | ($spl1 | $spl34)),
  inference(backward_demodulation,[],[f10297,f2537])).
fof(f2537,plain,(
  s(sK0,sK4('const.combin.o_2'(s(fun(sK0,sK0),sK1),s(fun(sK0,sK0),sK1)),'const.combin.I_0',sK0,sK0)) != s(sK0,happ(s(fun(sK0,sK0),'const.combin.o_2'(s(fun(sK0,sK0),sK1),s(fun(sK0,sK0),sK1))),s(sK0,sK4('const.combin.o_2'(s(fun(sK0,sK0),sK1),s(fun(sK0,sK0),sK1)),'const.combin.I_0',sK0,sK0)))) | $spl1),
  inference(forward_demodulation,[],[f2532,f68])).
fof(f2532,plain,(
  s(sK0,happ(s(fun(sK0,sK0),'const.combin.I_0'),s(sK0,sK4('const.combin.o_2'(s(fun(sK0,sK0),sK1),s(fun(sK0,sK0),sK1)),'const.combin.I_0',sK0,sK0)))) != s(sK0,happ(s(fun(sK0,sK0),'const.combin.o_2'(s(fun(sK0,sK0),sK1),s(fun(sK0,sK0),sK1))),s(sK0,sK4('const.combin.o_2'(s(fun(sK0,sK0),sK1),s(fun(sK0,sK0),sK1)),'const.combin.I_0',sK0,sK0)))) | $spl1),
  inference(unit_resulting_resolution,[],[f2353,f59])).
fof(f10296,plain,(
  $false | ($spl29 | $spl34)),
  inference(subsumption_resolution,[],[f10035,f9974])).
fof(f10035,plain,(
  s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | ($spl29 | $spl34)),
  inference(backward_demodulation,[],[f9974,f2142])).
fof(f10295,plain,(
  $false | $spl34),
  inference(subsumption_resolution,[],[f10026,f47])).
fof(f10026,plain,(
  ~p(s(bool,'T_0')) | $spl34),
  inference(backward_demodulation,[],[f9974,f49])).
fof(f10025,plain,(
  ( ! [X8,X7,X9] : (s(bool,'F_0') = s(bool,'const.combin.o_3'(s(fun(bool,bool),'const.combin.I_0'),s(fun(X7,bool),X8),s(X7,X9)))) ) | $spl46),
  inference(cnf_transformation,[],[f10025_D])).
fof(f10025_D,plain,(
  ( ! [X8,X7,X9] : (s(bool,'F_0') = s(bool,'const.combin.o_3'(s(fun(bool,bool),'const.combin.I_0'),s(fun(X7,bool),X8),s(X7,X9)))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f9979,plain,(
  ( ! [X10,X8,X7,X9] : (s(bool,'T_0') = s(bool,X10) | s(bool,'F_0') = s(bool,'const.combin.o_3'(s(fun(bool,bool),'const.combin.I_0'),s(fun(X7,bool),X8),s(X7,X9)))) ) | $spl34),
  inference(duplicate_literal_removal,[],[f9813])).
fof(f9813,plain,(
  ( ! [X10,X8,X7,X9] : (s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X10) | s(bool,'F_0') = s(bool,'const.combin.o_3'(s(fun(bool,bool),'const.combin.I_0'),s(fun(X7,bool),X8),s(X7,X9)))) ) | $spl34),
  inference(superposition,[],[f5535,f1307])).
fof(f84,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl6),
  inference(cnf_transformation,[],[f84_D])).
fof(f84_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f9980,plain,(
  ( ! [X6,X5] : (s(bool,'T_0') = s(bool,X6) | s(bool,'F_0') = s(bool,X5)) ) | ($spl1 | $spl9 | $spl13 | $spl34)),
  inference(duplicate_literal_removal,[],[f9812])).
fof(f9812,plain,(
  ( ! [X6,X5] : (s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,X6) | s(bool,'F_0') = s(bool,X5)) ) | ($spl1 | $spl9 | $spl13 | $spl34)),
  inference(superposition,[],[f5535,f6166])).
fof(f6166,plain,(
  ( ! [X34] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X34))) | s(bool,'F_0') = s(bool,X34)) ) | ($spl1 | $spl9 | $spl13)),
  inference(forward_demodulation,[],[f2464,f6130])).
fof(f2464,plain,(
  ( ! [X34] : (s(bool,X34) = s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X34)))) ) | $spl1),
  inference(superposition,[],[f68,f2359])).
fof(f7510,plain,(
  ( ! [X156,X157,X155,X160] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X160),s(bool,sK3(X155,X156,X157,bool)))) | s(X157,f(s(bool,'F_0'))) != s(X157,g(s(bool,'F_0'))) | s(fun(bool,X157),X155) = s(fun(bool,X157),X156)) ) | $spl42),
  inference(cnf_transformation,[],[f7510_D])).
fof(f7510_D,plain,(
  ( ! [X156,X157,X155,X160] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X160),s(bool,sK3(X155,X156,X157,bool)))) | s(X157,f(s(bool,'F_0'))) != s(X157,g(s(bool,'F_0'))) | s(fun(bool,X157),X155) = s(fun(bool,X157),X156)) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f7511,plain,(
  ( ! [X158,X159] : (s(X158,happ(s(fun(bool,X158),X159),s(bool,'T_0'))) = s(X158,happ(s(fun(bool,X158),X159),s(bool,'F_0')))) ) | $spl44),
  inference(cnf_transformation,[],[f7511_D])).
fof(f7511_D,plain,(
  ( ! [X158,X159] : (s(X158,happ(s(fun(bool,X158),X159),s(bool,'T_0'))) = s(X158,happ(s(fun(bool,X158),X159),s(bool,'F_0')))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f7509,plain,(
  ( ! [X158,X156,X159,X157,X155,X160] : (s(X158,happ(s(fun(bool,X158),X159),s(bool,'T_0'))) = s(X158,happ(s(fun(bool,X158),X159),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X160),s(bool,sK3(X155,X156,X157,bool)))) | s(fun(bool,X157),X155) = s(fun(bool,X157),X156) | s(X157,f(s(bool,'F_0'))) != s(X157,g(s(bool,'F_0')))) ) | $spl18),
  inference(forward_demodulation,[],[f7399,f6373])).
fof(f6373,plain,(
  ( ! [X68,X66,X67,X65] : (s(X66,happ(s(fun(bool,X66),X67),s(bool,X68))) = s(X66,'const.combin.o_3'(s(fun(bool,X66),X67),s(fun(bool,bool),X65),s(bool,X68)))) ) | $spl18),
  inference(superposition,[],[f1293,f1260])).
fof(f7399,plain,(
  ( ! [X158,X156,X159,X157,X155,X160] : (s(X158,happ(s(fun(bool,X158),X159),s(bool,'F_0'))) = s(X158,'const.combin.o_3'(s(fun(bool,X158),X159),s(fun(bool,bool),X160),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X160),s(bool,sK3(X155,X156,X157,bool)))) | s(fun(bool,X157),X155) = s(fun(bool,X157),X156) | s(X157,f(s(bool,'F_0'))) != s(X157,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f1296,f467])).
fof(f6076,plain,(
  $false | ($spl16 | $spl29 | $spl32)),
  inference(subsumption_resolution,[],[f2142,f6059])).
fof(f6059,plain,(
  ( ! [X4] : (s(bool,'T_0') = s(bool,X4)) ) | ($spl16 | $spl32)),
  inference(subsumption_resolution,[],[f1113,f5525])).
fof(f6075,plain,(
  $false | ($spl16 | $spl32)),
  inference(subsumption_resolution,[],[f6074,f47])).
fof(f6074,plain,(
  ~p(s(bool,'T_0')) | ($spl16 | $spl32)),
  inference(forward_demodulation,[],[f49,f6059])).
fof(f6073,plain,(
  $false | ($spl16 | $spl22 | $spl29 | $spl32)),
  inference(subsumption_resolution,[],[f6015,f6059])).
fof(f6015,plain,(
  s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | ($spl22 | $spl29)),
  inference(global_subsumption,[],[f3023,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f2142,f1956,f467,f49])).
fof(f3023,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X0) = s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | $spl22),
  inference(equality_factoring,[],[f2907])).
fof(f2907,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl22),
  inference(condensation,[],[f2906])).
fof(f2906,plain,(
  ( ! [X21,X20] : (s(bool,'T_0') = s(bool,X21) | s(bool,X21) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X20)) ) | $spl22),
  inference(trivial_inequality_removal,[],[f2894])).
fof(f2894,plain,(
  ( ! [X21,X20] : (s(bool,X20) != s(bool,X20) | s(bool,'T_0') = s(bool,X21) | s(bool,X21) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X20)) ) | $spl22),
  inference(superposition,[],[f2624,f50])).
fof(f2624,plain,(
  ( ! [X26,X25] : (s(bool,'F_0') != s(bool,X25) | s(bool,'T_0') = s(bool,X26) | s(bool,X26) = s(bool,'const.combin.I_1'(s(bool,X25)))) ) | $spl22),
  inference(superposition,[],[f2120,f51])).
fof(f6070,plain,(
  $false | ($spl1 | $spl16 | $spl32)),
  inference(subsumption_resolution,[],[f6068,f68])).
fof(f6068,plain,(
  s(sK0,sK4('const.combin.I_0','const.combin.o_2'(s(fun(sK0,sK0),sK1),s(fun(sK0,sK0),sK1)),sK0,sK0)) != s(sK0,happ(s(fun(sK0,sK0),'const.combin.I_0'),s(sK0,sK4('const.combin.I_0','const.combin.o_2'(s(fun(sK0,sK0),sK1),s(fun(sK0,sK0),sK1)),sK0,sK0)))) | ($spl1 | $spl16 | $spl32)),
  inference(backward_demodulation,[],[f6066,f2536])).
fof(f6066,plain,(
  ( ! [X0,X1] : (s(fun(X0,X0),'const.combin.I_0') = s(fun(X0,X0),'const.combin.o_2'(s(fun(X0,X0),X1),s(fun(X0,X0),X1)))) ) | ($spl16 | $spl32)),
  inference(subsumption_resolution,[],[f6065,f47])).
fof(f6065,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s(fun(X0,X0),'const.combin.I_0') = s(fun(X0,X0),'const.combin.o_2'(s(fun(X0,X0),X1),s(fun(X0,X0),X1)))) ) | ($spl16 | $spl32)),
  inference(forward_demodulation,[],[f55,f6059])).
fof(f6069,plain,(
  $false | ($spl1 | $spl16 | $spl32)),
  inference(subsumption_resolution,[],[f6067,f68])).
fof(f6067,plain,(
  s(sK0,sK4('const.combin.o_2'(s(fun(sK0,sK0),sK1),s(fun(sK0,sK0),sK1)),'const.combin.I_0',sK0,sK0)) != s(sK0,happ(s(fun(sK0,sK0),'const.combin.I_0'),s(sK0,sK4('const.combin.o_2'(s(fun(sK0,sK0),sK1),s(fun(sK0,sK0),sK1)),'const.combin.I_0',sK0,sK0)))) | ($spl1 | $spl16 | $spl32)),
  inference(backward_demodulation,[],[f6066,f2537])).
fof(f6063,plain,(
  $false | ($spl1 | $spl16 | $spl32)),
  inference(subsumption_resolution,[],[f6062,f47])).
fof(f6062,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl16 | $spl32)),
  inference(forward_demodulation,[],[f5659,f6059])).
fof(f5659,plain,(
  ~p(s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1)))) | $spl1),
  inference(global_subsumption,[],[f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f347,f356,f390,f463,f465,f486,f966,f1158,f1159,f1293,f1294,f1305,f784,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f49,f1330,f781,f1567,f1756,f1757,f1758,f1759,f1766,f1810,f1855,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1950,f1983,f1984,f1985,f1986,f1956,f1957,f1958,f1989,f1965,f1966,f1967,f1968,f1969,f1974,f2073,f2074,f2075,f2076,f2119,f2124,f2126,f2130,f2104,f2088,f1324,f473,f467,f64,f2353,f2359,f2434,f2535,f2536,f2533,f2537,f2594,f2659,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4556,f4051,f4557,f4559,f4172,f4176,f4177,f4182,f4675,f4824,f4825,f4900,f4914,f4929,f4930,f4911,f4855,f3857,f3812,f2724,f2723,f5658])).
fof(f5658,plain,(
  ( ! [X19,X18] : (s(bool,'T_0') = s(bool,'const.relation.INVOL_1'(s(fun(X18,X18),X19))) | ~p(s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1)))) | s(fun(X18,X18),'const.combin.I_0') = s(fun(X18,X18),'const.combin.o_2'(s(fun(X18,X18),X19),s(fun(X18,X18),X19)))) ) | $spl1),
  inference(forward_demodulation,[],[f2706,f51])).
fof(f2706,plain,(
  ( ! [X19,X18] : (~p(s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1)))) | s(fun(X18,X18),'const.combin.I_0') = s(fun(X18,X18),'const.combin.o_2'(s(fun(X18,X18),X19),s(fun(X18,X18),X19))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'const.relation.INVOL_1'(s(fun(X18,X18),X19)))))) ) | $spl1),
  inference(superposition,[],[f55,f2434])).
fof(f2723,plain,(
  ( ! [X64,X65,X63] : (s(X65,f(s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))))) != s(X65,g(s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))))) | s(fun(bool,X65),X63) = s(fun(bool,X65),X64) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,sK3(X63,X64,X65,bool))))) ) | $spl1),
  inference(superposition,[],[f57,f2434])).
fof(f2724,plain,(
  ( ! [X68,X66,X67] : (s(X68,happ(s(fun(bool,X68),X66),s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))))) != s(X68,happ(s(fun(bool,X68),X67),s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))))) | s(fun(bool,X68),X66) = s(fun(bool,X68),X67) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,sK4(X66,X67,bool,X68))))) ) | $spl1),
  inference(superposition,[],[f59,f2434])).
fof(f3812,plain,(
  ( ! [X78,X79,X77] : (s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) = s(bool,happ(s(fun(X77,bool),X78),s(X77,X79))) | s(bool,'T_0') = s(bool,'const.combin.o_3'(s(fun(X77,bool),X78),s(fun(X77,X77),'const.combin.I_0'),s(X77,X79)))) ) | $spl1),
  inference(superposition,[],[f2359,f1293])).
fof(f3857,plain,(
  ( ! [X85,X83,X84] : (s(bool,'T_0') = s(bool,happ(s(fun(X83,bool),X84),s(X83,X85))) | s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) = s(bool,'const.combin.o_3'(s(fun(X83,bool),X84),s(fun(X83,X83),'const.combin.I_0'),s(X83,X85)))) ) | $spl1),
  inference(forward_demodulation,[],[f3814,f51])).
fof(f3814,plain,(
  ( ! [X85,X83,X84] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,happ(s(fun(X83,bool),X84),s(X83,X85))))) | s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) = s(bool,'const.combin.o_3'(s(fun(X83,bool),X84),s(fun(X83,X83),'const.combin.I_0'),s(X83,X85)))) ) | $spl1),
  inference(superposition,[],[f2434,f1293])).
fof(f4855,plain,(
  ( ! [X88,X87,X86] : (s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) = s(bool,happ(s(fun(X86,bool),X87),s(X86,X88))) | s(bool,'T_0') = s(bool,'const.combin.o_3'(s(fun(bool,bool),'const.combin.I_0'),s(fun(X86,bool),X87),s(X86,X88)))) ) | $spl1),
  inference(superposition,[],[f2359,f1305])).
fof(f4911,plain,(
  ( ! [X94,X92,X93] : (s(bool,'T_0') = s(bool,happ(s(fun(X92,bool),X93),s(X92,X94))) | s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) = s(bool,'const.combin.o_3'(s(fun(bool,bool),'const.combin.I_0'),s(fun(X92,bool),X93),s(X92,X94)))) ) | $spl1),
  inference(forward_demodulation,[],[f4857,f51])).
fof(f4857,plain,(
  ( ! [X94,X92,X93] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,happ(s(fun(X92,bool),X93),s(X92,X94))))) | s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) = s(bool,'const.combin.o_3'(s(fun(bool,bool),'const.combin.I_0'),s(fun(X92,bool),X93),s(X92,X94)))) ) | $spl1),
  inference(superposition,[],[f2434,f1305])).
fof(f4914,plain,(
  ( ! [X101,X99,X102,X100] : (s(bool,'T_0') = s(bool,happ(s(fun(X99,bool),X100),s(X99,X101))) | s(bool,X102) = s(bool,'const.combin.o_3'(s(fun(bool,bool),'const.combin.I_0'),s(fun(X99,bool),X100),s(X99,X101))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X102)))) ) | $spl1),
  inference(forward_demodulation,[],[f4859,f51])).
fof(f4859,plain,(
  ( ! [X101,X99,X102,X100] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,happ(s(fun(X99,bool),X100),s(X99,X101))))) | s(bool,X102) = s(bool,'const.combin.o_3'(s(fun(bool,bool),'const.combin.I_0'),s(fun(X99,bool),X100),s(X99,X101))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X102)))) ) | $spl1),
  inference(superposition,[],[f2659,f1305])).
fof(f4675,plain,(
  ( ! [X198,X200,X199,X201] : (s(bool,'T_0') = s(bool,X201) | s(X200,happ(s(fun(bool,X200),X198),s(bool,X201))) != s(X200,happ(s(fun(bool,X200),X199),s(bool,X201))) | s(fun(bool,X200),X198) = s(fun(bool,X200),X199) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,sK4(X198,X199,bool,X200))))) ) | $spl1),
  inference(forward_demodulation,[],[f4183,f51])).
fof(f4183,plain,(
  ( ! [X198,X200,X199,X201] : (s(X200,happ(s(fun(bool,X200),X198),s(bool,X201))) != s(X200,happ(s(fun(bool,X200),X199),s(bool,X201))) | s(fun(bool,X200),X198) = s(fun(bool,X200),X199) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X201))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,sK4(X198,X199,bool,X200))))) ) | $spl1),
  inference(superposition,[],[f59,f2659])).
fof(f4182,plain,(
  ( ! [X196,X194,X197,X195] : (s(X196,f(s(bool,X197))) != s(X196,g(s(bool,X197))) | s(fun(bool,X196),X194) = s(fun(bool,X196),X195) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X197))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,sK3(X194,X195,X196,bool))))) ) | $spl1),
  inference(superposition,[],[f57,f2659])).
fof(f4177,plain,(
  ( ! [X177,X175,X178,X176] : (s(bool,X178) != s(bool,happ(s(fun(X175,bool),X177),s(X175,sK4(X176,X177,X175,bool)))) | s(fun(X175,bool),X176) = s(fun(X175,bool),X177) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X178))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,happ(s(fun(X175,bool),X176),s(X175,sK4(X176,X177,X175,bool))))))) ) | $spl1),
  inference(superposition,[],[f59,f2659])).
fof(f4176,plain,(
  ( ! [X173,X171,X174,X172] : (s(bool,X174) != s(bool,happ(s(fun(X171,bool),X173),s(X171,sK4(X173,X172,X171,bool)))) | s(fun(X171,bool),X172) = s(fun(X171,bool),X173) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X174))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,happ(s(fun(X171,bool),X172),s(X171,sK4(X173,X172,X171,bool))))))) ) | $spl1),
  inference(superposition,[],[f59,f2659])).
fof(f4172,plain,(
  ( ! [X156,X154,X155,X153] : (s(bool,X156) != s(bool,f(s(X153,sK3(X154,X155,bool,X153)))) | s(fun(X153,bool),X154) = s(fun(X153,bool),X155) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X156))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,g(s(X153,sK3(X154,X155,bool,X153))))))) ) | $spl1),
  inference(superposition,[],[f57,f2659])).
fof(f4559,plain,(
  ( ! [X80,X78,X81,X79] : (s(bool,'T_0') = s(bool,happ(s(fun(X78,bool),X79),s(X78,X80))) | s(bool,X81) = s(bool,'const.combin.o_3'(s(fun(X78,bool),X79),s(fun(X78,X78),'const.combin.I_0'),s(X78,X80))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X81)))) ) | $spl1),
  inference(forward_demodulation,[],[f4054,f51])).
fof(f4054,plain,(
  ( ! [X80,X78,X81,X79] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,happ(s(fun(X78,bool),X79),s(X78,X80))))) | s(bool,X81) = s(bool,'const.combin.o_3'(s(fun(X78,bool),X79),s(fun(X78,X78),'const.combin.I_0'),s(X78,X80))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X81)))) ) | $spl1),
  inference(superposition,[],[f2659,f1293])).
fof(f4557,plain,(
  ( ! [X70,X74,X72,X71,X69,X73] : (s(bool,'T_0') = s(bool,'const.combin.o_3'(s(fun(X70,bool),X71),s(fun(X69,X70),X72),s(X69,X73))) | s(bool,X74) = s(bool,happ(s(fun(X69,bool),'const.combin.o_2'(s(fun(X70,bool),X71),s(fun(X69,X70),X72))),s(X69,X73))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X74)))) ) | $spl1),
  inference(forward_demodulation,[],[f4052,f51])).
fof(f4052,plain,(
  ( ! [X70,X74,X72,X71,X69,X73] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'const.combin.o_3'(s(fun(X70,bool),X71),s(fun(X69,X70),X72),s(X69,X73))))) | s(bool,X74) = s(bool,happ(s(fun(X69,bool),'const.combin.o_2'(s(fun(X70,bool),X71),s(fun(X69,X70),X72))),s(X69,X73))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X74)))) ) | $spl1),
  inference(superposition,[],[f2659,f61])).
fof(f4051,plain,(
  ( ! [X68,X66,X64,X67,X65,X63] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'const.combin.o_3'(s(fun(X63,bool),X64),s(fun(X65,X63),X66),s(X65,X67))))) | s(bool,X68) = s(bool,happ(s(fun(X63,bool),X64),s(X63,happ(s(fun(X65,X63),X66),s(X65,X67))))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X68)))) ) | $spl1),
  inference(superposition,[],[f2659,f60])).
fof(f4556,plain,(
  ( ! [X61,X62] : (s(bool,'T_0') = s(bool,X61) | s(bool,X62) = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X61))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X62)))) ) | $spl1),
  inference(forward_demodulation,[],[f4050,f51])).
fof(f4050,plain,(
  ( ! [X61,X62] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X61))) | s(bool,X62) = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X61))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X62)))) ) | $spl1),
  inference(superposition,[],[f2659,f68])).
fof(f2533,plain,(
  s(sK0,f(s(sK0,sK3('const.combin.o_2'(s(fun(sK0,sK0),sK1),s(fun(sK0,sK0),sK1)),'const.combin.I_0',sK0,sK0)))) != s(sK0,g(s(sK0,sK3('const.combin.o_2'(s(fun(sK0,sK0),sK1),s(fun(sK0,sK0),sK1)),'const.combin.I_0',sK0,sK0)))) | $spl1),
  inference(unit_resulting_resolution,[],[f2353,f57])).
fof(f2535,plain,(
  s(sK0,f(s(sK0,sK3('const.combin.I_0','const.combin.o_2'(s(fun(sK0,sK0),sK1),s(fun(sK0,sK0),sK1)),sK0,sK0)))) != s(sK0,g(s(sK0,sK3('const.combin.I_0','const.combin.o_2'(s(fun(sK0,sK0),sK1),s(fun(sK0,sK0),sK1)),sK0,sK0)))) | $spl1),
  inference(unit_resulting_resolution,[],[f2353,f57])).
fof(f2076,plain,(
  ( ! [X80,X78,X76,X81,X79,X77,X82] : (s(bool,X81) != s(bool,'const.combin.o_3'(s(fun(X77,bool),X78),s(fun(X76,X77),X79),s(X76,X80))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X81))) | s(bool,X82) = s(bool,happ(s(fun(X76,bool),'const.combin.o_2'(s(fun(X77,bool),X78),s(fun(X76,X77),X79))),s(X76,X80))) | s(bool,'T_0') = s(bool,X82)) ) | $spl1),
  inference(superposition,[],[f1567,f61])).
fof(f2075,plain,(
  ( ! [X70,X74,X72,X71,X69,X75,X73] : (s(bool,X74) != s(bool,'const.combin.o_3'(s(fun(X69,bool),X70),s(fun(X71,X69),X72),s(X71,X73))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X74))) | s(bool,X75) = s(bool,happ(s(fun(X69,bool),X70),s(X69,happ(s(fun(X71,X69),X72),s(X71,X73))))) | s(bool,'T_0') = s(bool,X75)) ) | $spl1),
  inference(superposition,[],[f1567,f60])).
fof(f2074,plain,(
  ( ! [X68,X66,X67] : (s(bool,X66) != s(bool,X67) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X67))) | s(bool,X68) = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X66))) | s(bool,'T_0') = s(bool,X68)) ) | $spl1),
  inference(superposition,[],[f1567,f68])).
fof(f2073,plain,(
  ( ! [X64,X65,X63] : (s(bool,X63) != s(bool,X64) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X64))) | s(bool,X65) = s(bool,'const.combin.I_1'(s(bool,X63))) | s(bool,'T_0') = s(bool,X65)) ) | $spl1),
  inference(superposition,[],[f1567,f51])).
fof(f1974,plain,(
  ( ! [X35,X33,X36,X34] : (s(bool,X36) != s(bool,happ(s(fun(X33,bool),X35),s(X33,sK4(X34,X35,X33,bool)))) | s(fun(X33,bool),X34) = s(fun(X33,bool),X35) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X36))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,happ(s(fun(X33,bool),X34),s(X33,sK4(X34,X35,X33,bool))))))) ) | $spl1),
  inference(superposition,[],[f59,f486])).
fof(f1989,plain,(
  ( ! [X28,X26,X27,X25] : (s(bool,'T_0') = s(bool,X28) | s(X27,happ(s(fun(bool,X27),X25),s(bool,X28))) != s(X27,happ(s(fun(bool,X27),X26),s(bool,X28))) | s(fun(bool,X27),X25) = s(fun(bool,X27),X26) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,sK4(X25,X26,bool,X27))))) ) | $spl1),
  inference(forward_demodulation,[],[f1963,f51])).
fof(f1963,plain,(
  ( ! [X28,X26,X27,X25] : (s(X27,happ(s(fun(bool,X27),X25),s(bool,X28))) != s(X27,happ(s(fun(bool,X27),X26),s(bool,X28))) | s(fun(bool,X27),X25) = s(fun(bool,X27),X26) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X28))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,sK4(X25,X26,bool,X27))))) ) | $spl1),
  inference(superposition,[],[f59,f486])).
fof(f1950,plain,(
  ( ! [X35,X33,X36,X34] : (s(bool,X36) != s(bool,happ(s(fun(X33,bool),X35),s(X33,sK4(X35,X34,X33,bool)))) | s(fun(X33,bool),X34) = s(fun(X33,bool),X35) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X36))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,happ(s(fun(X33,bool),X34),s(X33,sK4(X35,X34,X33,bool))))))) ) | $spl1),
  inference(superposition,[],[f59,f486])).
fof(f1855,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) != s(bool,'const.combin.I_1'(s(bool,X0))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X1))) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl1),
  inference(superposition,[],[f1159,f73])).
fof(f1810,plain,(
  ( ! [X92,X90,X88,X87,X91,X89] : (s(bool,'T_0') = s(bool,'const.combin.o_3'(s(fun(X88,bool),X89),s(fun(X87,X88),X90),s(X87,X91))) | s(bool,X92) = s(bool,happ(s(fun(X87,bool),'const.combin.o_2'(s(fun(X88,bool),X89),s(fun(X87,X88),X90))),s(X87,X91))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X92)))) ) | $spl1),
  inference(forward_demodulation,[],[f1780,f51])).
fof(f1780,plain,(
  ( ! [X92,X90,X88,X87,X91,X89] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'const.combin.o_3'(s(fun(X88,bool),X89),s(fun(X87,X88),X90),s(X87,X91))))) | s(bool,X92) = s(bool,happ(s(fun(X87,bool),'const.combin.o_2'(s(fun(X88,bool),X89),s(fun(X87,X88),X90))),s(X87,X91))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X92)))) ) | $spl1),
  inference(superposition,[],[f486,f61])).
fof(f1567,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X1))) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl1),
  inference(superposition,[],[f1158,f73])).
fof(f781,plain,(
  ( ! [X64,X62,X65,X63] : (s(bool,X65) != s(bool,f(s(X62,sK3(X63,X64,bool,X62)))) | s(fun(X62,bool),X63) = s(fun(X62,bool),X64) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X65))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,g(s(X62,sK3(X63,X64,bool,X62))))))) ) | $spl1),
  inference(superposition,[],[f57,f486])).
fof(f1330,plain,(
  ( ! [X80,X78,X76,X81,X79,X77] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'const.combin.o_3'(s(fun(X76,bool),X77),s(fun(X78,X76),X79),s(X78,X80))))) | s(bool,X81) = s(bool,happ(s(fun(X76,bool),X77),s(X76,happ(s(fun(X78,X76),X79),s(X78,X80))))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X81)))) ) | $spl1),
  inference(superposition,[],[f486,f60])).
fof(f784,plain,(
  ( ! [X70,X72,X71,X73] : (s(X72,f(s(bool,X73))) != s(X72,g(s(bool,X73))) | s(fun(bool,X72),X70) = s(fun(bool,X72),X71) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X73))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,sK3(X70,X71,X72,bool))))) ) | $spl1),
  inference(superposition,[],[f57,f486])).
fof(f1159,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X0)))) ) | $spl1),
  inference(subsumption_resolution,[],[f842,f1158])).
fof(f842,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X0)))) ) | $spl1),
  inference(equality_factoring,[],[f486])).
fof(f1158,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X0)))) ) | $spl1),
  inference(duplicate_literal_removal,[],[f1157])).
fof(f1157,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X0))) | s(bool,'T_0') != s(bool,X0) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X0)))) ) | $spl1),
  inference(forward_demodulation,[],[f846,f51])).
fof(f846,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'const.combin.I_1'(s(bool,X0))))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X0)))) ) | $spl1),
  inference(equality_factoring,[],[f486])).
fof(f966,plain,(
  ( ! [X50,X49] : (s(bool,'T_0') = s(bool,X49) | s(bool,X50) = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X49))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X50)))) ) | $spl1),
  inference(forward_demodulation,[],[f716,f51])).
fof(f716,plain,(
  ( ! [X50,X49] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X49))) | s(bool,X50) = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X49))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X50)))) ) | $spl1),
  inference(superposition,[],[f486,f68])).
fof(f486,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X1))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X0)))) ) | $spl1),
  inference(superposition,[],[f356,f356])).
fof(f356,plain,(
  ( ! [X15] : (s(bool,X15) = s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X15)))) ) | $spl1),
  inference(superposition,[],[f347,f51])).
fof(f347,plain,(
  ( ! [X11] : (s(bool,X11) = s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,X11)) ) | $spl1),
  inference(subsumption_resolution,[],[f217,f47])).
fof(f217,plain,(
  ( ! [X11] : (~p(s(bool,'T_0')) | s(bool,X11) = s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,X11)) ) | $spl1),
  inference(superposition,[],[f64,f73])).
fof(f6061,plain,(
  $false | ($spl1 | $spl16 | $spl32)),
  inference(subsumption_resolution,[],[f6060,f47])).
fof(f6060,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl16 | $spl32)),
  inference(forward_demodulation,[],[f64,f6059])).
fof(f6058,plain,(
  $false | ($spl9 | $spl40)),
  inference(subsumption_resolution,[],[f86,f5547])).
fof(f6057,plain,(
  $false | ($spl13 | $spl40)),
  inference(subsumption_resolution,[],[f418,f5547])).
fof(f6055,plain,(
  $false | ($spl1 | $spl40)),
  inference(subsumption_resolution,[],[f6054,f47])).
fof(f6054,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl40)),
  inference(forward_demodulation,[],[f64,f5547])).
fof(f6053,plain,(
  $false | ($spl1 | $spl40)),
  inference(subsumption_resolution,[],[f6052,f47])).
fof(f6052,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl40)),
  inference(forward_demodulation,[],[f5659,f5547])).
fof(f6049,plain,(
  $false | ($spl1 | $spl40)),
  inference(subsumption_resolution,[],[f6047,f68])).
fof(f6047,plain,(
  s(sK0,sK4('const.combin.I_0','const.combin.o_2'(s(fun(sK0,sK0),sK1),s(fun(sK0,sK0),sK1)),sK0,sK0)) != s(sK0,happ(s(fun(sK0,sK0),'const.combin.I_0'),s(sK0,sK4('const.combin.I_0','const.combin.o_2'(s(fun(sK0,sK0),sK1),s(fun(sK0,sK0),sK1)),sK0,sK0)))) | ($spl1 | $spl40)),
  inference(backward_demodulation,[],[f6045,f2536])).
fof(f6045,plain,(
  ( ! [X0,X1] : (s(fun(X0,X0),'const.combin.I_0') = s(fun(X0,X0),'const.combin.o_2'(s(fun(X0,X0),X1),s(fun(X0,X0),X1)))) ) | $spl40),
  inference(subsumption_resolution,[],[f6044,f47])).
fof(f6044,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s(fun(X0,X0),'const.combin.I_0') = s(fun(X0,X0),'const.combin.o_2'(s(fun(X0,X0),X1),s(fun(X0,X0),X1)))) ) | $spl40),
  inference(forward_demodulation,[],[f55,f5547])).
fof(f6048,plain,(
  $false | ($spl1 | $spl40)),
  inference(subsumption_resolution,[],[f6046,f68])).
fof(f6046,plain,(
  s(sK0,sK4('const.combin.o_2'(s(fun(sK0,sK0),sK1),s(fun(sK0,sK0),sK1)),'const.combin.I_0',sK0,sK0)) != s(sK0,happ(s(fun(sK0,sK0),'const.combin.I_0'),s(sK0,sK4('const.combin.o_2'(s(fun(sK0,sK0),sK1),s(fun(sK0,sK0),sK1)),'const.combin.I_0',sK0,sK0)))) | ($spl1 | $spl40)),
  inference(backward_demodulation,[],[f6045,f2537])).
fof(f6041,plain,(
  $false | ($spl22 | $spl29 | $spl40)),
  inference(subsumption_resolution,[],[f6015,f5547])).
fof(f6040,plain,(
  $false | $spl40),
  inference(subsumption_resolution,[],[f6039,f47])).
fof(f6039,plain,(
  ~p(s(bool,'T_0')) | $spl40),
  inference(forward_demodulation,[],[f49,f5547])).
fof(f6038,plain,(
  $false | ($spl29 | $spl40)),
  inference(subsumption_resolution,[],[f2142,f5547])).
fof(f6037,plain,(
  $false | ($spl6 | $spl12 | $spl18 | $spl22 | $spl26)),
  inference(global_subsumption,[],[f1260,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f417,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f2131,f2120,f2594,f2624,f2625,f2626,f2627,f2881,f3322,f3323,f3332,f3350,f3352,f3353,f3701,f3784,f3785,f3786,f3846,f3805,f3809,f3870,f3871,f4824,f4825,f4900,f4852,f4929,f4930,f4853,f3810,f3704,f3703,f3702,f5533,f5534,f5553,f5661,f2634,f2637,f84,f2943,f3730,f3727,f5845,f5859,f5861,f2907,f5865,f5867,f5869,f1956,f5872,f467,f49,f5931,f5934,f5936,f5938,f5974,f5976,f5987,f5989,f5997,f6001,f6003,f3023,f2998,f6020,f6021,f6027])).
fof(f6027,plain,(
  ( ! [X0,X1] : (s(fun(X0,X0),'const.combin.I_0') = s(fun(X0,X0),'const.combin.o_2'(s(fun(X0,X0),X1),s(fun(X0,X0),X1)))) ) | ($spl6 | $spl12 | $spl22 | $spl26)),
  inference(subsumption_resolution,[],[f6026,f47])).
fof(f6026,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s(fun(X0,X0),'const.combin.I_0') = s(fun(X0,X0),'const.combin.o_2'(s(fun(X0,X0),X1),s(fun(X0,X0),X1)))) ) | ($spl6 | $spl12 | $spl22 | $spl26)),
  inference(forward_demodulation,[],[f5861,f6021])).
fof(f6021,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl6 | $spl12 | $spl22 | $spl26)),
  inference(global_subsumption,[],[f6020,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f417,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f2131,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3701,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f4853,f3810,f3704,f3703,f3702,f3730,f3727,f1956,f467,f49])).
fof(f6020,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,'const.combin.I_1'(s(bool,X1)))) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f6019,f84])).
fof(f6019,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') != s(bool,'F_0') | s(bool,X0) = s(bool,'const.combin.I_1'(s(bool,X1)))) ) | $spl22),
  inference(condensation,[],[f6018])).
fof(f6018,plain,(
  ( ! [X4,X2,X3] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X3) | s(bool,X3) = s(bool,'const.combin.I_1'(s(bool,X2))) | s(bool,'T_0') = s(bool,X4)) ) | $spl22),
  inference(subsumption_resolution,[],[f2888,f2881])).
fof(f2888,plain,(
  ( ! [X4,X2,X3] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X3) | s(bool,X3) = s(bool,'const.combin.I_1'(s(bool,X2))) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) ) | $spl22),
  inference(superposition,[],[f2624,f73])).
fof(f2998,plain,(
  ( ! [X70] : (s(bool,X70) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X70)))) ) | $spl22),
  inference(superposition,[],[f51,f2907])).
fof(f6003,plain,(
  ( ! [X21,X19,X20] : (s(fun(X19,bool),X20) = s(fun(X19,bool),X21) | s(bool,'T_0') = s(bool,g(s(X19,sK3(X20,X21,bool,X19))))) ) | $spl6),
  inference(subsumption_resolution,[],[f473,f84])).
fof(f6001,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,g(s(X0,sK3(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | $spl6),
  inference(condensation,[],[f6000])).
fof(f6000,plain,(
  ( ! [X14,X12,X13,X11] : (s(fun(X11,bool),X12) = s(fun(X11,bool),X13) | s(bool,'T_0') = s(bool,g(s(X11,sK3(X12,X13,bool,X11)))) | s(bool,'T_0') = s(bool,X14)) ) | $spl6),
  inference(subsumption_resolution,[],[f5999,f50])).
fof(f5999,plain,(
  ( ! [X14,X12,X13,X11] : (s(bool,'F_0') != s(bool,X14) | s(fun(X11,bool),X12) = s(fun(X11,bool),X13) | s(bool,'T_0') = s(bool,g(s(X11,sK3(X12,X13,bool,X11)))) | s(bool,'T_0') = s(bool,X14)) ) | $spl6),
  inference(forward_demodulation,[],[f471,f84])).
fof(f5997,plain,(
  ( ! [X6,X4,X5,X3] : (s(fun(X3,bool),X4) = s(fun(X3,bool),X5) | s(bool,X6) = s(bool,g(s(X3,sK3(X4,X5,bool,X3)))) | s(bool,'T_0') = s(bool,X6)) ) | $spl6),
  inference(subsumption_resolution,[],[f5996,f84])).
fof(f5996,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,'T_0') != s(bool,'F_0') | s(fun(X3,bool),X4) = s(fun(X3,bool),X5) | s(bool,X6) = s(bool,g(s(X3,sK3(X4,X5,bool,X3)))) | s(bool,'T_0') = s(bool,X6)) ) | $spl6),
  inference(forward_demodulation,[],[f469,f84])).
fof(f5989,plain,(
  ( ! [X68,X66,X69,X67] : (s(bool,'T_0') = s(bool,X69) | s(bool,X69) = s(bool,'const.combin.o_3'(s(fun(X66,bool),X67),s(fun(X66,X66),'const.combin.I_0'),s(X66,X68)))) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f3809,f84])).
fof(f5987,plain,(
  ( ! [X70,X72,X71,X73] : (s(bool,'T_0') = s(bool,X73) | s(bool,X73) = s(bool,'const.combin.o_3'(s(fun(X70,bool),X71),s(fun(X70,X70),'const.combin.I_0'),s(X70,X72)))) ) | ($spl6 | $spl26)),
  inference(subsumption_resolution,[],[f5986,f84])).
fof(f5986,plain,(
  ( ! [X70,X72,X71,X73] : (s(bool,'F_0') != s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,X73) | s(bool,X73) = s(bool,'const.combin.o_3'(s(fun(X70,bool),X71),s(fun(X70,X70),'const.combin.I_0'),s(X70,X72)))) ) | ($spl6 | $spl26)),
  inference(forward_demodulation,[],[f3810,f84])).
fof(f5976,plain,(
  ( ! [X78,X76,X77,X75] : (s(bool,'T_0') = s(bool,X78) | s(bool,X78) = s(bool,'const.combin.o_3'(s(fun(bool,bool),'const.combin.I_0'),s(fun(X75,bool),X76),s(X75,X77)))) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f4852,f84])).
fof(f5974,plain,(
  ( ! [X80,X81,X79,X82] : (s(bool,'T_0') = s(bool,X82) | s(bool,X82) = s(bool,'const.combin.o_3'(s(fun(bool,bool),'const.combin.I_0'),s(fun(X79,bool),X80),s(X79,X81)))) ) | ($spl6 | $spl26)),
  inference(subsumption_resolution,[],[f5973,f84])).
fof(f5973,plain,(
  ( ! [X80,X81,X79,X82] : (s(bool,'F_0') != s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,X82) | s(bool,X82) = s(bool,'const.combin.o_3'(s(fun(bool,bool),'const.combin.I_0'),s(fun(X79,bool),X80),s(X79,X81)))) ) | ($spl6 | $spl26)),
  inference(forward_demodulation,[],[f4853,f84])).
fof(f5938,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (s(bool,'T_0') = s(bool,X34) | s(bool,X34) = s(bool,happ(s(fun(X29,bool),X30),s(X29,happ(s(fun(X31,X29),X32),s(X31,X33)))))) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f2626,f84])).
fof(f5936,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (s(bool,'T_0') = s(bool,X40) | s(bool,X40) = s(bool,happ(s(fun(X35,bool),'const.combin.o_2'(s(fun(X36,bool),X37),s(fun(X35,X36),X38))),s(X35,X39)))) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f2627,f84])).
fof(f5934,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (s(bool,'T_0') = s(bool,X49) | s(bool,X49) = s(bool,happ(s(fun(X44,bool),X45),s(X44,happ(s(fun(X46,X44),X47),s(X46,X48)))))) ) | ($spl6 | $spl26)),
  inference(subsumption_resolution,[],[f5933,f84])).
fof(f5933,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (s(bool,'F_0') != s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,X49) | s(bool,X49) = s(bool,happ(s(fun(X44,bool),X45),s(X44,happ(s(fun(X46,X44),X47),s(X46,X48)))))) ) | ($spl6 | $spl26)),
  inference(forward_demodulation,[],[f3703,f84])).
fof(f5931,plain,(
  ( ! [X54,X52,X50,X55,X53,X51] : (s(bool,'T_0') = s(bool,X55) | s(bool,X55) = s(bool,happ(s(fun(X50,bool),'const.combin.o_2'(s(fun(X51,bool),X52),s(fun(X50,X51),X53))),s(X50,X54)))) ) | ($spl6 | $spl26)),
  inference(subsumption_resolution,[],[f5930,f84])).
fof(f5930,plain,(
  ( ! [X54,X52,X50,X55,X53,X51] : (s(bool,'F_0') != s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,X55) | s(bool,X55) = s(bool,happ(s(fun(X50,bool),'const.combin.o_2'(s(fun(X51,bool),X52),s(fun(X50,X51),X53))),s(X50,X54)))) ) | ($spl6 | $spl26)),
  inference(forward_demodulation,[],[f3704,f84])).
fof(f5872,plain,(
  ( ! [X10,X8,X9] : (s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,'T_0') = s(bool,happ(s(fun(X8,bool),X9),s(X8,sK4(X10,X9,X8,bool))))) ) | $spl6),
  inference(subsumption_resolution,[],[f1943,f84])).
fof(f5869,plain,(
  ( ! [X10,X8,X9] : (s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,'T_0') = s(bool,happ(s(fun(X8,bool),X9),s(X8,sK4(X9,X10,X8,bool))))) ) | $spl6),
  inference(subsumption_resolution,[],[f1967,f84])).
fof(f5867,plain,(
  ( ! [X26,X25] : (s(bool,'T_0') = s(bool,X26) | s(bool,X26) = s(bool,'const.combin.I_1'(s(bool,X25)))) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f2624,f84])).
fof(f5865,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') = s(bool,X28) | s(bool,X28) = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X27)))) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f2625,f84])).
fof(f5861,plain,(
  ( ! [X0,X1] : (~p(s(bool,'F_0')) | s(fun(X0,X0),'const.combin.I_0') = s(fun(X0,X0),'const.combin.o_2'(s(fun(X0,X0),X1),s(fun(X0,X0),X1)))) ) | $spl6),
  inference(forward_demodulation,[],[f55,f84])).
fof(f5859,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | s(fun(X0,X0),'const.combin.I_0') != s(fun(X0,X0),'const.combin.o_2'(s(fun(X0,X0),X1),s(fun(X0,X0),X1)))) ) | $spl6),
  inference(forward_demodulation,[],[f56,f84])).
fof(f5845,plain,(
  ( ! [X43,X42] : (s(bool,'T_0') = s(bool,X43) | s(bool,X43) = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X42)))) ) | ($spl6 | $spl26)),
  inference(subsumption_resolution,[],[f5844,f84])).
fof(f5844,plain,(
  ( ! [X43,X42] : (s(bool,'F_0') != s(bool,X42) | s(bool,'T_0') = s(bool,X43) | s(bool,X43) = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X42)))) ) | ($spl6 | $spl26)),
  inference(forward_demodulation,[],[f3702,f84])).
fof(f3727,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') != s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1)))) ) | $spl26),
  inference(condensation,[],[f3726])).
fof(f3726,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') != s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,X1)) ) | $spl26),
  inference(duplicate_literal_removal,[],[f3725])).
fof(f3725,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') != s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | $spl26),
  inference(condensation,[],[f3689])).
fof(f3689,plain,(
  ( ! [X12,X13,X11] : (s(bool,'T_0') != s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,X12) | s(bool,X11) = s(bool,X12) | s(bool,X11) = s(bool,X13) | s(bool,'T_0') = s(bool,X13)) ) | $spl26),
  inference(superposition,[],[f2131,f73])).
fof(f3730,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') != s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1)))) ) | $spl26),
  inference(condensation,[],[f3729])).
fof(f3729,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') != s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,X1)) ) | $spl26),
  inference(duplicate_literal_removal,[],[f3728])).
fof(f3728,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') != s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | $spl26),
  inference(condensation,[],[f3686])).
fof(f3686,plain,(
  ( ! [X4,X2,X3] : (s(bool,'T_0') != s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,X3) | s(bool,X2) = s(bool,X3) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) ) | $spl26),
  inference(superposition,[],[f2131,f73])).
fof(f2943,plain,(
  ( ! [X20] : (s(bool,X20) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X20)))) ) | $spl22),
  inference(superposition,[],[f2907,f51])).
fof(f5553,plain,(
  ( ! [X30,X29] : (s(bool,X30) = s(bool,'const.combin.I_1'(s(bool,X29))) | s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X30)) ) | $spl22),
  inference(subsumption_resolution,[],[f5552,f2624])).
fof(f5552,plain,(
  ( ! [X30,X29] : (s(bool,X30) = s(bool,'const.combin.I_1'(s(bool,X29))) | s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X30) | s(bool,'F_0') = s(bool,X29)) ) | $spl22),
  inference(forward_demodulation,[],[f2899,f51])).
fof(f2899,plain,(
  ( ! [X30,X29] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X30) | s(bool,X30) = s(bool,'const.combin.I_1'(s(bool,'const.combin.I_1'(s(bool,X29))))) | s(bool,'F_0') = s(bool,X29)) ) | $spl22),
  inference(superposition,[],[f2624,f75])).
fof(f3703,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) != s(bool,'const.combin.o_3'(s(fun(X44,bool),X45),s(fun(X46,X44),X47),s(X46,X48))) | s(bool,'T_0') = s(bool,X49) | s(bool,X49) = s(bool,happ(s(fun(X44,bool),X45),s(X44,happ(s(fun(X46,X44),X47),s(X46,X48)))))) ) | $spl26),
  inference(superposition,[],[f2131,f60])).
fof(f3704,plain,(
  ( ! [X54,X52,X50,X55,X53,X51] : (s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) != s(bool,'const.combin.o_3'(s(fun(X51,bool),X52),s(fun(X50,X51),X53),s(X50,X54))) | s(bool,'T_0') = s(bool,X55) | s(bool,X55) = s(bool,happ(s(fun(X50,bool),'const.combin.o_2'(s(fun(X51,bool),X52),s(fun(X50,X51),X53))),s(X50,X54)))) ) | $spl26),
  inference(superposition,[],[f2131,f61])).
fof(f3810,plain,(
  ( ! [X70,X72,X71,X73] : (s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) != s(bool,happ(s(fun(X70,bool),X71),s(X70,X72))) | s(bool,'T_0') = s(bool,X73) | s(bool,X73) = s(bool,'const.combin.o_3'(s(fun(X70,bool),X71),s(fun(X70,X70),'const.combin.I_0'),s(X70,X72)))) ) | $spl26),
  inference(superposition,[],[f2131,f1293])).
fof(f4853,plain,(
  ( ! [X80,X81,X79,X82] : (s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) != s(bool,happ(s(fun(X79,bool),X80),s(X79,X81))) | s(bool,'T_0') = s(bool,X82) | s(bool,X82) = s(bool,'const.combin.o_3'(s(fun(bool,bool),'const.combin.I_0'),s(fun(X79,bool),X80),s(X79,X81)))) ) | $spl26),
  inference(superposition,[],[f2131,f1305])).
fof(f4852,plain,(
  ( ! [X78,X76,X77,X75] : (s(bool,'F_0') != s(bool,happ(s(fun(X75,bool),X76),s(X75,X77))) | s(bool,'T_0') = s(bool,X78) | s(bool,X78) = s(bool,'const.combin.o_3'(s(fun(bool,bool),'const.combin.I_0'),s(fun(X75,bool),X76),s(X75,X77)))) ) | $spl22),
  inference(superposition,[],[f2120,f1305])).
fof(f3809,plain,(
  ( ! [X68,X66,X69,X67] : (s(bool,'F_0') != s(bool,happ(s(fun(X66,bool),X67),s(X66,X68))) | s(bool,'T_0') = s(bool,X69) | s(bool,X69) = s(bool,'const.combin.o_3'(s(fun(X66,bool),X67),s(fun(X66,X66),'const.combin.I_0'),s(X66,X68)))) ) | $spl22),
  inference(superposition,[],[f2120,f1293])).
fof(f3701,plain,(
  ( ! [X41,X40] : (s(bool,X40) != s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,X41) | s(bool,X41) = s(bool,'const.combin.I_1'(s(bool,X40)))) ) | $spl26),
  inference(superposition,[],[f2131,f51])).
fof(f2881,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X2) = s(bool,'const.combin.I_1'(s(bool,X1))) | s(bool,'T_0') = s(bool,X0)) ) | $spl22),
  inference(superposition,[],[f2624,f50])).
fof(f2627,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (s(bool,'F_0') != s(bool,'const.combin.o_3'(s(fun(X36,bool),X37),s(fun(X35,X36),X38),s(X35,X39))) | s(bool,'T_0') = s(bool,X40) | s(bool,X40) = s(bool,happ(s(fun(X35,bool),'const.combin.o_2'(s(fun(X36,bool),X37),s(fun(X35,X36),X38))),s(X35,X39)))) ) | $spl22),
  inference(superposition,[],[f2120,f61])).
fof(f2626,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (s(bool,'F_0') != s(bool,'const.combin.o_3'(s(fun(X29,bool),X30),s(fun(X31,X29),X32),s(X31,X33))) | s(bool,'T_0') = s(bool,X34) | s(bool,X34) = s(bool,happ(s(fun(X29,bool),X30),s(X29,happ(s(fun(X31,X29),X32),s(X31,X33)))))) ) | $spl22),
  inference(superposition,[],[f2120,f60])).
fof(f2625,plain,(
  ( ! [X28,X27] : (s(bool,'F_0') != s(bool,X27) | s(bool,'T_0') = s(bool,X28) | s(bool,X28) = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X27)))) ) | $spl22),
  inference(superposition,[],[f2120,f68])).
fof(f417,plain,(
  s(bool,'T_0') = s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | $spl12),
  inference(cnf_transformation,[],[f417_D])).
fof(f417_D,plain,(
  s(bool,'T_0') = s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f6036,plain,(
  $false | ($spl6 | $spl12 | $spl22 | $spl26 | $spl29)),
  inference(subsumption_resolution,[],[f2142,f6021])).
fof(f6035,plain,(
  $false | ($spl1 | $spl6 | $spl12 | $spl22 | $spl26)),
  inference(subsumption_resolution,[],[f6034,f47])).
fof(f6034,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl6 | $spl12 | $spl22 | $spl26)),
  inference(forward_demodulation,[],[f5829,f6021])).
fof(f5829,plain,(
  ~p(s(bool,'F_0')) | ($spl1 | $spl6)),
  inference(forward_demodulation,[],[f64,f84])).
fof(f6033,plain,(
  $false | ($spl1 | $spl6 | $spl12 | $spl22 | $spl26)),
  inference(global_subsumption,[],[f2536,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f417,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f2131,f2120,f2594,f2624,f2625,f2626,f2627,f2881,f3322,f3323,f3332,f3350,f3352,f3353,f3701,f3784,f3785,f3786,f3846,f3805,f3809,f3870,f3871,f4824,f4825,f4900,f4852,f4929,f4930,f4853,f3810,f3704,f3703,f3702,f5533,f5534,f5553,f5661,f2634,f2637,f84,f2943,f3730,f3727,f5845,f5859,f5861,f2907,f5865,f5867,f5869,f1956,f5872,f467,f49,f5931,f5934,f5936,f5938,f5974,f5976,f5987,f5989,f5997,f6001,f6003,f3023,f2998,f6020,f6021,f6027])).
fof(f6032,plain,(
  $false | ($spl1 | $spl6 | $spl12 | $spl22 | $spl26)),
  inference(global_subsumption,[],[f2537,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f417,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f2131,f2120,f2594,f2624,f2625,f2626,f2627,f2881,f3322,f3323,f3332,f3350,f3352,f3353,f3701,f3784,f3785,f3786,f3846,f3805,f3809,f3870,f3871,f4824,f4825,f4900,f4852,f4929,f4930,f4853,f3810,f3704,f3703,f3702,f5533,f5534,f5553,f5661,f2634,f2637,f84,f2943,f3730,f3727,f5845,f5859,f5861,f2907,f5865,f5867,f5869,f1956,f5872,f467,f49,f5931,f5934,f5936,f5938,f5974,f5976,f5987,f5989,f5997,f6001,f6003,f3023,f2998,f6020,f6021,f6027])).
fof(f6031,plain,(
  $false | ($spl1 | $spl6 | $spl12 | $spl22 | $spl26)),
  inference(subsumption_resolution,[],[f6030,f47])).
fof(f6030,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl6 | $spl12 | $spl22 | $spl26)),
  inference(forward_demodulation,[],[f5842,f6021])).
fof(f5842,plain,(
  ~p(s(bool,'F_0')) | ($spl1 | $spl6)),
  inference(forward_demodulation,[],[f5659,f84])).
fof(f6028,plain,(
  $false | ($spl6 | $spl12 | $spl22 | $spl26)),
  inference(global_subsumption,[],[f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f417,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f2131,f2120,f2594,f2624,f2625,f2626,f2627,f2881,f3322,f3323,f3332,f3350,f3352,f3353,f3701,f3784,f3785,f3786,f3846,f3805,f3809,f3870,f3871,f4824,f4825,f4900,f4852,f4929,f4930,f4853,f3810,f3704,f3703,f3702,f5533,f5534,f5553,f5661,f2634,f2637,f84,f2943,f3730,f3727,f5845,f5859,f5861,f2907,f5865,f5867,f5869,f1956,f5872,f467,f49,f5931,f5934,f5936,f5938,f5974,f5976,f5987,f5989,f5997,f6001,f6003,f3023,f2998,f6020,f6021,f6027])).
fof(f6025,plain,(
  $false | ($spl6 | $spl12 | $spl22 | $spl26)),
  inference(subsumption_resolution,[],[f6024,f47])).
fof(f6024,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl12 | $spl22 | $spl26)),
  inference(forward_demodulation,[],[f49,f6021])).
fof(f6023,plain,(
  $false | ($spl6 | $spl12 | $spl22 | $spl26 | $spl29)),
  inference(subsumption_resolution,[],[f6022,f6021])).
fof(f6022,plain,(
  s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | ($spl6 | $spl12 | $spl22 | $spl26 | $spl29)),
  inference(backward_demodulation,[],[f6021,f6015])).
fof(f6014,plain,(
  $false | ($spl6 | $spl9)),
  inference(subsumption_resolution,[],[f86,f84])).
fof(f6013,plain,(
  $false | ($spl6 | $spl13 | $spl20 | $spl29)),
  inference(global_subsumption,[],[f418,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142,f5859,f5861,f5869,f1956,f5872,f467,f49,f5997,f6001,f6003])).
fof(f5705,plain,(
  s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | ($spl6 | $spl20)),
  inference(forward_demodulation,[],[f1262,f84])).
fof(f6009,plain,(
  $false | ($spl1 | $spl6 | $spl20 | $spl29)),
  inference(global_subsumption,[],[f6008,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142,f5859,f5861,f5869,f1956,f5872,f467,f49,f5997,f6001,f6003])).
fof(f6008,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | ($spl1 | $spl6)),
  inference(forward_demodulation,[],[f6007,f84])).
fof(f6007,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,g(s(X0,sK3(X1,X2,bool,X0)))))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | ($spl1 | $spl6)),
  inference(condensation,[],[f6006])).
fof(f6006,plain,(
  ( ! [X156,X154,X155,X153] : (s(fun(X153,bool),X154) = s(fun(X153,bool),X155) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X156))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,g(s(X153,sK3(X154,X155,bool,X153))))))) ) | ($spl1 | $spl6)),
  inference(subsumption_resolution,[],[f6005,f84])).
fof(f6005,plain,(
  ( ! [X156,X154,X155,X153] : (s(bool,'F_0') != s(bool,X156) | s(fun(X153,bool),X154) = s(fun(X153,bool),X155) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X156))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,g(s(X153,sK3(X154,X155,bool,X153))))))) ) | ($spl1 | $spl6)),
  inference(forward_demodulation,[],[f4172,f84])).
fof(f6004,plain,(
  $false | ($spl6 | $spl20 | $spl29)),
  inference(global_subsumption,[],[f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142,f5859,f5861,f5869,f1956,f5872,f467,f49,f5997,f6001,f6003])).
fof(f6002,plain,(
  $false | ($spl6 | $spl20 | $spl29)),
  inference(global_subsumption,[],[f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142,f5859,f5861,f5869,f1956,f5872,f467,f49,f5997,f6001])).
fof(f5998,plain,(
  $false | ($spl6 | $spl20 | $spl29)),
  inference(global_subsumption,[],[f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142,f5859,f5861,f5869,f1956,f5872,f467,f49,f5997])).
fof(f5990,plain,(
  $false | ($spl6 | $spl20 | $spl22 | $spl29)),
  inference(global_subsumption,[],[f5989,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142,f5859,f5861,f5869,f1956,f5872,f467,f49])).
fof(f5988,plain,(
  $false | ($spl6 | $spl20 | $spl26 | $spl29)),
  inference(global_subsumption,[],[f5987,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142,f5859,f5861,f5869,f1956,f5872,f467,f49])).
fof(f5984,plain,(
  $false | ($spl6 | $spl16 | $spl20 | $spl29)),
  inference(global_subsumption,[],[f5983,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142,f5859,f5861,f5869,f1956,f5872,f467,f49])).
fof(f5983,plain,(
  ( ! [X59,X57,X58] : (s(bool,'T_0') = s(bool,'const.combin.o_3'(s(fun(X57,bool),X58),s(fun(X57,X57),'const.combin.I_0'),s(X57,X59)))) ) | ($spl6 | $spl16)),
  inference(subsumption_resolution,[],[f5982,f84])).
fof(f5982,plain,(
  ( ! [X59,X57,X58] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.combin.o_3'(s(fun(X57,bool),X58),s(fun(X57,X57),'const.combin.I_0'),s(X57,X59)))) ) | ($spl6 | $spl16)),
  inference(forward_demodulation,[],[f3848,f84])).
fof(f3848,plain,(
  ( ! [X59,X57,X58] : (s(bool,'T_0') != s(bool,happ(s(fun(X57,bool),X58),s(X57,X59))) | s(bool,'T_0') = s(bool,'const.combin.o_3'(s(fun(X57,bool),X58),s(fun(X57,X57),'const.combin.I_0'),s(X57,X59)))) ) | $spl16),
  inference(forward_demodulation,[],[f3806,f51])).
fof(f3806,plain,(
  ( ! [X59,X57,X58] : (s(bool,'const.combin.I_1'(s(bool,'T_0'))) != s(bool,happ(s(fun(X57,bool),X58),s(X57,X59))) | s(bool,'T_0') = s(bool,'const.combin.o_3'(s(fun(X57,bool),X58),s(fun(X57,X57),'const.combin.I_0'),s(X57,X59)))) ) | $spl16),
  inference(superposition,[],[f1113,f1293])).
fof(f5977,plain,(
  $false | ($spl6 | $spl20 | $spl22 | $spl29)),
  inference(global_subsumption,[],[f5976,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142,f5859,f5861,f5869,f1956,f5872,f467,f49])).
fof(f5975,plain,(
  $false | ($spl6 | $spl20 | $spl26 | $spl29)),
  inference(global_subsumption,[],[f5974,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142,f5859,f5861,f5869,f1956,f5872,f467,f49])).
fof(f5971,plain,(
  $false | ($spl6 | $spl16 | $spl20 | $spl29)),
  inference(global_subsumption,[],[f5970,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142,f5859,f5861,f5869,f1956,f5872,f467,f49])).
fof(f5970,plain,(
  ( ! [X68,X66,X67] : (s(bool,'T_0') = s(bool,'const.combin.o_3'(s(fun(bool,bool),'const.combin.I_0'),s(fun(X66,bool),X67),s(X66,X68)))) ) | ($spl6 | $spl16)),
  inference(subsumption_resolution,[],[f5969,f84])).
fof(f5969,plain,(
  ( ! [X68,X66,X67] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.combin.o_3'(s(fun(bool,bool),'const.combin.I_0'),s(fun(X66,bool),X67),s(X66,X68)))) ) | ($spl6 | $spl16)),
  inference(forward_demodulation,[],[f4902,f84])).
fof(f4902,plain,(
  ( ! [X68,X66,X67] : (s(bool,'T_0') != s(bool,happ(s(fun(X66,bool),X67),s(X66,X68))) | s(bool,'T_0') = s(bool,'const.combin.o_3'(s(fun(bool,bool),'const.combin.I_0'),s(fun(X66,bool),X67),s(X66,X68)))) ) | $spl16),
  inference(forward_demodulation,[],[f4849,f51])).
fof(f4849,plain,(
  ( ! [X68,X66,X67] : (s(bool,'const.combin.I_1'(s(bool,'T_0'))) != s(bool,happ(s(fun(X66,bool),X67),s(X66,X68))) | s(bool,'T_0') = s(bool,'const.combin.o_3'(s(fun(bool,bool),'const.combin.I_0'),s(fun(X66,bool),X67),s(X66,X68)))) ) | $spl16),
  inference(superposition,[],[f1113,f1305])).
fof(f5960,plain,(
  $false | ($spl4 | $spl6 | $spl20 | $spl22 | $spl29)),
  inference(global_subsumption,[],[f5959,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142,f5859,f5861,f5869,f1956,f5872,f467,f49])).
fof(f5959,plain,(
  ( ! [X50,X48,X49] : (s(bool,'T_0') = s(bool,X50) | s(bool,X50) = s(bool,'const.combin.o_3'(s(fun(sK0,bool),X48),s(fun(sK0,sK0),sK1),s(sK0,happ(s(fun(sK0,sK0),sK1),s(sK0,X49)))))) ) | ($spl4 | $spl6 | $spl22)),
  inference(subsumption_resolution,[],[f3126,f84])).
fof(f3126,plain,(
  ( ! [X50,X48,X49] : (s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),X48),s(sK0,X49))) | s(bool,'T_0') = s(bool,X50) | s(bool,X50) = s(bool,'const.combin.o_3'(s(fun(sK0,bool),X48),s(fun(sK0,sK0),sK1),s(sK0,happ(s(fun(sK0,sK0),sK1),s(sK0,X49)))))) ) | ($spl4 | $spl22)),
  inference(superposition,[],[f2120,f2368])).
fof(f2368,plain,(
  ( ! [X2,X3,X1] : (s(X2,happ(s(fun(sK0,X2),X3),s(sK0,X1))) = s(X2,'const.combin.o_3'(s(fun(sK0,X2),X3),s(fun(sK0,sK0),sK1),s(sK0,happ(s(fun(sK0,sK0),sK1),s(sK0,X1)))))) ) | $spl4),
  inference(superposition,[],[f60,f67])).
fof(f5957,plain,(
  $false | ($spl4 | $spl6 | $spl16 | $spl20 | $spl29)),
  inference(global_subsumption,[],[f5956,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142,f5859,f5861,f5869,f1956,f5872,f467,f49])).
fof(f5956,plain,(
  ( ! [X43,X42] : (s(bool,'T_0') = s(bool,'const.combin.o_3'(s(fun(sK0,bool),X42),s(fun(sK0,sK0),sK1),s(sK0,happ(s(fun(sK0,sK0),sK1),s(sK0,X43)))))) ) | ($spl4 | $spl6 | $spl16)),
  inference(subsumption_resolution,[],[f5955,f84])).
fof(f5955,plain,(
  ( ! [X43,X42] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.combin.o_3'(s(fun(sK0,bool),X42),s(fun(sK0,sK0),sK1),s(sK0,happ(s(fun(sK0,sK0),sK1),s(sK0,X43)))))) ) | ($spl4 | $spl6 | $spl16)),
  inference(forward_demodulation,[],[f3157,f84])).
fof(f3157,plain,(
  ( ! [X43,X42] : (s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),X42),s(sK0,X43))) | s(bool,'T_0') = s(bool,'const.combin.o_3'(s(fun(sK0,bool),X42),s(fun(sK0,sK0),sK1),s(sK0,happ(s(fun(sK0,sK0),sK1),s(sK0,X43)))))) ) | ($spl4 | $spl16)),
  inference(forward_demodulation,[],[f3123,f51])).
fof(f3123,plain,(
  ( ! [X43,X42] : (s(bool,'const.combin.I_1'(s(bool,'T_0'))) != s(bool,happ(s(fun(sK0,bool),X42),s(sK0,X43))) | s(bool,'T_0') = s(bool,'const.combin.o_3'(s(fun(sK0,bool),X42),s(fun(sK0,sK0),sK1),s(sK0,happ(s(fun(sK0,sK0),sK1),s(sK0,X43)))))) ) | ($spl4 | $spl16)),
  inference(superposition,[],[f1113,f2368])).
fof(f5953,plain,(
  $false | ($spl4 | $spl6 | $spl20 | $spl26 | $spl29)),
  inference(global_subsumption,[],[f5952,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142,f5859,f5861,f5869,f1956,f5872,f467,f49])).
fof(f5952,plain,(
  ( ! [X57,X58,X56] : (s(bool,'T_0') = s(bool,X58) | s(bool,X58) = s(bool,'const.combin.o_3'(s(fun(sK0,bool),X56),s(fun(sK0,sK0),sK1),s(sK0,happ(s(fun(sK0,sK0),sK1),s(sK0,X57)))))) ) | ($spl4 | $spl6 | $spl26)),
  inference(subsumption_resolution,[],[f5951,f84])).
fof(f5951,plain,(
  ( ! [X57,X58,X56] : (s(bool,'F_0') != s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,X58) | s(bool,X58) = s(bool,'const.combin.o_3'(s(fun(sK0,bool),X56),s(fun(sK0,sK0),sK1),s(sK0,happ(s(fun(sK0,sK0),sK1),s(sK0,X57)))))) ) | ($spl4 | $spl6 | $spl26)),
  inference(forward_demodulation,[],[f3705,f84])).
fof(f3705,plain,(
  ( ! [X57,X58,X56] : (s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) != s(bool,happ(s(fun(sK0,bool),X56),s(sK0,X57))) | s(bool,'T_0') = s(bool,X58) | s(bool,X58) = s(bool,'const.combin.o_3'(s(fun(sK0,bool),X56),s(fun(sK0,sK0),sK1),s(sK0,happ(s(fun(sK0,sK0),sK1),s(sK0,X57)))))) ) | ($spl4 | $spl26)),
  inference(superposition,[],[f2131,f2368])).
fof(f5947,plain,(
  $false | ($spl6 | $spl16 | $spl20 | $spl29)),
  inference(global_subsumption,[],[f5946,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142,f5859,f5861,f5869,f1956,f5872,f467,f49])).
fof(f5946,plain,(
  ( ! [X30,X28,X31,X29,X27] : (s(bool,'T_0') = s(bool,happ(s(fun(X27,bool),X28),s(X27,happ(s(fun(X29,X27),X30),s(X29,X31)))))) ) | ($spl6 | $spl16)),
  inference(subsumption_resolution,[],[f5945,f84])).
fof(f5945,plain,(
  ( ! [X30,X28,X31,X29,X27] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,happ(s(fun(X27,bool),X28),s(X27,happ(s(fun(X29,X27),X30),s(X29,X31)))))) ) | ($spl6 | $spl16)),
  inference(forward_demodulation,[],[f1700,f84])).
fof(f1700,plain,(
  ( ! [X30,X28,X31,X29,X27] : (s(bool,'T_0') != s(bool,'const.combin.o_3'(s(fun(X27,bool),X28),s(fun(X29,X27),X30),s(X29,X31))) | s(bool,'T_0') = s(bool,happ(s(fun(X27,bool),X28),s(X27,happ(s(fun(X29,X27),X30),s(X29,X31)))))) ) | $spl16),
  inference(forward_demodulation,[],[f1655,f51])).
fof(f1655,plain,(
  ( ! [X30,X28,X31,X29,X27] : (s(bool,'const.combin.I_1'(s(bool,'T_0'))) != s(bool,'const.combin.o_3'(s(fun(X27,bool),X28),s(fun(X29,X27),X30),s(X29,X31))) | s(bool,'T_0') = s(bool,happ(s(fun(X27,bool),X28),s(X27,happ(s(fun(X29,X27),X30),s(X29,X31)))))) ) | $spl16),
  inference(superposition,[],[f1113,f60])).
fof(f5942,plain,(
  $false | ($spl6 | $spl16 | $spl20 | $spl29)),
  inference(global_subsumption,[],[f5941,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142,f5859,f5861,f5869,f1956,f5872,f467,f49])).
fof(f5941,plain,(
  ( ! [X111,X114,X112,X115,X113] : (s(bool,'T_0') = s(bool,happ(s(fun(X111,bool),'const.combin.o_2'(s(fun(X112,bool),X113),s(fun(X111,X112),X114))),s(X111,X115)))) ) | ($spl6 | $spl16)),
  inference(subsumption_resolution,[],[f5940,f84])).
fof(f5940,plain,(
  ( ! [X111,X114,X112,X115,X113] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,happ(s(fun(X111,bool),'const.combin.o_2'(s(fun(X112,bool),X113),s(fun(X111,X112),X114))),s(X111,X115)))) ) | ($spl6 | $spl16)),
  inference(forward_demodulation,[],[f1814,f84])).
fof(f1814,plain,(
  ( ! [X111,X114,X112,X115,X113] : (s(bool,'T_0') != s(bool,'const.combin.o_3'(s(fun(X112,bool),X113),s(fun(X111,X112),X114),s(X111,X115))) | s(bool,'T_0') = s(bool,happ(s(fun(X111,bool),'const.combin.o_2'(s(fun(X112,bool),X113),s(fun(X111,X112),X114))),s(X111,X115)))) ) | $spl16),
  inference(forward_demodulation,[],[f1784,f51])).
fof(f1784,plain,(
  ( ! [X111,X114,X112,X115,X113] : (s(bool,'const.combin.I_1'(s(bool,'T_0'))) != s(bool,'const.combin.o_3'(s(fun(X112,bool),X113),s(fun(X111,X112),X114),s(X111,X115))) | s(bool,'T_0') = s(bool,happ(s(fun(X111,bool),'const.combin.o_2'(s(fun(X112,bool),X113),s(fun(X111,X112),X114))),s(X111,X115)))) ) | $spl16),
  inference(superposition,[],[f1113,f61])).
fof(f5939,plain,(
  $false | ($spl6 | $spl20 | $spl22 | $spl29)),
  inference(global_subsumption,[],[f5938,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142,f5859,f5861,f5869,f1956,f5872,f467,f49])).
fof(f5937,plain,(
  $false | ($spl6 | $spl20 | $spl22 | $spl29)),
  inference(global_subsumption,[],[f5936,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142,f5859,f5861,f5869,f1956,f5872,f467,f49])).
fof(f5935,plain,(
  $false | ($spl6 | $spl20 | $spl26 | $spl29)),
  inference(global_subsumption,[],[f5934,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142,f5859,f5861,f5869,f1956,f5872,f467,f49])).
fof(f5932,plain,(
  $false | ($spl6 | $spl20 | $spl26 | $spl29)),
  inference(global_subsumption,[],[f5931,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142,f5859,f5861,f5869,f1956,f5872,f467,f49])).
fof(f5924,plain,(
  $false | ($spl1 | $spl6 | $spl20 | $spl29)),
  inference(global_subsumption,[],[f5923,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142,f5859,f5861,f5869,f1956,f5872,f467,f49])).
fof(f5923,plain,(
  ( ! [X54,X56,X55] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(fun(X54,bool),X55) = s(fun(X54,bool),X56)) ) | ($spl1 | $spl6)),
  inference(forward_demodulation,[],[f5922,f84])).
fof(f5922,plain,(
  ( ! [X54,X56,X55] : (s(fun(X54,bool),X55) = s(fun(X54,bool),X56) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,happ(s(fun(X54,bool),X55),s(X54,sK4(X55,X56,X54,bool))))))) ) | ($spl1 | $spl6)),
  inference(subsumption_resolution,[],[f5921,f84])).
fof(f5921,plain,(
  ( ! [X54,X56,X55] : (s(bool,'F_0') != s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(fun(X54,bool),X55) = s(fun(X54,bool),X56) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,happ(s(fun(X54,bool),X55),s(X54,sK4(X55,X56,X54,bool))))))) ) | ($spl1 | $spl6)),
  inference(forward_demodulation,[],[f2720,f84])).
fof(f2720,plain,(
  ( ! [X54,X56,X55] : (s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) != s(bool,happ(s(fun(X54,bool),X56),s(X54,sK4(X55,X56,X54,bool)))) | s(fun(X54,bool),X55) = s(fun(X54,bool),X56) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,happ(s(fun(X54,bool),X55),s(X54,sK4(X55,X56,X54,bool))))))) ) | $spl1),
  inference(superposition,[],[f59,f2434])).
fof(f5920,plain,(
  $false | ($spl1 | $spl6 | $spl20 | $spl29)),
  inference(global_subsumption,[],[f5919,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142,f5859,f5861,f5869,f1956,f5872,f467,f49])).
fof(f5919,plain,(
  ( ! [X52,X53,X51] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(fun(X51,bool),X52) = s(fun(X51,bool),X53)) ) | ($spl1 | $spl6)),
  inference(forward_demodulation,[],[f5918,f84])).
fof(f5918,plain,(
  ( ! [X52,X53,X51] : (s(fun(X51,bool),X52) = s(fun(X51,bool),X53) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,happ(s(fun(X51,bool),X52),s(X51,sK4(X53,X52,X51,bool))))))) ) | ($spl1 | $spl6)),
  inference(subsumption_resolution,[],[f5917,f84])).
fof(f5917,plain,(
  ( ! [X52,X53,X51] : (s(bool,'F_0') != s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(fun(X51,bool),X52) = s(fun(X51,bool),X53) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,happ(s(fun(X51,bool),X52),s(X51,sK4(X53,X52,X51,bool))))))) ) | ($spl1 | $spl6)),
  inference(forward_demodulation,[],[f2719,f84])).
fof(f2719,plain,(
  ( ! [X52,X53,X51] : (s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) != s(bool,happ(s(fun(X51,bool),X53),s(X51,sK4(X53,X52,X51,bool)))) | s(fun(X51,bool),X52) = s(fun(X51,bool),X53) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,happ(s(fun(X51,bool),X52),s(X51,sK4(X53,X52,X51,bool))))))) ) | $spl1),
  inference(superposition,[],[f59,f2434])).
fof(f5913,plain,(
  $false | ($spl1 | $spl6 | $spl20 | $spl29)),
  inference(global_subsumption,[],[f5912,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142,f5859,f5861,f5869,f1956,f5872,f467,f49])).
fof(f5912,plain,(
  ( ! [X39,X37,X38] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(fun(X37,bool),X38) = s(fun(X37,bool),X39)) ) | ($spl1 | $spl6)),
  inference(forward_demodulation,[],[f5911,f84])).
fof(f5911,plain,(
  ( ! [X39,X37,X38] : (s(fun(X37,bool),X38) = s(fun(X37,bool),X39) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,g(s(X37,sK3(X38,X39,bool,X37))))))) ) | ($spl1 | $spl6)),
  inference(subsumption_resolution,[],[f5910,f84])).
fof(f5910,plain,(
  ( ! [X39,X37,X38] : (s(bool,'F_0') != s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(fun(X37,bool),X38) = s(fun(X37,bool),X39) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,g(s(X37,sK3(X38,X39,bool,X37))))))) ) | ($spl1 | $spl6)),
  inference(forward_demodulation,[],[f2715,f84])).
fof(f2715,plain,(
  ( ! [X39,X37,X38] : (s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) != s(bool,f(s(X37,sK3(X38,X39,bool,X37)))) | s(fun(X37,bool),X38) = s(fun(X37,bool),X39) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,g(s(X37,sK3(X38,X39,bool,X37))))))) ) | $spl1),
  inference(superposition,[],[f57,f2434])).
fof(f5909,plain,(
  $false | ($spl1 | $spl6 | $spl16 | $spl20 | $spl29)),
  inference(global_subsumption,[],[f5908,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142,f5859,f5861,f5869,f1956,f5872,f467,f49])).
fof(f5908,plain,(
  ( ! [X12] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X12)))) ) | ($spl1 | $spl6 | $spl16)),
  inference(subsumption_resolution,[],[f5907,f84])).
fof(f5907,plain,(
  ( ! [X12] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X12)))) ) | ($spl1 | $spl6 | $spl16)),
  inference(forward_demodulation,[],[f2771,f84])).
fof(f2771,plain,(
  ( ! [X12] : (s(bool,'const.combin.I_1'(s(bool,'T_0'))) != s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X12)))) ) | ($spl1 | $spl16)),
  inference(duplicate_literal_removal,[],[f2702])).
fof(f2702,plain,(
  ( ! [X12] : (s(bool,'const.combin.I_1'(s(bool,'T_0'))) != s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X12))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X12)))) ) | ($spl1 | $spl16)),
  inference(superposition,[],[f1653,f2434])).
fof(f1653,plain,(
  ( ! [X25] : (s(bool,X25) != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X25)))) ) | $spl16),
  inference(superposition,[],[f1113,f51])).
fof(f5906,plain,(
  $false | ($spl1 | $spl6 | $spl16 | $spl20 | $spl29)),
  inference(global_subsumption,[],[f5905,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142,f5859,f5861,f5869,f1956,f5872,f467,f49])).
fof(f5905,plain,(
  ( ! [X11] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X11)))) ) | ($spl1 | $spl6 | $spl16)),
  inference(subsumption_resolution,[],[f5904,f84])).
fof(f5904,plain,(
  ( ! [X11] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X11)))) ) | ($spl1 | $spl6 | $spl16)),
  inference(forward_demodulation,[],[f2804,f84])).
fof(f2804,plain,(
  ( ! [X11] : (s(bool,'const.combin.I_1'(s(bool,'T_0'))) != s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X11)))) ) | ($spl1 | $spl16)),
  inference(subsumption_resolution,[],[f2701,f2245])).
fof(f2245,plain,(
  ( ! [X14] : (s(bool,'T_0') != s(bool,X14) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X14)))) ) | $spl16),
  inference(forward_demodulation,[],[f2244,f51])).
fof(f2244,plain,(
  ( ! [X14] : (s(bool,'T_0') != s(bool,X14) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'const.combin.I_1'(s(bool,X14)))))) ) | $spl16),
  inference(forward_demodulation,[],[f2216,f51])).
fof(f2216,plain,(
  ( ! [X14] : (s(bool,X14) != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'const.combin.I_1'(s(bool,X14)))))) ) | $spl16),
  inference(superposition,[],[f1653,f51])).
fof(f2701,plain,(
  ( ! [X11] : (s(bool,'const.combin.I_1'(s(bool,'T_0'))) != s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X11)))) ) | ($spl1 | $spl16)),
  inference(superposition,[],[f1113,f2434])).
fof(f5894,plain,(
  $false | ($spl1 | $spl6 | $spl20 | $spl29)),
  inference(global_subsumption,[],[f5893,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142,f5859,f5861,f5869,f1956,f5872,f467,f49])).
fof(f5893,plain,(
  ( ! [X50,X48,X49] : (s(fun(X48,bool),X49) = s(fun(X48,bool),X50) | s(bool,'T_0') = s(bool,happ(s(fun(X48,bool),X49),s(X48,sK4(X49,X50,X48,bool))))) ) | ($spl1 | $spl6)),
  inference(subsumption_resolution,[],[f5892,f84])).
fof(f5892,plain,(
  ( ! [X50,X48,X49] : (s(bool,'F_0') != s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(fun(X48,bool),X49) = s(fun(X48,bool),X50) | s(bool,'T_0') = s(bool,happ(s(fun(X48,bool),X49),s(X48,sK4(X49,X50,X48,bool))))) ) | ($spl1 | $spl6)),
  inference(forward_demodulation,[],[f2468,f84])).
fof(f2468,plain,(
  ( ! [X50,X48,X49] : (s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) != s(bool,happ(s(fun(X48,bool),X50),s(X48,sK4(X49,X50,X48,bool)))) | s(fun(X48,bool),X49) = s(fun(X48,bool),X50) | s(bool,'T_0') = s(bool,happ(s(fun(X48,bool),X49),s(X48,sK4(X49,X50,X48,bool))))) ) | $spl1),
  inference(superposition,[],[f59,f2359])).
fof(f5891,plain,(
  $false | ($spl1 | $spl6 | $spl20 | $spl29)),
  inference(global_subsumption,[],[f5890,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142,f5859,f5861,f5869,f1956,f5872,f467,f49])).
fof(f5890,plain,(
  ( ! [X47,X45,X46] : (s(fun(X45,bool),X46) = s(fun(X45,bool),X47) | s(bool,'T_0') = s(bool,happ(s(fun(X45,bool),X46),s(X45,sK4(X47,X46,X45,bool))))) ) | ($spl1 | $spl6)),
  inference(subsumption_resolution,[],[f5889,f84])).
fof(f5889,plain,(
  ( ! [X47,X45,X46] : (s(bool,'F_0') != s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(fun(X45,bool),X46) = s(fun(X45,bool),X47) | s(bool,'T_0') = s(bool,happ(s(fun(X45,bool),X46),s(X45,sK4(X47,X46,X45,bool))))) ) | ($spl1 | $spl6)),
  inference(forward_demodulation,[],[f2467,f84])).
fof(f2467,plain,(
  ( ! [X47,X45,X46] : (s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) != s(bool,happ(s(fun(X45,bool),X47),s(X45,sK4(X47,X46,X45,bool)))) | s(fun(X45,bool),X46) = s(fun(X45,bool),X47) | s(bool,'T_0') = s(bool,happ(s(fun(X45,bool),X46),s(X45,sK4(X47,X46,X45,bool))))) ) | $spl1),
  inference(superposition,[],[f59,f2359])).
fof(f5884,plain,(
  $false | ($spl1 | $spl6 | $spl20 | $spl29)),
  inference(global_subsumption,[],[f5883,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142,f5859,f5861,f5869,f1956,f5872,f467,f49])).
fof(f5883,plain,(
  ( ! [X33,X31,X32] : (s(fun(X31,bool),X32) = s(fun(X31,bool),X33) | s(bool,'T_0') = s(bool,g(s(X31,sK3(X32,X33,bool,X31))))) ) | ($spl1 | $spl6)),
  inference(subsumption_resolution,[],[f5882,f84])).
fof(f5882,plain,(
  ( ! [X33,X31,X32] : (s(bool,'F_0') != s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(fun(X31,bool),X32) = s(fun(X31,bool),X33) | s(bool,'T_0') = s(bool,g(s(X31,sK3(X32,X33,bool,X31))))) ) | ($spl1 | $spl6)),
  inference(forward_demodulation,[],[f2463,f84])).
fof(f2463,plain,(
  ( ! [X33,X31,X32] : (s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) != s(bool,f(s(X31,sK3(X32,X33,bool,X31)))) | s(fun(X31,bool),X32) = s(fun(X31,bool),X33) | s(bool,'T_0') = s(bool,g(s(X31,sK3(X32,X33,bool,X31))))) ) | $spl1),
  inference(superposition,[],[f57,f2359])).
fof(f5881,plain,(
  $false | ($spl1 | $spl6 | $spl16 | $spl20 | $spl29)),
  inference(global_subsumption,[],[f5880,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142,f5859,f5861,f5869,f1956,f5872,f467,f49])).
fof(f5880,plain,(
  ( ! [X11] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X11)))) ) | ($spl1 | $spl6 | $spl16)),
  inference(subsumption_resolution,[],[f5879,f84])).
fof(f5879,plain,(
  ( ! [X11] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X11)))) ) | ($spl1 | $spl6 | $spl16)),
  inference(forward_demodulation,[],[f2518,f84])).
fof(f2518,plain,(
  ( ! [X11] : (s(bool,'const.combin.I_1'(s(bool,'T_0'))) != s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X11)))) ) | ($spl1 | $spl16)),
  inference(subsumption_resolution,[],[f2453,f2245])).
fof(f2453,plain,(
  ( ! [X11] : (s(bool,'const.combin.I_1'(s(bool,'T_0'))) != s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X11))) | s(bool,'T_0') = s(bool,X11)) ) | ($spl1 | $spl16)),
  inference(superposition,[],[f1653,f2359])).
fof(f5878,plain,(
  $false | ($spl1 | $spl6 | $spl16 | $spl20 | $spl29)),
  inference(global_subsumption,[],[f5877,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142,f5859,f5861,f5869,f1956,f5872,f467,f49])).
fof(f5877,plain,(
  ( ! [X10] : (s(bool,'T_0') = s(bool,X10)) ) | ($spl1 | $spl6 | $spl16)),
  inference(subsumption_resolution,[],[f5876,f84])).
fof(f5876,plain,(
  ( ! [X10] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X10)) ) | ($spl1 | $spl6 | $spl16)),
  inference(forward_demodulation,[],[f2496,f84])).
fof(f2496,plain,(
  ( ! [X10] : (s(bool,'const.combin.I_1'(s(bool,'T_0'))) != s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,X10)) ) | ($spl1 | $spl16)),
  inference(duplicate_literal_removal,[],[f2452])).
fof(f2452,plain,(
  ( ! [X10] : (s(bool,'const.combin.I_1'(s(bool,'T_0'))) != s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X10)) ) | ($spl1 | $spl16)),
  inference(superposition,[],[f1113,f2359])).
fof(f5875,plain,(
  $false | ($spl6 | $spl20 | $spl29)),
  inference(global_subsumption,[],[f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142,f5859,f5861,f5869,f1956,f5872,f467,f49])).
fof(f5874,plain,(
  $false | ($spl6 | $spl20 | $spl29)),
  inference(global_subsumption,[],[f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f49,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142,f5859,f5861,f5869,f1956,f5872,f467])).
fof(f5873,plain,(
  $false | ($spl6 | $spl20 | $spl29)),
  inference(global_subsumption,[],[f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f49,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f467,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142,f5859,f5861,f5869,f1956,f5872])).
fof(f5871,plain,(
  $false | ($spl6 | $spl20 | $spl29)),
  inference(global_subsumption,[],[f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f49,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f467,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142,f5859,f5861,f5869,f1956])).
fof(f5870,plain,(
  $false | ($spl6 | $spl20 | $spl29)),
  inference(global_subsumption,[],[f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f49,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1956,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f467,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142,f5859,f5861,f5869])).
fof(f5868,plain,(
  $false | ($spl6 | $spl20 | $spl22 | $spl29)),
  inference(global_subsumption,[],[f5867,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f49,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1956,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f467,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142,f5859,f5861])).
fof(f5866,plain,(
  $false | ($spl6 | $spl20 | $spl22 | $spl29)),
  inference(global_subsumption,[],[f5865,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f49,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1956,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f467,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142,f5859,f5861])).
fof(f5864,plain,(
  $false | ($spl6 | $spl20 | $spl22 | $spl29)),
  inference(global_subsumption,[],[f2907,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f49,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1956,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f467,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142,f5859,f5861])).
fof(f5862,plain,(
  $false | ($spl6 | $spl20 | $spl29)),
  inference(global_subsumption,[],[f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f49,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1956,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f467,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142,f5859,f5861])).
fof(f5860,plain,(
  $false | ($spl6 | $spl20 | $spl29)),
  inference(global_subsumption,[],[f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f49,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1956,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f467,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142,f5859])).
fof(f5858,plain,(
  $false | ($spl1 | $spl6 | $spl20 | $spl29)),
  inference(global_subsumption,[],[f5857,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f49,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1956,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f467,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142])).
fof(f5857,plain,(
  ( ! [X59,X57,X58] : (s(X59,f(s(bool,'F_0'))) != s(X59,g(s(bool,'F_0'))) | s(fun(bool,X59),X57) = s(fun(bool,X59),X58) | s(bool,'T_0') = s(bool,sK3(X57,X58,X59,bool))) ) | ($spl1 | $spl6)),
  inference(forward_demodulation,[],[f2471,f84])).
fof(f2471,plain,(
  ( ! [X59,X57,X58] : (s(X59,f(s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))))) != s(X59,g(s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))))) | s(fun(bool,X59),X57) = s(fun(bool,X59),X58) | s(bool,'T_0') = s(bool,sK3(X57,X58,X59,bool))) ) | $spl1),
  inference(superposition,[],[f57,f2359])).
fof(f5856,plain,(
  $false | ($spl1 | $spl6 | $spl20 | $spl29)),
  inference(global_subsumption,[],[f5855,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f49,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1956,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f467,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142])).
fof(f5855,plain,(
  ( ! [X61,X62,X60] : (s(X62,happ(s(fun(bool,X62),X60),s(bool,'F_0'))) != s(X62,happ(s(fun(bool,X62),X61),s(bool,'F_0'))) | s(fun(bool,X62),X60) = s(fun(bool,X62),X61) | s(bool,'T_0') = s(bool,sK4(X60,X61,bool,X62))) ) | ($spl1 | $spl6)),
  inference(forward_demodulation,[],[f2472,f84])).
fof(f2472,plain,(
  ( ! [X61,X62,X60] : (s(X62,happ(s(fun(bool,X62),X60),s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))))) != s(X62,happ(s(fun(bool,X62),X61),s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))))) | s(fun(bool,X62),X60) = s(fun(bool,X62),X61) | s(bool,'T_0') = s(bool,sK4(X60,X61,bool,X62))) ) | $spl1),
  inference(superposition,[],[f59,f2359])).
fof(f5854,plain,(
  $false | ($spl1 | $spl6 | $spl20 | $spl29)),
  inference(global_subsumption,[],[f5853,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f49,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1956,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f467,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142])).
fof(f5853,plain,(
  ( ! [X64,X65,X63] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(X65,f(s(bool,'F_0'))) != s(X65,g(s(bool,'F_0'))) | s(fun(bool,X65),X63) = s(fun(bool,X65),X64)) ) | ($spl1 | $spl6)),
  inference(forward_demodulation,[],[f5852,f84])).
fof(f5852,plain,(
  ( ! [X64,X65,X63] : (s(X65,f(s(bool,'F_0'))) != s(X65,g(s(bool,'F_0'))) | s(fun(bool,X65),X63) = s(fun(bool,X65),X64) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,sK3(X63,X64,X65,bool))))) ) | ($spl1 | $spl6)),
  inference(forward_demodulation,[],[f2723,f84])).
fof(f5848,plain,(
  $false | ($spl1 | $spl6 | $spl20 | $spl22 | $spl29)),
  inference(global_subsumption,[],[f5847,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f49,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1956,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f467,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142])).
fof(f5847,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl1 | $spl6 | $spl22)),
  inference(forward_demodulation,[],[f2921,f84])).
fof(f2921,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,'const.combin.I_1'(s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1)))))) ) | ($spl1 | $spl22)),
  inference(condensation,[],[f2920])).
fof(f2920,plain,(
  ( ! [X24,X25] : (s(bool,'T_0') = s(bool,X25) | s(bool,X25) = s(bool,'const.combin.I_1'(s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))))) | s(bool,'T_0') = s(bool,X24)) ) | ($spl1 | $spl22)),
  inference(subsumption_resolution,[],[f2896,f50])).
fof(f2896,plain,(
  ( ! [X24,X25] : (s(bool,'F_0') != s(bool,X24) | s(bool,'T_0') = s(bool,X25) | s(bool,X25) = s(bool,'const.combin.I_1'(s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))))) | s(bool,'T_0') = s(bool,X24)) ) | ($spl1 | $spl22)),
  inference(superposition,[],[f2624,f2359])).
fof(f5846,plain,(
  $false | ($spl6 | $spl20 | $spl26 | $spl29)),
  inference(global_subsumption,[],[f5845,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f49,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1956,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f467,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142])).
fof(f5843,plain,(
  $false | ($spl1 | $spl6 | $spl20 | $spl29)),
  inference(global_subsumption,[],[f5842,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f49,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1956,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f467,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142])).
fof(f5841,plain,(
  $false | ($spl1 | $spl6 | $spl20 | $spl29)),
  inference(global_subsumption,[],[f5840,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f49,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1956,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f467,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142])).
fof(f5840,plain,(
  ( ! [X196,X194,X197,X195] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(X196,f(s(bool,X197))) != s(X196,g(s(bool,X197))) | s(fun(bool,X196),X194) = s(fun(bool,X196),X195) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X197)))) ) | ($spl1 | $spl6)),
  inference(forward_demodulation,[],[f4182,f84])).
fof(f5837,plain,(
  $false | ($spl1 | $spl6 | $spl20 | $spl29)),
  inference(global_subsumption,[],[f2537,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f49,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1956,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f467,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142])).
fof(f5836,plain,(
  $false | ($spl1 | $spl6 | $spl20 | $spl29)),
  inference(global_subsumption,[],[f2536,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f49,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1956,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f467,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142])).
fof(f5835,plain,(
  $false | ($spl1 | $spl6 | $spl20 | $spl29)),
  inference(global_subsumption,[],[f5834,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f49,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1956,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f467,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142])).
fof(f5834,plain,(
  ( ! [X9] : (~p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,X9)) ) | ($spl1 | $spl6)),
  inference(forward_demodulation,[],[f2497,f84])).
fof(f2497,plain,(
  ( ! [X9] : (~p(s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1)))) | s(bool,'T_0') = s(bool,X9)) ) | $spl1),
  inference(duplicate_literal_removal,[],[f2451])).
fof(f2451,plain,(
  ( ! [X9] : (~p(s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1)))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) ) | $spl1),
  inference(superposition,[],[f76,f2359])).
fof(f5830,plain,(
  $false | ($spl1 | $spl6 | $spl20 | $spl29)),
  inference(global_subsumption,[],[f5829,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f49,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1956,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f467,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142])).
fof(f5828,plain,(
  $false | ($spl6 | $spl16 | $spl20 | $spl29)),
  inference(global_subsumption,[],[f5827,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f49,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1956,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f467,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142])).
fof(f5827,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X0)))) ) | ($spl6 | $spl16)),
  inference(subsumption_resolution,[],[f5826,f84])).
fof(f5826,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X0)))) ) | $spl16),
  inference(subsumption_resolution,[],[f2275,f2245])).
fof(f2275,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl16),
  inference(superposition,[],[f2245,f50])).
fof(f5825,plain,(
  $false | ($spl6 | $spl16 | $spl20 | $spl29)),
  inference(global_subsumption,[],[f5824,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f49,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1956,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f467,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142])).
fof(f5824,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X0)))) ) | ($spl6 | $spl16)),
  inference(subsumption_resolution,[],[f5823,f84])).
fof(f5823,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X0)))) ) | $spl16),
  inference(subsumption_resolution,[],[f2241,f2245])).
fof(f2241,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl16),
  inference(forward_demodulation,[],[f2208,f51])).
fof(f2208,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl16),
  inference(superposition,[],[f1653,f50])).
fof(f5822,plain,(
  $false | ($spl6 | $spl16 | $spl20 | $spl29)),
  inference(global_subsumption,[],[f5821,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f49,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1956,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f467,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142])).
fof(f5821,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X0)))) ) | ($spl6 | $spl16)),
  inference(subsumption_resolution,[],[f5820,f84])).
fof(f5820,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X0)))) ) | $spl16),
  inference(subsumption_resolution,[],[f2003,f1699])).
fof(f1699,plain,(
  ( ! [X26] : (s(bool,'T_0') != s(bool,X26) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X26)))) ) | $spl16),
  inference(forward_demodulation,[],[f1654,f51])).
fof(f1654,plain,(
  ( ! [X26] : (s(bool,X26) != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X26)))) ) | $spl16),
  inference(superposition,[],[f1113,f68])).
fof(f2003,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl16),
  inference(superposition,[],[f1699,f50])).
fof(f5819,plain,(
  $false | ($spl6 | $spl16 | $spl20 | $spl29)),
  inference(global_subsumption,[],[f5818,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f49,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1956,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f467,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142])).
fof(f5818,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl6 | $spl16)),
  inference(subsumption_resolution,[],[f1691,f84])).
fof(f1691,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X0)) ) | $spl16),
  inference(forward_demodulation,[],[f1685,f51])).
fof(f1685,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl16),
  inference(duplicate_literal_removal,[],[f1636])).
fof(f1636,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl16),
  inference(superposition,[],[f1113,f50])).
fof(f5814,plain,(
  $false | ($spl1 | $spl6 | $spl20 | $spl26 | $spl29)),
  inference(global_subsumption,[],[f4310,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f49,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1956,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f467,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142])).
fof(f4310,plain,(
  ( ! [X43,X44,X42] : (s(bool,'T_0') != s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,X43) | s(bool,X43) = s(bool,'const.combin.I_1'(s(bool,X42))) | s(bool,X42) = s(bool,X44) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X44)))) ) | ($spl1 | $spl26)),
  inference(superposition,[],[f2131,f2659])).
fof(f5812,plain,(
  $false | ($spl1 | $spl6 | $spl20 | $spl26 | $spl29)),
  inference(global_subsumption,[],[f4275,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f49,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1956,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f467,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142])).
fof(f4275,plain,(
  ( ! [X43,X44,X42] : (s(bool,'T_0') != s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,X43) | s(bool,X43) = s(bool,'const.combin.I_1'(s(bool,X42))) | s(bool,X42) = s(bool,X44) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X44)))) ) | ($spl1 | $spl26)),
  inference(superposition,[],[f2131,f2659])).
fof(f5808,plain,(
  $false | ($spl1 | $spl6 | $spl20 | $spl22 | $spl29)),
  inference(global_subsumption,[],[f4218,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f49,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1956,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f467,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142])).
fof(f4218,plain,(
  ( ! [X83,X81,X82] : (s(bool,X81) != s(bool,X82) | s(bool,'T_0') = s(bool,X83) | s(bool,X83) = s(bool,'const.combin.I_1'(s(bool,X82))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X81)))) ) | ($spl1 | $spl22)),
  inference(superposition,[],[f2624,f2659])).
fof(f5806,plain,(
  $false | ($spl1 | $spl6 | $spl16 | $spl20 | $spl22 | $spl29)),
  inference(global_subsumption,[],[f5805,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f49,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1956,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f467,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142])).
fof(f5805,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X0))) | s(bool,X0) = s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | ($spl1 | $spl16 | $spl22)),
  inference(condensation,[],[f5804])).
fof(f5804,plain,(
  ( ! [X43,X44] : (s(bool,X44) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X43))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X44)))) ) | ($spl1 | $spl16 | $spl22)),
  inference(subsumption_resolution,[],[f4202,f2245])).
fof(f4202,plain,(
  ( ! [X43,X44] : (s(bool,X44) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X43) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X43))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X44)))) ) | ($spl1 | $spl22)),
  inference(superposition,[],[f2907,f2659])).
fof(f5797,plain,(
  $false | ($spl1 | $spl6 | $spl20 | $spl22 | $spl29)),
  inference(global_subsumption,[],[f4144,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f49,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1956,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f467,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142])).
fof(f4144,plain,(
  ( ! [X83,X81,X82] : (s(bool,X81) != s(bool,X82) | s(bool,'T_0') = s(bool,X83) | s(bool,X83) = s(bool,'const.combin.I_1'(s(bool,X82))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X81))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | ($spl1 | $spl22)),
  inference(superposition,[],[f2624,f2659])).
fof(f5795,plain,(
  $false | ($spl1 | $spl6 | $spl16 | $spl20 | $spl22 | $spl29)),
  inference(global_subsumption,[],[f5794,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f49,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1956,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f467,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142])).
fof(f5794,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X0))) | s(bool,X0) = s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | ($spl1 | $spl16 | $spl22)),
  inference(condensation,[],[f5793])).
fof(f5793,plain,(
  ( ! [X43,X44] : (s(bool,X44) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X44))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X43)))) ) | ($spl1 | $spl16 | $spl22)),
  inference(subsumption_resolution,[],[f4128,f2245])).
fof(f4128,plain,(
  ( ! [X43,X44] : (s(bool,X44) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X43) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X44))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X43)))) ) | ($spl1 | $spl22)),
  inference(superposition,[],[f2907,f2659])).
fof(f5782,plain,(
  $false | ($spl1 | $spl6 | $spl16 | $spl20 | $spl22 | $spl29)),
  inference(global_subsumption,[],[f5781,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f49,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1956,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f467,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142])).
fof(f5781,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X0))) | s(bool,X0) = s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | ($spl1 | $spl16 | $spl22)),
  inference(condensation,[],[f5780])).
fof(f5780,plain,(
  ( ! [X57,X58] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X57))) | s(bool,X58) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X58)))) ) | ($spl1 | $spl16 | $spl22)),
  inference(subsumption_resolution,[],[f4091,f2245])).
fof(f4091,plain,(
  ( ! [X57,X58] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X57))) | s(bool,X58) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X58))) | s(bool,'T_0') = s(bool,X57)) ) | ($spl1 | $spl22)),
  inference(superposition,[],[f2659,f2907])).
fof(f5770,plain,(
  $false | ($spl1 | $spl6 | $spl16 | $spl20 | $spl22 | $spl29)),
  inference(global_subsumption,[],[f5769,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f49,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1956,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f467,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142])).
fof(f5769,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X0))) | s(bool,X0) = s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | ($spl1 | $spl16 | $spl22)),
  inference(condensation,[],[f5768])).
fof(f5768,plain,(
  ( ! [X57,X58] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X57))) | s(bool,X58) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X58)))) ) | ($spl1 | $spl16 | $spl22)),
  inference(subsumption_resolution,[],[f4048,f2245])).
fof(f4048,plain,(
  ( ! [X57,X58] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X57))) | s(bool,X58) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X58))) | s(bool,'T_0') = s(bool,X57)) ) | ($spl1 | $spl22)),
  inference(superposition,[],[f2659,f2907])).
fof(f5752,plain,(
  $false | ($spl1 | $spl6 | $spl16 | $spl20 | $spl22 | $spl29)),
  inference(global_subsumption,[],[f5751,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f49,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1956,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f467,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142])).
fof(f5751,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X0))) | s(bool,X0) = s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | ($spl1 | $spl16 | $spl22)),
  inference(condensation,[],[f5750])).
fof(f5750,plain,(
  ( ! [X24,X25] : (s(bool,X25) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X24))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X25)))) ) | ($spl1 | $spl16 | $spl22)),
  inference(subsumption_resolution,[],[f4002,f2245])).
fof(f4002,plain,(
  ( ! [X24,X25] : (s(bool,X25) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X24))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X25))) | s(bool,'T_0') = s(bool,X24)) ) | ($spl1 | $spl22)),
  inference(superposition,[],[f2659,f2907])).
fof(f5736,plain,(
  $false | ($spl1 | $spl6 | $spl16 | $spl20 | $spl22 | $spl29)),
  inference(global_subsumption,[],[f5735,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f49,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1956,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f467,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142])).
fof(f5735,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X0))) | s(bool,X0) = s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | ($spl1 | $spl16 | $spl22)),
  inference(condensation,[],[f5734])).
fof(f5734,plain,(
  ( ! [X24,X25] : (s(bool,X25) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X25))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X24)))) ) | ($spl1 | $spl16 | $spl22)),
  inference(subsumption_resolution,[],[f3971,f2245])).
fof(f3971,plain,(
  ( ! [X24,X25] : (s(bool,X25) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X25))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X24))) | s(bool,'T_0') = s(bool,X24)) ) | ($spl1 | $spl22)),
  inference(superposition,[],[f2659,f2907])).
fof(f5732,plain,(
  $false | ($spl1 | $spl6 | $spl20 | $spl26 | $spl29)),
  inference(global_subsumption,[],[f5731,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f49,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1956,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f467,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142])).
fof(f5731,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') != s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,X35) | s(bool,X35) = s(bool,'const.combin.I_1'(s(bool,X34)))) ) | ($spl1 | $spl26)),
  inference(subsumption_resolution,[],[f3699,f3701])).
fof(f3699,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') != s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,X35) | s(bool,X35) = s(bool,'const.combin.I_1'(s(bool,X34))) | s(bool,X34) = s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1)))) ) | ($spl1 | $spl26)),
  inference(superposition,[],[f2131,f2434])).
fof(f5730,plain,(
  $false | ($spl6 | $spl20 | $spl26 | $spl29)),
  inference(global_subsumption,[],[f3727,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f49,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1956,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f467,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142])).
fof(f5729,plain,(
  $false | ($spl6 | $spl20 | $spl26 | $spl29)),
  inference(global_subsumption,[],[f3730,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f49,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1956,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f467,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142])).
fof(f5728,plain,(
  $false | ($spl6 | $spl20 | $spl29)),
  inference(global_subsumption,[],[f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f49,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1956,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f467,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930,f84,f5705,f2142])).
fof(f5727,plain,(
  $false | ($spl6 | $spl13 | $spl20 | $spl22)),
  inference(global_subsumption,[],[f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f49,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1956,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1556,f1324,f473,f467,f2120,f418,f2594,f2624,f2625,f2626,f2627,f2881,f2907,f3070,f3215,f3266,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3809,f3870,f3871,f4824,f4825,f4900,f4852,f4929,f4930,f5533,f5534,f5553,f5661,f2634,f2637,f84,f5705,f2943,f2998,f3023])).
fof(f3266,plain,(
  ( ! [X26,X24,X23,X25,X22] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'const.combin.o_3'(s(fun(X22,bool),X23),s(fun(X24,X22),X25),s(X24,X26))))) | s(bool,'T_0') = s(bool,happ(s(fun(X22,bool),X23),s(X22,happ(s(fun(X24,X22),X25),s(X24,X26)))))) ) | ($spl13 | $spl22)),
  inference(superposition,[],[f3215,f60])).
fof(f3215,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl13 | $spl22)),
  inference(superposition,[],[f3070,f50])).
fof(f3070,plain,(
  s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | ($spl13 | $spl22)),
  inference(subsumption_resolution,[],[f2983,f418])).
fof(f2983,plain,(
  s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | ($spl13 | $spl22)),
  inference(superposition,[],[f418,f2907])).
fof(f5726,plain,(
  $false | ($spl6 | $spl13 | $spl20 | $spl22)),
  inference(global_subsumption,[],[f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f49,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1956,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1556,f1324,f473,f467,f2120,f418,f2594,f2624,f2625,f2626,f2627,f2881,f2907,f3070,f3215,f3266,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3809,f3870,f3871,f4824,f4825,f4900,f4852,f4929,f4930,f5533,f5534,f5553,f5661,f2634,f2637,f84,f5705,f2943,f2998])).
fof(f5724,plain,(
  $false | ($spl6 | $spl13 | $spl20 | $spl22)),
  inference(global_subsumption,[],[f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f49,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1956,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1556,f1324,f473,f467,f2120,f418,f2594,f2624,f2625,f2626,f2627,f2881,f2907,f3070,f3215,f3266,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3809,f3870,f3871,f4824,f4825,f4900,f4852,f4929,f4930,f5533,f5534,f5553,f5661,f2634,f2637,f84,f5705,f2943])).
fof(f5721,plain,(
  $false | ($spl1 | $spl6 | $spl13 | $spl16 | $spl20 | $spl22)),
  inference(global_subsumption,[],[f2770,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f49,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1956,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1556,f1324,f473,f467,f2120,f418,f2594,f2624,f2625,f2626,f2627,f2881,f2907,f3070,f3215,f3266,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3809,f3870,f3871,f4824,f4825,f4900,f4852,f4929,f4930,f5533,f5534,f5553,f5661,f2634,f2637,f84,f5705])).
fof(f2770,plain,(
  ( ! [X16] : (s(bool,'T_0') != s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X16)))) ) | ($spl1 | $spl16)),
  inference(duplicate_literal_removal,[],[f2705])).
fof(f2705,plain,(
  ( ! [X16] : (s(bool,'T_0') != s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X16))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X16)))) ) | ($spl1 | $spl16)),
  inference(superposition,[],[f2245,f2434])).
fof(f5720,plain,(
  $false | ($spl1 | $spl6 | $spl13 | $spl16 | $spl20 | $spl22)),
  inference(global_subsumption,[],[f5719,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f49,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1956,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1556,f1324,f473,f467,f2120,f418,f2594,f2624,f2625,f2626,f2627,f2881,f2907,f3070,f3215,f3266,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3809,f3870,f3871,f4824,f4825,f4900,f4852,f4929,f4930,f5533,f5534,f5553,f5661,f2634,f2637,f84,f5705])).
fof(f5719,plain,(
  ( ! [X13] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X13)))) ) | ($spl1 | $spl6 | $spl16)),
  inference(subsumption_resolution,[],[f5718,f1699])).
fof(f5718,plain,(
  ( ! [X13] : (s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X13)))) ) | ($spl1 | $spl6 | $spl16)),
  inference(forward_demodulation,[],[f5717,f51])).
fof(f5717,plain,(
  ( ! [X13] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X13))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X13)))) ) | ($spl1 | $spl6 | $spl16)),
  inference(subsumption_resolution,[],[f5716,f84])).
fof(f5716,plain,(
  ( ! [X13] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X13))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X13)))) ) | ($spl1 | $spl6 | $spl16)),
  inference(forward_demodulation,[],[f2703,f84])).
fof(f2703,plain,(
  ( ! [X13] : (s(bool,'T_0') != s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X13))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X13)))) ) | ($spl1 | $spl16)),
  inference(superposition,[],[f1699,f2434])).
fof(f5713,plain,(
  $false | ($spl1 | $spl6 | $spl13 | $spl20 | $spl22)),
  inference(global_subsumption,[],[f2490,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f49,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1956,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1556,f1324,f473,f467,f2120,f418,f2594,f2624,f2625,f2626,f2627,f2881,f2907,f3070,f3215,f3266,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3809,f3870,f3871,f4824,f4825,f4900,f4852,f4929,f4930,f5533,f5534,f5553,f5661,f2634,f2637,f84,f5705])).
fof(f2490,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,X0) = s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1)))) ) | $spl1),
  inference(equality_factoring,[],[f2359])).
fof(f5712,plain,(
  $false | ($spl1 | $spl6 | $spl13 | $spl16 | $spl20 | $spl22)),
  inference(global_subsumption,[],[f5711,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f49,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1956,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1556,f1324,f473,f467,f2120,f418,f2594,f2624,f2625,f2626,f2627,f2881,f2907,f3070,f3215,f3266,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3809,f3870,f3871,f4824,f4825,f4900,f4852,f4929,f4930,f5533,f5534,f5553,f5661,f2634,f2637,f84,f5705])).
fof(f5711,plain,(
  ( ! [X13] : (s(bool,'T_0') != s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X13)))) ) | ($spl1 | $spl16)),
  inference(subsumption_resolution,[],[f2455,f2245])).
fof(f2455,plain,(
  ( ! [X13] : (s(bool,'T_0') != s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X13))) | s(bool,'T_0') = s(bool,X13)) ) | ($spl1 | $spl16)),
  inference(superposition,[],[f2245,f2359])).
fof(f5710,plain,(
  $false | ($spl1 | $spl6 | $spl13 | $spl16 | $spl20 | $spl22)),
  inference(global_subsumption,[],[f5709,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f49,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1956,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1556,f1324,f473,f467,f2120,f418,f2594,f2624,f2625,f2626,f2627,f2881,f2907,f3070,f3215,f3266,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3809,f3870,f3871,f4824,f4825,f4900,f4852,f4929,f4930,f5533,f5534,f5553,f5661,f2634,f2637,f84,f5705])).
fof(f5709,plain,(
  ( ! [X12] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X12)))) ) | ($spl1 | $spl6 | $spl16)),
  inference(subsumption_resolution,[],[f5708,f84])).
fof(f5708,plain,(
  ( ! [X12] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X12)))) ) | ($spl1 | $spl6 | $spl16)),
  inference(forward_demodulation,[],[f5707,f84])).
fof(f5707,plain,(
  ( ! [X12] : (s(bool,'T_0') != s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X12)))) ) | ($spl1 | $spl16)),
  inference(subsumption_resolution,[],[f2454,f1699])).
fof(f2454,plain,(
  ( ! [X12] : (s(bool,'T_0') != s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X12))) | s(bool,'T_0') = s(bool,X12)) ) | ($spl1 | $spl16)),
  inference(superposition,[],[f1699,f2359])).
fof(f5706,plain,(
  $false | ($spl6 | $spl13 | $spl20 | $spl22)),
  inference(global_subsumption,[],[f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f49,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1956,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1556,f1324,f473,f467,f2120,f418,f2594,f2624,f2625,f2626,f2627,f2881,f2907,f3070,f3215,f3266,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3809,f3870,f3871,f4824,f4825,f4900,f4852,f4929,f4930,f5533,f5534,f5553,f5661,f2634,f2637,f84,f5705])).
fof(f5702,plain,(
  $false | ($spl1 | $spl10 | $spl16 | $spl29)),
  inference(subsumption_resolution,[],[f2142,f5665])).
fof(f5665,plain,(
  ( ! [X10] : (s(bool,'T_0') = s(bool,X10)) ) | ($spl1 | $spl10 | $spl16)),
  inference(subsumption_resolution,[],[f2496,f5565])).
fof(f5565,plain,(
  ( ! [X14] : (s(bool,X14) = s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1)))) ) | ($spl1 | $spl10 | $spl16)),
  inference(global_subsumption,[],[f2833,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f416,f1259,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f49,f1756,f1757,f1758,f1759,f1766,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1956,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1324,f473,f467,f2594,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3870,f3871,f4824,f4825,f4900,f4929,f4930])).
fof(f1259,plain,(
  ( ! [X6,X4,X5,X3] : (s(fun(X3,bool),X4) = s(fun(X3,bool),X5) | s(bool,X6) = s(bool,g(s(X3,sK3(X4,X5,bool,X3)))) | s(bool,'T_0') = s(bool,X6)) ) | $spl10),
  inference(subsumption_resolution,[],[f1247,f416])).
fof(f1247,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,'T_0') != s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(fun(X3,bool),X4) = s(fun(X3,bool),X5) | s(bool,X6) = s(bool,g(s(X3,sK3(X4,X5,bool,X3)))) | s(bool,'T_0') = s(bool,X6)) ) | $spl10),
  inference(backward_demodulation,[],[f416,f469])).
fof(f2833,plain,(
  ( ! [X14] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X14))) | s(bool,X14) = s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1)))) ) | ($spl1 | $spl16)),
  inference(forward_demodulation,[],[f2764,f51])).
fof(f2764,plain,(
  ( ! [X14] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'const.combin.I_1'(s(bool,X14))))) | s(bool,X14) = s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1)))) ) | ($spl1 | $spl16)),
  inference(trivial_inequality_removal,[],[f2754])).
fof(f2754,plain,(
  ( ! [X14] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'const.combin.I_1'(s(bool,X14))))) | s(bool,X14) = s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1)))) ) | ($spl1 | $spl16)),
  inference(superposition,[],[f1699,f2434])).
fof(f5701,plain,(
  $false | ($spl1 | $spl10 | $spl16)),
  inference(subsumption_resolution,[],[f5700,f47])).
fof(f5700,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl10 | $spl16)),
  inference(forward_demodulation,[],[f64,f5665])).
fof(f5699,plain,(
  $false | ($spl1 | $spl10 | $spl16)),
  inference(subsumption_resolution,[],[f5689,f47])).
fof(f5689,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl10 | $spl16)),
  inference(backward_demodulation,[],[f5665,f5659])).
fof(f5698,plain,(
  $false | ($spl1 | $spl10 | $spl16)),
  inference(subsumption_resolution,[],[f5696,f68])).
fof(f5696,plain,(
  s(sK0,sK4('const.combin.I_0','const.combin.o_2'(s(fun(sK0,sK0),sK1),s(fun(sK0,sK0),sK1)),sK0,sK0)) != s(sK0,happ(s(fun(sK0,sK0),'const.combin.I_0'),s(sK0,sK4('const.combin.I_0','const.combin.o_2'(s(fun(sK0,sK0),sK1),s(fun(sK0,sK0),sK1)),sK0,sK0)))) | ($spl1 | $spl10 | $spl16)),
  inference(backward_demodulation,[],[f5694,f2536])).
fof(f5694,plain,(
  ( ! [X0,X1] : (s(fun(X0,X0),'const.combin.I_0') = s(fun(X0,X0),'const.combin.o_2'(s(fun(X0,X0),X1),s(fun(X0,X0),X1)))) ) | ($spl1 | $spl10 | $spl16)),
  inference(subsumption_resolution,[],[f5681,f47])).
fof(f5681,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s(fun(X0,X0),'const.combin.I_0') = s(fun(X0,X0),'const.combin.o_2'(s(fun(X0,X0),X1),s(fun(X0,X0),X1)))) ) | ($spl1 | $spl10 | $spl16)),
  inference(backward_demodulation,[],[f5665,f55])).
fof(f5697,plain,(
  $false | ($spl1 | $spl10 | $spl16)),
  inference(subsumption_resolution,[],[f5695,f68])).
fof(f5695,plain,(
  s(sK0,sK4('const.combin.o_2'(s(fun(sK0,sK0),sK1),s(fun(sK0,sK0),sK1)),'const.combin.I_0',sK0,sK0)) != s(sK0,happ(s(fun(sK0,sK0),'const.combin.I_0'),s(sK0,sK4('const.combin.o_2'(s(fun(sK0,sK0),sK1),s(fun(sK0,sK0),sK1)),'const.combin.I_0',sK0,sK0)))) | ($spl1 | $spl10 | $spl16)),
  inference(backward_demodulation,[],[f5694,f2537])).
fof(f5693,plain,(
  $false | ($spl1 | $spl10 | $spl16)),
  inference(subsumption_resolution,[],[f5666,f47])).
fof(f5666,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl10 | $spl16)),
  inference(backward_demodulation,[],[f5665,f49])).
fof(f86,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl9),
  inference(cnf_transformation,[],[f86_D])).
fof(f86_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f2637,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') != s(bool,'F_0')) ) | $spl22),
  inference(condensation,[],[f2636])).
fof(f2636,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X1)) ) | $spl22),
  inference(duplicate_literal_removal,[],[f2635])).
fof(f2635,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | $spl22),
  inference(condensation,[],[f2615])).
fof(f2615,plain,(
  ( ! [X4,X2,X3] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X3) | s(bool,X2) = s(bool,X3) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) ) | $spl22),
  inference(superposition,[],[f2120,f73])).
fof(f2634,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') != s(bool,'F_0')) ) | $spl22),
  inference(condensation,[],[f2633])).
fof(f2633,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X1)) ) | $spl22),
  inference(duplicate_literal_removal,[],[f2632])).
fof(f2632,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | $spl22),
  inference(condensation,[],[f2618])).
fof(f2618,plain,(
  ( ! [X12,X13,X11] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X12) | s(bool,X11) = s(bool,X12) | s(bool,X11) = s(bool,X13) | s(bool,'T_0') = s(bool,X13)) ) | $spl22),
  inference(superposition,[],[f2120,f73])).
fof(f5525,plain,(
  ( ! [X54,X53] : (s(bool,X54) = s(bool,'const.combin.I_1'(s(bool,X53))) | s(bool,'T_0') = s(bool,X54)) ) | $spl32),
  inference(cnf_transformation,[],[f5525_D])).
fof(f5525_D,plain,(
  ( ! [X54,X53] : (s(bool,X54) = s(bool,'const.combin.I_1'(s(bool,X53))) | s(bool,'T_0') = s(bool,X54)) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f5661,plain,(
  ( ! [X24,X23] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X24) | s(bool,X24) = s(bool,'const.combin.I_1'(s(bool,X23)))) ) | $spl22),
  inference(subsumption_resolution,[],[f2623,f2624])).
fof(f2623,plain,(
  ( ! [X24,X23] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X24) | s(bool,X24) = s(bool,'const.combin.I_1'(s(bool,X23))) | s(bool,'F_0') = s(bool,X23)) ) | $spl22),
  inference(superposition,[],[f2120,f75])).
fof(f1260,plain,(
  ( ! [X23,X21,X22] : (s(fun(X21,bool),X22) = s(fun(X21,bool),X23)) ) | $spl18),
  inference(cnf_transformation,[],[f1260_D])).
fof(f1260_D,plain,(
  ( ! [X23,X21,X22] : (s(fun(X21,bool),X22) = s(fun(X21,bool),X23)) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f1262,plain,(
  s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))))) | $spl20),
  inference(cnf_transformation,[],[f1262_D])).
fof(f1262_D,plain,(
  s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))))) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f5657,plain,(
  ( ! [X39,X37,X38] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))))) | s(fun(X37,bool),X38) = s(fun(X37,bool),X39)) ) | ($spl1 | $spl10 | $spl16)),
  inference(forward_demodulation,[],[f5656,f5565])).
fof(f5656,plain,(
  ( ! [X39,X37,X38] : (s(fun(X37,bool),X38) = s(fun(X37,bool),X39) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,g(s(X37,sK3(X38,X39,bool,X37))))))) ) | ($spl1 | $spl10 | $spl16)),
  inference(subsumption_resolution,[],[f2715,f5565])).
fof(f5653,plain,(
  ( ! [X52,X53,X51] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))))) | s(fun(X51,bool),X52) = s(fun(X51,bool),X53)) ) | ($spl1 | $spl10 | $spl16)),
  inference(forward_demodulation,[],[f5652,f5565])).
fof(f5652,plain,(
  ( ! [X52,X53,X51] : (s(fun(X51,bool),X52) = s(fun(X51,bool),X53) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,happ(s(fun(X51,bool),X52),s(X51,sK4(X53,X52,X51,bool))))))) ) | ($spl1 | $spl10 | $spl16)),
  inference(subsumption_resolution,[],[f2719,f5565])).
fof(f5651,plain,(
  ( ! [X54,X56,X55] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))))) | s(fun(X54,bool),X55) = s(fun(X54,bool),X56)) ) | ($spl1 | $spl10 | $spl16)),
  inference(forward_demodulation,[],[f5650,f5565])).
fof(f5650,plain,(
  ( ! [X54,X56,X55] : (s(fun(X54,bool),X55) = s(fun(X54,bool),X56) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,happ(s(fun(X54,bool),X55),s(X54,sK4(X55,X56,X54,bool))))))) ) | ($spl1 | $spl10 | $spl16)),
  inference(subsumption_resolution,[],[f2720,f5565])).
fof(f5628,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | ($spl1 | $spl10 | $spl16)),
  inference(forward_demodulation,[],[f5627,f5565])).
fof(f5627,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,g(s(X0,sK3(X1,X2,bool,X0)))))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | ($spl1 | $spl10 | $spl16)),
  inference(condensation,[],[f5626])).
fof(f5626,plain,(
  ( ! [X156,X154,X155,X153] : (s(fun(X153,bool),X154) = s(fun(X153,bool),X155) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X156))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,g(s(X153,sK3(X154,X155,bool,X153))))))) ) | ($spl1 | $spl10 | $spl16)),
  inference(subsumption_resolution,[],[f5566,f2434])).
fof(f5566,plain,(
  ( ! [X156,X154,X155,X153] : (s(bool,X156) != s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(fun(X153,bool),X154) = s(fun(X153,bool),X155) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X156))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,g(s(X153,sK3(X154,X155,bool,X153))))))) ) | ($spl1 | $spl10 | $spl16)),
  inference(backward_demodulation,[],[f5565,f4172])).
fof(f5564,plain,(
  ( ! [X15,X16] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X16) | s(bool,X16) = s(bool,'const.combin.I_1'(s(bool,X15)))) ) | ($spl1 | $spl22 | $spl26)),
  inference(subsumption_resolution,[],[f2755,f3701])).
fof(f2755,plain,(
  ( ! [X15,X16] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X16) | s(bool,X16) = s(bool,'const.combin.I_1'(s(bool,X15))) | s(bool,X15) = s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1)))) ) | ($spl1 | $spl22)),
  inference(superposition,[],[f2120,f2434])).
fof(f5563,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') != s(bool,'F_0')) ) | ($spl16 | $spl22)),
  inference(duplicate_literal_removal,[],[f5562])).
fof(f5562,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X0)) ) | ($spl16 | $spl22)),
  inference(condensation,[],[f5561])).
fof(f5561,plain,(
  ( ! [X4,X2,X3] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X3) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) ) | ($spl16 | $spl22)),
  inference(subsumption_resolution,[],[f2888,f1656])).
fof(f1656,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) != s(bool,'const.combin.I_1'(s(bool,X0))) | s(bool,'T_0') = s(bool,X1) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl16),
  inference(superposition,[],[f1113,f73])).
fof(f5560,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') != s(bool,'F_0')) ) | ($spl16 | $spl22)),
  inference(duplicate_literal_removal,[],[f5559])).
fof(f5559,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X0)) ) | ($spl16 | $spl22)),
  inference(condensation,[],[f5558])).
fof(f5558,plain,(
  ( ! [X12,X13,X11] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X12) | s(bool,X11) = s(bool,X13) | s(bool,'T_0') = s(bool,X13)) ) | ($spl16 | $spl22)),
  inference(subsumption_resolution,[],[f2891,f1656])).
fof(f2891,plain,(
  ( ! [X12,X13,X11] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X12) | s(bool,X12) = s(bool,'const.combin.I_1'(s(bool,X11))) | s(bool,X11) = s(bool,X13) | s(bool,'T_0') = s(bool,X13)) ) | $spl22),
  inference(superposition,[],[f2624,f73])).
fof(f5557,plain,(
  ( ! [X28] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X28)) ) | ($spl1 | $spl16 | $spl22 | $spl26)),
  inference(global_subsumption,[],[f5556,f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f49,f1113,f1699,f1700,f1656,f1756,f1757,f1758,f1759,f1766,f1814,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1956,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1995,f1653,f1324,f473,f467,f2245,f2220,f2269,f2120,f2594,f2624,f2625,f2626,f2627,f2881,f2907,f3065,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3848,f3809,f3870,f3871,f4824,f4825,f4900,f4902,f4852,f4929,f4930,f5533,f5534,f3064,f3062,f3059,f5553,f5554])).
fof(f2269,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X1))) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl16),
  inference(superposition,[],[f2245,f73])).
fof(f2220,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) != s(bool,'const.combin.I_1'(s(bool,X0))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X1))) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl16),
  inference(superposition,[],[f1653,f73])).
fof(f1995,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X1))) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl16),
  inference(superposition,[],[f1699,f73])).
fof(f5556,plain,(
  ( ! [X28,X27] : (s(bool,X28) = s(bool,'const.combin.I_1'(s(bool,X27))) | s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X28)) ) | ($spl1 | $spl22 | $spl26)),
  inference(subsumption_resolution,[],[f5555,f3701])).
fof(f5555,plain,(
  ( ! [X28,X27] : (s(bool,X28) = s(bool,'const.combin.I_1'(s(bool,X27))) | s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X28) | s(bool,X27) = s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1)))) ) | ($spl1 | $spl22)),
  inference(forward_demodulation,[],[f2898,f51])).
fof(f2898,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X28) | s(bool,X28) = s(bool,'const.combin.I_1'(s(bool,'const.combin.I_1'(s(bool,X27))))) | s(bool,X27) = s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1)))) ) | ($spl1 | $spl22)),
  inference(superposition,[],[f2624,f2434])).
fof(f5554,plain,(
  ( ! [X30] : (s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') != s(bool,'F_0')) ) | ($spl16 | $spl22)),
  inference(global_subsumption,[],[f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f463,f465,f1293,f1294,f1305,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f49,f1113,f1699,f1700,f1656,f1756,f1757,f1758,f1759,f1766,f1814,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1983,f1984,f1985,f1986,f1956,f1957,f1958,f1965,f1966,f1967,f1968,f1969,f1995,f1653,f1324,f473,f467,f2245,f2220,f2269,f2120,f2594,f2624,f2625,f2626,f2627,f2881,f2907,f3065,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3848,f3809,f3870,f3871,f4824,f4825,f4900,f4902,f4852,f4929,f4930,f5533,f5534,f3064,f3062,f3059,f5553])).
fof(f5538,plain,(
  s(bool,'F_0') = s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | $spl36),
  inference(cnf_transformation,[],[f5538_D])).
fof(f5538_D,plain,(
  s(bool,'F_0') = s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f5551,plain,(
  ( ! [X9] : (s(bool,'F_0') = s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,X9)) ) | ($spl1 | $spl22)),
  inference(forward_demodulation,[],[f3043,f51])).
fof(f3043,plain,(
  ( ! [X9] : (s(bool,'const.combin.I_1'(s(bool,'F_0'))) = s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,X9)) ) | ($spl1 | $spl22)),
  inference(duplicate_literal_removal,[],[f2935])).
fof(f2935,plain,(
  ( ! [X9] : (s(bool,'const.combin.I_1'(s(bool,'F_0'))) = s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) ) | ($spl1 | $spl22)),
  inference(superposition,[],[f2907,f2359])).
fof(f1112,plain,(
  ( ! [X3] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X3)))) ) | $spl14),
  inference(cnf_transformation,[],[f1112_D])).
fof(f1112_D,plain,(
  ( ! [X3] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X3)))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f5550,plain,(
  ( ! [X10] : (s(bool,'F_0') = s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X10)))) ) | ($spl1 | $spl16 | $spl22)),
  inference(forward_demodulation,[],[f3051,f51])).
fof(f3051,plain,(
  ( ! [X10] : (s(bool,'const.combin.I_1'(s(bool,'F_0'))) = s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X10)))) ) | ($spl1 | $spl16 | $spl22)),
  inference(subsumption_resolution,[],[f2936,f2245])).
fof(f2936,plain,(
  ( ! [X10] : (s(bool,'const.combin.I_1'(s(bool,'F_0'))) = s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X10)))) ) | ($spl1 | $spl22)),
  inference(superposition,[],[f2907,f2434])).
fof(f3042,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,X2) | s(bool,'F_0') = s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1)))) ) | ($spl1 | $spl22)),
  inference(duplicate_literal_removal,[],[f2954])).
fof(f2954,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,'F_0') = s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1)))) ) | ($spl1 | $spl22)),
  inference(superposition,[],[f2907,f2434])).
fof(f3059,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X11)) ) | ($spl16 | $spl22)),
  inference(forward_demodulation,[],[f3058,f51])).
fof(f3058,plain,(
  ( ! [X11] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X11)) ) | ($spl16 | $spl22)),
  inference(forward_demodulation,[],[f3031,f51])).
fof(f3031,plain,(
  ( ! [X11] : (s(bool,'const.combin.I_1'(s(bool,'T_0'))) != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X11)) ) | ($spl16 | $spl22)),
  inference(duplicate_literal_removal,[],[f2970])).
fof(f2970,plain,(
  ( ! [X11] : (s(bool,'const.combin.I_1'(s(bool,'T_0'))) != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X11)) ) | ($spl16 | $spl22)),
  inference(superposition,[],[f1113,f2907])).
fof(f3062,plain,(
  ( ! [X12] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X12)))) ) | ($spl16 | $spl22)),
  inference(forward_demodulation,[],[f3061,f51])).
fof(f3061,plain,(
  ( ! [X12] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X12)))) ) | ($spl16 | $spl22)),
  inference(forward_demodulation,[],[f3060,f51])).
fof(f3060,plain,(
  ( ! [X12] : (s(bool,'const.combin.I_1'(s(bool,'T_0'))) != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X12)))) ) | ($spl16 | $spl22)),
  inference(subsumption_resolution,[],[f2971,f2245])).
fof(f2971,plain,(
  ( ! [X12] : (s(bool,'const.combin.I_1'(s(bool,'T_0'))) != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X12))) | s(bool,'T_0') = s(bool,X12)) ) | ($spl16 | $spl22)),
  inference(superposition,[],[f1653,f2907])).
fof(f3064,plain,(
  ( ! [X13] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X13)))) ) | ($spl16 | $spl22)),
  inference(forward_demodulation,[],[f3063,f51])).
fof(f3063,plain,(
  ( ! [X13] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X13)))) ) | ($spl16 | $spl22)),
  inference(subsumption_resolution,[],[f2972,f1699])).
fof(f2972,plain,(
  ( ! [X13] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X13))) | s(bool,'T_0') = s(bool,X13)) ) | ($spl16 | $spl22)),
  inference(superposition,[],[f1699,f2907])).
fof(f5549,plain,(
  ( ! [X17] : (s(bool,'F_0') = s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,X17)) ) | ($spl1 | $spl22)),
  inference(forward_demodulation,[],[f3030,f51])).
fof(f3030,plain,(
  ( ! [X17] : (s(bool,'const.combin.I_1'(s(bool,'F_0'))) = s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,X17)) ) | ($spl1 | $spl22)),
  inference(duplicate_literal_removal,[],[f2975])).
fof(f2975,plain,(
  ( ! [X17] : (s(bool,'const.combin.I_1'(s(bool,'F_0'))) = s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,X17) | s(bool,'T_0') = s(bool,X17)) ) | ($spl1 | $spl22)),
  inference(superposition,[],[f2359,f2907])).
fof(f5548,plain,(
  ( ! [X18] : (s(bool,'F_0') = s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X18)))) ) | ($spl1 | $spl16 | $spl22)),
  inference(forward_demodulation,[],[f3067,f51])).
fof(f3067,plain,(
  ( ! [X18] : (s(bool,'const.combin.I_1'(s(bool,'F_0'))) = s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X18)))) ) | ($spl1 | $spl16 | $spl22)),
  inference(subsumption_resolution,[],[f2976,f2245])).
fof(f2976,plain,(
  ( ! [X18] : (s(bool,'const.combin.I_1'(s(bool,'F_0'))) = s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X18))) | s(bool,'T_0') = s(bool,X18)) ) | ($spl1 | $spl22)),
  inference(superposition,[],[f2434,f2907])).
fof(f3029,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1)))) ) | ($spl1 | $spl22)),
  inference(duplicate_literal_removal,[],[f3001])).
fof(f3001,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl1 | $spl22)),
  inference(superposition,[],[f2434,f2907])).
fof(f5546,plain,(
  ( ! [X21] : (s(bool,'F_0') = s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X21)))) ) | ($spl1 | $spl16 | $spl22)),
  inference(forward_demodulation,[],[f3080,f51])).
fof(f3080,plain,(
  ( ! [X21] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X21))) | s(bool,'const.combin.I_1'(s(bool,'F_0'))) = s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1)))) ) | ($spl1 | $spl16 | $spl22)),
  inference(subsumption_resolution,[],[f3015,f2245])).
fof(f3015,plain,(
  ( ! [X21] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X21))) | s(bool,'const.combin.I_1'(s(bool,'F_0'))) = s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,X21)) ) | ($spl1 | $spl22)),
  inference(superposition,[],[f2434,f2907])).
fof(f5541,plain,(
  ( ! [X17] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X17))) | s(bool,'F_0') = s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1)))) ) | ($spl1 | $spl16)),
  inference(subsumption_resolution,[],[f5540,f1699])).
fof(f5540,plain,(
  ( ! [X17] : (s(bool,'T_0') = s(bool,X17) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X17))) | s(bool,'F_0') = s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1)))) ) | $spl1),
  inference(forward_demodulation,[],[f3313,f51])).
fof(f3313,plain,(
  ( ! [X17] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X17))) | s(bool,'F_0') = s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X17)))) ) | $spl1),
  inference(superposition,[],[f96,f2434])).
fof(f5536,plain,(
  ( ! [X18] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X18))) | s(bool,'F_0') = s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1)))) ) | ($spl1 | $spl16)),
  inference(subsumption_resolution,[],[f3314,f1699])).
fof(f5535,plain,(
  ( ! [X12,X13] : (s(bool,X13) = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X12))) | s(bool,'T_0') = s(bool,X13)) ) | $spl34),
  inference(cnf_transformation,[],[f5535_D])).
fof(f5535_D,plain,(
  ( ! [X12,X13] : (s(bool,X13) = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X12))) | s(bool,'T_0') = s(bool,X13)) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f5534,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X13) | s(bool,X13) = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X12)))) ) | $spl22),
  inference(subsumption_resolution,[],[f3341,f2625])).
fof(f3341,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X13) | s(bool,X13) = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X12))) | s(bool,'F_0') = s(bool,X12)) ) | $spl22),
  inference(superposition,[],[f2120,f96])).
fof(f5533,plain,(
  ( ! [X19,X18] : (s(bool,X19) = s(bool,'const.combin.I_1'(s(bool,X18))) | s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X19)) ) | $spl22),
  inference(subsumption_resolution,[],[f5532,f2624])).
fof(f5532,plain,(
  ( ! [X19,X18] : (s(bool,X19) = s(bool,'const.combin.I_1'(s(bool,X18))) | s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X19) | s(bool,'F_0') = s(bool,X18)) ) | $spl22),
  inference(forward_demodulation,[],[f3346,f68])).
fof(f3346,plain,(
  ( ! [X19,X18] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X19) | s(bool,X19) = s(bool,'const.combin.I_1'(s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X18))))) | s(bool,'F_0') = s(bool,X18)) ) | $spl22),
  inference(superposition,[],[f2624,f96])).
fof(f5529,plain,(
  ( ! [X39,X41,X40] : (s(bool,'T_0') = s(bool,X40) | s(bool,X40) = s(bool,'const.combin.I_1'(s(bool,X39))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X41))) | s(bool,'T_0') != s(bool,'F_0')) ) | ($spl1 | $spl22)),
  inference(global_subsumption,[],[f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f347,f356,f390,f463,f465,f486,f966,f1158,f1159,f1293,f1294,f1305,f784,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f49,f1330,f781,f1567,f1756,f1757,f1758,f1759,f1766,f1810,f1855,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1950,f1983,f1984,f1985,f1986,f1956,f1957,f1958,f1989,f1965,f1966,f1967,f1968,f1969,f1974,f2073,f2074,f2075,f2076,f2119,f2124,f2126,f2130,f2104,f2088,f1324,f473,f467,f2120,f64,f2353,f2359,f2434,f2535,f2536,f2533,f2537,f2594,f2624,f2625,f2626,f2627,f2659,f2881,f2907,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3809,f3870,f3871,f4556,f4051,f4557,f4559,f4172,f4176,f4177,f4182,f4675,f4824,f4825,f4900,f4852,f4914,f4929,f4930,f4911,f4855,f5523,f5524,f4309,f5526,f5527,f5528,f4274])).
fof(f4274,plain,(
  ( ! [X39,X41,X40] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X40) | s(bool,X40) = s(bool,'const.combin.I_1'(s(bool,X39))) | s(bool,X39) = s(bool,X41) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X41)))) ) | ($spl1 | $spl22)),
  inference(superposition,[],[f2120,f2659])).
fof(f5527,plain,(
  ( ! [X54,X55,X53] : (s(bool,X54) = s(bool,'const.combin.I_1'(s(bool,X53))) | s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X54) | s(bool,X53) = s(bool,X55) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X55)))) ) | ($spl1 | $spl22)),
  inference(forward_demodulation,[],[f4280,f51])).
fof(f4280,plain,(
  ( ! [X54,X55,X53] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X54) | s(bool,X54) = s(bool,'const.combin.I_1'(s(bool,'const.combin.I_1'(s(bool,X53))))) | s(bool,X53) = s(bool,X55) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X55)))) ) | ($spl1 | $spl22)),
  inference(superposition,[],[f2624,f2659])).
fof(f4309,plain,(
  ( ! [X39,X41,X40] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X40) | s(bool,X40) = s(bool,'const.combin.I_1'(s(bool,X39))) | s(bool,X39) = s(bool,X41) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X41)))) ) | ($spl1 | $spl22)),
  inference(superposition,[],[f2120,f2659])).
fof(f5523,plain,(
  ( ! [X54,X55,X53] : (s(bool,X54) = s(bool,'const.combin.I_1'(s(bool,X53))) | s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X54) | s(bool,X53) = s(bool,X55) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X55)))) ) | ($spl1 | $spl22)),
  inference(forward_demodulation,[],[f4315,f51])).
fof(f4315,plain,(
  ( ! [X54,X55,X53] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X54) | s(bool,X54) = s(bool,'const.combin.I_1'(s(bool,'const.combin.I_1'(s(bool,X53))))) | s(bool,X53) = s(bool,X55) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X55)))) ) | ($spl1 | $spl22)),
  inference(superposition,[],[f2624,f2659])).
fof(f5528,plain,(
  ( ! [X54,X55,X53] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X55))) | s(bool,X54) = s(bool,'const.combin.I_1'(s(bool,X53))) | s(bool,'T_0') = s(bool,X54) | s(bool,'T_0') != s(bool,'F_0')) ) | ($spl1 | $spl22)),
  inference(global_subsumption,[],[f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f347,f356,f390,f463,f465,f486,f966,f1158,f1159,f1293,f1294,f1305,f784,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f49,f1330,f781,f1567,f1756,f1757,f1758,f1759,f1766,f1810,f1855,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1950,f1983,f1984,f1985,f1986,f1956,f1957,f1958,f1989,f1965,f1966,f1967,f1968,f1969,f1974,f2073,f2074,f2075,f2076,f2119,f2124,f2126,f2130,f2104,f2088,f1324,f473,f467,f2120,f64,f2353,f2359,f2434,f2535,f2536,f2533,f2537,f2594,f2624,f2625,f2626,f2627,f2659,f2881,f2907,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3809,f3870,f3871,f4556,f4051,f4557,f4559,f4172,f4176,f4177,f4182,f4675,f4824,f4825,f4900,f4852,f4914,f4929,f4930,f4911,f4855,f5523,f5524,f4309,f5526,f5527])).
fof(f5526,plain,(
  ( ! [X39,X41,X40] : (s(bool,X40) = s(bool,'const.combin.I_1'(s(bool,X39))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X41))) | s(bool,'T_0') = s(bool,X40) | s(bool,'T_0') != s(bool,'F_0')) ) | ($spl1 | $spl22)),
  inference(global_subsumption,[],[f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f347,f356,f390,f463,f465,f486,f966,f1158,f1159,f1293,f1294,f1305,f784,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f49,f1330,f781,f1567,f1756,f1757,f1758,f1759,f1766,f1810,f1855,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1950,f1983,f1984,f1985,f1986,f1956,f1957,f1958,f1989,f1965,f1966,f1967,f1968,f1969,f1974,f2073,f2074,f2075,f2076,f2119,f2124,f2126,f2130,f2104,f2088,f1324,f473,f467,f2120,f64,f2353,f2359,f2434,f2535,f2536,f2533,f2537,f2594,f2624,f2625,f2626,f2627,f2659,f2881,f2907,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3809,f3870,f3871,f4556,f4051,f4557,f4559,f4172,f4176,f4177,f4182,f4675,f4824,f4825,f4900,f4852,f4914,f4929,f4930,f4911,f4855,f5523,f5524,f4309])).
fof(f5524,plain,(
  ( ! [X54,X55,X53] : (s(bool,X54) = s(bool,'const.combin.I_1'(s(bool,X53))) | s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X55))) | s(bool,'T_0') = s(bool,X54)) ) | ($spl1 | $spl22)),
  inference(global_subsumption,[],[f46,f45,f47,f50,f51,f68,f57,f59,f58,f60,f61,f73,f76,f83,f140,f141,f153,f163,f164,f205,f347,f356,f390,f463,f465,f486,f966,f1158,f1159,f1293,f1294,f1305,f784,f1308,f1309,f1297,f1298,f56,f55,f348,f469,f471,f1307,f1296,f122,f96,f75,f49,f1330,f781,f1567,f1756,f1757,f1758,f1759,f1766,f1810,f1855,f1977,f1978,f1979,f1980,f1941,f1942,f1943,f1944,f1945,f1950,f1983,f1984,f1985,f1986,f1956,f1957,f1958,f1989,f1965,f1966,f1967,f1968,f1969,f1974,f2073,f2074,f2075,f2076,f2119,f2124,f2126,f2130,f2104,f2088,f1324,f473,f467,f2120,f64,f2353,f2359,f2434,f2535,f2536,f2533,f2537,f2594,f2624,f2625,f2626,f2627,f2659,f2881,f2907,f3322,f3323,f3332,f3350,f3352,f3353,f3784,f3785,f3786,f3846,f3805,f3809,f3870,f3871,f4556,f4051,f4557,f4559,f4172,f4176,f4177,f4182,f4675,f4824,f4825,f4900,f4852,f4914,f4929,f4930,f4911,f4855,f5523])).
fof(f5312,plain,(
  ( ! [X30,X29] : (s(bool,'const.combin.I_1'(s(bool,X29))) = s(bool,'const.combin.I_1'(s(bool,X30))) | s(bool,'T_0') = s(bool,X29) | s(bool,'T_0') = s(bool,X30)) ) | $spl30),
  inference(cnf_transformation,[],[f5312_D])).
fof(f5312_D,plain,(
  ( ! [X30,X29] : (s(bool,'const.combin.I_1'(s(bool,X29))) = s(bool,'const.combin.I_1'(s(bool,X30))) | s(bool,'T_0') = s(bool,X29) | s(bool,'T_0') = s(bool,X30)) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f5271,plain,(
  ( ! [X61,X62,X63] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X62))) | s(bool,'const.combin.I_1'(s(bool,X61))) = s(bool,'const.combin.I_1'(s(bool,X63))) | s(bool,'T_0') = s(bool,X63) | s(bool,'T_0') = s(bool,X61)) ) | ($spl1 | $spl13 | $spl22)),
  inference(duplicate_literal_removal,[],[f5049])).
fof(f5049,plain,(
  ( ! [X61,X62,X63] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X62))) | s(bool,'const.combin.I_1'(s(bool,X61))) = s(bool,'const.combin.I_1'(s(bool,X63))) | s(bool,'T_0') = s(bool,X63) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X62))) | s(bool,'T_0') = s(bool,X61)) ) | ($spl1 | $spl13 | $spl22)),
  inference(superposition,[],[f3407,f3407])).
fof(f3407,plain,(
  ( ! [X0,X1] : (s(bool,X1) = s(bool,'const.combin.I_1'(s(bool,X0))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X1))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl1 | $spl13 | $spl22)),
  inference(superposition,[],[f3053,f50])).
fof(f3053,plain,(
  ( ! [X15] : (s(bool,X15) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X15)))) ) | ($spl1 | $spl13 | $spl22)),
  inference(subsumption_resolution,[],[f2938,f418])).
fof(f2938,plain,(
  ( ! [X15] : (s(bool,X15) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X15)))) ) | ($spl1 | $spl22)),
  inference(superposition,[],[f2907,f2434])).
fof(f5289,plain,(
  ( ! [X30,X28,X29] : (s(bool,'const.combin.I_1'(s(bool,X29))) = s(bool,'const.combin.I_1'(s(bool,X30))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X28))) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,X29)) ) | ($spl1 | $spl13 | $spl22)),
  inference(duplicate_literal_removal,[],[f4946])).
fof(f4946,plain,(
  ( ! [X30,X28,X29] : (s(bool,'const.combin.I_1'(s(bool,X29))) = s(bool,'const.combin.I_1'(s(bool,X30))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X28))) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X28))) | s(bool,'T_0') = s(bool,X29)) ) | ($spl1 | $spl13 | $spl22)),
  inference(superposition,[],[f3407,f3407])).
fof(f2142,plain,(
  s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | $spl29),
  inference(cnf_transformation,[],[f2142_D])).
fof(f2142_D,plain,(
  s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'F_0'))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f3065,plain,(
  ( ! [X16] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X16)))) ) | ($spl16 | $spl22)),
  inference(subsumption_resolution,[],[f2974,f2245])).
fof(f2974,plain,(
  ( ! [X16] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X16))) | s(bool,'T_0') = s(bool,X16)) ) | ($spl16 | $spl22)),
  inference(superposition,[],[f2245,f2907])).
fof(f2341,plain,(
  $false | ($spl0 | $spl3)),
  inference(unit_resulting_resolution,[],[f2265,f2332])).
fof(f2332,plain,(
  ( ! [X0] : (s(sK0,X0) = s(sK0,'const.combin.o_3'(s(fun(sK0,sK0),sK1),s(fun(sK0,sK0),sK1),s(sK0,X0)))) ) | $spl0),
  inference(forward_demodulation,[],[f2314,f68])).
fof(f2314,plain,(
  ( ! [X0] : (s(sK0,happ(s(fun(sK0,sK0),'const.combin.I_0'),s(sK0,X0))) = s(sK0,'const.combin.o_3'(s(fun(sK0,sK0),sK1),s(fun(sK0,sK0),sK1),s(sK0,X0)))) ) | $spl0),
  inference(superposition,[],[f61,f2167])).
fof(f2167,plain,(
  s(fun(sK0,sK0),'const.combin.I_0') = s(fun(sK0,sK0),'const.combin.o_2'(s(fun(sK0,sK0),sK1),s(fun(sK0,sK0),sK1))) | $spl0),
  inference(unit_resulting_resolution,[],[f63,f55])).
fof(f2265,plain,(
  s(sK0,sK2) != s(sK0,'const.combin.o_3'(s(fun(sK0,sK0),sK1),s(fun(sK0,sK0),sK1),s(sK0,sK2))) | $spl3),
  inference(superposition,[],[f66,f60])).
fof(f2352,plain,(
  $false | ($spl0 | $spl3)),
  inference(trivial_inequality_removal,[],[f2347])).
fof(f2347,plain,(
  s(sK0,sK2) != s(sK0,sK2) | ($spl0 | $spl3)),
  inference(superposition,[],[f2265,f2332])).
fof(f2166,plain,(
  $false | ($spl0 | $spl9 | $spl13)),
  inference(subsumption_resolution,[],[f2165,f49])).
fof(f2165,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl9 | $spl13)),
  inference(forward_demodulation,[],[f63,f1534])).
fof(f2164,plain,(
  $false | ($spl3 | $spl4)),
  inference(subsumption_resolution,[],[f66,f67])).
fof(f2125,plain,(
  ( ! [X10,X8,X9] : (s(bool,X8) != s(bool,X9) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X10) | s(bool,X9) = s(bool,X10)) ) | $spl24),
  inference(cnf_transformation,[],[f2125_D])).
fof(f2125_D,plain,(
  ( ! [X10,X8,X9] : (s(bool,X8) != s(bool,X9) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X10) | s(bool,X9) = s(bool,X10)) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f2143,plain,(
  s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | $spl28),
  inference(cnf_transformation,[],[f2143_D])).
fof(f2143_D,plain,(
  s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'F_0'))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f2145,plain,(
  ( ! [X50,X51,X49] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X49) != s(bool,X50) | s(bool,X50) = s(bool,X51) | s(bool,'T_0') = s(bool,X51) | s(bool,'T_0') = s(bool,X49)) ) | ($spl1 | $spl9)),
  inference(forward_demodulation,[],[f2091,f349])).
fof(f349,plain,(
  s(bool,'F_0') = s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | ($spl1 | $spl9)),
  inference(unit_resulting_resolution,[],[f86,f347])).
fof(f2091,plain,(
  ( ! [X50,X51,X49] : (s(bool,X49) != s(bool,X50) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))))) | s(bool,X50) = s(bool,X51) | s(bool,'T_0') = s(bool,X51) | s(bool,'T_0') = s(bool,X49)) ) | $spl1),
  inference(superposition,[],[f1567,f347])).
fof(f2088,plain,(
  ( ! [X43,X41,X42] : (s(bool,X41) != s(bool,X42) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X42) = s(bool,X43) | s(bool,'T_0') = s(bool,X43) | s(bool,'T_0') = s(bool,X41)) ) | $spl1),
  inference(superposition,[],[f1567,f50])).
fof(f2131,plain,(
  ( ! [X21,X20] : (s(bool,X20) != s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,X21) | s(bool,X20) = s(bool,X21)) ) | $spl26),
  inference(cnf_transformation,[],[f2131_D])).
fof(f2131_D,plain,(
  ( ! [X21,X20] : (s(bool,X20) != s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,X21) | s(bool,X20) = s(bool,X21)) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f2104,plain,(
  ( ! [X24,X23,X22] : (s(bool,X23) != s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X22))) | s(bool,X23) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) ) | $spl1),
  inference(duplicate_literal_removal,[],[f2083])).
fof(f2083,plain,(
  ( ! [X24,X23,X22] : (s(bool,X23) != s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X22))) | s(bool,X23) = s(bool,X24) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X22)))) ) | $spl1),
  inference(superposition,[],[f1567,f356])).
fof(f2130,plain,(
  ( ! [X21,X19,X20] : (s(bool,X20) != s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X19))) | s(bool,X20) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) ) | $spl1),
  inference(subsumption_resolution,[],[f2082,f1158])).
fof(f2082,plain,(
  ( ! [X21,X19,X20] : (s(bool,X20) != s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X19))) | s(bool,X20) = s(bool,X21) | s(bool,'T_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X19)) ) | $spl1),
  inference(superposition,[],[f1567,f347])).
fof(f2126,plain,(
  ( ! [X14,X12,X13,X11] : (s(bool,X12) != s(bool,X13) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X11))) | s(bool,X13) = s(bool,X14) | s(bool,'T_0') = s(bool,X14) | s(bool,'T_0') = s(bool,X12)) ) | $spl1),
  inference(subsumption_resolution,[],[f2080,f1158])).
fof(f2080,plain,(
  ( ! [X14,X12,X13,X11] : (s(bool,X12) != s(bool,X13) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X11))) | s(bool,X13) = s(bool,X14) | s(bool,'T_0') = s(bool,X14) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X12)) ) | $spl1),
  inference(superposition,[],[f1567,f73])).
fof(f2124,plain,(
  ( ! [X10,X8,X7,X9] : (s(bool,X8) != s(bool,X9) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X7))) | s(bool,X9) = s(bool,X10) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X8)) ) | $spl1),
  inference(subsumption_resolution,[],[f2079,f1158])).
fof(f2079,plain,(
  ( ! [X10,X8,X7,X9] : (s(bool,X8) != s(bool,X9) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X7))) | s(bool,X9) = s(bool,X10) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X7)) ) | $spl1),
  inference(superposition,[],[f1567,f73])).
fof(f2120,plain,(
  ( ! [X2,X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X1) = s(bool,X2)) ) | $spl22),
  inference(cnf_transformation,[],[f2120_D])).
fof(f2120_D,plain,(
  ( ! [X2,X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X1) = s(bool,X2)) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f2119,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X0))) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl1),
  inference(subsumption_resolution,[],[f2077,f1158])).
fof(f2077,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X0))) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) ) | $spl1),
  inference(superposition,[],[f1567,f50])).
fof(f1453,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1430,f47])).
fof(f1430,plain,(
  ~p(s(bool,'T_0')) | $spl14),
  inference(backward_demodulation,[],[f1383,f49])).
fof(f1383,plain,(
  ( ! [X18] : (s(bool,'T_0') = s(bool,X18)) ) | $spl14),
  inference(superposition,[],[f1112,f51])).
fof(f1251,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | ($spl1 | $spl10)),
  inference(forward_demodulation,[],[f1250,f416])).
fof(f1250,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,g(s(X0,sK3(X1,X2,bool,X0)))))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | ($spl1 | $spl10)),
  inference(condensation,[],[f1249])).
fof(f1249,plain,(
  ( ! [X64,X62,X65,X63] : (s(fun(X62,bool),X63) = s(fun(X62,bool),X64) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X65))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,g(s(X62,sK3(X63,X64,bool,X62))))))) ) | ($spl1 | $spl10)),
  inference(subsumption_resolution,[],[f1244,f356])).
fof(f1244,plain,(
  ( ! [X64,X62,X65,X63] : (s(bool,X65) != s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(fun(X62,bool),X63) = s(fun(X62,bool),X64) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X65))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,g(s(X62,sK3(X63,X64,bool,X62))))))) ) | ($spl1 | $spl10)),
  inference(backward_demodulation,[],[f416,f781])).
fof(f1254,plain,(
  ( ! [X23,X21,X22] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))))) | s(fun(X21,bool),X22) = s(fun(X21,bool),X23)) ) | ($spl1 | $spl9 | $spl10)),
  inference(forward_demodulation,[],[f1253,f416])).
fof(f1253,plain,(
  ( ! [X23,X21,X22] : (s(fun(X21,bool),X22) = s(fun(X21,bool),X23) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,g(s(X21,sK3(X22,X23,bool,X21))))))) ) | ($spl1 | $spl9 | $spl10)),
  inference(subsumption_resolution,[],[f1245,f416])).
fof(f1245,plain,(
  ( ! [X23,X21,X22] : (s(bool,'F_0') != s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(fun(X21,bool),X22) = s(fun(X21,bool),X23) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,g(s(X21,sK3(X22,X23,bool,X21))))))) ) | ($spl1 | $spl9 | $spl10)),
  inference(backward_demodulation,[],[f416,f585])).
fof(f585,plain,(
  ( ! [X23,X21,X22] : (s(bool,'F_0') != s(bool,f(s(X21,sK3(X22,X23,bool,X21)))) | s(fun(X21,bool),X22) = s(fun(X21,bool),X23) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,g(s(X21,sK3(X22,X23,bool,X21))))))) ) | ($spl1 | $spl9)),
  inference(forward_demodulation,[],[f533,f349])).
fof(f533,plain,(
  ( ! [X23,X21,X22] : (s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) != s(bool,f(s(X21,sK3(X22,X23,bool,X21)))) | s(fun(X21,bool),X22) = s(fun(X21,bool),X23) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,g(s(X21,sK3(X22,X23,bool,X21))))))) ) | $spl1),
  inference(superposition,[],[f57,f356])).
fof(f1243,plain,(
  $false | ($spl1 | $spl9 | $spl12)),
  inference(subsumption_resolution,[],[f1242,f86])).
fof(f1242,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl1 | $spl9 | $spl12)),
  inference(backward_demodulation,[],[f417,f349])).
fof(f1113,plain,(
  ( ! [X4] : (s(bool,X4) != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X4)) ) | $spl16),
  inference(cnf_transformation,[],[f1113_D])).
fof(f1113_D,plain,(
  ( ! [X4] : (s(bool,X4) != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X4)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f1114,plain,(
  ( ! [X6,X5] : (s(bool,X6) != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X5))) | s(bool,'T_0') = s(bool,X6)) ) | ($spl1 | $spl13)),
  inference(subsumption_resolution,[],[f1058,f537])).
fof(f537,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X0)))) ) | ($spl1 | $spl13)),
  inference(superposition,[],[f418,f356])).
fof(f1058,plain,(
  ( ! [X6,X5] : (s(bool,X6) != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X5))) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,X6)) ) | ($spl1 | $spl13)),
  inference(superposition,[],[f1049,f73])).
fof(f1049,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X0)))) ) | ($spl1 | $spl13)),
  inference(subsumption_resolution,[],[f841,f537])).
fof(f841,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X0)))) ) | $spl1),
  inference(equality_factoring,[],[f486])).
fof(f1111,plain,(
  ( ! [X4,X3] : (s(bool,X4) != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X3))) | s(bool,'T_0') = s(bool,X4)) ) | ($spl1 | $spl13)),
  inference(subsumption_resolution,[],[f1057,f537])).
fof(f1057,plain,(
  ( ! [X4,X3] : (s(bool,X4) != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X3))) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) ) | ($spl1 | $spl13)),
  inference(superposition,[],[f1049,f73])).
fof(f416,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1)))) ) | $spl10),
  inference(cnf_transformation,[],[f416_D])).
fof(f416_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1)))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f418,plain,(
  s(bool,'T_0') != s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | $spl13),
  inference(cnf_transformation,[],[f418_D])).
fof(f418_D,plain,(
  s(bool,'T_0') != s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f390,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))) | s(bool,X0) = s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1)))) ) | $spl1),
  inference(equality_factoring,[],[f347])).
fof(f83,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f50])).
fof(f63,plain,(
  p(s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1)))) | $spl0),
  inference(cnf_transformation,[],[f63_D])).
fof(f63_D,plain,(
  p(s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f67,plain,(
  ( ! [X3] : (s(sK0,X3) = s(sK0,happ(s(fun(sK0,sK0),sK1),s(sK0,happ(s(fun(sK0,sK0),sK1),s(sK0,X3)))))) ) | $spl4),
  inference(cnf_transformation,[],[f67_D])).
fof(f67_D,plain,(
  ( ! [X3] : (s(sK0,X3) = s(sK0,happ(s(fun(sK0,sK0),sK1),s(sK0,happ(s(fun(sK0,sK0),sK1),s(sK0,X3)))))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f45,plain,(
  ( ! [X3] : (s(sK0,X3) = s(sK0,happ(s(fun(sK0,sK0),sK1),s(sK0,happ(s(fun(sK0,sK0),sK1),s(sK0,X3))))) | p(s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))))) )),
  inference(cnf_transformation,[],[f38])).
fof(f38,plain,(
  (p(s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1)))) | ! [X3] : s(sK0,X3) = s(sK0,happ(s(fun(sK0,sK0),sK1),s(sK0,happ(s(fun(sK0,sK0),sK1),s(sK0,X3)))))) & (~p(s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1)))) | s(sK0,sK2) != s(sK0,happ(s(fun(sK0,sK0),sK1),s(sK0,happ(s(fun(sK0,sK0),sK1),s(sK0,sK2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f37])).
fof(f37,plain,(
  ? [X0,X1] : ((p(s(bool,'const.relation.INVOL_1'(s(fun(X0,X0),X1)))) | ! [X3] : s(X0,X3) = s(X0,happ(s(fun(X0,X0),X1),s(X0,happ(s(fun(X0,X0),X1),s(X0,X3)))))) & (~p(s(bool,'const.relation.INVOL_1'(s(fun(X0,X0),X1)))) | ? [X2] : s(X0,X2) != s(X0,happ(s(fun(X0,X0),X1),s(X0,happ(s(fun(X0,X0),X1),s(X0,X2)))))))),
  inference(rectify,[],[f36])).
fof(f36,plain,(
  ? [X0,X1] : ((p(s(bool,'const.relation.INVOL_1'(s(fun(X0,X0),X1)))) | ! [X2] : s(X0,X2) = s(X0,happ(s(fun(X0,X0),X1),s(X0,happ(s(fun(X0,X0),X1),s(X0,X2)))))) & (~p(s(bool,'const.relation.INVOL_1'(s(fun(X0,X0),X1)))) | ? [X2] : s(X0,X2) != s(X0,happ(s(fun(X0,X0),X1),s(X0,happ(s(fun(X0,X0),X1),s(X0,X2)))))))),
  inference(nnf_transformation,[],[f34])).
fof(f34,plain,(
  ? [X0,X1] : (p(s(bool,'const.relation.INVOL_1'(s(fun(X0,X0),X1)))) <~> ! [X2] : s(X0,X2) = s(X0,happ(s(fun(X0,X0),X1),s(X0,happ(s(fun(X0,X0),X1),s(X0,X2))))))),
  inference(ennf_transformation,[],[f16])).
fof(f16,plain,(
  ~! [X0,X1] : (p(s(bool,'const.relation.INVOL_1'(s(fun(X0,X0),X1)))) <=> ! [X2] : s(X0,X2) = s(X0,happ(s(fun(X0,X0),X1),s(X0,happ(s(fun(X0,X0),X1),s(X0,X2))))))),
  inference(flattening,[],[f15])).
fof(f15,plain,(
  ~! [X0] : ! [X1] : (p(s(bool,'const.relation.INVOL_1'(s(fun(X0,X0),X1)))) <=> ! [X2] : s(X0,X2) = s(X0,happ(s(fun(X0,X0),X1),s(X0,happ(s(fun(X0,X0),X1),s(X0,X2))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7] : (p(s(bool,'const.relation.INVOL_1'(s(fun(X6,X6),X7)))) <=> ! [X8] : s(X6,X8) = s(X6,happ(s(fun(X6,X6),X7),s(X6,happ(s(fun(X6,X6),X7),s(X6,X8))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7] : (p(s(bool,'const.relation.INVOL_1'(s(fun(X6,X6),X7)))) <=> ! [X8] : s(X6,X8) = s(X6,happ(s(fun(X6,X6),X7),s(X6,happ(s(fun(X6,X6),X7),s(X6,X8))))))),
  file('thm.relation.INVOL','thm.relation.INVOL')).
fof(f64,plain,(
  ~p(s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1)))) | $spl1),
  inference(cnf_transformation,[],[f64_D])).
fof(f64_D,plain,(
  ~p(s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f66,plain,(
  s(sK0,sK2) != s(sK0,happ(s(fun(sK0,sK0),sK1),s(sK0,happ(s(fun(sK0,sK0),sK1),s(sK0,sK2))))) | $spl3),
  inference(cnf_transformation,[],[f66_D])).
fof(f66_D,plain,(
  s(sK0,sK2) != s(sK0,happ(s(fun(sK0,sK0),sK1),s(sK0,happ(s(fun(sK0,sK0),sK1),s(sK0,sK2))))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f46,plain,(
  s(sK0,sK2) != s(sK0,happ(s(fun(sK0,sK0),sK1),s(sK0,happ(s(fun(sK0,sK0),sK1),s(sK0,sK2))))) | ~p(s(bool,'const.relation.INVOL_1'(s(fun(sK0,sK0),sK1))))),
  inference(cnf_transformation,[],[f38])).
% SZS output end Proof for thm.relation
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 149
% Passive clauses: 1376
% Generated clauses: 22414
% Final active clauses: 79
% Final passive clauses: 737
% Input formulas: 13
% Initial clauses: 17
% 
% Duplicate literals: 815
% Trivial inequalities: 239
% Fw subsumption resolutions: 1413
% Fw demodulations: 4712
% Bw demodulations: 374
% Condensations: 1512
% Global subsumptions: 99
% 
% Simple tautologies: 191
% Equational tautologies: 204
% Forward subsumptions: 11623
% Fw demodulations to eq. taut.: 51
% Bw demodulations to eq. taut.: 19
% 
% Unit resulting resolution: 892
% Forward superposition: 5747
% Backward superposition: 4598
% Self superposition: 247
% Equality factoring: 68
% Equality resolution: 32
% 
% Split clauses: 60
% Split components: 124
% SAT solver clauses: 1718
% SAT solver unit clauses: 25
% SAT solver binary clauses: 359
% 
% Sat splits: 60
% Sat splitting refutations: 121
% 
% Memory used [KB]: 15735
% Time elapsed: 0.942 s
% ------------------------------
% ------------------------------
% Success in time 2.267 s
