Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.arithmetic
% (146907)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 534
% Passive clauses: 8388
% Generated clauses: 57920
% Final active clauses: 446
% Final passive clauses: 7599
% Input formulas: 133
% Initial clauses: 271
% 
% Duplicate literals: 105
% Trivial inequalities: 318
% Fw subsumption resolutions: 6245
% Fw demodulations: 11080
% Bw demodulations: 84
% 
% Simple tautologies: 445
% Equational tautologies: 24
% Forward subsumptions: 31157
% Fw demodulations to eq. taut.: 13
% 
% Binary resolution: 2207
% Factoring: 2
% Forward superposition: 13457
% Backward superposition: 23920
% Self superposition: 79
% Equality factoring: 14
% Equality resolution: 11
% 
% Split clauses: 165
% Split components: 331
% SAT solver clauses: 324
% SAT solver unit clauses: 9
% SAT solver binary clauses: 99
% 
% Sat splits: 165
% Sat splitting refutations: 11
% 
% TWLsolver clauses: 321
% TWLsolver calls for satisfiability: 103
% 
% Memory used [KB]: 31598
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.arithmetic
% (146924)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 352
% Passive clauses: 3162
% Generated clauses: 26720
% Final active clauses: 148
% Final passive clauses: 1185
% Input formulas: 133
% Initial clauses: 271
% 
% Duplicate literals: 127
% Trivial inequalities: 107
% Fw subsumption resolutions: 1980
% Fw demodulations: 6689
% Bw demodulations: 1711
% Condensations: 298
% Global subsumptions: 20
% 
% Simple tautologies: 401
% Equational tautologies: 14
% Forward subsumptions: 9823
% Bw demodulations to eq. taut.: 7
% 
% Binary resolution: 1042
% Factoring: 1
% Forward superposition: 5238
% Backward superposition: 3882
% Self superposition: 85
% Equality factoring: 4
% Equality resolution: 12
% 
% Split clauses: 666
% Split components: 1397
% SAT solver clauses: 3729
% SAT solver unit clauses: 335
% SAT solver binary clauses: 124
% SAT solver learnt clauses: 49
% SAT solver learnt literals: 34
% 
% Sat splits: 666
% Sat splitting refutations: 110
% 
% TWLsolver clauses: 3640
% TWLsolver calls for satisfiability: 1836
% 
% Memory used [KB]: 15607
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.arithmetic
% (146944)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 470
% Passive clauses: 1706
% Generated clauses: 336742
% Final active clauses: 410
% Final passive clauses: 915
% Input formulas: 133
% Initial clauses: 271
% 
% Duplicate literals: 7
% Trivial inequalities: 4
% Fw subsumption resolutions: 126
% Fw demodulations: 226350
% Bw demodulations: 365
% Condensations: 20
% Global subsumptions: 3
% 
% Simple tautologies: 81
% Equational tautologies: 19
% Forward subsumptions: 108324
% Fw demodulations to eq. taut.: 78
% Bw demodulations to eq. taut.: 1
% 
% Unit resulting resolution: 102720
% Factoring: 1
% Forward superposition: 4231
% Backward superposition: 3160
% Self superposition: 70
% Equality factoring: 2
% Equality resolution: 5
% 
% Split clauses: 9
% Split components: 18
% SAT solver clauses: 1744
% SAT solver unit clauses: 874
% SAT solver binary clauses: 671
% 
% Sat splits: 9
% Sat splitting refutations: 15
% 
% Memory used [KB]: 255859
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.arithmetic
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.arithmetic
% SZS output start Proof for thm.arithmetic
fof(f5688,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f596,f597,f672,f598,f599,f673,f674,f677,f678,f681,f682,f685,f688,f680,f684,f689,f676,f693,f697,f701,f708,f712,f639,f716,f640,f641,f721,f642,f643,f722,f644,f645,f723,f720,f654,f705,f663,f730,f664,f666,f734,f801,f695,f710,f1461,f1487,f1462,f1463,f1488,f1464,f1465,f1489,f1493,f1623,f1647,f1624,f1625,f1648,f1626,f1627,f1649,f1653,f1655,f1646,f1659,f1960,f1962,f1946,f732,f2134,f2159,f2135,f2136,f2160,f2137,f2138,f2161,f2166,f1661,f1486,f2589,f2591,f2593,f2526,f2595,f2527,f2541,f1970,f2900,f2945,f714,f3357,f703,f3294,f718,f3362,f3364,f3363,f2553,f3365,f2551,f3366,f2486,f2485,f3371,f2484,f2483,f3372,f1949,f3375,f1947,f733,f1914,f1913,f3380,f1912,f1911,f3358,f671,f3382,f1495,f670,f4468,f4471,f4716,f4717,f4718,f4832,f4831,f4830,f4829,f4809,f4808,f4807,f4806,f4805,f4804,f4803,f4802,f4801,f4778,f4885,f4886,f4887,f4896,f4897,f5495,f5498,f5430,f1969,f728,f5505,f5501,f727,f687,f5687])).
fof(f5687,plain,(
  $false | ($spl7 | $spl50)),
  inference(subsumption_resolution,[],[f1657,f1485])).
fof(f1485,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK4)))) | $spl50),
  inference(cnf_transformation,[],[f1485_D])).
fof(f1485_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK4)))) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f1657,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK4)))) | $spl7),
  inference(trivial_inequality_removal,[],[f1656])).
fof(f1656,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK4)))) | $spl7),
  inference(forward_demodulation,[],[f1636,f433])).
fof(f433,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f163])).
fof(f163,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))),
  inference(rectify,[],[f115])).
fof(f115,axiom,(
  ! [X6,X7] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X7))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num',X6)))),
  file('thm.arithmetic.SUB_RIGHT_ADD','thm.arithmetic.ADD_SYM')).
fof(f1636,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num',sK4))),s('type.num.num',sK5))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK4)))) | $spl7),
  inference(superposition,[],[f676,f425])).
fof(f425,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X1))),s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num',X0))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f265])).
fof(f265,plain,(
  ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ! [X2] : s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X1))),s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num',X0))))))),
  inference(ennf_transformation,[],[f155])).
fof(f155,plain,(
  ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) => ! [X2] : s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X1))),s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num',X0))))))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X30,X31] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X30),s('type.num.num',X31)))) => ! [X32] : s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',X31))),s('type.num.num',X30))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X31),s('type.num.num',X30))))))),
  file('thm.arithmetic.SUB_RIGHT_ADD','thm.arithmetic.LESS_EQ_ADD_SUB')).
fof(f687,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))) | $spl12),
  inference(cnf_transformation,[],[f687_D])).
fof(f687_D,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f727,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))),s('type.num.num','const.num.|48|_0')))) | $spl40),
  inference(cnf_transformation,[],[f727_D])).
fof(f727_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))),s('type.num.num','const.num.|48|_0')))) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f5501,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))),s('type.num.num','const.num.|48|_0')))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))) | $spl42),
  inference(forward_demodulation,[],[f5381,f442])).
fof(f442,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f171])).
fof(f171,plain,(
  ! [X0,X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f170])).
fof(f170,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X9] : ! [X17,X18] : s(X9,X18) = s(X9,'const.bool.COND_3'(s(bool,'F_0'),s(X9,X17),s(X9,X18)))),
  file('thm.arithmetic.SUB_RIGHT_ADD','thm.bool.COND_CLAUSES|split|1')).
fof(f5381,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))),s('type.num.num','const.num.|48|_0')))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))) | $spl42),
  inference(superposition,[],[f729,f458])).
fof(f458,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.arithmetic.SUB_RIGHT_ADD',aHL_BOOL_CASES)).
fof(f729,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))),s('type.num.num','const.num.|48|_0')))) | $spl42),
  inference(cnf_transformation,[],[f729_D])).
fof(f729_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))),s('type.num.num','const.num.|48|_0')))) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f5505,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))),s('type.num.num','const.num.|48|_0')))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))) | $spl42),
  inference(forward_demodulation,[],[f5383,f442])).
fof(f5383,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))),s('type.num.num','const.num.|48|_0')))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))) | $spl42),
  inference(superposition,[],[f729,f459])).
fof(f459,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f187])).
fof(f187,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X14] : (s(bool,'T_0') = s(bool,X14) | s(bool,'F_0') = s(bool,X14))),
  file('thm.arithmetic.SUB_RIGHT_ADD','thm.bool.BOOL_CASES_AX')).
fof(f728,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))),s('type.num.num','const.num.|48|_0')))) | $spl41),
  inference(cnf_transformation,[],[f728_D])).
fof(f728_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))),s('type.num.num','const.num.|48|_0')))) <=> ~$spl41),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl41])])).
fof(f1969,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')))) | $spl72),
  inference(cnf_transformation,[],[f1969_D])).
fof(f1969_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')))) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f5430,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))),s('type.num.num','const.num.|48|_0')))) | $spl42),
  inference(superposition,[],[f729,f448])).
fof(f448,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))) )),
  inference(cnf_transformation,[],[f322])).
fof(f322,plain,(
  ! [X0] : ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) & (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))))),
  inference(nnf_transformation,[],[f180])).
fof(f180,plain,(
  ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X0))),
  inference(rectify,[],[f49])).
fof(f49,axiom,(
  ! [X7] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num','const.num.|48|_0')))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X7))),
  file('thm.arithmetic.SUB_RIGHT_ADD','thm.arithmetic.LESS_EQ_|48|')).
fof(f5498,plain,(
  $false | ($spl1 | $spl41 | $spl42)),
  inference(subsumption_resolution,[],[f5497,f728])).
fof(f5497,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))),s('type.num.num','const.num.|48|_0')))) | ($spl1 | $spl42)),
  inference(forward_demodulation,[],[f5496,f442])).
fof(f5496,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))),s('type.num.num','const.num.|48|_0')))) | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5377,f670])).
fof(f5377,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5)))) | $spl42),
  inference(superposition,[],[f729,f453])).
fof(f453,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f324])).
fof(f324,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f184])).
fof(f184,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f183])).
fof(f183,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X14] : (s(bool,'F_0') = s(bool,X14) <=> ~p(s(bool,X14)))),
  file('thm.arithmetic.SUB_RIGHT_ADD','thm.bool.EQ_CLAUSES|split|3')).
fof(f5495,plain,(
  $false | ($spl1 | $spl41 | $spl42)),
  inference(subsumption_resolution,[],[f5494,f728])).
fof(f5494,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))),s('type.num.num','const.num.|48|_0')))) | ($spl1 | $spl42)),
  inference(forward_demodulation,[],[f5493,f442])).
fof(f5493,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))),s('type.num.num','const.num.|48|_0')))) | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5376,f670])).
fof(f5376,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5)))) | $spl42),
  inference(superposition,[],[f729,f451])).
fof(f451,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f323])).
fof(f323,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f182])).
fof(f182,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f181])).
fof(f181,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X14] : (s(bool,'F_0') = s(bool,X14) <=> ~p(s(bool,X14)))),
  file('thm.arithmetic.SUB_RIGHT_ADD','thm.bool.EQ_CLAUSES|split|2')).
fof(f4897,plain,(
  $false | $spl86),
  inference(subsumption_resolution,[],[f4827,f3363])).
fof(f4827,plain,(
  ( ! [X235,X234] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X234),s('type.num.num',X235))))) ) | $spl86),
  inference(resolution,[],[f3363,f491])).
fof(f491,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f333])).
fof(f333,plain,(
  ! [X0,X1] : ((p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) & (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(nnf_transformation,[],[f204])).
fof(f204,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(flattening,[],[f203])).
fof(f203,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(rectify,[],[f116])).
fof(f116,axiom,(
  ! [X6,X7] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num',X7)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',X6)))))),
  file('thm.arithmetic.SUB_RIGHT_ADD','thm.arithmetic.NOT_LESS')).
fof(f4896,plain,(
  $false | $spl86),
  inference(subsumption_resolution,[],[f4798,f3363])).
fof(f4798,plain,(
  ( ! [X187,X186] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X186),s('type.num.num',X187))))) ) | $spl86),
  inference(resolution,[],[f3363,f491])).
fof(f4887,plain,(
  $false | $spl86),
  inference(subsumption_resolution,[],[f4780,f3363])).
fof(f4780,plain,(
  ~p(s(bool,'T_0')) | $spl86),
  inference(resolution,[],[f3363,f575])).
fof(f575,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f374])).
fof(f374,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f244])).
fof(f244,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f57])).
fof(f57,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.arithmetic.SUB_RIGHT_ADD','thm.bool.NOT_CLAUSES|split|1')).
fof(f4886,plain,(
  $false | $spl86),
  inference(subsumption_resolution,[],[f4779,f3363])).
fof(f4779,plain,(
  ~p(s(bool,'T_0')) | $spl86),
  inference(resolution,[],[f3363,f573])).
fof(f573,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f373])).
fof(f373,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f243])).
fof(f243,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f58])).
fof(f58,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.arithmetic.SUB_RIGHT_ADD','thm.bool.NOT_CLAUSES|split|2')).
fof(f4885,plain,(
  $false | $spl86),
  inference(subsumption_resolution,[],[f4771,f3363])).
fof(f4771,plain,(
  ( ! [X132] : (~p(s(bool,X132))) ) | $spl86),
  inference(resolution,[],[f3363,f522])).
fof(f522,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f342])).
fof(f342,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f341])).
fof(f341,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f282])).
fof(f282,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f221])).
fof(f221,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f220])).
fof(f220,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X14] : ((p(s(bool,X14)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X14)))),
  file('thm.arithmetic.SUB_RIGHT_ADD','thm.bool.IMP_CLAUSES|split|4')).
fof(f4778,plain,(
  $false | $spl86),
  inference(resolution,[],[f3363,f571])).
fof(f571,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f242])).
fof(f242,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.arithmetic.SUB_RIGHT_ADD',aHL_FALSITY)).
fof(f4801,plain,(
  $false | ($spl41 | $spl86)),
  inference(resolution,[],[f3363,f728])).
fof(f4802,plain,(
  $false | ($spl19 | $spl86)),
  inference(resolution,[],[f3363,f703])).
fof(f4803,plain,(
  $false | ($spl31 | $spl86)),
  inference(resolution,[],[f3363,f718])).
fof(f4804,plain,(
  $false | $spl86),
  inference(resolution,[],[f3363,f497])).
fof(f497,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0'))))) )),
  inference(cnf_transformation,[],[f211])).
fof(f211,plain,(
  ! [X0] : ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0'))))),
  inference(flattening,[],[f210])).
fof(f210,plain,(
  ! [X0] : ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0'))))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X7] : ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))),s('type.num.num','const.num.|48|_0'))))),
  file('thm.arithmetic.SUB_RIGHT_ADD','thm.arithmetic.NOT_SUC_LESS_EQ_|48|')).
fof(f4805,plain,(
  $false | ($spl15 | $spl86)),
  inference(resolution,[],[f3363,f695])).
fof(f4806,plain,(
  $false | ($spl23 | $spl86)),
  inference(resolution,[],[f3363,f710])).
fof(f4807,plain,(
  $false | ($spl33 | $spl86)),
  inference(resolution,[],[f3363,f720])).
fof(f4808,plain,(
  $false | $spl86),
  inference(resolution,[],[f3363,f699])).
fof(f699,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))))),
  inference(subsumption_resolution,[],[f698,f410])).
fof(f410,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num',X2))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num',X0),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X2)))))) )),
  inference(cnf_transformation,[],[f141])).
fof(f141,plain,(
  ! [X0,X1,X2] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num',X2))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num',X0),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X2)))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6,X7,X8] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X7),s('type.num.num',X8))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',X8))),s('type.num.num',X6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X7))),s('type.num.num',X8)))))),
  file('thm.arithmetic.SUB_RIGHT_ADD','thm.arithmetic.SUB_LEFT_ADD')).
fof(f698,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK4))),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))))),
  inference(forward_demodulation,[],[f621,f433])).
fof(f621,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))))),
  inference(superposition,[],[f595,f448])).
fof(f595,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5)))))),
  inference(forward_demodulation,[],[f379,f433])).
fof(f379,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5)))))),
  inference(cnf_transformation,[],[f299])).
fof(f299,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4,sK5,sK6])],[f256])).
fof(f256,plain,(
  ? [X0,X1,X2] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X2))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X2),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X2))),s('type.num.num',X1)))))),
  inference(ennf_transformation,[],[f135])).
fof(f135,plain,(
  ~! [X0,X1,X2] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X2))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X2),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X2))),s('type.num.num',X1)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X6),s('type.num.num',X7))),s('type.num.num',X8))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num',X7))),s('type.num.num',X8),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X8))),s('type.num.num',X7)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X6),s('type.num.num',X7))),s('type.num.num',X8))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num',X7))),s('type.num.num',X8),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X8))),s('type.num.num',X7)))))),
  file('thm.arithmetic.SUB_RIGHT_ADD','thm.arithmetic.SUB_RIGHT_ADD')).
fof(f4809,plain,(
  $false | ($spl27 | $spl86)),
  inference(resolution,[],[f3363,f714])).
fof(f4829,plain,(
  $false | ($spl15 | $spl86)),
  inference(resolution,[],[f3363,f822])).
fof(f822,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0')))) | $spl15),
  inference(resolution,[],[f695,f494])).
fof(f494,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f276])).
fof(f276,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))))),
  inference(ennf_transformation,[],[f207])).
fof(f207,plain,(
  ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) => p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))))),
  inference(rectify,[],[f124])).
fof(f124,axiom,(
  ! [X6,X7] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num',X7)))) => p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num',X7)))))),
  file('thm.arithmetic.SUB_RIGHT_ADD','thm.arithmetic.LESS_IMP_LESS_OR_EQ')).
fof(f4830,plain,(
  $false | ($spl23 | $spl86)),
  inference(resolution,[],[f3363,f1024])).
fof(f1024,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')))) | $spl23),
  inference(resolution,[],[f710,f494])).
fof(f4831,plain,(
  $false | $spl86),
  inference(resolution,[],[f3363,f923])).
fof(f923,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))))),
  inference(resolution,[],[f699,f494])).
fof(f4832,plain,(
  $false | ($spl27 | $spl86)),
  inference(resolution,[],[f3363,f2962])).
fof(f2962,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6),s('type.num.num',sK5)))) | $spl27),
  inference(resolution,[],[f714,f494])).
fof(f4718,plain,(
  $false | ($spl1 | $spl51)),
  inference(subsumption_resolution,[],[f1231,f1486])).
fof(f1231,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK4)))) | $spl1),
  inference(resolution,[],[f735,f490])).
fof(f490,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f333])).
fof(f735,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK4),s('type.num.num',sK5)))) | $spl1),
  inference(resolution,[],[f670,f494])).
fof(f4717,plain,(
  $false | ($spl1 | $spl51)),
  inference(subsumption_resolution,[],[f1315,f1486])).
fof(f1315,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK4)))) | $spl1),
  inference(resolution,[],[f736,f494])).
fof(f736,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK5),s('type.num.num',sK4)))) | $spl1),
  inference(resolution,[],[f670,f490])).
fof(f4716,plain,(
  $false | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f3383,f671])).
fof(f3383,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5)))) | $spl9),
  inference(subsumption_resolution,[],[f1460,f461])).
fof(f461,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f189])).
fof(f189,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X6] : s('type.num.num',X6) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0')))),
  file('thm.arithmetic.SUB_RIGHT_ADD','thm.arithmetic.ADD_|48|')).
fof(f1460,plain,(
  s('type.num.num',sK6) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5)))) | $spl9),
  inference(superposition,[],[f680,f417])).
fof(f417,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f317])).
fof(f317,plain,(
  ! [X0,X1] : ((s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))) & (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1)))))),
  inference(nnf_transformation,[],[f151])).
fof(f151,plain,(
  ! [X0,X1] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))))),
  inference(rectify,[],[f70])).
fof(f70,axiom,(
  ! [X6,X7] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X6),s('type.num.num',X7))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num',X7)))))),
  file('thm.arithmetic.SUB_RIGHT_ADD','thm.arithmetic.SUB_EQ_|48|')).
fof(f4471,plain,(
  $false | ($spl0 | $spl9 | $spl20)),
  inference(subsumption_resolution,[],[f4470,f680])).
fof(f4470,plain,(
  s('type.num.num',sK6) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | ($spl0 | $spl20)),
  inference(forward_demodulation,[],[f4469,f444])).
fof(f444,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f175])).
fof(f175,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f174])).
fof(f174,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X9] : ! [X17,X18] : s(X9,X17) = s(X9,'const.bool.COND_3'(s(bool,'T_0'),s(X9,X17),s(X9,X18)))),
  file('thm.arithmetic.SUB_RIGHT_ADD','thm.bool.COND_CLAUSES|split|0')).
fof(f4469,plain,(
  s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | ($spl0 | $spl20)),
  inference(subsumption_resolution,[],[f4216,f671])).
fof(f4216,plain,(
  s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5)))) | $spl20),
  inference(superposition,[],[f704,f457])).
fof(f457,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f326])).
fof(f326,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f186])).
fof(f186,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X14] : (s(bool,'T_0') = s(bool,X14) <=> p(s(bool,X14)))),
  file('thm.arithmetic.SUB_RIGHT_ADD','thm.bool.EQ_CLAUSES|split|0')).
fof(f704,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))) | $spl20),
  inference(cnf_transformation,[],[f704_D])).
fof(f704_D,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f4468,plain,(
  $false | ($spl0 | $spl9 | $spl20)),
  inference(subsumption_resolution,[],[f4467,f680])).
fof(f4467,plain,(
  s('type.num.num',sK6) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | ($spl0 | $spl20)),
  inference(forward_demodulation,[],[f4466,f444])).
fof(f4466,plain,(
  s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | ($spl0 | $spl20)),
  inference(subsumption_resolution,[],[f4215,f671])).
fof(f4215,plain,(
  s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5)))) | $spl20),
  inference(superposition,[],[f704,f455])).
fof(f455,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f325])).
fof(f325,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f185])).
fof(f185,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X14] : (s(bool,'T_0') = s(bool,X14) <=> p(s(bool,X14)))),
  file('thm.arithmetic.SUB_RIGHT_ADD','thm.bool.EQ_CLAUSES|split|1')).
fof(f670,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5)))) | $spl1),
  inference(cnf_transformation,[],[f670_D])).
fof(f670_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f1495,plain,(
  s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) | $spl59),
  inference(cnf_transformation,[],[f1495_D])).
fof(f1495_D,plain,(
  s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) <=> ~$spl59),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl59])])).
fof(f3382,plain,(
  s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5)))) | $spl7),
  inference(forward_demodulation,[],[f3381,f460])).
fof(f460,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f188])).
fof(f188,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X6] : s('type.num.num',X6) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X6)))),
  file('thm.arithmetic.SUB_RIGHT_ADD','thm.arithmetic.ADD_CLAUSES|split|0')).
fof(f3381,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5)))) | $spl7),
  inference(forward_demodulation,[],[f1622,f433])).
fof(f1622,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5)))) | $spl7),
  inference(superposition,[],[f676,f417])).
fof(f671,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5)))) | $spl0),
  inference(cnf_transformation,[],[f671_D])).
fof(f671_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f3358,plain,(
  ( ! [X24] : (s('type.num.num','const.bool.COND_3'(s(bool,X24),s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | p(s(bool,X24))) ) | $spl84),
  inference(cnf_transformation,[],[f3358_D])).
fof(f3358_D,plain,(
  ( ! [X24] : (s('type.num.num','const.bool.COND_3'(s(bool,X24),s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | p(s(bool,X24))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f1911,plain,(
  ( ! [X9] : (s('type.num.num','const.bool.COND_3'(s(bool,X9),s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | p(s(bool,X9)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) ) | $spl21),
  inference(superposition,[],[f705,f420])).
fof(f420,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f263])).
fof(f263,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f262])).
fof(f262,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f153])).
fof(f153,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X17,X18] : ((p(s(bool,X17)) => p(s(bool,X18))) => ((p(s(bool,X18)) => p(s(bool,X17))) => s(bool,X17) = s(bool,X18)))),
  file('thm.arithmetic.SUB_RIGHT_ADD','thm.bool.IMP_ANTISYM_AX')).
fof(f1912,plain,(
  ( ! [X10] : (s('type.num.num','const.bool.COND_3'(s(bool,X10),s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5)))) | p(s(bool,X10))) ) | $spl21),
  inference(superposition,[],[f705,f420])).
fof(f3380,plain,(
  ( ! [X12] : (s('type.num.num','const.bool.COND_3'(s(bool,X12),s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | ~p(s(bool,X12))) ) | $spl100),
  inference(cnf_transformation,[],[f3380_D])).
fof(f3380_D,plain,(
  ( ! [X12] : (s('type.num.num','const.bool.COND_3'(s(bool,X12),s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | ~p(s(bool,X12))) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f1913,plain,(
  ( ! [X11] : (s('type.num.num','const.bool.COND_3'(s(bool,X11),s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5)))) | ~p(s(bool,X11))) ) | $spl21),
  inference(superposition,[],[f705,f423])).
fof(f423,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f263])).
fof(f1914,plain,(
  ( ! [X12] : (s('type.num.num','const.bool.COND_3'(s(bool,X12),s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | ~p(s(bool,X12)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) ) | $spl21),
  inference(superposition,[],[f705,f423])).
fof(f733,plain,(
  ( ! [X14,X15,X16] : (s('type.num.num','const.bool.COND_3'(s(bool,X14),s('type.num.num',X15),s('type.num.num',X16))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | s(bool,X14) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))) | s('type.num.num',sK6) != s('type.num.num',X15)) ) | $spl46),
  inference(cnf_transformation,[],[f733_D])).
fof(f733_D,plain,(
  ( ! [X14,X15,X16] : (s('type.num.num','const.bool.COND_3'(s(bool,X14),s('type.num.num',X15),s('type.num.num',X16))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | s(bool,X14) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))) | s('type.num.num',sK6) != s('type.num.num',X15)) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f1947,plain,(
  ( ! [X14,X15,X16] : (s('type.num.num','const.bool.COND_3'(s(bool,X14),s('type.num.num',X15),s('type.num.num',X16))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5)))) | s('type.num.num',sK6) != s('type.num.num',X15) | s(bool,X14) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5)))) ) | $spl21),
  inference(superposition,[],[f705,f435])).
fof(f435,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))) | ~p(s(bool,X2)) | s(X0,X3) != s(X0,X4) | s(bool,X1) != s(bool,X2)) )),
  inference(cnf_transformation,[],[f269])).
fof(f269,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : (s(bool,X1) != s(bool,X2) | (p(s(bool,X2)) & s(X0,X3) != s(X0,X4)) | (~p(s(bool,X2)) & s(X0,X5) != s(X0,X6)) | s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))))),
  inference(flattening,[],[f268])).
fof(f268,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ((s(bool,X1) != s(bool,X2) | (p(s(bool,X2)) & s(X0,X3) != s(X0,X4)) | (~p(s(bool,X2)) & s(X0,X5) != s(X0,X6))) | s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))))),
  inference(ennf_transformation,[],[f165])).
fof(f165,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ((s(bool,X1) = s(bool,X2) & (p(s(bool,X2)) => s(X0,X3) = s(X0,X4)) & (~p(s(bool,X2)) => s(X0,X5) = s(X0,X6))) => s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))))),
  inference(flattening,[],[f164])).
fof(f164,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6] : ((s(bool,X1) = s(bool,X2) & (p(s(bool,X2)) => s(X0,X3) = s(X0,X4)) & (~p(s(bool,X2)) => s(X0,X5) = s(X0,X6))) => s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))))),
  inference(rectify,[],[f71])).
fof(f71,axiom,(
  ! [X9] : ! [X16,X47,X15,X48,X40,X49] : ((s(bool,X16) = s(bool,X47) & (p(s(bool,X47)) => s(X9,X15) = s(X9,X48)) & (~p(s(bool,X47)) => s(X9,X40) = s(X9,X49))) => s(X9,'const.bool.COND_3'(s(bool,X16),s(X9,X15),s(X9,X40))) = s(X9,'const.bool.COND_3'(s(bool,X47),s(X9,X48),s(X9,X49))))),
  file('thm.arithmetic.SUB_RIGHT_ADD','thm.bool.COND_CONG')).
fof(f3375,plain,(
  ( ! [X21,X22,X20] : (s('type.num.num','const.bool.COND_3'(s(bool,X20),s('type.num.num',X21),s('type.num.num',X22))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | s(bool,X20) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X22)) ) | $spl98),
  inference(cnf_transformation,[],[f3375_D])).
fof(f3375_D,plain,(
  ( ! [X21,X22,X20] : (s('type.num.num','const.bool.COND_3'(s(bool,X20),s('type.num.num',X21),s('type.num.num',X22))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | s(bool,X20) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X22)) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f1949,plain,(
  ( ! [X21,X22,X20] : (s('type.num.num','const.bool.COND_3'(s(bool,X20),s('type.num.num',X21),s('type.num.num',X22))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X22) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5)))) | s(bool,X20) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5)))) ) | $spl21),
  inference(superposition,[],[f705,f436])).
fof(f436,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))) | s(X0,X5) != s(X0,X6) | p(s(bool,X2)) | s(bool,X1) != s(bool,X2)) )),
  inference(cnf_transformation,[],[f269])).
fof(f3372,plain,(
  ( ! [X10] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bool.COND_3'(s(bool,X10),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X10))) ) | $spl96),
  inference(cnf_transformation,[],[f3372_D])).
fof(f3372_D,plain,(
  ( ! [X10] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bool.COND_3'(s(bool,X10),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X10))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f2483,plain,(
  ( ! [X9] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bool.COND_3'(s(bool,X9),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X9)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) ) | $spl43),
  inference(superposition,[],[f730,f420])).
fof(f2484,plain,(
  ( ! [X10] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bool.COND_3'(s(bool,X10),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5)))) | p(s(bool,X10))) ) | $spl43),
  inference(superposition,[],[f730,f420])).
fof(f3371,plain,(
  ( ! [X12] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bool.COND_3'(s(bool,X12),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X12))) ) | $spl94),
  inference(cnf_transformation,[],[f3371_D])).
fof(f3371_D,plain,(
  ( ! [X12] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bool.COND_3'(s(bool,X12),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X12))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f2485,plain,(
  ( ! [X11] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bool.COND_3'(s(bool,X11),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5)))) | ~p(s(bool,X11))) ) | $spl43),
  inference(superposition,[],[f730,f423])).
fof(f2486,plain,(
  ( ! [X12] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bool.COND_3'(s(bool,X12),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X12)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) ) | $spl43),
  inference(superposition,[],[f730,f423])).
fof(f3366,plain,(
  ( ! [X14,X15,X16] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bool.COND_3'(s(bool,X14),s('type.num.num',X15),s('type.num.num',X16))),s('type.num.num','const.num.|48|_0')))) | s(bool,X14) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))) | s('type.num.num',sK6) != s('type.num.num',X15)) ) | $spl92),
  inference(cnf_transformation,[],[f3366_D])).
fof(f3366_D,plain,(
  ( ! [X14,X15,X16] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bool.COND_3'(s(bool,X14),s('type.num.num',X15),s('type.num.num',X16))),s('type.num.num','const.num.|48|_0')))) | s(bool,X14) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))) | s('type.num.num',sK6) != s('type.num.num',X15)) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f2551,plain,(
  ( ! [X14,X15,X16] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bool.COND_3'(s(bool,X14),s('type.num.num',X15),s('type.num.num',X16))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5)))) | s('type.num.num',sK6) != s('type.num.num',X15) | s(bool,X14) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5)))) ) | $spl43),
  inference(superposition,[],[f730,f435])).
fof(f3365,plain,(
  ( ! [X21,X22,X20] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bool.COND_3'(s(bool,X20),s('type.num.num',X21),s('type.num.num',X22))),s('type.num.num','const.num.|48|_0')))) | s(bool,X20) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5)))) ) | $spl90),
  inference(cnf_transformation,[],[f3365_D])).
fof(f3365_D,plain,(
  ( ! [X21,X22,X20] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bool.COND_3'(s(bool,X20),s('type.num.num',X21),s('type.num.num',X22))),s('type.num.num','const.num.|48|_0')))) | s(bool,X20) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5)))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f2553,plain,(
  ( ! [X21,X22,X20] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bool.COND_3'(s(bool,X20),s('type.num.num',X21),s('type.num.num',X22))),s('type.num.num','const.num.|48|_0')))) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5)))) | s(bool,X20) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5)))) ) | $spl43),
  inference(superposition,[],[f730,f436])).
fof(f3363,plain,(
  ( ! [X35] : (p(s(bool,X35))) ) | $spl86),
  inference(cnf_transformation,[],[f3363_D])).
fof(f3363_D,plain,(
  ( ! [X35] : (p(s(bool,X35))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f3364,plain,(
  ( ! [X37,X38,X36] : (s('type.num.num','const.bool.COND_3'(s(bool,X36),s('type.num.num',X37),s('type.num.num',X38))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | p(s(bool,X36)) | s('type.num.num',X38) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5)))) ) | $spl88),
  inference(cnf_transformation,[],[f3364_D])).
fof(f3364_D,plain,(
  ( ! [X37,X38,X36] : (s('type.num.num','const.bool.COND_3'(s(bool,X36),s('type.num.num',X37),s('type.num.num',X38))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | p(s(bool,X36)) | s('type.num.num',X38) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5)))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f3362,plain,(
  ( ! [X37,X35,X38,X36] : (s('type.num.num','const.bool.COND_3'(s(bool,X36),s('type.num.num',X37),s('type.num.num',X38))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | p(s(bool,X35)) | s('type.num.num',X38) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) | p(s(bool,X36))) ) | $spl2),
  inference(subsumption_resolution,[],[f3321,f420])).
fof(f3321,plain,(
  ( ! [X37,X35,X38,X36] : (s('type.num.num','const.bool.COND_3'(s(bool,X36),s('type.num.num',X37),s('type.num.num',X38))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | p(s(bool,X35)) | s('type.num.num',X38) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) | p(s(bool,X36)) | s(bool,X35) != s(bool,X36)) ) | $spl2),
  inference(superposition,[],[f672,f436])).
fof(f718,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5)))) | $spl31),
  inference(cnf_transformation,[],[f718_D])).
fof(f718_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5)))) <=> ~$spl31),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl31])])).
fof(f3294,plain,(
  ( ! [X1] : (s('type.num.num','const.bool.COND_3'(s(bool,X1),s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | p(s(bool,X1)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) ) | $spl2),
  inference(superposition,[],[f672,f417])).
fof(f703,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0')))) | $spl19),
  inference(cnf_transformation,[],[f703_D])).
fof(f703_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0')))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f3357,plain,(
  ( ! [X24] : (s('type.num.num','const.bool.COND_3'(s(bool,X24),s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | p(s(bool,X24)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'))))) ) | $spl2),
  inference(forward_demodulation,[],[f3285,f463])).
fof(f463,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f191])).
fof(f191,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X6] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X6)))),
  file('thm.arithmetic.SUB_RIGHT_ADD','thm.arithmetic.SUB|split|0')).
fof(f3285,plain,(
  ( ! [X24] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) != s('type.num.num','const.bool.COND_3'(s(bool,X24),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK5))))) | p(s(bool,X24)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'))))) ) | $spl2),
  inference(superposition,[],[f672,f448])).
fof(f714,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK5)))) | $spl27),
  inference(cnf_transformation,[],[f714_D])).
fof(f714_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK5)))) <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f2945,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))),s('type.num.num','const.num.|48|_0')))) | $spl83),
  inference(cnf_transformation,[],[f2945_D])).
fof(f2945_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))),s('type.num.num','const.num.|48|_0')))) <=> ~$spl83),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl83])])).
fof(f2900,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK5)))) | $spl41),
  inference(superposition,[],[f728,f424])).
fof(f424,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X1))),s('type.num.num',X0))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f264])).
fof(f264,plain,(
  ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ! [X2] : s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X1))),s('type.num.num',X0))))),
  inference(ennf_transformation,[],[f154])).
fof(f154,plain,(
  ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) => ! [X2] : s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X1))),s('type.num.num',X0))))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X31,X30] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X30),s('type.num.num',X31)))) => ! [X32] : s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X32),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X31),s('type.num.num',X30))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',X30))),s('type.num.num',X31))))),
  file('thm.arithmetic.SUB_RIGHT_ADD','thm.arithmetic.SUB_SUB')).
fof(f1970,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')))) | $spl73),
  inference(cnf_transformation,[],[f1970_D])).
fof(f1970_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')))) <=> ~$spl73),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl73])])).
fof(f2541,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))),s('type.num.num','const.num.|48|_0')))) | $spl43),
  inference(superposition,[],[f730,f448])).
fof(f2527,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5)))) | $spl43),
  inference(superposition,[],[f730,f417])).
fof(f2595,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))))),s('type.num.num','const.num.|48|_0')))) | $spl81),
  inference(cnf_transformation,[],[f2595_D])).
fof(f2595_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))))),s('type.num.num','const.num.|48|_0')))) <=> ~$spl81),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl81])])).
fof(f2526,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK5)))) | $spl43),
  inference(superposition,[],[f730,f424])).
fof(f2593,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0')))) | $spl43),
  inference(forward_demodulation,[],[f2518,f463])).
fof(f2518,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK5))))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0')))) | $spl43),
  inference(superposition,[],[f730,f448])).
fof(f2591,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))),s('type.num.num','const.num.|48|_0')))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))) | $spl43),
  inference(forward_demodulation,[],[f2494,f442])).
fof(f2494,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))),s('type.num.num','const.num.|48|_0')))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))) | $spl43),
  inference(superposition,[],[f730,f459])).
fof(f2589,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))),s('type.num.num','const.num.|48|_0')))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))) | $spl43),
  inference(forward_demodulation,[],[f2492,f442])).
fof(f2492,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))),s('type.num.num','const.num.|48|_0')))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))) | $spl43),
  inference(superposition,[],[f730,f458])).
fof(f1486,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK4)))) | $spl51),
  inference(cnf_transformation,[],[f1486_D])).
fof(f1486_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK4)))) <=> ~$spl51),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl51])])).
fof(f1661,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) | $spl71),
  inference(cnf_transformation,[],[f1661_D])).
fof(f1661_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) <=> ~$spl71),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl71])])).
fof(f2166,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK4)))) | $spl45),
  inference(forward_demodulation,[],[f2147,f433])).
fof(f2147,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num',sK4))),s('type.num.num',sK5))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK4)))) | $spl45),
  inference(superposition,[],[f732,f425])).
fof(f2161,plain,(
  ( ! [X3] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X3),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',X3)))) | s('type.num.num',sK4) != s('type.num.num',X3)) ) | $spl78),
  inference(cnf_transformation,[],[f2161_D])).
fof(f2161_D,plain,(
  ( ! [X3] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X3),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',X3)))) | s('type.num.num',sK4) != s('type.num.num',X3)) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f2138,plain,(
  ( ! [X4] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X4),s('type.num.num',sK5))))) | s('type.num.num',sK4) != s('type.num.num',X4) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',X4)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK4))))) ) | $spl45),
  inference(superposition,[],[f732,f405])).
fof(f405,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num',X0))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num',X0))) | s('type.num.num',X1) != s('type.num.num',X2) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X2)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f314])).
fof(f314,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X2)))) | ((s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num',X0))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num',X0))) | s('type.num.num',X1) = s('type.num.num',X2)) & (s('type.num.num',X1) != s('type.num.num',X2) | s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num',X0))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num',X0))))))),
  inference(nnf_transformation,[],[f258])).
fof(f258,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X2)))) | (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num',X0))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num',X0))) <=> s('type.num.num',X1) = s('type.num.num',X2)))),
  inference(flattening,[],[f257])).
fof(f257,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X2))))) | (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num',X0))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num',X0))) <=> s('type.num.num',X1) = s('type.num.num',X2)))),
  inference(ennf_transformation,[],[f138])).
fof(f138,plain,(
  ! [X0,X1,X2] : ((p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X2))))) => (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num',X0))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num',X0))) <=> s('type.num.num',X1) = s('type.num.num',X2)))),
  inference(rectify,[],[f82])).
fof(f82,axiom,(
  ! [X8,X7,X6] : ((p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X8),s('type.num.num',X7)))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X8),s('type.num.num',X6))))) => (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X6),s('type.num.num',X8))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X7),s('type.num.num',X8))) <=> s('type.num.num',X6) = s('type.num.num',X7)))),
  file('thm.arithmetic.SUB_RIGHT_ADD','thm.arithmetic.CANCEL_SUB')).
fof(f2137,plain,(
  ( ! [X3] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X3),s('type.num.num',sK5))))) | s('type.num.num',sK4) != s('type.num.num',X3) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK4)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',X3))))) ) | $spl45),
  inference(superposition,[],[f732,f405])).
fof(f2160,plain,(
  ( ! [X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X1))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',sK4)))) | s('type.num.num',sK5) != s('type.num.num',X1)) ) | $spl76),
  inference(cnf_transformation,[],[f2160_D])).
fof(f2160_D,plain,(
  ( ! [X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X1))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',sK4)))) | s('type.num.num',sK5) != s('type.num.num',X1)) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f2136,plain,(
  ( ! [X2] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X2))))) | s('type.num.num',sK5) != s('type.num.num',X2) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',sK4)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK4))))) ) | $spl45),
  inference(superposition,[],[f732,f407])).
fof(f407,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X2))) | s('type.num.num',X1) != s('type.num.num',X2) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f315])).
fof(f315,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X0)))) | ((s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X2))) | s('type.num.num',X1) = s('type.num.num',X2)) & (s('type.num.num',X1) != s('type.num.num',X2) | s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X2))))))),
  inference(nnf_transformation,[],[f260])).
fof(f260,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X0)))) | (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X2))) <=> s('type.num.num',X1) = s('type.num.num',X2)))),
  inference(flattening,[],[f259])).
fof(f259,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X0))))) | (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X2))) <=> s('type.num.num',X1) = s('type.num.num',X2)))),
  inference(ennf_transformation,[],[f139])).
fof(f139,plain,(
  ! [X0,X1,X2] : ((p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X0))))) => (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X2))) <=> s('type.num.num',X1) = s('type.num.num',X2)))),
  inference(rectify,[],[f83])).
fof(f83,axiom,(
  ! [X8,X7,X6] : ((p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',X8)))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num',X8))))) => (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X8),s('type.num.num',X6))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X8),s('type.num.num',X7))) <=> s('type.num.num',X6) = s('type.num.num',X7)))),
  file('thm.arithmetic.SUB_RIGHT_ADD','thm.arithmetic.SUB_CANCEL')).
fof(f2135,plain,(
  ( ! [X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X1))))) | s('type.num.num',sK5) != s('type.num.num',X1) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK4)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',sK4))))) ) | $spl45),
  inference(superposition,[],[f732,f407])).
fof(f2159,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num',X0))) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',sK5)))) ) | $spl74),
  inference(cnf_transformation,[],[f2159_D])).
fof(f2159_D,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num',X0))) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',sK5)))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f2134,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num',X0))) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',sK5))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK4))))) ) | $spl45),
  inference(superposition,[],[f732,f408])).
fof(f408,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num',X1))) | s('type.num.num',X2) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2))))) )),
  inference(cnf_transformation,[],[f316])).
fof(f316,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2)))) | ((s('type.num.num',X2) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num',X1)))) & (s('type.num.num',X0) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num',X1))) | s('type.num.num',X2) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(nnf_transformation,[],[f261])).
fof(f261,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2)))) | (s('type.num.num',X2) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) <=> s('type.num.num',X0) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num',X1)))))),
  inference(ennf_transformation,[],[f140])).
fof(f140,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2)))) => (s('type.num.num',X2) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) <=> s('type.num.num',X0) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num',X1)))))),
  inference(rectify,[],[f55])).
fof(f55,axiom,(
  ! [X6,X7,X8] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',X8)))) => (s('type.num.num',X8) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X7))) <=> s('type.num.num',X6) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X8),s('type.num.num',X7)))))),
  file('thm.arithmetic.SUB_RIGHT_ADD','thm.arithmetic.ADD_EQ_SUB')).
fof(f732,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | $spl45),
  inference(cnf_transformation,[],[f732_D])).
fof(f732_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) <=> ~$spl45),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl45])])).
fof(f1946,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')))) | $spl21),
  inference(superposition,[],[f705,f448])).
fof(f1962,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))) | $spl21),
  inference(forward_demodulation,[],[f1922,f442])).
fof(f1922,plain,(
  s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))) | $spl21),
  inference(superposition,[],[f705,f459])).
fof(f1960,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))) | $spl21),
  inference(forward_demodulation,[],[f1920,f442])).
fof(f1920,plain,(
  s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))) | $spl21),
  inference(superposition,[],[f705,f458])).
fof(f1659,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))),s('type.num.num','const.num.|48|_0')))) | $spl69),
  inference(cnf_transformation,[],[f1659_D])).
fof(f1659_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))),s('type.num.num','const.num.|48|_0')))) <=> ~$spl69),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl69])])).
fof(f1646,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))),s('type.num.num','const.num.|48|_0')))) | $spl7),
  inference(superposition,[],[f676,f448])).
fof(f1655,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num','const.num.|48|_0')))) | $spl67),
  inference(cnf_transformation,[],[f1655_D])).
fof(f1655_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num','const.num.|48|_0')))) <=> ~$spl67),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl67])])).
fof(f1653,plain,(
  s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num','const.num.|48|_0')))) | $spl7),
  inference(forward_demodulation,[],[f1652,f460])).
fof(f1652,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num','const.num.|48|_0')))) | $spl7),
  inference(forward_demodulation,[],[f1635,f433])).
fof(f1635,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num','const.num.|48|_0')))) | $spl7),
  inference(superposition,[],[f676,f448])).
fof(f1649,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X3),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',X3)))) | s('type.num.num',sK4) != s('type.num.num',X3)) ) | $spl64),
  inference(cnf_transformation,[],[f1649_D])).
fof(f1649_D,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X3),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',X3)))) | s('type.num.num',sK4) != s('type.num.num',X3)) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f1627,plain,(
  ( ! [X4] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X4),s('type.num.num',sK5))))) | s('type.num.num',sK4) != s('type.num.num',X4) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',X4)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK4))))) ) | $spl7),
  inference(superposition,[],[f676,f405])).
fof(f1626,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X3),s('type.num.num',sK5))))) | s('type.num.num',sK4) != s('type.num.num',X3) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK4)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',X3))))) ) | $spl7),
  inference(superposition,[],[f676,f405])).
fof(f1648,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X1))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',sK4)))) | s('type.num.num',sK5) != s('type.num.num',X1)) ) | $spl62),
  inference(cnf_transformation,[],[f1648_D])).
fof(f1648_D,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X1))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',sK4)))) | s('type.num.num',sK5) != s('type.num.num',X1)) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f1625,plain,(
  ( ! [X2] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X2))))) | s('type.num.num',sK5) != s('type.num.num',X2) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',sK4)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK4))))) ) | $spl7),
  inference(superposition,[],[f676,f407])).
fof(f1624,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X1))))) | s('type.num.num',sK5) != s('type.num.num',X1) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK4)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',sK4))))) ) | $spl7),
  inference(superposition,[],[f676,f407])).
fof(f1647,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num',X0))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',sK5)))) ) | $spl60),
  inference(cnf_transformation,[],[f1647_D])).
fof(f1647_D,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num',X0))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',sK5)))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f1623,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num',X0))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',sK5))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK4))))) ) | $spl7),
  inference(superposition,[],[f676,f408])).
fof(f1493,plain,(
  s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK4)))) | $spl9),
  inference(forward_demodulation,[],[f1474,f433])).
fof(f1474,plain,(
  s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num',sK4))),s('type.num.num',sK5))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK4)))) | $spl9),
  inference(superposition,[],[f680,f425])).
fof(f1489,plain,(
  ( ! [X3] : (s('type.num.num',sK6) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X3),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',X3)))) | s('type.num.num',sK4) != s('type.num.num',X3)) ) | $spl56),
  inference(cnf_transformation,[],[f1489_D])).
fof(f1489_D,plain,(
  ( ! [X3] : (s('type.num.num',sK6) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X3),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',X3)))) | s('type.num.num',sK4) != s('type.num.num',X3)) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f1465,plain,(
  ( ! [X4] : (s('type.num.num',sK6) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X4),s('type.num.num',sK5))))) | s('type.num.num',sK4) != s('type.num.num',X4) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',X4)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK4))))) ) | $spl9),
  inference(superposition,[],[f680,f405])).
fof(f1464,plain,(
  ( ! [X3] : (s('type.num.num',sK6) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X3),s('type.num.num',sK5))))) | s('type.num.num',sK4) != s('type.num.num',X3) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK4)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',X3))))) ) | $spl9),
  inference(superposition,[],[f680,f405])).
fof(f1488,plain,(
  ( ! [X1] : (s('type.num.num',sK6) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X1))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',sK4)))) | s('type.num.num',sK5) != s('type.num.num',X1)) ) | $spl54),
  inference(cnf_transformation,[],[f1488_D])).
fof(f1488_D,plain,(
  ( ! [X1] : (s('type.num.num',sK6) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X1))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',sK4)))) | s('type.num.num',sK5) != s('type.num.num',X1)) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f1463,plain,(
  ( ! [X2] : (s('type.num.num',sK6) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X2))))) | s('type.num.num',sK5) != s('type.num.num',X2) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',sK4)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK4))))) ) | $spl9),
  inference(superposition,[],[f680,f407])).
fof(f1462,plain,(
  ( ! [X1] : (s('type.num.num',sK6) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X1))))) | s('type.num.num',sK5) != s('type.num.num',X1) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK4)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',sK4))))) ) | $spl9),
  inference(superposition,[],[f680,f407])).
fof(f1487,plain,(
  ( ! [X0] : (s('type.num.num',sK6) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num',X0))) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',sK5)))) ) | $spl52),
  inference(cnf_transformation,[],[f1487_D])).
fof(f1487_D,plain,(
  ( ! [X0] : (s('type.num.num',sK6) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num',X0))) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',sK5)))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f1461,plain,(
  ( ! [X0] : (s('type.num.num',sK6) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num',X0))) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',sK5))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK4))))) ) | $spl9),
  inference(superposition,[],[f680,f408])).
fof(f710,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')))) | $spl23),
  inference(cnf_transformation,[],[f710_D])).
fof(f710_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f695,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0')))) | $spl15),
  inference(cnf_transformation,[],[f695_D])).
fof(f695_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0')))) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f801,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')))) | $spl1),
  inference(superposition,[],[f670,f448])).
fof(f734,plain,(
  ( ! [X21,X22,X20] : (s('type.num.num','const.bool.COND_3'(s(bool,X20),s('type.num.num',X21),s('type.num.num',X22))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | s(bool,X20) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5)))) ) | $spl48),
  inference(cnf_transformation,[],[f734_D])).
fof(f734_D,plain,(
  ( ! [X21,X22,X20] : (s('type.num.num','const.bool.COND_3'(s(bool,X20),s('type.num.num',X21),s('type.num.num',X22))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | s(bool,X20) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5)))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f666,plain,(
  ( ! [X21,X22,X20] : (s('type.num.num','const.bool.COND_3'(s(bool,X20),s('type.num.num',X21),s('type.num.num',X22))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5)))) | s(bool,X20) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5)))) )),
  inference(superposition,[],[f595,f436])).
fof(f664,plain,(
  ( ! [X14,X15,X16] : (s('type.num.num','const.bool.COND_3'(s(bool,X14),s('type.num.num',X15),s('type.num.num',X16))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5)))) | s('type.num.num',sK6) != s('type.num.num',X15) | s(bool,X14) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5)))) )),
  inference(superposition,[],[f595,f435])).
fof(f730,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))),s('type.num.num','const.num.|48|_0')))) | $spl43),
  inference(cnf_transformation,[],[f730_D])).
fof(f730_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))),s('type.num.num','const.num.|48|_0')))) <=> ~$spl43),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl43])])).
fof(f663,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))),s('type.num.num','const.num.|48|_0'))))),
  inference(superposition,[],[f595,f448])).
fof(f705,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))) | $spl21),
  inference(cnf_transformation,[],[f705_D])).
fof(f705_D,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f654,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))),s('type.num.num','const.num.|48|_0'))))),
  inference(superposition,[],[f595,f448])).
fof(f720,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) | $spl33),
  inference(cnf_transformation,[],[f720_D])).
fof(f720_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f723,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X3),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',X3)))) | s('type.num.num',X3) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))) ) | $spl38),
  inference(cnf_transformation,[],[f723_D])).
fof(f723_D,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X3),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',X3)))) | s('type.num.num',X3) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f645,plain,(
  ( ! [X4] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X4),s('type.num.num',sK5))))) | s('type.num.num',X4) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',X4)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))) )),
  inference(superposition,[],[f595,f405])).
fof(f644,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X3),s('type.num.num',sK5))))) | s('type.num.num',X3) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',X3))))) )),
  inference(superposition,[],[f595,f405])).
fof(f722,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',X1))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) | s('type.num.num',sK5) != s('type.num.num',X1)) ) | $spl36),
  inference(cnf_transformation,[],[f722_D])).
fof(f722_D,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',X1))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) | s('type.num.num',sK5) != s('type.num.num',X1)) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f643,plain,(
  ( ! [X2] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',X2))))) | s('type.num.num',sK5) != s('type.num.num',X2) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))) )),
  inference(superposition,[],[f595,f407])).
fof(f642,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',X1))))) | s('type.num.num',sK5) != s('type.num.num',X1) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))) )),
  inference(superposition,[],[f595,f407])).
fof(f721,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6),s('type.num.num',X0))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',sK5)))) ) | $spl34),
  inference(cnf_transformation,[],[f721_D])).
fof(f721_D,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6),s('type.num.num',X0))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',sK5)))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f641,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6),s('type.num.num',X0))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',sK5))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))) )),
  inference(superposition,[],[f595,f408])).
fof(f640,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))),
  inference(superposition,[],[f595,f417])).
fof(f716,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))))) | $spl29),
  inference(cnf_transformation,[],[f716_D])).
fof(f716_D,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f639,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK5))))),
  inference(superposition,[],[f595,f424])).
fof(f712,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK6),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | $spl25),
  inference(cnf_transformation,[],[f712_D])).
fof(f712_D,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK6),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f708,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK6),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))),
  inference(forward_demodulation,[],[f707,f433])).
fof(f707,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num',sK4))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK6),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))),
  inference(forward_demodulation,[],[f706,f462])).
fof(f462,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f190])).
fof(f190,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X6] : s('type.num.num',X6) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0')))),
  file('thm.arithmetic.SUB_RIGHT_ADD','thm.arithmetic.SUB_|48||split|1')).
fof(f706,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK6),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))),
  inference(forward_demodulation,[],[f638,f462])).
fof(f638,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))),
  inference(superposition,[],[f595,f448])).
fof(f701,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'))))),
  inference(forward_demodulation,[],[f631,f463])).
fof(f631,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'))))),
  inference(superposition,[],[f595,f448])).
fof(f697,plain,(
  s('type.num.num',sK6) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK5))),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) | $spl17),
  inference(cnf_transformation,[],[f697_D])).
fof(f697_D,plain,(
  s('type.num.num',sK6) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK5))),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f693,plain,(
  s('type.num.num',sK6) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK5))),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))))),
  inference(forward_demodulation,[],[f692,f460])).
fof(f692,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK5))),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))))),
  inference(forward_demodulation,[],[f691,f433])).
fof(f691,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK5))),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))))),
  inference(forward_demodulation,[],[f690,f463])).
fof(f690,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK5))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK5))),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))))),
  inference(forward_demodulation,[],[f614,f460])).
fof(f614,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK5))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK5))),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))))),
  inference(superposition,[],[f595,f448])).
fof(f676,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | $spl7),
  inference(cnf_transformation,[],[f676_D])).
fof(f676_D,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f689,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5)))),
  inference(forward_demodulation,[],[f607,f442])).
fof(f607,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) != s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5)))),
  inference(superposition,[],[f595,f459])).
fof(f684,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))) | $spl10),
  inference(cnf_transformation,[],[f684_D])).
fof(f684_D,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f680,plain,(
  s('type.num.num',sK6) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | $spl9),
  inference(cnf_transformation,[],[f680_D])).
fof(f680_D,plain,(
  s('type.num.num',sK6) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f688,plain,(
  s('type.num.num',sK6) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5)))),
  inference(forward_demodulation,[],[f606,f444])).
fof(f606,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) != s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5)))),
  inference(superposition,[],[f595,f459])).
fof(f685,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5)))),
  inference(forward_demodulation,[],[f605,f442])).
fof(f605,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) != s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5)))),
  inference(superposition,[],[f595,f458])).
fof(f682,plain,(
  s('type.num.num',sK6) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5)))),
  inference(forward_demodulation,[],[f604,f444])).
fof(f604,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) != s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5)))),
  inference(superposition,[],[f595,f458])).
fof(f681,plain,(
  s('type.num.num',sK6) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))))),
  inference(forward_demodulation,[],[f603,f444])).
fof(f603,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) != s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))))),
  inference(superposition,[],[f595,f457])).
fof(f678,plain,(
  s('type.num.num',sK6) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))))),
  inference(forward_demodulation,[],[f602,f444])).
fof(f602,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) != s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))))),
  inference(superposition,[],[f595,f455])).
fof(f677,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))))),
  inference(forward_demodulation,[],[f601,f442])).
fof(f601,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) != s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))))),
  inference(superposition,[],[f595,f453])).
fof(f674,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))))),
  inference(forward_demodulation,[],[f600,f442])).
fof(f600,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) != s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))))),
  inference(superposition,[],[f595,f451])).
fof(f673,plain,(
  ( ! [X11] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) != s('type.num.num','const.bool.COND_3'(s(bool,X11),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | ~p(s(bool,X11))) ) | $spl4),
  inference(cnf_transformation,[],[f673_D])).
fof(f673_D,plain,(
  ( ! [X11] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) != s('type.num.num','const.bool.COND_3'(s(bool,X11),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | ~p(s(bool,X11))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f599,plain,(
  ( ! [X12] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) != s('type.num.num','const.bool.COND_3'(s(bool,X12),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | ~p(s(bool,X12)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) )),
  inference(superposition,[],[f595,f423])).
fof(f598,plain,(
  ( ! [X11] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) != s('type.num.num','const.bool.COND_3'(s(bool,X11),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5)))) | ~p(s(bool,X11))) )),
  inference(superposition,[],[f595,f423])).
fof(f672,plain,(
  ( ! [X9] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) != s('type.num.num','const.bool.COND_3'(s(bool,X9),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | p(s(bool,X9))) ) | $spl2),
  inference(cnf_transformation,[],[f672_D])).
fof(f672_D,plain,(
  ( ! [X9] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) != s('type.num.num','const.bool.COND_3'(s(bool,X9),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | p(s(bool,X9))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f597,plain,(
  ( ! [X10] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) != s('type.num.num','const.bool.COND_3'(s(bool,X10),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5)))) | p(s(bool,X10))) )),
  inference(superposition,[],[f595,f420])).
fof(f596,plain,(
  ( ! [X9] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) != s('type.num.num','const.bool.COND_3'(s(bool,X9),s('type.num.num',sK6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | p(s(bool,X9)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) )),
  inference(superposition,[],[f595,f420])).
% SZS output end Proof for thm.arithmetic
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 228
% Passive clauses: 2817
% Generated clauses: 5222
% Final active clauses: 210
% Final passive clauses: 2018
% Input formulas: 133
% Initial clauses: 211
% 
% Selected by SInE selection: 99
% SInE iterations: 1
% 
% Duplicate literals: 69
% Trivial inequalities: 6
% Fw subsumption resolutions: 381
% Fw demodulations: 482
% 
% Simple tautologies: 124
% Equational tautologies: 2
% Forward subsumptions: 1148
% 
% Binary resolution: 1751
% Forward superposition: 2119
% Backward superposition: 213
% 
% Split clauses: 75
% Split components: 150
% SAT solver clauses: 103
% SAT solver unit clauses: 9
% SAT solver binary clauses: 40
% 
% Sat splits: 75
% Sat splitting refutations: 27
% 
% Memory used [KB]: 3582
% Time elapsed: 0.133 s
% ------------------------------
% ------------------------------
% Success in time 11.895 s
