Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.pred_set
% (113057)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 131
% Passive clauses: 3265
% Generated clauses: 19982
% Final active clauses: 119
% Final passive clauses: 2890
% Input formulas: 20
% Initial clauses: 41
% 
% Duplicate literals: 808
% Trivial inequalities: 380
% Fw subsumption resolutions: 2940
% Fw demodulations: 3327
% Bw demodulations: 10
% 
% Simple tautologies: 402
% Equational tautologies: 48
% Forward subsumptions: 7994
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 529
% Factoring: 3
% Forward superposition: 5941
% Backward superposition: 5563
% Self superposition: 170
% Equality factoring: 111
% Equality resolution: 33
% 
% Split clauses: 450
% Split components: 918
% SAT solver clauses: 986
% SAT solver unit clauses: 23
% SAT solver binary clauses: 190
% 
% Sat splits: 450
% 
% TWLsolver clauses: 981
% TWLsolver calls for satisfiability: 139
% 
% Memory used [KB]: 10234
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pred_set
% (113084)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 110
% Passive clauses: 1129
% Generated clauses: 10726
% Final active clauses: 85
% Final passive clauses: 792
% Input formulas: 20
% Initial clauses: 41
% 
% Duplicate literals: 508
% Trivial inequalities: 107
% Fw subsumption resolutions: 784
% Fw demodulations: 883
% Bw demodulations: 202
% Condensations: 808
% Global subsumptions: 4
% 
% Simple tautologies: 183
% Equational tautologies: 40
% Forward subsumptions: 4626
% Bw demodulations to eq. taut.: 43
% 
% Binary resolution: 501
% Factoring: 3
% Forward superposition: 3668
% Backward superposition: 2018
% Self superposition: 125
% Equality factoring: 60
% Equality resolution: 21
% 
% Split clauses: 43
% Split components: 95
% SAT solver clauses: 1698
% SAT solver unit clauses: 45
% SAT solver binary clauses: 42
% SAT solver learnt clauses: 8
% 
% Sat splits: 43
% Sat splitting refutations: 8
% 
% TWLsolver clauses: 1682
% TWLsolver calls for satisfiability: 241
% 
% Memory used [KB]: 10106
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.pred_set
% (113115)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 458
% Passive clauses: 7086
% Generated clauses: 123912
% Final active clauses: 45
% Final passive clauses: 207
% Input formulas: 20
% Initial clauses: 41
% 
% Duplicate literals: 6363
% Trivial inequalities: 960
% Fw subsumption resolutions: 9070
% Fw demodulations: 17445
% Bw demodulations: 5574
% Condensations: 11876
% Global subsumptions: 2086
% 
% Simple tautologies: 1089
% Equational tautologies: 2959
% Forward subsumptions: 62228
% Fw demodulations to eq. taut.: 90
% Bw demodulations to eq. taut.: 89
% 
% Unit resulting resolution: 11892
% Factoring: 9
% Forward superposition: 26355
% Backward superposition: 19142
% Self superposition: 743
% Equality factoring: 489
% Equality resolution: 83
% 
% Split clauses: 304
% Split components: 617
% SAT solver clauses: 12011
% SAT solver unit clauses: 101
% SAT solver binary clauses: 1600
% 
% Sat splits: 304
% Sat splitting refutations: 2229
% 
% Memory used [KB]: 59487
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f4219,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f194,f195,f196,f197,f198,f199,f200,f201,f202,f203,f262,f266,f264,f255,f256,f268,f350,f351,f352,f348,f346,f625,f626,f637,f627,f628,f638,f629,f630,f636,f631,f632,f635,f633,f642,f644,f634,f640,f646,f693,f696,f695,f837,f838,f839,f840,f841,f842,f843,f844,f845,f846,f981,f982,f993,f983,f984,f994,f985,f986,f987,f988,f991,f989,f998,f1000,f990,f996,f1002,f1041,f1042,f1043,f1055,f1057,f1090,f1091,f1141,f213,f1142,f1457,f1458,f1459,f1508,f1510,f1511,f1513,f1133,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f211,f1550,f1551,f1552,f1572,f1573,f263,f1580,f1581,f206,f1582,f1583,f207,f204,f1584,f1585,f205,f1589,f209,f215,f1911,f1912,f1913,f1957,f1960,f1959,f2095,f2098,f2097,f2124,f2125,f2267,f2268,f992,f2323,f2324,f2325,f2326,f2327,f2070,f2068,f1958,f2336,f2372,f2373,f2374,f2395,f2396,f2397,f2398,f2399,f2400,f2401,f2402,f2403,f2404,f2573,f2574,f353,f2616,f2617,f2618,f2619,f2666,f2667,f2668,f2669,f2739,f2781,f2782,f849,f2783,f2784,f850,f2785,f2786,f2787,f2788,f847,f2789,f854,f856,f2790,f852,f858,f2930,f2931,f265,f848,f3378,f4209,f4213])).
fof(f4213,plain,(
  $false | ($spl38 | $spl44 | $spl50)),
  inference(subsumption_resolution,[],[f4212,f111])).
fof(f111,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.IMAGE_ID',aHL_TRUTH)).
fof(f4212,plain,(
  ~p(s(bool,'T_0')) | ($spl38 | $spl44 | $spl50)),
  inference(forward_demodulation,[],[f4211,f858])).
fof(f4211,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)))))) | ($spl38 | $spl44)),
  inference(subsumption_resolution,[],[f4210,f109])).
fof(f109,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f42])).
fof(f42,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.pred_set.IMAGE_ID',aHL_FALSITY)).
fof(f4210,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)))))) | ($spl38 | $spl44)),
  inference(forward_demodulation,[],[f4197,f851])).
fof(f851,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl44),
  inference(cnf_transformation,[],[f851_D])).
fof(f851_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f4197,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)))))) | $spl38),
  inference(superposition,[],[f91,f4169])).
fof(f4169,plain,(
  s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)) = s(sK0,sK2('_i_0_0',sK1,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0),sK0,sK0)) | $spl38),
  inference(superposition,[],[f2947,f93])).
fof(f93,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(flattening,[],[f35])).
fof(f35,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X15] : ! [X16] : s(X15,X16) = s(X15,happ(s(fun(X15,X15),'_i_0_0'),s(X15,X16)))),
  file('thm.pred_set.IMAGE_ID','_i0_ax')).
fof(f2947,plain,(
  s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)) = s(sK0,happ(s(fun(sK0,sK0),'_i_0_0'),s(sK0,sK2('_i_0_0',sK1,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0),sK0,sK0)))) | $spl38),
  inference(resolution,[],[f848,f90])).
fof(f90,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,X2) = s(X0,happ(s(fun(X1,X0),X4),s(X1,sK2(X4,X3,X2,X1,X0)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X4),s(fun(X1,bool),X3))))))) )),
  inference(cnf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X4),s(fun(X1,bool),X3)))))) | (s(X0,X2) = s(X0,happ(s(fun(X1,X0),X4),s(X1,sK2(X4,X3,X2,X1,X0)))) & p(s(bool,'const.bool.IN_2'(s(X1,sK2(X4,X3,X2,X1,X0)),s(fun(X1,bool),X3)))))) & (! [X5] : (s(X0,X2) != s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3))))) | p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X4),s(fun(X1,bool),X3))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2])],[f58])).
fof(f58,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X4),s(fun(X1,bool),X3)))))) | ? [X6] : (s(X0,X2) = s(X0,happ(s(fun(X1,X0),X4),s(X1,X6))) & p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3)))))) & (! [X5] : (s(X0,X2) != s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3))))) | p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X4),s(fun(X1,bool),X3))))))))),
  inference(rectify,[],[f57])).
fof(f57,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X4),s(fun(X1,bool),X3)))))) | ? [X5] : (s(X0,X2) = s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))) & p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))))) & (! [X5] : (s(X0,X2) != s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3))))) | p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X4),s(fun(X1,bool),X3))))))))),
  inference(nnf_transformation,[],[f34])).
fof(f34,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X4),s(fun(X1,bool),X3)))))) <=> ? [X5] : (s(X0,X2) = s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))) & p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3))))))),
  inference(flattening,[],[f33])).
fof(f33,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X4),s(fun(X1,bool),X3)))))) <=> ? [X5] : (s(X0,X2) = s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))) & p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3))))))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X12,X6] : ! [X13,X7,X14] : (p(s(bool,'const.bool.IN_2'(s(X12,X13),s(fun(X12,bool),'const.pred_set.IMAGE_2'(s(fun(X6,X12),X14),s(fun(X6,bool),X7)))))) <=> ? [X11] : (s(X12,X13) = s(X12,happ(s(fun(X6,X12),X14),s(X6,X11))) & p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X7))))))),
  file('thm.pred_set.IMAGE_ID','thm.pred_set.IN_IMAGE')).
fof(f91,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X1,sK2(X4,X3,X2,X1,X0)),s(fun(X1,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X4),s(fun(X1,bool),X3))))))) )),
  inference(cnf_transformation,[],[f59])).
fof(f4209,plain,(
  $false | ($spl38 | $spl47)),
  inference(subsumption_resolution,[],[f4189,f854])).
fof(f4189,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl38),
  inference(backward_demodulation,[],[f4169,f2990])).
fof(f2990,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('_i_0_0',sK1,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0),sK0,sK0)),s(fun(sK0,bool),sK1))) | $spl38),
  inference(resolution,[],[f2948,f86])).
fof(f86,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X10] : (s(bool,'T_0') = s(bool,X10) <=> p(s(bool,X10)))),
  file('thm.pred_set.IMAGE_ID','thm.bool.EQ_CLAUSES|split|1')).
fof(f2948,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2('_i_0_0',sK1,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0),sK0,sK0)),s(fun(sK0,bool),sK1)))) | $spl38),
  inference(resolution,[],[f848,f91])).
fof(f3378,plain,(
  $false | ($spl44 | $spl48)),
  inference(subsumption_resolution,[],[f856,f2791])).
fof(f2791,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1))))) | $spl44),
  inference(backward_demodulation,[],[f851,f393])).
fof(f393,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)))))),
  inference(equality_resolution,[],[f125])).
fof(f125,plain,(
  ( ! [X2] : (s(fun(sK0,bool),sK1) != s(fun(sK0,bool),X2) | s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),X2,sK0)),s(fun(sK0,bool),X2))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),X2,sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)))))) )),
  inference(superposition,[],[f74,f114])).
fof(f114,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X2)))) )),
  inference(cnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f72])).
fof(f72,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(rectify,[],[f71])).
fof(f71,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(nnf_transformation,[],[f46])).
fof(f46,plain,(
  ! [X0,X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(flattening,[],[f45])).
fof(f45,plain,(
  ! [X0] : ! [X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X6] : ! [X7,X10] : (s(fun(X6,bool),X7) = s(fun(X6,bool),X10) <=> ! [X11] : s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X7))) = s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X10))))),
  file('thm.pred_set.IMAGE_ID','thm.pred_set.EXTENSION')).
fof(f74,plain,(
  s(fun(sK0,bool),sK1) != s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)))),
  inference(cnf_transformation,[],[f51])).
fof(f51,plain,(
  s(fun(sK0,bool),sK1) != s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f47])).
fof(f47,plain,(
  ? [X0,X1] : s(fun(X0,bool),X1) != s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,bool),X1)))),
  inference(ennf_transformation,[],[f23])).
fof(f23,plain,(
  ~! [X0,X1] : s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,bool),X1)))),
  inference(flattening,[],[f22])).
fof(f22,plain,(
  ~! [X0] : ! [X1] : s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,bool),X1)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7] : s(fun(X6,bool),X7) = s(fun(X6,bool),'const.pred_set.IMAGE_2'(s(fun(X6,X6),'_i_0_0'),s(fun(X6,bool),X7)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7] : s(fun(X6,bool),X7) = s(fun(X6,bool),'const.pred_set.IMAGE_2'(s(fun(X6,X6),'_i_0_0'),s(fun(X6,bool),X7)))),
  file('thm.pred_set.IMAGE_ID','thm.pred_set.IMAGE_ID')).
fof(f848,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)))))) | $spl38),
  inference(cnf_transformation,[],[f848_D])).
fof(f848_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)))))) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f265,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | $spl16),
  inference(cnf_transformation,[],[f265_D])).
fof(f265_D,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f2931,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,X3) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1))))))) ) | $spl44),
  inference(subsumption_resolution,[],[f2918,f81])).
fof(f81,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f28])).
fof(f28,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f27])).
fof(f27,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X10] : (s(bool,'F_0') = s(bool,X10) <=> ~p(s(bool,X10)))),
  file('thm.pred_set.IMAGE_ID','thm.bool.EQ_CLAUSES|split|2')).
fof(f2918,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,X3) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)))))) | p(s(bool,X3))) ) | $spl44),
  inference(superposition,[],[f2791,f75])).
fof(f75,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f48])).
fof(f48,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f24])).
fof(f24,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X8,X9] : ((p(s(bool,X8)) => p(s(bool,X9))) => ((p(s(bool,X9)) => p(s(bool,X8))) => s(bool,X8) = s(bool,X9)))),
  file('thm.pred_set.IMAGE_ID','thm.bool.IMP_ANTISYM_AX')).
fof(f2930,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1))))))) ) | $spl44),
  inference(subsumption_resolution,[],[f2917,f81])).
fof(f2917,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | p(s(bool,X2)) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1))))))) ) | $spl44),
  inference(superposition,[],[f2791,f75])).
fof(f858,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1))))) | $spl50),
  inference(cnf_transformation,[],[f858_D])).
fof(f858_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1))))) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f852,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl45),
  inference(cnf_transformation,[],[f852_D])).
fof(f852_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl45),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl45])])).
fof(f2790,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)))))),
  inference(superposition,[],[f393,f89])).
fof(f89,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.pred_set.IMAGE_ID',aHL_BOOL_CASES)).
fof(f856,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1))))) | $spl48),
  inference(cnf_transformation,[],[f856_D])).
fof(f856_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1))))) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f854,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl47),
  inference(cnf_transformation,[],[f854_D])).
fof(f854_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl47),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl47])])).
fof(f2789,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)))))),
  inference(superposition,[],[f393,f89])).
fof(f847,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)))))) | $spl39),
  inference(cnf_transformation,[],[f847_D])).
fof(f847_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)))))) <=> ~$spl39),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl39])])).
fof(f2788,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1))))))),
  inference(superposition,[],[f393,f88])).
fof(f88,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f56])).
fof(f56,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f32])).
fof(f32,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X10] : (s(bool,'T_0') = s(bool,X10) <=> p(s(bool,X10)))),
  file('thm.pred_set.IMAGE_ID','thm.bool.EQ_CLAUSES|split|0')).
fof(f2787,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1))))))),
  inference(superposition,[],[f393,f86])).
fof(f2786,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1))))))),
  inference(superposition,[],[f393,f84])).
fof(f84,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f30])).
fof(f30,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f29])).
fof(f29,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X10] : (s(bool,'F_0') = s(bool,X10) <=> ~p(s(bool,X10)))),
  file('thm.pred_set.IMAGE_ID','thm.bool.EQ_CLAUSES|split|3')).
fof(f2785,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1))))))),
  inference(superposition,[],[f393,f82])).
fof(f82,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f53])).
fof(f850,plain,(
  ( ! [X4] : (s(bool,X4) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | ~p(s(bool,X4))) ) | $spl42),
  inference(cnf_transformation,[],[f850_D])).
fof(f850_D,plain,(
  ( ! [X4] : (s(bool,X4) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | ~p(s(bool,X4))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f2784,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | ~p(s(bool,X5)) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1))))))) )),
  inference(superposition,[],[f393,f78])).
fof(f78,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f49])).
fof(f2783,plain,(
  ( ! [X4] : (s(bool,X4) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)))))) | ~p(s(bool,X4))) )),
  inference(superposition,[],[f393,f78])).
fof(f849,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | p(s(bool,X2))) ) | $spl40),
  inference(cnf_transformation,[],[f849_D])).
fof(f849_D,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | p(s(bool,X2))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f2782,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)))))) | p(s(bool,X3))) )),
  inference(superposition,[],[f393,f75])).
fof(f2781,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | p(s(bool,X2)) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1))))))) )),
  inference(superposition,[],[f393,f75])).
fof(f2739,plain,(
  $false | ($spl39 | $spl46)),
  inference(subsumption_resolution,[],[f2738,f111])).
fof(f2738,plain,(
  ~p(s(bool,'T_0')) | ($spl39 | $spl46)),
  inference(forward_demodulation,[],[f2736,f853])).
fof(f853,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl46),
  inference(cnf_transformation,[],[f853_D])).
fof(f853_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f2736,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1)))) | $spl39),
  inference(equality_resolution,[],[f2726])).
fof(f2726,plain,(
  ( ! [X0] : (s(sK0,X0) != s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK1))))) ) | $spl39),
  inference(forward_demodulation,[],[f2693,f93])).
fof(f2693,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK1)))) | s(sK0,happ(s(fun(sK0,sK0),'_i_0_0'),s(sK0,X0))) != s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0))) ) | $spl39),
  inference(resolution,[],[f847,f92])).
fof(f92,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X4),s(fun(X1,bool),X3)))))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) | s(X0,X2) != s(X0,happ(s(fun(X1,X0),X4),s(X1,X5)))) )),
  inference(cnf_transformation,[],[f59])).
fof(f2669,plain,(
  $false | ($spl28 | $spl54 | $spl58)),
  inference(trivial_inequality_removal,[],[f2656])).
fof(f2656,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl28 | $spl54 | $spl58)),
  inference(superposition,[],[f2629,f639])).
fof(f639,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | $spl28),
  inference(cnf_transformation,[],[f639_D])).
fof(f639_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f2629,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2)) ) | ($spl54 | $spl58)),
  inference(subsumption_resolution,[],[f2628,f83])).
fof(f83,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f54])).
fof(f2628,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | p(s(bool,X2))) ) | ($spl54 | $spl58)),
  inference(forward_demodulation,[],[f993,f995])).
fof(f995,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | $spl58),
  inference(cnf_transformation,[],[f995_D])).
fof(f995_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f2668,plain,(
  $false | ($spl6 | $spl54 | $spl58)),
  inference(trivial_inequality_removal,[],[f2657])).
fof(f2657,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl6 | $spl54 | $spl58)),
  inference(superposition,[],[f2629,f208])).
fof(f208,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | $spl6),
  inference(cnf_transformation,[],[f208_D])).
fof(f208_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f2667,plain,(
  $false | ($spl54 | $spl58)),
  inference(trivial_inequality_removal,[],[f2662])).
fof(f2662,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl54 | $spl58)),
  inference(superposition,[],[f2629,f995])).
fof(f2666,plain,(
  $false | ($spl54 | $spl58)),
  inference(equality_resolution,[],[f2629])).
fof(f2619,plain,(
  $false | ($spl34 | $spl42 | $spl46)),
  inference(trivial_inequality_removal,[],[f2608])).
fof(f2608,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl34 | $spl42 | $spl46)),
  inference(superposition,[],[f2579,f646])).
fof(f2579,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,X4)) ) | ($spl42 | $spl46)),
  inference(subsumption_resolution,[],[f2578,f87])).
fof(f87,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f56])).
fof(f2578,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,X4) | ~p(s(bool,X4))) ) | ($spl42 | $spl46)),
  inference(forward_demodulation,[],[f850,f853])).
fof(f2618,plain,(
  $false | ($spl12 | $spl42 | $spl46)),
  inference(trivial_inequality_removal,[],[f2609])).
fof(f2609,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl12 | $spl42 | $spl46)),
  inference(superposition,[],[f2579,f215])).
fof(f2617,plain,(
  $false | ($spl42 | $spl46)),
  inference(trivial_inequality_removal,[],[f2613])).
fof(f2613,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl42 | $spl46)),
  inference(superposition,[],[f2579,f853])).
fof(f2616,plain,(
  $false | ($spl42 | $spl46)),
  inference(equality_resolution,[],[f2579])).
fof(f353,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3)) ) | $spl20),
  inference(cnf_transformation,[],[f353_D])).
fof(f353_D,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3)) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f2574,plain,(
  ( ! [X5] : (s(bool,'T_0') != s(bool,X5) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1))))))) ) | $spl46),
  inference(subsumption_resolution,[],[f2563,f85])).
fof(f85,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f55])).
fof(f2563,plain,(
  ( ! [X5] : (s(bool,'T_0') != s(bool,X5) | ~p(s(bool,X5)) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1))))))) ) | $spl46),
  inference(superposition,[],[f2460,f78])).
fof(f2460,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1))))) | $spl46),
  inference(backward_demodulation,[],[f853,f393])).
fof(f2573,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,X4) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1))))))) ) | $spl46),
  inference(subsumption_resolution,[],[f2562,f85])).
fof(f2562,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,X4) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)))))) | ~p(s(bool,X4))) ) | $spl46),
  inference(superposition,[],[f2460,f78])).
fof(f2404,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)))))),
  inference(superposition,[],[f393,f89])).
fof(f2403,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)))))),
  inference(superposition,[],[f393,f89])).
fof(f2402,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1))))))),
  inference(superposition,[],[f393,f88])).
fof(f2401,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1))))))),
  inference(superposition,[],[f393,f86])).
fof(f2400,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1))))))),
  inference(superposition,[],[f393,f84])).
fof(f2399,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1))))))),
  inference(superposition,[],[f393,f82])).
fof(f2398,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | ~p(s(bool,X5)) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1))))))) )),
  inference(superposition,[],[f393,f78])).
fof(f2397,plain,(
  ( ! [X4] : (s(bool,X4) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)))))) | ~p(s(bool,X4))) )),
  inference(superposition,[],[f393,f78])).
fof(f2396,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)))))) | p(s(bool,X3))) )),
  inference(superposition,[],[f393,f75])).
fof(f2395,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | p(s(bool,X2)) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1))))))) )),
  inference(superposition,[],[f393,f75])).
fof(f2374,plain,(
  $false | ($spl20 | $spl34)),
  inference(trivial_inequality_removal,[],[f2365])).
fof(f2365,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl20 | $spl34)),
  inference(superposition,[],[f353,f646])).
fof(f2373,plain,(
  $false | ($spl12 | $spl20)),
  inference(trivial_inequality_removal,[],[f2366])).
fof(f2366,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl12 | $spl20)),
  inference(superposition,[],[f353,f215])).
fof(f2372,plain,(
  $false | $spl20),
  inference(equality_resolution,[],[f353])).
fof(f2336,plain,(
  $false | ($spl45 | $spl47)),
  inference(subsumption_resolution,[],[f1954,f854])).
fof(f1954,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl45),
  inference(trivial_inequality_removal,[],[f1953])).
fof(f1953,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl45),
  inference(superposition,[],[f852,f89])).
fof(f1958,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1)))) | $spl67),
  inference(cnf_transformation,[],[f1958_D])).
fof(f1958_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1)))) <=> ~$spl67),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl67])])).
fof(f2068,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,X4) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))))) ) | $spl47),
  inference(subsumption_resolution,[],[f2057,f85])).
fof(f2057,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,X4) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1)))) | ~p(s(bool,X4))) ) | $spl47),
  inference(superposition,[],[f854,f78])).
fof(f2070,plain,(
  ( ! [X5] : (s(bool,'T_0') != s(bool,X5) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))))) ) | $spl47),
  inference(subsumption_resolution,[],[f2058,f85])).
fof(f2058,plain,(
  ( ! [X5] : (s(bool,'T_0') != s(bool,X5) | ~p(s(bool,X5)) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))))) ) | $spl47),
  inference(superposition,[],[f854,f78])).
fof(f2327,plain,(
  $false | ($spl28 | $spl40 | $spl44)),
  inference(trivial_inequality_removal,[],[f2313])).
fof(f2313,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl28 | $spl40 | $spl44)),
  inference(superposition,[],[f2286,f639])).
fof(f2286,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2)) ) | ($spl40 | $spl44)),
  inference(subsumption_resolution,[],[f2285,f83])).
fof(f2285,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | p(s(bool,X2))) ) | ($spl40 | $spl44)),
  inference(forward_demodulation,[],[f849,f851])).
fof(f2326,plain,(
  $false | ($spl6 | $spl40 | $spl44)),
  inference(trivial_inequality_removal,[],[f2314])).
fof(f2314,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl6 | $spl40 | $spl44)),
  inference(superposition,[],[f2286,f208])).
fof(f2325,plain,(
  $false | ($spl40 | $spl44)),
  inference(trivial_inequality_removal,[],[f2319])).
fof(f2319,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl40 | $spl44)),
  inference(superposition,[],[f2286,f851])).
fof(f2324,plain,(
  $false | ($spl40 | $spl44 | $spl58)),
  inference(trivial_inequality_removal,[],[f2320])).
fof(f2320,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl40 | $spl44 | $spl58)),
  inference(superposition,[],[f2286,f995])).
fof(f2323,plain,(
  $false | ($spl40 | $spl44)),
  inference(equality_resolution,[],[f2286])).
fof(f992,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)))))) | $spl52),
  inference(cnf_transformation,[],[f992_D])).
fof(f992_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)))))) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f2268,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,X3) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1))))))) ) | $spl58),
  inference(subsumption_resolution,[],[f2255,f81])).
fof(f2255,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,X3) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)))))) | p(s(bool,X3))) ) | $spl58),
  inference(superposition,[],[f2126,f75])).
fof(f2126,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1))))) | $spl58),
  inference(backward_demodulation,[],[f995,f531])).
fof(f531,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)))))),
  inference(equality_resolution,[],[f126])).
fof(f126,plain,(
  ( ! [X3] : (s(fun(sK0,bool),sK1) != s(fun(sK0,bool),X3) | s(bool,'const.bool.IN_2'(s(sK0,sK4(X3,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),X3))) != s(bool,'const.bool.IN_2'(s(sK0,sK4(X3,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)))))) )),
  inference(superposition,[],[f74,f114])).
fof(f2267,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1))))))) ) | $spl58),
  inference(subsumption_resolution,[],[f2254,f81])).
fof(f2254,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | p(s(bool,X2)) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1))))))) ) | $spl58),
  inference(superposition,[],[f2126,f75])).
fof(f2125,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,X3) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1))))))) ) | $spl44),
  inference(subsumption_resolution,[],[f2112,f81])).
fof(f2112,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,X3) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)))))) | p(s(bool,X3))) ) | $spl44),
  inference(superposition,[],[f1961,f75])).
fof(f1961,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1))))) | $spl44),
  inference(backward_demodulation,[],[f851,f393])).
fof(f2124,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1))))))) ) | $spl44),
  inference(subsumption_resolution,[],[f2111,f81])).
fof(f2111,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | p(s(bool,X2)) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1))))))) ) | $spl44),
  inference(superposition,[],[f1961,f75])).
fof(f2097,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1)))) | $spl68),
  inference(cnf_transformation,[],[f2097_D])).
fof(f2097_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1)))) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f2098,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,X3) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))))) ) | $spl59),
  inference(subsumption_resolution,[],[f2083,f81])).
fof(f2083,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,X3) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1)))) | p(s(bool,X3))) ) | $spl59),
  inference(superposition,[],[f996,f75])).
fof(f2095,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))))) ) | $spl59),
  inference(subsumption_resolution,[],[f2082,f81])).
fof(f2082,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | p(s(bool,X2)) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))))) ) | $spl59),
  inference(superposition,[],[f996,f75])).
fof(f1959,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1)))) | $spl66),
  inference(cnf_transformation,[],[f1959_D])).
fof(f1959_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1)))) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f1960,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,X3) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))))) ) | $spl45),
  inference(subsumption_resolution,[],[f1945,f81])).
fof(f1945,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,X3) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1)))) | p(s(bool,X3))) ) | $spl45),
  inference(superposition,[],[f852,f75])).
fof(f1957,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))))) ) | $spl45),
  inference(subsumption_resolution,[],[f1944,f81])).
fof(f1944,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | p(s(bool,X2)) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))))) ) | $spl45),
  inference(superposition,[],[f852,f75])).
fof(f1913,plain,(
  $false | ($spl24 | $spl28)),
  inference(trivial_inequality_removal,[],[f1903])).
fof(f1903,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl24 | $spl28)),
  inference(superposition,[],[f1876,f639])).
fof(f1876,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | ($spl24 | $spl28)),
  inference(subsumption_resolution,[],[f1875,f83])).
fof(f1875,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,X0))) ) | ($spl24 | $spl28)),
  inference(forward_demodulation,[],[f637,f639])).
fof(f1912,plain,(
  $false | ($spl6 | $spl24 | $spl28)),
  inference(trivial_inequality_removal,[],[f1904])).
fof(f1904,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl6 | $spl24 | $spl28)),
  inference(superposition,[],[f1876,f208])).
fof(f1911,plain,(
  $false | ($spl24 | $spl28)),
  inference(equality_resolution,[],[f1876])).
fof(f215,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | $spl12),
  inference(cnf_transformation,[],[f215_D])).
fof(f215_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f209,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | $spl7),
  inference(cnf_transformation,[],[f209_D])).
fof(f209_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f1589,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0))))),
  inference(superposition,[],[f187,f89])).
fof(f187,plain,(
  s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) != s(bool,g(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0))))),
  inference(equality_resolution,[],[f123])).
fof(f123,plain,(
  ( ! [X0] : (s(fun(sK0,bool),sK1) != s(fun(sK0,bool),X0) | s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),X0,bool,sK0)))) != s(bool,g(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),X0,bool,sK0))))) )),
  inference(superposition,[],[f74,f112])).
fof(f112,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun(X0,X1),X2) = s(fun(X0,X1),X3) | s(X1,f(s(X0,sK3(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK3(X3,X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f70])).
fof(f70,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK3(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK3(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3])],[f50])).
fof(f50,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f44])).
fof(f44,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f43])).
fof(f43,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.pred_set.IMAGE_ID',aHL_EXT)).
fof(f205,plain,(
  p(s(bool,g(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0))))) | $spl0),
  inference(cnf_transformation,[],[f205_D])).
fof(f205_D,plain,(
  p(s(bool,g(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f1585,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | p(s(bool,g(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))))),
  inference(superposition,[],[f187,f84])).
fof(f1584,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | p(s(bool,g(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))))),
  inference(superposition,[],[f187,f82])).
fof(f204,plain,(
  ~p(s(bool,g(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0))))) | $spl1),
  inference(cnf_transformation,[],[f204_D])).
fof(f204_D,plain,(
  ~p(s(bool,g(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f207,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | ~p(s(bool,X2))) ) | $spl4),
  inference(cnf_transformation,[],[f207_D])).
fof(f207_D,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | ~p(s(bool,X2))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f1583,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | ~p(s(bool,X3)) | ~p(s(bool,g(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))))) )),
  inference(superposition,[],[f187,f78])).
fof(f1582,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | ~p(s(bool,g(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0))))) | ~p(s(bool,X2))) )),
  inference(superposition,[],[f187,f78])).
fof(f206,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | p(s(bool,X0))) ) | $spl2),
  inference(cnf_transformation,[],[f206_D])).
fof(f206_D,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | p(s(bool,X0))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f1581,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | p(s(bool,g(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0))))) | p(s(bool,X1))) )),
  inference(superposition,[],[f187,f75])).
fof(f1580,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | p(s(bool,X0)) | p(s(bool,g(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))))) )),
  inference(superposition,[],[f187,f75])).
fof(f263,plain,(
  ~p(s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0))))) | $spl15),
  inference(cnf_transformation,[],[f263_D])).
fof(f263_D,plain,(
  ~p(s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0))))) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f1573,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))))) ) | $spl12),
  inference(subsumption_resolution,[],[f1562,f85])).
fof(f1562,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,X3)) | ~p(s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))))) ) | $spl12),
  inference(superposition,[],[f1052,f78])).
fof(f1052,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | $spl12),
  inference(forward_demodulation,[],[f187,f215])).
fof(f1572,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))))) ) | $spl12),
  inference(subsumption_resolution,[],[f1561,f85])).
fof(f1561,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0))))) | ~p(s(bool,X2))) ) | $spl12),
  inference(superposition,[],[f1052,f78])).
fof(f1552,plain,(
  $false | ($spl6 | $spl12 | $spl18)),
  inference(subsumption_resolution,[],[f267,f533])).
fof(f533,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ($spl6 | $spl12)),
  inference(backward_demodulation,[],[f215,f423])).
fof(f423,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | $spl6),
  inference(backward_demodulation,[],[f208,f187])).
fof(f267,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl18),
  inference(cnf_transformation,[],[f267_D])).
fof(f267_D,plain,(
  s(bool,'T_0') = s(bool,'F_0') <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f1551,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,g(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))))) ) | $spl6),
  inference(subsumption_resolution,[],[f1538,f81])).
fof(f1538,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,g(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0))))) | p(s(bool,X1))) ) | $spl6),
  inference(superposition,[],[f423,f75])).
fof(f1550,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,g(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))))) ) | $spl6),
  inference(subsumption_resolution,[],[f1537,f81])).
fof(f1537,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,X0)) | p(s(bool,g(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))))) ) | $spl6),
  inference(superposition,[],[f423,f75])).
fof(f211,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | $spl9),
  inference(cnf_transformation,[],[f211_D])).
fof(f211_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f1140,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | ~p(s(bool,g(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))))),
  inference(superposition,[],[f187,f88])).
fof(f1139,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | ~p(s(bool,g(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))))),
  inference(superposition,[],[f187,f86])).
fof(f1138,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | p(s(bool,g(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))))),
  inference(superposition,[],[f187,f84])).
fof(f1137,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | p(s(bool,g(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))))),
  inference(superposition,[],[f187,f82])).
fof(f1136,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | ~p(s(bool,X3)) | ~p(s(bool,g(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))))) )),
  inference(superposition,[],[f187,f78])).
fof(f1135,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | ~p(s(bool,g(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0))))) | ~p(s(bool,X2))) )),
  inference(superposition,[],[f187,f78])).
fof(f1134,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | p(s(bool,g(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0))))) | p(s(bool,X1))) )),
  inference(superposition,[],[f187,f75])).
fof(f1133,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | p(s(bool,X0)) | p(s(bool,g(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))))) )),
  inference(superposition,[],[f187,f75])).
fof(f1513,plain,(
  $false | ($spl0 | $spl10)),
  inference(subsumption_resolution,[],[f1512,f109])).
fof(f1512,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl10)),
  inference(forward_demodulation,[],[f205,f213])).
fof(f1511,plain,(
  $false | ($spl6 | $spl10)),
  inference(subsumption_resolution,[],[f423,f213])).
fof(f1510,plain,(
  $false | ($spl6 | $spl10)),
  inference(subsumption_resolution,[],[f1509,f109])).
fof(f1509,plain,(
  p(s(bool,'F_0')) | ($spl6 | $spl10)),
  inference(forward_demodulation,[],[f519,f213])).
fof(f519,plain,(
  p(s(bool,g(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0))))) | $spl6),
  inference(trivial_inequality_removal,[],[f513])).
fof(f513,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,g(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0))))) | $spl6),
  inference(superposition,[],[f423,f84])).
fof(f1508,plain,(
  $false | ($spl6 | $spl10)),
  inference(subsumption_resolution,[],[f1507,f109])).
fof(f1507,plain,(
  p(s(bool,'F_0')) | ($spl6 | $spl10)),
  inference(forward_demodulation,[],[f520,f213])).
fof(f520,plain,(
  p(s(bool,g(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0))))) | $spl6),
  inference(trivial_inequality_removal,[],[f512])).
fof(f512,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,g(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0))))) | $spl6),
  inference(superposition,[],[f423,f82])).
fof(f1459,plain,(
  $false | ($spl4 | $spl8)),
  inference(trivial_inequality_removal,[],[f1450])).
fof(f1450,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl4 | $spl8)),
  inference(superposition,[],[f1422,f210])).
fof(f210,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | $spl8),
  inference(cnf_transformation,[],[f210_D])).
fof(f210_D,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f1422,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2)) ) | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f1421,f87])).
fof(f1421,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,X2))) ) | ($spl4 | $spl8)),
  inference(forward_demodulation,[],[f207,f210])).
fof(f1458,plain,(
  $false | ($spl4 | $spl8 | $spl34)),
  inference(trivial_inequality_removal,[],[f1451])).
fof(f1451,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl4 | $spl8 | $spl34)),
  inference(superposition,[],[f1422,f646])).
fof(f1457,plain,(
  $false | ($spl4 | $spl8)),
  inference(equality_resolution,[],[f1422])).
fof(f1142,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0))))),
  inference(superposition,[],[f187,f89])).
fof(f213,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | $spl10),
  inference(cnf_transformation,[],[f213_D])).
fof(f213_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f1141,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0))))),
  inference(superposition,[],[f187,f89])).
fof(f1091,plain,(
  $false | ($spl16 | $spl28)),
  inference(trivial_inequality_removal,[],[f1084])).
fof(f1084,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl16 | $spl28)),
  inference(superposition,[],[f265,f639])).
fof(f1090,plain,(
  $false | $spl16),
  inference(equality_resolution,[],[f265])).
fof(f1057,plain,(
  $false | ($spl1 | $spl12)),
  inference(subsumption_resolution,[],[f1056,f111])).
fof(f1056,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl12)),
  inference(forward_demodulation,[],[f204,f215])).
fof(f1055,plain,(
  $false | ($spl10 | $spl12 | $spl19)),
  inference(subsumption_resolution,[],[f1054,f268])).
fof(f1054,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl10 | $spl12)),
  inference(forward_demodulation,[],[f213,f215])).
fof(f1043,plain,(
  $false | ($spl2 | $spl6)),
  inference(trivial_inequality_removal,[],[f1033])).
fof(f1033,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl2 | $spl6)),
  inference(superposition,[],[f1006,f208])).
fof(f1006,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | ($spl2 | $spl6)),
  inference(subsumption_resolution,[],[f1005,f83])).
fof(f1005,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,X0))) ) | ($spl2 | $spl6)),
  inference(forward_demodulation,[],[f206,f208])).
fof(f1042,plain,(
  $false | ($spl2 | $spl6 | $spl28)),
  inference(trivial_inequality_removal,[],[f1034])).
fof(f1034,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl2 | $spl6 | $spl28)),
  inference(superposition,[],[f1006,f639])).
fof(f1041,plain,(
  $false | ($spl2 | $spl6)),
  inference(equality_resolution,[],[f1006])).
fof(f1002,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1))))) | $spl64),
  inference(cnf_transformation,[],[f1002_D])).
fof(f1002_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1))))) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f996,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | $spl59),
  inference(cnf_transformation,[],[f996_D])).
fof(f996_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl59),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl59])])).
fof(f990,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)))))),
  inference(superposition,[],[f531,f89])).
fof(f1000,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1))))) | $spl62),
  inference(cnf_transformation,[],[f1000_D])).
fof(f1000_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1))))) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f998,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | $spl61),
  inference(cnf_transformation,[],[f998_D])).
fof(f998_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl61),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl61])])).
fof(f989,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)))))),
  inference(superposition,[],[f531,f89])).
fof(f991,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)))))) | $spl53),
  inference(cnf_transformation,[],[f991_D])).
fof(f991_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)))))) <=> ~$spl53),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl53])])).
fof(f988,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1))))))),
  inference(superposition,[],[f531,f88])).
fof(f987,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1))))))),
  inference(superposition,[],[f531,f86])).
fof(f986,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1))))))),
  inference(superposition,[],[f531,f84])).
fof(f985,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1))))))),
  inference(superposition,[],[f531,f82])).
fof(f994,plain,(
  ( ! [X4] : (s(bool,X4) != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | ~p(s(bool,X4))) ) | $spl56),
  inference(cnf_transformation,[],[f994_D])).
fof(f994_D,plain,(
  ( ! [X4] : (s(bool,X4) != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | ~p(s(bool,X4))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f984,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | ~p(s(bool,X5)) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1))))))) )),
  inference(superposition,[],[f531,f78])).
fof(f983,plain,(
  ( ! [X4] : (s(bool,X4) != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)))))) | ~p(s(bool,X4))) )),
  inference(superposition,[],[f531,f78])).
fof(f993,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | p(s(bool,X2))) ) | $spl54),
  inference(cnf_transformation,[],[f993_D])).
fof(f993_D,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | p(s(bool,X2))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f982,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)))))) | p(s(bool,X3))) )),
  inference(superposition,[],[f531,f75])).
fof(f981,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | p(s(bool,X2)) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1))))))) )),
  inference(superposition,[],[f531,f75])).
fof(f846,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)))))),
  inference(superposition,[],[f393,f89])).
fof(f845,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)))))),
  inference(superposition,[],[f393,f89])).
fof(f844,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1))))))),
  inference(superposition,[],[f393,f88])).
fof(f843,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1))))))),
  inference(superposition,[],[f393,f86])).
fof(f842,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1))))))),
  inference(superposition,[],[f393,f84])).
fof(f841,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1))))))),
  inference(superposition,[],[f393,f82])).
fof(f840,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | ~p(s(bool,X5)) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1))))))) )),
  inference(superposition,[],[f393,f78])).
fof(f839,plain,(
  ( ! [X4] : (s(bool,X4) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)))))) | ~p(s(bool,X4))) )),
  inference(superposition,[],[f393,f78])).
fof(f838,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)))))) | p(s(bool,X3))) )),
  inference(superposition,[],[f393,f75])).
fof(f837,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | p(s(bool,X2)) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1))))))) )),
  inference(superposition,[],[f393,f75])).
fof(f695,plain,(
  p(s(bool,f(s(sK0,sK3(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),bool,sK0))))) | $spl36),
  inference(cnf_transformation,[],[f695_D])).
fof(f695_D,plain,(
  p(s(bool,f(s(sK0,sK3(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),bool,sK0))))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f696,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,f(s(sK0,sK3(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),bool,sK0)))))) ) | $spl29),
  inference(subsumption_resolution,[],[f681,f81])).
fof(f681,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,f(s(sK0,sK3(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),bool,sK0))))) | p(s(bool,X1))) ) | $spl29),
  inference(superposition,[],[f640,f75])).
fof(f693,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,f(s(sK0,sK3(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),bool,sK0)))))) ) | $spl29),
  inference(subsumption_resolution,[],[f680,f81])).
fof(f680,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,X0)) | p(s(bool,f(s(sK0,sK3(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),bool,sK0)))))) ) | $spl29),
  inference(superposition,[],[f640,f75])).
fof(f646,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | $spl34),
  inference(cnf_transformation,[],[f646_D])).
fof(f646_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f640,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | $spl29),
  inference(cnf_transformation,[],[f640_D])).
fof(f640_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f634,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),bool,sK0))))),
  inference(superposition,[],[f297,f89])).
fof(f297,plain,(
  s(bool,f(s(sK0,sK3(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) != s(bool,g(s(sK0,sK3(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),bool,sK0))))),
  inference(equality_resolution,[],[f124])).
fof(f124,plain,(
  ( ! [X1] : (s(fun(sK0,bool),sK1) != s(fun(sK0,bool),X1) | s(bool,f(s(sK0,sK3(X1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) != s(bool,g(s(sK0,sK3(X1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),bool,sK0))))) )),
  inference(superposition,[],[f74,f112])).
fof(f644,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | $spl32),
  inference(cnf_transformation,[],[f644_D])).
fof(f644_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f642,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | $spl31),
  inference(cnf_transformation,[],[f642_D])).
fof(f642_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) <=> ~$spl31),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl31])])).
fof(f633,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),bool,sK0))))),
  inference(superposition,[],[f297,f89])).
fof(f635,plain,(
  ~p(s(bool,g(s(sK0,sK3(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),bool,sK0))))) | $spl23),
  inference(cnf_transformation,[],[f635_D])).
fof(f635_D,plain,(
  ~p(s(bool,g(s(sK0,sK3(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),bool,sK0))))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f632,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | ~p(s(bool,g(s(sK0,sK3(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),bool,sK0)))))),
  inference(superposition,[],[f297,f88])).
fof(f631,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | ~p(s(bool,g(s(sK0,sK3(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),bool,sK0)))))),
  inference(superposition,[],[f297,f86])).
fof(f636,plain,(
  p(s(bool,g(s(sK0,sK3(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),bool,sK0))))) | $spl22),
  inference(cnf_transformation,[],[f636_D])).
fof(f636_D,plain,(
  p(s(bool,g(s(sK0,sK3(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),bool,sK0))))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f630,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | p(s(bool,g(s(sK0,sK3(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),bool,sK0)))))),
  inference(superposition,[],[f297,f84])).
fof(f629,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | p(s(bool,g(s(sK0,sK3(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),bool,sK0)))))),
  inference(superposition,[],[f297,f82])).
fof(f638,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | ~p(s(bool,X2))) ) | $spl26),
  inference(cnf_transformation,[],[f638_D])).
fof(f638_D,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | ~p(s(bool,X2))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f628,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | ~p(s(bool,X3)) | ~p(s(bool,g(s(sK0,sK3(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),bool,sK0)))))) )),
  inference(superposition,[],[f297,f78])).
fof(f627,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | ~p(s(bool,g(s(sK0,sK3(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),bool,sK0))))) | ~p(s(bool,X2))) )),
  inference(superposition,[],[f297,f78])).
fof(f637,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | p(s(bool,X0))) ) | $spl24),
  inference(cnf_transformation,[],[f637_D])).
fof(f637_D,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | p(s(bool,X0))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f626,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | p(s(bool,g(s(sK0,sK3(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),bool,sK0))))) | p(s(bool,X1))) )),
  inference(superposition,[],[f297,f75])).
fof(f625,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | p(s(bool,X0)) | p(s(bool,g(s(sK0,sK3(sK1,'const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),bool,sK0)))))) )),
  inference(superposition,[],[f297,f75])).
fof(f346,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))))) ) | $spl9),
  inference(subsumption_resolution,[],[f335,f85])).
fof(f335,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0))))) | ~p(s(bool,X2))) ) | $spl9),
  inference(superposition,[],[f211,f78])).
fof(f348,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))))) ) | $spl9),
  inference(subsumption_resolution,[],[f336,f85])).
fof(f336,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,X3)) | ~p(s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))))) ) | $spl9),
  inference(superposition,[],[f211,f78])).
fof(f352,plain,(
  $false | ($spl7 | $spl9)),
  inference(subsumption_resolution,[],[f343,f209])).
fof(f343,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | $spl9),
  inference(trivial_inequality_removal,[],[f341])).
fof(f341,plain,(
  s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | $spl9),
  inference(superposition,[],[f211,f89])).
fof(f351,plain,(
  $false | ($spl7 | $spl9)),
  inference(subsumption_resolution,[],[f344,f260])).
fof(f260,plain,(
  p(s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0))))) | $spl7),
  inference(trivial_inequality_removal,[],[f254])).
fof(f254,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0))))) | $spl7),
  inference(superposition,[],[f209,f84])).
fof(f344,plain,(
  ~p(s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0))))) | $spl9),
  inference(trivial_inequality_removal,[],[f340])).
fof(f340,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ~p(s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0))))) | $spl9),
  inference(superposition,[],[f211,f88])).
fof(f350,plain,(
  $false | ($spl7 | $spl9)),
  inference(subsumption_resolution,[],[f345,f260])).
fof(f345,plain,(
  ~p(s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0))))) | $spl9),
  inference(trivial_inequality_removal,[],[f339])).
fof(f339,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ~p(s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0))))) | $spl9),
  inference(superposition,[],[f211,f86])).
fof(f268,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl19),
  inference(cnf_transformation,[],[f268_D])).
fof(f268_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f256,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0))))) | $spl7),
  inference(superposition,[],[f209,f88])).
fof(f255,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0))))) | $spl7),
  inference(superposition,[],[f209,f86])).
fof(f264,plain,(
  p(s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0))))) | $spl14),
  inference(cnf_transformation,[],[f264_D])).
fof(f264_D,plain,(
  p(s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0))))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f266,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))))) ) | $spl7),
  inference(subsumption_resolution,[],[f250,f81])).
fof(f250,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0))))) | p(s(bool,X1))) ) | $spl7),
  inference(superposition,[],[f209,f75])).
fof(f262,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))))) ) | $spl7),
  inference(subsumption_resolution,[],[f249,f81])).
fof(f249,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,X0)) | p(s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))))) ) | $spl7),
  inference(superposition,[],[f209,f75])).
fof(f203,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0))))),
  inference(superposition,[],[f187,f89])).
fof(f202,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0))))),
  inference(superposition,[],[f187,f89])).
fof(f201,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | ~p(s(bool,g(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))))),
  inference(superposition,[],[f187,f88])).
fof(f200,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | ~p(s(bool,g(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))))),
  inference(superposition,[],[f187,f86])).
fof(f199,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | p(s(bool,g(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))))),
  inference(superposition,[],[f187,f84])).
fof(f198,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | p(s(bool,g(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))))),
  inference(superposition,[],[f187,f82])).
fof(f197,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | ~p(s(bool,X3)) | ~p(s(bool,g(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))))) )),
  inference(superposition,[],[f187,f78])).
fof(f196,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | ~p(s(bool,g(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0))))) | ~p(s(bool,X2))) )),
  inference(superposition,[],[f187,f78])).
fof(f195,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | p(s(bool,g(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0))))) | p(s(bool,X1))) )),
  inference(superposition,[],[f187,f75])).
fof(f194,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,f(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | p(s(bool,X0)) | p(s(bool,g(s(sK0,sK3('const.pred_set.IMAGE_2'(s(fun(sK0,sK0),'_i_0_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))))) )),
  inference(superposition,[],[f187,f75])).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 209
% Passive clauses: 1698
% Generated clauses: 4242
% Final active clauses: 145
% Final passive clauses: 1044
% Input formulas: 20
% Initial clauses: 41
% 
% Selected by SInE selection: 20
% SInE iterations: 1
% 
% Duplicate literals: 35
% Trivial inequalities: 170
% Fw subsumption resolutions: 323
% Fw demodulations: 439
% Bw demodulations: 35
% 
% Simple tautologies: 95
% Equational tautologies: 1
% Forward subsumptions: 1415
% 
% Binary resolution: 146
% Forward superposition: 2231
% Backward superposition: 621
% Equality resolution: 37
% 
% Split clauses: 103
% Split components: 206
% SAT solver clauses: 146
% SAT solver unit clauses: 2
% SAT solver binary clauses: 104
% 
% Sat splits: 103
% Sat splitting refutations: 42
% 
% Memory used [KB]: 2558
% Time elapsed: 0.127 s
% ------------------------------
% ------------------------------
% Success in time 11.841 s
