Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.list
% (59767)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1068
% Passive clauses: 8259
% Generated clauses: 50947
% Final active clauses: 1025
% Final passive clauses: 6791
% Input formulas: 261
% Initial clauses: 511
% 
% Duplicate literals: 231
% Trivial inequalities: 201
% Fw subsumption resolutions: 2291
% Fw demodulations: 3814
% Bw demodulations: 12
% 
% Simple tautologies: 302
% Equational tautologies: 26
% Forward subsumptions: 35425
% 
% Binary resolution: 6289
% Factoring: 1010
% Forward superposition: 15272
% Backward superposition: 21066
% Self superposition: 61
% Equality factoring: 2
% Equality resolution: 18
% 
% Split clauses: 329
% Split components: 660
% SAT solver clauses: 498
% SAT solver unit clauses: 20
% SAT solver binary clauses: 75
% SAT solver learnt clauses: 1
% SAT solver learnt literals: 1
% 
% Sat splits: 329
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 497
% TWLsolver calls for satisfiability: 375
% 
% Memory used [KB]: 31214
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.list
% (59878)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 614
% Passive clauses: 3502
% Generated clauses: 36751
% Final active clauses: 572
% Final passive clauses: 2495
% Input formulas: 261
% Initial clauses: 511
% 
% Duplicate literals: 99
% Trivial inequalities: 34
% Fw subsumption resolutions: 1098
% Fw demodulations: 3606
% Bw demodulations: 405
% Condensations: 103
% Global subsumptions: 24
% 
% Simple tautologies: 266
% Equational tautologies: 10
% Forward subsumptions: 23669
% Fw demodulations to eq. taut.: 15
% 
% Binary resolution: 2489
% Factoring: 235
% Forward superposition: 12216
% Backward superposition: 12084
% Self superposition: 92
% Equality factoring: 2
% Equality resolution: 5
% 
% Split clauses: 26
% Split components: 52
% SAT solver clauses: 4557
% SAT solver unit clauses: 631
% SAT solver binary clauses: 126
% SAT solver learnt clauses: 49
% SAT solver learnt literals: 13
% 
% Sat splits: 26
% Sat splitting refutations: 22
% 
% TWLsolver clauses: 4460
% TWLsolver calls for satisfiability: 2998
% 
% Memory used [KB]: 32622
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.list
% (59903)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1547
% Passive clauses: 13774
% Generated clauses: 189181
% Final active clauses: 1240
% Final passive clauses: 3228
% Input formulas: 261
% Initial clauses: 511
% Discarded non-redundant clauses: 45893
% 
% Duplicate literals: 16
% Trivial inequalities: 17
% Fw subsumption resolutions: 2239
% Fw demodulations: 44055
% Bw demodulations: 4328
% Condensations: 185
% Global subsumptions: 430
% 
% Simple tautologies: 120
% Equational tautologies: 72
% Forward subsumptions: 81604
% Fw demodulations to eq. taut.: 20693
% Bw demodulations to eq. taut.: 18
% 
% Unit resulting resolution: 62885
% Factoring: 1
% Forward superposition: 36845
% Backward superposition: 33505
% Self superposition: 235
% Equality factoring: 2
% Equality resolution: 11
% 
% Split clauses: 61
% Split components: 122
% SAT solver clauses: 16312
% SAT solver unit clauses: 4760
% SAT solver binary clauses: 8010
% 
% Sat splits: 61
% Sat splitting refutations: 1984
% 
% Memory used [KB]: 180380
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.list
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.list
% SZS output start Proof for thm.list
fof(f17098,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f546,f784,f782,f545,f783,f781,f932,f1102,f1173,f1287,f1495,f1502,f1715,f1717,f1835,f1980,f1857,f1996,f1994,f2006,f1984,f2025,f2045,f2089,f2968,f2976,f3018,f2149,f2053,f3021,f1982,f3442,f3450,f4122,f4276,f4144,f4145,f4292,f4295,f4297,f4299,f4290,f4301,f4310,f4278,f3017,f4439,f4634,f4656,f4798,f4797,f4822,f4274,f4846,f4836,f3016,f4848,f5134,f5476,f5453,f5478,f5512,f5573,f5595,f5626,f5649,f5872,f5895,f7199,f7225,f7348,f7372,f4288,f7653,f6734,f7654,f4501,f7655,f3935,f7656,f3656,f7657,f1171,f8350,f8194,f3020,f7931,f8353,f8352,f7909,f1428,f1426,f8760,f2090,f8831,f8974,f8973,f8833,f8978,f8976,f9294,f9343,f9342,f9705,f9707,f9721,f9778,f9802,f10024,f10191,f10241,f10243,f10845,f10847,f10848,f11018,f10850,f11021,f10852,f11020,f10930,f11019,f11022,f11040,f11090,f11096,f11128,f1979,f1977,f11457,f3014,f3013,f11511,f4275,f4273,f11704,f10193,f11874,f11923,f11925,f12608,f12658,f11089,f13495,f13572,f13506,f13573,f13571,f14581,f14591,f1501,f16773,f5477,f16750,f16775,f16830,f16928,f16950,f17011,f16774,f17034,f17089])).
fof(f17089,plain,(
  $false | ($spl0 | $spl3)),
  inference(equality_resolution,[],[f17010])).
fof(f17010,plain,(
  ( ! [X0,X1] : (s('type.list.list'(sK11),'const.list.CONS_2'(s(sK11,X0),s('type.list.list'(sK11),'const.list.NIL_0'))) != s('type.list.list'(sK11),'const.list.CONS_2'(s(sK11,X1),s('type.list.list'(sK11),'const.list.NIL_0')))) ) | ($spl0 | $spl3)),
  inference(superposition,[],[f16739,f781])).
fof(f16739,plain,(
  ( ! [X92,X93,X91] : (s('type.list.list'(sK11),'const.list.CONS_2'(s(sK11,X91),s('type.list.list'(sK11),'const.list.NIL_0'))) != s('type.list.list'(sK11),'const.list.CONS_2'(s(sK11,X92),s('type.list.list'(sK11),'const.list.MAP_2'(s(fun(sK12,sK11),X93),s('type.list.list'(sK12),sK13)))))) ) | ($spl0 | $spl3)),
  inference(superposition,[],[f5820,f942])).
fof(f942,plain,(
  s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),sK68(sK13,sK12)))) | $spl3),
  inference(trivial_inequality_removal,[],[f933])).
fof(f933,plain,(
  s('type.list.list'(sK12),'const.list.NIL_0') != s('type.list.list'(sK12),'const.list.NIL_0') | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),sK68(sK13,sK12)))) | $spl3),
  inference(superposition,[],[f784,f772])).
fof(f772,plain,(
  ( ! [X0,X1] : (s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK67(X1,X0)),s('type.list.list'(X0),sK68(X1,X0)))) | s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X1)) )),
  inference(cnf_transformation,[],[f544])).
fof(f544,plain,(
  ! [X0,X1] : (s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X1) | s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK67(X1,X0)),s('type.list.list'(X0),sK68(X1,X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK67,sK68])],[f369])).
fof(f369,plain,(
  ! [X0,X1] : (s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X1) | ? [X2,X3] : s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3))))),
  inference(flattening,[],[f368])).
fof(f368,plain,(
  ! [X0] : ! [X1] : (s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X1) | ? [X2,X3] : s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3))))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X7] : ! [X8] : (s('type.list.list'(X7),'const.list.NIL_0') = s('type.list.list'(X7),X8) | ? [X10,X11] : s('type.list.list'(X7),X8) = s('type.list.list'(X7),'const.list.CONS_2'(s(X7,X10),s('type.list.list'(X7),X11))))),
  file('thm.list.MAP_EQ_NIL|split|0','thm.list.list_CASES')).
fof(f5820,plain,(
  ( ! [X80,X78,X81,X79,X77,X82] : (s('type.list.list'(sK11),'const.list.CONS_2'(s(sK11,X81),s('type.list.list'(sK11),'const.list.NIL_0'))) != s('type.list.list'(sK11),'const.list.CONS_2'(s(sK11,X82),s('type.list.list'(sK11),'const.list.MAP_2'(s(fun(X77,sK11),X78),s('type.list.list'(X77),'const.list.CONS_2'(s(X77,X79),s('type.list.list'(X77),X80)))))))) ) | $spl0),
  inference(subsumption_resolution,[],[f5819,f770])).
fof(f770,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X1)))) )),
  inference(cnf_transformation,[],[f365])).
fof(f365,plain,(
  ! [X0,X1,X2] : s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X1)))),
  inference(flattening,[],[f364])).
fof(f364,plain,(
  ! [X0] : ! [X1,X2] : ~s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X1)))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X7] : ! [X37,X38] : ~s('type.list.list'(X7),'const.list.NIL_0') = s('type.list.list'(X7),'const.list.CONS_2'(s(X7,X38),s('type.list.list'(X7),X37)))),
  file('thm.list.MAP_EQ_NIL|split|0','thm.list.list_distinct')).
fof(f5819,plain,(
  ( ! [X80,X78,X81,X79,X77,X82] : (s('type.list.list'(sK11),'const.list.NIL_0') = s('type.list.list'(sK11),'const.list.CONS_2'(s(sK11,happ(s(fun(X77,sK11),X78),s(X77,X79))),s('type.list.list'(sK11),'const.list.MAP_2'(s(fun(X77,sK11),X78),s('type.list.list'(X77),X80))))) | s('type.list.list'(sK11),'const.list.CONS_2'(s(sK11,X81),s('type.list.list'(sK11),'const.list.NIL_0'))) != s('type.list.list'(sK11),'const.list.CONS_2'(s(sK11,X82),s('type.list.list'(sK11),'const.list.MAP_2'(s(fun(X77,sK11),X78),s('type.list.list'(X77),'const.list.CONS_2'(s(X77,X79),s('type.list.list'(X77),X80)))))))) ) | $spl0),
  inference(forward_demodulation,[],[f5761,f781])).
fof(f5761,plain,(
  ( ! [X80,X78,X81,X79,X77,X82] : (s('type.list.list'(sK11),'const.list.CONS_2'(s(sK11,X81),s('type.list.list'(sK11),'const.list.NIL_0'))) != s('type.list.list'(sK11),'const.list.CONS_2'(s(sK11,X82),s('type.list.list'(sK11),'const.list.MAP_2'(s(fun(X77,sK11),X78),s('type.list.list'(X77),'const.list.CONS_2'(s(X77,X79),s('type.list.list'(X77),X80))))))) | s('type.list.list'(sK11),'const.list.MAP_2'(s(fun(sK12,sK11),sK14),s('type.list.list'(sK12),sK13))) = s('type.list.list'(sK11),'const.list.CONS_2'(s(sK11,happ(s(fun(X77,sK11),X78),s(X77,X79))),s('type.list.list'(sK11),'const.list.MAP_2'(s(fun(X77,sK11),X78),s('type.list.list'(X77),X80)))))) ) | $spl0),
  inference(superposition,[],[f972,f602])).
fof(f602,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X2),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X3),s('type.list.list'(X1),X4))))) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X3))),s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X2),s('type.list.list'(X1),X4)))))) )),
  inference(cnf_transformation,[],[f292])).
fof(f292,plain,(
  ! [X0,X1,X2,X3,X4] : s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X2),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X3),s('type.list.list'(X1),X4))))) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X3))),s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X2),s('type.list.list'(X1),X4)))))),
  inference(flattening,[],[f291])).
fof(f291,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X2),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X3),s('type.list.list'(X1),X4))))) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X3))),s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X2),s('type.list.list'(X1),X4)))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6,X7] : ! [X9,X10,X11] : s('type.list.list'(X6),'const.list.MAP_2'(s(fun(X7,X6),X9),s('type.list.list'(X7),'const.list.CONS_2'(s(X7,X10),s('type.list.list'(X7),X11))))) = s('type.list.list'(X6),'const.list.CONS_2'(s(X6,happ(s(fun(X7,X6),X9),s(X7,X10))),s('type.list.list'(X6),'const.list.MAP_2'(s(fun(X7,X6),X9),s('type.list.list'(X7),X11)))))),
  file('thm.list.MAP_EQ_NIL|split|0','thm.list.MAP|split|1')).
fof(f972,plain,(
  ( ! [X10,X11,X9] : (s('type.list.list'(sK11),'const.list.CONS_2'(s(sK11,X9),s('type.list.list'(sK11),'const.list.NIL_0'))) != s('type.list.list'(sK11),'const.list.CONS_2'(s(sK11,X10),s('type.list.list'(sK11),X11))) | s('type.list.list'(sK11),X11) = s('type.list.list'(sK11),'const.list.MAP_2'(s(fun(sK12,sK11),sK14),s('type.list.list'(sK12),sK13)))) ) | $spl0),
  inference(superposition,[],[f593,f781])).
fof(f593,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.list.list'(X0),X2) = s('type.list.list'(X0),X4) | s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2))) != s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X4)))) )),
  inference(cnf_transformation,[],[f450])).
fof(f450,plain,(
  ! [X0,X1,X2,X3,X4] : ((s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2))) != s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X4))) | (s(X0,X1) = s(X0,X3) & s('type.list.list'(X0),X2) = s('type.list.list'(X0),X4))) & (s(X0,X1) != s(X0,X3) | s('type.list.list'(X0),X2) != s('type.list.list'(X0),X4) | s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2))) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X4)))))),
  inference(flattening,[],[f449])).
fof(f449,plain,(
  ! [X0,X1,X2,X3,X4] : ((s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2))) != s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X4))) | (s(X0,X1) = s(X0,X3) & s('type.list.list'(X0),X2) = s('type.list.list'(X0),X4))) & ((s(X0,X1) != s(X0,X3) | s('type.list.list'(X0),X2) != s('type.list.list'(X0),X4)) | s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2))) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X4)))))),
  inference(nnf_transformation,[],[f278])).
fof(f278,plain,(
  ! [X0,X1,X2,X3,X4] : (s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2))) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X4))) <=> (s(X0,X1) = s(X0,X3) & s('type.list.list'(X0),X2) = s('type.list.list'(X0),X4)))),
  inference(flattening,[],[f277])).
fof(f277,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : (s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2))) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X4))) <=> (s(X0,X1) = s(X0,X3) & s('type.list.list'(X0),X2) = s('type.list.list'(X0),X4)))),
  inference(rectify,[],[f84])).
fof(f84,axiom,(
  ! [X7] : ! [X38,X37,X56,X57] : (s('type.list.list'(X7),'const.list.CONS_2'(s(X7,X38),s('type.list.list'(X7),X37))) = s('type.list.list'(X7),'const.list.CONS_2'(s(X7,X56),s('type.list.list'(X7),X57))) <=> (s(X7,X38) = s(X7,X56) & s('type.list.list'(X7),X37) = s('type.list.list'(X7),X57)))),
  file('thm.list.MAP_EQ_NIL|split|0','thm.list.list_|49||49|')).
fof(f17034,plain,(
  ( ! [X2] : (s('type.list.list'(sK11),'const.list.MAP_2'(s(fun(sK12,sK11),X2),s('type.list.list'(sK12),sK13))) = s('type.list.list'(sK11),'const.list.CONS_2'(s(sK11,sK67('const.list.MAP_2'(s(fun(sK12,sK11),X2),s('type.list.list'(sK12),sK13)),sK11)),s('type.list.list'(sK11),sK68('const.list.MAP_2'(s(fun(sK12,sK11),X2),s('type.list.list'(sK12),sK13)),sK11))))) ) | $spl144),
  inference(cnf_transformation,[],[f17034_D])).
fof(f17034_D,plain,(
  ( ! [X2] : (s('type.list.list'(sK11),'const.list.MAP_2'(s(fun(sK12,sK11),X2),s('type.list.list'(sK12),sK13))) = s('type.list.list'(sK11),'const.list.CONS_2'(s(sK11,sK67('const.list.MAP_2'(s(fun(sK12,sK11),X2),s('type.list.list'(sK12),sK13)),sK11)),s('type.list.list'(sK11),sK68('const.list.MAP_2'(s(fun(sK12,sK11),X2),s('type.list.list'(sK12),sK13)),sK11))))) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f16774,plain,(
  ( ! [X48,X49] : (s('type.list.list'(sK11),'const.list.CONS_2'(s(sK11,X48),s('type.list.list'(sK11),'const.list.NIL_0'))) != s('type.list.list'(sK11),'const.list.CONS_2'(s(sK11,X49),s('type.list.list'(sK11),'const.list.NIL_0')))) ) | $spl138),
  inference(cnf_transformation,[],[f16774_D])).
fof(f16774_D,plain,(
  ( ! [X48,X49] : (s('type.list.list'(sK11),'const.list.CONS_2'(s(sK11,X48),s('type.list.list'(sK11),'const.list.NIL_0'))) != s('type.list.list'(sK11),'const.list.CONS_2'(s(sK11,X49),s('type.list.list'(sK11),'const.list.NIL_0')))) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f17011,plain,(
  ( ! [X4,X2,X3] : (s('type.list.list'(sK11),'const.list.CONS_2'(s(sK11,X3),s('type.list.list'(sK11),'const.list.NIL_0'))) != s('type.list.list'(sK11),'const.list.CONS_2'(s(sK11,X4),s('type.list.list'(sK11),'const.list.NIL_0'))) | s('type.list.list'(sK11),'const.list.MAP_2'(s(fun(sK12,sK11),X2),s('type.list.list'(sK12),sK13))) = s('type.list.list'(sK11),'const.list.CONS_2'(s(sK11,sK67('const.list.MAP_2'(s(fun(sK12,sK11),X2),s('type.list.list'(sK12),sK13)),sK11)),s('type.list.list'(sK11),sK68('const.list.MAP_2'(s(fun(sK12,sK11),X2),s('type.list.list'(sK12),sK13)),sK11))))) ) | ($spl0 | $spl3)),
  inference(superposition,[],[f16739,f772])).
fof(f16950,plain,(
  ( ! [X37,X36] : (s('type.list.list'(sK11),'const.list.CONS_2'(s(sK11,X36),s('type.list.list'(sK11),X37))) = s('type.list.list'(sK11),'const.list.CONS_2'(s(sK11,sK67('const.list.CONS_2'(s(sK11,X36),s('type.list.list'(sK11),X37)),sK11)),s('type.list.list'(sK11),sK68('const.list.CONS_2'(s(sK11,X36),s('type.list.list'(sK11),X37)),sK11))))) ) | $spl142),
  inference(cnf_transformation,[],[f16950_D])).
fof(f16950_D,plain,(
  ( ! [X37,X36] : (s('type.list.list'(sK11),'const.list.CONS_2'(s(sK11,X36),s('type.list.list'(sK11),X37))) = s('type.list.list'(sK11),'const.list.CONS_2'(s(sK11,sK67('const.list.CONS_2'(s(sK11,X36),s('type.list.list'(sK11),X37)),sK11)),s('type.list.list'(sK11),sK68('const.list.CONS_2'(s(sK11,X36),s('type.list.list'(sK11),X37)),sK11))))) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f16928,plain,(
  ( ! [X39,X37,X38,X36] : (s('type.list.list'(sK11),'const.list.CONS_2'(s(sK11,X38),s('type.list.list'(sK11),'const.list.NIL_0'))) != s('type.list.list'(sK11),'const.list.CONS_2'(s(sK11,X39),s('type.list.list'(sK11),'const.list.NIL_0'))) | s('type.list.list'(sK11),'const.list.CONS_2'(s(sK11,X36),s('type.list.list'(sK11),X37))) = s('type.list.list'(sK11),'const.list.CONS_2'(s(sK11,sK67('const.list.CONS_2'(s(sK11,X36),s('type.list.list'(sK11),X37)),sK11)),s('type.list.list'(sK11),sK68('const.list.CONS_2'(s(sK11,X36),s('type.list.list'(sK11),X37)),sK11))))) ) | $spl0),
  inference(superposition,[],[f16748,f772])).
fof(f16748,plain,(
  ( ! [X6,X4,X5,X3] : (s('type.list.list'(sK11),'const.list.CONS_2'(s(sK11,X5),s('type.list.list'(sK11),'const.list.NIL_0'))) != s('type.list.list'(sK11),'const.list.CONS_2'(s(sK11,X6),s('type.list.list'(sK11),'const.list.CONS_2'(s(sK11,X3),s('type.list.list'(sK11),X4)))))) ) | $spl0),
  inference(superposition,[],[f5820,f605])).
fof(f605,plain,(
  ( ! [X0,X1] : (s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X0,X0),'const.combin.I_0'),s('type.list.list'(X0),X1)))) )),
  inference(cnf_transformation,[],[f298])).
fof(f298,plain,(
  ! [X0,X1] : s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X0,X0),'const.combin.I_0'),s('type.list.list'(X0),X1)))),
  inference(flattening,[],[f297])).
fof(f297,plain,(
  ! [X0] : ! [X1] : s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X0,X0),'const.combin.I_0'),s('type.list.list'(X0),X1)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X7] : ! [X8] : s('type.list.list'(X7),X8) = s('type.list.list'(X7),'const.list.MAP_2'(s(fun(X7,X7),'const.combin.I_0'),s('type.list.list'(X7),X8)))),
  file('thm.list.MAP_EQ_NIL|split|0','thm.list.MAP_ID|split|1')).
fof(f16830,plain,(
  $false | $spl138),
  inference(equality_resolution,[],[f16774])).
fof(f16775,plain,(
  ( ! [X14,X12,X13,X11] : (s('type.list.list'(sK11),'const.list.MAP_2'(s(fun(X11,sK11),X12),s('type.list.list'(X11),'const.list.CONS_2'(s(X11,X13),s('type.list.list'(X11),X14))))) = s('type.list.list'(sK11),'const.list.CONS_2'(s(sK11,sK67('const.list.MAP_2'(s(fun(X11,sK11),X12),s('type.list.list'(X11),'const.list.CONS_2'(s(X11,X13),s('type.list.list'(X11),X14)))),sK11)),s('type.list.list'(sK11),sK68('const.list.MAP_2'(s(fun(X11,sK11),X12),s('type.list.list'(X11),'const.list.CONS_2'(s(X11,X13),s('type.list.list'(X11),X14)))),sK11))))) ) | $spl140),
  inference(cnf_transformation,[],[f16775_D])).
fof(f16775_D,plain,(
  ( ! [X14,X12,X13,X11] : (s('type.list.list'(sK11),'const.list.MAP_2'(s(fun(X11,sK11),X12),s('type.list.list'(X11),'const.list.CONS_2'(s(X11,X13),s('type.list.list'(X11),X14))))) = s('type.list.list'(sK11),'const.list.CONS_2'(s(sK11,sK67('const.list.MAP_2'(s(fun(X11,sK11),X12),s('type.list.list'(X11),'const.list.CONS_2'(s(X11,X13),s('type.list.list'(X11),X14)))),sK11)),s('type.list.list'(sK11),sK68('const.list.MAP_2'(s(fun(X11,sK11),X12),s('type.list.list'(X11),'const.list.CONS_2'(s(X11,X13),s('type.list.list'(X11),X14)))),sK11))))) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f16750,plain,(
  ( ! [X14,X12,X15,X13,X11,X16] : (s('type.list.list'(sK11),'const.list.CONS_2'(s(sK11,X15),s('type.list.list'(sK11),'const.list.NIL_0'))) != s('type.list.list'(sK11),'const.list.CONS_2'(s(sK11,X16),s('type.list.list'(sK11),'const.list.NIL_0'))) | s('type.list.list'(sK11),'const.list.MAP_2'(s(fun(X11,sK11),X12),s('type.list.list'(X11),'const.list.CONS_2'(s(X11,X13),s('type.list.list'(X11),X14))))) = s('type.list.list'(sK11),'const.list.CONS_2'(s(sK11,sK67('const.list.MAP_2'(s(fun(X11,sK11),X12),s('type.list.list'(X11),'const.list.CONS_2'(s(X11,X13),s('type.list.list'(X11),X14)))),sK11)),s('type.list.list'(sK11),sK68('const.list.MAP_2'(s(fun(X11,sK11),X12),s('type.list.list'(X11),'const.list.CONS_2'(s(X11,X13),s('type.list.list'(X11),X14)))),sK11))))) ) | $spl0),
  inference(superposition,[],[f5820,f772])).
fof(f5477,plain,(
  ( ! [X39,X41,X40] : (s('type.list.list'(X39),'const.list.CONS_2'(s(X39,X40),s('type.list.list'(X39),X41))) = s('type.list.list'(X39),'const.list.CONS_2'(s(X39,sK67('const.list.CONS_2'(s(X39,X40),s('type.list.list'(X39),X41)),X39)),s('type.list.list'(X39),sK68('const.list.CONS_2'(s(X39,X40),s('type.list.list'(X39),X41)),X39))))) ) | $spl68),
  inference(cnf_transformation,[],[f5477_D])).
fof(f5477_D,plain,(
  ( ! [X39,X41,X40] : (s('type.list.list'(X39),'const.list.CONS_2'(s(X39,X40),s('type.list.list'(X39),X41))) = s('type.list.list'(X39),'const.list.CONS_2'(s(X39,sK67('const.list.CONS_2'(s(X39,X40),s('type.list.list'(X39),X41)),X39)),s('type.list.list'(X39),sK68('const.list.CONS_2'(s(X39,X40),s('type.list.list'(X39),X41)),X39))))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f16773,plain,(
  ( ! [X47,X45,X48,X46,X49] : (s('type.list.list'(sK11),'const.list.CONS_2'(s(sK11,X48),s('type.list.list'(sK11),'const.list.NIL_0'))) != s('type.list.list'(sK11),'const.list.CONS_2'(s(sK11,X49),s('type.list.list'(sK11),'const.list.NIL_0'))) | s('type.list.list'(X45),'const.list.CONS_2'(s(X45,X46),s('type.list.list'(X45),X47))) = s('type.list.list'(X45),'const.list.CONS_2'(s(X45,sK67('const.list.CONS_2'(s(X45,X46),s('type.list.list'(X45),X47)),X45)),s('type.list.list'(X45),sK68('const.list.CONS_2'(s(X45,X46),s('type.list.list'(X45),X47)),X45))))) ) | $spl0),
  inference(forward_demodulation,[],[f16733,f604])).
fof(f604,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'(X1),'const.list.NIL_0') = s('type.list.list'(X1),'const.list.MAP_2'(s(fun(X0,X1),X2),s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(cnf_transformation,[],[f296])).
fof(f296,plain,(
  ! [X0,X1,X2] : s('type.list.list'(X1),'const.list.NIL_0') = s('type.list.list'(X1),'const.list.MAP_2'(s(fun(X0,X1),X2),s('type.list.list'(X0),'const.list.NIL_0')))),
  inference(flattening,[],[f295])).
fof(f295,plain,(
  ! [X0,X1] : ! [X2] : s('type.list.list'(X1),'const.list.NIL_0') = s('type.list.list'(X1),'const.list.MAP_2'(s(fun(X0,X1),X2),s('type.list.list'(X0),'const.list.NIL_0')))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X7,X6] : ! [X9] : s('type.list.list'(X6),'const.list.NIL_0') = s('type.list.list'(X6),'const.list.MAP_2'(s(fun(X7,X6),X9),s('type.list.list'(X7),'const.list.NIL_0')))),
  file('thm.list.MAP_EQ_NIL|split|0','thm.list.MAP|split|0')).
fof(f16733,plain,(
  ( ! [X47,X45,X50,X48,X46,X49] : (s('type.list.list'(sK11),'const.list.CONS_2'(s(sK11,X48),s('type.list.list'(sK11),'const.list.NIL_0'))) != s('type.list.list'(sK11),'const.list.CONS_2'(s(sK11,X49),s('type.list.list'(sK11),'const.list.MAP_2'(s(fun(X45,sK11),X50),s('type.list.list'(X45),'const.list.NIL_0'))))) | s('type.list.list'(X45),'const.list.CONS_2'(s(X45,X46),s('type.list.list'(X45),X47))) = s('type.list.list'(X45),'const.list.CONS_2'(s(X45,sK67('const.list.CONS_2'(s(X45,X46),s('type.list.list'(X45),X47)),X45)),s('type.list.list'(X45),sK68('const.list.CONS_2'(s(X45,X46),s('type.list.list'(X45),X47)),X45))))) ) | $spl0),
  inference(superposition,[],[f5820,f772])).
fof(f1501,plain,(
  ( ! [X39] : (s('type.list.list'(sK12),X39) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(X39,sK12)),s('type.list.list'(sK12),sK68(X39,sK12))))) ) | $spl12),
  inference(cnf_transformation,[],[f1501_D])).
fof(f1501_D,plain,(
  ( ! [X39] : (s('type.list.list'(sK12),X39) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(X39,sK12)),s('type.list.list'(sK12),sK68(X39,sK12))))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f14591,plain,(
  ( ! [X50,X51] : (s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X50),s('type.list.list'(sK12),'const.list.NIL_0'))) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X51),s('type.list.list'(sK12),sK68(sK13,sK12))))) ) | $spl136),
  inference(cnf_transformation,[],[f14591_D])).
fof(f14591_D,plain,(
  ( ! [X50,X51] : (s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X50),s('type.list.list'(sK12),'const.list.NIL_0'))) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X51),s('type.list.list'(sK12),sK68(sK13,sK12))))) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f14581,plain,(
  ( ! [X50,X51,X49] : (s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X50),s('type.list.list'(sK12),'const.list.NIL_0'))) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X51),s('type.list.list'(sK12),sK68(sK13,sK12)))) | s('type.list.list'(sK12),X49) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(X49,sK12)),s('type.list.list'(sK12),sK68(X49,sK12))))) ) | $spl21),
  inference(trivial_inequality_removal,[],[f14562])).
fof(f14562,plain,(
  ( ! [X50,X51,X49] : (s('type.list.list'(sK12),X49) != s('type.list.list'(sK12),X49) | s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X50),s('type.list.list'(sK12),'const.list.NIL_0'))) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X51),s('type.list.list'(sK12),sK68(sK13,sK12)))) | s('type.list.list'(sK12),X49) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(X49,sK12)),s('type.list.list'(sK12),sK68(X49,sK12))))) ) | $spl21),
  inference(superposition,[],[f7659,f772])).
fof(f7659,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'(sK12),'const.list.NIL_0') != s('type.list.list'(sK12),X0) | s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X1),s('type.list.list'(sK12),X0))) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X2),s('type.list.list'(sK12),sK68(sK13,sK12))))) ) | $spl21),
  inference(superposition,[],[f1981,f597])).
fof(f597,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X1))) != s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X2))) | s('type.list.list'(X0),X1) = s('type.list.list'(X0),X2)) )),
  inference(cnf_transformation,[],[f379])).
fof(f379,plain,(
  ! [X0,X1,X2] : (s('type.list.list'(X0),X1) = s('type.list.list'(X0),X2) | ! [X3,X4] : s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X1))) != s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X2))))),
  inference(ennf_transformation,[],[f282])).
fof(f282,plain,(
  ! [X0,X1,X2] : (s('type.list.list'(X0),X1) != s('type.list.list'(X0),X2) => ! [X3,X4] : s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X1))) != s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X2))))),
  inference(flattening,[],[f281])).
fof(f281,plain,(
  ! [X0] : ! [X1,X2] : (~s('type.list.list'(X0),X1) = s('type.list.list'(X0),X2) => ! [X3,X4] : ~s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X1))) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X2))))),
  inference(rectify,[],[f97])).
fof(f97,axiom,(
  ! [X7] : ! [X14,X15] : (~s('type.list.list'(X7),X14) = s('type.list.list'(X7),X15) => ! [X49,X50] : ~s('type.list.list'(X7),'const.list.CONS_2'(s(X7,X49),s('type.list.list'(X7),X14))) = s('type.list.list'(X7),'const.list.CONS_2'(s(X7,X50),s('type.list.list'(X7),X15))))),
  file('thm.list.MAP_EQ_NIL|split|0','thm.list.LIST_NOT_EQ')).
fof(f1981,plain,(
  s('type.list.list'(sK12),'const.list.NIL_0') != s('type.list.list'(sK12),sK68(sK13,sK12)) | $spl21),
  inference(cnf_transformation,[],[f1981_D])).
fof(f1981_D,plain,(
  s('type.list.list'(sK12),'const.list.NIL_0') != s('type.list.list'(sK12),sK68(sK13,sK12)) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f13571,plain,(
  s(sK12,sK67(sK13,sK12)) = s(sK12,sK67(sK68('const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),sK68(sK13,sK12))),sK12),sK12)) | $spl130),
  inference(cnf_transformation,[],[f13571_D])).
fof(f13571_D,plain,(
  s(sK12,sK67(sK13,sK12)) = s(sK12,sK67(sK68('const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),sK68(sK13,sK12))),sK12),sK12)) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f13573,plain,(
  ( ! [X12] : (s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK68(sK13,sK12),sK12)),s('type.list.list'(sK12),X12)))) ) | $spl134),
  inference(cnf_transformation,[],[f13573_D])).
fof(f13573_D,plain,(
  ( ! [X12] : (s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK68(sK13,sK12),sK12)),s('type.list.list'(sK12),X12)))) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f13506,plain,(
  ( ! [X12] : (s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK68(sK13,sK12),sK12)),s('type.list.list'(sK12),X12))) | s(sK12,sK67(sK13,sK12)) = s(sK12,sK67(sK68('const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),sK68(sK13,sK12))),sK12),sK12))) ) | ($spl3 | $spl8 | $spl16 | $spl21 | $spl126)),
  inference(superposition,[],[f12181,f13460])).
fof(f13460,plain,(
  s(sK12,sK67(sK68(sK13,sK12),sK12)) = s(sK12,sK67(sK68('const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),sK68(sK13,sK12))),sK12),sK12)) | ($spl3 | $spl8 | $spl21)),
  inference(subsumption_resolution,[],[f13459,f1981])).
fof(f13459,plain,(
  s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),sK68(sK13,sK12)) | s(sK12,sK67(sK68(sK13,sK12),sK12)) = s(sK12,sK67(sK68('const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),sK68(sK13,sK12))),sK12),sK12)) | ($spl3 | $spl8)),
  inference(forward_demodulation,[],[f13456,f9841])).
fof(f9841,plain,(
  s('type.list.list'(sK12),sK68(sK13,sK12)) = s('type.list.list'(sK12),sK68('const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),sK68(sK13,sK12))),sK12)) | $spl3),
  inference(subsumption_resolution,[],[f9840,f770])).
fof(f9840,plain,(
  s('type.list.list'(sK12),sK68(sK13,sK12)) = s('type.list.list'(sK12),sK68('const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),sK68(sK13,sK12))),sK12)) | s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),sK68(sK13,sK12)))) | $spl3),
  inference(trivial_inequality_removal,[],[f9813])).
fof(f9813,plain,(
  s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),sK13) | s('type.list.list'(sK12),sK68(sK13,sK12)) = s('type.list.list'(sK12),sK68('const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),sK68(sK13,sK12))),sK12)) | s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),sK68(sK13,sK12)))) | $spl3),
  inference(superposition,[],[f2078,f942])).
fof(f2078,plain,(
  ( ! [X20] : (s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),X20) | s('type.list.list'(sK12),sK68(sK13,sK12)) = s('type.list.list'(sK12),sK68(X20,sK12)) | s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),X20)) ) | $spl3),
  inference(superposition,[],[f1311,f772])).
fof(f1311,plain,(
  ( ! [X4,X5] : (s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X4),s('type.list.list'(sK12),X5))) | s('type.list.list'(sK12),X5) = s('type.list.list'(sK12),sK68(sK13,sK12))) ) | $spl3),
  inference(superposition,[],[f593,f942])).
fof(f13456,plain,(
  s(sK12,sK67(sK68(sK13,sK12),sK12)) = s(sK12,sK67(sK68('const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),sK68(sK13,sK12))),sK12),sK12)) | s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),sK68('const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),sK68(sK13,sK12))),sK12)) | ($spl3 | $spl8)),
  inference(trivial_inequality_removal,[],[f13419])).
fof(f13419,plain,(
  s('type.list.list'(sK12),sK68(sK13,sK12)) != s('type.list.list'(sK12),sK68(sK13,sK12)) | s(sK12,sK67(sK68(sK13,sK12),sK12)) = s(sK12,sK67(sK68('const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),sK68(sK13,sK12))),sK12),sK12)) | s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),sK68('const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),sK68(sK13,sK12))),sK12)) | ($spl3 | $spl8)),
  inference(superposition,[],[f9429,f9841])).
fof(f9429,plain,(
  ( ! [X20] : (s('type.list.list'(sK12),X20) != s('type.list.list'(sK12),sK68(sK13,sK12)) | s(sK12,sK67(X20,sK12)) = s(sK12,sK67(sK68(sK13,sK12),sK12)) | s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),X20)) ) | $spl8),
  inference(superposition,[],[f8846,f772])).
fof(f8846,plain,(
  ( ! [X12,X13] : (s('type.list.list'(sK12),sK68(sK13,sK12)) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X12),s('type.list.list'(sK12),X13))) | s(sK12,X12) = s(sK12,sK67(sK68(sK13,sK12),sK12))) ) | $spl8),
  inference(superposition,[],[f592,f1426])).
fof(f592,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,X1) = s(X0,X3) | s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2))) != s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X4)))) )),
  inference(cnf_transformation,[],[f450])).
fof(f12181,plain,(
  ( ! [X10,X9] : (s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X9),s('type.list.list'(sK12),X10))) | s(sK12,X9) = s(sK12,sK67(sK13,sK12))) ) | ($spl16 | $spl126)),
  inference(backward_demodulation,[],[f12130,f5170])).
fof(f5170,plain,(
  ( ! [X10,X9] : (s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X9),s('type.list.list'(sK12),X10))) | s(sK12,X9) = s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12))) ) | $spl16),
  inference(superposition,[],[f592,f1977])).
fof(f12130,plain,(
  s(sK12,sK67(sK13,sK12)) = s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)) | ($spl16 | $spl126)),
  inference(trivial_inequality_removal,[],[f11994])).
fof(f11994,plain,(
  s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) != s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) | s(sK12,sK67(sK13,sK12)) = s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)) | ($spl16 | $spl126)),
  inference(superposition,[],[f5170,f11874])).
fof(f13572,plain,(
  ( ! [X5] : (s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK68(sK13,sK12),sK12)),s('type.list.list'(sK12),X5)))) ) | $spl132),
  inference(cnf_transformation,[],[f13572_D])).
fof(f13572_D,plain,(
  ( ! [X5] : (s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK68(sK13,sK12),sK12)),s('type.list.list'(sK12),X5)))) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f13495,plain,(
  ( ! [X5] : (s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK68(sK13,sK12),sK12)),s('type.list.list'(sK12),X5))) | s(sK12,sK67(sK13,sK12)) = s(sK12,sK67(sK68('const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),sK68(sK13,sK12))),sK12),sK12))) ) | ($spl3 | $spl8 | $spl21)),
  inference(superposition,[],[f1309,f13460])).
fof(f1309,plain,(
  ( ! [X0,X1] : (s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X0),s('type.list.list'(sK12),X1))) | s(sK12,X0) = s(sK12,sK67(sK13,sK12))) ) | $spl3),
  inference(superposition,[],[f592,f942])).
fof(f11089,plain,(
  s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),sK68('const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),sK68(sK13,sK12))),sK12)) | $spl122),
  inference(cnf_transformation,[],[f11089_D])).
fof(f11089_D,plain,(
  s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),sK68('const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),sK68(sK13,sK12))),sK12)) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f12658,plain,(
  ( ! [X4] : (s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X4),s('type.list.list'(sK12),sK68(sK13,sK12))))) ) | $spl128),
  inference(cnf_transformation,[],[f12658_D])).
fof(f12658_D,plain,(
  ( ! [X4] : (s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X4),s('type.list.list'(sK12),sK68(sK13,sK12))))) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f12608,plain,(
  ( ! [X4] : (s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X4),s('type.list.list'(sK12),sK68(sK13,sK12)))) | s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),sK68('const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),sK68(sK13,sK12))),sK12))) ) | ($spl3 | $spl126)),
  inference(superposition,[],[f12016,f9841])).
fof(f12016,plain,(
  ( ! [X21,X20] : (s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X20),s('type.list.list'(sK12),X21))) | s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),X21)) ) | $spl126),
  inference(superposition,[],[f593,f11874])).
fof(f11925,plain,(
  $false | ($spl32 | $spl37)),
  inference(subsumption_resolution,[],[f11918,f3016])).
fof(f11918,plain,(
  s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) | $spl32),
  inference(equality_resolution,[],[f11772])).
fof(f11772,plain,(
  ( ! [X94] : (s('type.list.list'(sK12),'const.list.NIL_0') != s('type.list.list'(sK12),X94) | s('type.list.list'(sK12),X94) = s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12))) ) | $spl32),
  inference(superposition,[],[f770,f3013])).
fof(f11923,plain,(
  $false | ($spl32 | $spl37)),
  inference(subsumption_resolution,[],[f11922,f3016])).
fof(f11922,plain,(
  s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) | $spl32),
  inference(forward_demodulation,[],[f11920,f604])).
fof(f11920,plain,(
  ( ! [X19,X18] : (s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) = s('type.list.list'(sK12),'const.list.MAP_2'(s(fun(X18,sK12),X19),s('type.list.list'(X18),'const.list.NIL_0')))) ) | $spl32),
  inference(trivial_inequality_removal,[],[f11897])).
fof(f11897,plain,(
  ( ! [X19,X18] : (s('type.list.list'(sK12),'const.list.NIL_0') != s('type.list.list'(sK12),'const.list.NIL_0') | s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) = s('type.list.list'(sK12),'const.list.MAP_2'(s(fun(X18,sK12),X19),s('type.list.list'(X18),'const.list.NIL_0')))) ) | $spl32),
  inference(superposition,[],[f11772,f604])).
fof(f11874,plain,(
  s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),'const.list.NIL_0'))) | $spl126),
  inference(cnf_transformation,[],[f11874_D])).
fof(f11874_D,plain,(
  s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),'const.list.NIL_0'))) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f10193,plain,(
  s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),'const.list.NIL_0'))) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),sK68('const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),'const.list.NIL_0')),sK12)))) | $spl110),
  inference(cnf_transformation,[],[f10193_D])).
fof(f10193_D,plain,(
  s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),'const.list.NIL_0'))) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),sK68('const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),'const.list.NIL_0')),sK12)))) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f11704,plain,(
  s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),'const.list.NIL_0'))) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),sK68('const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),'const.list.NIL_0')),sK12)))) | s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),'const.list.NIL_0'))) | ($spl32 | $spl40)),
  inference(superposition,[],[f3013,f3020])).
fof(f4273,plain,(
  s('type.list.list'(sK12),sK68(sK68(sK68(sK13,sK12),sK12),sK12)) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK68(sK13,sK12),sK12),sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK68(sK68(sK13,sK12),sK12),sK12),sK12)))) | $spl46),
  inference(cnf_transformation,[],[f4273_D])).
fof(f4273_D,plain,(
  s('type.list.list'(sK12),sK68(sK68(sK68(sK13,sK12),sK12),sK12)) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK68(sK13,sK12),sK12),sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK68(sK68(sK13,sK12),sK12),sK12),sK12)))) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f4275,plain,(
  s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)),s('type.list.list'(sK12),'const.list.NIL_0'))) | $spl48),
  inference(cnf_transformation,[],[f4275_D])).
fof(f4275_D,plain,(
  s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)),s('type.list.list'(sK12),'const.list.NIL_0'))) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f11511,plain,(
  s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK68(sK68(sK68(sK13,sK12),sK12),sK12)) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK68(sK13,sK12),sK12),sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK68(sK68(sK13,sK12),sK12),sK12),sK12)))) | $spl16),
  inference(superposition,[],[f1977,f772])).
fof(f3013,plain,(
  ( ! [X2] : (s('type.list.list'(sK12),X2) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(X2,sK12)),s('type.list.list'(sK12),sK68(X2,sK12)))) | s('type.list.list'(sK12),X2) = s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12))) ) | $spl32),
  inference(cnf_transformation,[],[f3013_D])).
fof(f3013_D,plain,(
  ( ! [X2] : (s('type.list.list'(sK12),X2) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(X2,sK12)),s('type.list.list'(sK12),sK68(X2,sK12)))) | s('type.list.list'(sK12),X2) = s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f3014,plain,(
  ( ! [X3] : (s('type.list.list'(sK12),sK68(sK13,sK12)) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X3),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl34),
  inference(cnf_transformation,[],[f3014_D])).
fof(f3014_D,plain,(
  ( ! [X3] : (s('type.list.list'(sK12),sK68(sK13,sK12)) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X3),s('type.list.list'(sK12),'const.list.NIL_0')))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f11457,plain,(
  ( ! [X4,X3] : (s('type.list.list'(sK12),sK68(sK13,sK12)) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X4),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),X3) = s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) | s('type.list.list'(sK12),X3) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(X3,sK12)),s('type.list.list'(sK12),sK68(X3,sK12))))) ) | $spl8),
  inference(superposition,[],[f8848,f772])).
fof(f8848,plain,(
  ( ! [X17,X16] : (s('type.list.list'(sK12),sK68(sK13,sK12)) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X16),s('type.list.list'(sK12),X17))) | s('type.list.list'(sK12),X17) = s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12))) ) | $spl8),
  inference(superposition,[],[f593,f1426])).
fof(f1977,plain,(
  s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK68(sK13,sK12),sK12),sK12)))) | $spl16),
  inference(cnf_transformation,[],[f1977_D])).
fof(f1977_D,plain,(
  s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK68(sK13,sK12),sK12),sK12)))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f1979,plain,(
  s('type.list.list'(sK12),sK68(sK13,sK12)) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK68(sK13,sK12),sK12)),s('type.list.list'(sK12),'const.list.NIL_0'))) | $spl18),
  inference(cnf_transformation,[],[f1979_D])).
fof(f1979_D,plain,(
  s('type.list.list'(sK12),sK68(sK13,sK12)) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK68(sK13,sK12),sK12)),s('type.list.list'(sK12),'const.list.NIL_0'))) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f11128,plain,(
  s('type.list.list'(sK12),sK68(sK13,sK12)) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK68(sK13,sK12),sK12)),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK68(sK13,sK12),sK12),sK12)))) | $spl8),
  inference(superposition,[],[f1426,f772])).
fof(f11096,plain,(
  $false | ($spl8 | $spl37 | $spl108)),
  inference(subsumption_resolution,[],[f11085,f3016])).
fof(f11085,plain,(
  s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) | ($spl8 | $spl108)),
  inference(trivial_inequality_removal,[],[f11068])).
fof(f11068,plain,(
  s('type.list.list'(sK12),sK68(sK13,sK12)) != s('type.list.list'(sK12),sK68(sK13,sK12)) | s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) | ($spl8 | $spl108)),
  inference(superposition,[],[f10296,f10452])).
fof(f10452,plain,(
  s('type.list.list'(sK12),sK68(sK13,sK12)) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)))) | ($spl8 | $spl108)),
  inference(backward_demodulation,[],[f10410,f1426])).
fof(f10410,plain,(
  s(sK12,sK67(sK13,sK12)) = s(sK12,sK67(sK68(sK13,sK12),sK12)) | ($spl8 | $spl108)),
  inference(trivial_inequality_removal,[],[f10278])).
fof(f10278,plain,(
  s('type.list.list'(sK12),sK68(sK13,sK12)) != s('type.list.list'(sK12),sK68(sK13,sK12)) | s(sK12,sK67(sK13,sK12)) = s(sK12,sK67(sK68(sK13,sK12),sK12)) | ($spl8 | $spl108)),
  inference(superposition,[],[f8846,f10191])).
fof(f10296,plain,(
  ( ! [X19,X20] : (s('type.list.list'(sK12),sK68(sK13,sK12)) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X19),s('type.list.list'(sK12),X20))) | s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),X20)) ) | $spl108),
  inference(superposition,[],[f593,f10191])).
fof(f11090,plain,(
  ( ! [X3] : (s('type.list.list'(sK12),sK68(sK13,sK12)) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X3),s('type.list.list'(sK12),sK68(sK13,sK12))))) ) | $spl124),
  inference(cnf_transformation,[],[f11090_D])).
fof(f11090_D,plain,(
  ( ! [X3] : (s('type.list.list'(sK12),sK68(sK13,sK12)) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X3),s('type.list.list'(sK12),sK68(sK13,sK12))))) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f11040,plain,(
  ( ! [X3] : (s('type.list.list'(sK12),sK68(sK13,sK12)) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X3),s('type.list.list'(sK12),sK68(sK13,sK12)))) | s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),sK68('const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),sK68(sK13,sK12))),sK12))) ) | ($spl3 | $spl108)),
  inference(superposition,[],[f10296,f9841])).
fof(f11022,plain,(
  ( ! [X93] : (p(s(bool,X93))) ) | $spl120),
  inference(cnf_transformation,[],[f11022_D])).
fof(f11022_D,plain,(
  ( ! [X93] : (p(s(bool,X93))) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f11019,plain,(
  ( ! [X61,X66,X64,X62] : (p(s(bool,happ(s(fun(X61,bool),happ(s(fun('type.list.list'(sK11),fun(X61,bool)),X62),s('type.list.list'(sK11),'const.list.NIL_0'))),s(X61,sK30('const.list.MAP_2'(s(fun(sK12,sK11),sK14),s('type.list.list'(sK12),sK13)),X62,'type.list.list'(sK11),X61))))) | s(bool,X64) != s(bool,happ(s(fun(X61,bool),happ(s(fun('type.list.list'(sK11),fun(X61,bool)),X62),s('type.list.list'(sK11),sK29(X66,X62,'type.list.list'(sK11),X61)))),s(X61,happ(s(fun('type.list.list'(sK11),X61),X66),s('type.list.list'(sK11),sK29(X66,X62,'type.list.list'(sK11),X61)))))) | ~p(s(bool,X64))) ) | $spl114),
  inference(cnf_transformation,[],[f11019_D])).
fof(f11019_D,plain,(
  ( ! [X61,X66,X64,X62] : (p(s(bool,happ(s(fun(X61,bool),happ(s(fun('type.list.list'(sK11),fun(X61,bool)),X62),s('type.list.list'(sK11),'const.list.NIL_0'))),s(X61,sK30('const.list.MAP_2'(s(fun(sK12,sK11),sK14),s('type.list.list'(sK12),sK13)),X62,'type.list.list'(sK11),X61))))) | s(bool,X64) != s(bool,happ(s(fun(X61,bool),happ(s(fun('type.list.list'(sK11),fun(X61,bool)),X62),s('type.list.list'(sK11),sK29(X66,X62,'type.list.list'(sK11),X61)))),s(X61,happ(s(fun('type.list.list'(sK11),X61),X66),s('type.list.list'(sK11),sK29(X66,X62,'type.list.list'(sK11),X61)))))) | ~p(s(bool,X64))) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f10930,plain,(
  ( ! [X484,X486,X485,X487,X483] : (p(s(bool,happ(s(fun(X483,bool),happ(s(fun('type.list.list'(sK11),fun(X483,bool)),X484),s('type.list.list'(sK11),'const.list.NIL_0'))),s(X483,sK30('const.list.MAP_2'(s(fun(sK12,sK11),sK14),s('type.list.list'(sK12),sK13)),X484,'type.list.list'(sK11),X483))))) | p(s(bool,X485)) | ~p(s(bool,X486)) | s(bool,X486) != s(bool,happ(s(fun(X483,bool),happ(s(fun('type.list.list'(sK11),fun(X483,bool)),X484),s('type.list.list'(sK11),sK29(X487,X484,'type.list.list'(sK11),X483)))),s(X483,happ(s(fun('type.list.list'(sK11),X483),X487),s('type.list.list'(sK11),sK29(X487,X484,'type.list.list'(sK11),X483))))))) ) | $spl0),
  inference(resolution,[],[f1038,f775])).
fof(f775,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f551])).
fof(f551,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f431])).
fof(f431,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f430])).
fof(f430,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f376])).
fof(f376,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f375])).
fof(f375,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f265])).
fof(f265,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f245])).
fof(f245,axiom,(
  ! [X18,X73,X20,X74] : ((s(bool,X18) = s(bool,X73) & (p(s(bool,X73)) => s(bool,X20) = s(bool,X74))) => ((p(s(bool,X18)) => p(s(bool,X20))) <=> (p(s(bool,X73)) => p(s(bool,X74)))))),
  file('thm.list.MAP_EQ_NIL|split|0','thm.bool.IMP_CONG')).
fof(f1038,plain,(
  ( ! [X204,X205,X203] : (~p(s(bool,happ(s(fun(X203,bool),happ(s(fun('type.list.list'(sK11),fun(X203,bool)),X204),s('type.list.list'(sK11),sK29(X205,X204,'type.list.list'(sK11),X203)))),s(X203,happ(s(fun('type.list.list'(sK11),X203),X205),s('type.list.list'(sK11),sK29(X205,X204,'type.list.list'(sK11),X203))))))) | p(s(bool,happ(s(fun(X203,bool),happ(s(fun('type.list.list'(sK11),fun(X203,bool)),X204),s('type.list.list'(sK11),'const.list.NIL_0'))),s(X203,sK30('const.list.MAP_2'(s(fun(sK12,sK11),sK14),s('type.list.list'(sK12),sK13)),X204,'type.list.list'(sK11),X203)))))) ) | $spl0),
  inference(superposition,[],[f612,f781])).
fof(f612,plain,(
  ( ! [X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X3))),s(X0,sK30(X3,X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,sK29(X5,X2,X1,X0)))),s(X0,happ(s(fun(X1,X0),X5),s(X1,sK29(X5,X2,X1,X0)))))))) )),
  inference(cnf_transformation,[],[f457])).
fof(f457,plain,(
  ! [X0,X1,X2] : ((! [X10] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,sK27(X2,X1,X0)))),s(X0,X10)))) | ! [X8] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X8))),s(X0,happ(s(fun(X1,X0),sK28(X2,X1,X0)),s(X1,X8))))))) & (! [X5] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,sK29(X5,X2,X1,X0)))),s(X0,happ(s(fun(X1,X0),X5),s(X1,sK29(X5,X2,X1,X0))))))) | ! [X3] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X3))),s(X0,sK30(X3,X2,X1,X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK27,sK28,sK29,sK30])],[f456])).
fof(f456,plain,(
  ! [X0,X1,X2] : ((? [X9] : ! [X10] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X9))),s(X0,X10)))) | ? [X7] : ! [X8] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X8))),s(X0,happ(s(fun(X1,X0),X7),s(X1,X8))))))) & (! [X5] : ? [X6] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X6))),s(X0,happ(s(fun(X1,X0),X5),s(X1,X6)))))) | ! [X3] : ? [X4] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X3))),s(X0,X4))))))),
  inference(rectify,[],[f455])).
fof(f455,plain,(
  ! [X0,X1,X2] : ((? [X3] : ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X3))),s(X0,X4)))) | ? [X5] : ! [X6] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X6))),s(X0,happ(s(fun(X1,X0),X5),s(X1,X6))))))) & (! [X5] : ? [X6] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X6))),s(X0,happ(s(fun(X1,X0),X5),s(X1,X6)))))) | ! [X3] : ? [X4] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X3))),s(X0,X4))))))),
  inference(nnf_transformation,[],[f306])).
fof(f306,plain,(
  ! [X0,X1,X2] : (! [X3] : ? [X4] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X3))),s(X0,X4)))) <=> ? [X5] : ! [X6] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X6))),s(X0,happ(s(fun(X1,X0),X5),s(X1,X6)))))))),
  inference(flattening,[],[f305])).
fof(f305,plain,(
  ! [X0,X1] : ! [X2] : (! [X3] : ? [X4] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X3))),s(X0,X4)))) <=> ? [X5] : ! [X6] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X6))),s(X0,happ(s(fun(X1,X0),X5),s(X1,X6)))))))),
  inference(rectify,[],[f218])).
fof(f218,axiom,(
  ! [X6,X7] : ! [X17] : (! [X18] : ? [X20] : p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X7,fun(X6,bool)),X17),s(X7,X18))),s(X6,X20)))) <=> ? [X9] : ! [X18] : p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X7,fun(X6,bool)),X17),s(X7,X18))),s(X6,happ(s(fun(X7,X6),X9),s(X7,X18)))))))),
  file('thm.list.MAP_EQ_NIL|split|0','thm.bool.SKOLEM_THM')).
fof(f11020,plain,(
  ( ! [X78,X75] : (~p(s(bool,X75)) | s(bool,X75) != s(bool,X78)) ) | $spl116),
  inference(cnf_transformation,[],[f11020_D])).
fof(f11020_D,plain,(
  ( ! [X78,X75] : (~p(s(bool,X75)) | s(bool,X75) != s(bool,X78)) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f10852,plain,(
  ( ! [X103,X107,X105,X102,X106,X104] : (p(s(bool,happ(s(fun(X102,bool),happ(s(fun('type.list.list'(sK11),fun(X102,bool)),X103),s('type.list.list'(sK11),'const.list.NIL_0'))),s(X102,sK30('const.list.MAP_2'(s(fun(sK12,sK11),sK14),s('type.list.list'(sK12),sK13)),X103,'type.list.list'(sK11),X102))))) | ~p(s(bool,X104)) | ~p(s(bool,X105)) | s(bool,X105) != s(bool,happ(s(fun(X102,bool),happ(s(fun('type.list.list'(sK11),fun(X102,bool)),X103),s('type.list.list'(sK11),sK29(X106,X103,'type.list.list'(sK11),X102)))),s(X102,happ(s(fun('type.list.list'(sK11),X102),X106),s('type.list.list'(sK11),sK29(X106,X103,'type.list.list'(sK11),X102)))))) | s(bool,X104) != s(bool,X107)) ) | $spl0),
  inference(resolution,[],[f1038,f554])).
fof(f554,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f431])).
fof(f11021,plain,(
  ( ! [X95] : (~p(s(bool,X95))) ) | $spl118),
  inference(cnf_transformation,[],[f11021_D])).
fof(f11021_D,plain,(
  ( ! [X95] : (~p(s(bool,X95))) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f10850,plain,(
  ( ! [X94,X92,X95,X93,X91,X96] : (p(s(bool,happ(s(fun(X91,bool),happ(s(fun('type.list.list'(sK11),fun(X91,bool)),X92),s('type.list.list'(sK11),'const.list.NIL_0'))),s(X91,sK30('const.list.MAP_2'(s(fun(sK12,sK11),sK14),s('type.list.list'(sK12),sK13)),X92,'type.list.list'(sK11),X91))))) | p(s(bool,X93)) | ~p(s(bool,X94)) | ~p(s(bool,X95)) | s(bool,X94) != s(bool,happ(s(fun(X91,bool),happ(s(fun('type.list.list'(sK11),fun(X91,bool)),X92),s('type.list.list'(sK11),sK29(X96,X92,'type.list.list'(sK11),X91)))),s(X91,happ(s(fun('type.list.list'(sK11),X91),X96),s('type.list.list'(sK11),sK29(X96,X92,'type.list.list'(sK11),X91))))))) ) | $spl0),
  inference(resolution,[],[f1038,f553])).
fof(f553,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f431])).
fof(f11018,plain,(
  ( ! [X65,X63] : (p(s(bool,X63)) | s(bool,X63) != s(bool,X65)) ) | $spl112),
  inference(cnf_transformation,[],[f11018_D])).
fof(f11018_D,plain,(
  ( ! [X65,X63] : (p(s(bool,X63)) | s(bool,X63) != s(bool,X65)) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f10848,plain,(
  ( ! [X80,X83,X81,X79,X84,X82] : (p(s(bool,happ(s(fun(X79,bool),happ(s(fun('type.list.list'(sK11),fun(X79,bool)),X80),s('type.list.list'(sK11),'const.list.NIL_0'))),s(X79,sK30('const.list.MAP_2'(s(fun(sK12,sK11),sK14),s('type.list.list'(sK12),sK13)),X80,'type.list.list'(sK11),X79))))) | p(s(bool,X81)) | ~p(s(bool,X82)) | s(bool,X81) != s(bool,X83) | s(bool,X82) != s(bool,happ(s(fun(X79,bool),happ(s(fun('type.list.list'(sK11),fun(X79,bool)),X80),s('type.list.list'(sK11),sK29(X84,X80,'type.list.list'(sK11),X79)))),s(X79,happ(s(fun('type.list.list'(sK11),X79),X84),s('type.list.list'(sK11),sK29(X84,X80,'type.list.list'(sK11),X79))))))) ) | $spl0),
  inference(resolution,[],[f1038,f552])).
fof(f552,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f431])).
fof(f10847,plain,(
  ( ! [X78,X76,X74,X77,X75,X73] : (p(s(bool,happ(s(fun(X73,bool),happ(s(fun('type.list.list'(sK11),fun(X73,bool)),X74),s('type.list.list'(sK11),'const.list.NIL_0'))),s(X73,sK30('const.list.MAP_2'(s(fun(sK12,sK11),sK14),s('type.list.list'(sK12),sK13)),X74,'type.list.list'(sK11),X73))))) | ~p(s(bool,X75)) | ~p(s(bool,X76)) | s(bool,X76) != s(bool,happ(s(fun(X73,bool),happ(s(fun('type.list.list'(sK11),fun(X73,bool)),X74),s('type.list.list'(sK11),sK29(X77,X74,'type.list.list'(sK11),X73)))),s(X73,happ(s(fun('type.list.list'(sK11),X73),X77),s('type.list.list'(sK11),sK29(X77,X74,'type.list.list'(sK11),X73)))))) | s(bool,X75) != s(bool,X78)) ) | $spl0),
  inference(resolution,[],[f1038,f550])).
fof(f550,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f431])).
fof(f10845,plain,(
  ( ! [X61,X66,X64,X62,X65,X63] : (p(s(bool,happ(s(fun(X61,bool),happ(s(fun('type.list.list'(sK11),fun(X61,bool)),X62),s('type.list.list'(sK11),'const.list.NIL_0'))),s(X61,sK30('const.list.MAP_2'(s(fun(sK12,sK11),sK14),s('type.list.list'(sK12),sK13)),X62,'type.list.list'(sK11),X61))))) | p(s(bool,X63)) | ~p(s(bool,X64)) | s(bool,X63) != s(bool,X65) | s(bool,X64) != s(bool,happ(s(fun(X61,bool),happ(s(fun('type.list.list'(sK11),fun(X61,bool)),X62),s('type.list.list'(sK11),sK29(X66,X62,'type.list.list'(sK11),X61)))),s(X61,happ(s(fun('type.list.list'(sK11),X61),X66),s('type.list.list'(sK11),sK29(X66,X62,'type.list.list'(sK11),X61))))))) ) | $spl0),
  inference(resolution,[],[f1038,f548])).
fof(f548,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f431])).
fof(f10243,plain,(
  $false | ($spl21 | $spl28)),
  inference(subsumption_resolution,[],[f10236,f1981])).
fof(f10236,plain,(
  s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),sK68(sK13,sK12)) | $spl28),
  inference(equality_resolution,[],[f10089])).
fof(f10089,plain,(
  ( ! [X90] : (s('type.list.list'(sK12),'const.list.NIL_0') != s('type.list.list'(sK12),X90) | s('type.list.list'(sK12),X90) = s('type.list.list'(sK12),sK68(sK13,sK12))) ) | $spl28),
  inference(superposition,[],[f770,f2089])).
fof(f10241,plain,(
  $false | ($spl21 | $spl28)),
  inference(subsumption_resolution,[],[f10240,f1981])).
fof(f10240,plain,(
  s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),sK68(sK13,sK12)) | $spl28),
  inference(forward_demodulation,[],[f10238,f604])).
fof(f10238,plain,(
  ( ! [X19,X18] : (s('type.list.list'(sK12),sK68(sK13,sK12)) = s('type.list.list'(sK12),'const.list.MAP_2'(s(fun(X18,sK12),X19),s('type.list.list'(X18),'const.list.NIL_0')))) ) | $spl28),
  inference(trivial_inequality_removal,[],[f10215])).
fof(f10215,plain,(
  ( ! [X19,X18] : (s('type.list.list'(sK12),'const.list.NIL_0') != s('type.list.list'(sK12),'const.list.NIL_0') | s('type.list.list'(sK12),sK68(sK13,sK12)) = s('type.list.list'(sK12),'const.list.MAP_2'(s(fun(X18,sK12),X19),s('type.list.list'(X18),'const.list.NIL_0')))) ) | $spl28),
  inference(superposition,[],[f10089,f604])).
fof(f10191,plain,(
  s('type.list.list'(sK12),sK68(sK13,sK12)) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),'const.list.NIL_0'))) | $spl108),
  inference(cnf_transformation,[],[f10191_D])).
fof(f10191_D,plain,(
  s('type.list.list'(sK12),sK68(sK13,sK12)) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),'const.list.NIL_0'))) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f10024,plain,(
  s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),'const.list.NIL_0'))) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),sK68('const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),'const.list.NIL_0')),sK12)))) | s('type.list.list'(sK12),sK68(sK13,sK12)) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),'const.list.NIL_0'))) | ($spl28 | $spl40)),
  inference(superposition,[],[f2089,f3020])).
fof(f9802,plain,(
  s('type.list.list'(sK12),sK68(sK13,sK12)) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)),s('type.list.list'(sK12),'const.list.NIL_0'))) | $spl107),
  inference(cnf_transformation,[],[f9802_D])).
fof(f9802_D,plain,(
  s('type.list.list'(sK12),sK68(sK13,sK12)) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)),s('type.list.list'(sK12),'const.list.NIL_0'))) <=> ~$spl107),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl107])])).
fof(f9778,plain,(
  s('type.list.list'(sK12),sK68(sK13,sK12)) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK68(sK68(sK68(sK13,sK12),sK12),sK12)) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK68(sK13,sK12),sK12),sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK68(sK68(sK13,sK12),sK12),sK12),sK12)))) | $spl103),
  inference(superposition,[],[f9341,f772])).
fof(f9341,plain,(
  s('type.list.list'(sK12),sK68(sK13,sK12)) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK68(sK13,sK12),sK12),sK12)))) | $spl103),
  inference(cnf_transformation,[],[f9341_D])).
fof(f9341_D,plain,(
  s('type.list.list'(sK12),sK68(sK13,sK12)) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK68(sK13,sK12),sK12),sK12)))) <=> ~$spl103),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl103])])).
fof(f9721,plain,(
  $false | ($spl37 | $spl101 | $spl102)),
  inference(subsumption_resolution,[],[f9720,f3016])).
fof(f9720,plain,(
  s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) | ($spl101 | $spl102)),
  inference(subsumption_resolution,[],[f9587,f8978])).
fof(f9587,plain,(
  s('type.list.list'(sK12),sK68(sK13,sK12)) = s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) | s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) | $spl102),
  inference(superposition,[],[f772,f9342])).
fof(f9707,plain,(
  $false | ($spl37 | $spl101 | $spl102)),
  inference(subsumption_resolution,[],[f9706,f3016])).
fof(f9706,plain,(
  s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) | ($spl101 | $spl102)),
  inference(subsumption_resolution,[],[f9554,f8978])).
fof(f9554,plain,(
  s('type.list.list'(sK12),sK68(sK13,sK12)) = s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) | s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) | $spl102),
  inference(superposition,[],[f9342,f772])).
fof(f9705,plain,(
  $false | ($spl16 | $spl101 | $spl102)),
  inference(subsumption_resolution,[],[f9540,f8978])).
fof(f9540,plain,(
  s('type.list.list'(sK12),sK68(sK13,sK12)) = s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) | ($spl16 | $spl102)),
  inference(backward_demodulation,[],[f9342,f1977])).
fof(f9342,plain,(
  s('type.list.list'(sK12),sK68(sK13,sK12)) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK68(sK13,sK12),sK12),sK12)))) | $spl102),
  inference(cnf_transformation,[],[f9342_D])).
fof(f9342_D,plain,(
  s('type.list.list'(sK12),sK68(sK13,sK12)) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK68(sK13,sK12),sK12),sK12)))) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f9343,plain,(
  ( ! [X0] : (s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X0),s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12))))) ) | $spl104),
  inference(cnf_transformation,[],[f9343_D])).
fof(f9343_D,plain,(
  ( ! [X0] : (s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X0),s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12))))) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f9294,plain,(
  ( ! [X0] : (s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X0),s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)))) | s('type.list.list'(sK12),sK68(sK13,sK12)) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK68(sK13,sK12),sK12),sK12))))) ) | ($spl3 | $spl16)),
  inference(superposition,[],[f1311,f1977])).
fof(f8976,plain,(
  s(sK12,sK67(sK68(sK13,sK12),sK12)) = s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)) | $spl98),
  inference(cnf_transformation,[],[f8976_D])).
fof(f8976_D,plain,(
  s(sK12,sK67(sK68(sK13,sK12),sK12)) = s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f8978,plain,(
  s('type.list.list'(sK12),sK68(sK13,sK12)) != s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) | $spl101),
  inference(cnf_transformation,[],[f8978_D])).
fof(f8978_D,plain,(
  s('type.list.list'(sK12),sK68(sK13,sK12)) != s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) <=> ~$spl101),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl101])])).
fof(f8833,plain,(
  s('type.list.list'(sK12),sK68(sK13,sK12)) != s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) | s(sK12,sK67(sK68(sK13,sK12),sK12)) = s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)) | ($spl8 | $spl16)),
  inference(superposition,[],[f5170,f1426])).
fof(f8973,plain,(
  s('type.list.list'(sK11),'const.list.NIL_0') = s('type.list.list'(sK11),'const.list.MAP_2'(s(fun(sK12,sK11),sK14),s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)))) | $spl94),
  inference(cnf_transformation,[],[f8973_D])).
fof(f8973_D,plain,(
  s('type.list.list'(sK11),'const.list.NIL_0') = s('type.list.list'(sK11),'const.list.MAP_2'(s(fun(sK12,sK11),sK14),s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)))) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f8974,plain,(
  ( ! [X0] : (s('type.list.list'(sK12),sK68(sK13,sK12)) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X0),s('type.list.list'(sK12),sK13)))) ) | $spl96),
  inference(cnf_transformation,[],[f8974_D])).
fof(f8974_D,plain,(
  ( ! [X0] : (s('type.list.list'(sK12),sK68(sK13,sK12)) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X0),s('type.list.list'(sK12),sK13)))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f8831,plain,(
  ( ! [X0] : (s('type.list.list'(sK12),sK68(sK13,sK12)) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X0),s('type.list.list'(sK12),sK13))) | s('type.list.list'(sK11),'const.list.NIL_0') = s('type.list.list'(sK11),'const.list.MAP_2'(s(fun(sK12,sK11),sK14),s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12))))) ) | ($spl0 | $spl8)),
  inference(superposition,[],[f950,f1426])).
fof(f950,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X2),s('type.list.list'(sK12),sK13))) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X1),s('type.list.list'(sK12),X0))) | s('type.list.list'(sK11),'const.list.NIL_0') = s('type.list.list'(sK11),'const.list.MAP_2'(s(fun(sK12,sK11),sK14),s('type.list.list'(sK12),X0)))) ) | $spl0),
  inference(superposition,[],[f781,f597])).
fof(f2090,plain,(
  ( ! [X3] : (s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X3),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl30),
  inference(cnf_transformation,[],[f2090_D])).
fof(f2090_D,plain,(
  ( ! [X3] : (s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X3),s('type.list.list'(sK12),'const.list.NIL_0')))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f8760,plain,(
  ( ! [X31,X32] : (s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X32),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),X31) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(X31,sK12)),s('type.list.list'(sK12),sK68(X31,sK12))))) ) | ($spl3 | $spl11)),
  inference(trivial_inequality_removal,[],[f8741])).
fof(f8741,plain,(
  ( ! [X31,X32] : (s('type.list.list'(sK12),X31) != s('type.list.list'(sK12),X31) | s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X32),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),X31) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(X31,sK12)),s('type.list.list'(sK12),sK68(X31,sK12))))) ) | ($spl3 | $spl11)),
  inference(superposition,[],[f8381,f772])).
fof(f8381,plain,(
  ( ! [X0,X1] : (s('type.list.list'(sK12),'const.list.NIL_0') != s('type.list.list'(sK12),X1) | s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X0),s('type.list.list'(sK12),X1)))) ) | ($spl3 | $spl11)),
  inference(subsumption_resolution,[],[f8368,f1309])).
fof(f8368,plain,(
  ( ! [X0,X1] : (s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X0),s('type.list.list'(sK12),X1))) | s('type.list.list'(sK12),'const.list.NIL_0') != s('type.list.list'(sK12),X1) | s(sK12,X0) != s(sK12,sK67(sK13,sK12))) ) | $spl11),
  inference(superposition,[],[f1427,f599])).
fof(f599,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X3))) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X4))) | s('type.list.list'(X0),X3) != s('type.list.list'(X0),X4) | s(X0,X1) != s(X0,X2)) )),
  inference(cnf_transformation,[],[f381])).
fof(f381,plain,(
  ! [X0,X1,X2] : (s(X0,X1) != s(X0,X2) | ! [X3,X4] : (s('type.list.list'(X0),X3) != s('type.list.list'(X0),X4) | s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X3))) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X4)))))),
  inference(ennf_transformation,[],[f286])).
fof(f286,plain,(
  ! [X0,X1,X2] : (s(X0,X1) = s(X0,X2) => ! [X3,X4] : (s('type.list.list'(X0),X3) = s('type.list.list'(X0),X4) => s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X3))) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X4)))))),
  inference(flattening,[],[f285])).
fof(f285,plain,(
  ! [X0] : ! [X1,X2] : (s(X0,X1) = s(X0,X2) => ! [X3,X4] : (s('type.list.list'(X0),X3) = s('type.list.list'(X0),X4) => s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X3))) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X4)))))),
  inference(rectify,[],[f100])).
fof(f100,axiom,(
  ! [X7] : ! [X49,X50] : (s(X7,X49) = s(X7,X50) => ! [X14,X15] : (s('type.list.list'(X7),X14) = s('type.list.list'(X7),X15) => s('type.list.list'(X7),'const.list.CONS_2'(s(X7,X49),s('type.list.list'(X7),X14))) = s('type.list.list'(X7),'const.list.CONS_2'(s(X7,X50),s('type.list.list'(X7),X15)))))),
  file('thm.list.MAP_EQ_NIL|split|0','thm.list.EQ_LIST')).
fof(f1427,plain,(
  s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),'const.list.NIL_0'))) | $spl11),
  inference(cnf_transformation,[],[f1427_D])).
fof(f1427_D,plain,(
  s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),'const.list.NIL_0'))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f1426,plain,(
  s('type.list.list'(sK12),sK68(sK13,sK12)) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK68(sK13,sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)))) | $spl8),
  inference(cnf_transformation,[],[f1426_D])).
fof(f1426_D,plain,(
  s('type.list.list'(sK12),sK68(sK13,sK12)) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK68(sK13,sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f1428,plain,(
  s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),'const.list.NIL_0'))) | $spl10),
  inference(cnf_transformation,[],[f1428_D])).
fof(f1428_D,plain,(
  s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),'const.list.NIL_0'))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f7909,plain,(
  s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK68(sK13,sK12)) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK68(sK13,sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)))) | $spl3),
  inference(superposition,[],[f942,f772])).
fof(f8352,plain,(
  s('type.list.list'(sK11),'const.list.NIL_0') = s('type.list.list'(sK11),'const.list.MAP_2'(s(fun(sK12,sK11),sK14),s('type.list.list'(sK12),sK68(sK13,sK12)))) | $spl90),
  inference(cnf_transformation,[],[f8352_D])).
fof(f8352_D,plain,(
  s('type.list.list'(sK11),'const.list.NIL_0') = s('type.list.list'(sK11),'const.list.MAP_2'(s(fun(sK12,sK11),sK14),s('type.list.list'(sK12),sK68(sK13,sK12)))) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f8353,plain,(
  ( ! [X0] : (s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X0),s('type.list.list'(sK12),sK13)))) ) | $spl92),
  inference(cnf_transformation,[],[f8353_D])).
fof(f8353_D,plain,(
  ( ! [X0] : (s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X0),s('type.list.list'(sK12),sK13)))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f7931,plain,(
  ( ! [X0] : (s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X0),s('type.list.list'(sK12),sK13))) | s('type.list.list'(sK11),'const.list.NIL_0') = s('type.list.list'(sK11),'const.list.MAP_2'(s(fun(sK12,sK11),sK14),s('type.list.list'(sK12),sK68(sK13,sK12))))) ) | ($spl0 | $spl3)),
  inference(superposition,[],[f950,f942])).
fof(f3020,plain,(
  s(sK12,sK67(sK13,sK12)) = s(sK12,sK67('const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),'const.list.NIL_0')),sK12)) | $spl40),
  inference(cnf_transformation,[],[f3020_D])).
fof(f3020_D,plain,(
  s(sK12,sK67(sK13,sK12)) = s(sK12,sK67('const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),'const.list.NIL_0')),sK12)) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f8194,plain,(
  s(sK12,sK67(sK13,sK12)) = s(sK12,sK67('const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),'const.list.NIL_0')),sK12)) | s('type.list.list'(sK12),sK68(sK13,sK12)) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK68(sK13,sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)))) | $spl3),
  inference(superposition,[],[f2131,f772])).
fof(f2131,plain,(
  s(sK12,sK67(sK13,sK12)) = s(sK12,sK67('const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),sK68(sK13,sK12))),sK12)) | $spl3),
  inference(subsumption_resolution,[],[f2130,f770])).
fof(f2130,plain,(
  s(sK12,sK67(sK13,sK12)) = s(sK12,sK67('const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),sK68(sK13,sK12))),sK12)) | s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),sK68(sK13,sK12)))) | $spl3),
  inference(trivial_inequality_removal,[],[f2102])).
fof(f2102,plain,(
  s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),sK13) | s(sK12,sK67(sK13,sK12)) = s(sK12,sK67('const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),sK68(sK13,sK12))),sK12)) | s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),sK68(sK13,sK12)))) | $spl3),
  inference(superposition,[],[f1821,f942])).
fof(f1821,plain,(
  ( ! [X20] : (s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),X20) | s(sK12,sK67(sK13,sK12)) = s(sK12,sK67(X20,sK12)) | s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),X20)) ) | $spl3),
  inference(superposition,[],[f1309,f772])).
fof(f8350,plain,(
  $false | ($spl3 | $spl10 | $spl21)),
  inference(subsumption_resolution,[],[f8343,f1981])).
fof(f8343,plain,(
  s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),sK68(sK13,sK12)) | ($spl3 | $spl10)),
  inference(trivial_inequality_removal,[],[f8327])).
fof(f8327,plain,(
  s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),sK13) | s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),sK68(sK13,sK12)) | ($spl3 | $spl10)),
  inference(superposition,[],[f7721,f942])).
fof(f7721,plain,(
  ( ! [X17,X16] : (s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X16),s('type.list.list'(sK12),X17))) | s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),X17)) ) | $spl10),
  inference(superposition,[],[f593,f1428])).
fof(f1171,plain,(
  s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),sK68(sK13,sK12)))) | $spl4),
  inference(cnf_transformation,[],[f1171_D])).
fof(f1171_D,plain,(
  s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),sK68(sK13,sK12)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f7657,plain,(
  ( ! [X0,X1] : (s('type.list.list'('type.list.list'(sK12)),'const.list.CONS_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'('type.list.list'(sK12)),X0))) != s('type.list.list'('type.list.list'(sK12)),'const.list.CONS_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'('type.list.list'(sK12)),X1)))) ) | $spl88),
  inference(cnf_transformation,[],[f7657_D])).
fof(f7657_D,plain,(
  ( ! [X0,X1] : (s('type.list.list'('type.list.list'(sK12)),'const.list.CONS_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'('type.list.list'(sK12)),X0))) != s('type.list.list'('type.list.list'(sK12)),'const.list.CONS_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'('type.list.list'(sK12)),X1)))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f3656,plain,(
  ( ! [X0,X1] : (s('type.list.list'('type.list.list'(sK12)),'const.list.CONS_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'('type.list.list'(sK12)),X0))) != s('type.list.list'('type.list.list'(sK12)),'const.list.CONS_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'('type.list.list'(sK12)),X1))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),sK68(sK13,sK12))))) ) | $spl3),
  inference(superposition,[],[f3441,f772])).
fof(f3441,plain,(
  ( ! [X0,X1] : (s('type.list.list'('type.list.list'(sK12)),'const.list.CONS_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'('type.list.list'(sK12)),X1))) != s('type.list.list'('type.list.list'(sK12)),'const.list.CONS_2'(s('type.list.list'(sK12),sK13),s('type.list.list'('type.list.list'(sK12)),X0)))) ) | $spl3),
  inference(equality_resolution,[],[f938])).
fof(f938,plain,(
  ( ! [X14,X12,X13] : (s('type.list.list'(sK12),'const.list.NIL_0') != s('type.list.list'(sK12),X12) | s('type.list.list'('type.list.list'(sK12)),'const.list.CONS_2'(s('type.list.list'(sK12),sK13),s('type.list.list'('type.list.list'(sK12)),X14))) != s('type.list.list'('type.list.list'(sK12)),'const.list.CONS_2'(s('type.list.list'(sK12),X12),s('type.list.list'('type.list.list'(sK12)),X13)))) ) | $spl3),
  inference(superposition,[],[f784,f598])).
fof(f598,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X3))) != s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X4))) | s(X0,X1) = s(X0,X2)) )),
  inference(cnf_transformation,[],[f380])).
fof(f380,plain,(
  ! [X0,X1,X2] : (s(X0,X1) = s(X0,X2) | ! [X3,X4] : s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X3))) != s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X4))))),
  inference(ennf_transformation,[],[f284])).
fof(f284,plain,(
  ! [X0,X1,X2] : (s(X0,X1) != s(X0,X2) => ! [X3,X4] : s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X3))) != s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X4))))),
  inference(flattening,[],[f283])).
fof(f283,plain,(
  ! [X0] : ! [X1,X2] : (~s(X0,X1) = s(X0,X2) => ! [X3,X4] : ~s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X3))) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X4))))),
  inference(rectify,[],[f96])).
fof(f96,axiom,(
  ! [X7] : ! [X49,X50] : (~s(X7,X49) = s(X7,X50) => ! [X14,X15] : ~s('type.list.list'(X7),'const.list.CONS_2'(s(X7,X49),s('type.list.list'(X7),X14))) = s('type.list.list'(X7),'const.list.CONS_2'(s(X7,X50),s('type.list.list'(X7),X15))))),
  file('thm.list.MAP_EQ_NIL|split|0','thm.list.NOT_EQ_LIST')).
fof(f7656,plain,(
  ( ! [X0] : (s('type.list.list'(sK11),'const.list.CONS_2'(s(sK11,happ(s(fun(sK12,sK11),sK14),s(sK12,X0))),s('type.list.list'(sK11),'const.list.NIL_0'))) = s('type.list.list'(sK11),'const.list.MAP_2'(s(fun(sK12,sK11),sK14),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X0),s('type.list.list'(sK12),'const.list.NIL_0')))))) ) | $spl86),
  inference(cnf_transformation,[],[f7656_D])).
fof(f7656_D,plain,(
  ( ! [X0] : (s('type.list.list'(sK11),'const.list.CONS_2'(s(sK11,happ(s(fun(sK12,sK11),sK14),s(sK12,X0))),s('type.list.list'(sK11),'const.list.NIL_0'))) = s('type.list.list'(sK11),'const.list.MAP_2'(s(fun(sK12,sK11),sK14),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X0),s('type.list.list'(sK12),'const.list.NIL_0')))))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f3935,plain,(
  ( ! [X0] : (s('type.list.list'(sK11),'const.list.CONS_2'(s(sK11,happ(s(fun(sK12,sK11),sK14),s(sK12,X0))),s('type.list.list'(sK11),'const.list.NIL_0'))) = s('type.list.list'(sK11),'const.list.MAP_2'(s(fun(sK12,sK11),sK14),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X0),s('type.list.list'(sK12),'const.list.NIL_0'))))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),sK68(sK13,sK12))))) ) | $spl0),
  inference(superposition,[],[f967,f772])).
fof(f967,plain,(
  ( ! [X0] : (s('type.list.list'(sK11),'const.list.CONS_2'(s(sK11,happ(s(fun(sK12,sK11),sK14),s(sK12,X0))),s('type.list.list'(sK11),'const.list.NIL_0'))) = s('type.list.list'(sK11),'const.list.MAP_2'(s(fun(sK12,sK11),sK14),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X0),s('type.list.list'(sK12),sK13)))))) ) | $spl0),
  inference(superposition,[],[f602,f781])).
fof(f7655,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X0),s('type.list.list'(sK12),'const.list.NIL_0'))) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X1),s('type.list.list'(sK12),X2))) | s('type.list.list'(sK11),'const.list.NIL_0') = s('type.list.list'(sK11),'const.list.MAP_2'(s(fun(sK12,sK11),sK14),s('type.list.list'(sK12),X2)))) ) | $spl84),
  inference(cnf_transformation,[],[f7655_D])).
fof(f7655_D,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X0),s('type.list.list'(sK12),'const.list.NIL_0'))) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X1),s('type.list.list'(sK12),X2))) | s('type.list.list'(sK11),'const.list.NIL_0') = s('type.list.list'(sK11),'const.list.MAP_2'(s(fun(sK12,sK11),sK14),s('type.list.list'(sK12),X2)))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f4501,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X0),s('type.list.list'(sK12),'const.list.NIL_0'))) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X1),s('type.list.list'(sK12),X2))) | s('type.list.list'(sK11),'const.list.NIL_0') = s('type.list.list'(sK11),'const.list.MAP_2'(s(fun(sK12,sK11),sK14),s('type.list.list'(sK12),X2))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),sK68(sK13,sK12))))) ) | $spl0),
  inference(superposition,[],[f950,f772])).
fof(f7654,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'('type.list.list'(sK12)),'const.list.CONS_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'('type.list.list'(sK12)),X0))) != s('type.list.list'('type.list.list'(sK12)),'const.list.CONS_2'(s('type.list.list'(sK12),X1),s('type.list.list'('type.list.list'(sK12)),X2))) | s('type.list.list'(sK11),'const.list.NIL_0') = s('type.list.list'(sK11),'const.list.MAP_2'(s(fun(sK12,sK11),sK14),s('type.list.list'(sK12),X1)))) ) | $spl82),
  inference(cnf_transformation,[],[f7654_D])).
fof(f7654_D,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'('type.list.list'(sK12)),'const.list.CONS_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'('type.list.list'(sK12)),X0))) != s('type.list.list'('type.list.list'(sK12)),'const.list.CONS_2'(s('type.list.list'(sK12),X1),s('type.list.list'('type.list.list'(sK12)),X2))) | s('type.list.list'(sK11),'const.list.NIL_0') = s('type.list.list'(sK11),'const.list.MAP_2'(s(fun(sK12,sK11),sK14),s('type.list.list'(sK12),X1)))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f6734,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'('type.list.list'(sK12)),'const.list.CONS_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'('type.list.list'(sK12)),X0))) != s('type.list.list'('type.list.list'(sK12)),'const.list.CONS_2'(s('type.list.list'(sK12),X1),s('type.list.list'('type.list.list'(sK12)),X2))) | s('type.list.list'(sK11),'const.list.NIL_0') = s('type.list.list'(sK11),'const.list.MAP_2'(s(fun(sK12,sK11),sK14),s('type.list.list'(sK12),X1))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),sK68(sK13,sK12))))) ) | $spl0),
  inference(superposition,[],[f954,f772])).
fof(f954,plain,(
  ( ! [X14,X12,X13] : (s('type.list.list'('type.list.list'(sK12)),'const.list.CONS_2'(s('type.list.list'(sK12),sK13),s('type.list.list'('type.list.list'(sK12)),X14))) != s('type.list.list'('type.list.list'(sK12)),'const.list.CONS_2'(s('type.list.list'(sK12),X12),s('type.list.list'('type.list.list'(sK12)),X13))) | s('type.list.list'(sK11),'const.list.NIL_0') = s('type.list.list'(sK11),'const.list.MAP_2'(s(fun(sK12,sK11),sK14),s('type.list.list'(sK12),X12)))) ) | $spl0),
  inference(superposition,[],[f781,f598])).
fof(f7653,plain,(
  $false | ($spl0 | $spl3 | $spl20)),
  inference(subsumption_resolution,[],[f7632,f604])).
fof(f7632,plain,(
  ( ! [X29] : (s('type.list.list'(sK12),'const.list.NIL_0') != s('type.list.list'(sK12),'const.list.MAP_2'(s(fun(sK11,sK12),X29),s('type.list.list'(sK11),'const.list.NIL_0')))) ) | ($spl0 | $spl3 | $spl20)),
  inference(superposition,[],[f7534,f781])).
fof(f7534,plain,(
  ( ! [X88,X87,X89] : (s('type.list.list'(sK12),'const.list.NIL_0') != s('type.list.list'(sK12),'const.list.MAP_2'(s(fun(X87,sK12),X88),s('type.list.list'(X87),'const.list.MAP_2'(s(fun(sK12,X87),X89),s('type.list.list'(sK12),sK13)))))) ) | ($spl3 | $spl20)),
  inference(superposition,[],[f6644,f3022])).
fof(f3022,plain,(
  s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),'const.list.NIL_0'))) | ($spl3 | $spl20)),
  inference(backward_demodulation,[],[f1982,f942])).
fof(f6644,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (s('type.list.list'(sK12),'const.list.NIL_0') != s('type.list.list'(sK12),'const.list.MAP_2'(s(fun(X24,sK12),X29),s('type.list.list'(X24),'const.list.MAP_2'(s(fun(X25,X24),X26),s('type.list.list'(X25),'const.list.CONS_2'(s(X25,X27),s('type.list.list'(X25),X28)))))))) ) | ($spl3 | $spl20)),
  inference(superposition,[],[f6524,f602])).
fof(f6524,plain,(
  ( ! [X39,X41,X38,X40] : (s('type.list.list'(sK12),'const.list.NIL_0') != s('type.list.list'(sK12),'const.list.MAP_2'(s(fun(X38,sK12),X39),s('type.list.list'(X38),'const.list.CONS_2'(s(X38,X40),s('type.list.list'(X38),X41)))))) ) | ($spl3 | $spl20)),
  inference(subsumption_resolution,[],[f6500,f5451])).
fof(f5451,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X2),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X0),s('type.list.list'(sK12),X1)))))) ) | ($spl3 | $spl20)),
  inference(superposition,[],[f3606,f605])).
fof(f3606,plain,(
  ( ! [X57,X54,X56,X55,X53] : (s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X57),s('type.list.list'(sK12),'const.list.MAP_2'(s(fun(X53,sK12),X54),s('type.list.list'(X53),'const.list.CONS_2'(s(X53,X55),s('type.list.list'(X53),X56)))))))) ) | ($spl3 | $spl20)),
  inference(subsumption_resolution,[],[f3573,f770])).
fof(f3573,plain,(
  ( ! [X57,X54,X56,X55,X53] : (s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X57),s('type.list.list'(sK12),'const.list.MAP_2'(s(fun(X53,sK12),X54),s('type.list.list'(X53),'const.list.CONS_2'(s(X53,X55),s('type.list.list'(X53),X56))))))) | s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,happ(s(fun(X53,sK12),X54),s(X53,X55))),s('type.list.list'(sK12),'const.list.MAP_2'(s(fun(X53,sK12),X54),s('type.list.list'(X53),X56)))))) ) | ($spl3 | $spl20)),
  inference(superposition,[],[f3023,f602])).
fof(f3023,plain,(
  ( ! [X4,X5] : (s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X4),s('type.list.list'(sK12),X5))) | s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),X5)) ) | ($spl3 | $spl20)),
  inference(backward_demodulation,[],[f1982,f1311])).
fof(f6500,plain,(
  ( ! [X39,X41,X38,X40] : (s('type.list.list'(sK12),'const.list.NIL_0') != s('type.list.list'(sK12),'const.list.MAP_2'(s(fun(X38,sK12),X39),s('type.list.list'(X38),'const.list.CONS_2'(s(X38,X40),s('type.list.list'(X38),X41))))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,happ(s(fun(X38,sK12),X39),s(X38,X40))),s('type.list.list'(sK12),'const.list.MAP_2'(s(fun(X38,sK12),X39),s('type.list.list'(X38),X41)))))))) ) | ($spl3 | $spl20)),
  inference(superposition,[],[f6468,f602])).
fof(f6468,plain,(
  ( ! [X0] : (s('type.list.list'(sK12),'const.list.NIL_0') != s('type.list.list'(sK12),X0) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),X0)))) ) | ($spl3 | $spl20)),
  inference(equality_resolution,[],[f3197])).
fof(f3197,plain,(
  ( ! [X0,X1] : (s(sK12,X0) != s(sK12,sK67(sK13,sK12)) | s('type.list.list'(sK12),'const.list.NIL_0') != s('type.list.list'(sK12),X1) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X0),s('type.list.list'(sK12),X1)))) ) | ($spl3 | $spl20)),
  inference(superposition,[],[f3022,f599])).
fof(f4288,plain,(
  s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),sK68(sK68(sK68(sK13,sK12),sK12),sK12)) | $spl52),
  inference(cnf_transformation,[],[f4288_D])).
fof(f4288_D,plain,(
  s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),sK68(sK68(sK68(sK13,sK12),sK12),sK12)) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f7372,plain,(
  ( ! [X66] : (s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),'const.list.MAP_2'(s(fun(sK12,sK12),X66),s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12))))) ) | $spl80),
  inference(cnf_transformation,[],[f7372_D])).
fof(f7372_D,plain,(
  ( ! [X66] : (s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),'const.list.MAP_2'(s(fun(sK12,sK12),X66),s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12))))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f7348,plain,(
  ( ! [X66] : (s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),'const.list.MAP_2'(s(fun(sK12,sK12),X66),s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)))) | s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),sK68(sK68(sK68(sK13,sK12),sK12),sK12))) ) | ($spl3 | $spl16 | $spl20)),
  inference(superposition,[],[f6201,f1977])).
fof(f6201,plain,(
  ( ! [X37,X35,X38,X36] : (s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),'const.list.MAP_2'(s(fun(X35,sK12),X37),s('type.list.list'(X35),'const.list.CONS_2'(s(X35,X38),s('type.list.list'(X35),X36))))) | s('type.list.list'(X35),'const.list.NIL_0') = s('type.list.list'(X35),X36)) ) | ($spl3 | $spl20)),
  inference(superposition,[],[f5466,f772])).
fof(f5466,plain,(
  ( ! [X30,X28,X31,X29,X32] : (s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),'const.list.MAP_2'(s(fun(X28,sK12),X29),s('type.list.list'(X28),'const.list.CONS_2'(s(X28,X30),s('type.list.list'(X28),'const.list.CONS_2'(s(X28,X31),s('type.list.list'(X28),X32)))))))) ) | ($spl3 | $spl20)),
  inference(superposition,[],[f3606,f602])).
fof(f7225,plain,(
  ( ! [X8,X7,X9] : (s('type.list.list'(X7),'const.list.NIL_0') = s('type.list.list'(X7),X9) | s('type.list.list'(sK12),'const.list.MAP_2'(s(fun(X7,sK12),X8),s('type.list.list'(X7),X9))) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67('const.list.MAP_2'(s(fun(X7,sK12),X8),s('type.list.list'(X7),X9)),sK12)),s('type.list.list'(sK12),sK68('const.list.MAP_2'(s(fun(X7,sK12),X8),s('type.list.list'(X7),X9)),sK12))))) ) | $spl78),
  inference(cnf_transformation,[],[f7225_D])).
fof(f7225_D,plain,(
  ( ! [X8,X7,X9] : (s('type.list.list'(X7),'const.list.NIL_0') = s('type.list.list'(X7),X9) | s('type.list.list'(sK12),'const.list.MAP_2'(s(fun(X7,sK12),X8),s('type.list.list'(X7),X9))) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67('const.list.MAP_2'(s(fun(X7,sK12),X8),s('type.list.list'(X7),X9)),sK12)),s('type.list.list'(sK12),sK68('const.list.MAP_2'(s(fun(X7,sK12),X8),s('type.list.list'(X7),X9)),sK12))))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f7199,plain,(
  ( ! [X10,X8,X7,X9] : (s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X10),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(X7),'const.list.NIL_0') = s('type.list.list'(X7),X9) | s('type.list.list'(sK12),'const.list.MAP_2'(s(fun(X7,sK12),X8),s('type.list.list'(X7),X9))) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67('const.list.MAP_2'(s(fun(X7,sK12),X8),s('type.list.list'(X7),X9)),sK12)),s('type.list.list'(sK12),sK68('const.list.MAP_2'(s(fun(X7,sK12),X8),s('type.list.list'(X7),X9)),sK12))))) ) | ($spl3 | $spl20)),
  inference(superposition,[],[f5439,f772])).
fof(f5439,plain,(
  ( ! [X37,X35,X38,X36] : (s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X37),s('type.list.list'(sK12),'const.list.MAP_2'(s(fun(X35,sK12),X38),s('type.list.list'(X35),X36))))) | s('type.list.list'(X35),'const.list.NIL_0') = s('type.list.list'(X35),X36)) ) | ($spl3 | $spl20)),
  inference(superposition,[],[f3606,f772])).
fof(f5895,plain,(
  ( ! [X2] : (s('type.list.list'(sK12),'const.list.MAP_2'(s(fun(sK12,sK12),X2),s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)))) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67('const.list.MAP_2'(s(fun(sK12,sK12),X2),s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12))),sK12)),s('type.list.list'(sK12),sK68('const.list.MAP_2'(s(fun(sK12,sK12),X2),s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12))),sK12))))) ) | $spl76),
  inference(cnf_transformation,[],[f5895_D])).
fof(f5895_D,plain,(
  ( ! [X2] : (s('type.list.list'(sK12),'const.list.MAP_2'(s(fun(sK12,sK12),X2),s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)))) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67('const.list.MAP_2'(s(fun(sK12,sK12),X2),s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12))),sK12)),s('type.list.list'(sK12),sK68('const.list.MAP_2'(s(fun(sK12,sK12),X2),s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12))),sK12))))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f5872,plain,(
  ( ! [X2,X3] : (s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X3),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),'const.list.MAP_2'(s(fun(sK12,sK12),X2),s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)))) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67('const.list.MAP_2'(s(fun(sK12,sK12),X2),s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12))),sK12)),s('type.list.list'(sK12),sK68('const.list.MAP_2'(s(fun(sK12,sK12),X2),s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12))),sK12))))) ) | ($spl3 | $spl16 | $spl20)),
  inference(superposition,[],[f5446,f772])).
fof(f5446,plain,(
  ( ! [X78,X79] : (s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X78),s('type.list.list'(sK12),'const.list.MAP_2'(s(fun(sK12,sK12),X79),s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12))))))) ) | ($spl3 | $spl16 | $spl20)),
  inference(superposition,[],[f3606,f1977])).
fof(f5649,plain,(
  ( ! [X2] : (s('type.list.list'(sK12),'const.list.MAP_2'(s(fun(sK12,sK12),X2),s('type.list.list'(sK12),sK13))) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67('const.list.MAP_2'(s(fun(sK12,sK12),X2),s('type.list.list'(sK12),sK13)),sK12)),s('type.list.list'(sK12),sK68('const.list.MAP_2'(s(fun(sK12,sK12),X2),s('type.list.list'(sK12),sK13)),sK12))))) ) | $spl74),
  inference(cnf_transformation,[],[f5649_D])).
fof(f5649_D,plain,(
  ( ! [X2] : (s('type.list.list'(sK12),'const.list.MAP_2'(s(fun(sK12,sK12),X2),s('type.list.list'(sK12),sK13))) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67('const.list.MAP_2'(s(fun(sK12,sK12),X2),s('type.list.list'(sK12),sK13)),sK12)),s('type.list.list'(sK12),sK68('const.list.MAP_2'(s(fun(sK12,sK12),X2),s('type.list.list'(sK12),sK13)),sK12))))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f5626,plain,(
  ( ! [X2,X3] : (s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X3),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),'const.list.MAP_2'(s(fun(sK12,sK12),X2),s('type.list.list'(sK12),sK13))) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67('const.list.MAP_2'(s(fun(sK12,sK12),X2),s('type.list.list'(sK12),sK13)),sK12)),s('type.list.list'(sK12),sK68('const.list.MAP_2'(s(fun(sK12,sK12),X2),s('type.list.list'(sK12),sK13)),sK12))))) ) | ($spl3 | $spl20)),
  inference(superposition,[],[f5445,f772])).
fof(f5445,plain,(
  ( ! [X76,X77] : (s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X76),s('type.list.list'(sK12),'const.list.MAP_2'(s(fun(sK12,sK12),X77),s('type.list.list'(sK12),sK13)))))) ) | ($spl3 | $spl20)),
  inference(superposition,[],[f3606,f3022])).
fof(f5595,plain,(
  ( ! [X30,X29] : (s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X29),s('type.list.list'(sK12),X30))) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67('const.list.CONS_2'(s(sK12,X29),s('type.list.list'(sK12),X30)),sK12)),s('type.list.list'(sK12),sK68('const.list.CONS_2'(s(sK12,X29),s('type.list.list'(sK12),X30)),sK12))))) ) | $spl72),
  inference(cnf_transformation,[],[f5595_D])).
fof(f5595_D,plain,(
  ( ! [X30,X29] : (s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X29),s('type.list.list'(sK12),X30))) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67('const.list.CONS_2'(s(sK12,X29),s('type.list.list'(sK12),X30)),sK12)),s('type.list.list'(sK12),sK68('const.list.CONS_2'(s(sK12,X29),s('type.list.list'(sK12),X30)),sK12))))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f5573,plain,(
  ( ! [X30,X31,X29] : (s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X31),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X29),s('type.list.list'(sK12),X30))) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67('const.list.CONS_2'(s(sK12,X29),s('type.list.list'(sK12),X30)),sK12)),s('type.list.list'(sK12),sK68('const.list.CONS_2'(s(sK12,X29),s('type.list.list'(sK12),X30)),sK12))))) ) | ($spl3 | $spl20)),
  inference(superposition,[],[f5451,f772])).
fof(f5512,plain,(
  $false | ($spl3 | $spl20 | $spl30)),
  inference(trivial_inequality_removal,[],[f5498])).
fof(f5498,plain,(
  s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),sK13) | ($spl3 | $spl20 | $spl30)),
  inference(superposition,[],[f2090,f3022])).
fof(f5478,plain,(
  ( ! [X6,X8,X7,X9] : (s('type.list.list'(sK12),'const.list.MAP_2'(s(fun(X6,sK12),X7),s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X8),s('type.list.list'(X6),X9))))) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67('const.list.MAP_2'(s(fun(X6,sK12),X7),s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X8),s('type.list.list'(X6),X9)))),sK12)),s('type.list.list'(sK12),sK68('const.list.MAP_2'(s(fun(X6,sK12),X7),s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X8),s('type.list.list'(X6),X9)))),sK12))))) ) | $spl70),
  inference(cnf_transformation,[],[f5478_D])).
fof(f5478_D,plain,(
  ( ! [X6,X8,X7,X9] : (s('type.list.list'(sK12),'const.list.MAP_2'(s(fun(X6,sK12),X7),s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X8),s('type.list.list'(X6),X9))))) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67('const.list.MAP_2'(s(fun(X6,sK12),X7),s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X8),s('type.list.list'(X6),X9)))),sK12)),s('type.list.list'(sK12),sK68('const.list.MAP_2'(s(fun(X6,sK12),X7),s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X8),s('type.list.list'(X6),X9)))),sK12))))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f5453,plain,(
  ( ! [X6,X10,X8,X7,X9] : (s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X10),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),'const.list.MAP_2'(s(fun(X6,sK12),X7),s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X8),s('type.list.list'(X6),X9))))) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67('const.list.MAP_2'(s(fun(X6,sK12),X7),s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X8),s('type.list.list'(X6),X9)))),sK12)),s('type.list.list'(sK12),sK68('const.list.MAP_2'(s(fun(X6,sK12),X7),s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X8),s('type.list.list'(X6),X9)))),sK12))))) ) | ($spl3 | $spl20)),
  inference(superposition,[],[f3606,f772])).
fof(f5476,plain,(
  ( ! [X39,X41,X42,X40] : (s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X42),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(X39),'const.list.CONS_2'(s(X39,X40),s('type.list.list'(X39),X41))) = s('type.list.list'(X39),'const.list.CONS_2'(s(X39,sK67('const.list.CONS_2'(s(X39,X40),s('type.list.list'(X39),X41)),X39)),s('type.list.list'(X39),sK68('const.list.CONS_2'(s(X39,X40),s('type.list.list'(X39),X41)),X39))))) ) | ($spl3 | $spl20)),
  inference(forward_demodulation,[],[f5440,f604])).
fof(f5440,plain,(
  ( ! [X39,X43,X41,X42,X40] : (s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X42),s('type.list.list'(sK12),'const.list.MAP_2'(s(fun(X39,sK12),X43),s('type.list.list'(X39),'const.list.NIL_0'))))) | s('type.list.list'(X39),'const.list.CONS_2'(s(X39,X40),s('type.list.list'(X39),X41))) = s('type.list.list'(X39),'const.list.CONS_2'(s(X39,sK67('const.list.CONS_2'(s(X39,X40),s('type.list.list'(X39),X41)),X39)),s('type.list.list'(X39),sK68('const.list.CONS_2'(s(X39,X40),s('type.list.list'(X39),X41)),X39))))) ) | ($spl3 | $spl20)),
  inference(superposition,[],[f3606,f772])).
fof(f5134,plain,(
  s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK68(sK68(sK68(sK13,sK12),sK12),sK12)) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK68(sK13,sK12),sK12),sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK68(sK68(sK13,sK12),sK12),sK12),sK12)))) | $spl16),
  inference(superposition,[],[f1977,f772])).
fof(f4848,plain,(
  s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK68(sK13,sK12),sK12),sK12)))) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67('const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK68(sK13,sK12),sK12),sK12))),sK12)),s('type.list.list'(sK12),sK68('const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK68(sK13,sK12),sK12),sK12))),sK12)))) | $spl66),
  inference(cnf_transformation,[],[f4848_D])).
fof(f4848_D,plain,(
  s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK68(sK13,sK12),sK12),sK12)))) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67('const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK68(sK13,sK12),sK12),sK12))),sK12)),s('type.list.list'(sK12),sK68('const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK68(sK13,sK12),sK12),sK12))),sK12)))) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f3016,plain,(
  s('type.list.list'(sK12),'const.list.NIL_0') != s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) | $spl37),
  inference(cnf_transformation,[],[f3016_D])).
fof(f3016_D,plain,(
  s('type.list.list'(sK12),'const.list.NIL_0') != s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) <=> ~$spl37),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl37])])).
fof(f4836,plain,(
  s('type.list.list'(sK12),'const.list.NIL_0') != s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) | s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK68(sK13,sK12),sK12),sK12)))) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67('const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK68(sK13,sK12),sK12),sK12))),sK12)),s('type.list.list'(sK12),sK68('const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK68(sK13,sK12),sK12),sK12))),sK12)))) | $spl17),
  inference(superposition,[],[f1976,f772])).
fof(f1976,plain,(
  s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK68(sK13,sK12),sK12),sK12)))) | $spl17),
  inference(cnf_transformation,[],[f1976_D])).
fof(f1976_D,plain,(
  s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK68(sK13,sK12),sK12),sK12)))) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f4846,plain,(
  $false | ($spl17 | $spl37)),
  inference(subsumption_resolution,[],[f4845,f3016])).
fof(f4845,plain,(
  s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) | $spl17),
  inference(trivial_inequality_removal,[],[f4831])).
fof(f4831,plain,(
  s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) != s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) | s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) | $spl17),
  inference(superposition,[],[f1976,f772])).
fof(f4274,plain,(
  s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)),s('type.list.list'(sK12),'const.list.NIL_0'))) | $spl49),
  inference(cnf_transformation,[],[f4274_D])).
fof(f4274_D,plain,(
  s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)),s('type.list.list'(sK12),'const.list.NIL_0'))) <=> ~$spl49),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl49])])).
fof(f4822,plain,(
  s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK68(sK68(sK68(sK13,sK12),sK12),sK12)) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK68(sK13,sK12),sK12),sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK68(sK68(sK13,sK12),sK12),sK12),sK12)))) | $spl17),
  inference(superposition,[],[f1976,f772])).
fof(f4797,plain,(
  s('type.list.list'(sK11),'const.list.NIL_0') = s('type.list.list'(sK11),'const.list.MAP_2'(s(fun(sK12,sK11),sK14),s('type.list.list'(sK12),sK68(sK68(sK68(sK13,sK12),sK12),sK12)))) | $spl62),
  inference(cnf_transformation,[],[f4797_D])).
fof(f4797_D,plain,(
  s('type.list.list'(sK11),'const.list.NIL_0') = s('type.list.list'(sK11),'const.list.MAP_2'(s(fun(sK12,sK11),sK14),s('type.list.list'(sK12),sK68(sK68(sK68(sK13,sK12),sK12),sK12)))) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f4798,plain,(
  ( ! [X0] : (s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X0),s('type.list.list'(sK12),sK13)))) ) | $spl64),
  inference(cnf_transformation,[],[f4798_D])).
fof(f4798_D,plain,(
  ( ! [X0] : (s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X0),s('type.list.list'(sK12),sK13)))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f4656,plain,(
  ( ! [X0] : (s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X0),s('type.list.list'(sK12),sK13))) | s('type.list.list'(sK11),'const.list.NIL_0') = s('type.list.list'(sK11),'const.list.MAP_2'(s(fun(sK12,sK11),sK14),s('type.list.list'(sK12),sK68(sK68(sK68(sK13,sK12),sK12),sK12))))) ) | ($spl0 | $spl16)),
  inference(superposition,[],[f950,f1977])).
fof(f4634,plain,(
  s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK68(sK68(sK68(sK13,sK12),sK12),sK12)) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK68(sK13,sK12),sK12),sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK68(sK68(sK13,sK12),sK12),sK12),sK12)))) | $spl16),
  inference(superposition,[],[f1977,f772])).
fof(f4439,plain,(
  $false | ($spl16 | $spl36)),
  inference(subsumption_resolution,[],[f4319,f771])).
fof(f771,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X1)))) )),
  inference(cnf_transformation,[],[f367])).
fof(f367,plain,(
  ! [X0,X1,X2] : s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X1)))),
  inference(flattening,[],[f366])).
fof(f366,plain,(
  ! [X0] : ! [X1,X2] : ~s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X1)))),
  inference(rectify,[],[f51])).
fof(f51,axiom,(
  ! [X7] : ! [X37,X38] : ~s('type.list.list'(X7),'const.list.NIL_0') = s('type.list.list'(X7),'const.list.CONS_2'(s(X7,X38),s('type.list.list'(X7),X37)))),
  file('thm.list.MAP_EQ_NIL|split|0','thm.list.NOT_CONS_NIL')).
fof(f4319,plain,(
  s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK68(sK13,sK12),sK12),sK12)))) | ($spl16 | $spl36)),
  inference(backward_demodulation,[],[f3017,f1977])).
fof(f3017,plain,(
  s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) | $spl36),
  inference(cnf_transformation,[],[f3017_D])).
fof(f3017_D,plain,(
  s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f4278,plain,(
  s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67('const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK68(sK13,sK12),sK12),sK12))),sK12)),s('type.list.list'(sK12),sK68('const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK68(sK13,sK12),sK12),sK12))),sK12)))) | $spl50),
  inference(cnf_transformation,[],[f4278_D])).
fof(f4278_D,plain,(
  s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67('const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK68(sK13,sK12),sK12),sK12))),sK12)),s('type.list.list'(sK12),sK68('const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK68(sK13,sK12),sK12),sK12))),sK12)))) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f4310,plain,(
  s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67('const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK68(sK13,sK12),sK12),sK12))),sK12)),s('type.list.list'(sK12),sK68('const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK68(sK13,sK12),sK12),sK12))),sK12)))) | s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) | $spl16),
  inference(forward_demodulation,[],[f4214,f1977])).
fof(f4214,plain,(
  s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) | s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK68(sK13,sK12),sK12),sK12)))) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67('const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK68(sK13,sK12),sK12),sK12))),sK12)),s('type.list.list'(sK12),sK68('const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK68(sK13,sK12),sK12),sK12))),sK12)))) | $spl16),
  inference(superposition,[],[f772,f1977])).
fof(f4301,plain,(
  s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),sK68('const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK68(sK13,sK12),sK12),sK12))),sK12)) | $spl60),
  inference(cnf_transformation,[],[f4301_D])).
fof(f4301_D,plain,(
  s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),sK68('const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK68(sK13,sK12),sK12),sK12))),sK12)) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f4290,plain,(
  s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) | $spl55),
  inference(cnf_transformation,[],[f4290_D])).
fof(f4290_D,plain,(
  s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) <=> ~$spl55),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl55])])).
fof(f4299,plain,(
  s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) | s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),sK68('const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK68(sK13,sK12),sK12),sK12))),sK12)) | ($spl3 | $spl16 | $spl20)),
  inference(subsumption_resolution,[],[f4155,f770])).
fof(f4155,plain,(
  s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) | s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),sK68('const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK68(sK13,sK12),sK12),sK12))),sK12)) | s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK68(sK13,sK12),sK12),sK12)))) | ($spl3 | $spl16 | $spl20)),
  inference(superposition,[],[f3592,f1977])).
fof(f3592,plain,(
  ( ! [X20] : (s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),X20) | s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),sK68(X20,sK12)) | s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),X20)) ) | ($spl3 | $spl20)),
  inference(superposition,[],[f3023,f772])).
fof(f4297,plain,(
  s(sK12,sK67(sK13,sK12)) = s(sK12,sK67('const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK68(sK13,sK12),sK12),sK12))),sK12)) | $spl58),
  inference(cnf_transformation,[],[f4297_D])).
fof(f4297_D,plain,(
  s(sK12,sK67(sK13,sK12)) = s(sK12,sK67('const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK68(sK13,sK12),sK12),sK12))),sK12)) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f4295,plain,(
  s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) | s(sK12,sK67(sK13,sK12)) = s(sK12,sK67('const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK68(sK13,sK12),sK12),sK12))),sK12)) | ($spl3 | $spl16)),
  inference(subsumption_resolution,[],[f4151,f770])).
fof(f4151,plain,(
  s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) | s(sK12,sK67(sK13,sK12)) = s(sK12,sK67('const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK68(sK13,sK12),sK12),sK12))),sK12)) | s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK68(sK13,sK12),sK12),sK12)))) | ($spl3 | $spl16)),
  inference(superposition,[],[f1821,f1977])).
fof(f4292,plain,(
  s(sK12,sK67(sK13,sK12)) = s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)) | $spl56),
  inference(cnf_transformation,[],[f4292_D])).
fof(f4292_D,plain,(
  s(sK12,sK67(sK13,sK12)) = s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f4145,plain,(
  s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) | s(sK12,sK67(sK13,sK12)) = s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)) | ($spl3 | $spl16)),
  inference(superposition,[],[f1309,f1977])).
fof(f4144,plain,(
  s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) | s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),sK68(sK68(sK68(sK13,sK12),sK12),sK12)) | ($spl3 | $spl16 | $spl20)),
  inference(superposition,[],[f3023,f1977])).
fof(f4276,plain,(
  s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67('const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK68(sK13,sK12),sK12),sK12))),sK12)),s('type.list.list'(sK12),sK68('const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK68(sK13,sK12),sK12),sK12))),sK12)))) | s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) | $spl16),
  inference(forward_demodulation,[],[f4135,f1977])).
fof(f4135,plain,(
  s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) | s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK68(sK13,sK12),sK12),sK12)))) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67('const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK68(sK13,sK12),sK12),sK12))),sK12)),s('type.list.list'(sK12),sK68('const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK68(sK13,sK12),sK12),sK12))),sK12)))) | $spl16),
  inference(superposition,[],[f1977,f772])).
fof(f4122,plain,(
  s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK68(sK68(sK68(sK13,sK12),sK12),sK12)) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK68(sK13,sK12),sK12),sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK68(sK68(sK13,sK12),sK12),sK12),sK12)))) | $spl16),
  inference(superposition,[],[f1977,f772])).
fof(f3450,plain,(
  ( ! [X43,X42] : (s('type.list.list'('type.list.list'(sK12)),'const.list.CONS_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'('type.list.list'(sK12)),X43))) != s('type.list.list'('type.list.list'(sK12)),'const.list.CONS_2'(s('type.list.list'(sK12),sK13),s('type.list.list'('type.list.list'(sK12)),X42)))) ) | $spl44),
  inference(cnf_transformation,[],[f3450_D])).
fof(f3450_D,plain,(
  ( ! [X43,X42] : (s('type.list.list'('type.list.list'(sK12)),'const.list.CONS_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'('type.list.list'(sK12)),X43))) != s('type.list.list'('type.list.list'(sK12)),'const.list.CONS_2'(s('type.list.list'(sK12),sK13),s('type.list.list'('type.list.list'(sK12)),X42)))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f3442,plain,(
  ( ! [X43,X41,X42] : (s('type.list.list'('type.list.list'(sK12)),'const.list.CONS_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'('type.list.list'(sK12)),X43))) != s('type.list.list'('type.list.list'(sK12)),'const.list.CONS_2'(s('type.list.list'(sK12),sK13),s('type.list.list'('type.list.list'(sK12)),X42))) | s('type.list.list'(sK12),X41) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(X41,sK12)),s('type.list.list'(sK12),sK68(X41,sK12))))) ) | $spl3),
  inference(trivial_inequality_removal,[],[f3423])).
fof(f3423,plain,(
  ( ! [X43,X41,X42] : (s('type.list.list'(sK12),X41) != s('type.list.list'(sK12),X41) | s('type.list.list'('type.list.list'(sK12)),'const.list.CONS_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'('type.list.list'(sK12)),X43))) != s('type.list.list'('type.list.list'(sK12)),'const.list.CONS_2'(s('type.list.list'(sK12),sK13),s('type.list.list'('type.list.list'(sK12)),X42))) | s('type.list.list'(sK12),X41) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(X41,sK12)),s('type.list.list'(sK12),sK68(X41,sK12))))) ) | $spl3),
  inference(superposition,[],[f938,f772])).
fof(f1982,plain,(
  s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),sK68(sK13,sK12)) | $spl20),
  inference(cnf_transformation,[],[f1982_D])).
fof(f1982_D,plain,(
  s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),sK68(sK13,sK12)) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f3021,plain,(
  ( ! [X31,X32] : (s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X32),s('type.list.list'(sK12),X31))) | s('type.list.list'(sK12),X31) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(X31,sK12)),s('type.list.list'(sK12),sK68(X31,sK12))))) ) | $spl42),
  inference(cnf_transformation,[],[f3021_D])).
fof(f3021_D,plain,(
  ( ! [X31,X32] : (s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X32),s('type.list.list'(sK12),X31))) | s('type.list.list'(sK12),X31) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(X31,sK12)),s('type.list.list'(sK12),sK68(X31,sK12))))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f2053,plain,(
  ( ! [X31,X32] : (s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X32),s('type.list.list'(sK12),X31))) | s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),sK68(sK13,sK12)) | s('type.list.list'(sK12),X31) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(X31,sK12)),s('type.list.list'(sK12),sK68(X31,sK12))))) ) | $spl3),
  inference(superposition,[],[f1311,f772])).
fof(f2149,plain,(
  s(sK12,sK67(sK13,sK12)) = s(sK12,sK67('const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),'const.list.NIL_0')),sK12)) | s('type.list.list'(sK12),sK68(sK13,sK12)) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK68(sK13,sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)))) | $spl3),
  inference(superposition,[],[f2131,f772])).
fof(f3018,plain,(
  ( ! [X31,X32] : (s('type.list.list'(sK12),sK68(sK13,sK12)) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X32),s('type.list.list'(sK12),X31))) | s('type.list.list'(sK12),X31) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(X31,sK12)),s('type.list.list'(sK12),sK68(X31,sK12))))) ) | $spl38),
  inference(cnf_transformation,[],[f3018_D])).
fof(f3018_D,plain,(
  ( ! [X31,X32] : (s('type.list.list'(sK12),sK68(sK13,sK12)) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X32),s('type.list.list'(sK12),X31))) | s('type.list.list'(sK12),X31) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(X31,sK12)),s('type.list.list'(sK12),sK68(X31,sK12))))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f2976,plain,(
  ( ! [X31,X32] : (s('type.list.list'(sK12),sK68(sK13,sK12)) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X32),s('type.list.list'(sK12),X31))) | s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) | s('type.list.list'(sK12),X31) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(X31,sK12)),s('type.list.list'(sK12),sK68(X31,sK12))))) ) | $spl8),
  inference(superposition,[],[f1862,f772])).
fof(f1862,plain,(
  ( ! [X8,X7] : (s('type.list.list'(sK12),sK68(sK13,sK12)) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X7),s('type.list.list'(sK12),X8))) | s('type.list.list'(sK12),X8) = s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12))) ) | $spl8),
  inference(superposition,[],[f593,f1426])).
fof(f2968,plain,(
  ( ! [X2,X3] : (s('type.list.list'(sK12),sK68(sK13,sK12)) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X3),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),X2) = s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) | s('type.list.list'(sK12),X2) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(X2,sK12)),s('type.list.list'(sK12),sK68(X2,sK12))))) ) | $spl8),
  inference(superposition,[],[f1862,f772])).
fof(f2089,plain,(
  ( ! [X2] : (s('type.list.list'(sK12),X2) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(X2,sK12)),s('type.list.list'(sK12),sK68(X2,sK12)))) | s('type.list.list'(sK12),X2) = s('type.list.list'(sK12),sK68(sK13,sK12))) ) | $spl28),
  inference(cnf_transformation,[],[f2089_D])).
fof(f2089_D,plain,(
  ( ! [X2] : (s('type.list.list'(sK12),X2) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(X2,sK12)),s('type.list.list'(sK12),sK68(X2,sK12)))) | s('type.list.list'(sK12),X2) = s('type.list.list'(sK12),sK68(sK13,sK12))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f2045,plain,(
  ( ! [X2,X3] : (s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X3),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),X2) = s('type.list.list'(sK12),sK68(sK13,sK12)) | s('type.list.list'(sK12),X2) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(X2,sK12)),s('type.list.list'(sK12),sK68(X2,sK12))))) ) | $spl3),
  inference(superposition,[],[f1311,f772])).
fof(f2025,plain,(
  $false | ($spl8 | $spl20)),
  inference(subsumption_resolution,[],[f1982,f1875])).
fof(f1875,plain,(
  s('type.list.list'(sK12),'const.list.NIL_0') != s('type.list.list'(sK12),sK68(sK13,sK12)) | $spl8),
  inference(superposition,[],[f770,f1426])).
fof(f1984,plain,(
  s('type.list.list'(sK12),sK68(sK13,sK12)) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67('const.list.CONS_2'(s(sK12,sK67(sK68(sK13,sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12))),sK12)),s('type.list.list'(sK12),sK68('const.list.CONS_2'(s(sK12,sK67(sK68(sK13,sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12))),sK12)))) | $spl22),
  inference(cnf_transformation,[],[f1984_D])).
fof(f1984_D,plain,(
  s('type.list.list'(sK12),sK68(sK13,sK12)) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67('const.list.CONS_2'(s(sK12,sK67(sK68(sK13,sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12))),sK12)),s('type.list.list'(sK12),sK68('const.list.CONS_2'(s(sK12,sK67(sK68(sK13,sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12))),sK12)))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f2006,plain,(
  s('type.list.list'(sK12),sK68(sK13,sK12)) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67('const.list.CONS_2'(s(sK12,sK67(sK68(sK13,sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12))),sK12)),s('type.list.list'(sK12),sK68('const.list.CONS_2'(s(sK12,sK67(sK68(sK13,sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12))),sK12)))) | s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),sK68(sK13,sK12)) | $spl8),
  inference(forward_demodulation,[],[f1918,f1426])).
fof(f1918,plain,(
  s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),sK68(sK13,sK12)) | s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK68(sK13,sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)))) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67('const.list.CONS_2'(s(sK12,sK67(sK68(sK13,sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12))),sK12)),s('type.list.list'(sK12),sK68('const.list.CONS_2'(s(sK12,sK67(sK68(sK13,sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12))),sK12)))) | $spl8),
  inference(superposition,[],[f772,f1426])).
fof(f1994,plain,(
  s(sK12,sK67(sK13,sK12)) = s(sK12,sK67(sK68(sK13,sK12),sK12)) | $spl24),
  inference(cnf_transformation,[],[f1994_D])).
fof(f1994_D,plain,(
  s(sK12,sK67(sK13,sK12)) = s(sK12,sK67(sK68(sK13,sK12),sK12)) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f1996,plain,(
  s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),sK68(sK13,sK12)) | $spl27),
  inference(cnf_transformation,[],[f1996_D])).
fof(f1996_D,plain,(
  s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),sK68(sK13,sK12)) <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f1857,plain,(
  s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),sK68(sK13,sK12)) | s(sK12,sK67(sK13,sK12)) = s(sK12,sK67(sK68(sK13,sK12),sK12)) | ($spl3 | $spl8)),
  inference(superposition,[],[f1309,f1426])).
fof(f1980,plain,(
  s('type.list.list'(sK12),sK68(sK13,sK12)) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67('const.list.CONS_2'(s(sK12,sK67(sK68(sK13,sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12))),sK12)),s('type.list.list'(sK12),sK68('const.list.CONS_2'(s(sK12,sK67(sK68(sK13,sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12))),sK12)))) | s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),sK68(sK13,sK12)) | $spl8),
  inference(forward_demodulation,[],[f1848,f1426])).
fof(f1848,plain,(
  s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),sK68(sK13,sK12)) | s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK68(sK13,sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)))) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67('const.list.CONS_2'(s(sK12,sK67(sK68(sK13,sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12))),sK12)),s('type.list.list'(sK12),sK68('const.list.CONS_2'(s(sK12,sK67(sK68(sK13,sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12))),sK12)))) | $spl8),
  inference(superposition,[],[f1426,f772])).
fof(f1835,plain,(
  s('type.list.list'(sK12),sK68(sK13,sK12)) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK68(sK13,sK12),sK12)),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK68(sK68(sK13,sK12),sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK68(sK13,sK12),sK12),sK12)))) | $spl8),
  inference(superposition,[],[f1426,f772])).
fof(f1717,plain,(
  $false | $spl12),
  inference(trivial_inequality_removal,[],[f1694])).
fof(f1694,plain,(
  s('type.list.list'(sK12),'const.list.NIL_0') != s('type.list.list'(sK12),'const.list.NIL_0') | $spl12),
  inference(superposition,[],[f1551,f604])).
fof(f1551,plain,(
  ( ! [X59] : (s('type.list.list'(sK12),'const.list.NIL_0') != s('type.list.list'(sK12),X59)) ) | $spl12),
  inference(superposition,[],[f770,f1501])).
fof(f1715,plain,(
  $false | $spl12),
  inference(equality_resolution,[],[f1551])).
fof(f1502,plain,(
  ( ! [X41,X40] : (s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X40),s('type.list.list'(sK12),sK13))) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X41),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl14),
  inference(cnf_transformation,[],[f1502_D])).
fof(f1502_D,plain,(
  ( ! [X41,X40] : (s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X40),s('type.list.list'(sK12),sK13))) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X41),s('type.list.list'(sK12),'const.list.NIL_0')))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f1495,plain,(
  ( ! [X39,X41,X40] : (s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X40),s('type.list.list'(sK12),sK13))) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X41),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),X39) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(X39,sK12)),s('type.list.list'(sK12),sK68(X39,sK12))))) ) | $spl3),
  inference(trivial_inequality_removal,[],[f1476])).
fof(f1476,plain,(
  ( ! [X39,X41,X40] : (s('type.list.list'(sK12),X39) != s('type.list.list'(sK12),X39) | s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X40),s('type.list.list'(sK12),sK13))) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X41),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),X39) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(X39,sK12)),s('type.list.list'(sK12),sK68(X39,sK12))))) ) | $spl3),
  inference(superposition,[],[f934,f772])).
fof(f934,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'(sK12),'const.list.NIL_0') != s('type.list.list'(sK12),X0) | s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X2),s('type.list.list'(sK12),sK13))) != s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X1),s('type.list.list'(sK12),X0)))) ) | $spl3),
  inference(superposition,[],[f784,f597])).
fof(f1287,plain,(
  s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK68(sK13,sK12)) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK68(sK13,sK12),sK12)),s('type.list.list'(sK12),sK68(sK68(sK13,sK12),sK12)))) | $spl3),
  inference(superposition,[],[f942,f772])).
fof(f1173,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK11),'const.list.NIL_0'))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))) | $spl6),
  inference(cnf_transformation,[],[f1173_D])).
fof(f1173_D,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK11),'const.list.NIL_0'))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f1102,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK11),'const.list.NIL_0'))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK67(sK13,sK12)),s('type.list.list'(sK12),sK68(sK13,sK12)))) | $spl0),
  inference(superposition,[],[f966,f772])).
fof(f966,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK11),'const.list.NIL_0'))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13))) | $spl0),
  inference(superposition,[],[f603,f781])).
fof(f603,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X2))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X3),s('type.list.list'(X1),X2)))))) )),
  inference(cnf_transformation,[],[f294])).
fof(f294,plain,(
  ! [X0,X1,X2,X3] : s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X2))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X3),s('type.list.list'(X1),X2)))))),
  inference(flattening,[],[f293])).
fof(f293,plain,(
  ! [X0,X1] : ! [X2,X3] : s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X2))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X3),s('type.list.list'(X1),X2)))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X6,X7] : ! [X8,X9] : s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X7),X8))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),'const.list.MAP_2'(s(fun(X7,X6),X9),s('type.list.list'(X7),X8)))))),
  file('thm.list.MAP_EQ_NIL|split|0','thm.list.LENGTH_MAP')).
fof(f932,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f931,f604])).
fof(f931,plain,(
  s('type.list.list'(sK11),'const.list.NIL_0') != s('type.list.list'(sK11),'const.list.MAP_2'(s(fun(sK12,sK11),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f782,f783])).
fof(f781,plain,(
  s('type.list.list'(sK11),'const.list.NIL_0') = s('type.list.list'(sK11),'const.list.MAP_2'(s(fun(sK12,sK11),sK14),s('type.list.list'(sK12),sK13))) | $spl0),
  inference(cnf_transformation,[],[f781_D])).
fof(f781_D,plain,(
  s('type.list.list'(sK11),'const.list.NIL_0') = s('type.list.list'(sK11),'const.list.MAP_2'(s(fun(sK12,sK11),sK14),s('type.list.list'(sK12),sK13))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f783,plain,(
  s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),sK13) | $spl2),
  inference(cnf_transformation,[],[f783_D])).
fof(f783_D,plain,(
  s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),sK13) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f545,plain,(
  s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),sK13) | s('type.list.list'(sK11),'const.list.NIL_0') = s('type.list.list'(sK11),'const.list.MAP_2'(s(fun(sK12,sK11),sK14),s('type.list.list'(sK12),sK13)))),
  inference(cnf_transformation,[],[f429])).
fof(f429,plain,(
  (s('type.list.list'(sK11),'const.list.NIL_0') = s('type.list.list'(sK11),'const.list.MAP_2'(s(fun(sK12,sK11),sK14),s('type.list.list'(sK12),sK13))) | s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),sK13)) & (s('type.list.list'(sK11),'const.list.NIL_0') != s('type.list.list'(sK11),'const.list.MAP_2'(s(fun(sK12,sK11),sK14),s('type.list.list'(sK12),sK13))) | s('type.list.list'(sK12),'const.list.NIL_0') != s('type.list.list'(sK12),sK13))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK11,sK12,sK13,sK14])],[f428])).
fof(f428,plain,(
  ? [X0,X1,X2,X3] : ((s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X3),s('type.list.list'(X1),X2))) | s('type.list.list'(X1),'const.list.NIL_0') = s('type.list.list'(X1),X2)) & (s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X3),s('type.list.list'(X1),X2))) | s('type.list.list'(X1),'const.list.NIL_0') != s('type.list.list'(X1),X2)))),
  inference(nnf_transformation,[],[f374])).
fof(f374,plain,(
  ? [X0,X1,X2,X3] : (s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X3),s('type.list.list'(X1),X2))) <~> s('type.list.list'(X1),'const.list.NIL_0') = s('type.list.list'(X1),X2))),
  inference(ennf_transformation,[],[f264])).
fof(f264,plain,(
  ~! [X0,X1,X2,X3] : (s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X3),s('type.list.list'(X1),X2))) <=> s('type.list.list'(X1),'const.list.NIL_0') = s('type.list.list'(X1),X2))),
  inference(flattening,[],[f263])).
fof(f263,plain,(
  ~! [X0,X1] : ! [X2,X3] : (s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X3),s('type.list.list'(X1),X2))) <=> s('type.list.list'(X1),'const.list.NIL_0') = s('type.list.list'(X1),X2))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9] : (s('type.list.list'(X6),'const.list.NIL_0') = s('type.list.list'(X6),'const.list.MAP_2'(s(fun(X7,X6),X9),s('type.list.list'(X7),X8))) <=> s('type.list.list'(X7),'const.list.NIL_0') = s('type.list.list'(X7),X8))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9] : (s('type.list.list'(X6),'const.list.NIL_0') = s('type.list.list'(X6),'const.list.MAP_2'(s(fun(X7,X6),X9),s('type.list.list'(X7),X8))) <=> s('type.list.list'(X7),'const.list.NIL_0') = s('type.list.list'(X7),X8))),
  file('thm.list.MAP_EQ_NIL|split|0','thm.list.MAP_EQ_NIL|split|0')).
fof(f782,plain,(
  s('type.list.list'(sK11),'const.list.NIL_0') != s('type.list.list'(sK11),'const.list.MAP_2'(s(fun(sK12,sK11),sK14),s('type.list.list'(sK12),sK13))) | $spl1),
  inference(cnf_transformation,[],[f782_D])).
fof(f782_D,plain,(
  s('type.list.list'(sK11),'const.list.NIL_0') != s('type.list.list'(sK11),'const.list.MAP_2'(s(fun(sK12,sK11),sK14),s('type.list.list'(sK12),sK13))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f784,plain,(
  s('type.list.list'(sK12),'const.list.NIL_0') != s('type.list.list'(sK12),sK13) | $spl3),
  inference(cnf_transformation,[],[f784_D])).
fof(f784_D,plain,(
  s('type.list.list'(sK12),'const.list.NIL_0') != s('type.list.list'(sK12),sK13) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f546,plain,(
  s('type.list.list'(sK12),'const.list.NIL_0') != s('type.list.list'(sK12),sK13) | s('type.list.list'(sK11),'const.list.NIL_0') != s('type.list.list'(sK11),'const.list.MAP_2'(s(fun(sK12,sK11),sK14),s('type.list.list'(sK12),sK13)))),
  inference(cnf_transformation,[],[f429])).
% SZS output end Proof for thm.list
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 323
% Passive clauses: 4951
% Generated clauses: 16288
% Final active clauses: 246
% Final passive clauses: 3082
% Input formulas: 261
% Initial clauses: 230
% 
% Selected by SInE selection: 61
% SInE iterations: 1
% 
% Duplicate literals: 10
% Trivial inequalities: 117
% Fw subsumption resolutions: 239
% Fw demodulations: 2070
% Bw demodulations: 16
% 
% Simple tautologies: 154
% Equational tautologies: 7
% Forward subsumptions: 8675
% Fw demodulations to eq. taut.: 13
% 
% Binary resolution: 105
% Forward superposition: 7324
% Backward superposition: 6252
% Equality resolution: 38
% 
% Split clauses: 68
% Split components: 137
% SAT solver clauses: 88
% SAT solver unit clauses: 3
% SAT solver binary clauses: 11
% 
% Sat splits: 68
% Sat splitting refutations: 19
% 
% Memory used [KB]: 14328
% Time elapsed: 0.546 s
% ------------------------------
% ------------------------------
% Success in time 12.27 s
