Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.quantHeuristics
% (66511)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 119
% Passive clauses: 4320
% Generated clauses: 12118
% Final active clauses: 115
% Final passive clauses: 4029
% Input formulas: 41
% Initial clauses: 104
% 
% Duplicate literals: 436
% Trivial inequalities: 212
% Fw subsumption resolutions: 1852
% Fw demodulations: 130
% Bw demodulations: 4
% 
% Simple tautologies: 299
% Equational tautologies: 71
% Forward subsumptions: 3644
% 
% Binary resolution: 950
% Factoring: 4
% Forward superposition: 3377
% Backward superposition: 4490
% Self superposition: 210
% Equality factoring: 168
% Equality resolution: 21
% 
% Split clauses: 444
% Split components: 932
% SAT solver clauses: 1044
% SAT solver unit clauses: 41
% SAT solver binary clauses: 388
% 
% Sat splits: 444
% 
% TWLsolver clauses: 1044
% TWLsolver calls for satisfiability: 208
% 
% Memory used [KB]: 8699
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.quantHeuristics
% (66535)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 197
% Passive clauses: 2025
% Generated clauses: 15256
% Final active clauses: 127
% Final passive clauses: 1075
% Input formulas: 41
% Initial clauses: 104
% 
% Duplicate literals: 406
% Trivial inequalities: 118
% Fw subsumption resolutions: 1654
% Fw demodulations: 1890
% Bw demodulations: 766
% Condensations: 668
% Global subsumptions: 13
% 
% Simple tautologies: 141
% Equational tautologies: 32
% Forward subsumptions: 6388
% Bw demodulations to eq. taut.: 53
% 
% Binary resolution: 766
% Factoring: 3
% Forward superposition: 3173
% Backward superposition: 3538
% Self superposition: 72
% Equality factoring: 52
% Equality resolution: 11
% 
% Split clauses: 58
% Split components: 126
% SAT solver clauses: 1974
% SAT solver unit clauses: 117
% SAT solver binary clauses: 50
% SAT solver learnt clauses: 26
% 
% Sat splits: 58
% Sat splitting refutations: 54
% 
% TWLsolver clauses: 1922
% TWLsolver calls for satisfiability: 578
% 
% Memory used [KB]: 13304
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.quantHeuristics
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.quantHeuristics
% SZS output start Proof for thm.quantHeuristics
fof(f178846,plain,(
  $false),
  inference(subsumption_resolution,[],[f178677,f4242])).
fof(f4242,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.bool.?|33|_1'(s(fun(sK2,bool),happ(s(fun(sK1,fun(sK2,bool)),sK4),s(sK1,happ(s(fun(sK0,sK1),sK3),s(sK0,X0)))))))) )),
  inference(unit_resulting_resolution,[],[f186,f4212,f299])).
fof(f299,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,X0) != s(bool,X1) | p(s(bool,X1))) )),
  inference(condensation,[],[f298])).
fof(f298,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2)) )),
  inference(condensation,[],[f287])).
fof(f287,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f183])).
fof(f183,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f182])).
fof(f182,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f118])).
fof(f118,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f98])).
fof(f98,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X15,X18,X11,X19] : ((s(bool,X15) = s(bool,X18) & (p(s(bool,X18)) => s(bool,X11) = s(bool,X19))) => ((p(s(bool,X15)) => p(s(bool,X11))) <=> (p(s(bool,X18)) => p(s(bool,X19)))))),
  file('thm.quantHeuristics.GUESS_RULES_EXISTS_UNIQUE|split|0','thm.bool.IMP_CONG')).
fof(f4212,plain,(
  ( ! [X24] : (~p(s(bool,'const.bool.?|33|_1'(s(fun(sK2,bool),happ(s(fun(sK1,fun(sK2,bool)),sK4),s(sK1,happ(s(fun(sK0,sK1),sK3),s(sK0,X24))))))))) )),
  inference(subsumption_resolution,[],[f4154,f305])).
fof(f305,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f189,f186])).
fof(f189,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f121])).
fof(f121,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f45])).
fof(f45,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f23])).
fof(f23,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.quantHeuristics.GUESS_RULES_EXISTS_UNIQUE|split|0','thm.bool.NOT_CLAUSES|split|1')).
fof(f4154,plain,(
  ( ! [X24] : (p(s(bool,'F_0')) | ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK2,bool),happ(s(fun(sK1,fun(sK2,bool)),sK4),s(sK1,happ(s(fun(sK0,sK1),sK3),s(sK0,X24))))))))) )),
  inference(superposition,[],[f236,f3783])).
fof(f3783,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,happ(s(fun(sK2,bool),happ(s(fun(sK1,fun(sK2,bool)),sK4),s(sK1,happ(s(fun(sK0,sK1),sK3),s(sK0,X1))))),s(sK2,X0)))) )),
  inference(superposition,[],[f279,f2518])).
fof(f2518,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),'_dst_x1x0_2'(s(fun(sK1,fun(sK2,bool)),sK4),s(sK2,X0))),s(sK1,happ(s(fun(sK0,sK1),sK3),s(sK0,X1)))))) )),
  inference(forward_demodulation,[],[f2510,f323])).
fof(f323,plain,(
  s(bool,'F_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK0,sK1),sK3),s(fun(sK1,bool),'_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4)))))),
  inference(unit_resulting_resolution,[],[f185,f200])).
fof(f200,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f51])).
fof(f51,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X14] : (s(bool,'F_0') = s(bool,X14) <=> ~p(s(bool,X14)))),
  file('thm.quantHeuristics.GUESS_RULES_EXISTS_UNIQUE|split|0','thm.bool.EQ_CLAUSES|split|2')).
fof(f185,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK0,sK1),sK3),s(fun(sK1,bool),'_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4))))))),
  inference(cnf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X5] : p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK0,sK1),sK3),s(fun(sK1,bool),'_dst_x1x0_2'(s(fun(sK1,fun(sK2,bool)),sK4),s(sK2,X5)))))) & ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK0,sK1),sK3),s(fun(sK1,bool),'_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4])],[f99])).
fof(f99,plain,(
  ? [X0,X1,X2,X3,X4] : (! [X5] : p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X0,X1),X3),s(fun(X1,bool),'_dst_x1x0_2'(s(fun(X1,fun(X2,bool)),X4),s(X2,X5)))))) & ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X0,X1),X3),s(fun(X1,bool),'_dst_x0x1_2'(s(fun(fun(X2,bool),bool),'const.bool.?|33|_0'),s(fun(X1,fun(X2,bool)),X4)))))))),
  inference(ennf_transformation,[],[f44])).
fof(f44,plain,(
  ~! [X0,X1,X2,X3,X4] : (! [X5] : p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X0,X1),X3),s(fun(X1,bool),'_dst_x1x0_2'(s(fun(X1,fun(X2,bool)),X4),s(X2,X5)))))) => p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X0,X1),X3),s(fun(X1,bool),'_dst_x0x1_2'(s(fun(fun(X2,bool),bool),'const.bool.?|33|_0'),s(fun(X1,fun(X2,bool)),X4)))))))),
  inference(flattening,[],[f43])).
fof(f43,plain,(
  ~! [X0,X1,X2] : ! [X3,X4] : (! [X5] : p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X0,X1),X3),s(fun(X1,bool),'_dst_x1x0_2'(s(fun(X1,fun(X2,bool)),X4),s(X2,X5)))))) => p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X0,X1),X3),s(fun(X1,bool),'_dst_x0x1_2'(s(fun(fun(X2,bool),bool),'const.bool.?|33|_0'),s(fun(X1,fun(X2,bool)),X4)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : ! [X9,X10] : (! [X11] : p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X6,X7),X9),s(fun(X7,bool),'_dst_x1x0_2'(s(fun(X7,fun(X8,bool)),X10),s(X8,X11)))))) => p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X6,X7),X9),s(fun(X7,bool),'_dst_x0x1_2'(s(fun(fun(X8,bool),bool),'const.bool.?|33|_0'),s(fun(X7,fun(X8,bool)),X10)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : ! [X9,X10] : (! [X11] : p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X6,X7),X9),s(fun(X7,bool),'_dst_x1x0_2'(s(fun(X7,fun(X8,bool)),X10),s(X8,X11)))))) => p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X6,X7),X9),s(fun(X7,bool),'_dst_x0x1_2'(s(fun(fun(X8,bool),bool),'const.bool.?|33|_0'),s(fun(X7,fun(X8,bool)),X10)))))))),
  file('thm.quantHeuristics.GUESS_RULES_EXISTS_UNIQUE|split|0','thm.quantHeuristics.GUESS_RULES_EXISTS_UNIQUE|split|0')).
fof(f2510,plain,(
  ( ! [X0,X1] : (s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK0,sK1),sK3),s(fun(sK1,bool),'_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4))))) = s(bool,happ(s(fun(sK1,bool),'_dst_x1x0_2'(s(fun(sK1,fun(sK2,bool)),sK4),s(sK2,X0))),s(sK1,happ(s(fun(sK0,sK1),sK3),s(sK0,X1)))))) )),
  inference(unit_resulting_resolution,[],[f324,f2497,f361])).
fof(f361,plain,(
  ( ! [X4,X5] : (p(s(bool,X5)) | s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f357,f305])).
fof(f357,plain,(
  ( ! [X4,X5] : (p(s(bool,'F_0')) | p(s(bool,X5)) | s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f251,f193])).
fof(f193,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.quantHeuristics.GUESS_RULES_EXISTS_UNIQUE|split|0',aHL_BOOL_CASES)).
fof(f251,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f110])).
fof(f110,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f109])).
fof(f109,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X12,X13] : ((p(s(bool,X12)) => p(s(bool,X13))) => ((p(s(bool,X13)) => p(s(bool,X12))) => s(bool,X12) = s(bool,X13)))),
  file('thm.quantHeuristics.GUESS_RULES_EXISTS_UNIQUE|split|0','thm.bool.IMP_ANTISYM_AX')).
fof(f2497,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(sK1,bool),'_dst_x1x0_2'(s(fun(sK1,fun(sK2,bool)),sK4),s(sK2,X0))),s(sK1,happ(s(fun(sK0,sK1),sK3),s(sK0,X1))))))) )),
  inference(subsumption_resolution,[],[f2487,f186])).
fof(f2487,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(sK1,bool),'_dst_x1x0_2'(s(fun(sK1,fun(sK2,bool)),sK4),s(sK2,X0))),s(sK1,happ(s(fun(sK0,sK1),sK3),s(sK0,X1))))))) )),
  inference(superposition,[],[f270,f314])).
fof(f314,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK0,sK1),sK3),s(fun(sK1,bool),'_dst_x1x0_2'(s(fun(sK1,fun(sK2,bool)),sK4),s(sK2,X0)))))) )),
  inference(unit_resulting_resolution,[],[f184,f196])).
fof(f196,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X14] : (s(bool,'T_0') = s(bool,X14) <=> p(s(bool,X14)))),
  file('thm.quantHeuristics.GUESS_RULES_EXISTS_UNIQUE|split|0','thm.bool.EQ_CLAUSES|split|1')).
fof(f184,plain,(
  ( ! [X5] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK0,sK1),sK3),s(fun(sK1,bool),'_dst_x1x0_2'(s(fun(sK1,fun(sK2,bool)),sK4),s(sK2,X5))))))) )),
  inference(cnf_transformation,[],[f120])).
fof(f270,plain,(
  ( ! [X2,X0,X5,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) )),
  inference(cnf_transformation,[],[f179])).
fof(f179,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) & (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK20(X3,X2,X1,X0))))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK20])],[f178])).
fof(f178,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) & (? [X4] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f177])).
fof(f177,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))))))) & (? [X4] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(nnf_transformation,[],[f92])).
fof(f92,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))))),
  inference(flattening,[],[f91])).
fof(f91,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X6,X8] : ! [X22,X23] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X8,X6),X22),s(fun(X6,bool),X23)))) <=> ! [X21] : ~p(s(bool,happ(s(fun(X6,bool),X23),s(X6,happ(s(fun(X8,X6),X22),s(X8,X21)))))))),
  file('thm.quantHeuristics.GUESS_RULES_EXISTS_UNIQUE|split|0','thm.quantHeuristics.GUESS_REWRITES|split|3')).
fof(f324,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK0,sK1),sK3),s(fun(sK1,bool),'_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4)))))),
  inference(unit_resulting_resolution,[],[f185,f195])).
fof(f195,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f123])).
fof(f279,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),'_dst_x1x0_2'(s(fun(X1,fun(X2,X0)),X3),s(X2,X4))),s(X1,X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X4)))) )),
  inference(cnf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X1,X0),'_dst_x1x0_2'(s(fun(X1,fun(X2,X0)),X3),s(X2,X4))),s(X1,X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X4)))),
  inference(flattening,[],[f96])).
fof(f96,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X1,X0),'_dst_x1x0_2'(s(fun(X1,fun(X2,X0)),X3),s(X2,X4))),s(X1,X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X4)))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X25,X26,X27] : ! [X28,X29,X30] : s(X25,happ(s(fun(X26,X25),'_dst_x1x0_2'(s(fun(X26,fun(X27,X25)),X28),s(X27,X29))),s(X26,X30))) = s(X25,happ(s(fun(X27,X25),happ(s(fun(X26,fun(X27,X25)),X28),s(X26,X30))),s(X27,X29)))),
  file('thm.quantHeuristics.GUESS_RULES_EXISTS_UNIQUE|split|0','_dst_x1x0_ax')).
fof(f236,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK5(X1,X0))))) | ~p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f151])).
fof(f151,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK5(X1,X0))))) & ! [X5,X6] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) | s(X0,X5) = s(X0,X6)))) & (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK6(X1,X0))))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK7(X1,X0))))) & s(X0,sK6(X1,X0)) != s(X0,sK7(X1,X0))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5,sK6,sK7])],[f150])).
fof(f150,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) | (? [X7] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X7)))) & ! [X5,X6] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) | s(X0,X5) = s(X0,X6)))) & (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ? [X2,X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & s(X0,X2) != s(X0,X3)) | p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1))))))),
  inference(rectify,[],[f149])).
fof(f149,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) | (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ! [X2,X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | s(X0,X2) = s(X0,X3)))) & (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ? [X2,X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & s(X0,X2) != s(X0,X3)) | p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1))))))),
  inference(flattening,[],[f148])).
fof(f148,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) | (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ! [X2,X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | s(X0,X2) = s(X0,X3)))) & ((! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ? [X2,X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & s(X0,X2) != s(X0,X3))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1))))))),
  inference(nnf_transformation,[],[f107])).
fof(f107,plain,(
  ! [X0,X1] : (p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) <=> (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ! [X2,X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | s(X0,X2) = s(X0,X3))))),
  inference(flattening,[],[f106])).
fof(f106,plain,(
  ! [X0,X1] : (p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) <=> (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ! [X2,X3] : ((~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | s(X0,X2) = s(X0,X3))))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0,X1] : (p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) <=> (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ! [X2,X3] : ((p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => s(X0,X2) = s(X0,X3))))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) <=> (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ! [X2,X3] : ((p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => s(X0,X2) = s(X0,X3))))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X8] : ! [X10] : (p(s(bool,'const.bool.?|33|_1'(s(fun(X8,bool),X10)))) <=> (? [X15] : p(s(bool,happ(s(fun(X8,bool),X10),s(X8,X15)))) & ! [X15,X11] : ((p(s(bool,happ(s(fun(X8,bool),X10),s(X8,X15)))) & p(s(bool,happ(s(fun(X8,bool),X10),s(X8,X11))))) => s(X8,X11) = s(X8,X15))))),
  file('thm.quantHeuristics.GUESS_RULES_EXISTS_UNIQUE|split|0','thm.bool.EXISTS_UNIQUE_THM')).
fof(f186,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.quantHeuristics.GUESS_RULES_EXISTS_UNIQUE|split|0',aHL_TRUTH)).
fof(f178677,plain,(
  s(bool,'T_0') = s(bool,'const.bool.?|33|_1'(s(fun(sK2,bool),happ(s(fun(sK1,fun(sK2,bool)),sK4),s(sK1,happ(s(fun(sK0,sK1),sK3),s(sK0,sK20('_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4)),sK3,sK0,sK1))))))))),
  inference(superposition,[],[f156326,f14077])).
fof(f14077,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),'_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4))),s(sK1,happ(s(fun(sK0,sK1),sK3),s(sK0,sK20('_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4)),sK3,sK0,sK1))))))),
  inference(subsumption_resolution,[],[f14062,f305])).
fof(f14062,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),'_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4))),s(sK1,happ(s(fun(sK0,sK1),sK3),s(sK0,sK20('_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4)),sK3,sK0,sK1))))))),
  inference(superposition,[],[f2930,f323])).
fof(f2930,plain,(
  ( ! [X6,X4,X5,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X5,X3),X6),s(fun(X3,bool),X4)))) | s(bool,'T_0') = s(bool,happ(s(fun(X3,bool),X4),s(X3,happ(s(fun(X5,X3),X6),s(X5,sK20(X4,X6,X5,X3))))))) )),
  inference(subsumption_resolution,[],[f2921,f305])).
fof(f2921,plain,(
  ( ! [X6,X4,X5,X3] : (p(s(bool,'F_0')) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X5,X3),X6),s(fun(X3,bool),X4)))) | s(bool,'T_0') = s(bool,happ(s(fun(X3,bool),X4),s(X3,happ(s(fun(X5,X3),X6),s(X5,sK20(X4,X6,X5,X3))))))) )),
  inference(superposition,[],[f271,f193])).
fof(f271,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK20(X3,X2,X1,X0))))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f179])).
fof(f156326,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'const.bool.?|33|_1'(s(fun(X28,bool),happ(s(fun(X29,fun(X28,bool)),X30),s(X29,X31))))) = s(bool,happ(s(fun(X29,bool),'_dst_x0x1_2'(s(fun(fun(X28,bool),bool),'const.bool.?|33|_0'),s(fun(X29,fun(X28,bool)),X30))),s(X29,X31)))) )),
  inference(superposition,[],[f235,f278])).
fof(f278,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))) )),
  inference(cnf_transformation,[],[f95])).
fof(f95,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(flattening,[],[f94])).
fof(f94,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X25,X27,X26] : ! [X28,X29,X30] : s(X25,happ(s(fun(X27,X25),X28),s(X27,happ(s(fun(X26,X27),X29),s(X26,X30))))) = s(X25,happ(s(fun(X26,X25),'_dst_x0x1_2'(s(fun(X27,X25),X28),s(fun(X26,X27),X29))),s(X26,X30)))),
  file('thm.quantHeuristics.GUESS_RULES_EXISTS_UNIQUE|split|0','_dst_x0x1_ax')).
fof(f235,plain,(
  ( ! [X0,X1] : (s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(fun(X0,bool),bool),'const.bool.?|33|_0'),s(fun(X0,bool),X1)))) )),
  inference(cnf_transformation,[],[f69])).
fof(f69,plain,(
  ! [X0,X1] : s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(fun(X0,bool),bool),'const.bool.?|33|_0'),s(fun(X0,bool),X1)))),
  inference(flattening,[],[f68])).
fof(f68,plain,(
  ! [X0] : ! [X1] : s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(fun(X0,bool),bool),'const.bool.?|33|_0'),s(fun(X0,bool),X1)))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X8] : ! [X24] : s(bool,'const.bool.?|33|_1'(s(fun(X8,bool),X24))) = s(bool,happ(s(fun(fun(X8,bool),bool),'const.bool.?|33|_0'),s(fun(X8,bool),X24)))),
  file('thm.quantHeuristics.GUESS_RULES_EXISTS_UNIQUE|split|0','ARITY_0_1_const.bool.?|33|')).
% SZS output end Proof for thm.quantHeuristics
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 1010
% Passive clauses: 14159
% Generated clauses: 185813
% Final active clauses: 190
% Final passive clauses: 1477
% Input formulas: 41
% Initial clauses: 104
% Discarded non-redundant clauses: 3938
% 
% Duplicate literals: 3559
% Trivial inequalities: 547
% Fw subsumption resolutions: 10749
% Fw demodulations: 45207
% Bw demodulations: 6241
% Condensations: 4948
% Global subsumptions: 110
% 
% Simple tautologies: 762
% Equational tautologies: 1210
% Forward subsumptions: 103692
% Fw demodulations to eq. taut.: 6
% Bw demodulations to eq. taut.: 376
% 
% Unit resulting resolution: 30870
% Factoring: 10
% Forward superposition: 43203
% Backward superposition: 31080
% Self superposition: 206
% Equality factoring: 108
% Equality resolution: 53
% 
% Split clauses: 393
% Split components: 795
% SAT solver clauses: 15327
% SAT solver unit clauses: 138
% SAT solver binary clauses: 3590
% 
% Sat splits: 393
% Sat splitting refutations: 270
% 
% Memory used [KB]: 85201
% Time elapsed: 8.510 s
% ------------------------------
% ------------------------------
% Success in time 9.846 s
