Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.arithmetic
% (89382)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 879
% Passive clauses: 8612
% Generated clauses: 25367
% Final active clauses: 844
% Final passive clauses: 7061
% Input formulas: 69
% Initial clauses: 245
% 
% Duplicate literals: 494
% Trivial inequalities: 160
% Fw subsumption resolutions: 2072
% Fw demodulations: 1458
% Bw demodulations: 24
% 
% Simple tautologies: 266
% Equational tautologies: 57
% Forward subsumptions: 11831
% 
% Binary resolution: 9666
% Factoring: 1098
% Forward superposition: 3975
% Backward superposition: 5880
% Self superposition: 40
% Equality factoring: 26
% Equality resolution: 18
% 
% Split clauses: 473
% Split components: 948
% SAT solver clauses: 912
% SAT solver unit clauses: 17
% SAT solver binary clauses: 248
% 
% Sat splits: 473
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 912
% TWLsolver calls for satisfiability: 460
% 
% Memory used [KB]: 13176
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.arithmetic
% (89512)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 494
% Passive clauses: 3762
% Generated clauses: 15666
% Final active clauses: 486
% Final passive clauses: 3259
% Input formulas: 69
% Initial clauses: 245
% 
% Duplicate literals: 205
% Trivial inequalities: 55
% Fw subsumption resolutions: 1146
% Fw demodulations: 905
% Bw demodulations: 17
% Condensations: 398
% Global subsumptions: 22
% 
% Simple tautologies: 271
% Equational tautologies: 17
% Forward subsumptions: 4617
% 
% Binary resolution: 4696
% Factoring: 149
% Forward superposition: 1910
% Backward superposition: 2083
% Self superposition: 35
% Equality factoring: 26
% Equality resolution: 9
% 
% Split clauses: 99
% Split components: 201
% SAT solver clauses: 4544
% SAT solver unit clauses: 238
% SAT solver binary clauses: 90
% SAT solver learnt clauses: 44
% SAT solver learnt literals: 10
% 
% Sat splits: 99
% 
% TWLsolver clauses: 4456
% TWLsolver calls for satisfiability: 1979
% 
% Memory used [KB]: 16375
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.arithmetic
% (89544)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 3245
% Passive clauses: 18878
% Generated clauses: 149262
% Final active clauses: 1342
% Final passive clauses: 6642
% Input formulas: 69
% Initial clauses: 245
% Discarded non-redundant clauses: 1440
% 
% Duplicate literals: 857
% Trivial inequalities: 69
% Fw subsumption resolutions: 1855
% Fw demodulations: 32678
% Bw demodulations: 1089
% Condensations: 2266
% Global subsumptions: 276
% 
% Simple tautologies: 173
% Equational tautologies: 394
% Forward subsumptions: 90928
% Bw demodulations to eq. taut.: 3
% 
% Unit resulting resolution: 55760
% Factoring: 1
% Forward superposition: 29385
% Backward superposition: 23440
% Self superposition: 101
% Equality factoring: 42
% Equality resolution: 4
% 
% Split clauses: 114
% Split components: 234
% SAT solver clauses: 19820
% SAT solver unit clauses: 1736
% SAT solver binary clauses: 8851
% 
% Sat splits: 114
% Sat splitting refutations: 435
% 
% Memory used [KB]: 95947
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.arithmetic
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.arithmetic
% SZS output start Proof for thm.arithmetic
fof(f11221,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f312,f546,f311,f548,f544,f310,f545,f547,f543,f549,f661,f812,f663,f924,f3914,f4075,f4106,f4218,f4359,f4604,f4220,f4716,f5099,f5212,f5475,f5593,f5591,f6033,f6034,f11212])).
fof(f11212,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f11211,f518])).
fof(f518,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.arithmetic.MAX_LE|split|0',aHL_TRUTH)).
fof(f11211,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl3 | $spl5)),
  inference(forward_demodulation,[],[f11210,f5435])).
fof(f5435,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK13),s('type.num.num',sK14))) | $spl5),
  inference(resolution,[],[f5102,f498])).
fof(f498,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f300])).
fof(f300,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X9] : (s(bool,'T_0') = s(bool,X9) <=> p(s(bool,X9)))),
  file('thm.arithmetic.MAX_LE|split|0','thm.bool.EQ_CLAUSES|split|0')).
fof(f5102,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK13),s('type.num.num',sK14)))) | $spl5),
  inference(resolution,[],[f548,f519])).
fof(f519,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f304])).
fof(f304,plain,(
  ! [X0,X1] : ((p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) & (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(nnf_transformation,[],[f144])).
fof(f144,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(flattening,[],[f143])).
fof(f143,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X7,X6] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num',X6)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num',X7)))))),
  file('thm.arithmetic.MAX_LE|split|0','thm.arithmetic.NOT_LESS')).
fof(f11210,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK13),s('type.num.num',sK14)))) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f11209,f539])).
fof(f539,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f444])).
fof(f444,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f266])).
fof(f266,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f101])).
fof(f101,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f60])).
fof(f60,axiom,(
  ! [X15] : ! [X19] : (s(X15,X19) = s(X15,X19) <=> p(s(bool,'T_0')))),
  file('thm.arithmetic.MAX_LE|split|0','thm.bool.REFL_CLAUSE')).
fof(f11209,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK13),s('type.num.num',sK14)))) | ($spl0 | $spl3)),
  inference(forward_demodulation,[],[f11090,f4460])).
fof(f4460,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK12),s('type.num.num',sK14))) | $spl3),
  inference(resolution,[],[f3964,f498])).
fof(f3964,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK12),s('type.num.num',sK14)))) | $spl3),
  inference(resolution,[],[f546,f519])).
fof(f11090,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK12),s('type.num.num',sK14)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK13),s('type.num.num',sK14)))) | $spl0),
  inference(resolution,[],[f5476,f527])).
fof(f527,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.num.num',X2)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X2)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X2))))) )),
  inference(cnf_transformation,[],[f308])).
fof(f308,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.num.num',X2)))) | (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X2)))) & p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X2)))))) & (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X2)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.num.num',X2))))))),
  inference(flattening,[],[f307])).
fof(f307,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.num.num',X2)))) | (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X2)))) & p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X2)))))) & ((~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X2)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X2))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.num.num',X2))))))),
  inference(nnf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.num.num',X2)))) <=> (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X2)))) & p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X2))))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6,X7,X8] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X7),s('type.num.num',X6))),s('type.num.num',X8)))) <=> (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num',X8)))) & p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num',X8))))))),
  file('thm.arithmetic.MAX_LE|split|0','thm.arithmetic.MAX_LT|split|1')).
fof(f5476,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',sK13),s('type.num.num',sK12))),s('type.num.num',sK14)))) | $spl0),
  inference(resolution,[],[f543,f520])).
fof(f520,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f304])).
fof(f6034,plain,(
  $false | ($spl13 | $spl17)),
  inference(subsumption_resolution,[],[f5923,f4217])).
fof(f4217,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK13)))) | $spl13),
  inference(cnf_transformation,[],[f4217_D])).
fof(f4217_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK13)))) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f5923,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK13)))) | $spl17),
  inference(resolution,[],[f4716,f519])).
fof(f6033,plain,(
  $false | ($spl5 | $spl17)),
  inference(subsumption_resolution,[],[f5921,f5102])).
fof(f5921,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK13),s('type.num.num',sK14)))) | $spl17),
  inference(resolution,[],[f4716,f522])).
fof(f522,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f306])).
fof(f306,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X0) = s('type.num.num',X1)) & ((~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) & s('type.num.num',X0) != s('type.num.num',X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(flattening,[],[f305])).
fof(f305,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X0) = s('type.num.num',X1))) & ((~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) & s('type.num.num',X0) != s('type.num.num',X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(nnf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X0) = s('type.num.num',X1)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X7,X6] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',X6)))) <=> (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num',X6)))) | s('type.num.num',X6) = s('type.num.num',X7)))),
  file('thm.arithmetic.MAX_LE|split|0','thm.arithmetic.LESS_OR_EQ')).
fof(f5591,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',sK13),s('type.num.num',sK12)))))) | $spl18),
  inference(cnf_transformation,[],[f5591_D])).
fof(f5591_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',sK13),s('type.num.num',sK12)))))) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f5593,plain,(
  s('type.num.num',sK14) = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl20),
  inference(cnf_transformation,[],[f5593_D])).
fof(f5593_D,plain,(
  s('type.num.num',sK14) = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',sK13),s('type.num.num',sK12))) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f5475,plain,(
  s('type.num.num',sK14) = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',sK13),s('type.num.num',sK12)))))) | $spl0),
  inference(resolution,[],[f543,f521])).
fof(f521,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) = s('type.num.num',X1) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f306])).
fof(f5212,plain,(
  $false | ($spl5 | $spl12)),
  inference(subsumption_resolution,[],[f5100,f4218])).
fof(f5100,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK13)))) | $spl5),
  inference(resolution,[],[f548,f522])).
fof(f5099,plain,(
  $false | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f3790,f4107])).
fof(f4107,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK13),s('type.num.num',sK14)))) | $spl4),
  inference(resolution,[],[f547,f520])).
fof(f3790,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK13),s('type.num.num',sK14)))) | $spl1),
  inference(resolution,[],[f671,f525])).
fof(f525,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X2)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.num.num',X2))))) )),
  inference(cnf_transformation,[],[f308])).
fof(f671,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',sK13),s('type.num.num',sK12))),s('type.num.num',sK14)))) | $spl1),
  inference(resolution,[],[f544,f519])).
fof(f4716,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK14)))) | $spl17),
  inference(cnf_transformation,[],[f4716_D])).
fof(f4716_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK14)))) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f4220,plain,(
  s('type.num.num',sK13) = s('type.num.num',sK14) | $spl14),
  inference(cnf_transformation,[],[f4220_D])).
fof(f4220_D,plain,(
  s('type.num.num',sK13) = s('type.num.num',sK14) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f4604,plain,(
  s('type.num.num',sK13) = s('type.num.num',sK14) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK14)))) | $spl4),
  inference(resolution,[],[f4107,f521])).
fof(f4359,plain,(
  $false | ($spl7 | $spl11)),
  inference(subsumption_resolution,[],[f4358,f512])).
fof(f512,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f140])).
fof(f140,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.arithmetic.MAX_LE|split|0',aHL_FALSITY)).
fof(f4358,plain,(
  p(s(bool,'F_0')) | ($spl7 | $spl11)),
  inference(forward_demodulation,[],[f4248,f1374])).
fof(f1374,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK14))) | $spl11),
  inference(resolution,[],[f924,f494])).
fof(f494,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f298])).
fof(f298,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f121])).
fof(f121,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X9] : (s(bool,'F_0') = s(bool,X9) <=> ~p(s(bool,X9)))),
  file('thm.arithmetic.MAX_LE|split|0','thm.bool.EQ_CLAUSES|split|2')).
fof(f4248,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK14)))) | $spl7),
  inference(resolution,[],[f660,f519])).
fof(f660,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK12)))) | $spl7),
  inference(cnf_transformation,[],[f660_D])).
fof(f660_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK12)))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f4218,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK13)))) | $spl12),
  inference(cnf_transformation,[],[f4218_D])).
fof(f4218_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK13)))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f4106,plain,(
  s('type.num.num',sK13) = s('type.num.num',sK14) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK13)))) | $spl4),
  inference(resolution,[],[f547,f521])).
fof(f4075,plain,(
  $false | ($spl3 | $spl6)),
  inference(subsumption_resolution,[],[f4074,f518])).
fof(f4074,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl6)),
  inference(forward_demodulation,[],[f3962,f1117])).
fof(f1117,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK12))) | $spl6),
  inference(resolution,[],[f661,f498])).
fof(f3962,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK12)))) | $spl3),
  inference(resolution,[],[f546,f522])).
fof(f3914,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f3913,f512])).
fof(f3913,plain,(
  p(s(bool,'F_0')) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f3791,f895])).
fof(f895,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK12),s('type.num.num',sK14))) | $spl2),
  inference(resolution,[],[f550,f494])).
fof(f550,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK12),s('type.num.num',sK14)))) | $spl2),
  inference(resolution,[],[f545,f520])).
fof(f3791,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK12),s('type.num.num',sK14)))) | $spl1),
  inference(resolution,[],[f671,f526])).
fof(f526,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X2)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.num.num',X2))))) )),
  inference(cnf_transformation,[],[f308])).
fof(f924,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK14)))) | $spl11),
  inference(cnf_transformation,[],[f924_D])).
fof(f924_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK14)))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f663,plain,(
  s('type.num.num',sK12) = s('type.num.num',sK14) | $spl8),
  inference(cnf_transformation,[],[f663_D])).
fof(f663_D,plain,(
  s('type.num.num',sK12) = s('type.num.num',sK14) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f812,plain,(
  s('type.num.num',sK12) = s('type.num.num',sK14) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK14)))) | $spl2),
  inference(resolution,[],[f550,f521])).
fof(f661,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK12)))) | $spl6),
  inference(cnf_transformation,[],[f661_D])).
fof(f661_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK12)))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f549,plain,(
  s('type.num.num',sK12) = s('type.num.num',sK14) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK12)))) | $spl2),
  inference(resolution,[],[f545,f521])).
fof(f543,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',sK13),s('type.num.num',sK12)))))) | $spl0),
  inference(cnf_transformation,[],[f543_D])).
fof(f543_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',sK13),s('type.num.num',sK12)))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f547,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK13)))) | $spl4),
  inference(cnf_transformation,[],[f547_D])).
fof(f547_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK13)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f545,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK12)))) | $spl2),
  inference(cnf_transformation,[],[f545_D])).
fof(f545_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK12)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f310,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK12)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK13)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',sK13),s('type.num.num',sK12))))))),
  inference(cnf_transformation,[],[f210])).
fof(f210,plain,(
  (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',sK13),s('type.num.num',sK12)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK13)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK12))))) & (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',sK13),s('type.num.num',sK12)))))) | (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK13)))) & ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK12))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13,sK14])],[f209])).
fof(f209,plain,(
  ? [X0,X1,X2] : ((p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X0)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X0))))) & (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X0)))))) | (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X1)))) & ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X0)))))))),
  inference(flattening,[],[f208])).
fof(f208,plain,(
  ? [X0,X1,X2] : ((p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X0)))))) | (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X0)))))) & (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X0)))))) | (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X1)))) & ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X0)))))))),
  inference(nnf_transformation,[],[f155])).
fof(f155,plain,(
  ? [X0,X1,X2] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X0)))))) <~> (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X0))))))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0,X1,X2] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X0)))))) <=> (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X0))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X8),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X7),s('type.num.num',X6)))))) <=> (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X8),s('type.num.num',X7)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X8),s('type.num.num',X6))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X8),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X7),s('type.num.num',X6)))))) <=> (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X8),s('type.num.num',X7)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X8),s('type.num.num',X6))))))),
  file('thm.arithmetic.MAX_LE|split|0','thm.arithmetic.MAX_LE|split|0')).
fof(f544,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',sK13),s('type.num.num',sK12)))))) | $spl1),
  inference(cnf_transformation,[],[f544_D])).
fof(f544_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',sK13),s('type.num.num',sK12)))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f548,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK13)))) | $spl5),
  inference(cnf_transformation,[],[f548_D])).
fof(f548_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK13)))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f311,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK13)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',sK13),s('type.num.num',sK12))))))),
  inference(cnf_transformation,[],[f210])).
fof(f546,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK12)))) | $spl3),
  inference(cnf_transformation,[],[f546_D])).
fof(f546_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK12)))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f312,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK12)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',sK13),s('type.num.num',sK12))))))),
  inference(cnf_transformation,[],[f210])).
% SZS output end Proof for thm.arithmetic
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 448
% Passive clauses: 3979
% Generated clauses: 11063
% Final active clauses: 321
% Final passive clauses: 1920
% Input formulas: 69
% Initial clauses: 226
% 
% Selected by SInE selection: 62
% SInE iterations: 1
% 
% Duplicate literals: 269
% Trivial inequalities: 81
% Fw subsumption resolutions: 1820
% Fw demodulations: 535
% Bw demodulations: 18
% 
% Simple tautologies: 646
% Equational tautologies: 1
% Forward subsumptions: 3497
% 
% Binary resolution: 3498
% Forward superposition: 1986
% Backward superposition: 2448
% 
% Split clauses: 8
% Split components: 17
% SAT solver clauses: 17
% SAT solver binary clauses: 9
% 
% Sat splits: 8
% Sat splitting refutations: 8
% 
% Memory used [KB]: 3582
% Time elapsed: 0.146 s
% ------------------------------
% ------------------------------
% Success in time 11.888 s
