Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.ring
% (110759)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 108
% Passive clauses: 3329
% Generated clauses: 15537
% Final active clauses: 85
% Final passive clauses: 2782
% Input formulas: 68
% Initial clauses: 77
% 
% Duplicate literals: 647
% Trivial inequalities: 97
% Fw subsumption resolutions: 1953
% Fw demodulations: 564
% Bw demodulations: 1
% 
% Simple tautologies: 226
% Equational tautologies: 77
% Forward subsumptions: 7278
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 274
% Factoring: 3
% Forward superposition: 6125
% Backward superposition: 5229
% Self superposition: 285
% Equality factoring: 187
% Equality resolution: 24
% 
% Split clauses: 875
% Split components: 1872
% SAT solver clauses: 1505
% SAT solver unit clauses: 16
% SAT solver binary clauses: 185
% 
% Sat splits: 875
% Sat splitting refutations: 7
% 
% TWLsolver clauses: 1482
% TWLsolver calls for satisfiability: 155
% 
% Memory used [KB]: 9722
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 144 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.ring
% (110783)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 84
% Passive clauses: 764
% Generated clauses: 10450
% Final active clauses: 55
% Final passive clauses: 366
% Input formulas: 68
% Initial clauses: 77
% 
% Duplicate literals: 713
% Trivial inequalities: 70
% Fw subsumption resolutions: 738
% Fw demodulations: 616
% Bw demodulations: 248
% Condensations: 1820
% 
% Simple tautologies: 121
% Equational tautologies: 54
% Forward subsumptions: 4545
% Bw demodulations to eq. taut.: 83
% 
% Binary resolution: 263
% Factoring: 2
% Forward superposition: 2616
% Backward superposition: 2173
% Self superposition: 183
% Equality factoring: 140
% Equality resolution: 12
% 
% Split clauses: 32
% Split components: 64
% SAT solver clauses: 1010
% SAT solver unit clauses: 130
% SAT solver binary clauses: 33
% 
% Sat splits: 32
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 1010
% TWLsolver calls for satisfiability: 206
% 
% Memory used [KB]: 8571
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.ring
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.ring
% SZS output start Proof for thm.ring
fof(f44501,plain,(
  $false),
  inference(subsumption_resolution,[],[f44500,f213])).
fof(f213,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.ring.ring_fupdfupds_comp|split|5',aHL_TRUTH)).
fof(f44500,plain,(
  ~p(s(bool,'T_0'))),
  inference(forward_demodulation,[],[f44498,f800])).
fof(f800,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X1)))) )),
  inference(unit_resulting_resolution,[],[f791,f293])).
fof(f293,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f215,f216])).
fof(f216,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.ring.ring_fupdfupds_comp|split|5',aHL_BOOL_CASES)).
fof(f215,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.ring.ring_fupdfupds_comp|split|5',aHL_FALSITY)).
fof(f791,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X1))))) )),
  inference(equality_resolution,[],[f268])).
fof(f268,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2))))) )),
  inference(cnf_transformation,[],[f211])).
fof(f211,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2))))))),
  inference(nnf_transformation,[],[f153])).
fof(f153,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f53])).
fof(f53,axiom,(
  ! [X43,X44,X45] : (p(s(bool,happ(s(fun(X43,bool),happ(s(fun(X43,fun(X43,bool)),'=_0'),s(X43,X44))),s(X43,X45)))) <=> s(X43,X44) = s(X43,X45))),
  file('thm.ring.ring_fupdfupds_comp|split|5',arity_0_HL_EQ)).
fof(f44498,plain,(
  ~p(s(bool,happ(s(fun(fun(sK0,'type.ring.ring'(sK1)),bool),happ(s(fun(fun(sK0,'type.ring.ring'(sK1)),fun(fun(sK0,'type.ring.ring'(sK1)),bool)),'=_0'),s(fun(sK0,'type.ring.ring'(sK1)),'const.combin.o_2'(s(fun('type.ring.ring'(sK1),'type.ring.ring'(sK1)),'const.ring.ring_RP_fupd_1'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK4))),s(fun(sK0,'type.ring.ring'(sK1)),'const.combin.o_2'(s(fun('type.ring.ring'(sK1),'type.ring.ring'(sK1)),'const.ring.ring_RP_fupd_1'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK3))),s(fun(sK0,'type.ring.ring'(sK1)),sK2))))))),s(fun(sK0,'type.ring.ring'(sK1)),'const.combin.o_2'(s(fun('type.ring.ring'(sK1),'type.ring.ring'(sK1)),'const.ring.ring_RP_fupd_1'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK4))),s(fun(sK0,'type.ring.ring'(sK1)),'const.combin.o_2'(s(fun('type.ring.ring'(sK1),'type.ring.ring'(sK1)),'const.ring.ring_RP_fupd_1'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK3))),s(fun(sK0,'type.ring.ring'(sK1)),sK2))))))))),
  inference(backward_demodulation,[],[f44420,f12155])).
fof(f12155,plain,(
  ~p(s(bool,happ(s(fun(fun(sK0,'type.ring.ring'(sK1)),bool),happ(s(fun(fun(sK0,'type.ring.ring'(sK1)),fun(fun(sK0,'type.ring.ring'(sK1)),bool)),'=_0'),s(fun(sK0,'type.ring.ring'(sK1)),'const.combin.o_2'(s(fun('type.ring.ring'(sK1),'type.ring.ring'(sK1)),'const.ring.ring_RP_fupd_1'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK4))),s(fun(sK0,'type.ring.ring'(sK1)),'const.combin.o_2'(s(fun('type.ring.ring'(sK1),'type.ring.ring'(sK1)),'const.ring.ring_RP_fupd_1'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK3))),s(fun(sK0,'type.ring.ring'(sK1)),sK2))))))),s(fun(sK0,'type.ring.ring'(sK1)),'const.combin.o_2'(s(fun('type.ring.ring'(sK1),'type.ring.ring'(sK1)),'const.ring.ring_RP_fupd_1'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),'const.combin.o_2'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK4),s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK3))))),s(fun(sK0,'type.ring.ring'(sK1)),sK2))))))),
  inference(unit_resulting_resolution,[],[f212,f267])).
fof(f267,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) )),
  inference(cnf_transformation,[],[f211])).
fof(f212,plain,(
  s(fun(sK0,'type.ring.ring'(sK1)),'const.combin.o_2'(s(fun('type.ring.ring'(sK1),'type.ring.ring'(sK1)),'const.ring.ring_RP_fupd_1'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),'const.combin.o_2'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK4),s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK3))))),s(fun(sK0,'type.ring.ring'(sK1)),sK2))) != s(fun(sK0,'type.ring.ring'(sK1)),'const.combin.o_2'(s(fun('type.ring.ring'(sK1),'type.ring.ring'(sK1)),'const.ring.ring_RP_fupd_1'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK4))),s(fun(sK0,'type.ring.ring'(sK1)),'const.combin.o_2'(s(fun('type.ring.ring'(sK1),'type.ring.ring'(sK1)),'const.ring.ring_RP_fupd_1'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK3))),s(fun(sK0,'type.ring.ring'(sK1)),sK2)))))),
  inference(cnf_transformation,[],[f199])).
fof(f199,plain,(
  s(fun(sK0,'type.ring.ring'(sK1)),'const.combin.o_2'(s(fun('type.ring.ring'(sK1),'type.ring.ring'(sK1)),'const.ring.ring_RP_fupd_1'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),'const.combin.o_2'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK4),s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK3))))),s(fun(sK0,'type.ring.ring'(sK1)),sK2))) != s(fun(sK0,'type.ring.ring'(sK1)),'const.combin.o_2'(s(fun('type.ring.ring'(sK1),'type.ring.ring'(sK1)),'const.ring.ring_RP_fupd_1'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK4))),s(fun(sK0,'type.ring.ring'(sK1)),'const.combin.o_2'(s(fun('type.ring.ring'(sK1),'type.ring.ring'(sK1)),'const.ring.ring_RP_fupd_1'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK3))),s(fun(sK0,'type.ring.ring'(sK1)),sK2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4])],[f194])).
fof(f194,plain,(
  ? [X0,X1,X2,X3,X4] : s(fun(X0,'type.ring.ring'(X1)),'const.combin.o_2'(s(fun('type.ring.ring'(X1),'type.ring.ring'(X1)),'const.ring.ring_RP_fupd_1'(s(fun(fun(X1,fun(X1,X1)),fun(X1,fun(X1,X1))),X4))),s(fun(X0,'type.ring.ring'(X1)),'const.combin.o_2'(s(fun('type.ring.ring'(X1),'type.ring.ring'(X1)),'const.ring.ring_RP_fupd_1'(s(fun(fun(X1,fun(X1,X1)),fun(X1,fun(X1,X1))),X3))),s(fun(X0,'type.ring.ring'(X1)),X2))))) != s(fun(X0,'type.ring.ring'(X1)),'const.combin.o_2'(s(fun('type.ring.ring'(X1),'type.ring.ring'(X1)),'const.ring.ring_RP_fupd_1'(s(fun(fun(X1,fun(X1,X1)),fun(X1,fun(X1,X1))),'const.combin.o_2'(s(fun(fun(X1,fun(X1,X1)),fun(X1,fun(X1,X1))),X4),s(fun(fun(X1,fun(X1,X1)),fun(X1,fun(X1,X1))),X3))))),s(fun(X0,'type.ring.ring'(X1)),X2)))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0,X1,X2,X3,X4] : s(fun(X0,'type.ring.ring'(X1)),'const.combin.o_2'(s(fun('type.ring.ring'(X1),'type.ring.ring'(X1)),'const.ring.ring_RP_fupd_1'(s(fun(fun(X1,fun(X1,X1)),fun(X1,fun(X1,X1))),X4))),s(fun(X0,'type.ring.ring'(X1)),'const.combin.o_2'(s(fun('type.ring.ring'(X1),'type.ring.ring'(X1)),'const.ring.ring_RP_fupd_1'(s(fun(fun(X1,fun(X1,X1)),fun(X1,fun(X1,X1))),X3))),s(fun(X0,'type.ring.ring'(X1)),X2))))) = s(fun(X0,'type.ring.ring'(X1)),'const.combin.o_2'(s(fun('type.ring.ring'(X1),'type.ring.ring'(X1)),'const.ring.ring_RP_fupd_1'(s(fun(fun(X1,fun(X1,X1)),fun(X1,fun(X1,X1))),'const.combin.o_2'(s(fun(fun(X1,fun(X1,X1)),fun(X1,fun(X1,X1))),X4),s(fun(fun(X1,fun(X1,X1)),fun(X1,fun(X1,X1))),X3))))),s(fun(X0,'type.ring.ring'(X1)),X2)))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ~! [X0,X1] : ! [X2,X3,X4] : s(fun(X0,'type.ring.ring'(X1)),'const.combin.o_2'(s(fun('type.ring.ring'(X1),'type.ring.ring'(X1)),'const.ring.ring_RP_fupd_1'(s(fun(fun(X1,fun(X1,X1)),fun(X1,fun(X1,X1))),X4))),s(fun(X0,'type.ring.ring'(X1)),'const.combin.o_2'(s(fun('type.ring.ring'(X1),'type.ring.ring'(X1)),'const.ring.ring_RP_fupd_1'(s(fun(fun(X1,fun(X1,X1)),fun(X1,fun(X1,X1))),X3))),s(fun(X0,'type.ring.ring'(X1)),X2))))) = s(fun(X0,'type.ring.ring'(X1)),'const.combin.o_2'(s(fun('type.ring.ring'(X1),'type.ring.ring'(X1)),'const.ring.ring_RP_fupd_1'(s(fun(fun(X1,fun(X1,X1)),fun(X1,fun(X1,X1))),'const.combin.o_2'(s(fun(fun(X1,fun(X1,X1)),fun(X1,fun(X1,X1))),X4),s(fun(fun(X1,fun(X1,X1)),fun(X1,fun(X1,X1))),X3))))),s(fun(X0,'type.ring.ring'(X1)),X2)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10] : s(fun(X6,'type.ring.ring'(X7)),'const.combin.o_2'(s(fun('type.ring.ring'(X7),'type.ring.ring'(X7)),'const.ring.ring_RP_fupd_1'(s(fun(fun(X7,fun(X7,X7)),fun(X7,fun(X7,X7))),X10))),s(fun(X6,'type.ring.ring'(X7)),'const.combin.o_2'(s(fun('type.ring.ring'(X7),'type.ring.ring'(X7)),'const.ring.ring_RP_fupd_1'(s(fun(fun(X7,fun(X7,X7)),fun(X7,fun(X7,X7))),X9))),s(fun(X6,'type.ring.ring'(X7)),X8))))) = s(fun(X6,'type.ring.ring'(X7)),'const.combin.o_2'(s(fun('type.ring.ring'(X7),'type.ring.ring'(X7)),'const.ring.ring_RP_fupd_1'(s(fun(fun(X7,fun(X7,X7)),fun(X7,fun(X7,X7))),'const.combin.o_2'(s(fun(fun(X7,fun(X7,X7)),fun(X7,fun(X7,X7))),X10),s(fun(fun(X7,fun(X7,X7)),fun(X7,fun(X7,X7))),X9))))),s(fun(X6,'type.ring.ring'(X7)),X8)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10] : s(fun(X6,'type.ring.ring'(X7)),'const.combin.o_2'(s(fun('type.ring.ring'(X7),'type.ring.ring'(X7)),'const.ring.ring_RP_fupd_1'(s(fun(fun(X7,fun(X7,X7)),fun(X7,fun(X7,X7))),X10))),s(fun(X6,'type.ring.ring'(X7)),'const.combin.o_2'(s(fun('type.ring.ring'(X7),'type.ring.ring'(X7)),'const.ring.ring_RP_fupd_1'(s(fun(fun(X7,fun(X7,X7)),fun(X7,fun(X7,X7))),X9))),s(fun(X6,'type.ring.ring'(X7)),X8))))) = s(fun(X6,'type.ring.ring'(X7)),'const.combin.o_2'(s(fun('type.ring.ring'(X7),'type.ring.ring'(X7)),'const.ring.ring_RP_fupd_1'(s(fun(fun(X7,fun(X7,X7)),fun(X7,fun(X7,X7))),'const.combin.o_2'(s(fun(fun(X7,fun(X7,X7)),fun(X7,fun(X7,X7))),X10),s(fun(fun(X7,fun(X7,X7)),fun(X7,fun(X7,X7))),X9))))),s(fun(X6,'type.ring.ring'(X7)),X8)))),
  file('thm.ring.ring_fupdfupds_comp|split|5','thm.ring.ring_fupdfupds_comp|split|5')).
fof(f44420,plain,(
  ( ! [X171,X169,X172,X170,X168] : (s(fun(X171,'type.ring.ring'(X168)),'const.combin.o_2'(s(fun('type.ring.ring'(X168),'type.ring.ring'(X168)),'const.ring.ring_RP_fupd_1'(s(fun(fun(X168,fun(X168,X168)),fun(X168,fun(X168,X168))),X169))),s(fun(X171,'type.ring.ring'(X168)),'const.combin.o_2'(s(fun('type.ring.ring'(X168),'type.ring.ring'(X168)),'const.ring.ring_RP_fupd_1'(s(fun(fun(X168,fun(X168,X168)),fun(X168,fun(X168,X168))),X170))),s(fun(X171,'type.ring.ring'(X168)),X172))))) = s(fun(X171,'type.ring.ring'(X168)),'const.combin.o_2'(s(fun('type.ring.ring'(X168),'type.ring.ring'(X168)),'const.ring.ring_RP_fupd_1'(s(fun(fun(X168,fun(X168,X168)),fun(X168,fun(X168,X168))),'const.combin.o_2'(s(fun(fun(X168,fun(X168,X168)),fun(X168,fun(X168,X168))),X169),s(fun(fun(X168,fun(X168,X168)),fun(X168,fun(X168,X168))),X170))))),s(fun(X171,'type.ring.ring'(X168)),X172)))) )),
  inference(superposition,[],[f272,f227])).
fof(f227,plain,(
  ( ! [X2,X0,X1] : (s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_RP_fupd_1'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.o_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X1))))) = s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.combin.o_2'(s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_RP_fupd_1'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2))),s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_RP_fupd_1'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X1)))))) )),
  inference(cnf_transformation,[],[f92])).
fof(f92,plain,(
  ! [X0,X1,X2] : s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_RP_fupd_1'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.o_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X1))))) = s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.combin.o_2'(s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_RP_fupd_1'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2))),s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_RP_fupd_1'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X1)))))),
  inference(flattening,[],[f91])).
fof(f91,plain,(
  ! [X0] : ! [X1,X2] : s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_RP_fupd_1'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.o_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X1))))) = s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.combin.o_2'(s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_RP_fupd_1'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2))),s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_RP_fupd_1'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X1)))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X7] : ! [X9,X10] : s(fun('type.ring.ring'(X7),'type.ring.ring'(X7)),'const.ring.ring_RP_fupd_1'(s(fun(fun(X7,fun(X7,X7)),fun(X7,fun(X7,X7))),'const.combin.o_2'(s(fun(fun(X7,fun(X7,X7)),fun(X7,fun(X7,X7))),X10),s(fun(fun(X7,fun(X7,X7)),fun(X7,fun(X7,X7))),X9))))) = s(fun('type.ring.ring'(X7),'type.ring.ring'(X7)),'const.combin.o_2'(s(fun('type.ring.ring'(X7),'type.ring.ring'(X7)),'const.ring.ring_RP_fupd_1'(s(fun(fun(X7,fun(X7,X7)),fun(X7,fun(X7,X7))),X10))),s(fun('type.ring.ring'(X7),'type.ring.ring'(X7)),'const.ring.ring_RP_fupd_1'(s(fun(fun(X7,fun(X7,X7)),fun(X7,fun(X7,X7))),X9)))))),
  file('thm.ring.ring_fupdfupds_comp|split|5','thm.ring.ring_fupdfupds_comp|split|4')).
fof(f272,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))) )),
  inference(cnf_transformation,[],[f161])).
fof(f161,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))),
  inference(flattening,[],[f160])).
fof(f160,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6] : s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X6,X7,X16,X13] : ! [X10,X9,X8] : s(fun(X16,X6),'const.combin.o_2'(s(fun(X7,X6),X10),s(fun(X16,X7),'const.combin.o_2'(s(fun(X13,X7),X9),s(fun(X16,X13),X8))))) = s(fun(X16,X6),'const.combin.o_2'(s(fun(X13,X6),'const.combin.o_2'(s(fun(X7,X6),X10),s(fun(X13,X7),X9))),s(fun(X16,X13),X8)))),
  file('thm.ring.ring_fupdfupds_comp|split|5','thm.combin.o_ASSOC')).
% SZS output end Proof for thm.ring
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 242
% Passive clauses: 3493
% Generated clauses: 45605
% Final active clauses: 138
% Final passive clauses: 1720
% Input formulas: 68
% Initial clauses: 77
% 
% Duplicate literals: 1996
% Trivial inequalities: 252
% Fw subsumption resolutions: 2369
% Fw demodulations: 3673
% Bw demodulations: 1361
% Condensations: 4491
% Global subsumptions: 54
% 
% Simple tautologies: 415
% Equational tautologies: 222
% Forward subsumptions: 28443
% Fw demodulations to eq. taut.: 9
% Bw demodulations to eq. taut.: 122
% 
% Unit resulting resolution: 952
% Factoring: 3
% Forward superposition: 14538
% Backward superposition: 13513
% Self superposition: 420
% Equality factoring: 197
% Equality resolution: 25
% 
% Split clauses: 94
% Split components: 188
% SAT solver clauses: 3799
% SAT solver unit clauses: 606
% SAT solver binary clauses: 762
% 
% Sat splits: 94
% Sat splitting refutations: 62
% 
% Memory used [KB]: 51939
% Time elapsed: 3.051 s
% ------------------------------
% ------------------------------
% Success in time 4.38 s
