Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.rat
% (91978)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 817
% Passive clauses: 8746
% Generated clauses: 21576
% Final active clauses: 760
% Final passive clauses: 7335
% Input formulas: 54
% Initial clauses: 210
% 
% Duplicate literals: 643
% Trivial inequalities: 133
% Fw subsumption resolutions: 2224
% Fw demodulations: 405
% Bw demodulations: 55
% 
% Simple tautologies: 270
% Equational tautologies: 65
% Forward subsumptions: 8424
% 
% Binary resolution: 10601
% Factoring: 1492
% Forward superposition: 3459
% Backward superposition: 1952
% Self superposition: 32
% Equality factoring: 28
% Equality resolution: 50
% 
% Split clauses: 656
% Split components: 1319
% SAT solver clauses: 1186
% SAT solver unit clauses: 28
% SAT solver binary clauses: 189
% 
% Sat splits: 656
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 1184
% TWLsolver calls for satisfiability: 490
% 
% Memory used [KB]: 8443
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.rat
% (91992)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 386
% Passive clauses: 2273
% Generated clauses: 10098
% Final active clauses: 378
% Final passive clauses: 1867
% Input formulas: 54
% Initial clauses: 210
% 
% Duplicate literals: 239
% Trivial inequalities: 20
% Fw subsumption resolutions: 811
% Fw demodulations: 201
% Bw demodulations: 23
% Condensations: 260
% Global subsumptions: 22
% 
% Simple tautologies: 255
% Equational tautologies: 20
% Forward subsumptions: 3250
% 
% Binary resolution: 3442
% Factoring: 339
% Forward superposition: 1119
% Backward superposition: 1063
% Self superposition: 32
% Equality factoring: 26
% Equality resolution: 6
% 
% Split clauses: 23
% Split components: 51
% SAT solver clauses: 2920
% SAT solver unit clauses: 215
% SAT solver binary clauses: 91
% SAT solver learnt clauses: 45
% SAT solver learnt literals: 11
% 
% Sat splits: 23
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 2831
% TWLsolver calls for satisfiability: 1749
% 
% Memory used [KB]: 12409
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.rat
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.rat
% SZS output start Proof for thm.rat
fof(f8950,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f276,f500,f498,f275,f499,f497,f506,f1779,f1783,f1787,f1790,f1793,f1796,f1799,f1802,f1805,f1808,f1811,f1814,f1817,f1820,f1823,f1826,f1829,f1832,f1835,f1838,f1873,f1780,f1875,f1784,f1878,f1788,f1880,f1791,f1882,f1794,f1884,f1797,f1886,f1800,f1888,f1803,f1890,f1806,f1892,f1809,f1894,f1812,f1896,f1815,f1898,f1818,f1900,f1821,f1902,f1824,f1904,f1827,f1906,f1830,f1908,f1833,f1910,f1836,f1912,f1781,f1839,f1943,f1944,f1948,f2417,f2418,f2425,f2422,f2419,f2420,f2428,f2421,f2430,f2423,f2427,f2447,f2448,f2460,f2458,f3055,f3058,f3056,f3060,f3061,f3062,f3064,f3069,f3071,f3075,f3077,f3079,f3081,f3089,f3091,f3099,f3101,f3108,f3115,f3127,f3144,f3172,f3174,f3226,f3230,f3250,f3257,f3265,f3272,f3274,f3281,f3287,f3294,f3300,f3302,f3304,f3306,f3308,f3879,f3880,f3321,f3323,f3325,f3327,f3344,f3346,f3352,f3358,f3367,f3374,f3376,f3387,f3395,f3397,f3406,f3412,f3419,f3425,f3427,f3437,f3439,f3445,f3447,f3458,f3460,f3466,f3468,f3479,f3481,f3487,f3489,f3496,f3502,f3509,f3515,f3524,f3533,f3543,f3555,f3558,f3566,f3568,f3570,f3575,f3577,f3582,f3584,f3589,f3591,f3596,f3598,f3608,f3616,f3625,f3632,f3634,f3650,f3652,f3654,f3656,f3666,f3668,f3678,f3680,f3691,f3693,f3704,f3706,f3708,f3714,f3719,f3724,f3726,f3731,f3733,f3741,f3749,f3758,f3766,f3771,f3773,f3775,f3777,f3782,f3784,f3786,f3788,f3790,f3798,f3803,f3808,f3811,f3816,f3819,f3826,f3831,f3836,f3841,f3843,f3845,f3850,f3852,f3857,f3881,f3882,f3883,f3884,f3885,f3886,f3887,f3888,f3889,f3890,f3891,f3892,f3893,f3894,f3895,f3896,f3897,f3898,f3899,f3900,f3901,f3902,f3903,f3904,f3905,f3906,f3907,f3908,f3909,f3910,f3911,f3912,f3913,f3914,f3915,f3916,f3917,f3918,f3919,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3927,f3928,f3929,f3930,f3931,f3932,f3933,f3934,f3935,f3936,f3937,f3938,f3939,f3940,f3941,f3942,f3943,f3944,f3945,f3946,f3947,f3948,f3949,f3950,f3951,f3952,f3953,f3954,f3955,f3956,f3957,f3958,f3960,f3962,f3964,f3966,f3967,f3968,f3969,f3970,f3971,f3972,f3976,f3977,f3978,f3979,f3980,f3981,f3982,f3983,f3984,f3985,f3986,f3987,f3988,f3989,f3990,f3991,f3992,f3993,f3994,f3995,f3996,f3997,f3998,f3999,f4000,f4001,f4002,f4003,f4004,f4005,f4006,f4007,f4008,f4009,f4038,f4040,f4039,f4042,f4043,f4320,f4323,f4321,f5074,f5076,f5075,f8475,f8477,f8479,f8481,f8500,f8502,f8506,f8507,f8511,f8512,f8590,f8476,f8591,f8478,f8592,f8480,f8593,f8482,f8611,f8501,f8612,f8503,f8912,f8917,f3057,f8913,f8945,f8934])).
fof(f8934,plain,(
  $false | ($spl0 | $spl3)),
  inference(unit_resulting_resolution,[],[f508,f8923])).
fof(f8923,plain,(
  ( ! [X0] : (s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',sK12),s('type.rat.rat',X0))) != s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',sK13),s('type.rat.rat',X0)))) ) | $spl3),
  inference(superposition,[],[f8889,f358])).
fof(f358,plain,(
  ( ! [X0,X1] : (s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))) = s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',X1),s('type.rat.rat',X0)))) )),
  inference(cnf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0,X1] : s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))) = s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',X1),s('type.rat.rat',X0)))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X20,X21] : s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',X20),s('type.rat.rat',X21))) = s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',X21),s('type.rat.rat',X20)))),
  file('thm.rat.RAT_EQ_LADD','thm.rat.RAT_ADD_COMM')).
fof(f8889,plain,(
  ( ! [X0] : (s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',X0),s('type.rat.rat',sK12))) != s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',sK13),s('type.rat.rat',X0)))) ) | $spl3),
  inference(superposition,[],[f8842,f358])).
fof(f8842,plain,(
  ( ! [X0] : (s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',X0),s('type.rat.rat',sK12))) != s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',X0),s('type.rat.rat',sK13)))) ) | $spl3),
  inference(forward_demodulation,[],[f8841,f359])).
fof(f359,plain,(
  ( ! [X0,X1] : (s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))) = s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0))),s('type.rat.rat',X1)))) )),
  inference(cnf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0,X1] : s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))) = s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0))),s('type.rat.rat',X1)))),
  inference(rectify,[],[f55])).
fof(f55,axiom,(
  ! [X25,X26] : s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',X25),s('type.rat.rat',X26))) = s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X25))),s('type.rat.rat',X26)))),
  file('thm.rat.RAT_EQ_LADD','ARITY_1_2_const.rat.rat_add')).
fof(f8841,plain,(
  ( ! [X0] : (s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',X0),s('type.rat.rat',sK13))) != s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0))),s('type.rat.rat',sK12)))) ) | $spl3),
  inference(forward_demodulation,[],[f8820,f359])).
fof(f8820,plain,(
  ( ! [X0] : (s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0))),s('type.rat.rat',sK12))) != s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0))),s('type.rat.rat',sK13)))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f500,f1336])).
fof(f1336,plain,(
  ( ! [X2,X0,X1] : (s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0))),s('type.rat.rat',X1))) != s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0))),s('type.rat.rat',X2))) | s('type.rat.rat',X1) = s('type.rat.rat',X2)) )),
  inference(subsumption_resolution,[],[f1334,f277])).
fof(f277,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.rat.RAT_EQ_LADD',aHL_TRUTH)).
fof(f1334,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0))),s('type.rat.rat',X1))) != s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0))),s('type.rat.rat',X2))) | s('type.rat.rat',X1) = s('type.rat.rat',X2)) )),
  inference(superposition,[],[f401,f551])).
fof(f551,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))))) )),
  inference(unit_resulting_resolution,[],[f284,f298])).
fof(f298,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f178])).
fof(f178,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X9] : (s(bool,'T_0') = s(bool,X9) <=> p(s(bool,X9)))),
  file('thm.rat.RAT_EQ_LADD','thm.bool.EQ_CLAUSES|split|0')).
fof(f284,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0))))))) )),
  inference(cnf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0] : p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0))))))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X6] : p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6))))))),
  file('thm.rat.RAT_EQ_LADD','thm.rat.RAT_ADD_ONE_ONE')).
fof(f401,plain,(
  ( ! [X6,X2,X0,X5,X1] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X0),X2)))) | s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) | s(X1,X5) = s(X1,X6)) )),
  inference(cnf_transformation,[],[f236])).
fof(f236,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X0),X2)))) | ! [X5,X6] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) | s(X1,X5) = s(X1,X6))) & ((s(X0,happ(s(fun(X1,X0),X2),s(X1,sK17(X2,X1,X0)))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,sK18(X2,X1,X0)))) & s(X1,sK17(X2,X1,X0)) != s(X1,sK18(X2,X1,X0))) | p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X0),X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK17,sK18])],[f235])).
fof(f235,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X0),X2)))) | ! [X5,X6] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) | s(X1,X5) = s(X1,X6))) & (? [X3,X4] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) & s(X1,X3) != s(X1,X4)) | p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X0),X2))))))),
  inference(rectify,[],[f234])).
fof(f234,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X0),X2)))) | ! [X3,X4] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,X3))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) | s(X1,X3) = s(X1,X4))) & (? [X3,X4] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) & s(X1,X3) != s(X1,X4)) | p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X0),X2))))))),
  inference(nnf_transformation,[],[f153])).
fof(f153,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X0),X2)))) <=> ! [X3,X4] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,X3))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) | s(X1,X3) = s(X1,X4)))),
  inference(ennf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X0),X2)))) <=> ! [X3,X4] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) => s(X1,X3) = s(X1,X4)))),
  inference(flattening,[],[f117])).
fof(f117,plain,(
  ! [X0,X1] : ! [X2] : (p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X0),X2)))) <=> ! [X3,X4] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) => s(X1,X3) = s(X1,X4)))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X15,X12] : ! [X17] : (p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X12,X15),X17)))) <=> ! [X18,X19] : (s(X15,happ(s(fun(X12,X15),X17),s(X12,X18))) = s(X15,happ(s(fun(X12,X15),X17),s(X12,X19))) => s(X12,X18) = s(X12,X19)))),
  file('thm.rat.RAT_EQ_LADD','thm.bool.ONE_ONE_THM')).
fof(f508,plain,(
  s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',sK12),s('type.rat.rat',sK14))) = s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',sK13),s('type.rat.rat',sK14))) | $spl0),
  inference(forward_demodulation,[],[f507,f358])).
fof(f507,plain,(
  s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',sK13),s('type.rat.rat',sK14))) = s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',sK14),s('type.rat.rat',sK12))) | $spl0),
  inference(forward_demodulation,[],[f497,f358])).
fof(f8945,plain,(
  $false | ($spl0 | $spl3)),
  inference(trivial_inequality_removal,[],[f8942])).
fof(f8942,plain,(
  s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',sK12),s('type.rat.rat',sK14))) != s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',sK12),s('type.rat.rat',sK14))) | ($spl0 | $spl3)),
  inference(superposition,[],[f8923,f508])).
fof(f8913,plain,(
  ( ! [X35,X31,X36,X34,X32] : (s(bool,X34) != s(bool,sK17(X32,bool,X31)) | s(bool,'T_0') = s(bool,X36) | s(bool,X35) = s(bool,X36) | s(bool,X35) = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X31),X32))) | s(bool,'T_0') = s(bool,sK18(X32,bool,X31)) | s(bool,'T_0') = s(bool,X34)) ) | $spl94),
  inference(cnf_transformation,[],[f8913_D])).
fof(f8913_D,plain,(
  ( ! [X35,X31,X36,X34,X32] : (s(bool,X34) != s(bool,sK17(X32,bool,X31)) | s(bool,'T_0') = s(bool,X36) | s(bool,X35) = s(bool,X36) | s(bool,X35) = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X31),X32))) | s(bool,'T_0') = s(bool,sK18(X32,bool,X31)) | s(bool,'T_0') = s(bool,X34)) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f3057,plain,(
  ( ! [X9] : (p(s(bool,X9))) ) | $spl64),
  inference(cnf_transformation,[],[f3057_D])).
fof(f3057_D,plain,(
  ( ! [X9] : (p(s(bool,X9))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f8917,plain,(
  ( ! [X59,X57,X60,X58,X56,X55] : (p(s(bool,X57)) | s(bool,X58) != s(bool,sK17(X56,bool,X55)) | s(bool,'T_0') = s(bool,X58) | s(bool,'T_0') = s(bool,sK18(X56,bool,X55)) | s(bool,X59) = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X55),X56))) | s(bool,X59) = s(bool,X60) | s(bool,'T_0') = s(bool,X60)) )),
  inference(subsumption_resolution,[],[f8907,f297])).
fof(f297,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f178])).
fof(f8907,plain,(
  ( ! [X59,X57,X60,X58,X56,X55] : (p(s(bool,X57)) | s(bool,X58) != s(bool,sK17(X56,bool,X55)) | s(bool,'T_0') = s(bool,X58) | s(bool,'T_0') = s(bool,sK18(X56,bool,X55)) | s(bool,X59) = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X55),X56))) | s(bool,X59) = s(bool,X60) | s(bool,'T_0') = s(bool,X60) | s(bool,'T_0') = s(bool,X57)) )),
  inference(superposition,[],[f1441,f1376])).
fof(f1376,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f657,f657])).
fof(f657,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f286,f286])).
fof(f286,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.rat.RAT_EQ_LADD',aHL_BOOL_CASES)).
fof(f1441,plain,(
  ( ! [X156,X157,X155] : (p(s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X156),X155)))) | s(bool,X157) != s(bool,sK17(X155,bool,X156)) | s(bool,'T_0') = s(bool,X157) | s(bool,'T_0') = s(bool,sK18(X155,bool,X156))) )),
  inference(superposition,[],[f403,f657])).
fof(f403,plain,(
  ( ! [X2,X0,X1] : (s(X1,sK17(X2,X1,X0)) != s(X1,sK18(X2,X1,X0)) | p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X0),X2))))) )),
  inference(cnf_transformation,[],[f236])).
fof(f8912,plain,(
  ( ! [X35,X33,X31,X36,X34,X32] : (p(s(bool,X33)) | s(bool,X34) != s(bool,sK17(X32,bool,X31)) | s(bool,'T_0') = s(bool,X34) | s(bool,'T_0') = s(bool,sK18(X32,bool,X31)) | s(bool,X35) = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X31),X32))) | s(bool,X35) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(subsumption_resolution,[],[f8903,f297])).
fof(f8903,plain,(
  ( ! [X35,X33,X31,X36,X34,X32] : (p(s(bool,X33)) | s(bool,X34) != s(bool,sK17(X32,bool,X31)) | s(bool,'T_0') = s(bool,X34) | s(bool,'T_0') = s(bool,sK18(X32,bool,X31)) | s(bool,X35) = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X31),X32))) | s(bool,'T_0') = s(bool,X33) | s(bool,X35) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f1441,f1376])).
fof(f8503,plain,(
  ( ! [X565,X563,X566,X562] : (s(bool,'T_0') = s(bool,sK18(X563,bool,X562)) | s(bool,'T_0') = s(bool,X566) | s(bool,X565) = s(bool,X566) | s(bool,X565) = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X562),X563))) | s(bool,'T_0') = s(bool,sK17(X563,bool,X562))) ) | $spl88),
  inference(cnf_transformation,[],[f8503_D])).
fof(f8503_D,plain,(
  ( ! [X565,X563,X566,X562] : (s(bool,'T_0') = s(bool,sK18(X563,bool,X562)) | s(bool,'T_0') = s(bool,X566) | s(bool,X565) = s(bool,X566) | s(bool,X565) = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X562),X563))) | s(bool,'T_0') = s(bool,sK17(X563,bool,X562))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f8612,plain,(
  ( ! [X565,X563,X564,X566,X562] : (p(s(bool,X564)) | s(bool,'T_0') = s(bool,sK18(X563,bool,X562)) | s(bool,'T_0') = s(bool,sK17(X563,bool,X562)) | s(bool,X565) = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X562),X563))) | s(bool,X565) = s(bool,X566) | s(bool,'T_0') = s(bool,X566)) )),
  inference(subsumption_resolution,[],[f6348,f297])).
fof(f6348,plain,(
  ( ! [X565,X563,X564,X566,X562] : (p(s(bool,X564)) | s(bool,'T_0') = s(bool,sK18(X563,bool,X562)) | s(bool,'T_0') = s(bool,sK17(X563,bool,X562)) | s(bool,X565) = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X562),X563))) | s(bool,X565) = s(bool,X566) | s(bool,'T_0') = s(bool,X566) | s(bool,'T_0') = s(bool,X564)) )),
  inference(superposition,[],[f2344,f1376])).
fof(f2344,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X1),X0)))) | s(bool,'T_0') = s(bool,sK18(X0,bool,X1)) | s(bool,'T_0') = s(bool,sK17(X0,bool,X1))) )),
  inference(condensation,[],[f2343])).
fof(f2343,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X7),X6)))) | s(bool,'T_0') = s(bool,sK18(X6,bool,X7)) | s(bool,'T_0') = s(bool,sK17(X6,bool,X7)) | s(bool,'T_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f2339,f286])).
fof(f2339,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') != s(bool,X8) | p(s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X7),X6)))) | s(bool,'T_0') = s(bool,sK18(X6,bool,X7)) | s(bool,'T_0') = s(bool,sK17(X6,bool,X7)) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f1178,f657])).
fof(f1178,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,sK17(X0,bool,X1)) | p(s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X1),X0)))) | s(bool,'T_0') = s(bool,sK18(X0,bool,X1))) )),
  inference(superposition,[],[f403,f286])).
fof(f8501,plain,(
  ( ! [X561,X557,X560,X556,X558] : (s(X557,happ(s(fun(X556,X557),X558),s(X556,sK17(X558,X556,X557)))) = s(X557,happ(s(fun(X556,X557),X558),s(X556,sK18(X558,X556,X557)))) | s(bool,'T_0') = s(bool,X561) | s(bool,X560) = s(bool,X561) | s(bool,X560) = s(bool,'const.bool.ONE_ONE_1'(s(fun(X556,X557),X558)))) ) | $spl86),
  inference(cnf_transformation,[],[f8501_D])).
fof(f8501_D,plain,(
  ( ! [X561,X557,X560,X556,X558] : (s(X557,happ(s(fun(X556,X557),X558),s(X556,sK17(X558,X556,X557)))) = s(X557,happ(s(fun(X556,X557),X558),s(X556,sK18(X558,X556,X557)))) | s(bool,'T_0') = s(bool,X561) | s(bool,X560) = s(bool,X561) | s(bool,X560) = s(bool,'const.bool.ONE_ONE_1'(s(fun(X556,X557),X558)))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f8611,plain,(
  ( ! [X561,X557,X559,X560,X556,X558] : (p(s(bool,X559)) | s(X557,happ(s(fun(X556,X557),X558),s(X556,sK17(X558,X556,X557)))) = s(X557,happ(s(fun(X556,X557),X558),s(X556,sK18(X558,X556,X557)))) | s(bool,X560) = s(bool,'const.bool.ONE_ONE_1'(s(fun(X556,X557),X558))) | s(bool,X560) = s(bool,X561) | s(bool,'T_0') = s(bool,X561)) )),
  inference(subsumption_resolution,[],[f6347,f297])).
fof(f6347,plain,(
  ( ! [X561,X557,X559,X560,X556,X558] : (p(s(bool,X559)) | s(X557,happ(s(fun(X556,X557),X558),s(X556,sK17(X558,X556,X557)))) = s(X557,happ(s(fun(X556,X557),X558),s(X556,sK18(X558,X556,X557)))) | s(bool,X560) = s(bool,'const.bool.ONE_ONE_1'(s(fun(X556,X557),X558))) | s(bool,X560) = s(bool,X561) | s(bool,'T_0') = s(bool,X561) | s(bool,'T_0') = s(bool,X559)) )),
  inference(superposition,[],[f402,f1376])).
fof(f402,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X0),X2)))) | s(X0,happ(s(fun(X1,X0),X2),s(X1,sK17(X2,X1,X0)))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,sK18(X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f236])).
fof(f8482,plain,(
  ( ! [X231,X235,X233,X236,X234] : (sP1(X231,X233) | s(bool,'T_0') = s(bool,X236) | s(bool,X235) = s(bool,X236) | s(bool,X231) = s(bool,X235) | s(bool,'T_0') = s(bool,X234) | s(bool,X233) = s(bool,X234)) ) | $spl84),
  inference(cnf_transformation,[],[f8482_D])).
fof(f8482_D,plain,(
  ( ! [X231,X235,X233,X236,X234] : (sP1(X231,X233) | s(bool,'T_0') = s(bool,X236) | s(bool,X235) = s(bool,X236) | s(bool,X231) = s(bool,X235) | s(bool,'T_0') = s(bool,X234) | s(bool,X233) = s(bool,X234)) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f8593,plain,(
  ( ! [X231,X235,X233,X236,X234,X232] : (p(s(bool,X232)) | sP1(X231,X233) | s(bool,X233) = s(bool,X234) | s(bool,'T_0') = s(bool,X234) | s(bool,X231) = s(bool,X235) | s(bool,X235) = s(bool,X236) | s(bool,'T_0') = s(bool,X236)) )),
  inference(subsumption_resolution,[],[f6265,f297])).
fof(f6265,plain,(
  ( ! [X231,X235,X233,X236,X234,X232] : (p(s(bool,X232)) | sP1(X231,X233) | s(bool,X233) = s(bool,X234) | s(bool,'T_0') = s(bool,X234) | s(bool,X231) = s(bool,X235) | s(bool,X235) = s(bool,X236) | s(bool,'T_0') = s(bool,X236) | s(bool,'T_0') = s(bool,X232)) )),
  inference(superposition,[],[f1845,f1376])).
fof(f1845,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X29)) | sP1(X29,X28) | s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(subsumption_resolution,[],[f1534,f277])).
fof(f1534,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'T_0')) | p(s(bool,X29)) | sP1(X29,X28) | s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f388,f657])).
fof(f388,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f228])).
fof(f228,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1)))),
  inference(rectify,[],[f227])).
fof(f227,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0)))),
  inference(flattening,[],[f226])).
fof(f226,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP1(X1,X0)))),
  inference(nnf_transformation,[],[f158])).
fof(f158,plain,(
  ! [X1,X0] : (sP1(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f8480,plain,(
  ( ! [X229,X227,X225,X230,X228] : (sP1(X227,X225) | s(bool,'T_0') = s(bool,X230) | s(bool,X229) = s(bool,X230) | s(bool,X225) = s(bool,X229) | s(bool,'T_0') = s(bool,X228) | s(bool,X227) = s(bool,X228)) ) | $spl82),
  inference(cnf_transformation,[],[f8480_D])).
fof(f8480_D,plain,(
  ( ! [X229,X227,X225,X230,X228] : (sP1(X227,X225) | s(bool,'T_0') = s(bool,X230) | s(bool,X229) = s(bool,X230) | s(bool,X225) = s(bool,X229) | s(bool,'T_0') = s(bool,X228) | s(bool,X227) = s(bool,X228)) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f8592,plain,(
  ( ! [X229,X227,X225,X230,X228,X226] : (p(s(bool,X226)) | sP1(X227,X225) | s(bool,X227) = s(bool,X228) | s(bool,'T_0') = s(bool,X228) | s(bool,X225) = s(bool,X229) | s(bool,X229) = s(bool,X230) | s(bool,'T_0') = s(bool,X230)) )),
  inference(subsumption_resolution,[],[f6264,f297])).
fof(f6264,plain,(
  ( ! [X229,X227,X225,X230,X228,X226] : (p(s(bool,X226)) | sP1(X227,X225) | s(bool,X227) = s(bool,X228) | s(bool,'T_0') = s(bool,X228) | s(bool,X225) = s(bool,X229) | s(bool,X229) = s(bool,X230) | s(bool,'T_0') = s(bool,X230) | s(bool,'T_0') = s(bool,X226)) )),
  inference(superposition,[],[f1844,f1376])).
fof(f1844,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X26)) | sP1(X25,X26) | s(bool,X25) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(subsumption_resolution,[],[f1533,f277])).
fof(f1533,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'T_0')) | p(s(bool,X26)) | sP1(X25,X26) | s(bool,X25) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f385,f657])).
fof(f385,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f228])).
fof(f8478,plain,(
  ( ! [X212,X210,X213] : (sP1(X210,'T_0') | s(bool,'T_0') = s(bool,X213) | s(bool,X212) = s(bool,X213) | s(bool,X210) = s(bool,X212)) ) | $spl80),
  inference(cnf_transformation,[],[f8478_D])).
fof(f8478_D,plain,(
  ( ! [X212,X210,X213] : (sP1(X210,'T_0') | s(bool,'T_0') = s(bool,X213) | s(bool,X212) = s(bool,X213) | s(bool,X210) = s(bool,X212)) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f8591,plain,(
  ( ! [X212,X210,X213,X211] : (p(s(bool,X211)) | sP1(X210,'T_0') | s(bool,X210) = s(bool,X212) | s(bool,X212) = s(bool,X213) | s(bool,'T_0') = s(bool,X213)) ) | $spl8),
  inference(subsumption_resolution,[],[f6261,f297])).
fof(f6261,plain,(
  ( ! [X212,X210,X213,X211] : (p(s(bool,X211)) | sP1(X210,'T_0') | s(bool,X210) = s(bool,X212) | s(bool,X212) = s(bool,X213) | s(bool,'T_0') = s(bool,X213) | s(bool,'T_0') = s(bool,X211)) ) | $spl8),
  inference(superposition,[],[f1784,f1376])).
fof(f8476,plain,(
  ( ! [X208,X206,X209] : (sP1('T_0',X206) | s(bool,'T_0') = s(bool,X209) | s(bool,X208) = s(bool,X209) | s(bool,X206) = s(bool,X208)) ) | $spl78),
  inference(cnf_transformation,[],[f8476_D])).
fof(f8476_D,plain,(
  ( ! [X208,X206,X209] : (sP1('T_0',X206) | s(bool,'T_0') = s(bool,X209) | s(bool,X208) = s(bool,X209) | s(bool,X206) = s(bool,X208)) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f8590,plain,(
  ( ! [X208,X206,X209,X207] : (p(s(bool,X207)) | sP1('T_0',X206) | s(bool,X206) = s(bool,X208) | s(bool,X208) = s(bool,X209) | s(bool,'T_0') = s(bool,X209)) ) | $spl4),
  inference(subsumption_resolution,[],[f6260,f297])).
fof(f6260,plain,(
  ( ! [X208,X206,X209,X207] : (p(s(bool,X207)) | sP1('T_0',X206) | s(bool,X206) = s(bool,X208) | s(bool,X208) = s(bool,X209) | s(bool,'T_0') = s(bool,X209) | s(bool,'T_0') = s(bool,X207)) ) | $spl4),
  inference(superposition,[],[f1780,f1376])).
fof(f8512,plain,(
  ( ! [X204,X202,X205,X203,X201] : (s(bool,X201) != s(bool,X202) | s(bool,'T_0') = s(bool,X205) | s(bool,X204) = s(bool,X205) | s(bool,X201) = s(bool,X204) | s(bool,'T_0') = s(bool,X203) | s(bool,X202) = s(bool,X203)) ) | $spl92),
  inference(cnf_transformation,[],[f8512_D])).
fof(f8512_D,plain,(
  ( ! [X204,X202,X205,X203,X201] : (s(bool,X201) != s(bool,X202) | s(bool,'T_0') = s(bool,X205) | s(bool,X204) = s(bool,X205) | s(bool,X201) = s(bool,X204) | s(bool,'T_0') = s(bool,X203) | s(bool,X202) = s(bool,X203)) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f8511,plain,(
  ( ! [X204,X202,X200,X205,X203,X201] : (s(bool,X201) != s(bool,X202) | p(s(bool,X200)) | s(bool,X202) = s(bool,X203) | s(bool,'T_0') = s(bool,X203) | s(bool,X201) = s(bool,X204) | s(bool,X204) = s(bool,X205) | s(bool,'T_0') = s(bool,X205)) )),
  inference(subsumption_resolution,[],[f5884,f297])).
fof(f5884,plain,(
  ( ! [X204,X202,X200,X205,X203,X201] : (s(bool,X201) != s(bool,X202) | p(s(bool,X200)) | s(bool,X202) = s(bool,X203) | s(bool,'T_0') = s(bool,X203) | s(bool,X201) = s(bool,X204) | s(bool,'T_0') = s(bool,X200) | s(bool,X204) = s(bool,X205) | s(bool,'T_0') = s(bool,X205)) )),
  inference(superposition,[],[f1490,f1376])).
fof(f1490,plain,(
  ( ! [X2,X3,X1] : (s(bool,X1) != s(bool,X2) | p(s(bool,X2)) | s(bool,X1) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f297,f657])).
fof(f8507,plain,(
  ( ! [X6,X7,X5] : (s(bool,'T_0') != s(bool,X5) | s(bool,'T_0') = s(bool,X7) | s(bool,X6) = s(bool,X7) | s(bool,X5) = s(bool,X6)) ) | $spl90),
  inference(cnf_transformation,[],[f8507_D])).
fof(f8507_D,plain,(
  ( ! [X6,X7,X5] : (s(bool,'T_0') != s(bool,X5) | s(bool,'T_0') = s(bool,X7) | s(bool,X6) = s(bool,X7) | s(bool,X5) = s(bool,X6)) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f8506,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'T_0') != s(bool,X5) | p(s(bool,X4)) | s(bool,X5) = s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f5848,f297])).
fof(f5848,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'T_0') != s(bool,X5) | p(s(bool,X4)) | s(bool,X5) = s(bool,X6) | s(bool,'T_0') = s(bool,X4) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f297,f1376])).
fof(f8502,plain,(
  ( ! [X565,X563,X564,X566,X562] : (p(s(bool,X564)) | s(bool,'T_0') = s(bool,sK18(X563,bool,X562)) | s(bool,'T_0') = s(bool,sK17(X563,bool,X562)) | s(bool,X565) = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X562),X563))) | s(bool,X565) = s(bool,X566) | s(bool,'T_0') = s(bool,X566)) )),
  inference(subsumption_resolution,[],[f5835,f297])).
fof(f5835,plain,(
  ( ! [X565,X563,X564,X566,X562] : (p(s(bool,X564)) | s(bool,'T_0') = s(bool,sK18(X563,bool,X562)) | s(bool,'T_0') = s(bool,sK17(X563,bool,X562)) | s(bool,X565) = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X562),X563))) | s(bool,'T_0') = s(bool,X564) | s(bool,X565) = s(bool,X566) | s(bool,'T_0') = s(bool,X566)) )),
  inference(superposition,[],[f2344,f1376])).
fof(f8500,plain,(
  ( ! [X561,X557,X559,X560,X556,X558] : (p(s(bool,X559)) | s(X557,happ(s(fun(X556,X557),X558),s(X556,sK17(X558,X556,X557)))) = s(X557,happ(s(fun(X556,X557),X558),s(X556,sK18(X558,X556,X557)))) | s(bool,X560) = s(bool,'const.bool.ONE_ONE_1'(s(fun(X556,X557),X558))) | s(bool,X560) = s(bool,X561) | s(bool,'T_0') = s(bool,X561)) )),
  inference(subsumption_resolution,[],[f5834,f297])).
fof(f5834,plain,(
  ( ! [X561,X557,X559,X560,X556,X558] : (p(s(bool,X559)) | s(X557,happ(s(fun(X556,X557),X558),s(X556,sK17(X558,X556,X557)))) = s(X557,happ(s(fun(X556,X557),X558),s(X556,sK18(X558,X556,X557)))) | s(bool,X560) = s(bool,'const.bool.ONE_ONE_1'(s(fun(X556,X557),X558))) | s(bool,'T_0') = s(bool,X559) | s(bool,X560) = s(bool,X561) | s(bool,'T_0') = s(bool,X561)) )),
  inference(superposition,[],[f402,f1376])).
fof(f8481,plain,(
  ( ! [X231,X235,X233,X236,X234,X232] : (p(s(bool,X232)) | sP1(X231,X233) | s(bool,X233) = s(bool,X234) | s(bool,'T_0') = s(bool,X234) | s(bool,X231) = s(bool,X235) | s(bool,X235) = s(bool,X236) | s(bool,'T_0') = s(bool,X236)) )),
  inference(subsumption_resolution,[],[f5752,f297])).
fof(f5752,plain,(
  ( ! [X231,X235,X233,X236,X234,X232] : (p(s(bool,X232)) | sP1(X231,X233) | s(bool,X233) = s(bool,X234) | s(bool,'T_0') = s(bool,X234) | s(bool,X231) = s(bool,X235) | s(bool,'T_0') = s(bool,X232) | s(bool,X235) = s(bool,X236) | s(bool,'T_0') = s(bool,X236)) )),
  inference(superposition,[],[f1845,f1376])).
fof(f8479,plain,(
  ( ! [X229,X227,X225,X230,X228,X226] : (p(s(bool,X226)) | sP1(X227,X225) | s(bool,X227) = s(bool,X228) | s(bool,'T_0') = s(bool,X228) | s(bool,X225) = s(bool,X229) | s(bool,X229) = s(bool,X230) | s(bool,'T_0') = s(bool,X230)) )),
  inference(subsumption_resolution,[],[f5751,f297])).
fof(f5751,plain,(
  ( ! [X229,X227,X225,X230,X228,X226] : (p(s(bool,X226)) | sP1(X227,X225) | s(bool,X227) = s(bool,X228) | s(bool,'T_0') = s(bool,X228) | s(bool,X225) = s(bool,X229) | s(bool,'T_0') = s(bool,X226) | s(bool,X229) = s(bool,X230) | s(bool,'T_0') = s(bool,X230)) )),
  inference(superposition,[],[f1844,f1376])).
fof(f8477,plain,(
  ( ! [X212,X210,X213,X211] : (p(s(bool,X211)) | sP1(X210,'T_0') | s(bool,X210) = s(bool,X212) | s(bool,X212) = s(bool,X213) | s(bool,'T_0') = s(bool,X213)) ) | $spl8),
  inference(subsumption_resolution,[],[f5748,f297])).
fof(f5748,plain,(
  ( ! [X212,X210,X213,X211] : (p(s(bool,X211)) | sP1(X210,'T_0') | s(bool,X210) = s(bool,X212) | s(bool,'T_0') = s(bool,X211) | s(bool,X212) = s(bool,X213) | s(bool,'T_0') = s(bool,X213)) ) | $spl8),
  inference(superposition,[],[f1784,f1376])).
fof(f8475,plain,(
  ( ! [X208,X206,X209,X207] : (p(s(bool,X207)) | sP1('T_0',X206) | s(bool,X206) = s(bool,X208) | s(bool,X208) = s(bool,X209) | s(bool,'T_0') = s(bool,X209)) ) | $spl4),
  inference(subsumption_resolution,[],[f5747,f297])).
fof(f5747,plain,(
  ( ! [X208,X206,X209,X207] : (p(s(bool,X207)) | sP1('T_0',X206) | s(bool,X206) = s(bool,X208) | s(bool,'T_0') = s(bool,X207) | s(bool,X208) = s(bool,X209) | s(bool,'T_0') = s(bool,X209)) ) | $spl4),
  inference(superposition,[],[f1780,f1376])).
fof(f5075,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') = s(bool,sK18(X7,bool,X6)) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X6),X7))) | s(bool,'T_0') = s(bool,sK17(X7,bool,X6))) ) | $spl76),
  inference(cnf_transformation,[],[f5075_D])).
fof(f5075_D,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') = s(bool,sK18(X7,bool,X6)) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X6),X7))) | s(bool,'T_0') = s(bool,sK17(X7,bool,X6))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f5076,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X11)) | s(bool,'T_0') = s(bool,sK18(X10,bool,X9)) | s(bool,'T_0') = s(bool,sK17(X10,bool,X9)) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X9),X10)))) )),
  inference(subsumption_resolution,[],[f5072,f297])).
fof(f5072,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X11)) | s(bool,'T_0') = s(bool,sK18(X10,bool,X9)) | s(bool,'T_0') = s(bool,sK17(X10,bool,X9)) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X9),X10)))) )),
  inference(superposition,[],[f2344,f657])).
fof(f5074,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X8)) | s(bool,'T_0') = s(bool,sK18(X7,bool,X6)) | s(bool,'T_0') = s(bool,sK17(X7,bool,X6)) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X6),X7)))) )),
  inference(subsumption_resolution,[],[f5071,f297])).
fof(f5071,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X8)) | s(bool,'T_0') = s(bool,sK18(X7,bool,X6)) | s(bool,'T_0') = s(bool,sK17(X7,bool,X6)) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X6),X7))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f2344,f657])).
fof(f4321,plain,(
  ( ! [X138,X137] : (s(fun(bool,bool),X137) = s(fun(bool,bool),X138)) ) | $spl72),
  inference(cnf_transformation,[],[f4321_D])).
fof(f4321_D,plain,(
  ( ! [X138,X137] : (s(fun(bool,bool),X137) = s(fun(bool,bool),X138)) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f4323,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | $spl75),
  inference(cnf_transformation,[],[f4323_D])).
fof(f4323_D,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) <=> ~$spl75),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl75])])).
fof(f4320,plain,(
  ( ! [X138,X137] : (s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X137) = s(fun(bool,bool),X138)) ) | ($spl48 | $spl61)),
  inference(forward_demodulation,[],[f4283,f2473])).
fof(f2473,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | $spl61),
  inference(unit_resulting_resolution,[],[f501,f2460,f705])).
fof(f705,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f702,f501])).
fof(f702,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f360,f286])).
fof(f360,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f145])).
fof(f145,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X10,X11] : ((p(s(bool,X10)) => p(s(bool,X11))) => ((p(s(bool,X11)) => p(s(bool,X10))) => s(bool,X10) = s(bool,X11)))),
  file('thm.rat.RAT_EQ_LADD','thm.bool.IMP_ANTISYM_AX')).
fof(f501,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f280,f277])).
fof(f280,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f176])).
fof(f176,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f57])).
fof(f57,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f30])).
fof(f30,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.rat.RAT_EQ_LADD','thm.bool.NOT_CLAUSES|split|1')).
fof(f4283,plain,(
  ( ! [X138,X137] : (s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X137) = s(fun(bool,bool),X138)) ) | $spl48),
  inference(duplicate_literal_removal,[],[f4282])).
fof(f4282,plain,(
  ( ! [X138,X137] : (s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X137) = s(fun(bool,bool),X138) | s(fun(bool,bool),X137) = s(fun(bool,bool),X138)) ) | $spl48),
  inference(superposition,[],[f404,f2423])).
fof(f404,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK19(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK19(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f237])).
fof(f237,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK19(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK19(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK19])],[f154])).
fof(f154,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f119])).
fof(f119,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.rat.RAT_EQ_LADD',aHL_EXT)).
fof(f4043,plain,(
  ( ! [X28,X29] : (sP1('F_0',X28) | s(bool,'T_0') = s(bool,X29) | s(bool,X28) = s(bool,X29)) ) | $spl70),
  inference(cnf_transformation,[],[f4043_D])).
fof(f4043_D,plain,(
  ( ! [X28,X29] : (sP1('F_0',X28) | s(bool,'T_0') = s(bool,X29) | s(bool,X28) = s(bool,X29)) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f4042,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | sP1('F_0',X28) | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29)) )),
  inference(subsumption_resolution,[],[f4033,f297])).
fof(f4033,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | sP1('F_0',X28) | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f1845,f286])).
fof(f4039,plain,(
  ( ! [X10,X8,X11] : (sP1(X8,X10) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X11) | s(bool,X10) = s(bool,X11)) ) | $spl68),
  inference(cnf_transformation,[],[f4039_D])).
fof(f4039_D,plain,(
  ( ! [X10,X8,X11] : (sP1(X8,X10) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X11) | s(bool,X10) = s(bool,X11)) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f4040,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X13)) | sP1(X12,X14) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f4029,f297])).
fof(f4029,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X13)) | sP1(X12,X14) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f1845,f657])).
fof(f4038,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X9)) | sP1(X8,X10) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f4028,f297])).
fof(f4028,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X9)) | sP1(X8,X10) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f1845,f657])).
fof(f4009,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3065,f3073])).
fof(f3073,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1)) ) | $spl64),
  inference(unit_resulting_resolution,[],[f3057,f3057,f363])).
fof(f363,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f146])).
fof(f3065,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | $spl64),
  inference(unit_resulting_resolution,[],[f3057,f301])).
fof(f301,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f180])).
fof(f180,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f75])).
fof(f75,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X9] : (s(bool,'F_0') = s(bool,X9) <=> ~p(s(bool,X9)))),
  file('thm.rat.RAT_EQ_LADD','thm.bool.EQ_CLAUSES|split|3')).
fof(f4008,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3068,f277])).
fof(f3068,plain,(
  ~p(s(bool,'T_0')) | $spl64),
  inference(unit_resulting_resolution,[],[f529,f3057,f363])).
fof(f529,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f501,f297])).
fof(f4007,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3074,f277])).
fof(f3074,plain,(
  ~p(s(bool,'T_0')) | $spl64),
  inference(unit_resulting_resolution,[],[f529,f3057,f363])).
fof(f4006,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3083,f277])).
fof(f3083,plain,(
  ~p(s(bool,'T_0')) | $spl64),
  inference(unit_resulting_resolution,[],[f545,f3057,f384])).
fof(f384,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f228])).
fof(f545,plain,(
  sP1('F_0','T_0')),
  inference(unit_resulting_resolution,[],[f501,f277,f388])).
fof(f4005,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3084,f3057])).
fof(f3084,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0))))))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f560,f3057,f384])).
fof(f560,plain,(
  ( ! [X0] : (sP1('F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))))) )),
  inference(unit_resulting_resolution,[],[f501,f284,f388])).
fof(f4004,plain,(
  $false | ($spl61 | $spl64)),
  inference(subsumption_resolution,[],[f3085,f3057])).
fof(f3085,plain,(
  ~p(s(bool,f(s(bool,'T_0')))) | ($spl61 | $spl64)),
  inference(unit_resulting_resolution,[],[f3019,f3057,f384])).
fof(f3019,plain,(
  sP1('F_0',f(s(bool,'T_0'))) | $spl61),
  inference(unit_resulting_resolution,[],[f501,f2783])).
fof(f2783,plain,(
  ( ! [X6] : (sP1(X6,f(s(bool,'T_0'))) | p(s(bool,X6))) ) | $spl61),
  inference(subsumption_resolution,[],[f2736,f277])).
fof(f2736,plain,(
  ( ! [X6] : (~p(s(bool,'T_0')) | p(s(bool,X6)) | sP1(X6,f(s(bool,'T_0')))) ) | $spl61),
  inference(superposition,[],[f388,f2473])).
fof(f4003,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3092,f277])).
fof(f3092,plain,(
  ~p(s(bool,'T_0')) | $spl64),
  inference(unit_resulting_resolution,[],[f539,f3057,f384])).
fof(f539,plain,(
  sP1('T_0','F_0')),
  inference(unit_resulting_resolution,[],[f501,f277,f385])).
fof(f4002,plain,(
  $false | ($spl61 | $spl64)),
  inference(subsumption_resolution,[],[f3096,f3057])).
fof(f3096,plain,(
  ~p(s(bool,f(s(bool,'T_0')))) | ($spl61 | $spl64)),
  inference(unit_resulting_resolution,[],[f2968,f3057,f384])).
fof(f2968,plain,(
  sP1(f(s(bool,'T_0')),'F_0') | $spl61),
  inference(unit_resulting_resolution,[],[f501,f2782])).
fof(f2782,plain,(
  ( ! [X5] : (sP1(f(s(bool,'T_0')),X5) | p(s(bool,X5))) ) | $spl61),
  inference(subsumption_resolution,[],[f2735,f277])).
fof(f2735,plain,(
  ( ! [X5] : (~p(s(bool,'T_0')) | p(s(bool,X5)) | sP1(f(s(bool,'T_0')),X5)) ) | $spl61),
  inference(superposition,[],[f385,f2473])).
fof(f4001,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3097,f3057])).
fof(f3097,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0))))))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f559,f3057,f384])).
fof(f559,plain,(
  ( ! [X0] : (sP1('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f501,f284,f385])).
fof(f4000,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3116,f277])).
fof(f3116,plain,(
  ~p(s(bool,'T_0')) | $spl64),
  inference(unit_resulting_resolution,[],[f720,f501,f3057,f414])).
fof(f414,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f242])).
fof(f242,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | (sP2(X2,X1,X0) & sP3(X1,X2,X0) & sP4(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP2(X2,X1,X0) | ~sP3(X1,X2,X0) | ~sP4(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f241])).
fof(f241,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP5(X2,X1,X0)))),
  inference(flattening,[],[f240])).
fof(f240,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP5(X2,X1,X0)))),
  inference(nnf_transformation,[],[f163])).
fof(f163,plain,(
  ! [X2,X1,X0] : (sP5(X2,X1,X0) <=> (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f720,plain,(
  sP5('F_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f501,f529,f430])).
fof(f430,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f251])).
fof(f251,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP5(X2,X1,X0)) & (~sP5(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f164])).
fof(f164,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP5(X2,X1,X0))),
  inference(definition_folding,[],[f123,f163,f162,f161,f160])).
fof(f160,plain,(
  ! [X0,X1,X2] : (sP2(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f161,plain,(
  ! [X1,X0,X2] : (sP3(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f162,plain,(
  ! [X2,X0,X1] : (sP4(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f123,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f122])).
fof(f122,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f50])).
fof(f50,axiom,(
  ! [X22,X23,X24] : ((p(s(bool,X24)) <=> s(bool,X22) = s(bool,X23)) <=> ((p(s(bool,X24)) | p(s(bool,X23)) | p(s(bool,X22))) & (p(s(bool,X24)) | ~p(s(bool,X22)) | ~p(s(bool,X23))) & (p(s(bool,X23)) | ~p(s(bool,X22)) | ~p(s(bool,X24))) & (p(s(bool,X22)) | ~p(s(bool,X23)) | ~p(s(bool,X24)))))),
  file('thm.rat.RAT_EQ_LADD','thm.sat.sat_2|49|_prove_ltb')).
fof(f3999,plain,(
  $false | ($spl61 | $spl64)),
  inference(subsumption_resolution,[],[f3117,f3057])).
fof(f3117,plain,(
  ~p(s(bool,f(s(bool,'T_0')))) | ($spl61 | $spl64)),
  inference(unit_resulting_resolution,[],[f2470,f501,f3057,f414])).
fof(f2470,plain,(
  sP5('F_0',f(s(bool,'T_0')),'F_0') | $spl61),
  inference(unit_resulting_resolution,[],[f501,f2460,f430])).
fof(f3998,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3118,f3057])).
fof(f3118,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0))))))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f1160,f501,f3057,f414])).
fof(f1160,plain,(
  ( ! [X0] : (sP5('F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f564,f515,f512,f501,f417])).
fof(f417,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X2,X1) | p(s(bool,X0)) | sP5(X0,X1,X2) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f242])).
fof(f512,plain,(
  ( ! [X0,X1] : (sP4(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f501,f420])).
fof(f420,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f245])).
fof(f245,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f244])).
fof(f244,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(flattening,[],[f243])).
fof(f243,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(nnf_transformation,[],[f162])).
fof(f515,plain,(
  ( ! [X0,X1] : (sP3(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f501,f424])).
fof(f424,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f248])).
fof(f248,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f247])).
fof(f247,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(flattening,[],[f246])).
fof(f246,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(nnf_transformation,[],[f161])).
fof(f564,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X1)))),X2)) )),
  inference(unit_resulting_resolution,[],[f284,f428])).
fof(f428,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f250])).
fof(f250,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(flattening,[],[f249])).
fof(f249,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(nnf_transformation,[],[f160])).
fof(f3997,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3169,f277])).
fof(f3169,plain,(
  ~p(s(bool,'T_0')) | $spl64),
  inference(unit_resulting_resolution,[],[f1123,f501,f3057,f414])).
fof(f1123,plain,(
  sP5('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f517,f514,f512,f501,f416])).
fof(f416,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X2,X1) | p(s(bool,X1)) | sP5(X0,X1,X2) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f242])).
fof(f514,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f277,f423])).
fof(f423,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f248])).
fof(f517,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f277,f427])).
fof(f427,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f250])).
fof(f3996,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3170,f3057])).
fof(f3170,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0))))))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f1124,f501,f3057,f414])).
fof(f1124,plain,(
  ( ! [X0] : (sP5('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f563,f562,f512,f501,f416])).
fof(f562,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X2)))))) )),
  inference(unit_resulting_resolution,[],[f284,f423])).
fof(f563,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X2)))))) )),
  inference(unit_resulting_resolution,[],[f284,f427])).
fof(f3995,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3183,f3343])).
fof(f3343,plain,(
  ( ! [X0] : (~sP5(X0,'T_0','F_0')) ) | $spl64),
  inference(unit_resulting_resolution,[],[f529,f3057,f432])).
fof(f432,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f251])).
fof(f3183,plain,(
  ( ! [X0] : (sP5(X0,'T_0','F_0')) ) | $spl64),
  inference(unit_resulting_resolution,[],[f518,f515,f511,f3057,f415])).
fof(f415,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP5(X0,X1,X2) | ~sP4(X0,X2,X1) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f242])).
fof(f511,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f277,f419])).
fof(f419,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f245])).
fof(f518,plain,(
  ( ! [X0,X1] : (sP2(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f277,f428])).
fof(f3994,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3200,f3343])).
fof(f3200,plain,(
  ( ! [X0] : (sP5(X0,'T_0','F_0')) ) | $spl64),
  inference(unit_resulting_resolution,[],[f518,f515,f512,f3057,f415])).
fof(f3993,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3203,f736])).
fof(f736,plain,(
  ~sP5('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f277,f529,f432])).
fof(f3203,plain,(
  sP5('T_0','T_0','F_0') | $spl64),
  inference(unit_resulting_resolution,[],[f518,f514,f512,f3057,f415])).
fof(f3992,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3206,f3343])).
fof(f3206,plain,(
  ( ! [X0] : (sP5('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))),'T_0','F_0')) ) | $spl64),
  inference(unit_resulting_resolution,[],[f518,f562,f512,f3057,f415])).
fof(f3991,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3223,f3273])).
fof(f3273,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2)) ) | $spl64),
  inference(unit_resulting_resolution,[],[f3057,f419])).
fof(f3223,plain,(
  ~sP4('T_0','F_0','T_0') | $spl64),
  inference(unit_resulting_resolution,[],[f518,f515,f736,f3057,f415])).
fof(f3990,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3224,f3301])).
fof(f3301,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2)) ) | $spl64),
  inference(unit_resulting_resolution,[],[f3057,f423])).
fof(f3224,plain,(
  ~sP3('T_0','F_0','T_0') | $spl64),
  inference(unit_resulting_resolution,[],[f518,f512,f736,f3057,f415])).
fof(f3989,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3225,f3307])).
fof(f3307,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2)) ) | $spl64),
  inference(unit_resulting_resolution,[],[f3057,f429])).
fof(f429,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f250])).
fof(f3225,plain,(
  ~sP2('F_0','T_0','T_0') | $spl64),
  inference(unit_resulting_resolution,[],[f515,f512,f736,f3057,f415])).
fof(f3988,plain,(
  $false | ($spl61 | $spl64)),
  inference(subsumption_resolution,[],[f3227,f3273])).
fof(f3227,plain,(
  ~sP4('T_0','F_0',f(s(bool,'T_0'))) | ($spl61 | $spl64)),
  inference(unit_resulting_resolution,[],[f517,f515,f2469,f3057,f415])).
fof(f2469,plain,(
  ~sP5('T_0',f(s(bool,'T_0')),'F_0') | $spl61),
  inference(unit_resulting_resolution,[],[f277,f2460,f432])).
fof(f3987,plain,(
  $false | ($spl61 | $spl64)),
  inference(subsumption_resolution,[],[f3228,f3301])).
fof(f3228,plain,(
  ~sP3(f(s(bool,'T_0')),'F_0','T_0') | ($spl61 | $spl64)),
  inference(unit_resulting_resolution,[],[f517,f512,f2469,f3057,f415])).
fof(f3986,plain,(
  $false | ($spl61 | $spl64)),
  inference(subsumption_resolution,[],[f3229,f3307])).
fof(f3229,plain,(
  ~sP2('F_0',f(s(bool,'T_0')),'T_0') | ($spl61 | $spl64)),
  inference(unit_resulting_resolution,[],[f515,f512,f2469,f3057,f415])).
fof(f3985,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3244,f3057])).
fof(f3244,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f513,f501,f3057,f418])).
fof(f418,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f245])).
fof(f513,plain,(
  ( ! [X0,X1] : (sP4('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f501,f421])).
fof(f421,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f245])).
fof(f3984,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3256,f3273])).
fof(f3256,plain,(
  ( ! [X0,X1] : (~sP4(X0,X1,'F_0')) ) | $spl64),
  inference(unit_resulting_resolution,[],[f501,f3057,f3057,f418])).
fof(f3983,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3258,f3057])).
fof(f3258,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f512,f501,f3057,f418])).
fof(f3982,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3264,f3273])).
fof(f3264,plain,(
  ( ! [X0] : (~sP4('T_0',X0,'F_0')) ) | $spl64),
  inference(unit_resulting_resolution,[],[f501,f277,f3057,f418])).
fof(f3981,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3271,f3273])).
fof(f3271,plain,(
  ( ! [X0,X1] : (~sP4(X0,X1,'F_0')) ) | $spl64),
  inference(unit_resulting_resolution,[],[f501,f3057,f3057,f418])).
fof(f3980,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3276,f3057])).
fof(f3276,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f516,f501,f3057,f422])).
fof(f422,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f248])).
fof(f516,plain,(
  ( ! [X0,X1] : (sP3('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f501,f425])).
fof(f425,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f248])).
fof(f3979,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3286,f3301])).
fof(f3286,plain,(
  ( ! [X0,X1] : (~sP3(X0,X1,'F_0')) ) | $spl64),
  inference(unit_resulting_resolution,[],[f501,f3057,f3057,f422])).
fof(f3978,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3288,f3057])).
fof(f3288,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f515,f501,f3057,f422])).
fof(f3977,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3293,f3301])).
fof(f3293,plain,(
  ( ! [X0] : (~sP3('T_0',X0,'F_0')) ) | $spl64),
  inference(unit_resulting_resolution,[],[f501,f277,f3057,f422])).
fof(f3976,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3299,f3301])).
fof(f3299,plain,(
  ( ! [X0,X1] : (~sP3(X0,X1,'F_0')) ) | $spl64),
  inference(unit_resulting_resolution,[],[f501,f3057,f3057,f422])).
fof(f3972,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3314,f551])).
fof(f3314,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f755,f3057,f431])).
fof(f431,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f251])).
fof(f755,plain,(
  ( ! [X0] : (~sP5('F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f501,f551,f433])).
fof(f433,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f251])).
fof(f3971,plain,(
  $false | ($spl61 | $spl64)),
  inference(subsumption_resolution,[],[f3316,f2473])).
fof(f3316,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl61 | $spl64)),
  inference(unit_resulting_resolution,[],[f2712,f3057,f431])).
fof(f2712,plain,(
  ~sP5('F_0',f(s(bool,'T_0')),'T_0') | $spl61),
  inference(unit_resulting_resolution,[],[f501,f2473,f433])).
fof(f3970,plain,(
  $false | ($spl61 | $spl64)),
  inference(subsumption_resolution,[],[f3318,f2473])).
fof(f3318,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl61 | $spl64)),
  inference(unit_resulting_resolution,[],[f2714,f3057,f431])).
fof(f2714,plain,(
  ~sP5('F_0','T_0',f(s(bool,'T_0'))) | $spl61),
  inference(unit_resulting_resolution,[],[f501,f2473,f433])).
fof(f3969,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3319,f551])).
fof(f3319,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f757,f3057,f431])).
fof(f757,plain,(
  ( ! [X0] : (~sP5('F_0','T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))))) )),
  inference(unit_resulting_resolution,[],[f501,f551,f433])).
fof(f3968,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3328,f529])).
fof(f3328,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl64),
  inference(unit_resulting_resolution,[],[f720,f3057,f432])).
fof(f3967,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3329,f529])).
fof(f3329,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl64),
  inference(unit_resulting_resolution,[],[f719,f3057,f432])).
fof(f719,plain,(
  sP5('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f501,f529,f430])).
fof(f3966,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3965,f529])).
fof(f3965,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl64),
  inference(forward_demodulation,[],[f3330,f551])).
fof(f3330,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f1127,f3057,f432])).
fof(f1127,plain,(
  ( ! [X0] : (sP5('F_0','F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))))) )),
  inference(unit_resulting_resolution,[],[f565,f516,f513,f501,f416])).
fof(f565,plain,(
  ( ! [X2,X0,X1] : (sP2('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))),X1,X2)) )),
  inference(unit_resulting_resolution,[],[f284,f429])).
fof(f3964,plain,(
  $false | ($spl61 | $spl64)),
  inference(subsumption_resolution,[],[f3963,f529])).
fof(f3963,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl61 | $spl64)),
  inference(forward_demodulation,[],[f3331,f2473])).
fof(f3331,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | ($spl61 | $spl64)),
  inference(unit_resulting_resolution,[],[f2475,f3057,f432])).
fof(f2475,plain,(
  sP5('F_0','F_0',f(s(bool,'T_0'))) | $spl61),
  inference(unit_resulting_resolution,[],[f501,f2460,f430])).
fof(f3962,plain,(
  $false | ($spl61 | $spl64)),
  inference(subsumption_resolution,[],[f3961,f529])).
fof(f3961,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl61 | $spl64)),
  inference(forward_demodulation,[],[f3332,f2473])).
fof(f3332,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | ($spl61 | $spl64)),
  inference(unit_resulting_resolution,[],[f2470,f3057,f432])).
fof(f3960,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3959,f529])).
fof(f3959,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl64),
  inference(forward_demodulation,[],[f3333,f551])).
fof(f3333,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f1160,f3057,f432])).
fof(f3958,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3351,f3488])).
fof(f3488,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2)) ) | $spl64),
  inference(unit_resulting_resolution,[],[f3057,f3057,f3057,f447])).
fof(f447,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f258])).
fof(f258,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f257])).
fof(f257,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f167])).
fof(f167,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP7(X2,X0,X1))),
  inference(definition_folding,[],[f125,f166,f165])).
fof(f165,plain,(
  ! [X0,X1,X2] : (sP6(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f166,plain,(
  ! [X2,X0,X1] : (sP7(X2,X0,X1) <=> (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f125,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f124])).
fof(f124,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f51])).
fof(f51,axiom,(
  ! [X22,X23,X24] : ((p(s(bool,X24)) <=> (p(s(bool,X23)) & p(s(bool,X22)))) <=> ((p(s(bool,X24)) | ~p(s(bool,X23)) | ~p(s(bool,X22))) & (p(s(bool,X23)) | ~p(s(bool,X24))) & (p(s(bool,X22)) | ~p(s(bool,X24)))))),
  file('thm.rat.RAT_EQ_LADD','thm.sat.sat_23_prove_ltb')).
fof(f3351,plain,(
  ( ! [X0,X1] : (~sP7(X0,X1,'F_0')) ) | $spl64),
  inference(unit_resulting_resolution,[],[f501,f3057,f435])).
fof(f435,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f254])).
fof(f254,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | (sP6(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP6(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f253])).
fof(f253,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP7(X2,X0,X1)))),
  inference(flattening,[],[f252])).
fof(f252,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP7(X2,X0,X1)))),
  inference(nnf_transformation,[],[f166])).
fof(f3957,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3357,f3488])).
fof(f3357,plain,(
  ( ! [X0,X1] : (~sP7(X0,'F_0',X1)) ) | $spl64),
  inference(unit_resulting_resolution,[],[f501,f3057,f436])).
fof(f436,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f254])).
fof(f3956,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3360,f3057])).
fof(f3360,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f522,f583,f3057,f437])).
fof(f437,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f254])).
fof(f583,plain,(
  ( ! [X0] : (~sP7('T_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f277,f501,f436])).
fof(f522,plain,(
  ( ! [X0,X1] : (sP6('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f501,f444])).
fof(f444,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f256])).
fof(f256,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(flattening,[],[f255])).
fof(f255,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(nnf_transformation,[],[f165])).
fof(f3955,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3361,f3057])).
fof(f3361,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f520,f583,f3057,f437])).
fof(f520,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f277,f442])).
fof(f442,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f256])).
fof(f3954,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3366,f3426])).
fof(f3426,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2)) ) | $spl64),
  inference(unit_resulting_resolution,[],[f3057,f442])).
fof(f3366,plain,(
  ~sP6('F_0','T_0','T_0') | $spl64),
  inference(unit_resulting_resolution,[],[f583,f277,f3057,f437])).
fof(f3953,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3373,f3426])).
fof(f3373,plain,(
  ( ! [X0] : (~sP6(X0,'F_0','T_0')) ) | $spl64),
  inference(unit_resulting_resolution,[],[f579,f3057,f3057,f437])).
fof(f579,plain,(
  ( ! [X0] : (~sP7('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f277,f501,f435])).
fof(f3952,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3375,f3426])).
fof(f3375,plain,(
  ( ! [X0] : (~sP6('F_0',X0,'T_0')) ) | $spl64),
  inference(unit_resulting_resolution,[],[f583,f3057,f3057,f437])).
fof(f3951,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3379,f3057])).
fof(f3379,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f521,f579,f3057,f437])).
fof(f521,plain,(
  ( ! [X0,X1] : (sP6(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f501,f443])).
fof(f443,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f256])).
fof(f3950,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3380,f3057])).
fof(f3380,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f520,f579,f3057,f437])).
fof(f3949,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3386,f3426])).
fof(f3386,plain,(
  ~sP6('T_0','F_0','T_0') | $spl64),
  inference(unit_resulting_resolution,[],[f579,f277,f3057,f437])).
fof(f3948,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3394,f3426])).
fof(f3394,plain,(
  ( ! [X0] : (~sP6(X0,'F_0','T_0')) ) | $spl64),
  inference(unit_resulting_resolution,[],[f579,f3057,f3057,f437])).
fof(f3947,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3396,f3426])).
fof(f3396,plain,(
  ( ! [X0] : (~sP6('F_0',X0,'T_0')) ) | $spl64),
  inference(unit_resulting_resolution,[],[f583,f3057,f3057,f437])).
fof(f3946,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3401,f3057])).
fof(f3401,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f522,f501,f3057,f441])).
fof(f441,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f256])).
fof(f3945,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3411,f3426])).
fof(f3411,plain,(
  ( ! [X0,X1] : (~sP6(X0,X1,'F_0')) ) | $spl64),
  inference(unit_resulting_resolution,[],[f501,f3057,f3057,f441])).
fof(f3944,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3413,f3057])).
fof(f3413,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f521,f501,f3057,f441])).
fof(f3943,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3418,f3426])).
fof(f3418,plain,(
  ( ! [X0] : (~sP6('T_0',X0,'F_0')) ) | $spl64),
  inference(unit_resulting_resolution,[],[f501,f277,f3057,f441])).
fof(f3942,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3424,f3426])).
fof(f3424,plain,(
  ( ! [X0,X1] : (~sP6(X0,X1,'F_0')) ) | $spl64),
  inference(unit_resulting_resolution,[],[f501,f3057,f3057,f441])).
fof(f3941,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3429,f3057])).
fof(f3429,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f583,f277,f3057,f447])).
fof(f3940,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3431,f3057])).
fof(f3431,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f583,f3057,f3057,f447])).
fof(f3939,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3432,f277])).
fof(f3432,plain,(
  ~p(s(bool,'T_0')) | $spl64),
  inference(unit_resulting_resolution,[],[f867,f3057,f3057,f447])).
fof(f867,plain,(
  ~sP7('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f858,f434])).
fof(f434,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f254])).
fof(f858,plain,(
  ~sP6('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f501,f277,f277,f441])).
fof(f3938,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3433,f3057])).
fof(f3433,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0))))))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f1349,f3057,f3057,f447])).
fof(f1349,plain,(
  ( ! [X0] : (~sP7('F_0','T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))))) )),
  inference(unit_resulting_resolution,[],[f1337,f434])).
fof(f1337,plain,(
  ( ! [X0] : (~sP6('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f277,f501,f863])).
fof(f863,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X7)) | p(s(bool,X8)) | ~sP6(X7,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6)))),X8)) )),
  inference(subsumption_resolution,[],[f861,f277])).
fof(f861,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'T_0')) | ~p(s(bool,X7)) | p(s(bool,X8)) | ~sP6(X7,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6)))),X8)) )),
  inference(superposition,[],[f441,f551])).
fof(f3937,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3434,f277])).
fof(f3434,plain,(
  ~p(s(bool,'T_0')) | $spl64),
  inference(unit_resulting_resolution,[],[f583,f277,f3057,f447])).
fof(f3936,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3440,f277])).
fof(f3440,plain,(
  ~p(s(bool,'T_0')) | $spl64),
  inference(unit_resulting_resolution,[],[f579,f3057,f3057,f447])).
fof(f3935,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3441,f277])).
fof(f3441,plain,(
  ~p(s(bool,'T_0')) | $spl64),
  inference(unit_resulting_resolution,[],[f583,f3057,f3057,f447])).
fof(f3934,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3448,f3057])).
fof(f3448,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f579,f277,f3057,f447])).
fof(f3933,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3450,f3057])).
fof(f3450,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f579,f3057,f3057,f447])).
fof(f3932,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3452,f277])).
fof(f3452,plain,(
  ~p(s(bool,'T_0')) | $spl64),
  inference(unit_resulting_resolution,[],[f867,f3057,f3057,f447])).
fof(f3931,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3453,f277])).
fof(f3453,plain,(
  ~p(s(bool,'T_0')) | $spl64),
  inference(unit_resulting_resolution,[],[f1349,f3057,f3057,f447])).
fof(f3930,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3454,f277])).
fof(f3454,plain,(
  ~p(s(bool,'T_0')) | $spl64),
  inference(unit_resulting_resolution,[],[f579,f277,f3057,f447])).
fof(f3929,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3461,f277])).
fof(f3461,plain,(
  ~p(s(bool,'T_0')) | $spl64),
  inference(unit_resulting_resolution,[],[f579,f3057,f3057,f447])).
fof(f3928,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3462,f277])).
fof(f3462,plain,(
  ~p(s(bool,'T_0')) | $spl64),
  inference(unit_resulting_resolution,[],[f583,f3057,f3057,f447])).
fof(f3927,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3470,f277])).
fof(f3470,plain,(
  ~p(s(bool,'T_0')) | $spl64),
  inference(unit_resulting_resolution,[],[f867,f277,f3057,f447])).
fof(f3926,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3471,f3057])).
fof(f3471,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f579,f3057,f3057,f447])).
fof(f3925,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3473,f277])).
fof(f3473,plain,(
  ~p(s(bool,'T_0')) | $spl64),
  inference(unit_resulting_resolution,[],[f867,f3057,f3057,f447])).
fof(f3924,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3474,f277])).
fof(f3474,plain,(
  ~p(s(bool,'T_0')) | $spl64),
  inference(unit_resulting_resolution,[],[f1349,f3057,f3057,f447])).
fof(f3923,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3476,f277])).
fof(f3476,plain,(
  ~p(s(bool,'T_0')) | $spl64),
  inference(unit_resulting_resolution,[],[f867,f277,f3057,f447])).
fof(f3922,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3477,f3057])).
fof(f3477,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0))))))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f1349,f277,f3057,f447])).
fof(f3921,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3483,f3057])).
fof(f3483,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f583,f3057,f3057,f447])).
fof(f3920,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3484,f277])).
fof(f3484,plain,(
  ~p(s(bool,'T_0')) | $spl64),
  inference(unit_resulting_resolution,[],[f867,f3057,f3057,f447])).
fof(f3919,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3485,f3057])).
fof(f3485,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0))))))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f1349,f3057,f3057,f447])).
fof(f3918,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3491,f3057])).
fof(f3491,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f598,f501,f3057,f450])).
fof(f450,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP7(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f258])).
fof(f598,plain,(
  ( ! [X0] : (sP7('F_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f501,f501,f446])).
fof(f446,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f258])).
fof(f3917,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3503,f3057])).
fof(f3503,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f589,f501,f3057,f450])).
fof(f589,plain,(
  ( ! [X0] : (sP7('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f501,f501,f445])).
fof(f445,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f258])).
fof(f3916,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3523,f3789])).
fof(f3789,plain,(
  ( ! [X0,X1] : (sP9('F_0',X0,X1)) ) | $spl64),
  inference(unit_resulting_resolution,[],[f525,f3057,f486])).
fof(f486,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f454])).
fof(f454,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f261])).
fof(f261,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP8(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f260])).
fof(f260,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f259])).
fof(f259,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1)) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f169])).
fof(f169,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f525,plain,(
  ( ! [X0,X1] : (sP8('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f501,f461])).
fof(f461,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f264])).
fof(f264,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f263])).
fof(f263,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f262])).
fof(f262,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f168])).
fof(f168,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f3523,plain,(
  ( ! [X0,X1] : (~sP9('F_0',X0,X1)) ) | $spl64),
  inference(unit_resulting_resolution,[],[f501,f3057,f451])).
fof(f451,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f261])).
fof(f3915,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3532,f3789])).
fof(f3532,plain,(
  ( ! [X0,X1] : (~sP9('F_0',X0,X1)) ) | $spl64),
  inference(unit_resulting_resolution,[],[f501,f3057,f452])).
fof(f452,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f261])).
fof(f3914,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3542,f3569])).
fof(f3569,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2)) ) | $spl64),
  inference(unit_resulting_resolution,[],[f3057,f460])).
fof(f460,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f264])).
fof(f3542,plain,(
  ~sP8('F_0','T_0','F_0') | $spl64),
  inference(unit_resulting_resolution,[],[f609,f501,f3057,f455])).
fof(f455,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f261])).
fof(f609,plain,(
  ( ! [X0] : (~sP9('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f277,f501,f452])).
fof(f3913,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3552,f3569])).
fof(f3552,plain,(
  ~sP8('F_0','F_0','T_0') | $spl64),
  inference(unit_resulting_resolution,[],[f605,f501,f3057,f456])).
fof(f456,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f261])).
fof(f605,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f277,f501,f451])).
fof(f3912,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3565,f3569])).
fof(f3565,plain,(
  ( ! [X0] : (~sP8(X0,'F_0','F_0')) ) | $spl64),
  inference(unit_resulting_resolution,[],[f501,f501,f3057,f458])).
fof(f458,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f264])).
fof(f3911,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3573,f277])).
fof(f3573,plain,(
  ~p(s(bool,'T_0')) | $spl64),
  inference(unit_resulting_resolution,[],[f935,f3057,f463])).
fof(f463,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f266])).
fof(f266,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f265])).
fof(f265,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f170])).
fof(f170,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP9(X2,X0,X1))),
  inference(definition_folding,[],[f127,f169,f168])).
fof(f127,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f126])).
fof(f126,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X22,X23,X24] : ((p(s(bool,X24)) <=> (p(s(bool,X23)) | p(s(bool,X22)))) <=> ((p(s(bool,X24)) | ~p(s(bool,X23))) & (p(s(bool,X24)) | ~p(s(bool,X22))) & (p(s(bool,X23)) | p(s(bool,X22)) | ~p(s(bool,X24)))))),
  file('thm.rat.RAT_EQ_LADD','thm.sat.sat_25_prove_ltb')).
fof(f935,plain,(
  ~sP9('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f922,f453])).
fof(f453,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f261])).
fof(f922,plain,(
  ~sP8('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f501,f501,f277,f458])).
fof(f3910,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3578,f277])).
fof(f3578,plain,(
  ~p(s(bool,'T_0')) | $spl64),
  inference(unit_resulting_resolution,[],[f605,f3057,f463])).
fof(f3909,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3579,f3057])).
fof(f3579,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f609,f3057,f463])).
fof(f3908,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3587,f277])).
fof(f3587,plain,(
  ~p(s(bool,'T_0')) | $spl64),
  inference(unit_resulting_resolution,[],[f935,f3057,f464])).
fof(f464,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f266])).
fof(f3907,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3592,f3057])).
fof(f3592,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f605,f3057,f464])).
fof(f3906,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3593,f277])).
fof(f3593,plain,(
  ~p(s(bool,'T_0')) | $spl64),
  inference(unit_resulting_resolution,[],[f609,f3057,f464])).
fof(f3905,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3615,f3776])).
fof(f3776,plain,(
  ( ! [X0,X1] : (sP11('F_0',X0,X1)) ) | $spl64),
  inference(unit_resulting_resolution,[],[f528,f3057,f485])).
fof(f485,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP10(X0,X1,X2) | sP11(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f471])).
fof(f471,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f269])).
fof(f269,plain,(
  ! [X0,X1,X2] : ((~sP11(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP10(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP10(X0,X1,X2) | sP11(X0,X1,X2)))),
  inference(rectify,[],[f268])).
fof(f268,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1) | sP11(X2,X0,X1)))),
  inference(flattening,[],[f267])).
fof(f267,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1)) | sP11(X2,X0,X1)))),
  inference(nnf_transformation,[],[f172])).
fof(f172,plain,(
  ! [X2,X0,X1] : (sP11(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f528,plain,(
  ( ! [X0,X1] : (sP10('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f501,f478])).
fof(f478,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f272])).
fof(f272,plain,(
  ! [X0,X1,X2] : ((~sP10(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP10(X0,X1,X2)))),
  inference(rectify,[],[f271])).
fof(f271,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(flattening,[],[f270])).
fof(f270,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(nnf_transformation,[],[f171])).
fof(f171,plain,(
  ! [X2,X0,X1] : (sP10(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f3615,plain,(
  ( ! [X0,X1] : (~sP11('F_0',X0,X1)) ) | $spl64),
  inference(unit_resulting_resolution,[],[f501,f3057,f469])).
fof(f469,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f269])).
fof(f3904,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3618,f3057])).
fof(f3618,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f528,f623,f3057,f472])).
fof(f472,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f269])).
fof(f623,plain,(
  ( ! [X0] : (~sP11('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f277,f501,f469])).
fof(f3903,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3619,f3057])).
fof(f3619,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f527,f623,f3057,f472])).
fof(f527,plain,(
  ( ! [X0,X1] : (sP10(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f277,f477])).
fof(f477,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f272])).
fof(f3902,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3624,f3707])).
fof(f3707,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2)) ) | $spl64),
  inference(unit_resulting_resolution,[],[f3057,f477])).
fof(f3624,plain,(
  ~sP10('F_0','T_0','T_0') | $spl64),
  inference(unit_resulting_resolution,[],[f623,f277,f3057,f472])).
fof(f3901,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3631,f3707])).
fof(f3631,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'F_0')) ) | $spl64),
  inference(unit_resulting_resolution,[],[f622,f3057,f3057,f472])).
fof(f622,plain,(
  ( ! [X0] : (~sP11('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f501,f501,f468])).
fof(f468,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f269])).
fof(f3900,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3633,f3707])).
fof(f3633,plain,(
  ( ! [X0] : (~sP10('F_0','T_0',X0)) ) | $spl64),
  inference(unit_resulting_resolution,[],[f623,f3057,f3057,f472])).
fof(f3899,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3646,f3707])).
fof(f3646,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'F_0')) ) | $spl64),
  inference(unit_resulting_resolution,[],[f622,f3057,f3057,f472])).
fof(f3898,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3647,f3707])).
fof(f3647,plain,(
  ( ! [X0] : (~sP10('F_0','T_0',X0)) ) | $spl64),
  inference(unit_resulting_resolution,[],[f623,f3057,f3057,f472])).
fof(f3897,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3663,f3707])).
fof(f3663,plain,(
  ~sP10('F_0','F_0','F_0') | $spl64),
  inference(unit_resulting_resolution,[],[f622,f501,f3057,f473])).
fof(f473,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f269])).
fof(f3896,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3675,f3707])).
fof(f3675,plain,(
  ~sP10('F_0','F_0','F_0') | $spl64),
  inference(unit_resulting_resolution,[],[f622,f501,f3057,f474])).
fof(f474,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f269])).
fof(f3895,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3689,f3057])).
fof(f3689,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f528,f501,f3057,f475])).
fof(f475,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f272])).
fof(f3894,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3692,f3707])).
fof(f3692,plain,(
  ( ! [X0,X1] : (~sP10(X0,'F_0',X1)) ) | $spl64),
  inference(unit_resulting_resolution,[],[f3057,f501,f3057,f475])).
fof(f3893,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3701,f3057])).
fof(f3701,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f526,f501,f3057,f475])).
fof(f526,plain,(
  ( ! [X0,X1] : (sP10(X0,X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f501,f476])).
fof(f476,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f272])).
fof(f3892,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3703,f3707])).
fof(f3703,plain,(
  ( ! [X0] : (~sP10('T_0','F_0',X0)) ) | $spl64),
  inference(unit_resulting_resolution,[],[f277,f501,f3057,f475])).
fof(f3891,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3705,f3707])).
fof(f3705,plain,(
  ( ! [X0,X1] : (~sP10(X0,'F_0',X1)) ) | $spl64),
  inference(unit_resulting_resolution,[],[f3057,f501,f3057,f475])).
fof(f3890,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3722,f277])).
fof(f3722,plain,(
  ~p(s(bool,'T_0')) | $spl64),
  inference(unit_resulting_resolution,[],[f982,f3057,f481])).
fof(f481,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f274])).
fof(f274,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f273])).
fof(f273,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f173])).
fof(f173,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP11(X2,X0,X1))),
  inference(definition_folding,[],[f155,f172,f171])).
fof(f155,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f129])).
fof(f129,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f128])).
fof(f128,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f53])).
fof(f53,axiom,(
  ! [X22,X23,X24] : ((p(s(bool,X24)) <=> (p(s(bool,X23)) => p(s(bool,X22)))) <=> ((p(s(bool,X24)) | p(s(bool,X23))) & (p(s(bool,X24)) | ~p(s(bool,X22))) & (~p(s(bool,X23)) | p(s(bool,X22)) | ~p(s(bool,X24)))))),
  file('thm.rat.RAT_EQ_LADD','thm.sat.sat_27_prove_ltb')).
fof(f982,plain,(
  ~sP11('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f970,f470])).
fof(f470,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f269])).
fof(f970,plain,(
  ~sP10('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f277,f501,f277,f475])).
fof(f3889,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3727,f3057])).
fof(f3727,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f622,f3057,f481])).
fof(f3888,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3728,f277])).
fof(f3728,plain,(
  ~p(s(bool,'T_0')) | $spl64),
  inference(unit_resulting_resolution,[],[f623,f3057,f481])).
fof(f3887,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3734,f277])).
fof(f3734,plain,(
  ~p(s(bool,'T_0')) | $spl64),
  inference(unit_resulting_resolution,[],[f625,f501,f3057,f482])).
fof(f482,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f274])).
fof(f625,plain,(
  ( ! [X0] : (sP11('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f501,f277,f480])).
fof(f480,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f274])).
fof(f3886,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3736,f3057])).
fof(f3736,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0))))))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f1241,f501,f3057,f482])).
fof(f1241,plain,(
  ( ! [X0,X1] : (sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))),X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f501,f790])).
fof(f790,plain,(
  ( ! [X26,X27,X25] : (sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X25)))),X27,X26) | p(s(bool,X26))) )),
  inference(subsumption_resolution,[],[f775,f277])).
fof(f775,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'T_0')) | p(s(bool,X26)) | sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X25)))),X27,X26)) )),
  inference(superposition,[],[f480,f551])).
fof(f3885,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3751,f277])).
fof(f3751,plain,(
  ~p(s(bool,'T_0')) | $spl64),
  inference(unit_resulting_resolution,[],[f962,f501,f3057,f482])).
fof(f962,plain,(
  sP11('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f528,f277,f501,f474])).
fof(f3884,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3767,f526])).
fof(f3767,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'F_0')) ) | $spl64),
  inference(unit_resulting_resolution,[],[f622,f3057,f485])).
fof(f3883,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3768,f527])).
fof(f3768,plain,(
  ( ! [X0] : (~sP10('F_0','T_0',X0)) ) | $spl64),
  inference(unit_resulting_resolution,[],[f623,f3057,f485])).
fof(f3882,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3778,f523])).
fof(f523,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f277,f459])).
fof(f459,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f264])).
fof(f3778,plain,(
  ( ! [X0] : (~sP8('F_0',X0,'T_0')) ) | $spl64),
  inference(unit_resulting_resolution,[],[f605,f3057,f486])).
fof(f3881,plain,(
  $false | $spl64),
  inference(subsumption_resolution,[],[f3779,f524])).
fof(f524,plain,(
  ( ! [X0,X1] : (sP8(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f277,f460])).
fof(f3779,plain,(
  ( ! [X0] : (~sP8('F_0','T_0',X0)) ) | $spl64),
  inference(unit_resulting_resolution,[],[f609,f3057,f486])).
fof(f3857,plain,(
  $false | ($spl34 | $spl64)),
  inference(unit_resulting_resolution,[],[f625,f501,f3057,f1824])).
fof(f3852,plain,(
  $false | ($spl32 | $spl64)),
  inference(unit_resulting_resolution,[],[f982,f3057,f1821])).
fof(f3850,plain,(
  $false | ($spl28 | $spl64)),
  inference(unit_resulting_resolution,[],[f528,f501,f3057,f1815])).
fof(f3845,plain,(
  $false | ($spl24 | $spl64)),
  inference(unit_resulting_resolution,[],[f935,f3057,f1809])).
fof(f3843,plain,(
  $false | ($spl22 | $spl64)),
  inference(unit_resulting_resolution,[],[f935,f3057,f1806])).
fof(f3841,plain,(
  $false | ($spl18 | $spl64)),
  inference(unit_resulting_resolution,[],[f598,f501,f3057,f1800])).
fof(f3836,plain,(
  $false | ($spl16 | $spl64)),
  inference(unit_resulting_resolution,[],[f522,f501,f3057,f1797])).
fof(f3831,plain,(
  $false | ($spl14 | $spl64)),
  inference(unit_resulting_resolution,[],[f516,f501,f3057,f1794])).
fof(f3826,plain,(
  $false | ($spl12 | $spl64)),
  inference(unit_resulting_resolution,[],[f513,f501,f3057,f1791])).
fof(f3819,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f501,f1241,f3057,f999])).
fof(f999,plain,(
  ( ! [X6,X8,X7] : (~sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6)))),X8,X7) | p(s(bool,X8)) | ~p(s(bool,X7))) )),
  inference(subsumption_resolution,[],[f997,f277])).
fof(f997,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'T_0')) | ~p(s(bool,X7)) | p(s(bool,X8)) | ~sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6)))),X8,X7)) )),
  inference(superposition,[],[f482,f551])).
fof(f3816,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f501,f528,f3057,f975])).
fof(f975,plain,(
  ( ! [X6,X8,X7] : (~sP10(X8,X7,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6))))) | ~p(s(bool,X8)) | p(s(bool,X7))) )),
  inference(subsumption_resolution,[],[f973,f277])).
fof(f973,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'T_0')) | p(s(bool,X7)) | ~p(s(bool,X8)) | ~sP10(X8,X7,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6)))))) )),
  inference(superposition,[],[f475,f551])).
fof(f3811,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f501,f598,f3057,f894])).
fof(f894,plain,(
  ( ! [X6,X8,X7] : (~sP7(X8,X7,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6))))) | p(s(bool,X8)) | ~p(s(bool,X7))) )),
  inference(subsumption_resolution,[],[f893,f277])).
fof(f893,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'T_0')) | ~p(s(bool,X7)) | p(s(bool,X8)) | ~sP7(X8,X7,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6)))))) )),
  inference(superposition,[],[f450,f551])).
fof(f3808,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f501,f522,f3057,f863])).
fof(f3803,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f501,f516,f3057,f820])).
fof(f820,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X7)) | p(s(bool,X8)) | ~sP3(X7,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6)))),X8)) )),
  inference(subsumption_resolution,[],[f818,f277])).
fof(f818,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'T_0')) | ~p(s(bool,X7)) | p(s(bool,X8)) | ~sP3(X7,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6)))),X8)) )),
  inference(superposition,[],[f422,f551])).
fof(f3798,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f501,f513,f3057,f803])).
fof(f803,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X7)) | p(s(bool,X8)) | ~sP4(X7,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6)))),X8)) )),
  inference(subsumption_resolution,[],[f801,f277])).
fof(f801,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'T_0')) | ~p(s(bool,X7)) | p(s(bool,X8)) | ~sP4(X7,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6)))),X8)) )),
  inference(superposition,[],[f418,f551])).
fof(f3790,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f605,f525,f3057,f486])).
fof(f3788,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f605,f568,f3057,f486])).
fof(f568,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X1)))),X2)) )),
  inference(unit_resulting_resolution,[],[f284,f460])).
fof(f3786,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f609,f524,f3057,f486])).
fof(f3784,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f609,f567,f3057,f486])).
fof(f567,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X2)))))) )),
  inference(unit_resulting_resolution,[],[f284,f459])).
fof(f3782,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f605,f523,f3057,f486])).
fof(f3777,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f622,f528,f3057,f485])).
fof(f3775,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f622,f569,f3057,f485])).
fof(f569,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X1)))),X2)) )),
  inference(unit_resulting_resolution,[],[f284,f477])).
fof(f3773,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f623,f527,f3057,f485])).
fof(f3771,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f622,f526,f3057,f485])).
fof(f3766,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f625,f501,f3057,f3057,f482])).
fof(f3758,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f962,f501,f277,f3057,f482])).
fof(f3749,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f625,f501,f3057,f3057,f482])).
fof(f3741,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f625,f501,f277,f3057,f482])).
fof(f3733,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f622,f3057,f3057,f481])).
fof(f3731,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f623,f277,f3057,f481])).
fof(f3726,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f622,f3057,f3057,f481])).
fof(f3724,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f982,f277,f3057,f481])).
fof(f3719,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f622,f501,f3057,f480])).
fof(f3714,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f622,f501,f501,f3057,f479])).
fof(f479,plain,(
  ( ! [X2,X0,X1] : (sP11(X2,X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f274])).
fof(f3708,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f970,f3057,f477])).
fof(f3706,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f526,f3057,f501,f3057,f475])).
fof(f3704,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f526,f277,f501,f3057,f475])).
fof(f3693,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f526,f3057,f501,f3057,f475])).
fof(f3691,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f528,f277,f501,f3057,f475])).
fof(f3680,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f622,f528,f501,f3057,f474])).
fof(f3678,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f622,f526,f501,f3057,f474])).
fof(f3668,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f622,f528,f501,f3057,f473])).
fof(f3666,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f622,f526,f501,f3057,f473])).
fof(f3656,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f622,f528,f3057,f3057,f472])).
fof(f3654,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f622,f569,f3057,f3057,f472])).
fof(f3652,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f623,f527,f3057,f3057,f472])).
fof(f3650,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f622,f526,f3057,f3057,f472])).
fof(f3634,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f528,f623,f3057,f3057,f472])).
fof(f3632,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f528,f622,f3057,f3057,f472])).
fof(f3625,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f528,f623,f277,f3057,f472])).
fof(f3616,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f962,f501,f3057,f469])).
fof(f3608,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f913,f501,f501,f3057,f465])).
fof(f465,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)) | ~sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f266])).
fof(f913,plain,(
  sP9('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f525,f501,f501,f457])).
fof(f457,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f261])).
fof(f3598,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f605,f3057,f3057,f464])).
fof(f3596,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f609,f277,f3057,f464])).
fof(f3591,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f605,f3057,f3057,f464])).
fof(f3589,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f935,f277,f3057,f464])).
fof(f3584,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f605,f3057,f3057,f463])).
fof(f3582,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f605,f277,f3057,f463])).
fof(f3577,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f605,f3057,f3057,f463])).
fof(f3575,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f935,f277,f3057,f463])).
fof(f3570,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f922,f3057,f460])).
fof(f3568,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f922,f3057,f459])).
fof(f3566,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f525,f501,f501,f3057,f458])).
fof(f3558,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f605,f525,f501,f3057,f456])).
fof(f3555,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f605,f523,f501,f3057,f456])).
fof(f3543,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f525,f609,f501,f3057,f455])).
fof(f3533,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f913,f501,f3057,f452])).
fof(f3524,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f913,f501,f3057,f451])).
fof(f3515,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f589,f501,f3057,f3057,f450])).
fof(f3509,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f589,f501,f277,f3057,f450])).
fof(f3502,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f589,f501,f3057,f3057,f450])).
fof(f3496,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f598,f501,f277,f3057,f450])).
fof(f3489,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f579,f3057,f3057,f3057,f447])).
fof(f3487,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f583,f277,f3057,f3057,f447])).
fof(f3481,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f579,f3057,f277,f3057,f447])).
fof(f3479,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f867,f277,f277,f3057,f447])).
fof(f3468,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f579,f3057,f3057,f3057,f447])).
fof(f3466,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f579,f277,f3057,f3057,f447])).
fof(f3460,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f579,f3057,f277,f3057,f447])).
fof(f3458,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f579,f277,f277,f3057,f447])).
fof(f3447,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f579,f3057,f3057,f3057,f447])).
fof(f3445,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f579,f277,f3057,f3057,f447])).
fof(f3439,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f583,f3057,f277,f3057,f447])).
fof(f3437,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f583,f277,f277,f3057,f447])).
fof(f3427,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f858,f3057,f442])).
fof(f3425,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f521,f501,f3057,f3057,f441])).
fof(f3419,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f521,f501,f277,f3057,f441])).
fof(f3412,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f521,f501,f3057,f3057,f441])).
fof(f3406,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f522,f501,f277,f3057,f441])).
fof(f3397,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f522,f583,f3057,f3057,f437])).
fof(f3395,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f521,f579,f3057,f3057,f437])).
fof(f3387,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f521,f579,f277,f3057,f437])).
fof(f3376,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f522,f583,f3057,f3057,f437])).
fof(f3374,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f521,f579,f3057,f3057,f437])).
fof(f3367,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f522,f583,f277,f3057,f437])).
fof(f3358,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f598,f501,f3057,f436])).
fof(f3352,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f589,f501,f3057,f435])).
fof(f3346,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f719,f529,f3057,f432])).
fof(f3344,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f720,f529,f3057,f432])).
fof(f3327,plain,(
  $false | ($spl61 | $spl64)),
  inference(unit_resulting_resolution,[],[f2714,f2473,f3057,f431])).
fof(f3325,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f757,f551,f3057,f431])).
fof(f3323,plain,(
  $false | ($spl61 | $spl64)),
  inference(unit_resulting_resolution,[],[f2712,f2473,f3057,f431])).
fof(f3321,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f755,f551,f3057,f431])).
fof(f3880,plain,(
  $false | $spl64),
  inference(trivial_inequality_removal,[],[f3317])).
fof(f3317,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl64),
  inference(unit_resulting_resolution,[],[f839,f3057,f431])).
fof(f839,plain,(
  ~sP5('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f830,f411])).
fof(f411,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f242])).
fof(f830,plain,(
  ~sP2('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f501,f501,f501,f426])).
fof(f426,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f250])).
fof(f3879,plain,(
  $false | $spl64),
  inference(trivial_inequality_removal,[],[f3315])).
fof(f3315,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl64),
  inference(unit_resulting_resolution,[],[f824,f3057,f431])).
fof(f824,plain,(
  ~sP5('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f815,f412])).
fof(f412,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP3(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f242])).
fof(f815,plain,(
  ~sP3('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f501,f277,f277,f422])).
fof(f3308,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f830,f3057,f429])).
fof(f3306,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f830,f3057,f428])).
fof(f3304,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f830,f3057,f427])).
fof(f3302,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f815,f3057,f423])).
fof(f3300,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f515,f501,f3057,f3057,f422])).
fof(f3294,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f515,f501,f277,f3057,f422])).
fof(f3287,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f515,f501,f3057,f3057,f422])).
fof(f3281,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f516,f501,f277,f3057,f422])).
fof(f3274,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f798,f3057,f419])).
fof(f798,plain,(
  ~sP4('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f501,f277,f277,f418])).
fof(f3272,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f512,f501,f3057,f3057,f418])).
fof(f3265,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f512,f501,f277,f3057,f418])).
fof(f3257,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f512,f501,f3057,f3057,f418])).
fof(f3250,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f513,f501,f277,f3057,f418])).
fof(f3230,plain,(
  $false | ($spl61 | $spl64)),
  inference(unit_resulting_resolution,[],[f517,f515,f512,f2469,f3057,f415])).
fof(f3226,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f518,f515,f512,f736,f3057,f415])).
fof(f3174,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f720,f3057,f501,f3057,f414])).
fof(f3172,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f1123,f277,f501,f3057,f414])).
fof(f3144,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f720,f501,f3057,f3057,f414])).
fof(f3127,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f720,f501,f277,f3057,f414])).
fof(f3115,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f536,f501,f3057,f388])).
fof(f536,plain,(
  ~sP1('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f501,f501,f383])).
fof(f383,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f228])).
fof(f3108,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f536,f501,f3057,f385])).
fof(f3101,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f539,f3057,f3057,f384])).
fof(f3099,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f539,f277,f3057,f384])).
fof(f3091,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f539,f3057,f3057,f384])).
fof(f3089,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f545,f277,f3057,f384])).
fof(f3081,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f550,f3057,f370])).
fof(f370,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f221])).
fof(f221,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f220])).
fof(f220,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f156])).
fof(f156,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f550,plain,(
  ~sP0('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f501,f501,f502])).
fof(f502,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f368,f501])).
fof(f368,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f221])).
fof(f3079,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f550,f3057,f369])).
fof(f369,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f221])).
fof(f3077,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f3057,f529,f3057,f363])).
fof(f3075,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f277,f529,f3057,f363])).
fof(f3071,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f3057,f529,f3057,f363])).
fof(f3069,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f277,f529,f3057,f363])).
fof(f3064,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f529,f3057,f298])).
fof(f3062,plain,(
  $false | $spl64),
  inference(unit_resulting_resolution,[],[f501,f3057])).
fof(f3061,plain,(
  ( ! [X28,X29] : (sP1(X28,'F_0') | s(bool,'T_0') = s(bool,X29) | s(bool,X28) = s(bool,X29)) ) | $spl66),
  inference(cnf_transformation,[],[f3061_D])).
fof(f3061_D,plain,(
  ( ! [X28,X29] : (sP1(X28,'F_0') | s(bool,'T_0') = s(bool,X29) | s(bool,X28) = s(bool,X29)) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f3060,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | sP1(X28,'F_0') | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29)) )),
  inference(subsumption_resolution,[],[f3050,f297])).
fof(f3050,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | sP1(X28,'F_0') | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f1844,f286])).
fof(f3056,plain,(
  ( ! [X10,X8,X11] : (sP1(X10,X8) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X11) | s(bool,X10) = s(bool,X11)) ) | $spl62),
  inference(cnf_transformation,[],[f3056_D])).
fof(f3056_D,plain,(
  ( ! [X10,X8,X11] : (sP1(X10,X8) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X11) | s(bool,X10) = s(bool,X11)) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f3058,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X13)) | sP1(X14,X12) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f3046,f297])).
fof(f3046,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X13)) | sP1(X14,X12) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f1844,f657])).
fof(f3055,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X9)) | sP1(X10,X8) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f3045,f297])).
fof(f3045,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X9)) | sP1(X10,X8) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f1844,f657])).
fof(f2458,plain,(
  ( ! [X2,X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X2) = s(bool,sK19(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK19(X0,X1,bool,bool))))) ) | $spl58),
  inference(cnf_transformation,[],[f2458_D])).
fof(f2458_D,plain,(
  ( ! [X2,X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X2) = s(bool,sK19(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK19(X0,X1,bool,bool))))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f2460,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | $spl61),
  inference(cnf_transformation,[],[f2460_D])).
fof(f2460_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl61),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl61])])).
fof(f2448,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,g(s(bool,sK19(X3,X4,bool,bool)))) | s(bool,X5) = s(bool,sK19(X3,X4,bool,bool)) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1293,f657])).
fof(f1293,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK19(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK19(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f404,f286])).
fof(f2447,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,g(s(bool,sK19(X0,X1,bool,bool)))) | s(bool,X2) = s(bool,sK19(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1293,f657])).
fof(f2427,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl52),
  inference(cnf_transformation,[],[f2427_D])).
fof(f2427_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f2423,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK19(X0,X1,bool,bool)) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1)) ) | $spl48),
  inference(cnf_transformation,[],[f2423_D])).
fof(f2423_D,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK19(X0,X1,bool,bool)) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1)) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f2430,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl57),
  inference(cnf_transformation,[],[f2430_D])).
fof(f2430_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl57),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl57])])).
fof(f2421,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X12) = s(fun(bool,bool),X13) | s(bool,'T_0') = s(bool,sK19(X12,X13,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f1292,f286])).
fof(f1292,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK19(X0,X1,X2,bool))) )),
  inference(superposition,[],[f404,f286])).
fof(f2428,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl54),
  inference(cnf_transformation,[],[f2428_D])).
fof(f2428_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f2420,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X10) = s(fun(bool,bool),X11) | s(bool,'T_0') = s(bool,sK19(X10,X11,bool,bool)) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f1292,f657])).
fof(f2419,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X7) = s(fun(bool,bool),X8) | s(bool,'T_0') = s(bool,sK19(X7,X8,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f1292,f657])).
fof(f2422,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl46),
  inference(cnf_transformation,[],[f2422_D])).
fof(f2422_D,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f2425,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl51),
  inference(cnf_transformation,[],[f2425_D])).
fof(f2425_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl51),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl51])])).
fof(f2418,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,sK19(X3,X4,bool,bool)) | s(bool,X5) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1292,f657])).
fof(f2417,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK19(X0,X1,bool,bool)) | s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1292,f657])).
fof(f1948,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f276,f275,f277,f501,f284,f286,f298,f297,f302,f301,f357,f356,f358,f359,f363,f360,f370,f369,f502,f388,f385,f384,f383,f403,f402,f401,f404,f417,f416,f415,f414,f413,f412,f411,f421,f420,f419,f418,f425,f424,f423,f422,f429,f428,f427,f426,f433,f432,f431,f430,f487,f437,f436,f435,f434,f444,f443,f442,f441,f450,f447,f446,f445,f457,f456,f455,f486,f453,f452,f451,f461,f460,f459,f458,f465,f464,f463,f462,f474,f473,f472,f485,f470,f469,f468,f478,f477,f476,f475,f482,f481,f480,f479,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f523,f524,f525,f526,f527,f528,f529,f536,f537,f539,f545,f550,f569,f568,f567,f566,f565,f564,f563,f562,f561,f560,f559,f558,f557,f555,f554,f553,f552,f551,f579,f583,f589,f598,f605,f609,f612,f615,f622,f623,f625,f632,f657,f705,f716,f720,f719,f727,f730,f736,f734,f758,f757,f756,f755,f786,f787,f788,f789,f790,f791,f792,f793,f798,f803,f815,f820,f824,f830,f839,f844,f843,f858,f863,f867,f894,f913,f922,f928,f935,f954,f962,f970,f975,f982,f999,f1000,f1070,f1098,f1092,f1086,f1083,f1104,f1127,f1124,f1123,f1160,f1155,f1178,f1241,f1255,f1257,f1261,f1271,f1292,f1293,f1294,f1302,f1306,f1314,f1333,f1336,f1337,f1345,f1349,f1376,f1438,f1439,f1440,f1441,f1442,f1490,f1779,f1783,f1787,f1790,f1793,f1796,f1799,f1802,f1805,f1808,f1811,f1814,f1817,f1820,f1823,f1826,f1829,f1832,f1835,f1838,f1840,f1841,f1842,f1843,f1844,f1845,f1846,f1847,f1848,f1849,f1850,f1851,f1852,f1853,f1854,f1855,f1856,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1564,f1565,f1566,f1567,f1568,f1873,f1875,f1878,f1880,f1882,f1884,f1886,f1888,f1890,f1892,f1894,f1896,f1898,f1900,f1902,f1904,f1906,f1908,f1910,f1912,f1781,f1942,f1947])).
fof(f1947,plain,(
  ( ! [X6,X2,X0,X5,X1] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) | s(X1,X5) = s(X1,X6)) ) | $spl6),
  inference(subsumption_resolution,[],[f1932,f277])).
fof(f1932,plain,(
  ( ! [X6,X2,X0,X5,X1] : (~p(s(bool,'T_0')) | s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) | s(X1,X5) = s(X1,X6)) ) | $spl6),
  inference(backward_demodulation,[],[f1781,f401])).
fof(f1942,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl6),
  inference(subsumption_resolution,[],[f1913,f1781])).
fof(f1913,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl6),
  inference(backward_demodulation,[],[f1781,f301])).
fof(f1568,plain,(
  ( ! [X156,X154,X157,X155] : (s(X156,f(s(bool,'T_0'))) != s(X156,g(s(bool,'T_0'))) | s(fun(bool,X156),X154) = s(fun(bool,X156),X155) | s(bool,X157) = s(bool,sK19(X154,X155,X156,bool)) | s(bool,'T_0') = s(bool,X157)) )),
  inference(superposition,[],[f404,f657])).
fof(f1567,plain,(
  ( ! [X152,X151,X153] : (s(bool,'T_0') != s(bool,sK17(X151,bool,X152)) | p(s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X152),X151)))) | s(bool,X153) = s(bool,sK18(X151,bool,X152)) | s(bool,'T_0') = s(bool,X153)) )),
  inference(superposition,[],[f403,f657])).
fof(f1566,plain,(
  ( ! [X146,X147,X148] : (~p(s(bool,happ(s(fun(bool,bool),X146),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X146),s(bool,X147)))) | s(bool,X148) = s(bool,sK16(X146,bool)) | s(bool,'T_0') = s(bool,X148)) )),
  inference(superposition,[],[f357,f657])).
fof(f1565,plain,(
  ( ! [X144,X145,X143] : (~p(s(bool,happ(s(fun(bool,bool),X143),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X143),s(bool,X144)))) | s(bool,X145) = s(bool,sK15(X143,bool)) | s(bool,'T_0') = s(bool,X145)) )),
  inference(superposition,[],[f356,f657])).
fof(f1564,plain,(
  ( ! [X142,X140,X141,X139] : (s(bool,'T_0') != s(bool,f(s(X139,sK19(X140,X141,bool,X139)))) | s(fun(X139,bool),X140) = s(fun(X139,bool),X141) | s(bool,X142) = s(bool,g(s(X139,sK19(X140,X141,bool,X139)))) | s(bool,'T_0') = s(bool,X142)) )),
  inference(superposition,[],[f404,f657])).
fof(f1867,plain,(
  ( ! [X127,X130,X128,X126,X131,X129] : (s(X127,happ(s(fun(X126,X127),X128),s(X126,X129))) != s(X127,happ(s(fun(X126,X127),X128),s(X126,X130))) | s(X126,X129) = s(X126,X130) | s(bool,X131) = s(bool,'const.bool.ONE_ONE_1'(s(fun(X126,X127),X128))) | s(bool,'T_0') = s(bool,X131)) )),
  inference(subsumption_resolution,[],[f1562,f277])).
fof(f1562,plain,(
  ( ! [X127,X130,X128,X126,X131,X129] : (~p(s(bool,'T_0')) | s(X127,happ(s(fun(X126,X127),X128),s(X126,X129))) != s(X127,happ(s(fun(X126,X127),X128),s(X126,X130))) | s(X126,X129) = s(X126,X130) | s(bool,X131) = s(bool,'const.bool.ONE_ONE_1'(s(fun(X126,X127),X128))) | s(bool,'T_0') = s(bool,X131)) )),
  inference(superposition,[],[f401,f657])).
fof(f1866,plain,(
  ( ! [X123,X121,X124,X122] : (p(s(bool,happ(s(fun(X121,bool),X122),s(X121,X123)))) | s(bool,X124) = s(bool,happ(s(fun(X121,bool),X122),s(X121,sK16(X122,X121)))) | s(bool,'T_0') = s(bool,X124)) )),
  inference(subsumption_resolution,[],[f1561,f277])).
fof(f1561,plain,(
  ( ! [X123,X121,X124,X122] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X121,bool),X122),s(X121,X123)))) | s(bool,X124) = s(bool,happ(s(fun(X121,bool),X122),s(X121,sK16(X122,X121)))) | s(bool,'T_0') = s(bool,X124)) )),
  inference(superposition,[],[f357,f657])).
fof(f1865,plain,(
  ( ! [X118,X120,X119,X117] : (p(s(bool,happ(s(fun(X117,bool),X118),s(X117,X119)))) | s(bool,X120) = s(bool,happ(s(fun(X117,bool),X118),s(X117,sK15(X118,X117)))) | s(bool,'T_0') = s(bool,X120)) )),
  inference(subsumption_resolution,[],[f1560,f277])).
fof(f1560,plain,(
  ( ! [X118,X120,X119,X117] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X117,bool),X118),s(X117,X119)))) | s(bool,X120) = s(bool,happ(s(fun(X117,bool),X118),s(X117,sK15(X118,X117)))) | s(bool,'T_0') = s(bool,X120)) )),
  inference(superposition,[],[f356,f657])).
fof(f1864,plain,(
  ( ! [X109,X107,X108,X106] : (p(s(bool,X107)) | ~sP6(X106,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X108)))),X107) | s(bool,X106) = s(bool,X109) | s(bool,'T_0') = s(bool,X109)) )),
  inference(subsumption_resolution,[],[f1554,f277])).
fof(f1554,plain,(
  ( ! [X109,X107,X108,X106] : (~p(s(bool,'T_0')) | p(s(bool,X107)) | ~sP6(X106,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X108)))),X107) | s(bool,X106) = s(bool,X109) | s(bool,'T_0') = s(bool,X109)) )),
  inference(superposition,[],[f863,f657])).
fof(f1863,plain,(
  ( ! [X103,X105,X102,X104] : (p(s(bool,X103)) | ~sP3(X102,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X104)))),X103) | s(bool,X102) = s(bool,X105) | s(bool,'T_0') = s(bool,X105)) )),
  inference(subsumption_resolution,[],[f1553,f277])).
fof(f1553,plain,(
  ( ! [X103,X105,X102,X104] : (~p(s(bool,'T_0')) | p(s(bool,X103)) | ~sP3(X102,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X104)))),X103) | s(bool,X102) = s(bool,X105) | s(bool,'T_0') = s(bool,X105)) )),
  inference(superposition,[],[f820,f657])).
fof(f1862,plain,(
  ( ! [X101,X99,X100,X98] : (p(s(bool,X99)) | ~sP4(X98,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X100)))),X99) | s(bool,X98) = s(bool,X101) | s(bool,'T_0') = s(bool,X101)) )),
  inference(subsumption_resolution,[],[f1552,f277])).
fof(f1552,plain,(
  ( ! [X101,X99,X100,X98] : (~p(s(bool,'T_0')) | p(s(bool,X99)) | ~sP4(X98,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X100)))),X99) | s(bool,X98) = s(bool,X101) | s(bool,'T_0') = s(bool,X101)) )),
  inference(superposition,[],[f803,f657])).
fof(f1861,plain,(
  ( ! [X94,X97,X95,X96] : (sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X95)))),X94,X96) | s(bool,X94) = s(bool,X97) | s(bool,'T_0') = s(bool,X97)) )),
  inference(subsumption_resolution,[],[f1551,f277])).
fof(f1551,plain,(
  ( ! [X94,X97,X95,X96] : (~p(s(bool,'T_0')) | sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X95)))),X94,X96) | s(bool,X94) = s(bool,X97) | s(bool,'T_0') = s(bool,X97)) )),
  inference(superposition,[],[f791,f657])).
fof(f1860,plain,(
  ( ! [X90,X88,X87,X89] : (~sP8(X87,X88,X89) | sP9(X87,X88,X89) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(subsumption_resolution,[],[f1549,f277])).
fof(f1549,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,'T_0')) | ~sP8(X87,X88,X89) | sP9(X87,X88,X89) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f486,f657])).
fof(f1859,plain,(
  ( ! [X85,X83,X86,X84] : (~sP10(X83,X84,X85) | sP11(X83,X84,X85) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(subsumption_resolution,[],[f1548,f277])).
fof(f1548,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,'T_0')) | ~sP10(X83,X84,X85) | sP11(X83,X84,X85) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f485,f657])).
fof(f1858,plain,(
  ( ! [X80,X81,X79,X82] : (~sP11(X79,X81,X80) | p(s(bool,X81)) | ~p(s(bool,X80)) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(subsumption_resolution,[],[f1547,f277])).
fof(f1547,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,'T_0')) | ~p(s(bool,X80)) | p(s(bool,X81)) | ~sP11(X79,X81,X80) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(superposition,[],[f482,f657])).
fof(f1857,plain,(
  ( ! [X78,X76,X77,X75] : (sP11(X75,X76,X77) | ~p(s(bool,X76)) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(subsumption_resolution,[],[f1546,f277])).
fof(f1546,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,'T_0')) | ~p(s(bool,X76)) | sP11(X75,X76,X77) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f481,f657])).
fof(f1856,plain,(
  ( ! [X74,X72,X71,X73] : (sP11(X71,X73,X72) | p(s(bool,X72)) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(subsumption_resolution,[],[f1545,f277])).
fof(f1545,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,'T_0')) | p(s(bool,X72)) | sP11(X71,X73,X72) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(superposition,[],[f480,f657])).
fof(f1855,plain,(
  ( ! [X70,X68,X69,X67] : (~sP10(X69,X68,X67) | ~p(s(bool,X69)) | p(s(bool,X68)) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f1544,f277])).
fof(f1544,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,'T_0')) | p(s(bool,X68)) | ~p(s(bool,X69)) | ~sP10(X69,X68,X67) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f475,f657])).
fof(f1854,plain,(
  ( ! [X66,X64,X65,X63] : (~sP9(X63,X65,X64) | p(s(bool,X65)) | p(s(bool,X64)) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(subsumption_resolution,[],[f1543,f277])).
fof(f1543,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,'T_0')) | p(s(bool,X64)) | p(s(bool,X65)) | ~sP9(X63,X65,X64) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f465,f657])).
fof(f1853,plain,(
  ( ! [X61,X59,X62,X60] : (sP9(X59,X60,X61) | ~p(s(bool,X60)) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(subsumption_resolution,[],[f1542,f277])).
fof(f1542,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,'T_0')) | ~p(s(bool,X60)) | sP9(X59,X60,X61) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f464,f657])).
fof(f1852,plain,(
  ( ! [X57,X58,X56,X55] : (sP9(X55,X57,X56) | ~p(s(bool,X56)) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(subsumption_resolution,[],[f1541,f277])).
fof(f1541,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,'T_0')) | ~p(s(bool,X56)) | sP9(X55,X57,X56) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f463,f657])).
fof(f1851,plain,(
  ( ! [X54,X52,X53,X51] : (~sP8(X51,X52,X53) | p(s(bool,X53)) | p(s(bool,X52)) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(subsumption_resolution,[],[f1540,f277])).
fof(f1540,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,'T_0')) | p(s(bool,X52)) | p(s(bool,X53)) | ~sP8(X51,X52,X53) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(superposition,[],[f458,f657])).
fof(f1850,plain,(
  ( ! [X47,X50,X48,X49] : (~sP7(X49,X48,X47) | p(s(bool,X49)) | ~p(s(bool,X48)) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f1539,f277])).
fof(f1539,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,X48)) | p(s(bool,X49)) | ~sP7(X49,X48,X47) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f450,f657])).
fof(f1849,plain,(
  ( ! [X45,X43,X46,X44] : (~sP6(X44,X43,X45) | p(s(bool,X45)) | ~p(s(bool,X44)) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f1538,f277])).
fof(f1538,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,'T_0')) | ~p(s(bool,X44)) | p(s(bool,X45)) | ~sP6(X44,X43,X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f441,f657])).
fof(f1848,plain,(
  ( ! [X39,X41,X42,X40] : (~sP3(X40,X39,X41) | p(s(bool,X41)) | ~p(s(bool,X40)) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(subsumption_resolution,[],[f1537,f277])).
fof(f1537,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X40)) | p(s(bool,X41)) | ~sP3(X40,X39,X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f422,f657])).
fof(f1847,plain,(
  ( ! [X37,X35,X38,X36] : (~sP4(X36,X35,X37) | p(s(bool,X37)) | ~p(s(bool,X36)) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(subsumption_resolution,[],[f1536,f277])).
fof(f1536,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'T_0')) | ~p(s(bool,X36)) | p(s(bool,X37)) | ~sP4(X36,X35,X37) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f418,f657])).
fof(f1846,plain,(
  ( ! [X33,X31,X34,X32] : (~sP4(X32,X31,X33) | sP5(X32,X33,X31) | ~sP3(X33,X31,X32) | ~sP2(X31,X33,X32) | s(bool,X31) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f1535,f277])).
fof(f1535,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'T_0')) | sP5(X32,X33,X31) | ~sP4(X32,X31,X33) | ~sP3(X33,X31,X32) | ~sP2(X31,X33,X32) | s(bool,X31) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f415,f657])).
fof(f1843,plain,(
  ( ! [X24,X23,X22] : (sP0(X22,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f1532,f277])).
fof(f1532,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | sP0(X22,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f370,f657])).
fof(f1842,plain,(
  ( ! [X21,X19,X20] : (sP0(X20,X19) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(subsumption_resolution,[],[f1531,f277])).
fof(f1531,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | sP0(X20,X19) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f369,f657])).
fof(f1841,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f1530,f277])).
fof(f1530,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f363,f657])).
fof(f1840,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f1526,f277])).
fof(f1526,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f301,f657])).
fof(f1442,plain,(
  ( ! [X158,X161,X159,X160] : (s(X160,f(s(bool,X161))) != s(X160,g(s(bool,X161))) | s(fun(bool,X160),X158) = s(fun(bool,X160),X159) | s(bool,'T_0') = s(bool,X161) | s(bool,'T_0') = s(bool,sK19(X158,X159,X160,bool))) )),
  inference(superposition,[],[f404,f657])).
fof(f1440,plain,(
  ( ! [X151,X149,X150] : (~p(s(bool,happ(s(fun(bool,bool),X149),s(bool,X150)))) | p(s(bool,happ(s(fun(bool,bool),X149),s(bool,X151)))) | s(bool,'T_0') = s(bool,X150) | s(bool,'T_0') = s(bool,sK16(X149,bool))) )),
  inference(superposition,[],[f357,f657])).
fof(f1439,plain,(
  ( ! [X146,X147,X148] : (~p(s(bool,happ(s(fun(bool,bool),X146),s(bool,X147)))) | p(s(bool,happ(s(fun(bool,bool),X146),s(bool,X148)))) | s(bool,'T_0') = s(bool,X147) | s(bool,'T_0') = s(bool,sK15(X146,bool))) )),
  inference(superposition,[],[f356,f657])).
fof(f1438,plain,(
  ( ! [X144,X142,X145,X143] : (s(bool,X145) != s(bool,f(s(X142,sK19(X143,X144,bool,X142)))) | s(fun(X142,bool),X143) = s(fun(X142,bool),X144) | s(bool,'T_0') = s(bool,X145) | s(bool,'T_0') = s(bool,g(s(X142,sK19(X143,X144,bool,X142))))) )),
  inference(superposition,[],[f404,f657])).
fof(f1345,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X7)) | ~sP6('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6)))),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X8)))),X7)) )),
  inference(subsumption_resolution,[],[f1343,f277])).
fof(f1343,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'T_0')) | p(s(bool,X7)) | ~sP6('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6)))),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X8)))),X7)) )),
  inference(superposition,[],[f863,f551])).
fof(f1333,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X7)) | ~sP3('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6)))),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X8)))),X7)) )),
  inference(subsumption_resolution,[],[f1331,f277])).
fof(f1331,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'T_0')) | p(s(bool,X7)) | ~sP3('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6)))),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X8)))),X7)) )),
  inference(superposition,[],[f820,f551])).
fof(f1314,plain,(
  ( ! [X2,X3,X1] : (s(X2,happ(s(fun(X1,X2),X3),s(X1,sK17(X3,X1,X2)))) = s(X2,happ(s(fun(X1,X2),X3),s(X1,sK18(X3,X1,X2)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X2),X3)))) )),
  inference(subsumption_resolution,[],[f1313,f501])).
fof(f1313,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,'F_0')) | s(X2,happ(s(fun(X1,X2),X3),s(X1,sK17(X3,X1,X2)))) = s(X2,happ(s(fun(X1,X2),X3),s(X1,sK18(X3,X1,X2)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X2),X3)))) )),
  inference(superposition,[],[f402,f286])).
fof(f1306,plain,(
  ( ! [X0] : (~sP5('T_0','F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))))) )),
  inference(unit_resulting_resolution,[],[f1294,f413])).
fof(f1302,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X7)) | ~sP4('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6)))),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X8)))),X7)) )),
  inference(subsumption_resolution,[],[f1300,f277])).
fof(f1300,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'T_0')) | p(s(bool,X7)) | ~sP4('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6)))),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X8)))),X7)) )),
  inference(superposition,[],[f803,f551])).
fof(f1294,plain,(
  ( ! [X0] : (~sP4('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f277,f501,f803])).
fof(f1271,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X1)))) | s(bool,'T_0') = s(bool,sK16(X0,bool))) )),
  inference(superposition,[],[f357,f286])).
fof(f1261,plain,(
  ( ! [X6,X8,X7] : (sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X7)))),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6)))),X8)) )),
  inference(subsumption_resolution,[],[f1260,f277])).
fof(f1260,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'T_0')) | sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X7)))),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6)))),X8)) )),
  inference(superposition,[],[f791,f551])).
fof(f1257,plain,(
  ( ! [X0,X1] : (sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))),'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f277,f791])).
fof(f1255,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X1)))) | s(bool,'T_0') = s(bool,sK15(X0,bool))) )),
  inference(superposition,[],[f356,f286])).
fof(f1155,plain,(
  ( ! [X0] : (~sP3('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f518,f757,f511,f501,f417])).
fof(f1104,plain,(
  ( ! [X6,X8,X7] : (~sP4(X7,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6)))),X8) | sP5(X7,X8,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6))))) | ~sP3(X8,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6)))),X7)) )),
  inference(subsumption_resolution,[],[f1103,f565])).
fof(f1103,plain,(
  ( ! [X6,X8,X7] : (sP5(X7,X8,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6))))) | ~sP4(X7,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6)))),X8) | ~sP3(X8,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6)))),X7) | ~sP2('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6)))),X8,X7)) )),
  inference(subsumption_resolution,[],[f1102,f277])).
fof(f1102,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'T_0')) | sP5(X7,X8,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6))))) | ~sP4(X7,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6)))),X8) | ~sP3(X8,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6)))),X7) | ~sP2('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6)))),X8,X7)) )),
  inference(superposition,[],[f415,f551])).
fof(f1083,plain,(
  sP5('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f519,f514,f511,f277,f415])).
fof(f1086,plain,(
  ( ! [X0] : (sP5('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))),'T_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f563,f562,f511,f277,f415])).
fof(f1092,plain,(
  ( ! [X0,X1] : (sP5('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X1)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f563,f562,f561,f277,f415])).
fof(f1098,plain,(
  ( ! [X0] : (~sP3('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f564,f561,f755,f277,f415])).
fof(f1070,plain,(
  ( ! [X0,X1] : (sP9('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))),'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f277,f789])).
fof(f1000,plain,(
  ( ! [X0,X1] : (sP9('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))),X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f277,f788])).
fof(f954,plain,(
  ( ! [X6,X8,X7] : (~sP9('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6)))),X8,X7) | p(s(bool,X8)) | p(s(bool,X7))) )),
  inference(subsumption_resolution,[],[f951,f277])).
fof(f951,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'T_0')) | p(s(bool,X7)) | p(s(bool,X8)) | ~sP9('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6)))),X8,X7)) )),
  inference(superposition,[],[f465,f551])).
fof(f928,plain,(
  ( ! [X6,X8,X7] : (~sP8('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6)))),X7,X8) | p(s(bool,X8)) | p(s(bool,X7))) )),
  inference(subsumption_resolution,[],[f925,f277])).
fof(f925,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'T_0')) | p(s(bool,X7)) | p(s(bool,X8)) | ~sP8('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X6)))),X7,X8)) )),
  inference(superposition,[],[f458,f551])).
fof(f843,plain,(
  sP7('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f520,f277,f277,f437])).
fof(f844,plain,(
  ( ! [X0] : (sP7('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))),'T_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f566,f277,f277,f437])).
fof(f793,plain,(
  ( ! [X35,X36,X34] : (sP9('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X34)))),X35,X36) | ~sP8('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X34)))),X35,X36)) )),
  inference(subsumption_resolution,[],[f778,f277])).
fof(f778,plain,(
  ( ! [X35,X36,X34] : (~p(s(bool,'T_0')) | ~sP8('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X34)))),X35,X36) | sP9('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X34)))),X35,X36)) )),
  inference(superposition,[],[f486,f551])).
fof(f792,plain,(
  ( ! [X33,X31,X32] : (sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X31)))),X32,X33) | ~sP10('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X31)))),X32,X33)) )),
  inference(subsumption_resolution,[],[f777,f277])).
fof(f777,plain,(
  ( ! [X33,X31,X32] : (~p(s(bool,'T_0')) | ~sP10('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X31)))),X32,X33) | sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X31)))),X32,X33)) )),
  inference(superposition,[],[f485,f551])).
fof(f791,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,X29)) | sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X28)))),X29,X30)) )),
  inference(subsumption_resolution,[],[f776,f277])).
fof(f776,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'T_0')) | ~p(s(bool,X29)) | sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X28)))),X29,X30)) )),
  inference(superposition,[],[f481,f551])).
fof(f789,plain,(
  ( ! [X24,X23,X22] : (sP9('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X22)))),X23,X24) | ~p(s(bool,X23))) )),
  inference(subsumption_resolution,[],[f774,f277])).
fof(f774,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | ~p(s(bool,X23)) | sP9('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X22)))),X23,X24)) )),
  inference(superposition,[],[f464,f551])).
fof(f788,plain,(
  ( ! [X21,X19,X20] : (sP9('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X19)))),X21,X20) | ~p(s(bool,X20))) )),
  inference(subsumption_resolution,[],[f773,f277])).
fof(f773,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | ~p(s(bool,X20)) | sP9('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X19)))),X21,X20)) )),
  inference(superposition,[],[f463,f551])).
fof(f787,plain,(
  ( ! [X17,X18] : (sP1(X18,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X17))))) | p(s(bool,X18))) )),
  inference(subsumption_resolution,[],[f772,f277])).
fof(f772,plain,(
  ( ! [X17,X18] : (~p(s(bool,'T_0')) | p(s(bool,X18)) | sP1(X18,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X17)))))) )),
  inference(superposition,[],[f388,f551])).
fof(f786,plain,(
  ( ! [X15,X16] : (sP1('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X15)))),X16) | p(s(bool,X16))) )),
  inference(subsumption_resolution,[],[f771,f277])).
fof(f771,plain,(
  ( ! [X15,X16] : (~p(s(bool,'T_0')) | p(s(bool,X16)) | sP1('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X15)))),X16)) )),
  inference(superposition,[],[f385,f551])).
fof(f756,plain,(
  ( ! [X0] : (sP5('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f277,f551,f431])).
fof(f758,plain,(
  ( ! [X0] : (sP5('T_0','T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))))) )),
  inference(unit_resulting_resolution,[],[f277,f551,f431])).
fof(f734,plain,(
  ~sP5('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f277,f529,f432])).
fof(f730,plain,(
  ( ! [X0,X1] : (sP11('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))),X1)) )),
  inference(unit_resulting_resolution,[],[f277,f569,f485])).
fof(f727,plain,(
  ( ! [X0,X1] : (sP9('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))),X1)) )),
  inference(unit_resulting_resolution,[],[f277,f568,f486])).
fof(f716,plain,(
  ( ! [X0,X1] : (sP9('T_0',X0,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X1)))))) )),
  inference(unit_resulting_resolution,[],[f277,f567,f486])).
fof(f632,plain,(
  ( ! [X0] : (sP11('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f277,f277,f481])).
fof(f615,plain,(
  ( ! [X0] : (sP9('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f277,f277,f464])).
fof(f612,plain,(
  ( ! [X0] : (sP9('T_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f277,f277,f463])).
fof(f552,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))))) )),
  inference(unit_resulting_resolution,[],[f284,f301])).
fof(f553,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X1)))))) )),
  inference(unit_resulting_resolution,[],[f284,f369])).
fof(f554,plain,(
  ( ! [X0,X1] : (sP0('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))),X1)) )),
  inference(unit_resulting_resolution,[],[f284,f370])).
fof(f555,plain,(
  ( ! [X0] : (~sP1('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f277,f284,f384])).
fof(f557,plain,(
  ( ! [X0] : (~sP1('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))))) )),
  inference(unit_resulting_resolution,[],[f277,f284,f384])).
fof(f558,plain,(
  ( ! [X0,X1] : (~sP1('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X0)))),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X1)))))) )),
  inference(unit_resulting_resolution,[],[f284,f284,f384])).
fof(f561,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X2)))))) )),
  inference(unit_resulting_resolution,[],[f284,f419])).
fof(f566,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X2)))))) )),
  inference(unit_resulting_resolution,[],[f284,f442])).
fof(f537,plain,(
  ~sP1('T_0','T_0')),
  inference(unit_resulting_resolution,[],[f277,f277,f384])).
fof(f519,plain,(
  ( ! [X0,X1] : (sP2('T_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f277,f429])).
fof(f510,plain,(
  ( ! [X0] : (sP0('T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f277,f370])).
fof(f509,plain,(
  ( ! [X0] : (sP0(X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f277,f369])).
fof(f462,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f266])).
fof(f487,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(duplicate_literal_removal,[],[f440])).
fof(f440,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f254])).
fof(f413,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP4(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f242])).
fof(f356,plain,(
  ( ! [X0,X5,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK15(X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))) )),
  inference(cnf_transformation,[],[f217])).
fof(f217,plain,(
  ! [X0,X1] : ((! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK15(X1,X0)))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK16(X1,X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK15,sK16])],[f216])).
fof(f216,plain,(
  ! [X0,X1] : ((! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(rectify,[],[f215])).
fof(f215,plain,(
  ! [X0,X1] : ((! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(nnf_transformation,[],[f144])).
fof(f144,plain,(
  ! [X0,X1] : (? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(ennf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0,X1] : (~! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(flattening,[],[f101])).
fof(f101,plain,(
  ! [X0] : ! [X1] : (~! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X12] : ! [X14] : (~! [X13] : p(s(bool,happ(s(fun(X12,bool),X14),s(X12,X13)))) <=> ? [X13] : ~p(s(bool,happ(s(fun(X12,bool),X14),s(X12,X13)))))),
  file('thm.rat.RAT_EQ_LADD','thm.bool.NOT_FORALL_THM')).
fof(f357,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK16(X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f217])).
fof(f302,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f180])).
fof(f1944,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f1915,f277])).
fof(f1915,plain,(
  ~p(s(bool,'T_0')) | $spl6),
  inference(backward_demodulation,[],[f1781,f501])).
fof(f1943,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f276,f275,f277,f501,f284,f286,f298,f297,f302,f301,f357,f356,f358,f359,f363,f360,f370,f369,f502,f388,f385,f384,f383,f403,f402,f401,f404,f417,f416,f415,f414,f413,f412,f411,f421,f420,f419,f418,f425,f424,f423,f422,f429,f428,f427,f426,f433,f432,f431,f430,f487,f437,f436,f435,f434,f444,f443,f442,f441,f450,f447,f446,f445,f457,f456,f455,f486,f453,f452,f451,f461,f460,f459,f458,f465,f464,f463,f462,f474,f473,f472,f485,f470,f469,f468,f478,f477,f476,f475,f482,f481,f480,f479,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f523,f524,f525,f526,f527,f528,f529,f536,f537,f539,f545,f550,f569,f568,f567,f566,f565,f564,f563,f562,f561,f560,f559,f558,f557,f555,f554,f553,f552,f551,f579,f583,f589,f598,f605,f609,f612,f615,f622,f623,f625,f632,f657,f705,f716,f720,f719,f727,f730,f736,f734,f758,f757,f756,f755,f786,f787,f788,f789,f790,f791,f792,f793,f798,f803,f815,f820,f824,f830,f839,f844,f843,f858,f863,f867,f894,f913,f922,f928,f935,f954,f962,f970,f975,f982,f999,f1000,f1070,f1098,f1092,f1086,f1083,f1104,f1127,f1124,f1123,f1160,f1155,f1178,f1241,f1255,f1257,f1261,f1271,f1292,f1293,f1294,f1302,f1306,f1314,f1333,f1336,f1337,f1345,f1349,f1376,f1438,f1439,f1440,f1441,f1442,f1490,f1779,f1783,f1787,f1790,f1793,f1796,f1799,f1802,f1805,f1808,f1811,f1814,f1817,f1820,f1823,f1826,f1829,f1832,f1835,f1838,f1840,f1841,f1842,f1843,f1844,f1845,f1846,f1847,f1848,f1849,f1850,f1851,f1852,f1853,f1854,f1855,f1856,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1564,f1565,f1566,f1567,f1568,f1873,f1875,f1878,f1880,f1882,f1884,f1886,f1888,f1890,f1892,f1894,f1896,f1898,f1900,f1902,f1904,f1906,f1908,f1910,f1912,f1781,f1942])).
fof(f1839,plain,(
  ( ! [X2,X1] : (~sP6('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X2)))),X1) | p(s(bool,X1))) ) | $spl44),
  inference(cnf_transformation,[],[f1839_D])).
fof(f1839_D,plain,(
  ( ! [X2,X1] : (~sP6('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X2)))),X1) | p(s(bool,X1))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f1781,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl6),
  inference(cnf_transformation,[],[f1781_D])).
fof(f1781_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f1912,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP6('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X2)))),X1)) )),
  inference(condensation,[],[f1911])).
fof(f1911,plain,(
  ( ! [X111,X112,X110,X113] : (p(s(bool,X111)) | ~sP6('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X112)))),X111) | s(bool,X110) = s(bool,X113) | s(bool,'T_0') = s(bool,X113)) )),
  inference(subsumption_resolution,[],[f1601,f705])).
fof(f1601,plain,(
  ( ! [X111,X112,X110,X113] : (~p(s(bool,X110)) | p(s(bool,X111)) | ~sP6('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X112)))),X111) | s(bool,X110) = s(bool,X113) | s(bool,'T_0') = s(bool,X113)) )),
  inference(superposition,[],[f863,f657])).
fof(f1836,plain,(
  ( ! [X2,X1] : (~sP3('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X2)))),X1) | p(s(bool,X1))) ) | $spl42),
  inference(cnf_transformation,[],[f1836_D])).
fof(f1836_D,plain,(
  ( ! [X2,X1] : (~sP3('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X2)))),X1) | p(s(bool,X1))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f1910,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP3('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X2)))),X1)) )),
  inference(condensation,[],[f1909])).
fof(f1909,plain,(
  ( ! [X109,X107,X108,X106] : (p(s(bool,X107)) | ~sP3('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X108)))),X107) | s(bool,X106) = s(bool,X109) | s(bool,'T_0') = s(bool,X109)) )),
  inference(subsumption_resolution,[],[f1600,f705])).
fof(f1600,plain,(
  ( ! [X109,X107,X108,X106] : (~p(s(bool,X106)) | p(s(bool,X107)) | ~sP3('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X108)))),X107) | s(bool,X106) = s(bool,X109) | s(bool,'T_0') = s(bool,X109)) )),
  inference(superposition,[],[f820,f657])).
fof(f1833,plain,(
  ( ! [X2,X1] : (~sP4('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X2)))),X1) | p(s(bool,X1))) ) | $spl40),
  inference(cnf_transformation,[],[f1833_D])).
fof(f1833_D,plain,(
  ( ! [X2,X1] : (~sP4('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X2)))),X1) | p(s(bool,X1))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f1908,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP4('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X2)))),X1)) )),
  inference(condensation,[],[f1907])).
fof(f1907,plain,(
  ( ! [X103,X105,X102,X104] : (p(s(bool,X103)) | ~sP4('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X104)))),X103) | s(bool,X102) = s(bool,X105) | s(bool,'T_0') = s(bool,X105)) )),
  inference(subsumption_resolution,[],[f1599,f705])).
fof(f1599,plain,(
  ( ! [X103,X105,X102,X104] : (~p(s(bool,X102)) | p(s(bool,X103)) | ~sP4('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X104)))),X103) | s(bool,X102) = s(bool,X105) | s(bool,'T_0') = s(bool,X105)) )),
  inference(superposition,[],[f803,f657])).
fof(f1830,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~sP8('T_0',X1,X2)) ) | $spl38),
  inference(cnf_transformation,[],[f1830_D])).
fof(f1830_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~sP8('T_0',X1,X2)) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f1906,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('T_0',X1,X2) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f1905])).
fof(f1905,plain,(
  ( ! [X94,X92,X93,X91] : (~sP8('T_0',X92,X93) | sP9('T_0',X92,X93) | s(bool,X91) = s(bool,X94) | s(bool,'T_0') = s(bool,X94)) )),
  inference(subsumption_resolution,[],[f1596,f705])).
fof(f1596,plain,(
  ( ! [X94,X92,X93,X91] : (~p(s(bool,X91)) | ~sP8('T_0',X92,X93) | sP9('T_0',X92,X93) | s(bool,X91) = s(bool,X94) | s(bool,'T_0') = s(bool,X94)) )),
  inference(superposition,[],[f486,f657])).
fof(f1827,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~sP10('T_0',X1,X2)) ) | $spl36),
  inference(cnf_transformation,[],[f1827_D])).
fof(f1827_D,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~sP10('T_0',X1,X2)) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f1904,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP10('T_0',X1,X2) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f1903])).
fof(f1903,plain,(
  ( ! [X90,X88,X87,X89] : (~sP10('T_0',X88,X89) | sP11('T_0',X88,X89) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(subsumption_resolution,[],[f1595,f705])).
fof(f1595,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,X87)) | ~sP10('T_0',X88,X89) | sP11('T_0',X88,X89) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f485,f657])).
fof(f1824,plain,(
  ( ! [X2,X1] : (~sP11('T_0',X2,X1) | p(s(bool,X2)) | ~p(s(bool,X1))) ) | $spl34),
  inference(cnf_transformation,[],[f1824_D])).
fof(f1824_D,plain,(
  ( ! [X2,X1] : (~sP11('T_0',X2,X1) | p(s(bool,X2)) | ~p(s(bool,X1))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f1902,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f1901])).
fof(f1901,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,X84)) | p(s(bool,X85)) | ~sP11('T_0',X85,X84) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(subsumption_resolution,[],[f1594,f705])).
fof(f1594,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,X83)) | ~p(s(bool,X84)) | p(s(bool,X85)) | ~sP11('T_0',X85,X84) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f482,f657])).
fof(f1821,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~p(s(bool,X1))) ) | $spl32),
  inference(cnf_transformation,[],[f1821_D])).
fof(f1821_D,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~p(s(bool,X1))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f1900,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f1899])).
fof(f1899,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,X80)) | sP11('T_0',X80,X81) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(subsumption_resolution,[],[f1593,f705])).
fof(f1593,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,X79)) | ~p(s(bool,X80)) | sP11('T_0',X80,X81) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(superposition,[],[f481,f657])).
fof(f1818,plain,(
  ( ! [X2,X1] : (sP11('T_0',X2,X1) | p(s(bool,X1))) ) | $spl30),
  inference(cnf_transformation,[],[f1818_D])).
fof(f1818_D,plain,(
  ( ! [X2,X1] : (sP11('T_0',X2,X1) | p(s(bool,X1))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f1898,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f1897])).
fof(f1897,plain,(
  ( ! [X78,X76,X77,X75] : (p(s(bool,X76)) | sP11('T_0',X77,X76) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(subsumption_resolution,[],[f1592,f705])).
fof(f1592,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,X75)) | p(s(bool,X76)) | sP11('T_0',X77,X76) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f480,f657])).
fof(f1815,plain,(
  ( ! [X2,X1] : (~sP10(X2,X1,'T_0') | ~p(s(bool,X2)) | p(s(bool,X1))) ) | $spl28),
  inference(cnf_transformation,[],[f1815_D])).
fof(f1815_D,plain,(
  ( ! [X2,X1] : (~sP10(X2,X1,'T_0') | ~p(s(bool,X2)) | p(s(bool,X1))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f1896,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP10(X2,X1,'T_0')) )),
  inference(condensation,[],[f1895])).
fof(f1895,plain,(
  ( ! [X74,X72,X71,X73] : (p(s(bool,X72)) | ~p(s(bool,X73)) | ~sP10(X73,X72,'T_0') | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(subsumption_resolution,[],[f1591,f705])).
fof(f1591,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,X71)) | p(s(bool,X72)) | ~p(s(bool,X73)) | ~sP10(X73,X72,'T_0') | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(superposition,[],[f475,f657])).
fof(f1812,plain,(
  ( ! [X2,X1] : (~sP9('T_0',X2,X1) | p(s(bool,X2)) | p(s(bool,X1))) ) | $spl26),
  inference(cnf_transformation,[],[f1812_D])).
fof(f1812_D,plain,(
  ( ! [X2,X1] : (~sP9('T_0',X2,X1) | p(s(bool,X2)) | p(s(bool,X1))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f1894,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f1893])).
fof(f1893,plain,(
  ( ! [X70,X68,X69,X67] : (p(s(bool,X68)) | p(s(bool,X69)) | ~sP9('T_0',X69,X68) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f1590,f705])).
fof(f1590,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,X67)) | p(s(bool,X68)) | p(s(bool,X69)) | ~sP9('T_0',X69,X68) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f465,f657])).
fof(f1809,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~p(s(bool,X1))) ) | $spl24),
  inference(cnf_transformation,[],[f1809_D])).
fof(f1809_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~p(s(bool,X1))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f1892,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f1891])).
fof(f1891,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X64)) | sP9('T_0',X64,X65) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(subsumption_resolution,[],[f1589,f705])).
fof(f1589,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X63)) | ~p(s(bool,X64)) | sP9('T_0',X64,X65) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f464,f657])).
fof(f1806,plain,(
  ( ! [X2,X1] : (sP9('T_0',X2,X1) | ~p(s(bool,X1))) ) | $spl22),
  inference(cnf_transformation,[],[f1806_D])).
fof(f1806_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X2,X1) | ~p(s(bool,X1))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f1890,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f1889])).
fof(f1889,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X60)) | sP9('T_0',X61,X60) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(subsumption_resolution,[],[f1588,f705])).
fof(f1588,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X59)) | ~p(s(bool,X60)) | sP9('T_0',X61,X60) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f463,f657])).
fof(f1803,plain,(
  ( ! [X2,X1] : (~sP8('T_0',X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) ) | $spl20),
  inference(cnf_transformation,[],[f1803_D])).
fof(f1803_D,plain,(
  ( ! [X2,X1] : (~sP8('T_0',X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f1888,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP8('T_0',X1,X2)) )),
  inference(condensation,[],[f1887])).
fof(f1887,plain,(
  ( ! [X57,X58,X56,X55] : (p(s(bool,X56)) | p(s(bool,X57)) | ~sP8('T_0',X56,X57) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(subsumption_resolution,[],[f1587,f705])).
fof(f1587,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X55)) | p(s(bool,X56)) | p(s(bool,X57)) | ~sP8('T_0',X56,X57) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f458,f657])).
fof(f1800,plain,(
  ( ! [X2,X1] : (~sP7(X2,X1,'T_0') | p(s(bool,X2)) | ~p(s(bool,X1))) ) | $spl18),
  inference(cnf_transformation,[],[f1800_D])).
fof(f1800_D,plain,(
  ( ! [X2,X1] : (~sP7(X2,X1,'T_0') | p(s(bool,X2)) | ~p(s(bool,X1))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f1886,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,X1,'T_0')) )),
  inference(condensation,[],[f1885])).
fof(f1885,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X52)) | p(s(bool,X53)) | ~sP7(X53,X52,'T_0') | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(subsumption_resolution,[],[f1586,f705])).
fof(f1586,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X51)) | ~p(s(bool,X52)) | p(s(bool,X53)) | ~sP7(X53,X52,'T_0') | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(superposition,[],[f450,f657])).
fof(f1797,plain,(
  ( ! [X2,X1] : (~sP6(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl16),
  inference(cnf_transformation,[],[f1797_D])).
fof(f1797_D,plain,(
  ( ! [X2,X1] : (~sP6(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f1884,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X1,'T_0',X2)) )),
  inference(condensation,[],[f1883])).
fof(f1883,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X48)) | p(s(bool,X49)) | ~sP6(X48,'T_0',X49) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f1585,f705])).
fof(f1585,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X47)) | ~p(s(bool,X48)) | p(s(bool,X49)) | ~sP6(X48,'T_0',X49) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f441,f657])).
fof(f1794,plain,(
  ( ! [X2,X1] : (~sP3(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl14),
  inference(cnf_transformation,[],[f1794_D])).
fof(f1794_D,plain,(
  ( ! [X2,X1] : (~sP3(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f1882,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP3(X1,'T_0',X2)) )),
  inference(condensation,[],[f1881])).
fof(f1881,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X44)) | p(s(bool,X45)) | ~sP3(X44,'T_0',X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f1584,f705])).
fof(f1584,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X43)) | ~p(s(bool,X44)) | p(s(bool,X45)) | ~sP3(X44,'T_0',X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f422,f657])).
fof(f1791,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl12),
  inference(cnf_transformation,[],[f1791_D])).
fof(f1791_D,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f1880,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X1,'T_0',X2)) )),
  inference(condensation,[],[f1879])).
fof(f1879,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X40)) | p(s(bool,X41)) | ~sP4(X40,'T_0',X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(subsumption_resolution,[],[f1583,f705])).
fof(f1583,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X39)) | ~p(s(bool,X40)) | p(s(bool,X41)) | ~sP4(X40,'T_0',X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f418,f657])).
fof(f1788,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~sP3(X2,'T_0',X1) | sP5(X1,X2,'T_0')) ) | $spl10),
  inference(cnf_transformation,[],[f1788_D])).
fof(f1788_D,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~sP3(X2,'T_0',X1) | sP5(X1,X2,'T_0')) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f1878,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP5(X1,X2,'T_0') | ~sP4(X1,'T_0',X2) | ~sP3(X2,'T_0',X1)) )),
  inference(subsumption_resolution,[],[f1877,f519])).
fof(f1877,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP5(X1,X2,'T_0') | ~sP4(X1,'T_0',X2) | ~sP3(X2,'T_0',X1) | ~sP2('T_0',X2,X1)) )),
  inference(condensation,[],[f1876])).
fof(f1876,plain,(
  ( ! [X37,X35,X38,X36] : (sP5(X36,X37,'T_0') | ~sP4(X36,'T_0',X37) | ~sP3(X37,'T_0',X36) | ~sP2('T_0',X37,X36) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(subsumption_resolution,[],[f1582,f705])).
fof(f1582,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X35)) | sP5(X36,X37,'T_0') | ~sP4(X36,'T_0',X37) | ~sP3(X37,'T_0',X36) | ~sP2('T_0',X37,X36) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f415,f657])).
fof(f1784,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) | $spl8),
  inference(cnf_transformation,[],[f1784_D])).
fof(f1784_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f1875,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f1874])).
fof(f1874,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,X33)) | sP1(X33,'T_0') | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f1581,f705])).
fof(f1581,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,X32)) | p(s(bool,X33)) | sP1(X33,'T_0') | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f388,f657])).
fof(f1780,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) | $spl4),
  inference(cnf_transformation,[],[f1780_D])).
fof(f1780_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f1873,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f1872])).
fof(f1872,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X30)) | sP1('T_0',X30) | s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(subsumption_resolution,[],[f1580,f705])).
fof(f1580,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,X29)) | p(s(bool,X30)) | sP1('T_0',X30) | s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(superposition,[],[f385,f657])).
fof(f1838,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP6('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X2)))),X1)) )),
  inference(condensation,[],[f1837])).
fof(f1837,plain,(
  ( ! [X111,X112,X110,X113] : (p(s(bool,X111)) | ~sP6('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X112)))),X111) | s(bool,X110) = s(bool,X113) | s(bool,'T_0') = s(bool,X113)) )),
  inference(subsumption_resolution,[],[f1522,f705])).
fof(f1522,plain,(
  ( ! [X111,X112,X110,X113] : (~p(s(bool,X110)) | p(s(bool,X111)) | ~sP6('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X112)))),X111) | s(bool,X110) = s(bool,X113) | s(bool,'T_0') = s(bool,X113)) )),
  inference(superposition,[],[f863,f657])).
fof(f1835,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP3('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X2)))),X1)) )),
  inference(condensation,[],[f1834])).
fof(f1834,plain,(
  ( ! [X109,X107,X108,X106] : (p(s(bool,X107)) | ~sP3('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X108)))),X107) | s(bool,X106) = s(bool,X109) | s(bool,'T_0') = s(bool,X109)) )),
  inference(subsumption_resolution,[],[f1521,f705])).
fof(f1521,plain,(
  ( ! [X109,X107,X108,X106] : (~p(s(bool,X106)) | p(s(bool,X107)) | ~sP3('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X108)))),X107) | s(bool,X106) = s(bool,X109) | s(bool,'T_0') = s(bool,X109)) )),
  inference(superposition,[],[f820,f657])).
fof(f1832,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP4('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X2)))),X1)) )),
  inference(condensation,[],[f1831])).
fof(f1831,plain,(
  ( ! [X103,X105,X102,X104] : (p(s(bool,X103)) | ~sP4('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X104)))),X103) | s(bool,X102) = s(bool,X105) | s(bool,'T_0') = s(bool,X105)) )),
  inference(subsumption_resolution,[],[f1520,f705])).
fof(f1520,plain,(
  ( ! [X103,X105,X102,X104] : (~p(s(bool,X102)) | p(s(bool,X103)) | ~sP4('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_add_1'(s('type.rat.rat',X104)))),X103) | s(bool,X102) = s(bool,X105) | s(bool,'T_0') = s(bool,X105)) )),
  inference(superposition,[],[f803,f657])).
fof(f1829,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('T_0',X1,X2) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f1828])).
fof(f1828,plain,(
  ( ! [X94,X92,X93,X91] : (~sP8('T_0',X92,X93) | sP9('T_0',X92,X93) | s(bool,X91) = s(bool,X94) | s(bool,'T_0') = s(bool,X94)) )),
  inference(subsumption_resolution,[],[f1517,f705])).
fof(f1517,plain,(
  ( ! [X94,X92,X93,X91] : (~p(s(bool,X91)) | ~sP8('T_0',X92,X93) | sP9('T_0',X92,X93) | s(bool,X91) = s(bool,X94) | s(bool,'T_0') = s(bool,X94)) )),
  inference(superposition,[],[f486,f657])).
fof(f1826,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP10('T_0',X1,X2) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f1825])).
fof(f1825,plain,(
  ( ! [X90,X88,X87,X89] : (~sP10('T_0',X88,X89) | sP11('T_0',X88,X89) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(subsumption_resolution,[],[f1516,f705])).
fof(f1516,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,X87)) | ~sP10('T_0',X88,X89) | sP11('T_0',X88,X89) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f485,f657])).
fof(f1823,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f1822])).
fof(f1822,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,X84)) | p(s(bool,X85)) | ~sP11('T_0',X85,X84) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(subsumption_resolution,[],[f1515,f705])).
fof(f1515,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,X83)) | ~p(s(bool,X84)) | p(s(bool,X85)) | ~sP11('T_0',X85,X84) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f482,f657])).
fof(f1820,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f1819])).
fof(f1819,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,X80)) | sP11('T_0',X80,X81) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(subsumption_resolution,[],[f1514,f705])).
fof(f1514,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,X79)) | ~p(s(bool,X80)) | sP11('T_0',X80,X81) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(superposition,[],[f481,f657])).
fof(f1817,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f1816])).
fof(f1816,plain,(
  ( ! [X78,X76,X77,X75] : (p(s(bool,X76)) | sP11('T_0',X77,X76) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(subsumption_resolution,[],[f1513,f705])).
fof(f1513,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,X75)) | p(s(bool,X76)) | sP11('T_0',X77,X76) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f480,f657])).
fof(f1814,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP10(X2,X1,'T_0')) )),
  inference(condensation,[],[f1813])).
fof(f1813,plain,(
  ( ! [X74,X72,X71,X73] : (p(s(bool,X72)) | ~p(s(bool,X73)) | ~sP10(X73,X72,'T_0') | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(subsumption_resolution,[],[f1512,f705])).
fof(f1512,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,X71)) | p(s(bool,X72)) | ~p(s(bool,X73)) | ~sP10(X73,X72,'T_0') | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(superposition,[],[f475,f657])).
fof(f1811,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f1810])).
fof(f1810,plain,(
  ( ! [X70,X68,X69,X67] : (p(s(bool,X68)) | p(s(bool,X69)) | ~sP9('T_0',X69,X68) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f1511,f705])).
fof(f1511,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,X67)) | p(s(bool,X68)) | p(s(bool,X69)) | ~sP9('T_0',X69,X68) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f465,f657])).
fof(f1808,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f1807])).
fof(f1807,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X64)) | sP9('T_0',X64,X65) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(subsumption_resolution,[],[f1510,f705])).
fof(f1510,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X63)) | ~p(s(bool,X64)) | sP9('T_0',X64,X65) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f464,f657])).
fof(f1805,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f1804])).
fof(f1804,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X60)) | sP9('T_0',X61,X60) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(subsumption_resolution,[],[f1509,f705])).
fof(f1509,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X59)) | ~p(s(bool,X60)) | sP9('T_0',X61,X60) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f463,f657])).
fof(f1802,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP8('T_0',X1,X2)) )),
  inference(condensation,[],[f1801])).
fof(f1801,plain,(
  ( ! [X57,X58,X56,X55] : (p(s(bool,X56)) | p(s(bool,X57)) | ~sP8('T_0',X56,X57) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(subsumption_resolution,[],[f1508,f705])).
fof(f1508,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X55)) | p(s(bool,X56)) | p(s(bool,X57)) | ~sP8('T_0',X56,X57) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f458,f657])).
fof(f1799,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,X1,'T_0')) )),
  inference(condensation,[],[f1798])).
fof(f1798,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X52)) | p(s(bool,X53)) | ~sP7(X53,X52,'T_0') | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(subsumption_resolution,[],[f1507,f705])).
fof(f1507,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X51)) | ~p(s(bool,X52)) | p(s(bool,X53)) | ~sP7(X53,X52,'T_0') | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(superposition,[],[f450,f657])).
fof(f1796,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X1,'T_0',X2)) )),
  inference(condensation,[],[f1795])).
fof(f1795,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X48)) | p(s(bool,X49)) | ~sP6(X48,'T_0',X49) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f1506,f705])).
fof(f1506,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X47)) | ~p(s(bool,X48)) | p(s(bool,X49)) | ~sP6(X48,'T_0',X49) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f441,f657])).
fof(f1793,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP3(X1,'T_0',X2)) )),
  inference(condensation,[],[f1792])).
fof(f1792,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X44)) | p(s(bool,X45)) | ~sP3(X44,'T_0',X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f1505,f705])).
fof(f1505,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X43)) | ~p(s(bool,X44)) | p(s(bool,X45)) | ~sP3(X44,'T_0',X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f422,f657])).
fof(f1790,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X1,'T_0',X2)) )),
  inference(condensation,[],[f1789])).
fof(f1789,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X40)) | p(s(bool,X41)) | ~sP4(X40,'T_0',X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(subsumption_resolution,[],[f1504,f705])).
fof(f1504,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X39)) | ~p(s(bool,X40)) | p(s(bool,X41)) | ~sP4(X40,'T_0',X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f418,f657])).
fof(f1787,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP5(X1,X2,'T_0') | ~sP4(X1,'T_0',X2) | ~sP3(X2,'T_0',X1)) )),
  inference(subsumption_resolution,[],[f1786,f519])).
fof(f1786,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP5(X1,X2,'T_0') | ~sP4(X1,'T_0',X2) | ~sP3(X2,'T_0',X1) | ~sP2('T_0',X2,X1)) )),
  inference(condensation,[],[f1785])).
fof(f1785,plain,(
  ( ! [X37,X35,X38,X36] : (sP5(X36,X37,'T_0') | ~sP4(X36,'T_0',X37) | ~sP3(X37,'T_0',X36) | ~sP2('T_0',X37,X36) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(subsumption_resolution,[],[f1503,f705])).
fof(f1503,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X35)) | sP5(X36,X37,'T_0') | ~sP4(X36,'T_0',X37) | ~sP3(X37,'T_0',X36) | ~sP2('T_0',X37,X36) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f415,f657])).
fof(f1783,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f1782])).
fof(f1782,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,X33)) | sP1(X33,'T_0') | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f1502,f705])).
fof(f1502,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,X32)) | p(s(bool,X33)) | sP1(X33,'T_0') | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f388,f657])).
fof(f1779,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f1778])).
fof(f1778,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X30)) | sP1('T_0',X30) | s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(subsumption_resolution,[],[f1501,f705])).
fof(f1501,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,X29)) | p(s(bool,X30)) | sP1('T_0',X30) | s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(superposition,[],[f385,f657])).
fof(f506,plain,(
  $false | ($spl1 | $spl2)),
  inference(trivial_inequality_removal,[],[f505])).
fof(f505,plain,(
  s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',sK12),s('type.rat.rat',sK14))) != s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',sK12),s('type.rat.rat',sK14))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f504,f499])).
fof(f504,plain,(
  s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',sK12),s('type.rat.rat',sK14))) != s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',sK13),s('type.rat.rat',sK14))) | $spl1),
  inference(forward_demodulation,[],[f503,f358])).
fof(f503,plain,(
  s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',sK13),s('type.rat.rat',sK14))) != s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',sK14),s('type.rat.rat',sK12))) | $spl1),
  inference(forward_demodulation,[],[f498,f358])).
fof(f497,plain,(
  s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',sK14),s('type.rat.rat',sK12))) = s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',sK14),s('type.rat.rat',sK13))) | $spl0),
  inference(cnf_transformation,[],[f497_D])).
fof(f497_D,plain,(
  s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',sK14),s('type.rat.rat',sK12))) = s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',sK14),s('type.rat.rat',sK13))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f499,plain,(
  s('type.rat.rat',sK12) = s('type.rat.rat',sK13) | $spl2),
  inference(cnf_transformation,[],[f499_D])).
fof(f499_D,plain,(
  s('type.rat.rat',sK12) = s('type.rat.rat',sK13) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f275,plain,(
  s('type.rat.rat',sK12) = s('type.rat.rat',sK13) | s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',sK14),s('type.rat.rat',sK12))) = s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',sK14),s('type.rat.rat',sK13)))),
  inference(cnf_transformation,[],[f175])).
fof(f175,plain,(
  (s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',sK14),s('type.rat.rat',sK12))) = s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',sK14),s('type.rat.rat',sK13))) | s('type.rat.rat',sK12) = s('type.rat.rat',sK13)) & (s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',sK14),s('type.rat.rat',sK12))) != s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',sK14),s('type.rat.rat',sK13))) | s('type.rat.rat',sK12) != s('type.rat.rat',sK13))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13,sK14])],[f174])).
fof(f174,plain,(
  ? [X0,X1,X2] : ((s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',X2),s('type.rat.rat',X0))) = s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',X2),s('type.rat.rat',X1))) | s('type.rat.rat',X0) = s('type.rat.rat',X1)) & (s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',X2),s('type.rat.rat',X0))) != s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',X2),s('type.rat.rat',X1))) | s('type.rat.rat',X0) != s('type.rat.rat',X1)))),
  inference(nnf_transformation,[],[f130])).
fof(f130,plain,(
  ? [X0,X1,X2] : (s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',X2),s('type.rat.rat',X0))) = s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',X2),s('type.rat.rat',X1))) <~> s('type.rat.rat',X0) = s('type.rat.rat',X1))),
  inference(ennf_transformation,[],[f56])).
fof(f56,plain,(
  ~! [X0,X1,X2] : (s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',X2),s('type.rat.rat',X0))) = s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',X2),s('type.rat.rat',X1))) <=> s('type.rat.rat',X0) = s('type.rat.rat',X1))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : (s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',X8),s('type.rat.rat',X6))) = s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',X8),s('type.rat.rat',X7))) <=> s('type.rat.rat',X6) = s('type.rat.rat',X7))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : (s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',X8),s('type.rat.rat',X6))) = s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',X8),s('type.rat.rat',X7))) <=> s('type.rat.rat',X6) = s('type.rat.rat',X7))),
  file('thm.rat.RAT_EQ_LADD','thm.rat.RAT_EQ_LADD')).
fof(f498,plain,(
  s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',sK14),s('type.rat.rat',sK12))) != s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',sK14),s('type.rat.rat',sK13))) | $spl1),
  inference(cnf_transformation,[],[f498_D])).
fof(f498_D,plain,(
  s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',sK14),s('type.rat.rat',sK12))) != s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',sK14),s('type.rat.rat',sK13))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f500,plain,(
  s('type.rat.rat',sK12) != s('type.rat.rat',sK13) | $spl3),
  inference(cnf_transformation,[],[f500_D])).
fof(f500_D,plain,(
  s('type.rat.rat',sK12) != s('type.rat.rat',sK13) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f276,plain,(
  s('type.rat.rat',sK12) != s('type.rat.rat',sK13) | s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',sK14),s('type.rat.rat',sK12))) != s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',sK14),s('type.rat.rat',sK13)))),
  inference(cnf_transformation,[],[f175])).
% SZS output end Proof for thm.rat
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 349
% Passive clauses: 782
% Generated clauses: 8380
% Final active clauses: 338
% Final passive clauses: 349
% Input formulas: 54
% Initial clauses: 210
% 
% Duplicate literals: 868
% Trivial inequalities: 43
% Fw subsumption resolutions: 521
% Fw demodulations: 241
% Bw demodulations: 29
% Condensations: 1213
% Global subsumptions: 4
% 
% Simple tautologies: 118
% Equational tautologies: 164
% Forward subsumptions: 4126
% Bw demodulations to eq. taut.: 1
% 
% Unit resulting resolution: 2263
% Factoring: 1
% Forward superposition: 961
% Backward superposition: 1775
% Self superposition: 113
% Equality factoring: 106
% Equality resolution: 4
% 
% Split clauses: 74
% Split components: 153
% SAT solver clauses: 1202
% SAT solver unit clauses: 358
% SAT solver binary clauses: 269
% 
% Sat splits: 74
% Sat splitting refutations: 266
% 
% Memory used [KB]: 5117
% Time elapsed: 0.508 s
% ------------------------------
% ------------------------------
% Success in time 1.847 s
