Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.relation
% (114921)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 909
% Passive clauses: 7638
% Generated clauses: 23878
% Final active clauses: 886
% Final passive clauses: 6180
% Input formulas: 48
% Initial clauses: 197
% 
% Duplicate literals: 519
% Trivial inequalities: 75
% Fw subsumption resolutions: 2544
% Fw demodulations: 1491
% Bw demodulations: 9
% 
% Simple tautologies: 295
% Equational tautologies: 57
% Forward subsumptions: 10475
% 
% Binary resolution: 11203
% Factoring: 1499
% Forward superposition: 2000
% Backward superposition: 3981
% Self superposition: 28
% Equality factoring: 28
% Equality resolution: 15
% 
% Split clauses: 675
% Split components: 1353
% SAT solver clauses: 1102
% SAT solver unit clauses: 8
% SAT solver binary clauses: 172
% SAT solver learnt clauses: 1
% SAT solver learnt literals: 1
% 
% Sat splits: 675
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 1101
% TWLsolver calls for satisfiability: 573
% 
% Memory used [KB]: 10490
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.relation
% (114943)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 449
% Passive clauses: 2933
% Generated clauses: 13543
% Final active clauses: 437
% Final passive clauses: 2473
% Input formulas: 48
% Initial clauses: 197
% 
% Duplicate literals: 279
% Trivial inequalities: 66
% Fw subsumption resolutions: 999
% Fw demodulations: 864
% Bw demodulations: 22
% Condensations: 348
% Global subsumptions: 13
% 
% Simple tautologies: 312
% Equational tautologies: 19
% Forward subsumptions: 4181
% 
% Binary resolution: 3627
% Factoring: 86
% Forward superposition: 1995
% Backward superposition: 2041
% Self superposition: 31
% Equality factoring: 30
% Equality resolution: 9
% 
% Split clauses: 102
% Split components: 208
% SAT solver clauses: 3673
% SAT solver unit clauses: 186
% SAT solver binary clauses: 74
% SAT solver learnt clauses: 26
% SAT solver learnt literals: 10
% 
% Sat splits: 102
% 
% TWLsolver clauses: 3621
% TWLsolver calls for satisfiability: 1848
% 
% Memory used [KB]: 15351
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.relation
% (114966)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 3725
% Passive clauses: 19114
% Generated clauses: 98339
% Final active clauses: 2995
% Final passive clauses: 7587
% Input formulas: 48
% Initial clauses: 197
% Discarded non-redundant clauses: 3856
% 
% Duplicate literals: 1517
% Trivial inequalities: 50
% Fw subsumption resolutions: 1910
% Fw demodulations: 20986
% Bw demodulations: 2478
% Condensations: 837
% Global subsumptions: 29
% 
% Simple tautologies: 105
% Equational tautologies: 538
% Forward subsumptions: 51801
% Bw demodulations to eq. taut.: 3
% 
% Unit resulting resolution: 52077
% Factoring: 1
% Forward superposition: 10512
% Backward superposition: 7236
% Self superposition: 35
% Equality factoring: 34
% Equality resolution: 1
% 
% Split clauses: 200
% Split components: 407
% SAT solver clauses: 20235
% SAT solver unit clauses: 589
% SAT solver binary clauses: 5478
% 
% Sat splits: 200
% Sat splitting refutations: 689
% 
% Memory used [KB]: 57824
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.relation
% (115191)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 2601
% Passive clauses: 51117
% Generated clauses: 207532
% Final active clauses: 2484
% Final passive clauses: 46203
% Input formulas: 48
% Initial clauses: 194
% 
% Selected by SInE selection: 45
% SInE iterations: 1
% 
% Duplicate literals: 6904
% Trivial inequalities: 377
% Fw subsumption resolutions: 20872
% Fw demodulations: 1236
% Bw demodulations: 47
% 
% Simple tautologies: 4483
% Equational tautologies: 751
% Forward subsumptions: 115363
% 
% Binary resolution: 164857
% Factoring: 896
% Forward superposition: 8201
% Backward superposition: 2847
% Self superposition: 5
% Equality resolution: 37
% 
% Split clauses: 7728
% Split components: 18260
% SAT solver clauses: 8357
% SAT solver unit clauses: 31
% SAT solver binary clauses: 1607
% 
% Sat splits: 7728
% Sat splitting refutations: 51
% 
% Memory used [KB]: 52962
% Time elapsed: 4.0000 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs+1002_2:3_br=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:sas=minisat:stl=30:sac=on:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none:sp=reverse_arity:urr=on_4 on thm.relation
% (115330)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 599
% Passive clauses: 4288
% Generated clauses: 15939
% Final active clauses: 585
% Final passive clauses: 3547
% Input formulas: 48
% Initial clauses: 197
% 
% Duplicate literals: 364
% Trivial inequalities: 43
% Fw subsumption resolutions: 872
% Fw demodulations: 1286
% Bw demodulations: 106
% Global subsumptions: 5
% 
% Simple tautologies: 396
% Equational tautologies: 60
% Forward subsumptions: 8540
% 
% Unit resulting resolution: 5088
% Factoring: 2
% Forward superposition: 3312
% Backward superposition: 4523
% Self superposition: 67
% Equality factoring: 6
% Equality resolution: 8
% 
% Split clauses: 13
% Split components: 26
% SAT solver clauses: 4364
% SAT solver unit clauses: 244
% SAT solver binary clauses: 995
% 
% Sat splits: 13
% Sat splitting refutations: 56
% 
% Memory used [KB]: 11641
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 15
lrs+1011_12_bs=on:bsr=unit_only:cond=on:gs=on:gsaa=from_current:gsssp=full:nwc=1.1:sas=minisat:stl=60:sos=all:sac=on:sdd=off:sser=off:sfr=on:ssfp=100000:ssfq=1.2:ssnc=none:sp=reverse_arity:updr=off_14 on thm.relation
% (115355)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1494
% Passive clauses: 8034
% Generated clauses: 23779
% Final active clauses: 636
% Final passive clauses: 3439
% Input formulas: 48
% Initial clauses: 197
% 
% Duplicate literals: 860
% Trivial inequalities: 84
% Fw subsumption resolutions: 1110
% Bw subsumption resolutions: 1685
% Fw demodulations: 1070
% Bw demodulations: 180
% Condensations: 1496
% Global subsumptions: 150
% 
% Simple tautologies: 1266
% Equational tautologies: 35
% Forward subsumptions: 9342
% Backward subsumptions: 1844
% 
% Binary resolution: 7260
% Factoring: 69
% Forward superposition: 2993
% Backward superposition: 3248
% Self superposition: 68
% Equality factoring: 4
% Equality resolution: 12
% 
% Split clauses: 23
% Split components: 46
% SAT solver clauses: 8532
% SAT solver unit clauses: 220
% SAT solver binary clauses: 1418
% 
% Sat splits: 23
% Sat splitting refutations: 322
% 
% Memory used [KB]: 13816
% Time elapsed: 1.600 s
% ------------------------------
% ------------------------------
% remaining time: 120 next slice time: 9
lrs+10_8:1_bd=preordered:bs=on:ccuc=first:fsr=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:nicw=on:sas=minisat:stl=120:sos=on:sscc=on:sser=off:sfr=on:ssfp=4000:ssfq=1.0:ssnc=none:sp=reverse_arity:urr=on_7 on thm.relation
% (115491)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1664
% Passive clauses: 3983
% Generated clauses: 26110
% Final active clauses: 681
% Final passive clauses: 1208
% Input formulas: 48
% Initial clauses: 197
% 
% Duplicate literals: 389
% Trivial inequalities: 29
% Fw demodulations: 4124
% Bw demodulations: 51
% Global subsumptions: 318
% 
% Simple tautologies: 995
% Equational tautologies: 72
% Forward subsumptions: 10625
% Backward subsumptions: 113
% 
% Binary resolution: 5009
% Unit resulting resolution: 11693
% Forward superposition: 134
% Backward superposition: 1169
% 
% Split clauses: 221
% Split components: 469
% SAT solver clauses: 15075
% SAT solver unit clauses: 238
% SAT solver binary clauses: 12085
% 
% Sat splits: 221
% Sat splitting refutations: 319
% 
% Memory used [KB]: 15991
% Time elapsed: 1.0000 s
% ------------------------------
% ------------------------------
% remaining time: 110 next slice time: 85
lrs+1_5:4_cond=on:fsr=off:fde=none:gs=on:gsem=on:lwlo=on:nm=64:nwc=1:stl=60:sos=all:spl=off_80 on thm.relation
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.relation
% SZS output start Proof for thm.relation
fof(f15039,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f14965])).
fof(f14965,plain,(
  s(bool,'F_0') != s(bool,'F_0')),
  inference(resolution,[],[f14757,f254])).
fof(f254,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f167])).
fof(f167,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f68])).
fof(f68,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f67])).
fof(f67,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X13] : (s(bool,'F_0') = s(bool,X13) <=> ~p(s(bool,X13)))),
  file('thm.relation.TC_CASES|49|','thm.bool.EQ_CLAUSES|split|3')).
fof(f14757,plain,(
  p(s(bool,'F_0'))),
  inference(resolution,[],[f14756,f232])).
fof(f232,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'T_0'))),
  file('thm.relation.TC_CASES|49|','thm.bool.TRUTH')).
fof(f14756,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(forward_demodulation,[],[f14717,f7911])).
fof(f7911,plain,(
  s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK4),s(sK1,sK3),s(sK1,sK2)))),
  inference(resolution,[],[f7903,f253])).
fof(f253,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f166])).
fof(f166,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f65])).
fof(f65,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X13] : (s(bool,'F_0') = s(bool,X13) <=> ~p(s(bool,X13)))),
  file('thm.relation.TC_CASES|49|','thm.bool.EQ_CLAUSES|split|2')).
fof(f7903,plain,(
  ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK4),s(sK1,sK3),s(sK1,sK2))))),
  inference(global_subsumption,[],[f7901,f230])).
fof(f230,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK4),s(sK1,sK3))),s(sK1,sK2)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK4),s(sK1,sK3),s(sK1,sK2))))),
  inference(cnf_transformation,[],[f161])).
fof(f161,plain,(
  (p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK4),s(sK1,sK3),s(sK1,sK2)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK4),s(sK1,sK3))),s(sK1,sK2)))) | (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK4),s(sK1,sK3))),s(sK1,sK5)))) & p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK4),s(sK1,sK5),s(sK1,sK2)))))) & (~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK4),s(sK1,sK3),s(sK1,sK2)))) | (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK4),s(sK1,sK3))),s(sK1,sK2)))) & ! [X4] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK4),s(sK1,sK3))),s(sK1,X4)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK4),s(sK1,X4),s(sK1,sK2)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1,sK2,sK3,sK4,sK5])],[f160])).
fof(f160,plain,(
  ? [X0,X1,X2,X3] : ((p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X3),s(X0,X2),s(X0,X1)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,X2))),s(X0,X1)))) | ? [X5] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,X2))),s(X0,X5)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X3),s(X0,X5),s(X0,X1)))))) & (~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X3),s(X0,X2),s(X0,X1)))) | (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,X2))),s(X0,X1)))) & ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,X2))),s(X0,X4)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X3),s(X0,X4),s(X0,X1))))))))),
  inference(rectify,[],[f159])).
fof(f159,plain,(
  ? [X0,X1,X2,X3] : ((p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X3),s(X0,X2),s(X0,X1)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,X2))),s(X0,X1)))) | ? [X4] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,X2))),s(X0,X4)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X3),s(X0,X4),s(X0,X1)))))) & (~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X3),s(X0,X2),s(X0,X1)))) | (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,X2))),s(X0,X1)))) & ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,X2))),s(X0,X4)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X3),s(X0,X4),s(X0,X1))))))))),
  inference(flattening,[],[f158])).
fof(f158,plain,(
  ? [X0,X1,X2,X3] : ((p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X3),s(X0,X2),s(X0,X1)))) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,X2))),s(X0,X1)))) | ? [X4] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,X2))),s(X0,X4)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X3),s(X0,X4),s(X0,X1))))))) & (~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X3),s(X0,X2),s(X0,X1)))) | (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,X2))),s(X0,X1)))) & ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,X2))),s(X0,X4)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X3),s(X0,X4),s(X0,X1))))))))),
  inference(nnf_transformation,[],[f127])).
fof(f127,plain,(
  ? [X0,X1,X2,X3] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X3),s(X0,X2),s(X0,X1)))) <~> (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,X2))),s(X0,X1)))) | ? [X4] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,X2))),s(X0,X4)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X3),s(X0,X4),s(X0,X1)))))))),
  inference(ennf_transformation,[],[f51])).
fof(f51,plain,(
  ~! [X0,X1,X2,X3] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X3),s(X0,X2),s(X0,X1)))) <=> (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,X2))),s(X0,X1)))) | ? [X4] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,X2))),s(X0,X4)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X3),s(X0,X4),s(X0,X1)))))))),
  inference(flattening,[],[f50])).
fof(f50,plain,(
  ~! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X3),s(X0,X2),s(X0,X1)))) <=> (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,X2))),s(X0,X1)))) | ? [X4] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,X2))),s(X0,X4)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X3),s(X0,X4),s(X0,X1)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : (p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X9),s(X6,X8),s(X6,X7)))) <=> (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X9),s(X6,X8))),s(X6,X7)))) | ? [X10] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X9),s(X6,X8))),s(X6,X10)))) & p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X9),s(X6,X10),s(X6,X7)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : (p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X9),s(X6,X8),s(X6,X7)))) <=> (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X9),s(X6,X8))),s(X6,X7)))) | ? [X10] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X9),s(X6,X8))),s(X6,X10)))) & p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X9),s(X6,X10),s(X6,X7)))))))),
  file('thm.relation.TC_CASES|49|','thm.relation.TC_CASES|49|')).
fof(f7901,plain,(
  ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK4),s(sK1,sK3),s(sK1,sK2)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK4),s(sK1,sK3))),s(sK1,sK2))))),
  inference(duplicate_literal_removal,[],[f7816])).
fof(f7816,plain,(
  ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK4),s(sK1,sK3),s(sK1,sK2)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK4),s(sK1,sK3))),s(sK1,sK2)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK4),s(sK1,sK3),s(sK1,sK2)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK4),s(sK1,sK3))),s(sK1,sK2)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK4),s(sK1,sK3),s(sK1,sK2))))),
  inference(resolution,[],[f1043,f298])).
fof(f298,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,sK15(X3,X2,X1,X0)),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f199])).
fof(f199,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,sK15(X3,X2,X1,X0))))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,sK15(X3,X2,X1,X0)),s(X0,X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK15])],[f142])).
fof(f142,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | ? [X4] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X4)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X3))))))),
  inference(flattening,[],[f141])).
fof(f141,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | ? [X4] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X4)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X3)))))))),
  inference(ennf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) => (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | ? [X4] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X4)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X3)))))))),
  inference(flattening,[],[f96])).
fof(f96,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) => (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | ? [X4] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X4)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X3)))))))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X6] : ! [X9,X8,X7] : (p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X9),s(X6,X8),s(X6,X7)))) => (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X9),s(X6,X8))),s(X6,X7)))) | ? [X10] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X9),s(X6,X8))),s(X6,X10)))) & p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X9),s(X6,X10),s(X6,X7)))))))),
  file('thm.relation.TC_CASES|49|','thm.relation.TC_CASES|49|_E')).
fof(f1043,plain,(
  ( ! [X2] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK4),s(sK1,sK15(X2,sK3,sK4,sK1)),s(sK1,sK2)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK4),s(sK1,sK3),s(sK1,sK2)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK4),s(sK1,sK3))),s(sK1,X2)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK4),s(sK1,sK3),s(sK1,X2))))) )),
  inference(resolution,[],[f231,f297])).
fof(f297,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,sK15(X3,X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f199])).
fof(f231,plain,(
  ( ! [X4] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK4),s(sK1,sK3))),s(sK1,X4)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK4),s(sK1,X4),s(sK1,sK2)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK4),s(sK1,sK3),s(sK1,sK2))))) )),
  inference(cnf_transformation,[],[f161])).
fof(f14717,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK4),s(sK1,sK3),s(sK1,sK2))))),
  inference(superposition,[],[f8780,f8985])).
fof(f8985,plain,(
  s(bool,'T_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK4),s(sK1,sK3),s(sK1,sK5)))),
  inference(resolution,[],[f8792,f249])).
fof(f249,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f164])).
fof(f164,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f63])).
fof(f63,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X13] : (s(bool,'T_0') = s(bool,X13) <=> p(s(bool,X13)))),
  file('thm.relation.TC_CASES|49|','thm.bool.EQ_CLAUSES|split|1')).
fof(f8792,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK4),s(sK1,sK3),s(sK1,sK5))))),
  inference(global_subsumption,[],[f8069,f2378,f7903])).
fof(f2378,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK4),s(sK1,sK3),s(sK1,sK2)))) | p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK4),s(sK1,sK3),s(sK1,sK5))))),
  inference(duplicate_literal_removal,[],[f2307])).
fof(f2307,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK4),s(sK1,sK3),s(sK1,sK2)))) | p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK4),s(sK1,sK3),s(sK1,sK5)))) | p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK4),s(sK1,sK3),s(sK1,sK2))))),
  inference(resolution,[],[f1170,f299])).
fof(f299,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))))),
  inference(ennf_transformation,[],[f99])).
fof(f99,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) => p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))))),
  inference(flattening,[],[f98])).
fof(f98,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) => p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X6] : ! [X9,X8,X10] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X9),s(X6,X8))),s(X6,X10)))) => p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X9),s(X6,X8),s(X6,X10)))))),
  file('thm.relation.TC_CASES|49|','thm.relation.TC_RULES|split|0')).
fof(f1170,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK4),s(sK1,sK3))),s(sK1,sK2)))) | p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK4),s(sK1,sK3),s(sK1,sK2)))) | p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK4),s(sK1,sK3),s(sK1,sK5))))),
  inference(resolution,[],[f228,f299])).
fof(f228,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK4),s(sK1,sK3))),s(sK1,sK5)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK4),s(sK1,sK3))),s(sK1,sK2)))) | p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK4),s(sK1,sK3),s(sK1,sK2))))),
  inference(cnf_transformation,[],[f161])).
fof(f8069,plain,(
  p(s(bool,'F_0')) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK4),s(sK1,sK3))),s(sK1,sK2)))) | p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK4),s(sK1,sK3),s(sK1,sK5))))),
  inference(backward_demodulation,[],[f7911,f1170])).
fof(f8780,plain,(
  ( ! [X1] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK4),s(sK1,X1),s(sK1,sK5)))) | p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK4),s(sK1,X1),s(sK1,sK2))))) )),
  inference(global_subsumption,[],[f8028,f975,f7903])).
fof(f975,plain,(
  ( ! [X1] : (p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK4),s(sK1,sK3),s(sK1,sK2)))) | p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK4),s(sK1,X1),s(sK1,sK2)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK4),s(sK1,X1),s(sK1,sK5))))) )),
  inference(resolution,[],[f973,f300])).
fof(f300,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4)))))),
  inference(flattening,[],[f144])).
fof(f144,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4))))) | p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4)))))),
  inference(ennf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0,X1,X2,X3,X4] : ((p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4))))) => p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4)))))),
  inference(flattening,[],[f100])).
fof(f100,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : ((p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4))))) => p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X4)))))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X6] : ! [X9,X8,X10,X7] : ((p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X9),s(X6,X8),s(X6,X10)))) & p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X9),s(X6,X10),s(X6,X7))))) => p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X9),s(X6,X8),s(X6,X7)))))),
  file('thm.relation.TC_CASES|49|','thm.relation.TC_RULES|split|1')).
fof(f973,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK4),s(sK1,sK5),s(sK1,sK2)))) | p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK4),s(sK1,sK3),s(sK1,sK2))))),
  inference(duplicate_literal_removal,[],[f902])).
fof(f902,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK4),s(sK1,sK5),s(sK1,sK2)))) | p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK4),s(sK1,sK3),s(sK1,sK2)))) | p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK4),s(sK1,sK3),s(sK1,sK2))))),
  inference(resolution,[],[f229,f299])).
fof(f229,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK4),s(sK1,sK3))),s(sK1,sK2)))) | p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK4),s(sK1,sK5),s(sK1,sK2)))) | p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK4),s(sK1,sK3),s(sK1,sK2))))),
  inference(cnf_transformation,[],[f161])).
fof(f8028,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK4),s(sK1,X1),s(sK1,sK2)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK4),s(sK1,X1),s(sK1,sK5))))) )),
  inference(backward_demodulation,[],[f7911,f975])).
% SZS output end Proof for thm.relation
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 243
% Passive clauses: 1781
% Generated clauses: 14298
% Final active clauses: 162
% Final passive clauses: 857
% Input formulas: 48
% Initial clauses: 326
% 
% Duplicate literals: 1284
% Trivial inequalities: 193
% Fw demodulations: 685
% Bw demodulations: 791
% Condensations: 1356
% Global subsumptions: 159
% 
% Simple tautologies: 1258
% Equational tautologies: 32
% Forward subsumptions: 8031
% Bw demodulations to eq. taut.: 41
% 
% Binary resolution: 5625
% Factoring: 6
% Forward superposition: 3889
% Backward superposition: 575
% Equality resolution: 22
% 
% SAT solver clauses: 2656
% SAT solver unit clauses: 553
% SAT solver binary clauses: 280
% SAT solver learnt clauses: 372
% SAT solver learnt literals: 610
% 
% TWLsolver clauses: 1940
% TWLsolver calls for satisfiability: 49019
% 
% Memory used [KB]: 7931
% Time elapsed: 0.528 s
% ------------------------------
% ------------------------------
% Success in time 19.559 s
