Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.ringNorm
% (77896)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 115
% Passive clauses: 3648
% Generated clauses: 12986
% Final active clauses: 108
% Final passive clauses: 3257
% Input formulas: 69
% Initial clauses: 125
% 
% Duplicate literals: 498
% Trivial inequalities: 211
% Fw subsumption resolutions: 1423
% Fw demodulations: 438
% Bw demodulations: 5
% 
% Simple tautologies: 290
% Equational tautologies: 89
% Forward subsumptions: 5480
% 
% Binary resolution: 282
% Factoring: 2
% Forward superposition: 4143
% Backward superposition: 5350
% Self superposition: 286
% Equality factoring: 190
% Equality resolution: 34
% 
% Split clauses: 475
% Split components: 994
% SAT solver clauses: 1066
% SAT solver unit clauses: 42
% SAT solver binary clauses: 379
% 
% Sat splits: 475
% 
% TWLsolver clauses: 1036
% TWLsolver calls for satisfiability: 134
% 
% Memory used [KB]: 11129
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.ringNorm
% (77915)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 134
% Passive clauses: 1133
% Generated clauses: 12038
% Final active clauses: 125
% Final passive clauses: 975
% Input formulas: 69
% Initial clauses: 125
% 
% Duplicate literals: 176
% Trivial inequalities: 86
% Fw subsumption resolutions: 216
% Fw demodulations: 1276
% Bw demodulations: 31
% Condensations: 346
% Global subsumptions: 12
% 
% Simple tautologies: 90
% Equational tautologies: 41
% Forward subsumptions: 6715
% Fw demodulations to eq. taut.: 6
% 
% Binary resolution: 51
% Factoring: 1
% Forward superposition: 3820
% Backward superposition: 4746
% Self superposition: 183
% Equality factoring: 50
% Equality resolution: 21
% 
% Split clauses: 4
% Split components: 8
% SAT solver clauses: 2078
% SAT solver unit clauses: 158
% SAT solver binary clauses: 59
% SAT solver learnt clauses: 24
% 
% Sat splits: 4
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 2030
% TWLsolver calls for satisfiability: 422
% 
% Memory used [KB]: 25074
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.ringNorm
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.ringNorm
% SZS output start Proof for thm.ringNorm
fof(f48392,plain,(
  $false),
  inference(subsumption_resolution,[],[f48391,f325])).
fof(f325,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.ringNorm.polynom'(X0),'const.ringNorm.Pvar_1'(s('type.quote.index',X3))) != s('type.ringNorm.polynom'(X0),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X0),X2),s('type.ringNorm.polynom'(X0),X1)))) )),
  inference(cnf_transformation,[],[f141])).
fof(f141,plain,(
  ! [X0,X1,X2,X3] : s('type.ringNorm.polynom'(X0),'const.ringNorm.Pvar_1'(s('type.quote.index',X3))) != s('type.ringNorm.polynom'(X0),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X0),X2),s('type.ringNorm.polynom'(X0),X1)))),
  inference(flattening,[],[f140])).
fof(f140,plain,(
  ! [X0] : ! [X1,X2,X3] : ~s('type.ringNorm.polynom'(X0),'const.ringNorm.Pvar_1'(s('type.quote.index',X3))) = s('type.ringNorm.polynom'(X0),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X0),X2),s('type.ringNorm.polynom'(X0),X1)))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X6] : ! [X38,X39,X8] : ~s('type.ringNorm.polynom'(X6),'const.ringNorm.Pvar_1'(s('type.quote.index',X8))) = s('type.ringNorm.polynom'(X6),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X6),X39),s('type.ringNorm.polynom'(X6),X38)))),
  file('thm.ringNorm.polynom_distinct|split|3','thm.ringNorm.polynom_distinct|split|2')).
fof(f48391,plain,(
  ( ! [X81] : (s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pvar_1'(s('type.quote.index',sK2))) = s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(sK0),X81),s('type.ringNorm.polynom'(sK0),sK1)))) )),
  inference(forward_demodulation,[],[f47391,f48390])).
fof(f48390,plain,(
  ( ! [X79] : (s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pvar_1'(s('type.quote.index',sK2))) = s('type.ringNorm.polynom'(sK0),happ(s(fun('type.quote.index','type.ringNorm.polynom'(sK0)),X79),s('type.quote.index',sK2)))) )),
  inference(forward_demodulation,[],[f47389,f245])).
fof(f245,plain,(
  s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pvar_1'(s('type.quote.index',sK2))) = s('type.ringNorm.polynom'(sK0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(sK0),sK1)))),
  inference(cnf_transformation,[],[f201])).
fof(f201,plain,(
  s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pvar_1'(s('type.quote.index',sK2))) = s('type.ringNorm.polynom'(sK0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(sK0),sK1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f184])).
fof(f184,plain,(
  ? [X0,X1,X2] : s('type.ringNorm.polynom'(X0),'const.ringNorm.Pvar_1'(s('type.quote.index',X2))) = s('type.ringNorm.polynom'(X0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X0),X1)))),
  inference(ennf_transformation,[],[f72])).
fof(f72,plain,(
  ~! [X0,X1,X2] : s('type.ringNorm.polynom'(X0),'const.ringNorm.Pvar_1'(s('type.quote.index',X2))) != s('type.ringNorm.polynom'(X0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X0),X1)))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ~! [X0] : ! [X1,X2] : ~s('type.ringNorm.polynom'(X0),'const.ringNorm.Pvar_1'(s('type.quote.index',X2))) = s('type.ringNorm.polynom'(X0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X0),X1)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : ~s('type.ringNorm.polynom'(X6),'const.ringNorm.Pvar_1'(s('type.quote.index',X8))) = s('type.ringNorm.polynom'(X6),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X6),X7)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : ~s('type.ringNorm.polynom'(X6),'const.ringNorm.Pvar_1'(s('type.quote.index',X8))) = s('type.ringNorm.polynom'(X6),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X6),X7)))),
  file('thm.ringNorm.polynom_distinct|split|3','thm.ringNorm.polynom_distinct|split|3')).
fof(f47389,plain,(
  ( ! [X79] : (s('type.ringNorm.polynom'(sK0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(sK0),sK1))) = s('type.ringNorm.polynom'(sK0),happ(s(fun('type.quote.index','type.ringNorm.polynom'(sK0)),X79),s('type.quote.index',sK2)))) )),
  inference(superposition,[],[f23879,f298])).
fof(f298,plain,(
  ( ! [X0,X1] : (s('type.ringNorm.polynom'(X0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X0),X1))) = s('type.ringNorm.polynom'(X0),happ(s(fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0)),'const.ringNorm.Popp_0'),s('type.ringNorm.polynom'(X0),X1)))) )),
  inference(cnf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0,X1] : s('type.ringNorm.polynom'(X0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X0),X1))) = s('type.ringNorm.polynom'(X0),happ(s(fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0)),'const.ringNorm.Popp_0'),s('type.ringNorm.polynom'(X0),X1)))),
  inference(flattening,[],[f101])).
fof(f101,plain,(
  ! [X0] : ! [X1] : s('type.ringNorm.polynom'(X0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X0),X1))) = s('type.ringNorm.polynom'(X0),happ(s(fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0)),'const.ringNorm.Popp_0'),s('type.ringNorm.polynom'(X0),X1)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6] : ! [X9] : s('type.ringNorm.polynom'(X6),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X6),X9))) = s('type.ringNorm.polynom'(X6),happ(s(fun('type.ringNorm.polynom'(X6),'type.ringNorm.polynom'(X6)),'const.ringNorm.Popp_0'),s('type.ringNorm.polynom'(X6),X9)))),
  file('thm.ringNorm.polynom_distinct|split|3','ARITY_0_1_const.ringNorm.Popp')).
fof(f23879,plain,(
  ( ! [X6,X7,X11] : (s(X6,happ(s(fun('type.quote.index',X6),X7),s('type.quote.index',sK2))) = s(X6,happ(s(fun('type.ringNorm.polynom'(sK0),X6),X11),s('type.ringNorm.polynom'(sK0),sK1)))) )),
  inference(superposition,[],[f8582,f344])).
fof(f344,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X1,happ(s(fun('type.quote.index',X1),X3),s('type.quote.index',X2))) = s(X1,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X0),'const.ringNorm.Pvar_1'(s('type.quote.index',X2))),s(fun('type.quote.index',X1),X3),s(fun(X0,X1),X4),s(fun('type.ringNorm.polynom'(X0),fun('type.ringNorm.polynom'(X0),X1)),X5),s(fun('type.ringNorm.polynom'(X0),fun('type.ringNorm.polynom'(X0),X1)),X6),s(fun('type.ringNorm.polynom'(X0),X1),X7)))) )),
  inference(cnf_transformation,[],[f163])).
fof(f163,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X1,happ(s(fun('type.quote.index',X1),X3),s('type.quote.index',X2))) = s(X1,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X0),'const.ringNorm.Pvar_1'(s('type.quote.index',X2))),s(fun('type.quote.index',X1),X3),s(fun(X0,X1),X4),s(fun('type.ringNorm.polynom'(X0),fun('type.ringNorm.polynom'(X0),X1)),X5),s(fun('type.ringNorm.polynom'(X0),fun('type.ringNorm.polynom'(X0),X1)),X6),s(fun('type.ringNorm.polynom'(X0),X1),X7)))),
  inference(flattening,[],[f162])).
fof(f162,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6,X7] : s(X1,happ(s(fun('type.quote.index',X1),X3),s('type.quote.index',X2))) = s(X1,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X0),'const.ringNorm.Pvar_1'(s('type.quote.index',X2))),s(fun('type.quote.index',X1),X3),s(fun(X0,X1),X4),s(fun('type.ringNorm.polynom'(X0),fun('type.ringNorm.polynom'(X0),X1)),X5),s(fun('type.ringNorm.polynom'(X0),fun('type.ringNorm.polynom'(X0),X1)),X6),s(fun('type.ringNorm.polynom'(X0),X1),X7)))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X6,X10] : ! [X8,X11,X12,X13,X14,X15] : s(X10,happ(s(fun('type.quote.index',X10),X11),s('type.quote.index',X8))) = s(X10,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X6),'const.ringNorm.Pvar_1'(s('type.quote.index',X8))),s(fun('type.quote.index',X10),X11),s(fun(X6,X10),X12),s(fun('type.ringNorm.polynom'(X6),fun('type.ringNorm.polynom'(X6),X10)),X13),s(fun('type.ringNorm.polynom'(X6),fun('type.ringNorm.polynom'(X6),X10)),X14),s(fun('type.ringNorm.polynom'(X6),X10),X15)))),
  file('thm.ringNorm.polynom_distinct|split|3','thm.ringNorm.polynom_case_def|split|0')).
fof(f8582,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.ringNorm.polynom'(sK0),X0),X1),s('type.ringNorm.polynom'(sK0),sK1))) = s(X0,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pvar_1'(s('type.quote.index',sK2))),s(fun('type.quote.index',X0),X2),s(fun(sK0,X0),X3),s(fun('type.ringNorm.polynom'(sK0),fun('type.ringNorm.polynom'(sK0),X0)),X4),s(fun('type.ringNorm.polynom'(sK0),fun('type.ringNorm.polynom'(sK0),X0)),X5),s(fun('type.ringNorm.polynom'(sK0),X0),X1)))) )),
  inference(superposition,[],[f346,f245])).
fof(f346,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun('type.ringNorm.polynom'(X1),X0),X7),s('type.ringNorm.polynom'(X1),X2))) = s(X0,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X1),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X1),X2))),s(fun('type.quote.index',X0),X3),s(fun(X1,X0),X4),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X5),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X6),s(fun('type.ringNorm.polynom'(X1),X0),X7)))) )),
  inference(cnf_transformation,[],[f167])).
fof(f167,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun('type.ringNorm.polynom'(X1),X0),X7),s('type.ringNorm.polynom'(X1),X2))) = s(X0,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X1),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X1),X2))),s(fun('type.quote.index',X0),X3),s(fun(X1,X0),X4),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X5),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X6),s(fun('type.ringNorm.polynom'(X1),X0),X7)))),
  inference(flattening,[],[f166])).
fof(f166,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun('type.ringNorm.polynom'(X1),X0),X7),s('type.ringNorm.polynom'(X1),X2))) = s(X0,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X1),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X1),X2))),s(fun('type.quote.index',X0),X3),s(fun(X1,X0),X4),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X5),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X6),s(fun('type.ringNorm.polynom'(X1),X0),X7)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X10,X6] : ! [X8,X11,X12,X13,X14,X15] : s(X10,happ(s(fun('type.ringNorm.polynom'(X6),X10),X15),s('type.ringNorm.polynom'(X6),X8))) = s(X10,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X6),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X6),X8))),s(fun('type.quote.index',X10),X11),s(fun(X6,X10),X12),s(fun('type.ringNorm.polynom'(X6),fun('type.ringNorm.polynom'(X6),X10)),X13),s(fun('type.ringNorm.polynom'(X6),fun('type.ringNorm.polynom'(X6),X10)),X14),s(fun('type.ringNorm.polynom'(X6),X10),X15)))),
  file('thm.ringNorm.polynom_distinct|split|3','thm.ringNorm.polynom_case_def|split|4')).
fof(f47391,plain,(
  ( ! [X81,X82] : (s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(sK0),X81),s('type.ringNorm.polynom'(sK0),sK1))) = s('type.ringNorm.polynom'(sK0),happ(s(fun('type.quote.index','type.ringNorm.polynom'(sK0)),X82),s('type.quote.index',sK2)))) )),
  inference(superposition,[],[f23879,f307])).
fof(f307,plain,(
  ( ! [X2,X0,X1] : (s('type.ringNorm.polynom'(X0),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X0),X1),s('type.ringNorm.polynom'(X0),X2))) = s('type.ringNorm.polynom'(X0),happ(s(fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0)),happ(s(fun('type.ringNorm.polynom'(X0),fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0))),'const.ringNorm.Pmult_0'),s('type.ringNorm.polynom'(X0),X1))),s('type.ringNorm.polynom'(X0),X2)))) )),
  inference(cnf_transformation,[],[f115])).
fof(f115,plain,(
  ! [X0,X1,X2] : s('type.ringNorm.polynom'(X0),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X0),X1),s('type.ringNorm.polynom'(X0),X2))) = s('type.ringNorm.polynom'(X0),happ(s(fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0)),happ(s(fun('type.ringNorm.polynom'(X0),fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0))),'const.ringNorm.Pmult_0'),s('type.ringNorm.polynom'(X0),X1))),s('type.ringNorm.polynom'(X0),X2)))),
  inference(flattening,[],[f114])).
fof(f114,plain,(
  ! [X0] : ! [X1,X2] : s('type.ringNorm.polynom'(X0),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X0),X1),s('type.ringNorm.polynom'(X0),X2))) = s('type.ringNorm.polynom'(X0),happ(s(fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0)),happ(s(fun('type.ringNorm.polynom'(X0),fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0))),'const.ringNorm.Pmult_0'),s('type.ringNorm.polynom'(X0),X1))),s('type.ringNorm.polynom'(X0),X2)))),
  inference(rectify,[],[f51])).
fof(f51,axiom,(
  ! [X6] : ! [X9,X36] : s('type.ringNorm.polynom'(X6),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X6),X9),s('type.ringNorm.polynom'(X6),X36))) = s('type.ringNorm.polynom'(X6),happ(s(fun('type.ringNorm.polynom'(X6),'type.ringNorm.polynom'(X6)),happ(s(fun('type.ringNorm.polynom'(X6),fun('type.ringNorm.polynom'(X6),'type.ringNorm.polynom'(X6))),'const.ringNorm.Pmult_0'),s('type.ringNorm.polynom'(X6),X9))),s('type.ringNorm.polynom'(X6),X36)))),
  file('thm.ringNorm.polynom_distinct|split|3','ARITY_0_2_const.ringNorm.Pmult')).
% SZS output end Proof for thm.ringNorm
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 315
% Passive clauses: 5372
% Generated clauses: 53604
% Final active clauses: 153
% Final passive clauses: 1248
% Input formulas: 69
% Initial clauses: 125
% 
% Duplicate literals: 925
% Trivial inequalities: 322
% Fw subsumption resolutions: 1601
% Fw demodulations: 9452
% Bw demodulations: 2973
% Condensations: 1661
% Global subsumptions: 572
% 
% Simple tautologies: 294
% Equational tautologies: 348
% Forward subsumptions: 28782
% Fw demodulations to eq. taut.: 2817
% Bw demodulations to eq. taut.: 376
% 
% Unit resulting resolution: 8116
% Factoring: 2
% Forward superposition: 11782
% Backward superposition: 9774
% Self superposition: 319
% Equality factoring: 172
% Equality resolution: 39
% 
% Split clauses: 106
% Split components: 213
% SAT solver clauses: 6824
% SAT solver unit clauses: 476
% SAT solver binary clauses: 2464
% 
% Sat splits: 106
% Sat splitting refutations: 666
% 
% Memory used [KB]: 53218
% Time elapsed: 3.536 s
% ------------------------------
% ------------------------------
% Success in time 4.85 s
