Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.extreal
% (46667)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 456
% Passive clauses: 4148
% Generated clauses: 25141
% Final active clauses: 453
% Final passive clauses: 3367
% Input formulas: 40
% Initial clauses: 118
% 
% Duplicate literals: 712
% Trivial inequalities: 71
% Fw subsumption resolutions: 2741
% Fw demodulations: 3227
% Bw demodulations: 2
% 
% Simple tautologies: 331
% Equational tautologies: 48
% Forward subsumptions: 12871
% 
% Binary resolution: 6506
% Factoring: 23
% Forward superposition: 5027
% Backward superposition: 6340
% Self superposition: 71
% Equality factoring: 60
% Equality resolution: 11
% 
% Split clauses: 633
% Split components: 1276
% SAT solver clauses: 1163
% SAT solver unit clauses: 19
% SAT solver binary clauses: 202
% 
% Sat splits: 633
% 
% TWLsolver clauses: 1160
% TWLsolver calls for satisfiability: 245
% 
% Memory used [KB]: 11001
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.extreal
% (46812)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 281
% Passive clauses: 1445
% Generated clauses: 18345
% Final active clauses: 266
% Final passive clauses: 1059
% Input formulas: 40
% Initial clauses: 118
% 
% Duplicate literals: 790
% Trivial inequalities: 109
% Fw subsumption resolutions: 1445
% Fw demodulations: 1317
% Bw demodulations: 119
% Condensations: 1641
% Global subsumptions: 13
% 
% Simple tautologies: 259
% Equational tautologies: 39
% Forward subsumptions: 9233
% 
% Binary resolution: 1247
% Factoring: 9
% Forward superposition: 5103
% Backward superposition: 4836
% Self superposition: 87
% Equality factoring: 60
% Equality resolution: 15
% 
% Split clauses: 200
% Split components: 405
% SAT solver clauses: 2193
% SAT solver unit clauses: 135
% SAT solver binary clauses: 56
% SAT solver learnt clauses: 26
% SAT solver learnt literals: 10
% 
% Sat splits: 200
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 2141
% TWLsolver calls for satisfiability: 1024
% 
% Memory used [KB]: 13176
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.extreal
% (46888)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1730
% Passive clauses: 9504
% Generated clauses: 235152
% Final active clauses: 660
% Final passive clauses: 1506
% Input formulas: 40
% Initial clauses: 118
% 
% Duplicate literals: 7847
% Trivial inequalities: 387
% Fw subsumption resolutions: 10687
% Fw demodulations: 21517
% Bw demodulations: 5196
% Condensations: 26159
% Global subsumptions: 1265
% 
% Simple tautologies: 437
% Equational tautologies: 781
% Forward subsumptions: 154430
% Bw demodulations to eq. taut.: 33
% 
% Unit resulting resolution: 13549
% Factoring: 4
% Forward superposition: 71183
% Backward superposition: 69017
% Self superposition: 323
% Equality factoring: 123
% Equality resolution: 15
% 
% Split clauses: 323
% Split components: 653
% SAT solver clauses: 13224
% SAT solver unit clauses: 697
% SAT solver binary clauses: 3658
% 
% Sat splits: 323
% Sat splitting refutations: 1807
% 
% Memory used [KB]: 106949
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.extreal
% (47024)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1811
% Passive clauses: 27234
% Generated clauses: 187887
% Final active clauses: 1591
% Final passive clauses: 23534
% Input formulas: 40
% Initial clauses: 115
% 
% Selected by SInE selection: 37
% SInE iterations: 1
% 
% Duplicate literals: 7348
% Trivial inequalities: 3790
% Fw subsumption resolutions: 20265
% Fw demodulations: 2703
% Bw demodulations: 154
% 
% Simple tautologies: 6947
% Equational tautologies: 7780
% Forward subsumptions: 104064
% 
% Binary resolution: 134962
% Factoring: 79
% Forward superposition: 12472
% Backward superposition: 4841
% Self superposition: 83
% Equality factoring: 68
% Equality resolution: 129
% 
% Split clauses: 8889
% Split components: 18825
% SAT solver clauses: 9209
% SAT solver unit clauses: 33
% SAT solver binary clauses: 2044
% 
% Sat splits: 8889
% Sat splitting refutations: 60
% 
% Memory used [KB]: 42728
% Time elapsed: 4.0000 s
% ------------------------------
% ------------------------------
% remaining time: 142 next slice time: 6
lrs+1002_2:3_br=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:sas=minisat:stl=30:sac=on:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none:sp=reverse_arity:urr=on_4 on thm.extreal
% (47174)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 261
% Passive clauses: 2211
% Generated clauses: 18871
% Final active clauses: 255
% Final passive clauses: 1918
% Input formulas: 40
% Initial clauses: 118
% 
% Duplicate literals: 333
% Trivial inequalities: 231
% Fw subsumption resolutions: 1716
% Fw demodulations: 734
% Bw demodulations: 6
% Global subsumptions: 258
% 
% Simple tautologies: 242
% Equational tautologies: 49
% Forward subsumptions: 11776
% 
% Unit resulting resolution: 3623
% Factoring: 1
% Forward superposition: 6856
% Backward superposition: 5015
% Self superposition: 83
% Equality factoring: 6
% Equality resolution: 23
% 
% Split clauses: 9
% Split components: 18
% SAT solver clauses: 3801
% SAT solver unit clauses: 1353
% SAT solver binary clauses: 569
% 
% Sat splits: 9
% Sat splitting refutations: 1321
% 
% Memory used [KB]: 15095
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 135 next slice time: 15
lrs+1011_12_bs=on:bsr=unit_only:cond=on:gs=on:gsaa=from_current:gsssp=full:nwc=1.1:sas=minisat:stl=60:sos=all:sac=on:sdd=off:sser=off:sfr=on:ssfp=100000:ssfq=1.2:ssnc=none:sp=reverse_arity:updr=off_14 on thm.extreal
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.extreal
% SZS output start Proof for thm.extreal
fof(f3142,plain,(
  $false),
  inference(subsumption_resolution,[],[f3141,f200])).
fof(f200,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(cnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0] : s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X12] : s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X12),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.extreal.EXTREAL_SUM_IMAGE_ZERO','thm.extreal.mul_rzero')).
fof(f3141,plain,(
  s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK7,bool),sK8))))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(forward_demodulation,[],[f3140,f200])).
fof(f3140,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK7,bool),sK8))))),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(trivial_inequality_removal,[],[f3133])).
fof(f3133,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK7,bool),sK8))))),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f398,f200])).
fof(f398,plain,(
  ( ! [X0] : (s('type.extreal.extreal',X0) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK7,bool),sK8))))),s('type.extreal.extreal',X0)))) )),
  inference(subsumption_resolution,[],[f397,f193])).
fof(f193,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.extreal.EXTREAL_SUM_IMAGE_ZERO',aHL_TRUTH)).
fof(f397,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.extreal.extreal',X0) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK7,bool),sK8))))),s('type.extreal.extreal',X0)))) )),
  inference(forward_demodulation,[],[f396,f317])).
fof(f317,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK7,bool),sK8)))),
  inference(resolution,[],[f191,f209])).
fof(f209,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f144])).
fof(f144,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X9] : (s(bool,'T_0') = s(bool,X9) <=> p(s(bool,X9)))),
  file('thm.extreal.EXTREAL_SUM_IMAGE_ZERO','thm.bool.EQ_CLAUSES|split|1')).
fof(f191,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK7,bool),sK8))))),
  inference(cnf_transformation,[],[f141])).
fof(f141,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK7,bool),sK8)))) & s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(sK7,'type.extreal.extreal'),'_dst_x0_1'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(fun(sK7,bool),sK8)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7,sK8])],[f103])).
fof(f103,plain,(
  ? [X0,X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) & s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(X0,'type.extreal.extreal'),'_dst_x0_1'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(fun(X0,bool),X1))))),
  inference(ennf_transformation,[],[f43])).
fof(f43,plain,(
  ~! [X0,X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) => s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(X0,'type.extreal.extreal'),'_dst_x0_1'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(fun(X0,bool),X1))))),
  inference(flattening,[],[f42])).
fof(f42,plain,(
  ~! [X0] : ! [X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) => s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(X0,'type.extreal.extreal'),'_dst_x0_1'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(fun(X0,bool),X1))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X7)))) => s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(X6,'type.extreal.extreal'),'_dst_x0_1'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(fun(X6,bool),X7))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X7)))) => s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(X6,'type.extreal.extreal'),'_dst_x0_1'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(fun(X6,bool),X7))))),
  file('thm.extreal.EXTREAL_SUM_IMAGE_ZERO','thm.extreal.EXTREAL_SUM_IMAGE_ZERO')).
fof(f396,plain,(
  ( ! [X0] : (s('type.extreal.extreal',X0) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK7,bool),sK8))))),s('type.extreal.extreal',X0))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK7,bool),sK8))))) )),
  inference(forward_demodulation,[],[f395,f265])).
fof(f265,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f92])).
fof(f92,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))),
  inference(flattening,[],[f91])).
fof(f91,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X20,X21] : ! [X22,X23] : s(X21,X22) = s(X21,happ(s(fun(X20,X21),'_dst_x0_1'(s(X21,X22))),s(X20,X23)))),
  file('thm.extreal.EXTREAL_SUM_IMAGE_ZERO','_dst_x0_ax')).
fof(f395,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK7,bool),sK8))))),s('type.extreal.extreal',X0))) | s('type.extreal.extreal',X0) != s('type.extreal.extreal',happ(s(fun(sK7,'type.extreal.extreal'),'_dst_x0_1'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(sK7,sK9(X0,'_dst_x0_1'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),sK8,sK7)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK7,bool),sK8))))) )),
  inference(superposition,[],[f192,f236])).
fof(f236,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(X0,'type.extreal.extreal'),X2),s(fun(X0,bool),X1))) = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))))),s('type.extreal.extreal',X3))) | s('type.extreal.extreal',X3) != s('type.extreal.extreal',happ(s(fun(X0,'type.extreal.extreal'),X2),s(X0,sK9(X3,X2,X1,X0)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f160])).
fof(f160,plain,(
  ! [X0,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) | ! [X2,X3] : (s('type.extreal.extreal',X3) != s('type.extreal.extreal',happ(s(fun(X0,'type.extreal.extreal'),X2),s(X0,sK9(X3,X2,X1,X0)))) | s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(X0,'type.extreal.extreal'),X2),s(fun(X0,bool),X1))) = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))))),s('type.extreal.extreal',X3)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK9])],[f114])).
fof(f114,plain,(
  ! [X0,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) | ! [X2,X3] : (? [X4] : s('type.extreal.extreal',X3) != s('type.extreal.extreal',happ(s(fun(X0,'type.extreal.extreal'),X2),s(X0,X4))) | s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(X0,'type.extreal.extreal'),X2),s(fun(X0,bool),X1))) = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))))),s('type.extreal.extreal',X3)))))),
  inference(ennf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0,X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) => ! [X2,X3] : (! [X4] : s('type.extreal.extreal',X3) = s('type.extreal.extreal',happ(s(fun(X0,'type.extreal.extreal'),X2),s(X0,X4))) => s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(X0,'type.extreal.extreal'),X2),s(fun(X0,bool),X1))) = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))))),s('type.extreal.extreal',X3)))))),
  inference(flattening,[],[f72])).
fof(f72,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) => ! [X2,X3] : (! [X4] : s('type.extreal.extreal',X3) = s('type.extreal.extreal',happ(s(fun(X0,'type.extreal.extreal'),X2),s(X0,X4))) => s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(X0,'type.extreal.extreal'),X2),s(fun(X0,bool),X1))) = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))))),s('type.extreal.extreal',X3)))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X6] : ! [X14] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X14)))) => ! [X15,X12] : (! [X16] : s('type.extreal.extreal',X12) = s('type.extreal.extreal',happ(s(fun(X6,'type.extreal.extreal'),X15),s(X6,X16))) => s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(X6,'type.extreal.extreal'),X15),s(fun(X6,bool),X14))) = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X6,bool),X14))))),s('type.extreal.extreal',X12)))))),
  file('thm.extreal.EXTREAL_SUM_IMAGE_ZERO','thm.extreal.EXTREAL_SUM_IMAGE_FINITE_CONST')).
fof(f192,plain,(
  s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(sK7,'type.extreal.extreal'),'_dst_x0_1'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(fun(sK7,bool),sK8)))),
  inference(cnf_transformation,[],[f141])).
% SZS output end Proof for thm.extreal
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 200
% Passive clauses: 437
% Generated clauses: 2821
% Final active clauses: 197
% Final passive clauses: 306
% Input formulas: 40
% Initial clauses: 118
% 
% Duplicate literals: 181
% Trivial inequalities: 31
% Fw subsumption resolutions: 304
% Bw subsumption resolutions: 1
% Fw demodulations: 317
% Bw demodulations: 1
% Condensations: 225
% 
% Simple tautologies: 305
% Equational tautologies: 10
% Forward subsumptions: 1052
% Backward subsumptions: 13
% 
% Binary resolution: 955
% Forward superposition: 453
% Backward superposition: 354
% Self superposition: 2
% Equality factoring: 2
% Equality resolution: 6
% 
% SAT solver clauses: 439
% SAT solver unit clauses: 14
% SAT solver binary clauses: 66
% 
% Memory used [KB]: 2174
% Time elapsed: 0.093 s
% ------------------------------
% ------------------------------
% Success in time 16.528 s
