Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.numeral
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.numeral
% SZS output start Proof for thm.numeral
fof(f7714,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f7678,f455,f656,f654,f379,f655])).
fof(f655,plain,(
  ~p(s(bool,'F_0')) | $spl1),
  inference(cnf_transformation,[],[f655_D])).
fof(f655_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f379,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f92])).
fof(f92,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.numeral.numeral_distrib|split|23',aHL_FALSITY)).
fof(f654,plain,(
  p(s(bool,'F_0')) | $spl0),
  inference(cnf_transformation,[],[f654_D])).
fof(f654_D,plain,(
  p(s(bool,'F_0')) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f656,plain,(
  p(s(bool,'T_0')) | $spl2),
  inference(cnf_transformation,[],[f656_D])).
fof(f656_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f455,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f281])).
fof(f281,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f280])).
fof(f280,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f207])).
fof(f207,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f142])).
fof(f142,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f53])).
fof(f53,axiom,(
  ! [X9] : ((p(s(bool,'F_0')) => p(s(bool,X9))) <=> p(s(bool,'T_0')))),
  file('thm.numeral.numeral_distrib|split|23','thm.bool.IMP_CLAUSES|split|2')).
fof(f7678,plain,(
  $false | $spl2),
  inference(subsumption_resolution,[],[f7593,f695])).
fof(f695,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl2),
  inference(resolution,[],[f411,f656])).
fof(f411,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f252])).
fof(f252,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f122])).
fof(f122,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f64])).
fof(f64,axiom,(
  ! [X9] : (s(bool,'F_0') = s(bool,X9) <=> ~p(s(bool,X9)))),
  file('thm.numeral.numeral_distrib|split|23','thm.bool.EQ_CLAUSES|split|3')).
fof(f7593,plain,(
  s(bool,'T_0') = s(bool,'F_0')),
  inference(superposition,[],[f7481,f1930])).
fof(f1930,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)))),
  inference(trivial_inequality_removal,[],[f1926])).
fof(f1926,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)))),
  inference(superposition,[],[f368,f395])).
fof(f395,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X9] : (s(bool,'T_0') = s(bool,X9) | s(bool,'F_0') = s(bool,X9))),
  file('thm.numeral.numeral_distrib|split|23','thm.bool.BOOL_CASES_AX')).
fof(f368,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)))),
  inference(cnf_transformation,[],[f247])).
fof(f247,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13])],[f196])).
fof(f196,plain,(
  ? [X0] : s(bool,'F_0') != s(bool,'const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(ennf_transformation,[],[f89])).
fof(f89,plain,(
  ~! [X0] : s(bool,'F_0') = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : s(bool,'F_0') = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X6)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : s(bool,'F_0') = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X6)))),
  file('thm.numeral.numeral_distrib|split|23','thm.numeral.numeral_distrib|split|23')).
fof(f7481,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(superposition,[],[f481,f705])).
fof(f705,plain,(
  ( ! [X6] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0')))) )),
  inference(resolution,[],[f412,f468])).
fof(f468,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))) )),
  inference(cnf_transformation,[],[f151])).
fof(f151,plain,(
  ! [X0] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))),
  inference(flattening,[],[f150])).
fof(f150,plain,(
  ! [X0] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))),
  inference(rectify,[],[f78])).
fof(f78,axiom,(
  ! [X6] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0'))))),
  file('thm.numeral.numeral_distrib|split|23','thm.prim_rec.NOT_LESS_|48|')).
fof(f412,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f252])).
fof(f481,plain,(
  ( ! [X0,X1] : (s(bool,'const.arithmetic.>_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f161])).
fof(f161,plain,(
  ! [X0,X1] : s(bool,'const.arithmetic.>_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X8,X6] : s(bool,'const.arithmetic.>_2'(s('type.num.num',X8),s('type.num.num',X6))) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num',X8)))),
  file('thm.numeral.numeral_distrib|split|23','thm.arithmetic.GREATER_DEF')).
% SZS output end Proof for thm.numeral
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 464
% Passive clauses: 2029
% Generated clauses: 7337
% Final active clauses: 458
% Final passive clauses: 1426
% Input formulas: 87
% Initial clauses: 277
% 
% Duplicate literals: 169
% Trivial inequalities: 27
% Fw subsumption resolutions: 603
% Fw demodulations: 290
% Bw demodulations: 5
% 
% Simple tautologies: 193
% Equational tautologies: 21
% Forward subsumptions: 3929
% 
% Binary resolution: 3237
% Factoring: 788
% Forward superposition: 1173
% Backward superposition: 695
% Self superposition: 14
% Equality factoring: 2
% Equality resolution: 8
% 
% Split clauses: 124
% Split components: 250
% SAT solver clauses: 222
% SAT solver unit clauses: 15
% SAT solver binary clauses: 47
% SAT solver learnt clauses: 1
% 
% Sat splits: 124
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 220
% TWLsolver calls for satisfiability: 183
% 
% Memory used [KB]: 3582
% Time elapsed: 0.137 s
% ------------------------------
% ------------------------------
% Success in time 0.132 s
