Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.extreal
% (89601)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 813
% Passive clauses: 8427
% Generated clauses: 22300
% Final active clauses: 693
% Final passive clauses: 6962
% Input formulas: 69
% Initial clauses: 239
% 
% Duplicate literals: 201
% Trivial inequalities: 146
% Fw subsumption resolutions: 1219
% Fw demodulations: 1815
% Bw demodulations: 91
% 
% Simple tautologies: 227
% Equational tautologies: 21
% Forward subsumptions: 9784
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 7526
% Factoring: 726
% Forward superposition: 3507
% Backward superposition: 6514
% Self superposition: 55
% Equality factoring: 2
% Equality resolution: 22
% 
% Split clauses: 423
% Split components: 847
% SAT solver clauses: 708
% SAT solver unit clauses: 32
% SAT solver binary clauses: 102
% 
% Sat splits: 423
% Sat splitting refutations: 22
% 
% TWLsolver clauses: 702
% TWLsolver calls for satisfiability: 367
% 
% Memory used [KB]: 15735
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.extreal
% (89617)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 372
% Passive clauses: 2554
% Generated clauses: 14265
% Final active clauses: 265
% Final passive clauses: 1005
% Input formulas: 69
% Initial clauses: 239
% 
% Duplicate literals: 205
% Trivial inequalities: 13
% Fw subsumption resolutions: 885
% Fw demodulations: 1320
% Bw demodulations: 1268
% Condensations: 199
% Global subsumptions: 13
% 
% Simple tautologies: 202
% Equational tautologies: 4
% Forward subsumptions: 5350
% Bw demodulations to eq. taut.: 3
% 
% Binary resolution: 1149
% Factoring: 45
% Forward superposition: 2313
% Backward superposition: 3282
% Self superposition: 59
% Equality factoring: 2
% Equality resolution: 8
% 
% Split clauses: 57
% Split components: 115
% SAT solver clauses: 2522
% SAT solver unit clauses: 286
% SAT solver binary clauses: 120
% SAT solver learnt clauses: 28
% SAT solver learnt literals: 12
% 
% Sat splits: 57
% Sat splitting refutations: 42
% 
% TWLsolver clauses: 2468
% TWLsolver calls for satisfiability: 3766
% 
% Memory used [KB]: 17654
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.extreal
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.extreal
% SZS output start Proof for thm.extreal
fof(f74424,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f3413,f3423,f3422,f3797,f3800,f3801,f3803,f3805,f3809,f3811,f3815,f3817,f3832,f3834,f3849,f3851,f3867,f3883,f3903,f3957,f3977,f3984,f3986,f4054,f4079,f4092,f4094,f4098,f4105,f4107,f4124,f4199,f4219,f4226,f4228,f4296,f4321,f4334,f4336,f4340,f4347,f4349,f4366,f4404,f4406,f4408,f4420,f4422,f4424,f4454,f4456,f4470,f4472,f4502,f4504,f4518,f4520,f4550,f4552,f4566,f4568,f4585,f4587,f5142,f5143,f5144,f5145,f5146,f5147,f4613,f4615,f4617,f4625,f4627,f4629,f4631,f4639,f4646,f4653,f4663,f4665,f4667,f4674,f4676,f4683,f4685,f4692,f4694,f4701,f4703,f4713,f4722,f4731,f4768,f4778,f4780,f4794,f4798,f4803,f4805,f4807,f4809,f4811,f4813,f4815,f4834,f4836,f4848,f4857,f4859,f4868,f4870,f4875,f4877,f4886,f4888,f4893,f4895,f4904,f4906,f4911,f4913,f4925,f4934,f4953,f4955,f4962,f4964,f4971,f4973,f4980,f4997,f5008,f5010,f5012,f5021,f5023,f5029,f5034,f5039,f5049,f5056,f5058,f5060,f5062,f5089,f5091,f5117,f5119,f5122,f5127,f5129,f5157,f5158,f5159,f5160,f5161,f5162,f5163,f5164,f5165,f5166,f5167,f5168,f5169,f5170,f5171,f5172,f5173,f5174,f5175,f5176,f5177,f5178,f5179,f5180,f5181,f5182,f5183,f5184,f5185,f5186,f5187,f5188,f5189,f5190,f5191,f5192,f5193,f5194,f5195,f5196,f5197,f5198,f5199,f5200,f5201,f5202,f5203,f5204,f5205,f5206,f5207,f5208,f5209,f5210,f5211,f5212,f5213,f5214,f5215,f5216,f5217,f5218,f5219,f5220,f5221,f5222,f5223,f5224,f5225,f5226,f5227,f5228,f5229,f5230,f5231,f5232,f5233,f5234,f5235,f5236,f5237,f5238,f5239,f5240,f5241,f5242,f5243,f5244,f5245,f5246,f5247,f5248,f5249,f5250,f5251,f5252,f5253,f5254,f5255,f5256,f5257,f5258,f5259,f5260,f5261,f5262,f5263,f5264,f5265,f5266,f5267,f5268,f5269,f5270,f5271,f5272,f5273,f5274,f5275,f5276,f5277,f5278,f5279,f5280,f5281,f5282,f5283,f5284,f5285,f5286,f5287,f5288,f5289,f5290,f5291,f5292,f5293,f5294,f5295,f5296,f5297,f5298,f5299,f5300,f5301,f5302,f5303,f5304,f5305,f5306,f5307,f5308,f5309,f5310,f5311,f5312,f5313,f5314,f5315,f5316,f5317,f5318,f5319,f5320,f5321,f5322,f5323,f5324,f5325,f5326,f5327,f5328,f5329,f5330,f5331,f5332,f5333,f5334,f5335,f5336,f5337,f5338,f5339,f5340,f5341,f5342,f5343,f5344,f5345,f5346,f5347,f5348,f5349,f5350,f5351,f5352,f5353,f5354,f5355,f5356,f5357,f5358,f5359,f5360,f5361,f5362,f5363,f5364,f5365,f5366,f5367,f5368,f5369,f5370,f5371,f5372,f5373,f5374,f5375,f5376,f5377,f5378,f5379,f5380,f5381,f5382,f5383,f5384,f5385,f5386,f5387,f5388,f5389,f5390,f5391,f5392,f5393,f5394,f5395,f5396,f5397,f5398,f5399,f5400,f5401,f5402,f5403,f5404,f5405,f5406,f5407,f5408,f5409,f5410,f5411,f5412,f5413,f5414,f5415,f5416,f5417,f5418,f5419,f5420,f5421,f5422,f5423,f5424,f5425,f5426,f5427,f5428,f5429,f5430,f5431,f5432,f5433,f5434,f5435,f5436,f5437,f5438,f5439,f5440,f5441,f5442,f5443,f5444,f5445,f5446,f5447,f5448,f5449,f5450,f5451,f5452,f5453,f5454,f5455,f5456,f5457,f5458,f5459,f5460,f5461,f5462,f5463,f5464,f5465,f5466,f5467,f5468,f5469,f5470,f5471,f5472,f5473,f5474,f5475,f5476,f5477,f5478,f5479,f5480,f5481,f5482,f5483,f5484,f5485,f5824,f5826,f5865,f5867,f6426,f6427,f6474,f6475,f7340,f7341,f7406,f7407,f7471,f7472,f7540,f7541,f3798,f8808,f8811,f8810,f9799,f9801,f26708,f26710,f26722,f26711,f26724,f72083,f74058,f74080,f74105,f74260,f74261,f74266,f74280,f74298,f74312,f74326,f74341])).
fof(f74341,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f72833,f26711])).
fof(f72833,plain,(
  ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK15))))),
  inference(unit_resulting_resolution,[],[f593,f691,f67423,f513])).
fof(f513,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f315])).
fof(f315,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f314])).
fof(f314,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP5(X2,X0,X1)))),
  inference(flattening,[],[f313])).
fof(f313,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP5(X2,X0,X1)))),
  inference(nnf_transformation,[],[f229])).
fof(f229,plain,(
  ! [X2,X0,X1] : (sP5(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f67423,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))))))),
  inference(unit_resulting_resolution,[],[f26908,f435])).
fof(f435,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK51(X2,X1,X0))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f277])).
fof(f277,plain,(
  ! [X0,X1,X2] : ((s(X0,X2) = s(X0,sK51(X2,X1,X0)) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK51(X2,X1,X0)))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK51])],[f208])).
fof(f208,plain,(
  ! [X0,X1,X2] : (? [X3] : (s(X0,X2) = s(X0,X3) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(ennf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0,X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f121])).
fof(f121,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X17] : ! [X6,X18] : (! [X7] : (s(X17,X7) = s(X17,X18) => p(s(bool,happ(s(fun(X17,bool),X6),s(X17,X7))))) => p(s(bool,'|dollar|exists_1'(s(fun(X17,bool),X6)))))),
  file('thm.extreal.extreal_mul_ind','thm.bool.JRH_INDUCT_UTIL')).
fof(f26908,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',X0))))) )),
  inference(subsumption_resolution,[],[f26907,f355])).
fof(f355,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.extreal.extreal_mul_ind',aHL_TRUTH)).
fof(f26907,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',X0))))) )),
  inference(forward_demodulation,[],[f26906,f726])).
fof(f726,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))),
  inference(unit_resulting_resolution,[],[f345,f404])).
fof(f404,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f259])).
fof(f259,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f94])).
fof(f94,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X18] : (s(bool,'T_0') = s(bool,X18) <=> p(s(bool,X18)))),
  file('thm.extreal.extreal_mul_ind','thm.bool.EQ_CLAUSES|split|1')).
fof(f345,plain,(
  p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.NegInf_0'))))),
  inference(cnf_transformation,[],[f242])).
fof(f242,plain,(
  p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) & p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) & p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) & p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) & ! [X8] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X8))))),s('type.extreal.extreal','const.extreal.NegInf_0')))) & ! [X7] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X7)))))) & ! [X6] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X6))))),s('type.extreal.extreal','const.extreal.PosInf_0')))) & ! [X5] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))))) & ! [X3,X4] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4)))))) & ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK15))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13,sK14,sK15])],[f241])).
fof(f241,plain,(
  ? [X0] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) & p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) & p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) & p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) & ! [X8] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X8))))),s('type.extreal.extreal','const.extreal.NegInf_0')))) & ! [X7] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X7)))))) & ! [X6] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X6))))),s('type.extreal.extreal','const.extreal.PosInf_0')))) & ! [X5] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))))) & ! [X3,X4] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4)))))) & ? [X1,X2] : ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal',X2)))))),
  inference(rectify,[],[f181])).
fof(f181,plain,(
  ? [X0] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) & p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) & p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) & p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) & ! [X6] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X6))))),s('type.extreal.extreal','const.extreal.NegInf_0')))) & ! [X5] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))))) & ! [X4] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4))))),s('type.extreal.extreal','const.extreal.PosInf_0')))) & ! [X3] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3)))))) & ! [X1,X2] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1))))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))))) & ? [X7,X8] : ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal',X7))),s('type.extreal.extreal',X8)))))),
  inference(flattening,[],[f180])).
fof(f180,plain,(
  ? [X0] : ((p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) & p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) & p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) & p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) & ! [X6] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X6))))),s('type.extreal.extreal','const.extreal.NegInf_0')))) & ! [X5] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))))) & ! [X4] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4))))),s('type.extreal.extreal','const.extreal.PosInf_0')))) & ! [X3] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3)))))) & ! [X1,X2] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1))))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))))))) & ? [X7,X8] : ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal',X7))),s('type.extreal.extreal',X8)))))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0] : ((p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) & p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) & p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) & p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) & ! [X6] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X6))))),s('type.extreal.extreal','const.extreal.NegInf_0')))) & ! [X5] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))))) & ! [X4] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4))))),s('type.extreal.extreal','const.extreal.PosInf_0')))) & ! [X3] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3)))))) & ! [X1,X2] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1))))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))))))) => ! [X7,X8] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal',X7))),s('type.extreal.extreal',X8)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ((p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X6),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) & p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X6),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) & p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X6),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) & p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X6),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) & ! [X7] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X6),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X7))))),s('type.extreal.extreal','const.extreal.NegInf_0')))) & ! [X8] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X6),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X8)))))) & ! [X7] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X6),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X7))))),s('type.extreal.extreal','const.extreal.PosInf_0')))) & ! [X8] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X6),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X8)))))) & ! [X7,X8] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X6),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X7))))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X8))))))) => ! [X9,X10] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X6),s('type.extreal.extreal',X9))),s('type.extreal.extreal',X10)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ((p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X6),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) & p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X6),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) & p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X6),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) & p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X6),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) & ! [X7] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X6),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X7))))),s('type.extreal.extreal','const.extreal.NegInf_0')))) & ! [X8] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X6),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X8)))))) & ! [X7] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X6),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X7))))),s('type.extreal.extreal','const.extreal.PosInf_0')))) & ! [X8] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X6),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X8)))))) & ! [X7,X8] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X6),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X7))))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X8))))))) => ! [X9,X10] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X6),s('type.extreal.extreal',X9))),s('type.extreal.extreal',X10)))))),
  file('thm.extreal.extreal_mul_ind','thm.extreal.extreal_mul_ind')).
fof(f26906,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',X0)))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.NegInf_0'))))) )),
  inference(subsumption_resolution,[],[f26905,f355])).
fof(f26905,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',X0)))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.NegInf_0'))))) )),
  inference(forward_demodulation,[],[f26904,f727])).
fof(f727,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))),
  inference(unit_resulting_resolution,[],[f346,f404])).
fof(f346,plain,(
  p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.PosInf_0'))))),
  inference(cnf_transformation,[],[f242])).
fof(f26904,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',X0)))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.NegInf_0'))))) )),
  inference(subsumption_resolution,[],[f26795,f355])).
fof(f26795,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',X0)))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.NegInf_0'))))) )),
  inference(superposition,[],[f369,f1041])).
fof(f1041,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))))) )),
  inference(unit_resulting_resolution,[],[f350,f404])).
fof(f350,plain,(
  ( ! [X7] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X7))))))) )),
  inference(cnf_transformation,[],[f242])).
fof(f369,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(X0))))))) | p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',X1)))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal','const.extreal.NegInf_0'))))) )),
  inference(cnf_transformation,[],[f247])).
fof(f247,plain,(
  ! [X0] : (~p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal','const.extreal.NegInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK17(X0))))))) | ! [X1] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',X1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK17])],[f246])).
fof(f246,plain,(
  ! [X0] : (~p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal','const.extreal.NegInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ? [X2] : ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))))) | ! [X1] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',X1)))))),
  inference(rectify,[],[f185])).
fof(f185,plain,(
  ! [X0] : (~p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal','const.extreal.NegInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ? [X1] : ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))))) | ! [X2] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',X2)))))),
  inference(flattening,[],[f184])).
fof(f184,plain,(
  ! [X0] : ((~p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal','const.extreal.NegInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ? [X1] : ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1))))))) | ! [X2] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',X2)))))),
  inference(ennf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0] : ((p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal','const.extreal.NegInf_0')))) & p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal','const.extreal.PosInf_0')))) & ! [X1] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1))))))) => ! [X2] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',X2)))))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X6] : ((p(s(bool,happ(s(fun('type.extreal.extreal',bool),X6),s('type.extreal.extreal','const.extreal.NegInf_0')))) & p(s(bool,happ(s(fun('type.extreal.extreal',bool),X6),s('type.extreal.extreal','const.extreal.PosInf_0')))) & ! [X24] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),X6),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X24))))))) => ! [X27] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),X6),s('type.extreal.extreal',X27)))))),
  file('thm.extreal.extreal_mul_ind','thm.extreal.extreal_induction')).
fof(f691,plain,(
  ( ! [X0,X1] : (sP5(X0,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK15)),X1)) )),
  inference(unit_resulting_resolution,[],[f354,f515])).
fof(f515,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f315])).
fof(f354,plain,(
  ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK15))))),
  inference(cnf_transformation,[],[f242])).
fof(f593,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f358,f355])).
fof(f358,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f243])).
fof(f243,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f72])).
fof(f72,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f63])).
fof(f63,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.extreal.extreal_mul_ind','thm.bool.NOT_CLAUSES|split|2')).
fof(f74326,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f72866,f26711])).
fof(f72866,plain,(
  ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK15))))),
  inference(unit_resulting_resolution,[],[f593,f692,f67423,f513])).
fof(f692,plain,(
  ( ! [X0,X1] : (sP5(happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK15)),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f354,f516])).
fof(f516,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f315])).
fof(f74312,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f72899,f26711])).
fof(f72899,plain,(
  ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK15))))),
  inference(unit_resulting_resolution,[],[f593,f693,f67423,f517])).
fof(f517,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f318])).
fof(f318,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f317])).
fof(f317,plain,(
  ! [X1,X0,X2] : ((~sP4(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP4(X1,X0,X2)))),
  inference(flattening,[],[f316])).
fof(f316,plain,(
  ! [X1,X0,X2] : ((~sP4(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP4(X1,X0,X2)))),
  inference(nnf_transformation,[],[f228])).
fof(f228,plain,(
  ! [X1,X0,X2] : (sP4(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f693,plain,(
  ( ! [X0,X1] : (sP4(X0,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK15)),X1)) )),
  inference(unit_resulting_resolution,[],[f354,f519])).
fof(f519,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f318])).
fof(f74298,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f72931,f26711])).
fof(f72931,plain,(
  ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK15))))),
  inference(unit_resulting_resolution,[],[f593,f694,f67423,f517])).
fof(f694,plain,(
  ( ! [X0,X1] : (sP4(happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK15)),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f354,f520])).
fof(f520,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f318])).
fof(f74280,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f73159,f26711])).
fof(f73159,plain,(
  ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK15))))),
  inference(unit_resulting_resolution,[],[f593,f695,f67423,f536])).
fof(f536,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f326])).
fof(f326,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(flattening,[],[f325])).
fof(f325,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(nnf_transformation,[],[f232])).
fof(f232,plain,(
  ! [X0,X1,X2] : (sP7(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f695,plain,(
  ( ! [X0,X1] : (sP7(X0,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK15)),X1)) )),
  inference(unit_resulting_resolution,[],[f354,f538])).
fof(f538,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f326])).
fof(f74266,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f73191,f26711])).
fof(f73191,plain,(
  ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK15))))),
  inference(unit_resulting_resolution,[],[f593,f696,f67423,f536])).
fof(f696,plain,(
  ( ! [X0,X1] : (sP7(happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK15)),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f354,f539])).
fof(f539,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f326])).
fof(f74261,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f73214,f26711])).
fof(f73214,plain,(
  ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK15))))),
  inference(unit_resulting_resolution,[],[f593,f15407,f67423,f545])).
fof(f545,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP8(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f328])).
fof(f328,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP8(X2,X0,X1)) & (~sP8(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f327])).
fof(f327,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP8(X2,X0,X1)) & (~sP8(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f234])).
fof(f234,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP8(X2,X0,X1))),
  inference(definition_folding,[],[f167,f233,f232])).
fof(f233,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (sP7(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f167,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f166])).
fof(f166,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f53])).
fof(f53,axiom,(
  ! [X24,X23,X55] : ((p(s(bool,X55)) <=> (p(s(bool,X23)) & p(s(bool,X24)))) <=> ((p(s(bool,X55)) | ~p(s(bool,X23)) | ~p(s(bool,X24))) & (p(s(bool,X23)) | ~p(s(bool,X55))) & (p(s(bool,X24)) | ~p(s(bool,X55)))))),
  file('thm.extreal.extreal_mul_ind','thm.sat.sat_23_prove_ltb')).
fof(f15407,plain,(
  ( ! [X0] : (sP8('F_0',X0,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK15)))) )),
  inference(unit_resulting_resolution,[],[f593,f695,f586])).
fof(f586,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0)) | ~sP7(X1,X2,X0)) )),
  inference(duplicate_literal_removal,[],[f535])).
fof(f535,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X0)) | ~sP7(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f324])).
fof(f324,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | (sP7(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP7(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f323])).
fof(f323,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (sP7(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP7(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f322])).
fof(f322,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (sP7(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP7(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f233])).
fof(f74260,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f73218,f26711])).
fof(f73218,plain,(
  ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK15))))),
  inference(unit_resulting_resolution,[],[f593,f15438,f67423,f545])).
fof(f15438,plain,(
  ( ! [X0] : (sP8('F_0',happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK15)),X0)) )),
  inference(unit_resulting_resolution,[],[f593,f696,f586])).
fof(f74105,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f73642,f26711])).
fof(f73642,plain,(
  ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK15))))),
  inference(unit_resulting_resolution,[],[f593,f698,f67423,f570])).
fof(f570,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f342])).
fof(f342,plain,(
  ! [X0,X1,X2] : ((~sP11(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP11(X0,X1,X2)))),
  inference(rectify,[],[f341])).
fof(f341,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP11(X2,X0,X1)))),
  inference(flattening,[],[f340])).
fof(f340,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP11(X2,X0,X1)))),
  inference(nnf_transformation,[],[f238])).
fof(f238,plain,(
  ! [X2,X0,X1] : (sP11(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f698,plain,(
  ( ! [X0,X1] : (sP11(X0,X1,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK15)))) )),
  inference(unit_resulting_resolution,[],[f354,f571])).
fof(f571,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f342])).
fof(f74080,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f73700,f26711])).
fof(f73700,plain,(
  ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK15))))),
  inference(unit_resulting_resolution,[],[f593,f699,f67423,f570])).
fof(f699,plain,(
  ( ! [X0,X1] : (sP11(happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK15)),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f354,f573])).
fof(f573,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f342])).
fof(f74058,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f73923,f26711])).
fof(f73923,plain,(
  ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK15))))),
  inference(unit_resulting_resolution,[],[f593,f15438,f67423,f1651])).
fof(f1651,plain,(
  ( ! [X12,X10,X11] : (~sP8(X12,X11,'const.combin.I_1'(s(bool,X10))) | ~p(s(bool,X11)) | p(s(bool,X12)) | ~p(s(bool,X10))) )),
  inference(superposition,[],[f545,f426])).
fof(f426,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))) )),
  inference(cnf_transformation,[],[f110])).
fof(f110,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))),
  inference(flattening,[],[f109])).
fof(f109,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))),
  inference(rectify,[],[f48])).
fof(f48,axiom,(
  ! [X17] : ! [X7] : s(X17,X7) = s(X17,'const.combin.I_1'(s(X17,X7)))),
  file('thm.extreal.extreal_mul_ind','thm.combin.I_THM')).
fof(f72083,plain,(
  $false | $spl33),
  inference(subsumption_resolution,[],[f72019,f355])).
fof(f72019,plain,(
  ~p(s(bool,'T_0')) | $spl33),
  inference(backward_demodulation,[],[f67433,f26724])).
fof(f67433,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',X0)))) )),
  inference(unit_resulting_resolution,[],[f355,f26908,f474])).
fof(f474,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f212])).
fof(f212,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f211])).
fof(f211,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f148])).
fof(f148,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X11,X12] : ((p(s(bool,X11)) => p(s(bool,X12))) => ((p(s(bool,X12)) => p(s(bool,X11))) => s(bool,X11) = s(bool,X12)))),
  file('thm.extreal.extreal_mul_ind','thm.bool.IMP_ANTISYM_AX')).
fof(f26724,plain,(
  ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',sK26(sK13))))) | $spl33),
  inference(cnf_transformation,[],[f26724_D])).
fof(f26724_D,plain,(
  ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',sK26(sK13))))) <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f26711,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal',X0))),s('type.extreal.extreal',X1))))) ) | $spl30),
  inference(cnf_transformation,[],[f26711_D])).
fof(f26711_D,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal',X0))),s('type.extreal.extreal',X1))))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f26722,plain,(
  ( ! [X2,X3] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal',X2))),s('type.extreal.extreal',X3)))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',sK26(sK13)))))) )),
  inference(subsumption_resolution,[],[f26721,f355])).
fof(f26721,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal',X2))),s('type.extreal.extreal',X3)))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',sK26(sK13)))))) )),
  inference(forward_demodulation,[],[f26720,f1215])).
fof(f1215,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))))) )),
  inference(unit_resulting_resolution,[],[f353,f404])).
fof(f353,plain,(
  ( ! [X4,X3] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4))))))) )),
  inference(cnf_transformation,[],[f242])).
fof(f26720,plain,(
  ( ! [X2,X3] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal',X2))),s('type.extreal.extreal',X3)))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',sK26(sK13))))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK24(sK13)))))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK25(sK13)))))))) )),
  inference(subsumption_resolution,[],[f26719,f355])).
fof(f26719,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal',X2))),s('type.extreal.extreal',X3)))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',sK26(sK13))))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK24(sK13)))))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK25(sK13)))))))) )),
  inference(forward_demodulation,[],[f26718,f729])).
fof(f729,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))),
  inference(unit_resulting_resolution,[],[f348,f404])).
fof(f348,plain,(
  p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.PosInf_0'))))),
  inference(cnf_transformation,[],[f242])).
fof(f26718,plain,(
  ( ! [X2,X3] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal',X2))),s('type.extreal.extreal',X3)))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',sK26(sK13))))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK24(sK13)))))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK25(sK13)))))))) )),
  inference(subsumption_resolution,[],[f26717,f355])).
fof(f26717,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal',X2))),s('type.extreal.extreal',X3)))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',sK26(sK13))))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK24(sK13)))))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK25(sK13)))))))) )),
  inference(forward_demodulation,[],[f26716,f1099])).
fof(f1099,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))))),s('type.extreal.extreal','const.extreal.PosInf_0')))) )),
  inference(unit_resulting_resolution,[],[f351,f404])).
fof(f351,plain,(
  ( ! [X6] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X6))))),s('type.extreal.extreal','const.extreal.PosInf_0'))))) )),
  inference(cnf_transformation,[],[f242])).
fof(f26716,plain,(
  ( ! [X2,X3] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal',X2))),s('type.extreal.extreal',X3)))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK27(sK13)))))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',sK26(sK13))))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK24(sK13)))))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK25(sK13)))))))) )),
  inference(subsumption_resolution,[],[f26715,f355])).
fof(f26715,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal',X2))),s('type.extreal.extreal',X3)))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK27(sK13)))))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',sK26(sK13))))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK24(sK13)))))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK25(sK13)))))))) )),
  inference(forward_demodulation,[],[f26714,f728])).
fof(f728,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))),
  inference(unit_resulting_resolution,[],[f347,f404])).
fof(f347,plain,(
  p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.NegInf_0'))))),
  inference(cnf_transformation,[],[f242])).
fof(f26714,plain,(
  ( ! [X2,X3] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal',X2))),s('type.extreal.extreal',X3)))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK27(sK13)))))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',sK26(sK13))))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK24(sK13)))))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK25(sK13)))))))) )),
  inference(subsumption_resolution,[],[f26713,f355])).
fof(f26713,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal',X2))),s('type.extreal.extreal',X3)))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK27(sK13)))))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',sK26(sK13))))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK24(sK13)))))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK25(sK13)))))))) )),
  inference(forward_demodulation,[],[f26712,f1157])).
fof(f1157,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))))) )),
  inference(unit_resulting_resolution,[],[f352,f404])).
fof(f352,plain,(
  ( ! [X5] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5))))))) )),
  inference(cnf_transformation,[],[f242])).
fof(f26712,plain,(
  ( ! [X2,X3] : (~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK29(sK13))))))) | p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal',X2))),s('type.extreal.extreal',X3)))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK27(sK13)))))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',sK26(sK13))))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK24(sK13)))))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK25(sK13)))))))) )),
  inference(subsumption_resolution,[],[f26589,f355])).
fof(f26589,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK29(sK13))))))) | p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal',X2))),s('type.extreal.extreal',X3)))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK27(sK13)))))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',sK26(sK13))))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK24(sK13)))))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK25(sK13)))))))) )),
  inference(superposition,[],[f371,f981])).
fof(f981,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))))),s('type.extreal.extreal','const.extreal.NegInf_0')))) )),
  inference(unit_resulting_resolution,[],[f349,f404])).
fof(f349,plain,(
  ( ! [X8] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X8))))),s('type.extreal.extreal','const.extreal.NegInf_0'))))) )),
  inference(cnf_transformation,[],[f242])).
fof(f371,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK28(X0)))))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK29(X0))))))) | p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal',X2)))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK27(X0)))))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',sK26(X0))))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK24(X0)))))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK25(X0)))))))) )),
  inference(cnf_transformation,[],[f251])).
fof(f251,plain,(
  ! [X0] : (~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK24(X0)))))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK25(X0))))))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',sK26(X0))))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK27(X0)))))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK28(X0)))))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK29(X0))))))) | ! [X1,X2] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal',X2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK24,sK25,sK26,sK27,sK28,sK29])],[f250])).
fof(f250,plain,(
  ! [X0] : (? [X7,X8] : ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X7))))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X8)))))) | ? [X6] : ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',X6)))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ? [X5] : ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5))))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | ? [X4] : ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4))))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | ? [X3] : ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3)))))) | ! [X1,X2] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal',X2)))))),
  inference(rectify,[],[f189])).
fof(f189,plain,(
  ! [X0] : (? [X5,X6] : ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5))))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X6)))))) | ? [X4] : ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',X4)))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ? [X3] : ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | ? [X2] : ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | ? [X1] : ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))))) | ! [X7,X8] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal',X7))),s('type.extreal.extreal',X8)))))),
  inference(flattening,[],[f188])).
fof(f188,plain,(
  ! [X0] : ((? [X5,X6] : ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5))))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X6)))))) | ? [X4] : ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',X4)))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ? [X3] : ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | ? [X2] : ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | ? [X1] : ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1))))))) | ! [X7,X8] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal',X7))),s('type.extreal.extreal',X8)))))),
  inference(ennf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0] : ((! [X5,X6] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5))))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X6)))))) & ! [X4] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',X4)))) & p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) & ! [X3] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))))),s('type.extreal.extreal','const.extreal.PosInf_0')))) & p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) & ! [X2] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))))),s('type.extreal.extreal','const.extreal.NegInf_0')))) & ! [X1] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1))))))) => ! [X7,X8] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal',X7))),s('type.extreal.extreal',X8)))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6] : ((! [X7,X8] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X6),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X7))))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X8)))))) & ! [X13] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X6),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',X13)))) & p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X6),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) & ! [X14] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X6),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X14))))),s('type.extreal.extreal','const.extreal.PosInf_0')))) & p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X6),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) & ! [X16] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X6),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X16))))),s('type.extreal.extreal','const.extreal.NegInf_0')))) & ! [X15] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X6),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X15))))))) => ! [X9,X10] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X6),s('type.extreal.extreal',X9))),s('type.extreal.extreal',X10)))))),
  file('thm.extreal.extreal_mul_ind','thm.extreal.extreal_le_ind')).
fof(f26710,plain,(
  ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal',sK20(sK13))))) | $spl29),
  inference(cnf_transformation,[],[f26710_D])).
fof(f26710_D,plain,(
  ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal',sK20(sK13))))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f26708,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal',X0))),s('type.extreal.extreal',X1)))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal',sK20(sK13)))))) )),
  inference(subsumption_resolution,[],[f26707,f355])).
fof(f26707,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal',X0))),s('type.extreal.extreal',X1)))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal',sK20(sK13)))))) )),
  inference(forward_demodulation,[],[f26706,f1215])).
fof(f26706,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal',X0))),s('type.extreal.extreal',X1)))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal',sK20(sK13))))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK18(sK13)))))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK19(sK13)))))))) )),
  inference(subsumption_resolution,[],[f26705,f355])).
fof(f26705,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal',X0))),s('type.extreal.extreal',X1)))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal',sK20(sK13))))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK18(sK13)))))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK19(sK13)))))))) )),
  inference(forward_demodulation,[],[f26704,f727])).
fof(f26704,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal',X0))),s('type.extreal.extreal',X1)))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal',sK20(sK13))))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK18(sK13)))))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK19(sK13)))))))) )),
  inference(subsumption_resolution,[],[f26703,f355])).
fof(f26703,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal',X0))),s('type.extreal.extreal',X1)))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal',sK20(sK13))))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK18(sK13)))))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK19(sK13)))))))) )),
  inference(forward_demodulation,[],[f26702,f1099])).
fof(f26702,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal',X0))),s('type.extreal.extreal',X1)))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK21(sK13)))))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal',sK20(sK13))))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK18(sK13)))))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK19(sK13)))))))) )),
  inference(subsumption_resolution,[],[f26701,f355])).
fof(f26701,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal',X0))),s('type.extreal.extreal',X1)))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK21(sK13)))))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal',sK20(sK13))))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK18(sK13)))))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK19(sK13)))))))) )),
  inference(forward_demodulation,[],[f26700,f726])).
fof(f26700,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal',X0))),s('type.extreal.extreal',X1)))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK21(sK13)))))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal',sK20(sK13))))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK18(sK13)))))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK19(sK13)))))))) )),
  inference(subsumption_resolution,[],[f26699,f355])).
fof(f26699,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal',X0))),s('type.extreal.extreal',X1)))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK21(sK13)))))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal',sK20(sK13))))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK18(sK13)))))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK19(sK13)))))))) )),
  inference(forward_demodulation,[],[f26698,f1041])).
fof(f26698,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal',X0))),s('type.extreal.extreal',X1)))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK22(sK13))))))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK21(sK13)))))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal',sK20(sK13))))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK18(sK13)))))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK19(sK13)))))))) )),
  inference(subsumption_resolution,[],[f26588,f355])).
fof(f26588,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal',X0))),s('type.extreal.extreal',X1)))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK22(sK13))))))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK21(sK13)))))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal',sK20(sK13))))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK18(sK13)))))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK19(sK13)))))))) )),
  inference(superposition,[],[f370,f981])).
fof(f370,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK23(X0)))))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal',X2)))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK22(X0))))))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK21(X0)))))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal',sK20(X0))))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK18(X0)))))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK19(X0)))))))) )),
  inference(cnf_transformation,[],[f249])).
fof(f249,plain,(
  ! [X0] : (~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK18(X0)))))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK19(X0))))))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal',sK20(X0))))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK21(X0)))))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK22(X0))))))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK23(X0)))))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | ! [X1,X2] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal',X2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK18,sK19,sK20,sK21,sK22,sK23])],[f248])).
fof(f248,plain,(
  ! [X0] : (? [X7,X8] : ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X7))))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X8)))))) | ? [X6] : ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal',X6)))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ? [X5] : ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5))))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | ? [X4] : ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4)))))) | ? [X3] : ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | ! [X1,X2] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal',X2)))))),
  inference(rectify,[],[f187])).
fof(f187,plain,(
  ! [X0] : (? [X5,X6] : ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5))))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X6)))))) | ? [X4] : ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal',X4)))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ? [X3] : ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | ? [X2] : ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))))) | ? [X1] : ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1))))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | ! [X7,X8] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal',X7))),s('type.extreal.extreal',X8)))))),
  inference(flattening,[],[f186])).
fof(f186,plain,(
  ! [X0] : ((? [X5,X6] : ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5))))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X6)))))) | ? [X4] : ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal',X4)))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ? [X3] : ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | ? [X2] : ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))))) | ? [X1] : ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1))))),s('type.extreal.extreal','const.extreal.NegInf_0'))))) | ! [X7,X8] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal',X7))),s('type.extreal.extreal',X8)))))),
  inference(ennf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0] : ((! [X5,X6] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5))))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X6)))))) & ! [X4] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal',X4)))) & p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) & ! [X3] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))))),s('type.extreal.extreal','const.extreal.PosInf_0')))) & p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) & ! [X2] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))))) & ! [X1] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1))))),s('type.extreal.extreal','const.extreal.NegInf_0'))))) => ! [X7,X8] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X0),s('type.extreal.extreal',X7))),s('type.extreal.extreal',X8)))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : ((! [X7,X8] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X6),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X7))))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X8)))))) & ! [X13] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X6),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal',X13)))) & p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X6),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) & ! [X14] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X6),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X14))))),s('type.extreal.extreal','const.extreal.PosInf_0')))) & p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X6),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) & ! [X15] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X6),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X15)))))) & ! [X16] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X6),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X16))))),s('type.extreal.extreal','const.extreal.NegInf_0'))))) => ! [X9,X10] : p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),X6),s('type.extreal.extreal',X9))),s('type.extreal.extreal',X10)))))),
  file('thm.extreal.extreal_mul_ind','thm.extreal.extreal_add_ind')).
fof(f9801,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f9800,f9798])).
fof(f9798,plain,(
  ( ! [X1] : (~p(s(bool,X1))) ) | $spl24),
  inference(forward_demodulation,[],[f9664,f455])).
fof(f455,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X1) = s(X0,happ(s(fun('type.extreal.extreal',X0),sK59(X3,X2,X1,X0)),s('type.extreal.extreal','const.extreal.NegInf_0')))) )),
  inference(cnf_transformation,[],[f291])).
fof(f291,plain,(
  ! [X0,X1,X2,X3] : (s(X0,X1) = s(X0,happ(s(fun('type.extreal.extreal',X0),sK59(X3,X2,X1,X0)),s('type.extreal.extreal','const.extreal.NegInf_0'))) & s(X0,X2) = s(X0,happ(s(fun('type.extreal.extreal',X0),sK59(X3,X2,X1,X0)),s('type.extreal.extreal','const.extreal.PosInf_0'))) & ! [X5] : s(X0,happ(s(fun('type.realax.real',X0),X3),s('type.realax.real',X5))) = s(X0,happ(s(fun('type.extreal.extreal',X0),sK59(X3,X2,X1,X0)),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK59])],[f142])).
fof(f142,plain,(
  ! [X0,X1,X2,X3] : ? [X4] : (s(X0,X1) = s(X0,happ(s(fun('type.extreal.extreal',X0),X4),s('type.extreal.extreal','const.extreal.NegInf_0'))) & s(X0,X2) = s(X0,happ(s(fun('type.extreal.extreal',X0),X4),s('type.extreal.extreal','const.extreal.PosInf_0'))) & ! [X5] : s(X0,happ(s(fun('type.realax.real',X0),X3),s('type.realax.real',X5))) = s(X0,happ(s(fun('type.extreal.extreal',X0),X4),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5))))))),
  inference(flattening,[],[f141])).
fof(f141,plain,(
  ! [X0] : ! [X1,X2,X3] : ? [X4] : (s(X0,X1) = s(X0,happ(s(fun('type.extreal.extreal',X0),X4),s('type.extreal.extreal','const.extreal.NegInf_0'))) & s(X0,X2) = s(X0,happ(s(fun('type.extreal.extreal',X0),X4),s('type.extreal.extreal','const.extreal.PosInf_0'))) & ! [X5] : s(X0,happ(s(fun('type.realax.real',X0),X3),s('type.realax.real',X5))) = s(X0,happ(s(fun('type.extreal.extreal',X0),X4),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5))))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X17] : ! [X29,X30,X31] : ? [X32] : (s(X17,X29) = s(X17,happ(s(fun('type.extreal.extreal',X17),X32),s('type.extreal.extreal','const.extreal.NegInf_0'))) & s(X17,X30) = s(X17,happ(s(fun('type.extreal.extreal',X17),X32),s('type.extreal.extreal','const.extreal.PosInf_0'))) & ! [X13] : s(X17,happ(s(fun('type.realax.real',X17),X31),s('type.realax.real',X13))) = s(X17,happ(s(fun('type.extreal.extreal',X17),X32),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13))))))),
  file('thm.extreal.extreal_mul_ind','thm.extreal.extreal_Axiom')).
fof(f9664,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK59(X0,'F_0',X1,bool)),s('type.extreal.extreal','const.extreal.NegInf_0'))))) ) | $spl24),
  inference(unit_resulting_resolution,[],[f8879,f407])).
fof(f407,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f261])).
fof(f261,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f96])).
fof(f96,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X18] : (s(bool,'F_0') = s(bool,X18) <=> ~p(s(bool,X18)))),
  file('thm.extreal.extreal_mul_ind','thm.bool.EQ_CLAUSES|split|2')).
fof(f8879,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X2) = s(X0,happ(s(fun('type.extreal.extreal',X0),sK59(X3,X2,X1,X0)),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl24),
  inference(backward_demodulation,[],[f8810,f456])).
fof(f456,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X2) = s(X0,happ(s(fun('type.extreal.extreal',X0),sK59(X3,X2,X1,X0)),s('type.extreal.extreal','const.extreal.PosInf_0')))) )),
  inference(cnf_transformation,[],[f291])).
fof(f9800,plain,(
  ( ! [X265,X267,X266] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK59(X265,X266,X267,bool)),s('type.extreal.extreal','const.extreal.NegInf_0'))))) ) | $spl24),
  inference(subsumption_resolution,[],[f9730,f9676])).
fof(f9676,plain,(
  ( ! [X6,X4,X7] : (s(X4,X6) = s(X4,X7)) ) | $spl24),
  inference(superposition,[],[f8879,f455])).
fof(f9730,plain,(
  ( ! [X265,X267,X266] : (s(bool,'T_0') != s(bool,X266) | p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK59(X265,X266,X267,bool)),s('type.extreal.extreal','const.extreal.NegInf_0'))))) ) | $spl24),
  inference(superposition,[],[f403,f8879])).
fof(f403,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f259])).
fof(f9799,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f9663,f9798])).
fof(f9663,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK59(X0,'T_0',X1,bool)),s('type.extreal.extreal','const.extreal.NegInf_0'))))) ) | $spl24),
  inference(unit_resulting_resolution,[],[f8879,f403])).
fof(f8810,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.PosInf_0') | $spl24),
  inference(cnf_transformation,[],[f8810_D])).
fof(f8810_D,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.PosInf_0') <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f8811,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(X8,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(X8,X6),s(X8,X7),s(fun('type.realax.real',X8),X5))) = s(X8,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal',X1),s(X8,X4),s(X8,X3),s(fun('type.realax.real',X8),X2))) | s(X8,X4) != s(X8,X6) | s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X1) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK60(X2,X3,X4,'const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0')),X1,X5,X6,X7,X8))))) ) | $spl26),
  inference(cnf_transformation,[],[f8811_D])).
fof(f8811_D,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(X8,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(X8,X6),s(X8,X7),s(fun('type.realax.real',X8),X5))) = s(X8,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal',X1),s(X8,X4),s(X8,X3),s(fun('type.realax.real',X8),X2))) | s(X8,X4) != s(X8,X6) | s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X1) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK60(X2,X3,X4,'const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0')),X1,X5,X6,X7,X8))))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f8808,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(X8,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(X8,X6),s(X8,X7),s(fun('type.realax.real',X8),X5))) = s(X8,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal',X1),s(X8,X4),s(X8,X3),s(fun('type.realax.real',X8),X2))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.PosInf_0') | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK60(X2,X3,X4,'const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0')),X1,X5,X6,X7,X8)))) | s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X1) | s(X8,X4) != s(X8,X6)) )),
  inference(forward_demodulation,[],[f8807,f363])).
fof(f363,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0')))) )),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0] : s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0')))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X14] : s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X14))),s('type.extreal.extreal','const.extreal.PosInf_0')))),
  file('thm.extreal.extreal_mul_ind','thm.extreal.extreal_sub_def|split|3')).
fof(f8807,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.PosInf_0') | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK60(X2,X3,X4,'const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0')),X1,X5,X6,X7,X8)))) | s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X1) | s(X8,X4) != s(X8,X6) | s(X8,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal',X1),s(X8,X4),s(X8,X3),s(fun('type.realax.real',X8),X2))) = s(X8,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0'))),s(X8,X6),s(X8,X7),s(fun('type.realax.real',X8),X5)))) )),
  inference(forward_demodulation,[],[f8806,f363])).
fof(f8806,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK60(X2,X3,X4,'const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0')),X1,X5,X6,X7,X8)))) | s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X1) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s(X8,X4) != s(X8,X6) | s(X8,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal',X1),s(X8,X4),s(X8,X3),s(fun('type.realax.real',X8),X2))) = s(X8,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0'))),s(X8,X6),s(X8,X7),s(fun('type.realax.real',X8),X5)))) )),
  inference(forward_demodulation,[],[f8770,f363])).
fof(f8770,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X1) | s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK60(X2,X3,X4,'const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0')),X1,X5,X6,X7,X8)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s(X8,X4) != s(X8,X6) | s(X8,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal',X1),s(X8,X4),s(X8,X3),s(fun('type.realax.real',X8),X2))) = s(X8,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0'))),s(X8,X6),s(X8,X7),s(fun('type.realax.real',X8),X5)))) )),
  inference(superposition,[],[f460,f363])).
fof(f460,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s('type.extreal.extreal',X4) != s('type.extreal.extreal',X5) | s('type.extreal.extreal',X5) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK60(X8,X7,X6,X5,X4,X3,X2,X1,X0)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X5) | s(X0,X2) != s(X0,X6) | s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal',X4),s(X0,X6),s(X0,X7),s(fun('type.realax.real',X0),X8))) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal',X5),s(X0,X2),s(X0,X1),s(fun('type.realax.real',X0),X3)))) )),
  inference(cnf_transformation,[],[f292])).
fof(f292,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : (s('type.extreal.extreal',X4) != s('type.extreal.extreal',X5) | (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X5) & s(X0,X2) != s(X0,X6)) | (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X5) & s(X0,X1) != s(X0,X7)) | (s('type.extreal.extreal',X5) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK60(X8,X7,X6,X5,X4,X3,X2,X1,X0)))) & s(X0,happ(s(fun('type.realax.real',X0),X3),s('type.realax.real',sK60(X8,X7,X6,X5,X4,X3,X2,X1,X0)))) != s(X0,happ(s(fun('type.realax.real',X0),X8),s('type.realax.real',sK60(X8,X7,X6,X5,X4,X3,X2,X1,X0))))) | s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal',X4),s(X0,X6),s(X0,X7),s(fun('type.realax.real',X0),X8))) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal',X5),s(X0,X2),s(X0,X1),s(fun('type.realax.real',X0),X3))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK60])],[f210])).
fof(f210,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : (s('type.extreal.extreal',X4) != s('type.extreal.extreal',X5) | (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X5) & s(X0,X2) != s(X0,X6)) | (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X5) & s(X0,X1) != s(X0,X7)) | ? [X9] : (s('type.extreal.extreal',X5) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X9))) & s(X0,happ(s(fun('type.realax.real',X0),X3),s('type.realax.real',X9))) != s(X0,happ(s(fun('type.realax.real',X0),X8),s('type.realax.real',X9)))) | s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal',X4),s(X0,X6),s(X0,X7),s(fun('type.realax.real',X0),X8))) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal',X5),s(X0,X2),s(X0,X1),s(fun('type.realax.real',X0),X3))))),
  inference(flattening,[],[f209])).
fof(f209,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : ((s('type.extreal.extreal',X4) != s('type.extreal.extreal',X5) | (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X5) & s(X0,X2) != s(X0,X6)) | (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X5) & s(X0,X1) != s(X0,X7)) | ? [X9] : (s('type.extreal.extreal',X5) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X9))) & s(X0,happ(s(fun('type.realax.real',X0),X3),s('type.realax.real',X9))) != s(X0,happ(s(fun('type.realax.real',X0),X8),s('type.realax.real',X9))))) | s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal',X4),s(X0,X6),s(X0,X7),s(fun('type.realax.real',X0),X8))) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal',X5),s(X0,X2),s(X0,X1),s(fun('type.realax.real',X0),X3))))),
  inference(ennf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : ((s('type.extreal.extreal',X4) = s('type.extreal.extreal',X5) & (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X5) => s(X0,X2) = s(X0,X6)) & (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X5) => s(X0,X1) = s(X0,X7)) & ! [X9] : (s('type.extreal.extreal',X5) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X9))) => s(X0,happ(s(fun('type.realax.real',X0),X3),s('type.realax.real',X9))) = s(X0,happ(s(fun('type.realax.real',X0),X8),s('type.realax.real',X9))))) => s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal',X4),s(X0,X6),s(X0,X7),s(fun('type.realax.real',X0),X8))) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal',X5),s(X0,X2),s(X0,X1),s(fun('type.realax.real',X0),X3))))),
  inference(flattening,[],[f145])).
fof(f145,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8] : ((s('type.extreal.extreal',X4) = s('type.extreal.extreal',X5) & (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X5) => s(X0,X2) = s(X0,X6)) & (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X5) => s(X0,X1) = s(X0,X7)) & ! [X9] : (s('type.extreal.extreal',X5) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X9))) => s(X0,happ(s(fun('type.realax.real',X0),X3),s('type.realax.real',X9))) = s(X0,happ(s(fun('type.realax.real',X0),X8),s('type.realax.real',X9))))) => s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal',X4),s(X0,X6),s(X0,X7),s(fun('type.realax.real',X0),X8))) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal',X5),s(X0,X2),s(X0,X1),s(fun('type.realax.real',X0),X3))))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X17] : ! [X49,X50,X51,X52,X53,X9,X10,X25] : ((s('type.extreal.extreal',X52) = s('type.extreal.extreal',X53) & (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X53) => s(X17,X9) = s(X17,X50)) & (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X53) => s(X17,X10) = s(X17,X49)) & ! [X13] : (s('type.extreal.extreal',X53) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13))) => s(X17,happ(s(fun('type.realax.real',X17),X25),s('type.realax.real',X13))) = s(X17,happ(s(fun('type.realax.real',X17),X51),s('type.realax.real',X13))))) => s(X17,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal',X52),s(X17,X9),s(X17,X10),s(fun('type.realax.real',X17),X25))) = s(X17,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal',X53),s(X17,X50),s(X17,X49),s(fun('type.realax.real',X17),X51))))),
  file('thm.extreal.extreal_mul_ind','thm.extreal.extreal_case_cong')).
fof(f3798,plain,(
  ( ! [X3] : (~p(s(bool,X3))) ) | $spl4),
  inference(cnf_transformation,[],[f3798_D])).
fof(f3798_D,plain,(
  ( ! [X3] : (~p(s(bool,X3))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f7541,plain,(
  ( ! [X5] : (p(s(bool,sK58('const.combin.I_0',X5,bool))) | p(s(bool,happ(s(fun(bool,bool),X5),s(bool,sK58('const.combin.I_0',X5,bool)))))) ) | $spl22),
  inference(cnf_transformation,[],[f7541_D])).
fof(f7541_D,plain,(
  ( ! [X5] : (p(s(bool,sK58('const.combin.I_0',X5,bool))) | p(s(bool,happ(s(fun(bool,bool),X5),s(bool,sK58('const.combin.I_0',X5,bool)))))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f7540,plain,(
  ( ! [X4,X5] : (p(s(bool,sK58('const.combin.I_0',X5,bool))) | ~p(s(bool,X4)) | p(s(bool,happ(s(fun(bool,bool),X5),s(bool,sK58('const.combin.I_0',X5,bool)))))) )),
  inference(forward_demodulation,[],[f7522,f595])).
fof(f595,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))) )),
  inference(forward_demodulation,[],[f427,f426])).
fof(f427,plain,(
  ( ! [X0,X1] : (s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f112])).
fof(f112,plain,(
  ! [X0,X1] : s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))),
  inference(flattening,[],[f111])).
fof(f111,plain,(
  ! [X0] : ! [X1] : s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))),
  inference(rectify,[],[f49])).
fof(f49,axiom,(
  ! [X17] : ! [X21] : s(X17,'const.combin.I_1'(s(X17,X21))) = s(X17,happ(s(fun(X17,X17),'const.combin.I_0'),s(X17,X21)))),
  file('thm.extreal.extreal_mul_ind','ARITY_0_1_const.combin.I')).
fof(f7522,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,happ(s(fun(bool,bool),X5),s(bool,sK58('const.combin.I_0',X5,bool))))) | p(s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,sK58('const.combin.I_0',X5,bool)))))) )),
  inference(superposition,[],[f449,f595])).
fof(f449,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK58(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK58(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f290])).
fof(f290,plain,(
  ! [X0,X1,X2] : ((! [X8] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X8)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X8))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK56(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK57(X2,X1,X0)))))) & ((! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK58(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK58(X2,X1,X0))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK56,sK57,sK58])],[f289])).
fof(f289,plain,(
  ! [X0,X1,X2] : ((! [X8] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X8)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X8))))) | ? [X7] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X7)))) | ? [X6] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) & ((! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(rectify,[],[f288])).
fof(f288,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) & ((! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(flattening,[],[f287])).
fof(f287,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & ((! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f129])).
fof(f129,plain,(
  ! [X0] : ! [X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f59])).
fof(f59,axiom,(
  ! [X17] : ! [X6,X56] : (? [X7] : (p(s(bool,happ(s(fun(X17,bool),X6),s(X17,X7)))) | p(s(bool,happ(s(fun(X17,bool),X56),s(X17,X7))))) <=> (? [X7] : p(s(bool,happ(s(fun(X17,bool),X6),s(X17,X7)))) | ? [X7] : p(s(bool,happ(s(fun(X17,bool),X56),s(X17,X7))))))),
  file('thm.extreal.extreal_mul_ind','thm.bool.EXISTS_OR_THM')).
fof(f7472,plain,(
  ( ! [X5] : (p(s(bool,sK58(X5,'const.combin.I_0',bool))) | p(s(bool,happ(s(fun(bool,bool),X5),s(bool,sK58(X5,'const.combin.I_0',bool)))))) ) | $spl20),
  inference(cnf_transformation,[],[f7472_D])).
fof(f7472_D,plain,(
  ( ! [X5] : (p(s(bool,sK58(X5,'const.combin.I_0',bool))) | p(s(bool,happ(s(fun(bool,bool),X5),s(bool,sK58(X5,'const.combin.I_0',bool)))))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f7471,plain,(
  ( ! [X4,X5] : (p(s(bool,sK58(X5,'const.combin.I_0',bool))) | ~p(s(bool,X4)) | p(s(bool,happ(s(fun(bool,bool),X5),s(bool,sK58(X5,'const.combin.I_0',bool)))))) )),
  inference(forward_demodulation,[],[f7453,f595])).
fof(f7453,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,sK58(X5,'const.combin.I_0',bool))))) | p(s(bool,happ(s(fun(bool,bool),X5),s(bool,sK58(X5,'const.combin.I_0',bool)))))) )),
  inference(superposition,[],[f448,f595])).
fof(f448,plain,(
  ( ! [X2,X0,X5,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK58(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK58(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f290])).
fof(f7407,plain,(
  ( ! [X5] : (p(s(bool,sK57('const.combin.I_0',X5,bool))) | p(s(bool,happ(s(fun(bool,bool),X5),s(bool,sK56('const.combin.I_0',X5,bool)))))) ) | $spl18),
  inference(cnf_transformation,[],[f7407_D])).
fof(f7407_D,plain,(
  ( ! [X5] : (p(s(bool,sK57('const.combin.I_0',X5,bool))) | p(s(bool,happ(s(fun(bool,bool),X5),s(bool,sK56('const.combin.I_0',X5,bool)))))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f7406,plain,(
  ( ! [X4,X5] : (p(s(bool,sK57('const.combin.I_0',X5,bool))) | ~p(s(bool,X4)) | p(s(bool,happ(s(fun(bool,bool),X5),s(bool,sK56('const.combin.I_0',X5,bool)))))) )),
  inference(forward_demodulation,[],[f7388,f595])).
fof(f7388,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,happ(s(fun(bool,bool),X5),s(bool,sK56('const.combin.I_0',X5,bool))))) | p(s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,sK57('const.combin.I_0',X5,bool)))))) )),
  inference(superposition,[],[f447,f595])).
fof(f447,plain,(
  ( ! [X2,X0,X8,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X8)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK56(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK57(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f290])).
fof(f7341,plain,(
  ( ! [X5] : (p(s(bool,sK56(X5,'const.combin.I_0',bool))) | p(s(bool,happ(s(fun(bool,bool),X5),s(bool,sK57(X5,'const.combin.I_0',bool)))))) ) | $spl16),
  inference(cnf_transformation,[],[f7341_D])).
fof(f7341_D,plain,(
  ( ! [X5] : (p(s(bool,sK56(X5,'const.combin.I_0',bool))) | p(s(bool,happ(s(fun(bool,bool),X5),s(bool,sK57(X5,'const.combin.I_0',bool)))))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f7340,plain,(
  ( ! [X4,X5] : (p(s(bool,sK56(X5,'const.combin.I_0',bool))) | ~p(s(bool,X4)) | p(s(bool,happ(s(fun(bool,bool),X5),s(bool,sK57(X5,'const.combin.I_0',bool)))))) )),
  inference(forward_demodulation,[],[f7322,f595])).
fof(f7322,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,sK56(X5,'const.combin.I_0',bool))))) | p(s(bool,happ(s(fun(bool,bool),X5),s(bool,sK57(X5,'const.combin.I_0',bool)))))) )),
  inference(superposition,[],[f446,f595])).
fof(f446,plain,(
  ( ! [X2,X0,X8,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X8)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK56(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK57(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f290])).
fof(f6475,plain,(
  ( ! [X5] : (p(s(bool,sK53('const.combin.I_0',X5,bool))) | p(s(bool,X5))) ) | $spl14),
  inference(cnf_transformation,[],[f6475_D])).
fof(f6475_D,plain,(
  ( ! [X5] : (p(s(bool,sK53('const.combin.I_0',X5,bool))) | p(s(bool,X5))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f6474,plain,(
  ( ! [X4,X5] : (p(s(bool,sK53('const.combin.I_0',X5,bool))) | ~p(s(bool,X4)) | p(s(bool,X5))) )),
  inference(forward_demodulation,[],[f6461,f595])).
fof(f6461,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,X5)) | p(s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,sK53('const.combin.I_0',X5,bool)))))) )),
  inference(superposition,[],[f441,f595])).
fof(f441,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK53(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f282])).
fof(f282,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK52(X2,X1,X0))))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK53(X2,X1,X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK52,sK53])],[f281])).
fof(f281,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | ? [X5] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(rectify,[],[f280])).
fof(f280,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(flattening,[],[f279])).
fof(f279,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f125])).
fof(f125,plain,(
  ! [X0] : ! [X1,X2] : ((p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X17] : ! [X6,X56] : ((p(s(bool,X6)) | ? [X7] : p(s(bool,happ(s(fun(X17,bool),X56),s(X17,X7))))) <=> ? [X7] : (p(s(bool,X6)) | p(s(bool,happ(s(fun(X17,bool),X56),s(X17,X7))))))),
  file('thm.extreal.extreal_mul_ind','thm.bool.RIGHT_OR_EXISTS_THM')).
fof(f6427,plain,(
  ( ! [X5] : (p(s(bool,sK52('const.combin.I_0',X5,bool))) | p(s(bool,X5))) ) | $spl12),
  inference(cnf_transformation,[],[f6427_D])).
fof(f6427_D,plain,(
  ( ! [X5] : (p(s(bool,sK52('const.combin.I_0',X5,bool))) | p(s(bool,X5))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f6426,plain,(
  ( ! [X4,X5] : (p(s(bool,sK52('const.combin.I_0',X5,bool))) | ~p(s(bool,X4)) | p(s(bool,X5))) )),
  inference(forward_demodulation,[],[f6413,f595])).
fof(f6413,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,X5)) | p(s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,sK52('const.combin.I_0',X5,bool)))))) )),
  inference(superposition,[],[f439,f595])).
fof(f439,plain,(
  ( ! [X6,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK52(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f282])).
fof(f5867,plain,(
  p(s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0')))) | $spl10),
  inference(cnf_transformation,[],[f5867_D])).
fof(f5867_D,plain,(
  p(s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0')))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f5865,plain,(
  ( ! [X3] : (p(s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0')))) | ~p(s(bool,X3))) )),
  inference(forward_demodulation,[],[f5855,f595])).
fof(f5855,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | p(s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))))))) )),
  inference(superposition,[],[f433,f595])).
fof(f433,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1))))))) )),
  inference(cnf_transformation,[],[f207])).
fof(f207,plain,(
  ! [X0,X1,X2] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(ennf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(flattening,[],[f119])).
fof(f119,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X17] : ! [X6,X7] : (p(s(bool,happ(s(fun(X17,bool),X6),s(X17,X7)))) => p(s(bool,happ(s(fun(X17,bool),X6),s(X17,'const.min.@_1'(s(fun(X17,bool),X6)))))))),
  file('thm.extreal.extreal_mul_ind','thm.bool.SELECT_AX')).
fof(f5826,plain,(
  p(s(bool,sK50('const.combin.I_0',bool))) | $spl8),
  inference(cnf_transformation,[],[f5826_D])).
fof(f5826_D,plain,(
  p(s(bool,sK50('const.combin.I_0',bool))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f5824,plain,(
  ( ! [X3] : (p(s(bool,sK50('const.combin.I_0',bool))) | ~p(s(bool,X3))) )),
  inference(forward_demodulation,[],[f5814,f595])).
fof(f5814,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | p(s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,sK50('const.combin.I_0',bool)))))) )),
  inference(superposition,[],[f431,f595])).
fof(f431,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK50(X1,X0)))))) )),
  inference(cnf_transformation,[],[f276])).
fof(f276,plain,(
  ! [X0,X1] : ((p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK49(X1,X0))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) & (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK50(X1,X0))))) | ! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK49,sK50])],[f275])).
fof(f275,plain,(
  ! [X0,X1] : ((? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) & (? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(rectify,[],[f274])).
fof(f274,plain,(
  ! [X0,X1] : ((? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(nnf_transformation,[],[f206])).
fof(f206,plain,(
  ! [X0,X1] : (! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(ennf_transformation,[],[f116])).
fof(f116,plain,(
  ! [X0,X1] : (~? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(flattening,[],[f115])).
fof(f115,plain,(
  ! [X0] : ! [X1] : (~? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(rectify,[],[f62])).
fof(f62,axiom,(
  ! [X17] : ! [X6] : (~? [X7] : p(s(bool,happ(s(fun(X17,bool),X6),s(X17,X7)))) <=> ! [X7] : ~p(s(bool,happ(s(fun(X17,bool),X6),s(X17,X7)))))),
  file('thm.extreal.extreal_mul_ind','thm.bool.NOT_EXISTS_THM')).
fof(f5485,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f3804,f5121])).
fof(f5121,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) != s(X0,X2)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f3798,f505])).
fof(f505,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,X1))),s(X0,X2)))) | s(X0,X1) != s(X0,X2)) )),
  inference(cnf_transformation,[],[f309])).
fof(f309,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,X1))),s(X0,X2))))))),
  inference(nnf_transformation,[],[f163])).
fof(f163,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,X1))),s(X0,X2)))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X41,X42,X43] : (p(s(bool,happ(s(fun(X41,bool),'=_1'(s(X41,X42))),s(X41,X43)))) <=> s(X41,X42) = s(X41,X43))),
  file('thm.extreal.extreal_mul_ind',arity_1_HL_EQ)).
fof(f3804,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f3798,f408])).
fof(f408,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f261])).
fof(f5484,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f3806,f5121])).
fof(f3806,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f593,f3798,f471])).
fof(f471,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f212])).
fof(f5483,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f3807,f5121])).
fof(f3807,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f3798,f3798,f471])).
fof(f5482,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f3810,f593])).
fof(f3810,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f713,f3798,f471])).
fof(f713,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f593,f403])).
fof(f5481,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f3812,f5121])).
fof(f3812,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f593,f3798,f471])).
fof(f5480,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f3813,f5121])).
fof(f3813,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f3798,f3798,f471])).
fof(f5479,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f3816,f593])).
fof(f3816,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f713,f3798,f471])).
fof(f5478,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f3818,f593])).
fof(f3818,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f769,f3798,f490])).
fof(f490,plain,(
  ( ! [X0,X1] : (~sP2(X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f304])).
fof(f304,plain,(
  ! [X0,X1] : ((~sP2(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP2(X0,X1)))),
  inference(rectify,[],[f303])).
fof(f303,plain,(
  ! [X1,X0] : ((~sP2(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP2(X1,X0)))),
  inference(flattening,[],[f302])).
fof(f302,plain,(
  ! [X1,X0] : ((~sP2(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP2(X1,X0)))),
  inference(nnf_transformation,[],[f225])).
fof(f225,plain,(
  ! [X1,X0] : (sP2(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f769,plain,(
  sP2('T_0','F_0')),
  inference(unit_resulting_resolution,[],[f593,f355,f492])).
fof(f492,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP2(X0,X1)) )),
  inference(cnf_transformation,[],[f304])).
fof(f5477,plain,(
  $false | ($spl2 | $spl4)),
  inference(subsumption_resolution,[],[f3819,f3798])).
fof(f3819,plain,(
  ( ! [X0] : (p(s(bool,sK51(X0,'const.combin.I_0',bool)))) ) | ($spl2 | $spl4)),
  inference(unit_resulting_resolution,[],[f3434,f3798,f490])).
fof(f3434,plain,(
  ( ! [X0] : (sP2('T_0',sK51(X0,'const.combin.I_0',bool))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f355,f3423,f492])).
fof(f5476,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f3821,f3798])).
fof(f3821,plain,(
  p(s(bool,sK61('const.extreal.NegInf_0'))) | $spl4),
  inference(unit_resulting_resolution,[],[f831,f3798,f490])).
fof(f831,plain,(
  sP2('F_0',sK61('const.extreal.NegInf_0'))),
  inference(unit_resulting_resolution,[],[f593,f821,f495])).
fof(f495,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP2(X0,X1)) )),
  inference(cnf_transformation,[],[f304])).
fof(f821,plain,(
  p(s(bool,sK61('const.extreal.NegInf_0')))),
  inference(equality_resolution,[],[f468])).
fof(f468,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) | p(s(bool,sK61(X0)))) )),
  inference(cnf_transformation,[],[f295])).
fof(f295,plain,(
  ! [X0] : ((~p(s(bool,sK61(X0))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)) & ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0)) | p(s(bool,sK61(X0)))) & s('type.realax.real','const.extreal.real_1'(s('type.extreal.extreal',X0))) = s('type.realax.real','const.bool.COND_3'(s(bool,sK61(X0)),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.min.@_1'(s(fun('type.realax.real',bool),'_dst_x0x1_2'(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal',X0))),s(fun('type.realax.real','type.extreal.extreal'),'const.extreal.Normal_0'))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK61])],[f294])).
fof(f294,plain,(
  ! [X0] : ? [X1] : ((~p(s(bool,X1)) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)) & ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0)) | p(s(bool,X1))) & s('type.realax.real','const.extreal.real_1'(s('type.extreal.extreal',X0))) = s('type.realax.real','const.bool.COND_3'(s(bool,X1),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.min.@_1'(s(fun('type.realax.real',bool),'_dst_x0x1_2'(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal',X0))),s(fun('type.realax.real','type.extreal.extreal'),'const.extreal.Normal_0'))))))))),
  inference(flattening,[],[f293])).
fof(f293,plain,(
  ! [X0] : ? [X1] : (((~p(s(bool,X1)) | (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0))) & ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0)) | p(s(bool,X1)))) & s('type.realax.real','const.extreal.real_1'(s('type.extreal.extreal',X0))) = s('type.realax.real','const.bool.COND_3'(s(bool,X1),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.min.@_1'(s(fun('type.realax.real',bool),'_dst_x0x1_2'(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal',X0))),s(fun('type.realax.real','type.extreal.extreal'),'const.extreal.Normal_0'))))))))),
  inference(nnf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0] : ? [X1] : ((p(s(bool,X1)) <=> (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0))) & s('type.realax.real','const.extreal.real_1'(s('type.extreal.extreal',X0))) = s('type.realax.real','const.bool.COND_3'(s(bool,X1),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.min.@_1'(s(fun('type.realax.real',bool),'_dst_x0x1_2'(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal',X0))),s(fun('type.realax.real','type.extreal.extreal'),'const.extreal.Normal_0'))))))))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X7] : ? [X40] : ((p(s(bool,X40)) <=> (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X7) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X7))) & s('type.realax.real','const.extreal.real_1'(s('type.extreal.extreal',X7))) = s('type.realax.real','const.bool.COND_3'(s(bool,X40),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.min.@_1'(s(fun('type.realax.real',bool),'_dst_x0x1_2'(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal',X7))),s(fun('type.realax.real','type.extreal.extreal'),'const.extreal.Normal_0'))))))))),
  file('thm.extreal.extreal_mul_ind','thm.extreal.real_def')).
fof(f5475,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f3822,f3798])).
fof(f3822,plain,(
  p(s(bool,sK61('const.extreal.PosInf_0'))) | $spl4),
  inference(unit_resulting_resolution,[],[f885,f3798,f490])).
fof(f885,plain,(
  sP2('F_0',sK61('const.extreal.PosInf_0'))),
  inference(unit_resulting_resolution,[],[f593,f851,f495])).
fof(f851,plain,(
  p(s(bool,sK61('const.extreal.PosInf_0')))),
  inference(equality_resolution,[],[f469])).
fof(f469,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0) | p(s(bool,sK61(X0)))) )),
  inference(cnf_transformation,[],[f295])).
fof(f5474,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f3823,f3798])).
fof(f3823,plain,(
  p(s(bool,sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl4),
  inference(unit_resulting_resolution,[],[f1706,f3798,f490])).
fof(f1706,plain,(
  sP2('F_0',sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))),
  inference(unit_resulting_resolution,[],[f593,f819,f495])).
fof(f819,plain,(
  p(s(bool,sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0')))))),
  inference(unit_resulting_resolution,[],[f426,f468])).
fof(f5473,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f3824,f3798])).
fof(f3824,plain,(
  p(s(bool,sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))))) | $spl4),
  inference(unit_resulting_resolution,[],[f1952,f3798,f490])).
fof(f1952,plain,(
  sP2('F_0',sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))))),
  inference(unit_resulting_resolution,[],[f593,f849,f495])).
fof(f849,plain,(
  p(s(bool,sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.PosInf_0')))))),
  inference(unit_resulting_resolution,[],[f426,f469])).
fof(f5472,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f3825,f593])).
fof(f3825,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f816,f3798,f490])).
fof(f816,plain,(
  ( ! [X0,X1] : (sP2('const.bool.DATATYPE_1'(s(X0,X1)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f593,f772])).
fof(f772,plain,(
  ( ! [X2,X0,X1] : (sP2('const.bool.DATATYPE_1'(s(X0,X1)),X2) | p(s(bool,X2))) )),
  inference(subsumption_resolution,[],[f770,f355])).
fof(f770,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X2)) | sP2('const.bool.DATATYPE_1'(s(X0,X1)),X2)) )),
  inference(superposition,[],[f492,f425])).
fof(f425,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.DATATYPE_1'(s(X0,X1)))) )),
  inference(cnf_transformation,[],[f108])).
fof(f108,plain,(
  ! [X0,X1] : s(bool,'T_0') = s(bool,'const.bool.DATATYPE_1'(s(X0,X1)))),
  inference(flattening,[],[f107])).
fof(f107,plain,(
  ! [X0] : ! [X1] : s(bool,'T_0') = s(bool,'const.bool.DATATYPE_1'(s(X0,X1)))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X17] : ! [X7] : s(bool,'T_0') = s(bool,'const.bool.DATATYPE_1'(s(X17,X7)))),
  file('thm.extreal.extreal_mul_ind','thm.bool.DATATYPE_TAG_THM')).
fof(f5471,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f3827,f593])).
fof(f3827,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f830,f3798,f490])).
fof(f830,plain,(
  sP2(sK61('const.extreal.NegInf_0'),'F_0')),
  inference(unit_resulting_resolution,[],[f593,f821,f492])).
fof(f5470,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f3828,f593])).
fof(f3828,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f884,f3798,f490])).
fof(f884,plain,(
  sP2(sK61('const.extreal.PosInf_0'),'F_0')),
  inference(unit_resulting_resolution,[],[f593,f851,f492])).
fof(f5469,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f3829,f593])).
fof(f3829,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f1705,f3798,f490])).
fof(f1705,plain,(
  sP2(sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),'F_0')),
  inference(unit_resulting_resolution,[],[f593,f819,f492])).
fof(f5468,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f3830,f593])).
fof(f3830,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f1951,f3798,f490])).
fof(f1951,plain,(
  sP2(sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))),'F_0')),
  inference(unit_resulting_resolution,[],[f593,f849,f492])).
fof(f5467,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f3837,f593])).
fof(f3837,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f775,f3798,f490])).
fof(f775,plain,(
  sP2('F_0','T_0')),
  inference(unit_resulting_resolution,[],[f593,f355,f495])).
fof(f5466,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f3838,f593])).
fof(f3838,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f831,f3798,f490])).
fof(f5465,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f3839,f593])).
fof(f3839,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f885,f3798,f490])).
fof(f5464,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f3840,f593])).
fof(f3840,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f1706,f3798,f490])).
fof(f5463,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f3841,f593])).
fof(f3841,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f1952,f3798,f490])).
fof(f5462,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f3842,f3798])).
fof(f3842,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.DATATYPE_1'(s(X0,X1))))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f816,f3798,f490])).
fof(f5461,plain,(
  $false | ($spl2 | $spl4)),
  inference(subsumption_resolution,[],[f3843,f3798])).
fof(f3843,plain,(
  ( ! [X0] : (p(s(bool,sK51(X0,'const.combin.I_0',bool)))) ) | ($spl2 | $spl4)),
  inference(unit_resulting_resolution,[],[f3435,f3798,f490])).
fof(f3435,plain,(
  ( ! [X0] : (sP2(sK51(X0,'const.combin.I_0',bool),'T_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f355,f3423,f495])).
fof(f5460,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f3844,f3798])).
fof(f3844,plain,(
  p(s(bool,sK61('const.extreal.NegInf_0'))) | $spl4),
  inference(unit_resulting_resolution,[],[f830,f3798,f490])).
fof(f5459,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f3845,f3798])).
fof(f3845,plain,(
  p(s(bool,sK61('const.extreal.PosInf_0'))) | $spl4),
  inference(unit_resulting_resolution,[],[f884,f3798,f490])).
fof(f5458,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f3846,f3798])).
fof(f3846,plain,(
  p(s(bool,sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl4),
  inference(unit_resulting_resolution,[],[f1705,f3798,f490])).
fof(f5457,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f3847,f3798])).
fof(f3847,plain,(
  p(s(bool,sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))))) | $spl4),
  inference(unit_resulting_resolution,[],[f1951,f3798,f490])).
fof(f5456,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f3902,f4193])).
fof(f4193,plain,(
  ( ! [X0] : (sP6('T_0','T_0',X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f607,f601,f604,f3798,f512])).
fof(f512,plain,(
  ( ! [X2,X0,X1] : (~sP4(X1,X2,X0) | p(s(bool,X0)) | ~sP5(X0,X2,X1) | sP6(X0,X1,X2) | ~sP3(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f312])).
fof(f312,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (sP3(X2,X1,X0) & sP4(X1,X2,X0) & sP5(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP3(X2,X1,X0) | ~sP4(X1,X2,X0) | ~sP5(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(rectify,[],[f311])).
fof(f311,plain,(
  ! [X2,X1,X0] : ((~sP6(X2,X1,X0) | (sP3(X0,X1,X2) & sP4(X1,X0,X2) & sP5(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP3(X0,X1,X2) | ~sP4(X1,X0,X2) | ~sP5(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP6(X2,X1,X0)))),
  inference(flattening,[],[f310])).
fof(f310,plain,(
  ! [X2,X1,X0] : ((~sP6(X2,X1,X0) | (sP3(X0,X1,X2) & sP4(X1,X0,X2) & sP5(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP3(X0,X1,X2) | ~sP4(X1,X0,X2) | ~sP5(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP6(X2,X1,X0)))),
  inference(nnf_transformation,[],[f230])).
fof(f230,plain,(
  ! [X2,X1,X0] : (sP6(X2,X1,X0) <=> (sP3(X0,X1,X2) & sP4(X1,X0,X2) & sP5(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f604,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f355,f518])).
fof(f518,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f318])).
fof(f601,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f355,f514])).
fof(f514,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f315])).
fof(f607,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f355,f522])).
fof(f522,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f320])).
fof(f320,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(flattening,[],[f319])).
fof(f319,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(nnf_transformation,[],[f227])).
fof(f227,plain,(
  ! [X0,X1,X2] : (sP3(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f3902,plain,(
  ( ! [X0] : (~sP6('T_0','T_0',X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f355,f355,f3798,f509])).
fof(f509,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f312])).
fof(f5455,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f3904,f4423])).
fof(f4423,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f3798,f520])).
fof(f3904,plain,(
  ~sP4('T_0','F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f608,f1404,f601,f3798,f511])).
fof(f511,plain,(
  ( ! [X2,X0,X1] : (~sP4(X1,X2,X0) | p(s(bool,X1)) | ~sP5(X0,X2,X1) | sP6(X0,X1,X2) | ~sP3(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f312])).
fof(f1404,plain,(
  ~sP6('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f355,f713,f527])).
fof(f527,plain,(
  ( ! [X2,X0,X1] : (~sP6(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f321])).
fof(f321,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP6(X2,X1,X0)) & (~sP6(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f231])).
fof(f231,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP6(X2,X1,X0))),
  inference(definition_folding,[],[f165,f230,f229,f228,f227])).
fof(f165,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f164])).
fof(f164,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f60])).
fof(f60,axiom,(
  ! [X24,X23,X55] : ((p(s(bool,X55)) <=> s(bool,X23) = s(bool,X24)) <=> ((p(s(bool,X55)) | p(s(bool,X23)) | p(s(bool,X24))) & (p(s(bool,X55)) | ~p(s(bool,X24)) | ~p(s(bool,X23))) & (p(s(bool,X23)) | ~p(s(bool,X24)) | ~p(s(bool,X55))) & (p(s(bool,X24)) | ~p(s(bool,X23)) | ~p(s(bool,X55)))))),
  file('thm.extreal.extreal_mul_ind','thm.sat.sat_2|49|_prove_ltb')).
fof(f608,plain,(
  ( ! [X0,X1] : (sP3(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f355,f523])).
fof(f523,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f320])).
fof(f5454,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f3910,f4423])).
fof(f3910,plain,(
  ~sP4('T_0','F_0',sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f608,f1748,f601,f3798,f511])).
fof(f1748,plain,(
  ~sP6(sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),'T_0','F_0')),
  inference(unit_resulting_resolution,[],[f713,f819,f527])).
fof(f5453,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f3912,f4423])).
fof(f3912,plain,(
  ~sP4(sK61('const.extreal.NegInf_0'),'const.combin.I_1'(s(bool,sK61('const.extreal.NegInf_0'))),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f835,f1414,f832,f3798,f511])).
fof(f832,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,sK61('const.extreal.NegInf_0'))) )),
  inference(unit_resulting_resolution,[],[f821,f514])).
fof(f1414,plain,(
  ( ! [X0] : (~sP6('F_0',X0,'const.combin.I_1'(s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f593,f426,f528])).
fof(f528,plain,(
  ( ! [X2,X0,X1] : (~sP6(X2,X1,X0) | s(bool,X0) != s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f321])).
fof(f835,plain,(
  ( ! [X0,X1] : (sP3(X0,sK61('const.extreal.NegInf_0'),X1)) )),
  inference(unit_resulting_resolution,[],[f821,f523])).
fof(f5452,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f3914,f4423])).
fof(f3914,plain,(
  ~sP4(sK61('const.extreal.PosInf_0'),'const.combin.I_1'(s(bool,sK61('const.extreal.PosInf_0'))),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f889,f1414,f886,f3798,f511])).
fof(f886,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,sK61('const.extreal.PosInf_0'))) )),
  inference(unit_resulting_resolution,[],[f851,f514])).
fof(f889,plain,(
  ( ! [X0,X1] : (sP3(X0,sK61('const.extreal.PosInf_0'),X1)) )),
  inference(unit_resulting_resolution,[],[f851,f523])).
fof(f5451,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f3915,f4423])).
fof(f3915,plain,(
  ~sP4(sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),'F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f1740,f1709,f1725,f3798,f511])).
fof(f1725,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) )),
  inference(unit_resulting_resolution,[],[f819,f514])).
fof(f1709,plain,(
  ~sP6('T_0',sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),'F_0')),
  inference(unit_resulting_resolution,[],[f355,f593,f819,f509])).
fof(f1740,plain,(
  ( ! [X0,X1] : (sP3(X0,sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),X1)) )),
  inference(unit_resulting_resolution,[],[f819,f523])).
fof(f5450,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f3916,f4423])).
fof(f3916,plain,(
  ~sP4(sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),'const.combin.I_1'(s(bool,sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f1740,f1414,f1725,f3798,f511])).
fof(f5449,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f3917,f4423])).
fof(f3917,plain,(
  ~sP4('T_0','F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f608,f1404,f602,f3798,f511])).
fof(f602,plain,(
  ( ! [X0,X1] : (sP5(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f593,f515])).
fof(f5448,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f3918,f4423])).
fof(f3918,plain,(
  ~sP4(sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),'F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f1740,f1709,f602,f3798,f511])).
fof(f5447,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f3919,f4423])).
fof(f3919,plain,(
  ~sP4('T_0','F_0',sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f608,f1748,f602,f3798,f511])).
fof(f5446,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f3921,f4423])).
fof(f3921,plain,(
  ~sP4(sK61('const.extreal.NegInf_0'),'const.combin.I_1'(s(bool,sK61('const.extreal.NegInf_0'))),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f835,f1414,f603,f3798,f511])).
fof(f603,plain,(
  ( ! [X0,X1] : (sP5('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f593,f516])).
fof(f5445,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f3922,f4423])).
fof(f3922,plain,(
  ~sP4(sK61('const.extreal.PosInf_0'),'const.combin.I_1'(s(bool,sK61('const.extreal.PosInf_0'))),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f889,f1414,f603,f3798,f511])).
fof(f5444,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f3923,f4423])).
fof(f3923,plain,(
  ~sP4(sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),'const.combin.I_1'(s(bool,sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f1740,f1414,f603,f3798,f511])).
fof(f5443,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f3930,f4423])).
fof(f3930,plain,(
  ~sP4('const.combin.I_1'(s(bool,sK61('const.extreal.NegInf_0'))),sK61('const.extreal.NegInf_0'),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f836,f1412,f603,f3798,f511])).
fof(f1412,plain,(
  ( ! [X0] : (~sP6('F_0','const.combin.I_1'(s(bool,X0)),X0)) )),
  inference(unit_resulting_resolution,[],[f593,f426,f528])).
fof(f836,plain,(
  ( ! [X0,X1] : (sP3(sK61('const.extreal.NegInf_0'),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f821,f524])).
fof(f524,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f320])).
fof(f5442,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f3931,f4423])).
fof(f3931,plain,(
  ~sP4('const.combin.I_1'(s(bool,sK61('const.extreal.PosInf_0'))),sK61('const.extreal.PosInf_0'),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f890,f1412,f603,f3798,f511])).
fof(f890,plain,(
  ( ! [X0,X1] : (sP3(sK61('const.extreal.PosInf_0'),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f851,f524])).
fof(f5441,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f3932,f4423])).
fof(f3932,plain,(
  ~sP4('const.combin.I_1'(s(bool,sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))),sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f1741,f1412,f603,f3798,f511])).
fof(f1741,plain,(
  ( ! [X0,X1] : (sP3(sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f819,f524])).
fof(f5440,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f3936,f4407])).
fof(f4407,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f3798,f516])).
fof(f3936,plain,(
  ~sP5('T_0','F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f608,f1404,f604,f3798,f511])).
fof(f5439,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f3937,f4407])).
fof(f3937,plain,(
  ~sP5('T_0','F_0',sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f1740,f1709,f604,f3798,f511])).
fof(f5438,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4045,f4407])).
fof(f4045,plain,(
  ~sP5(sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),'F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f608,f1748,f1738,f3798,f511])).
fof(f1738,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) )),
  inference(unit_resulting_resolution,[],[f819,f518])).
fof(f5437,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4080,f4407])).
fof(f4080,plain,(
  ~sP5('T_0','F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f608,f1404,f605,f3798,f511])).
fof(f605,plain,(
  ( ! [X0,X1] : (sP4(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f593,f519])).
fof(f5436,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4081,f4407])).
fof(f4081,plain,(
  ~sP5('T_0','F_0',sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f1740,f1709,f605,f3798,f511])).
fof(f5435,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4082,f4407])).
fof(f4082,plain,(
  ~sP5(sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),'F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f608,f1748,f605,f3798,f511])).
fof(f5434,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4146,f4423])).
fof(f4146,plain,(
  ~sP4('T_0','F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f608,f1404,f601,f3798,f512])).
fof(f5433,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4152,f4423])).
fof(f4152,plain,(
  ~sP4('T_0','F_0',sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f608,f1748,f601,f3798,f512])).
fof(f5432,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4154,f4423])).
fof(f4154,plain,(
  ~sP4(sK61('const.extreal.NegInf_0'),'const.combin.I_1'(s(bool,sK61('const.extreal.NegInf_0'))),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f835,f1414,f832,f3798,f512])).
fof(f5431,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4156,f4423])).
fof(f4156,plain,(
  ~sP4(sK61('const.extreal.PosInf_0'),'const.combin.I_1'(s(bool,sK61('const.extreal.PosInf_0'))),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f889,f1414,f886,f3798,f512])).
fof(f5430,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4157,f4423])).
fof(f4157,plain,(
  ~sP4(sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),'F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f1740,f1709,f1725,f3798,f512])).
fof(f5429,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4158,f4423])).
fof(f4158,plain,(
  ~sP4(sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),'const.combin.I_1'(s(bool,sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f1740,f1414,f1725,f3798,f512])).
fof(f5428,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4159,f4423])).
fof(f4159,plain,(
  ~sP4('T_0','F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f608,f1404,f602,f3798,f512])).
fof(f5427,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4160,f4423])).
fof(f4160,plain,(
  ~sP4(sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),'F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f1740,f1709,f602,f3798,f512])).
fof(f5426,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4161,f4423])).
fof(f4161,plain,(
  ~sP4('T_0','F_0',sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f608,f1748,f602,f3798,f512])).
fof(f5425,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4163,f4423])).
fof(f4163,plain,(
  ~sP4(sK61('const.extreal.NegInf_0'),'const.combin.I_1'(s(bool,sK61('const.extreal.NegInf_0'))),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f835,f1414,f603,f3798,f512])).
fof(f5424,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4164,f4423])).
fof(f4164,plain,(
  ~sP4(sK61('const.extreal.PosInf_0'),'const.combin.I_1'(s(bool,sK61('const.extreal.PosInf_0'))),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f889,f1414,f603,f3798,f512])).
fof(f5423,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4165,f4423])).
fof(f4165,plain,(
  ~sP4(sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),'const.combin.I_1'(s(bool,sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f1740,f1414,f603,f3798,f512])).
fof(f5422,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4172,f4423])).
fof(f4172,plain,(
  ~sP4('const.combin.I_1'(s(bool,sK61('const.extreal.NegInf_0'))),sK61('const.extreal.NegInf_0'),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f836,f1412,f603,f3798,f512])).
fof(f5421,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4173,f4423])).
fof(f4173,plain,(
  ~sP4('const.combin.I_1'(s(bool,sK61('const.extreal.PosInf_0'))),sK61('const.extreal.PosInf_0'),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f890,f1412,f603,f3798,f512])).
fof(f5420,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4174,f4423])).
fof(f4174,plain,(
  ~sP4('const.combin.I_1'(s(bool,sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))),sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f1741,f1412,f603,f3798,f512])).
fof(f5419,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4178,f4407])).
fof(f4178,plain,(
  ~sP5('T_0','F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f608,f1404,f604,f3798,f512])).
fof(f5418,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4179,f4407])).
fof(f4179,plain,(
  ~sP5('T_0','F_0',sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f1740,f1709,f604,f3798,f512])).
fof(f5417,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4232,f4628])).
fof(f4628,plain,(
  ( ! [X0] : (~sP6(X0,'T_0',sK61('const.extreal.NegInf_0'))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f822,f3798,f528])).
fof(f822,plain,(
  s(bool,'T_0') = s(bool,sK61('const.extreal.NegInf_0'))),
  inference(unit_resulting_resolution,[],[f821,f404])).
fof(f4232,plain,(
  sP6(sK61('const.extreal.NegInf_0'),'T_0',sK61('const.extreal.NegInf_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f836,f601,f833,f3798,f512])).
fof(f833,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,sK61('const.extreal.NegInf_0'))) )),
  inference(unit_resulting_resolution,[],[f821,f518])).
fof(f5416,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4233,f4630])).
fof(f4630,plain,(
  ( ! [X0] : (~sP6(X0,'T_0',sK61('const.extreal.PosInf_0'))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f876,f3798,f528])).
fof(f876,plain,(
  s(bool,'T_0') = s(bool,sK61('const.extreal.PosInf_0'))),
  inference(unit_resulting_resolution,[],[f851,f404])).
fof(f4233,plain,(
  sP6(sK61('const.extreal.NegInf_0'),'T_0',sK61('const.extreal.PosInf_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f890,f601,f833,f3798,f512])).
fof(f5415,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4261,f4628])).
fof(f4261,plain,(
  sP6(sK61('const.extreal.PosInf_0'),'T_0',sK61('const.extreal.NegInf_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f836,f601,f887,f3798,f512])).
fof(f887,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,sK61('const.extreal.PosInf_0'))) )),
  inference(unit_resulting_resolution,[],[f851,f518])).
fof(f5414,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4262,f4630])).
fof(f4262,plain,(
  sP6(sK61('const.extreal.PosInf_0'),'T_0',sK61('const.extreal.PosInf_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f890,f601,f887,f3798,f512])).
fof(f5413,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4287,f4407])).
fof(f4287,plain,(
  ~sP5(sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),'F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f608,f1748,f1738,f3798,f512])).
fof(f5412,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4292,f4628])).
fof(f4292,plain,(
  sP6(sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),'T_0',sK61('const.extreal.NegInf_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f836,f601,f1738,f3798,f512])).
fof(f5411,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4293,f4630])).
fof(f4293,plain,(
  sP6(sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),'T_0',sK61('const.extreal.PosInf_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f890,f601,f1738,f3798,f512])).
fof(f5410,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4322,f4407])).
fof(f4322,plain,(
  ~sP5('T_0','F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f608,f1404,f605,f3798,f512])).
fof(f5409,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4323,f4407])).
fof(f4323,plain,(
  ~sP5('T_0','F_0',sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f1740,f1709,f605,f3798,f512])).
fof(f5408,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4324,f4407])).
fof(f4324,plain,(
  ~sP5(sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),'F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f608,f1748,f605,f3798,f512])).
fof(f5407,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4361,f1709])).
fof(f4361,plain,(
  sP6('T_0',sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f607,f1725,f605,f3798,f512])).
fof(f5406,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4364,f1710])).
fof(f1710,plain,(
  ~sP6(sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),'F_0')),
  inference(unit_resulting_resolution,[],[f819,f593,f819,f509])).
fof(f4364,plain,(
  sP6(sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f1739,f1725,f605,f3798,f512])).
fof(f1739,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) )),
  inference(unit_resulting_resolution,[],[f819,f522])).
fof(f5405,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4403,f4407])).
fof(f4403,plain,(
  ( ! [X0] : (~sP5('T_0','T_0',X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f355,f355,f3798,f513])).
fof(f5404,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4419,f4423])).
fof(f4419,plain,(
  ( ! [X0] : (~sP4('T_0','T_0',X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f355,f355,f3798,f517])).
fof(f5403,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4426,f3798])).
fof(f4426,plain,(
  p(s(bool,sK61('const.extreal.NegInf_0'))) | $spl4),
  inference(unit_resulting_resolution,[],[f835,f593,f3798,f521])).
fof(f521,plain,(
  ( ! [X2,X0,X1] : (~sP3(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f320])).
fof(f5402,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4427,f3798])).
fof(f4427,plain,(
  p(s(bool,sK61('const.extreal.PosInf_0'))) | $spl4),
  inference(unit_resulting_resolution,[],[f889,f593,f3798,f521])).
fof(f5401,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4428,f3798])).
fof(f4428,plain,(
  p(s(bool,sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl4),
  inference(unit_resulting_resolution,[],[f1740,f593,f3798,f521])).
fof(f5400,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4429,f3798])).
fof(f4429,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f609,f593,f3798,f521])).
fof(f609,plain,(
  ( ! [X0,X1] : (sP3('T_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f355,f524])).
fof(f5399,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4430,f3798])).
fof(f4430,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f836,f593,f3798,f521])).
fof(f5398,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4431,f3798])).
fof(f4431,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f890,f593,f3798,f521])).
fof(f5397,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4432,f3798])).
fof(f4432,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1741,f593,f3798,f521])).
fof(f5396,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4433,f3798])).
fof(f4433,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f607,f3798,f3798,f521])).
fof(f5395,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4434,f3798])).
fof(f4434,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f834,f3798,f3798,f521])).
fof(f834,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,sK61('const.extreal.NegInf_0'))) )),
  inference(unit_resulting_resolution,[],[f821,f522])).
fof(f5394,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4435,f3798])).
fof(f4435,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f888,f3798,f3798,f521])).
fof(f888,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,sK61('const.extreal.PosInf_0'))) )),
  inference(unit_resulting_resolution,[],[f851,f522])).
fof(f5393,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4436,f3798])).
fof(f4436,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1739,f3798,f3798,f521])).
fof(f5392,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4438,f3798])).
fof(f4438,plain,(
  p(s(bool,sK61('const.extreal.NegInf_0'))) | $spl4),
  inference(unit_resulting_resolution,[],[f835,f3798,f3798,f521])).
fof(f5391,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4439,f3798])).
fof(f4439,plain,(
  p(s(bool,sK61('const.extreal.PosInf_0'))) | $spl4),
  inference(unit_resulting_resolution,[],[f889,f3798,f3798,f521])).
fof(f5390,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4440,f3798])).
fof(f4440,plain,(
  p(s(bool,sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl4),
  inference(unit_resulting_resolution,[],[f1740,f3798,f3798,f521])).
fof(f5389,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4441,f3798])).
fof(f4441,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f609,f3798,f3798,f521])).
fof(f5388,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4442,f3798])).
fof(f4442,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f836,f3798,f3798,f521])).
fof(f5387,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4443,f3798])).
fof(f4443,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f890,f3798,f3798,f521])).
fof(f5386,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4444,f3798])).
fof(f4444,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1741,f3798,f3798,f521])).
fof(f5385,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4446,f3798])).
fof(f4446,plain,(
  p(s(bool,sK61('const.extreal.NegInf_0'))) | $spl4),
  inference(unit_resulting_resolution,[],[f834,f593,f3798,f521])).
fof(f5384,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4447,f3798])).
fof(f4447,plain,(
  p(s(bool,sK61('const.extreal.PosInf_0'))) | $spl4),
  inference(unit_resulting_resolution,[],[f888,f593,f3798,f521])).
fof(f5383,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4448,f3798])).
fof(f4448,plain,(
  p(s(bool,sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl4),
  inference(unit_resulting_resolution,[],[f1739,f593,f3798,f521])).
fof(f5382,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4449,f3798])).
fof(f4449,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f609,f593,f3798,f521])).
fof(f5381,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4450,f3798])).
fof(f4450,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f836,f593,f3798,f521])).
fof(f5380,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4451,f3798])).
fof(f4451,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f890,f593,f3798,f521])).
fof(f5379,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4452,f3798])).
fof(f4452,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1741,f593,f3798,f521])).
fof(f5378,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4458,f3798])).
fof(f4458,plain,(
  p(s(bool,sK61('const.extreal.NegInf_0'))) | $spl4),
  inference(unit_resulting_resolution,[],[f834,f3798,f3798,f521])).
fof(f5377,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4459,f3798])).
fof(f4459,plain,(
  p(s(bool,sK61('const.extreal.PosInf_0'))) | $spl4),
  inference(unit_resulting_resolution,[],[f888,f3798,f3798,f521])).
fof(f5376,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4460,f3798])).
fof(f4460,plain,(
  p(s(bool,sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl4),
  inference(unit_resulting_resolution,[],[f1739,f3798,f3798,f521])).
fof(f5375,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4461,f3798])).
fof(f4461,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f608,f3798,f3798,f521])).
fof(f5374,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4462,f3798])).
fof(f4462,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f835,f3798,f3798,f521])).
fof(f5373,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4463,f3798])).
fof(f4463,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f889,f3798,f3798,f521])).
fof(f5372,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4464,f3798])).
fof(f4464,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1740,f3798,f3798,f521])).
fof(f5371,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4465,f3798])).
fof(f4465,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f609,f3798,f3798,f521])).
fof(f5370,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4466,f3798])).
fof(f4466,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f836,f3798,f3798,f521])).
fof(f5369,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4467,f3798])).
fof(f4467,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f890,f3798,f3798,f521])).
fof(f5368,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4468,f3798])).
fof(f4468,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1741,f3798,f3798,f521])).
fof(f5367,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4474,f3798])).
fof(f4474,plain,(
  p(s(bool,sK61('const.extreal.NegInf_0'))) | $spl4),
  inference(unit_resulting_resolution,[],[f834,f593,f3798,f521])).
fof(f5366,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4475,f3798])).
fof(f4475,plain,(
  p(s(bool,sK61('const.extreal.PosInf_0'))) | $spl4),
  inference(unit_resulting_resolution,[],[f888,f593,f3798,f521])).
fof(f5365,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4476,f3798])).
fof(f4476,plain,(
  p(s(bool,sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl4),
  inference(unit_resulting_resolution,[],[f1739,f593,f3798,f521])).
fof(f5364,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4477,f3798])).
fof(f4477,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f608,f593,f3798,f521])).
fof(f5363,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4478,f3798])).
fof(f4478,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f835,f593,f3798,f521])).
fof(f5362,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4479,f3798])).
fof(f4479,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f889,f593,f3798,f521])).
fof(f5361,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4480,f3798])).
fof(f4480,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1740,f593,f3798,f521])).
fof(f5360,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4482,f3798])).
fof(f4482,plain,(
  p(s(bool,sK61('const.extreal.NegInf_0'))) | $spl4),
  inference(unit_resulting_resolution,[],[f834,f3798,f3798,f521])).
fof(f5359,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4483,f3798])).
fof(f4483,plain,(
  p(s(bool,sK61('const.extreal.PosInf_0'))) | $spl4),
  inference(unit_resulting_resolution,[],[f888,f3798,f3798,f521])).
fof(f5358,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4484,f3798])).
fof(f4484,plain,(
  p(s(bool,sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl4),
  inference(unit_resulting_resolution,[],[f1739,f3798,f3798,f521])).
fof(f5357,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4485,f3798])).
fof(f4485,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f608,f3798,f3798,f521])).
fof(f5356,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4486,f3798])).
fof(f4486,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f835,f3798,f3798,f521])).
fof(f5355,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4487,f3798])).
fof(f4487,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f889,f3798,f3798,f521])).
fof(f5354,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4488,f3798])).
fof(f4488,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1740,f3798,f3798,f521])).
fof(f5353,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4489,f3798])).
fof(f4489,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f609,f3798,f3798,f521])).
fof(f5352,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4490,f3798])).
fof(f4490,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f836,f3798,f3798,f521])).
fof(f5351,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4491,f3798])).
fof(f4491,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f890,f3798,f3798,f521])).
fof(f5350,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4492,f3798])).
fof(f4492,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1741,f3798,f3798,f521])).
fof(f5349,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4493,f3798])).
fof(f4493,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f608,f593,f3798,f521])).
fof(f5348,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4494,f3798])).
fof(f4494,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f835,f593,f3798,f521])).
fof(f5347,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4495,f3798])).
fof(f4495,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f889,f593,f3798,f521])).
fof(f5346,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4496,f3798])).
fof(f4496,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1740,f593,f3798,f521])).
fof(f5345,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4498,f3798])).
fof(f4498,plain,(
  p(s(bool,sK61('const.extreal.NegInf_0'))) | $spl4),
  inference(unit_resulting_resolution,[],[f836,f593,f3798,f521])).
fof(f5344,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4499,f3798])).
fof(f4499,plain,(
  p(s(bool,sK61('const.extreal.PosInf_0'))) | $spl4),
  inference(unit_resulting_resolution,[],[f890,f593,f3798,f521])).
fof(f5343,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4500,f3798])).
fof(f4500,plain,(
  p(s(bool,sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl4),
  inference(unit_resulting_resolution,[],[f1741,f593,f3798,f521])).
fof(f5342,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4505,f3798])).
fof(f4505,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f607,f3798,f3798,f521])).
fof(f5341,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4506,f3798])).
fof(f4506,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f834,f3798,f3798,f521])).
fof(f5340,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4507,f3798])).
fof(f4507,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f888,f3798,f3798,f521])).
fof(f5339,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4508,f3798])).
fof(f4508,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1739,f3798,f3798,f521])).
fof(f5338,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4509,f3798])).
fof(f4509,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f608,f3798,f3798,f521])).
fof(f5337,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4510,f3798])).
fof(f4510,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f835,f3798,f3798,f521])).
fof(f5336,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4511,f3798])).
fof(f4511,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f889,f3798,f3798,f521])).
fof(f5335,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4512,f3798])).
fof(f4512,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1740,f3798,f3798,f521])).
fof(f5334,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4514,f3798])).
fof(f4514,plain,(
  p(s(bool,sK61('const.extreal.NegInf_0'))) | $spl4),
  inference(unit_resulting_resolution,[],[f836,f3798,f3798,f521])).
fof(f5333,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4515,f3798])).
fof(f4515,plain,(
  p(s(bool,sK61('const.extreal.PosInf_0'))) | $spl4),
  inference(unit_resulting_resolution,[],[f890,f3798,f3798,f521])).
fof(f5332,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4516,f3798])).
fof(f4516,plain,(
  p(s(bool,sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl4),
  inference(unit_resulting_resolution,[],[f1741,f3798,f3798,f521])).
fof(f5331,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4521,f3798])).
fof(f4521,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f607,f593,f3798,f521])).
fof(f5330,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4522,f3798])).
fof(f4522,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f834,f593,f3798,f521])).
fof(f5329,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4523,f3798])).
fof(f4523,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f888,f593,f3798,f521])).
fof(f5328,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4524,f3798])).
fof(f4524,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1739,f593,f3798,f521])).
fof(f5327,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4526,f3798])).
fof(f4526,plain,(
  p(s(bool,sK61('const.extreal.NegInf_0'))) | $spl4),
  inference(unit_resulting_resolution,[],[f835,f593,f3798,f521])).
fof(f5326,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4527,f3798])).
fof(f4527,plain,(
  p(s(bool,sK61('const.extreal.PosInf_0'))) | $spl4),
  inference(unit_resulting_resolution,[],[f889,f593,f3798,f521])).
fof(f5325,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4528,f3798])).
fof(f4528,plain,(
  p(s(bool,sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl4),
  inference(unit_resulting_resolution,[],[f1740,f593,f3798,f521])).
fof(f5324,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4529,f3798])).
fof(f4529,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f607,f3798,f3798,f521])).
fof(f5323,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4530,f3798])).
fof(f4530,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f834,f3798,f3798,f521])).
fof(f5322,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4531,f3798])).
fof(f4531,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f888,f3798,f3798,f521])).
fof(f5321,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4532,f3798])).
fof(f4532,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1739,f3798,f3798,f521])).
fof(f5320,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4534,f3798])).
fof(f4534,plain,(
  p(s(bool,sK61('const.extreal.NegInf_0'))) | $spl4),
  inference(unit_resulting_resolution,[],[f835,f3798,f3798,f521])).
fof(f5319,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4535,f3798])).
fof(f4535,plain,(
  p(s(bool,sK61('const.extreal.PosInf_0'))) | $spl4),
  inference(unit_resulting_resolution,[],[f889,f3798,f3798,f521])).
fof(f5318,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4536,f3798])).
fof(f4536,plain,(
  p(s(bool,sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl4),
  inference(unit_resulting_resolution,[],[f1740,f3798,f3798,f521])).
fof(f5317,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4537,f3798])).
fof(f4537,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f609,f3798,f3798,f521])).
fof(f5316,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4538,f3798])).
fof(f4538,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f836,f3798,f3798,f521])).
fof(f5315,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4539,f3798])).
fof(f4539,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f890,f3798,f3798,f521])).
fof(f5314,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4540,f3798])).
fof(f4540,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1741,f3798,f3798,f521])).
fof(f5313,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4541,f3798])).
fof(f4541,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f607,f593,f3798,f521])).
fof(f5312,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4542,f3798])).
fof(f4542,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f834,f593,f3798,f521])).
fof(f5311,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4543,f3798])).
fof(f4543,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f888,f593,f3798,f521])).
fof(f5310,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4544,f3798])).
fof(f4544,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1739,f593,f3798,f521])).
fof(f5309,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4546,f3798])).
fof(f4546,plain,(
  p(s(bool,sK61('const.extreal.NegInf_0'))) | $spl4),
  inference(unit_resulting_resolution,[],[f836,f593,f3798,f521])).
fof(f5308,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4547,f3798])).
fof(f4547,plain,(
  p(s(bool,sK61('const.extreal.PosInf_0'))) | $spl4),
  inference(unit_resulting_resolution,[],[f890,f593,f3798,f521])).
fof(f5307,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4548,f3798])).
fof(f4548,plain,(
  p(s(bool,sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl4),
  inference(unit_resulting_resolution,[],[f1741,f593,f3798,f521])).
fof(f5306,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4553,f3798])).
fof(f4553,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f607,f3798,f3798,f521])).
fof(f5305,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4554,f3798])).
fof(f4554,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f834,f3798,f3798,f521])).
fof(f5304,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4555,f3798])).
fof(f4555,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f888,f3798,f3798,f521])).
fof(f5303,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4556,f3798])).
fof(f4556,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1739,f3798,f3798,f521])).
fof(f5302,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4557,f3798])).
fof(f4557,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f608,f3798,f3798,f521])).
fof(f5301,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4558,f3798])).
fof(f4558,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f835,f3798,f3798,f521])).
fof(f5300,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4559,f3798])).
fof(f4559,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f889,f3798,f3798,f521])).
fof(f5299,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4560,f3798])).
fof(f4560,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1740,f3798,f3798,f521])).
fof(f5298,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4562,f3798])).
fof(f4562,plain,(
  p(s(bool,sK61('const.extreal.NegInf_0'))) | $spl4),
  inference(unit_resulting_resolution,[],[f836,f3798,f3798,f521])).
fof(f5297,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4563,f3798])).
fof(f4563,plain,(
  p(s(bool,sK61('const.extreal.PosInf_0'))) | $spl4),
  inference(unit_resulting_resolution,[],[f890,f3798,f3798,f521])).
fof(f5296,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4564,f3798])).
fof(f4564,plain,(
  p(s(bool,sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl4),
  inference(unit_resulting_resolution,[],[f1741,f3798,f3798,f521])).
fof(f5295,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4569,f5121])).
fof(f4569,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f1402,f3798,f525])).
fof(f525,plain,(
  ( ! [X2,X0,X1] : (sP6(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f321])).
fof(f1402,plain,(
  ~sP6('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f355,f713,f527])).
fof(f5294,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4570,f5121])).
fof(f4570,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f1404,f3798,f525])).
fof(f5293,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4571,f5121])).
fof(f4571,plain,(
  s(bool,'F_0') = s(bool,sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f1709,f3798,f525])).
fof(f5292,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4582,f5121])).
fof(f4582,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f1748,f3798,f525])).
fof(f5291,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4583,f5121])).
fof(f4583,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f1749,f3798,f525])).
fof(f1749,plain,(
  ~sP6(sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f713,f819,f527])).
fof(f5290,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4645,f5050])).
fof(f5050,plain,(
  ( ! [X0,X1] : (sP8('T_0',X0,X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f610,f3798,f586])).
fof(f610,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f355,f537])).
fof(f537,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f326])).
fof(f4645,plain,(
  ( ! [X0,X1] : (~sP8('T_0',X0,X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f355,f3798,f530])).
fof(f530,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f324])).
fof(f5289,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4652,f5050])).
fof(f4652,plain,(
  ( ! [X0,X1] : (~sP8('T_0',X0,X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f355,f3798,f531])).
fof(f531,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f324])).
fof(f5288,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4662,f4666])).
fof(f4666,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f3798,f539])).
fof(f4662,plain,(
  ( ! [X0] : (~sP7('T_0','T_0',X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f355,f355,f3798,f536])).
fof(f5287,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4670,f3798])).
fof(f4670,plain,(
  p(s(bool,sK61('const.extreal.PosInf_0'))) | $spl4),
  inference(unit_resulting_resolution,[],[f891,f3798,f540])).
fof(f540,plain,(
  ( ! [X2,X0,X1] : (sP8(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f328])).
fof(f891,plain,(
  ( ! [X0] : (~sP8(sK61('const.extreal.PosInf_0'),X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f593,f851,f530])).
fof(f5286,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4671,f3798])).
fof(f4671,plain,(
  p(s(bool,sK61('const.extreal.PosInf_0'))) | $spl4),
  inference(unit_resulting_resolution,[],[f892,f3798,f540])).
fof(f892,plain,(
  ( ! [X0] : (~sP8(sK61('const.extreal.PosInf_0'),'F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f593,f851,f531])).
fof(f5285,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4672,f593])).
fof(f4672,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f1624,f3798,f540])).
fof(f1624,plain,(
  ~sP8('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f1617,f529])).
fof(f529,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | sP7(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f324])).
fof(f1617,plain,(
  ~sP7('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f355,f593,f355,f536])).
fof(f5284,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4677,f593])).
fof(f4677,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f856,f3798,f540])).
fof(f856,plain,(
  ( ! [X0] : (~sP8('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f355,f593,f530])).
fof(f5283,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4678,f3798])).
fof(f4678,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f859,f3798,f540])).
fof(f859,plain,(
  ( ! [X0] : (~sP8('T_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f355,f593,f531])).
fof(f5282,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4679,f593])).
fof(f4679,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f891,f3798,f540])).
fof(f5281,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4680,f3798])).
fof(f4680,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f892,f3798,f540])).
fof(f5280,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4688,f3798])).
fof(f4688,plain,(
  p(s(bool,sK61('const.extreal.PosInf_0'))) | $spl4),
  inference(unit_resulting_resolution,[],[f891,f3798,f541])).
fof(f541,plain,(
  ( ! [X2,X0,X1] : (sP8(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f328])).
fof(f5279,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4689,f3798])).
fof(f4689,plain,(
  p(s(bool,sK61('const.extreal.PosInf_0'))) | $spl4),
  inference(unit_resulting_resolution,[],[f892,f3798,f541])).
fof(f5278,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4690,f593])).
fof(f4690,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f1624,f3798,f541])).
fof(f5277,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4695,f3798])).
fof(f4695,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f856,f3798,f541])).
fof(f5276,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4696,f593])).
fof(f4696,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f859,f3798,f541])).
fof(f5275,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4697,f3798])).
fof(f4697,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f891,f3798,f541])).
fof(f5274,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4698,f593])).
fof(f4698,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f892,f3798,f541])).
fof(f5273,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4721,f4912])).
fof(f4912,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f3798,f3798,f3798,f557])).
fof(f557,plain,(
  ( ! [X2,X0,X1] : (sP10(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f336])).
fof(f336,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP10(X2,X0,X1)) & (~sP10(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f335])).
fof(f335,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP10(X2,X0,X1)) & (~sP10(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f237])).
fof(f237,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP10(X2,X0,X1))),
  inference(definition_folding,[],[f169,f236,f235])).
fof(f235,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f236,plain,(
  ! [X2,X0,X1] : (sP10(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP9(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f169,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f168])).
fof(f168,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f51])).
fof(f51,axiom,(
  ! [X24,X23,X55] : ((p(s(bool,X55)) <=> (p(s(bool,X23)) | p(s(bool,X24)))) <=> ((p(s(bool,X55)) | ~p(s(bool,X23))) & (p(s(bool,X55)) | ~p(s(bool,X24))) & (p(s(bool,X23)) | p(s(bool,X24)) | ~p(s(bool,X55)))))),
  file('thm.extreal.extreal_mul_ind','thm.sat.sat_25_prove_ltb')).
fof(f4721,plain,(
  ( ! [X0,X1] : (~sP10(X0,X1,'T_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f355,f3798,f546])).
fof(f546,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f331])).
fof(f331,plain,(
  ! [X0,X1,X2] : ((~sP10(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP9(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP9(X0,X1,X2) | sP10(X0,X1,X2)))),
  inference(rectify,[],[f330])).
fof(f330,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP9(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP9(X2,X0,X1) | sP10(X2,X0,X1)))),
  inference(flattening,[],[f329])).
fof(f329,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP9(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP9(X2,X0,X1)) | sP10(X2,X0,X1)))),
  inference(nnf_transformation,[],[f236])).
fof(f5272,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4730,f4912])).
fof(f4730,plain,(
  ( ! [X0,X1] : (~sP10(X0,'T_0',X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f355,f3798,f547])).
fof(f547,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f331])).
fof(f5271,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4758,f3798])).
fof(f4758,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f615,f906,f3798,f552])).
fof(f552,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~sP9(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f331])).
fof(f906,plain,(
  ( ! [X0] : (~sP10('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f355,f593,f547])).
fof(f615,plain,(
  ( ! [X0,X1] : (sP9('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f593,f556])).
fof(f556,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f334])).
fof(f334,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f333])).
fof(f333,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f332])).
fof(f332,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f235])).
fof(f5270,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4759,f3798])).
fof(f4759,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f614,f906,f3798,f552])).
fof(f614,plain,(
  ( ! [X0,X1] : (sP9(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f355,f555])).
fof(f555,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f334])).
fof(f5269,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4761,f3798])).
fof(f4761,plain,(
  p(s(bool,sK61('const.extreal.NegInf_0'))) | $spl4),
  inference(unit_resulting_resolution,[],[f838,f906,f3798,f552])).
fof(f838,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,sK61('const.extreal.NegInf_0'))) )),
  inference(unit_resulting_resolution,[],[f821,f554])).
fof(f554,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f334])).
fof(f5268,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4762,f3798])).
fof(f4762,plain,(
  p(s(bool,sK61('const.extreal.PosInf_0'))) | $spl4),
  inference(unit_resulting_resolution,[],[f894,f906,f3798,f552])).
fof(f894,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,sK61('const.extreal.PosInf_0'))) )),
  inference(unit_resulting_resolution,[],[f851,f554])).
fof(f5267,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4767,f4858])).
fof(f4858,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f3798,f556])).
fof(f4767,plain,(
  ~sP9('F_0','T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f906,f593,f3798,f552])).
fof(f5266,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4777,f4858])).
fof(f4777,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'T_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f903,f3798,f3798,f552])).
fof(f903,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f355,f593,f546])).
fof(f5265,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4779,f4858])).
fof(f4779,plain,(
  ( ! [X0] : (~sP9('F_0','T_0',X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f906,f3798,f3798,f552])).
fof(f5264,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4782,f3798])).
fof(f4782,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f903,f613,f3798,f552])).
fof(f613,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f355,f554])).
fof(f5263,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4787,f3798])).
fof(f4787,plain,(
  p(s(bool,sK61('const.extreal.NegInf_0'))) | $spl4),
  inference(unit_resulting_resolution,[],[f903,f839,f3798,f552])).
fof(f839,plain,(
  ( ! [X0,X1] : (sP9(X0,sK61('const.extreal.NegInf_0'),X1)) )),
  inference(unit_resulting_resolution,[],[f821,f555])).
fof(f5262,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4788,f3798])).
fof(f4788,plain,(
  p(s(bool,sK61('const.extreal.PosInf_0'))) | $spl4),
  inference(unit_resulting_resolution,[],[f903,f895,f3798,f552])).
fof(f895,plain,(
  ( ! [X0,X1] : (sP9(X0,sK61('const.extreal.PosInf_0'),X1)) )),
  inference(unit_resulting_resolution,[],[f851,f555])).
fof(f5261,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4789,f3798])).
fof(f4789,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f903,f615,f3798,f552])).
fof(f5260,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4791,f4858])).
fof(f4791,plain,(
  ~sP9('F_0','F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f903,f593,f3798,f552])).
fof(f5259,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4799,f4858])).
fof(f4799,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'T_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f903,f3798,f3798,f552])).
fof(f5258,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4800,f4858])).
fof(f4800,plain,(
  ( ! [X0] : (~sP9('F_0','T_0',X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f906,f3798,f3798,f552])).
fof(f5257,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4828,f3798])).
fof(f4828,plain,(
  p(s(bool,sK61('const.extreal.NegInf_0'))) | $spl4),
  inference(unit_resulting_resolution,[],[f838,f355,f3798,f553])).
fof(f553,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f334])).
fof(f5256,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4829,f3798])).
fof(f4829,plain,(
  p(s(bool,sK61('const.extreal.PosInf_0'))) | $spl4),
  inference(unit_resulting_resolution,[],[f894,f355,f3798,f553])).
fof(f5255,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4830,f3798])).
fof(f4830,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f614,f355,f3798,f553])).
fof(f5254,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4831,f3798])).
fof(f4831,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f839,f355,f3798,f553])).
fof(f5253,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4832,f3798])).
fof(f4832,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f895,f355,f3798,f553])).
fof(f5252,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4833,f4858])).
fof(f4833,plain,(
  ( ! [X0] : (~sP9('T_0',X0,'F_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f593,f355,f3798,f553])).
fof(f5251,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4835,f4858])).
fof(f4835,plain,(
  ( ! [X0,X1] : (~sP9('T_0',X0,X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f3798,f355,f3798,f553])).
fof(f5250,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4837,f3798])).
fof(f4837,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f613,f355,f3798,f553])).
fof(f5249,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4838,f3798])).
fof(f4838,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f838,f355,f3798,f553])).
fof(f5248,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4839,f3798])).
fof(f4839,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f894,f355,f3798,f553])).
fof(f5247,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4841,f3798])).
fof(f4841,plain,(
  p(s(bool,sK61('const.extreal.NegInf_0'))) | $spl4),
  inference(unit_resulting_resolution,[],[f839,f355,f3798,f553])).
fof(f5246,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4842,f3798])).
fof(f4842,plain,(
  p(s(bool,sK61('const.extreal.PosInf_0'))) | $spl4),
  inference(unit_resulting_resolution,[],[f895,f355,f3798,f553])).
fof(f5245,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4847,f4858])).
fof(f4847,plain,(
  ( ! [X0] : (~sP9('T_0','F_0',X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f355,f593,f3798,f553])).
fof(f5244,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4856,f4858])).
fof(f4856,plain,(
  ( ! [X0,X1] : (~sP9('T_0',X0,X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f355,f3798,f3798,f553])).
fof(f5243,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4861,f3798])).
fof(f4861,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f906,f593,f3798,f557])).
fof(f5242,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4863,f3798])).
fof(f4863,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f906,f3798,f3798,f557])).
fof(f5241,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4864,f593])).
fof(f4864,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f1898,f3798,f3798,f557])).
fof(f1898,plain,(
  ~sP10('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1884,f548])).
fof(f548,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f331])).
fof(f1884,plain,(
  ~sP9('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f593,f593,f355,f553])).
fof(f5240,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4865,f593])).
fof(f4865,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f906,f593,f3798,f557])).
fof(f5239,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4871,f593])).
fof(f4871,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f903,f3798,f3798,f557])).
fof(f5238,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4872,f593])).
fof(f4872,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f906,f3798,f3798,f557])).
fof(f5237,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4878,f3798])).
fof(f4878,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f903,f593,f3798,f557])).
fof(f5236,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4880,f3798])).
fof(f4880,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f903,f3798,f3798,f557])).
fof(f5235,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4882,f593])).
fof(f4882,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f1898,f3798,f3798,f557])).
fof(f5234,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4883,f593])).
fof(f4883,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f903,f593,f3798,f557])).
fof(f5233,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4889,f593])).
fof(f4889,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f903,f3798,f3798,f557])).
fof(f5232,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4890,f593])).
fof(f4890,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f906,f3798,f3798,f557])).
fof(f5231,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4897,f593])).
fof(f4897,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f1898,f593,f3798,f557])).
fof(f5230,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4898,f3798])).
fof(f4898,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f903,f3798,f3798,f557])).
fof(f5229,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4900,f593])).
fof(f4900,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f1898,f3798,f3798,f557])).
fof(f5228,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4902,f593])).
fof(f4902,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f1898,f593,f3798,f557])).
fof(f5227,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4908,f3798])).
fof(f4908,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f906,f3798,f3798,f557])).
fof(f5226,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4909,f593])).
fof(f4909,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f1898,f3798,f3798,f557])).
fof(f5225,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4915,f3798])).
fof(f4915,plain,(
  p(s(bool,sK61('const.extreal.NegInf_0'))) | $spl4),
  inference(unit_resulting_resolution,[],[f980,f355,f3798,f560])).
fof(f560,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)) | ~sP10(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f336])).
fof(f980,plain,(
  ( ! [X0] : (sP10('T_0',X0,sK61('const.extreal.NegInf_0'))) )),
  inference(unit_resulting_resolution,[],[f355,f838,f585])).
fof(f585,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP9(X0,X1,X2) | sP10(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f549])).
fof(f549,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~sP9(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f331])).
fof(f5224,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4916,f3798])).
fof(f4916,plain,(
  p(s(bool,sK61('const.extreal.PosInf_0'))) | $spl4),
  inference(unit_resulting_resolution,[],[f1349,f355,f3798,f560])).
fof(f1349,plain,(
  ( ! [X0] : (sP10('T_0',X0,sK61('const.extreal.PosInf_0'))) )),
  inference(unit_resulting_resolution,[],[f355,f894,f585])).
fof(f5223,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4917,f3798])).
fof(f4917,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f922,f355,f3798,f560])).
fof(f922,plain,(
  ( ! [X0] : (sP10('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f355,f355,f559])).
fof(f559,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP10(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f336])).
fof(f5222,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4918,f3798])).
fof(f4918,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1038,f355,f3798,f560])).
fof(f1038,plain,(
  ( ! [X0] : (sP10('T_0',sK61('const.extreal.NegInf_0'),X0)) )),
  inference(unit_resulting_resolution,[],[f355,f839,f585])).
fof(f5221,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4919,f3798])).
fof(f4919,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1366,f355,f3798,f560])).
fof(f1366,plain,(
  ( ! [X0] : (sP10('T_0',sK61('const.extreal.PosInf_0'),X0)) )),
  inference(unit_resulting_resolution,[],[f355,f895,f585])).
fof(f5220,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4946,f3798])).
fof(f4946,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f911,f355,f3798,f560])).
fof(f911,plain,(
  ( ! [X0] : (sP10('T_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f355,f355,f558])).
fof(f558,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP10(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f336])).
fof(f5219,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4947,f3798])).
fof(f4947,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f980,f355,f3798,f560])).
fof(f5218,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4948,f3798])).
fof(f4948,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1349,f355,f3798,f560])).
fof(f5217,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4950,f3798])).
fof(f4950,plain,(
  p(s(bool,sK61('const.extreal.NegInf_0'))) | $spl4),
  inference(unit_resulting_resolution,[],[f1038,f355,f3798,f560])).
fof(f5216,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4951,f3798])).
fof(f4951,plain,(
  p(s(bool,sK61('const.extreal.PosInf_0'))) | $spl4),
  inference(unit_resulting_resolution,[],[f1366,f355,f3798,f560])).
fof(f5215,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4956,f593])).
fof(f4956,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f939,f3798,f563])).
fof(f563,plain,(
  ( ! [X2,X0,X1] : (~sP12(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f339])).
fof(f339,plain,(
  ! [X0,X1,X2] : ((~sP12(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP11(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP11(X0,X1,X2) | sP12(X0,X1,X2)))),
  inference(rectify,[],[f338])).
fof(f338,plain,(
  ! [X2,X0,X1] : ((~sP12(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP11(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP11(X2,X0,X1) | sP12(X2,X0,X1)))),
  inference(flattening,[],[f337])).
fof(f337,plain,(
  ! [X2,X0,X1] : ((~sP12(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP11(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP11(X2,X0,X1)) | sP12(X2,X0,X1)))),
  inference(nnf_transformation,[],[f239])).
fof(f239,plain,(
  ! [X2,X0,X1] : (sP12(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP11(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP12])])).
fof(f939,plain,(
  ( ! [X0] : (sP12('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f593,f355,f575])).
fof(f575,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | sP12(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f344])).
fof(f344,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP12(X2,X0,X1)) & (~sP12(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f343])).
fof(f343,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP12(X2,X0,X1)) & (~sP12(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f240])).
fof(f240,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP12(X2,X0,X1))),
  inference(definition_folding,[],[f220,f239,f238])).
fof(f220,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f171])).
fof(f171,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f170])).
fof(f170,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X24,X23,X55] : ((p(s(bool,X55)) <=> (p(s(bool,X23)) => p(s(bool,X24)))) <=> ((p(s(bool,X55)) | p(s(bool,X23))) & (p(s(bool,X55)) | ~p(s(bool,X24))) & (~p(s(bool,X23)) | p(s(bool,X24)) | ~p(s(bool,X55)))))),
  file('thm.extreal.extreal_mul_ind','thm.sat.sat_27_prove_ltb')).
fof(f5214,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4957,f3798])).
fof(f4957,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f948,f3798,f563])).
fof(f948,plain,(
  ( ! [X0] : (sP12('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f355,f355,f576])).
fof(f576,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP12(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f344])).
fof(f5213,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4958,f3798])).
fof(f4958,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1040,f3798,f563])).
fof(f1040,plain,(
  ( ! [X0] : (sP12('T_0',sK61('const.extreal.NegInf_0'),X0)) )),
  inference(unit_resulting_resolution,[],[f355,f840,f584])).
fof(f584,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP11(X0,X1,X2) | sP12(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f566])).
fof(f566,plain,(
  ( ! [X2,X0,X1] : (sP12(X0,X1,X2) | ~sP11(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f339])).
fof(f840,plain,(
  ( ! [X0,X1] : (sP11(X0,sK61('const.extreal.NegInf_0'),X1)) )),
  inference(unit_resulting_resolution,[],[f821,f572])).
fof(f572,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f342])).
fof(f5212,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4959,f3798])).
fof(f4959,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1368,f3798,f563])).
fof(f1368,plain,(
  ( ! [X0] : (sP12('T_0',sK61('const.extreal.PosInf_0'),X0)) )),
  inference(unit_resulting_resolution,[],[f355,f896,f584])).
fof(f896,plain,(
  ( ! [X0,X1] : (sP11(X0,sK61('const.extreal.PosInf_0'),X1)) )),
  inference(unit_resulting_resolution,[],[f851,f572])).
fof(f5211,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4969,f593])).
fof(f4969,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f2149,f3798,f563])).
fof(f2149,plain,(
  sP12('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f618,f355,f593,f569])).
fof(f569,plain,(
  ( ! [X2,X0,X1] : (sP12(X0,X1,X2) | ~sP11(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f339])).
fof(f618,plain,(
  ( ! [X0,X1] : (sP11('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f593,f573])).
fof(f5210,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f4996,f5011])).
fof(f5011,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f3798,f573])).
fof(f4996,plain,(
  ~sP11('F_0','T_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f934,f355,f3798,f569])).
fof(f934,plain,(
  ( ! [X0] : (~sP12('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f355,f593,f564])).
fof(f564,plain,(
  ( ! [X2,X0,X1] : (~sP12(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f339])).
fof(f5209,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f5007,f5011])).
fof(f5007,plain,(
  ( ! [X0] : (~sP11('T_0',X0,'T_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f355,f355,f3798,f570])).
fof(f5208,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f5018,f593])).
fof(f5018,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f934,f355,f3798,f574])).
fof(f574,plain,(
  ( ! [X2,X0,X1] : (sP12(X2,X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f344])).
fof(f5207,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f5025,f593])).
fof(f5025,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f2166,f355,f3798,f574])).
fof(f2166,plain,(
  ~sP12('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f2159,f565])).
fof(f565,plain,(
  ( ! [X2,X0,X1] : (~sP12(X0,X1,X2) | sP11(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f339])).
fof(f2159,plain,(
  ~sP11('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f355,f355,f593,f570])).
fof(f5206,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f5055,f610])).
fof(f5055,plain,(
  ( ! [X0] : (~sP7(X0,'F_0','T_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f856,f3798,f586])).
fof(f5205,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f5057,f610])).
fof(f5057,plain,(
  ( ! [X0] : (~sP7('F_0',X0,'T_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f859,f3798,f586])).
fof(f5204,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f5059,f893])).
fof(f893,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,sK61('const.extreal.PosInf_0'))) )),
  inference(unit_resulting_resolution,[],[f851,f537])).
fof(f5059,plain,(
  ( ! [X0] : (~sP7(X0,'F_0',sK61('const.extreal.PosInf_0'))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f891,f3798,f586])).
fof(f5203,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f5061,f893])).
fof(f5061,plain,(
  ( ! [X0] : (~sP7('F_0',X0,sK61('const.extreal.PosInf_0'))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f892,f3798,f586])).
fof(f5202,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f5065,f3798])).
fof(f5065,plain,(
  p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f633,f3798,f596])).
fof(f596,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f475,f593])).
fof(f475,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f297])).
fof(f297,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP1(X0,X1)))),
  inference(flattening,[],[f296])).
fof(f296,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP1(X0,X1)))),
  inference(nnf_transformation,[],[f223])).
fof(f223,plain,(
  ! [X0,X1] : (sP1(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f633,plain,(
  ( ! [X0] : (sP1(X0,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) )),
  inference(unit_resulting_resolution,[],[f345,f476])).
fof(f476,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f297])).
fof(f5201,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f5066,f3798])).
fof(f5066,plain,(
  p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f648,f3798,f596])).
fof(f648,plain,(
  ( ! [X0] : (sP1(X0,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) )),
  inference(unit_resulting_resolution,[],[f346,f476])).
fof(f5200,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f5067,f3798])).
fof(f5067,plain,(
  p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f665,f3798,f596])).
fof(f665,plain,(
  ( ! [X0] : (sP1(X0,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) )),
  inference(unit_resulting_resolution,[],[f347,f476])).
fof(f5199,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f5068,f3798])).
fof(f5068,plain,(
  p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f678,f3798,f596])).
fof(f678,plain,(
  ( ! [X0] : (sP1(X0,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) )),
  inference(unit_resulting_resolution,[],[f348,f476])).
fof(f5198,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f5069,f3798])).
fof(f5069,plain,(
  ( ! [X0] : (p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'))))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f700,f3798,f596])).
fof(f700,plain,(
  ( ! [X0,X1] : (sP1(X0,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),'const.relation.EMPTY_REL_0')))) )),
  inference(unit_resulting_resolution,[],[f362,f476])).
fof(f362,plain,(
  ( ! [X0] : (p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'))))) )),
  inference(cnf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0] : p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X17] : p(s(bool,'const.relation.WF_1'(s(fun(X17,fun(X17,bool)),'const.relation.EMPTY_REL_0'))))),
  file('thm.extreal.extreal_mul_ind','thm.relation.WF_EMPTY_REL')).
fof(f5197,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f5070,f3798])).
fof(f5070,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.DATATYPE_1'(s(X0,X1))))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f623,f3798,f596])).
fof(f623,plain,(
  ( ! [X4,X5,X3] : (sP1(X5,'const.bool.DATATYPE_1'(s(X3,X4)))) )),
  inference(subsumption_resolution,[],[f621,f355])).
fof(f621,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'T_0')) | sP1(X5,'const.bool.DATATYPE_1'(s(X3,X4)))) )),
  inference(superposition,[],[f476,f425])).
fof(f5196,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f5071,f3798])).
fof(f5071,plain,(
  p(s(bool,'const.combin.I_1'(s(bool,'T_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f660,f3798,f596])).
fof(f660,plain,(
  ( ! [X0] : (sP1(X0,'const.combin.I_1'(s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f355,f627])).
fof(f627,plain,(
  ( ! [X2,X3] : (sP1(X3,'const.combin.I_1'(s(bool,X2))) | ~p(s(bool,X2))) )),
  inference(superposition,[],[f476,f426])).
fof(f5195,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f5072,f3798])).
fof(f5072,plain,(
  p(s(bool,sK61('const.extreal.NegInf_0'))) | $spl4),
  inference(unit_resulting_resolution,[],[f824,f3798,f596])).
fof(f824,plain,(
  ( ! [X0] : (sP1(X0,sK61('const.extreal.NegInf_0'))) )),
  inference(unit_resulting_resolution,[],[f821,f476])).
fof(f5194,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f5073,f3798])).
fof(f5073,plain,(
  p(s(bool,sK61('const.extreal.PosInf_0'))) | $spl4),
  inference(unit_resulting_resolution,[],[f878,f3798,f596])).
fof(f878,plain,(
  ( ! [X0] : (sP1(X0,sK61('const.extreal.PosInf_0'))) )),
  inference(unit_resulting_resolution,[],[f851,f476])).
fof(f5193,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f5074,f3798])).
fof(f5074,plain,(
  p(s(bool,sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl4),
  inference(unit_resulting_resolution,[],[f1699,f3798,f596])).
fof(f1699,plain,(
  ( ! [X0] : (sP1(X0,sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) )),
  inference(unit_resulting_resolution,[],[f819,f476])).
fof(f5192,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f5075,f3798])).
fof(f5075,plain,(
  p(s(bool,sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))))) | $spl4),
  inference(unit_resulting_resolution,[],[f1945,f3798,f596])).
fof(f1945,plain,(
  ( ! [X0] : (sP1(X0,sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))))) )),
  inference(unit_resulting_resolution,[],[f849,f476])).
fof(f5191,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f5076,f3798])).
fof(f5076,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f600,f3798,f596])).
fof(f600,plain,(
  ( ! [X0] : (sP1('T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f355,f477])).
fof(f477,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f297])).
fof(f5190,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f5077,f3798])).
fof(f5077,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f634,f3798,f596])).
fof(f634,plain,(
  ( ! [X0] : (sP1(happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')),X0)) )),
  inference(unit_resulting_resolution,[],[f345,f477])).
fof(f5189,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f5078,f3798])).
fof(f5078,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f649,f3798,f596])).
fof(f649,plain,(
  ( ! [X0] : (sP1(happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')),X0)) )),
  inference(unit_resulting_resolution,[],[f346,f477])).
fof(f5188,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f5079,f3798])).
fof(f5079,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f666,f3798,f596])).
fof(f666,plain,(
  ( ! [X0] : (sP1(happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')),X0)) )),
  inference(unit_resulting_resolution,[],[f347,f477])).
fof(f5187,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f5080,f3798])).
fof(f5080,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f679,f3798,f596])).
fof(f679,plain,(
  ( ! [X0] : (sP1(happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')),X0)) )),
  inference(unit_resulting_resolution,[],[f348,f477])).
fof(f5186,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f5081,f3798])).
fof(f5081,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f701,f3798,f596])).
fof(f701,plain,(
  ( ! [X0,X1] : (sP1('const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f362,f477])).
fof(f5185,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f5082,f3798])).
fof(f5082,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f622,f3798,f596])).
fof(f622,plain,(
  ( ! [X2,X0,X1] : (sP1('const.bool.DATATYPE_1'(s(X0,X1)),X2)) )),
  inference(subsumption_resolution,[],[f620,f355])).
fof(f620,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | sP1('const.bool.DATATYPE_1'(s(X0,X1)),X2)) )),
  inference(superposition,[],[f477,f425])).
fof(f5184,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f5083,f3798])).
fof(f5083,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f644,f3798,f596])).
fof(f644,plain,(
  ( ! [X0] : (sP1('const.combin.I_1'(s(bool,'T_0')),X0)) )),
  inference(unit_resulting_resolution,[],[f355,f626])).
fof(f626,plain,(
  ( ! [X0,X1] : (sP1('const.combin.I_1'(s(bool,X0)),X1) | ~p(s(bool,X0))) )),
  inference(superposition,[],[f477,f426])).
fof(f5183,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f5084,f3798])).
fof(f5084,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f825,f3798,f596])).
fof(f825,plain,(
  ( ! [X0] : (sP1(sK61('const.extreal.NegInf_0'),X0)) )),
  inference(unit_resulting_resolution,[],[f821,f477])).
fof(f5182,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f5085,f3798])).
fof(f5085,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f879,f3798,f596])).
fof(f879,plain,(
  ( ! [X0] : (sP1(sK61('const.extreal.PosInf_0'),X0)) )),
  inference(unit_resulting_resolution,[],[f851,f477])).
fof(f5181,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f5086,f3798])).
fof(f5086,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1700,f3798,f596])).
fof(f1700,plain,(
  ( ! [X0] : (sP1(sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),X0)) )),
  inference(unit_resulting_resolution,[],[f819,f477])).
fof(f5180,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f5087,f3798])).
fof(f5087,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1946,f3798,f596])).
fof(f1946,plain,(
  ( ! [X0] : (sP1(sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))),X0)) )),
  inference(unit_resulting_resolution,[],[f849,f477])).
fof(f5179,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f5092,f3798])).
fof(f5092,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f599,f3798,f596])).
fof(f599,plain,(
  ( ! [X0] : (sP1(X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f355,f476])).
fof(f5178,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f5093,f3798])).
fof(f5093,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f633,f3798,f596])).
fof(f5177,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f5094,f3798])).
fof(f5094,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f648,f3798,f596])).
fof(f5176,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f5095,f3798])).
fof(f5095,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f665,f3798,f596])).
fof(f5175,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f5096,f3798])).
fof(f5096,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f678,f3798,f596])).
fof(f5174,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f5097,f3798])).
fof(f5097,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f700,f3798,f596])).
fof(f5173,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f5098,f3798])).
fof(f5098,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f623,f3798,f596])).
fof(f5172,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f5099,f3798])).
fof(f5099,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f660,f3798,f596])).
fof(f5171,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f5100,f3798])).
fof(f5100,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f824,f3798,f596])).
fof(f5170,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f5101,f3798])).
fof(f5101,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f878,f3798,f596])).
fof(f5169,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f5102,f3798])).
fof(f5102,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1699,f3798,f596])).
fof(f5168,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f5103,f3798])).
fof(f5103,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1945,f3798,f596])).
fof(f5167,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f5105,f3798])).
fof(f5105,plain,(
  p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f634,f3798,f596])).
fof(f5166,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f5106,f3798])).
fof(f5106,plain,(
  p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f649,f3798,f596])).
fof(f5165,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f5107,f3798])).
fof(f5107,plain,(
  p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f666,f3798,f596])).
fof(f5164,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f5108,f3798])).
fof(f5108,plain,(
  p(s(bool,happ(s(fun('type.extreal.extreal',bool),happ(s(fun('type.extreal.extreal',fun('type.extreal.extreal',bool)),sK13),s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f679,f3798,f596])).
fof(f5163,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f5109,f3798])).
fof(f5109,plain,(
  ( ! [X0] : (p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'))))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f701,f3798,f596])).
fof(f5162,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f5110,f3798])).
fof(f5110,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.DATATYPE_1'(s(X0,X1))))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f622,f3798,f596])).
fof(f5161,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f5111,f3798])).
fof(f5111,plain,(
  p(s(bool,'const.combin.I_1'(s(bool,'T_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f644,f3798,f596])).
fof(f5160,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f5112,f3798])).
fof(f5112,plain,(
  p(s(bool,sK61('const.extreal.NegInf_0'))) | $spl4),
  inference(unit_resulting_resolution,[],[f825,f3798,f596])).
fof(f5159,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f5113,f3798])).
fof(f5113,plain,(
  p(s(bool,sK61('const.extreal.PosInf_0'))) | $spl4),
  inference(unit_resulting_resolution,[],[f879,f3798,f596])).
fof(f5158,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f5114,f3798])).
fof(f5114,plain,(
  p(s(bool,sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl4),
  inference(unit_resulting_resolution,[],[f1700,f3798,f596])).
fof(f5157,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f5115,f3798])).
fof(f5115,plain,(
  p(s(bool,sK61('const.combin.I_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))))) | $spl4),
  inference(unit_resulting_resolution,[],[f1946,f3798,f596])).
fof(f5129,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f363,f3798,f468])).
fof(f5127,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f595,f3798,f469])).
fof(f5122,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f595,f3798,f505])).
fof(f5119,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f599,f3798,f3798,f596])).
fof(f5117,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f599,f593,f3798,f596])).
fof(f5091,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f599,f3798,f3798,f596])).
fof(f5089,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f600,f593,f3798,f596])).
fof(f5062,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f612,f892,f3798,f586])).
fof(f612,plain,(
  ( ! [X0,X1] : (sP7('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f593,f539])).
fof(f5060,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f611,f891,f3798,f586])).
fof(f611,plain,(
  ( ! [X0,X1] : (sP7(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f593,f538])).
fof(f5058,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f612,f859,f3798,f586])).
fof(f5056,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f611,f856,f3798,f586])).
fof(f5049,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f948,f355,f355,f3798,f577])).
fof(f577,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~sP12(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f344])).
fof(f5039,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f2166,f355,f3798,f575])).
fof(f5034,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f934,f355,f3798,f3798,f574])).
fof(f5029,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f2166,f355,f593,f3798,f574])).
fof(f5023,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f934,f3798,f355,f3798,f574])).
fof(f5021,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f934,f593,f355,f3798,f574])).
fof(f5012,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f2159,f3798,f573])).
fof(f5010,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f2159,f3798,f571])).
fof(f5008,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f617,f355,f355,f3798,f570])).
fof(f617,plain,(
  ( ! [X0,X1] : (sP11(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f355,f572])).
fof(f4997,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f618,f934,f355,f3798,f569])).
fof(f4980,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f948,f355,f3798,f564])).
fof(f4973,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f939,f3798,f3798,f563])).
fof(f4971,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f2149,f593,f3798,f563])).
fof(f4964,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f939,f3798,f3798,f563])).
fof(f4962,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f939,f593,f3798,f563])).
fof(f4955,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f911,f3798,f355,f3798,f560])).
fof(f4953,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f911,f593,f355,f3798,f560])).
fof(f4934,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f911,f355,f3798,f3798,f560])).
fof(f4925,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f922,f355,f593,f3798,f560])).
fof(f4913,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f903,f3798,f3798,f3798,f557])).
fof(f4911,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f906,f593,f3798,f3798,f557])).
fof(f4906,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f903,f3798,f593,f3798,f557])).
fof(f4904,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1898,f593,f593,f3798,f557])).
fof(f4895,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f903,f3798,f3798,f3798,f557])).
fof(f4893,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f903,f593,f3798,f3798,f557])).
fof(f4888,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f903,f3798,f593,f3798,f557])).
fof(f4886,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f903,f593,f593,f3798,f557])).
fof(f4877,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f903,f3798,f3798,f3798,f557])).
fof(f4875,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f903,f593,f3798,f3798,f557])).
fof(f4870,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f906,f3798,f593,f3798,f557])).
fof(f4868,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f906,f593,f593,f3798,f557])).
fof(f4859,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1884,f3798,f556])).
fof(f4857,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f613,f355,f3798,f3798,f553])).
fof(f4848,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f613,f355,f593,f3798,f553])).
fof(f4836,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f613,f3798,f355,f3798,f553])).
fof(f4834,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f614,f593,f355,f3798,f553])).
fof(f4815,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f903,f615,f3798,f3798,f552])).
fof(f4813,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f903,f895,f3798,f3798,f552])).
fof(f4811,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f903,f839,f3798,f3798,f552])).
fof(f4809,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f906,f614,f3798,f3798,f552])).
fof(f4807,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f906,f894,f3798,f3798,f552])).
fof(f4805,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f906,f838,f3798,f3798,f552])).
fof(f4803,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f903,f613,f3798,f3798,f552])).
fof(f4798,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f903,f615,f593,f3798,f552])).
fof(f4794,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f903,f613,f593,f3798,f552])).
fof(f4780,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f615,f906,f3798,f3798,f552])).
fof(f4778,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f615,f903,f3798,f3798,f552])).
fof(f4768,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f615,f906,f593,f3798,f552])).
fof(f4731,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f922,f355,f3798,f547])).
fof(f4722,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f911,f355,f3798,f546])).
fof(f4713,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1606,f355,f355,f3798,f545])).
fof(f1606,plain,(
  sP8('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f610,f355,f355,f532])).
fof(f532,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP7(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f324])).
fof(f4703,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f856,f3798,f3798,f541])).
fof(f4701,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f859,f593,f3798,f541])).
fof(f4694,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f856,f3798,f3798,f541])).
fof(f4692,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1624,f593,f3798,f541])).
fof(f4685,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f856,f3798,f3798,f540])).
fof(f4683,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f856,f593,f3798,f540])).
fof(f4676,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f856,f3798,f3798,f540])).
fof(f4674,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1624,f593,f3798,f540])).
fof(f4667,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1617,f3798,f539])).
fof(f4665,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1617,f3798,f538])).
fof(f4663,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f610,f355,f355,f3798,f536])).
fof(f4653,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1606,f355,f3798,f531])).
fof(f4646,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1606,f355,f3798,f530])).
fof(f4639,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1396,f426,f3798,f528])).
fof(f1396,plain,(
  ( ! [X0] : (sP6('T_0',X0,'const.combin.I_1'(s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f355,f426,f526])).
fof(f526,plain,(
  ( ! [X2,X0,X1] : (sP6(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f321])).
fof(f4631,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f2376,f876,f3798,f528])).
fof(f2376,plain,(
  sP6('T_0','T_0',sK61('const.extreal.PosInf_0'))),
  inference(unit_resulting_resolution,[],[f355,f876,f526])).
fof(f4629,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f2199,f822,f3798,f528])).
fof(f2199,plain,(
  sP6('T_0','T_0',sK61('const.extreal.NegInf_0'))),
  inference(unit_resulting_resolution,[],[f355,f822,f526])).
fof(f4627,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1395,f425,f3798,f528])).
fof(f1395,plain,(
  ( ! [X0,X1] : (sP6('T_0','T_0','const.bool.DATATYPE_1'(s(X0,X1)))) )),
  inference(unit_resulting_resolution,[],[f355,f425,f526])).
fof(f4625,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1394,f426,f3798,f528])).
fof(f1394,plain,(
  ( ! [X0] : (sP6('T_0','const.combin.I_1'(s(bool,X0)),X0)) )),
  inference(unit_resulting_resolution,[],[f355,f426,f526])).
fof(f4617,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f2374,f876,f3798,f528])).
fof(f2374,plain,(
  sP6('T_0',sK61('const.extreal.PosInf_0'),'T_0')),
  inference(unit_resulting_resolution,[],[f355,f876,f526])).
fof(f4615,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f2197,f822,f3798,f528])).
fof(f2197,plain,(
  sP6('T_0',sK61('const.extreal.NegInf_0'),'T_0')),
  inference(unit_resulting_resolution,[],[f355,f822,f526])).
fof(f4613,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1393,f425,f3798,f528])).
fof(f1393,plain,(
  ( ! [X0,X1] : (sP6('T_0','const.bool.DATATYPE_1'(s(X0,X1)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f355,f425,f526])).
fof(f5147,plain,(
  $false | $spl4),
  inference(trivial_inequality_removal,[],[f4607])).
fof(f4607,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f2745,f3798,f528])).
fof(f2745,plain,(
  sP6(sK61('const.extreal.PosInf_0'),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f888,f602,f605,f593,f511])).
fof(f5146,plain,(
  $false | $spl4),
  inference(trivial_inequality_removal,[],[f4606])).
fof(f4606,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f2744,f3798,f528])).
fof(f2744,plain,(
  sP6(sK61('const.extreal.NegInf_0'),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f834,f602,f605,f593,f511])).
fof(f5145,plain,(
  $false | $spl4),
  inference(trivial_inequality_removal,[],[f4605])).
fof(f4605,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f2690,f3798,f528])).
fof(f2690,plain,(
  sP6(sK61('const.extreal.PosInf_0'),'T_0','T_0')),
  inference(unit_resulting_resolution,[],[f609,f601,f887,f355,f510])).
fof(f510,plain,(
  ( ! [X2,X0,X1] : (~sP4(X1,X2,X0) | ~p(s(bool,X2)) | ~sP5(X0,X2,X1) | sP6(X0,X1,X2) | ~sP3(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f312])).
fof(f5144,plain,(
  $false | $spl4),
  inference(trivial_inequality_removal,[],[f4604])).
fof(f4604,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f2686,f3798,f528])).
fof(f2686,plain,(
  sP6(sK61('const.extreal.NegInf_0'),'T_0','T_0')),
  inference(unit_resulting_resolution,[],[f609,f601,f833,f355,f510])).
fof(f5143,plain,(
  $false | $spl4),
  inference(trivial_inequality_removal,[],[f4599])).
fof(f4599,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f2743,f3798,f528])).
fof(f2743,plain,(
  sP6('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f607,f602,f605,f593,f511])).
fof(f5142,plain,(
  $false | $spl4),
  inference(trivial_inequality_removal,[],[f4598])).
fof(f4598,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f2682,f3798,f528])).
fof(f2682,plain,(
  sP6('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f609,f601,f604,f355,f510])).
fof(f4587,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1404,f713,f3798,f525])).
fof(f4585,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1402,f713,f3798,f525])).
fof(f4568,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f607,f3798,f3798,f3798,f521])).
fof(f4566,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f607,f593,f3798,f3798,f521])).
fof(f4552,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f607,f3798,f593,f3798,f521])).
fof(f4550,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f607,f593,f593,f3798,f521])).
fof(f4520,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f607,f3798,f3798,f3798,f521])).
fof(f4518,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f607,f593,f3798,f3798,f521])).
fof(f4504,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f608,f3798,f593,f3798,f521])).
fof(f4502,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f608,f593,f593,f3798,f521])).
fof(f4472,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f607,f3798,f3798,f3798,f521])).
fof(f4470,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f608,f593,f3798,f3798,f521])).
fof(f4456,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f607,f3798,f593,f3798,f521])).
fof(f4454,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f609,f593,f593,f3798,f521])).
fof(f4424,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1572,f3798,f520])).
fof(f1572,plain,(
  ~sP4('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f355,f593,f355,f517])).
fof(f4422,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1572,f3798,f519])).
fof(f4420,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f604,f355,f355,f3798,f517])).
fof(f4408,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1553,f3798,f516])).
fof(f1553,plain,(
  ~sP5('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f355,f593,f355,f513])).
fof(f4406,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1553,f3798,f515])).
fof(f4404,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1480,f355,f355,f3798,f513])).
fof(f1480,plain,(
  ( ! [X0,X1] : (sP5('T_0','T_0','const.bool.DATATYPE_1'(s(X0,X1)))) )),
  inference(unit_resulting_resolution,[],[f1393,f508])).
fof(f508,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | sP5(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f312])).
fof(f4366,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1740,f1709,f1725,f605,f3798,f512])).
fof(f4349,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f608,f1748,f601,f605,f3798,f512])).
fof(f4347,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f608,f1404,f601,f605,f3798,f512])).
fof(f4340,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f608,f1748,f602,f605,f3798,f512])).
fof(f4336,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1740,f1709,f602,f605,f3798,f512])).
fof(f4334,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f608,f1404,f602,f605,f3798,f512])).
fof(f4321,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f608,f1748,f602,f1738,f3798,f512])).
fof(f4296,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f608,f1748,f601,f1738,f3798,f512])).
fof(f4228,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1740,f1709,f602,f604,f3798,f512])).
fof(f4226,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f608,f1404,f602,f604,f3798,f512])).
fof(f4219,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1740,f1709,f1725,f604,f3798,f512])).
fof(f4199,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f608,f1404,f601,f604,f3798,f512])).
fof(f4124,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1740,f1709,f1725,f605,f3798,f511])).
fof(f4107,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f608,f1748,f601,f605,f3798,f511])).
fof(f4105,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f608,f1404,f601,f605,f3798,f511])).
fof(f4098,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f608,f1748,f602,f605,f3798,f511])).
fof(f4094,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1740,f1709,f602,f605,f3798,f511])).
fof(f4092,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f608,f1404,f602,f605,f3798,f511])).
fof(f4079,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f608,f1748,f602,f1738,f3798,f511])).
fof(f4054,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f608,f1748,f601,f1738,f3798,f511])).
fof(f3986,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1740,f1709,f602,f604,f3798,f511])).
fof(f3984,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f608,f1404,f602,f604,f3798,f511])).
fof(f3977,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1740,f1709,f1725,f604,f3798,f511])).
fof(f3957,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f608,f1404,f601,f604,f3798,f511])).
fof(f3903,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1395,f355,f355,f3798,f509])).
fof(f3883,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f766,f355,f3798,f495])).
fof(f766,plain,(
  ~sP2('T_0','T_0')),
  inference(unit_resulting_resolution,[],[f355,f355,f491])).
fof(f491,plain,(
  ( ! [X0,X1] : (~sP2(X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f304])).
fof(f3867,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f766,f355,f3798,f492])).
fof(f3851,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f769,f3798,f3798,f490])).
fof(f3849,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f775,f593,f3798,f490])).
fof(f3834,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f769,f3798,f3798,f490])).
fof(f3832,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f769,f593,f3798,f490])).
fof(f3817,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f593,f713,f3798,f471])).
fof(f3815,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f3798,f713,f3798,f471])).
fof(f3811,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f593,f713,f3798,f471])).
fof(f3809,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f3798,f713,f3798,f471])).
fof(f3805,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f713,f3798,f408])).
fof(f3803,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f425,f3798,f403])).
fof(f3801,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f355,f3798])).
fof(f3800,plain,(
  p(s(bool,sK49('const.combin.I_0',bool))) | $spl6),
  inference(cnf_transformation,[],[f3800_D])).
fof(f3800_D,plain,(
  p(s(bool,sK49('const.combin.I_0',bool))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f3797,plain,(
  ( ! [X3] : (p(s(bool,sK49('const.combin.I_0',bool))) | ~p(s(bool,X3))) )),
  inference(forward_demodulation,[],[f3787,f595])).
fof(f3787,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | p(s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,sK49('const.combin.I_0',bool)))))) )),
  inference(superposition,[],[f430,f595])).
fof(f430,plain,(
  ( ! [X4,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK49(X1,X0)))))) )),
  inference(cnf_transformation,[],[f276])).
fof(f3422,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'const.combin.I_0')))) | $spl0),
  inference(cnf_transformation,[],[f3422_D])).
fof(f3422_D,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'const.combin.I_0')))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f3423,plain,(
  ( ! [X3] : (~p(s(bool,sK51(X3,'const.combin.I_0',bool)))) ) | $spl2),
  inference(cnf_transformation,[],[f3423_D])).
fof(f3423_D,plain,(
  ( ! [X3] : (~p(s(bool,sK51(X3,'const.combin.I_0',bool)))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f3413,plain,(
  ( ! [X3] : (~p(s(bool,sK51(X3,'const.combin.I_0',bool))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'const.combin.I_0'))))) )),
  inference(superposition,[],[f435,f595])).
% SZS output end Proof for thm.extreal
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 1638
% Passive clauses: 17485
% Generated clauses: 74403
% Final active clauses: 1538
% Final passive clauses: 7041
% Input formulas: 69
% Initial clauses: 239
% Discarded non-redundant clauses: 15193
% 
% Duplicate literals: 90
% Trivial inequalities: 27
% Fw subsumption resolutions: 840
% Fw demodulations: 10984
% Bw demodulations: 1065
% Condensations: 9
% 
% Simple tautologies: 49
% Equational tautologies: 50
% Forward subsumptions: 37582
% Bw demodulations to eq. taut.: 8
% 
% Unit resulting resolution: 47243
% Factoring: 1
% Forward superposition: 9548
% Backward superposition: 3750
% Self superposition: 59
% Equality factoring: 2
% Equality resolution: 28
% 
% Split clauses: 13
% Split components: 26
% SAT solver clauses: 18003
% SAT solver unit clauses: 9631
% SAT solver binary clauses: 5778
% 
% Sat splits: 13
% Sat splitting refutations: 490
% 
% Memory used [KB]: 55393
% Time elapsed: 6.034 s
% ------------------------------
% ------------------------------
% Success in time 7.371 s
