Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.ringNorm
% (79890)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 120
% Passive clauses: 3991
% Generated clauses: 14280
% Final active clauses: 112
% Final passive clauses: 3588
% Input formulas: 69
% Initial clauses: 125
% 
% Duplicate literals: 521
% Trivial inequalities: 211
% Fw subsumption resolutions: 1487
% Fw demodulations: 477
% Bw demodulations: 6
% 
% Simple tautologies: 297
% Equational tautologies: 90
% Forward subsumptions: 6185
% 
% Binary resolution: 282
% Factoring: 2
% Forward superposition: 4337
% Backward superposition: 6223
% Self superposition: 291
% Equality factoring: 197
% Equality resolution: 34
% 
% Split clauses: 487
% Split components: 1022
% SAT solver clauses: 1067
% SAT solver unit clauses: 42
% SAT solver binary clauses: 379
% 
% Sat splits: 487
% 
% TWLsolver clauses: 1041
% TWLsolver calls for satisfiability: 139
% 
% Memory used [KB]: 11897
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.ringNorm
% (79903)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 133
% Passive clauses: 1073
% Generated clauses: 11794
% Final active clauses: 123
% Final passive clauses: 917
% Input formulas: 69
% Initial clauses: 125
% 
% Duplicate literals: 201
% Trivial inequalities: 121
% Fw subsumption resolutions: 285
% Fw demodulations: 1527
% Bw demodulations: 31
% Condensations: 436
% Global subsumptions: 12
% 
% Simple tautologies: 90
% Equational tautologies: 41
% Forward subsumptions: 6192
% 
% Binary resolution: 56
% Factoring: 1
% Forward superposition: 3480
% Backward superposition: 4218
% Self superposition: 132
% Equality factoring: 50
% Equality resolution: 22
% 
% Split clauses: 15
% Split components: 30
% SAT solver clauses: 1923
% SAT solver unit clauses: 158
% SAT solver binary clauses: 54
% SAT solver learnt clauses: 24
% 
% Sat splits: 15
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 1875
% TWLsolver calls for satisfiability: 410
% 
% Memory used [KB]: 27632
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.ringNorm
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.ringNorm
% SZS output start Proof for thm.ringNorm
fof(f46238,plain,(
  $false),
  inference(subsumption_resolution,[],[f46237,f326])).
fof(f326,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.ringNorm.polynom'(X0),'const.ringNorm.Pconst_1'(s(X0,X3))) != s('type.ringNorm.polynom'(X0),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X0),X2),s('type.ringNorm.polynom'(X0),X1)))) )),
  inference(cnf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X0,X1,X2,X3] : s('type.ringNorm.polynom'(X0),'const.ringNorm.Pconst_1'(s(X0,X3))) != s('type.ringNorm.polynom'(X0),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X0),X2),s('type.ringNorm.polynom'(X0),X1)))),
  inference(flattening,[],[f142])).
fof(f142,plain,(
  ! [X0] : ! [X1,X2,X3] : ~s('type.ringNorm.polynom'(X0),'const.ringNorm.Pconst_1'(s(X0,X3))) = s('type.ringNorm.polynom'(X0),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X0),X2),s('type.ringNorm.polynom'(X0),X1)))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X6] : ! [X38,X39,X8] : ~s('type.ringNorm.polynom'(X6),'const.ringNorm.Pconst_1'(s(X6,X8))) = s('type.ringNorm.polynom'(X6),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X6),X39),s('type.ringNorm.polynom'(X6),X38)))),
  file('thm.ringNorm.polynom_distinct|split|6','thm.ringNorm.polynom_distinct|split|5')).
fof(f46237,plain,(
  ( ! [X60] : (s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pconst_1'(s(sK0,sK2))) = s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(sK0),X60),s('type.ringNorm.polynom'(sK0),sK1)))) )),
  inference(forward_demodulation,[],[f45247,f46236])).
fof(f46236,plain,(
  ( ! [X58] : (s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pconst_1'(s(sK0,sK2))) = s('type.ringNorm.polynom'(sK0),happ(s(fun(sK0,'type.ringNorm.polynom'(sK0)),X58),s(sK0,sK2)))) )),
  inference(forward_demodulation,[],[f45245,f245])).
fof(f245,plain,(
  s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pconst_1'(s(sK0,sK2))) = s('type.ringNorm.polynom'(sK0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(sK0),sK1)))),
  inference(cnf_transformation,[],[f201])).
fof(f201,plain,(
  s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pconst_1'(s(sK0,sK2))) = s('type.ringNorm.polynom'(sK0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(sK0),sK1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f184])).
fof(f184,plain,(
  ? [X0,X1,X2] : s('type.ringNorm.polynom'(X0),'const.ringNorm.Pconst_1'(s(X0,X2))) = s('type.ringNorm.polynom'(X0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X0),X1)))),
  inference(ennf_transformation,[],[f72])).
fof(f72,plain,(
  ~! [X0,X1,X2] : s('type.ringNorm.polynom'(X0),'const.ringNorm.Pconst_1'(s(X0,X2))) != s('type.ringNorm.polynom'(X0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X0),X1)))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ~! [X0] : ! [X1,X2] : ~s('type.ringNorm.polynom'(X0),'const.ringNorm.Pconst_1'(s(X0,X2))) = s('type.ringNorm.polynom'(X0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X0),X1)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : ~s('type.ringNorm.polynom'(X6),'const.ringNorm.Pconst_1'(s(X6,X8))) = s('type.ringNorm.polynom'(X6),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X6),X7)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : ~s('type.ringNorm.polynom'(X6),'const.ringNorm.Pconst_1'(s(X6,X8))) = s('type.ringNorm.polynom'(X6),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X6),X7)))),
  file('thm.ringNorm.polynom_distinct|split|6','thm.ringNorm.polynom_distinct|split|6')).
fof(f45245,plain,(
  ( ! [X58] : (s('type.ringNorm.polynom'(sK0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(sK0),sK1))) = s('type.ringNorm.polynom'(sK0),happ(s(fun(sK0,'type.ringNorm.polynom'(sK0)),X58),s(sK0,sK2)))) )),
  inference(superposition,[],[f23383,f298])).
fof(f298,plain,(
  ( ! [X0,X1] : (s('type.ringNorm.polynom'(X0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X0),X1))) = s('type.ringNorm.polynom'(X0),happ(s(fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0)),'const.ringNorm.Popp_0'),s('type.ringNorm.polynom'(X0),X1)))) )),
  inference(cnf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0,X1] : s('type.ringNorm.polynom'(X0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X0),X1))) = s('type.ringNorm.polynom'(X0),happ(s(fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0)),'const.ringNorm.Popp_0'),s('type.ringNorm.polynom'(X0),X1)))),
  inference(flattening,[],[f101])).
fof(f101,plain,(
  ! [X0] : ! [X1] : s('type.ringNorm.polynom'(X0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X0),X1))) = s('type.ringNorm.polynom'(X0),happ(s(fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0)),'const.ringNorm.Popp_0'),s('type.ringNorm.polynom'(X0),X1)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6] : ! [X9] : s('type.ringNorm.polynom'(X6),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X6),X9))) = s('type.ringNorm.polynom'(X6),happ(s(fun('type.ringNorm.polynom'(X6),'type.ringNorm.polynom'(X6)),'const.ringNorm.Popp_0'),s('type.ringNorm.polynom'(X6),X9)))),
  file('thm.ringNorm.polynom_distinct|split|6','ARITY_0_1_const.ringNorm.Popp')).
fof(f23383,plain,(
  ( ! [X6,X8,X11] : (s(X6,happ(s(fun(sK0,X6),X8),s(sK0,sK2))) = s(X6,happ(s(fun('type.ringNorm.polynom'(sK0),X6),X11),s('type.ringNorm.polynom'(sK0),sK1)))) )),
  inference(superposition,[],[f8396,f348])).
fof(f348,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),X4),s(X1,X2))) = s(X0,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X1),'const.ringNorm.Pconst_1'(s(X1,X2))),s(fun('type.quote.index',X0),X3),s(fun(X1,X0),X4),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X5),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X6),s(fun('type.ringNorm.polynom'(X1),X0),X7)))) )),
  inference(cnf_transformation,[],[f171])).
fof(f171,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun(X1,X0),X4),s(X1,X2))) = s(X0,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X1),'const.ringNorm.Pconst_1'(s(X1,X2))),s(fun('type.quote.index',X0),X3),s(fun(X1,X0),X4),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X5),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X6),s(fun('type.ringNorm.polynom'(X1),X0),X7)))),
  inference(flattening,[],[f170])).
fof(f170,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun(X1,X0),X4),s(X1,X2))) = s(X0,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X1),'const.ringNorm.Pconst_1'(s(X1,X2))),s(fun('type.quote.index',X0),X3),s(fun(X1,X0),X4),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X5),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X6),s(fun('type.ringNorm.polynom'(X1),X0),X7)))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X25,X6] : ! [X8,X26,X34,X35,X36,X37] : s(X25,happ(s(fun(X6,X25),X34),s(X6,X8))) = s(X25,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X6),'const.ringNorm.Pconst_1'(s(X6,X8))),s(fun('type.quote.index',X25),X26),s(fun(X6,X25),X34),s(fun('type.ringNorm.polynom'(X6),fun('type.ringNorm.polynom'(X6),X25)),X35),s(fun('type.ringNorm.polynom'(X6),fun('type.ringNorm.polynom'(X6),X25)),X36),s(fun('type.ringNorm.polynom'(X6),X25),X37)))),
  file('thm.ringNorm.polynom_distinct|split|6','thm.ringNorm.polynom_case_def|split|1')).
fof(f8396,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.ringNorm.polynom'(sK0),X0),X1),s('type.ringNorm.polynom'(sK0),sK1))) = s(X0,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pconst_1'(s(sK0,sK2))),s(fun('type.quote.index',X0),X2),s(fun(sK0,X0),X3),s(fun('type.ringNorm.polynom'(sK0),fun('type.ringNorm.polynom'(sK0),X0)),X4),s(fun('type.ringNorm.polynom'(sK0),fun('type.ringNorm.polynom'(sK0),X0)),X5),s(fun('type.ringNorm.polynom'(sK0),X0),X1)))) )),
  inference(superposition,[],[f349,f245])).
fof(f349,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun('type.ringNorm.polynom'(X1),X0),X7),s('type.ringNorm.polynom'(X1),X2))) = s(X0,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X1),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X1),X2))),s(fun('type.quote.index',X0),X3),s(fun(X1,X0),X4),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X5),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X6),s(fun('type.ringNorm.polynom'(X1),X0),X7)))) )),
  inference(cnf_transformation,[],[f173])).
fof(f173,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun('type.ringNorm.polynom'(X1),X0),X7),s('type.ringNorm.polynom'(X1),X2))) = s(X0,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X1),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X1),X2))),s(fun('type.quote.index',X0),X3),s(fun(X1,X0),X4),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X5),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X6),s(fun('type.ringNorm.polynom'(X1),X0),X7)))),
  inference(flattening,[],[f172])).
fof(f172,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun('type.ringNorm.polynom'(X1),X0),X7),s('type.ringNorm.polynom'(X1),X2))) = s(X0,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X1),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X1),X2))),s(fun('type.quote.index',X0),X3),s(fun(X1,X0),X4),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X5),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X6),s(fun('type.ringNorm.polynom'(X1),X0),X7)))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X25,X6] : ! [X8,X26,X34,X35,X36,X37] : s(X25,happ(s(fun('type.ringNorm.polynom'(X6),X25),X37),s('type.ringNorm.polynom'(X6),X8))) = s(X25,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X6),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X6),X8))),s(fun('type.quote.index',X25),X26),s(fun(X6,X25),X34),s(fun('type.ringNorm.polynom'(X6),fun('type.ringNorm.polynom'(X6),X25)),X35),s(fun('type.ringNorm.polynom'(X6),fun('type.ringNorm.polynom'(X6),X25)),X36),s(fun('type.ringNorm.polynom'(X6),X25),X37)))),
  file('thm.ringNorm.polynom_distinct|split|6','thm.ringNorm.polynom_case_def|split|4')).
fof(f45247,plain,(
  ( ! [X61,X60] : (s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(sK0),X60),s('type.ringNorm.polynom'(sK0),sK1))) = s('type.ringNorm.polynom'(sK0),happ(s(fun(sK0,'type.ringNorm.polynom'(sK0)),X61),s(sK0,sK2)))) )),
  inference(superposition,[],[f23383,f307])).
fof(f307,plain,(
  ( ! [X2,X0,X1] : (s('type.ringNorm.polynom'(X0),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X0),X1),s('type.ringNorm.polynom'(X0),X2))) = s('type.ringNorm.polynom'(X0),happ(s(fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0)),happ(s(fun('type.ringNorm.polynom'(X0),fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0))),'const.ringNorm.Pmult_0'),s('type.ringNorm.polynom'(X0),X1))),s('type.ringNorm.polynom'(X0),X2)))) )),
  inference(cnf_transformation,[],[f115])).
fof(f115,plain,(
  ! [X0,X1,X2] : s('type.ringNorm.polynom'(X0),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X0),X1),s('type.ringNorm.polynom'(X0),X2))) = s('type.ringNorm.polynom'(X0),happ(s(fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0)),happ(s(fun('type.ringNorm.polynom'(X0),fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0))),'const.ringNorm.Pmult_0'),s('type.ringNorm.polynom'(X0),X1))),s('type.ringNorm.polynom'(X0),X2)))),
  inference(flattening,[],[f114])).
fof(f114,plain,(
  ! [X0] : ! [X1,X2] : s('type.ringNorm.polynom'(X0),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X0),X1),s('type.ringNorm.polynom'(X0),X2))) = s('type.ringNorm.polynom'(X0),happ(s(fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0)),happ(s(fun('type.ringNorm.polynom'(X0),fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0))),'const.ringNorm.Pmult_0'),s('type.ringNorm.polynom'(X0),X1))),s('type.ringNorm.polynom'(X0),X2)))),
  inference(rectify,[],[f59])).
fof(f59,axiom,(
  ! [X6] : ! [X9,X32] : s('type.ringNorm.polynom'(X6),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X6),X9),s('type.ringNorm.polynom'(X6),X32))) = s('type.ringNorm.polynom'(X6),happ(s(fun('type.ringNorm.polynom'(X6),'type.ringNorm.polynom'(X6)),happ(s(fun('type.ringNorm.polynom'(X6),fun('type.ringNorm.polynom'(X6),'type.ringNorm.polynom'(X6))),'const.ringNorm.Pmult_0'),s('type.ringNorm.polynom'(X6),X9))),s('type.ringNorm.polynom'(X6),X32)))),
  file('thm.ringNorm.polynom_distinct|split|6','ARITY_0_2_const.ringNorm.Pmult')).
% SZS output end Proof for thm.ringNorm
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 313
% Passive clauses: 5189
% Generated clauses: 51271
% Final active clauses: 153
% Final passive clauses: 1200
% Input formulas: 69
% Initial clauses: 125
% 
% Duplicate literals: 919
% Trivial inequalities: 311
% Fw subsumption resolutions: 1489
% Fw demodulations: 9001
% Bw demodulations: 2852
% Condensations: 1628
% Global subsumptions: 561
% 
% Simple tautologies: 291
% Equational tautologies: 330
% Forward subsumptions: 27348
% Fw demodulations to eq. taut.: 2768
% Bw demodulations to eq. taut.: 365
% 
% Unit resulting resolution: 7919
% Factoring: 2
% Forward superposition: 11160
% Backward superposition: 9210
% Self superposition: 294
% Equality factoring: 172
% Equality resolution: 39
% 
% Split clauses: 106
% Split components: 213
% SAT solver clauses: 6618
% SAT solver unit clauses: 476
% SAT solver binary clauses: 2314
% 
% Sat splits: 106
% Sat splitting refutations: 654
% 
% Memory used [KB]: 51043
% Time elapsed: 3.287 s
% ------------------------------
% ------------------------------
% Success in time 4.62 s
