Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.real_sigma
% (114563)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 732
% Passive clauses: 8721
% Generated clauses: 20167
% Final active clauses: 723
% Final passive clauses: 7599
% Input formulas: 42
% Initial clauses: 183
% 
% Duplicate literals: 529
% Trivial inequalities: 19
% Fw subsumption resolutions: 1992
% Fw demodulations: 501
% Bw demodulations: 8
% 
% Simple tautologies: 242
% Equational tautologies: 58
% Forward subsumptions: 7295
% 
% Binary resolution: 10355
% Factoring: 1429
% Forward superposition: 2727
% Backward superposition: 2103
% Self superposition: 26
% Equality factoring: 30
% Equality resolution: 5
% 
% Split clauses: 594
% Split components: 1195
% SAT solver clauses: 1045
% SAT solver unit clauses: 22
% SAT solver binary clauses: 170
% 
% Sat splits: 594
% 
% TWLsolver clauses: 1044
% TWLsolver calls for satisfiability: 435
% 
% Memory used [KB]: 9466
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.real_sigma
% (114576)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 559
% Passive clauses: 2405
% Generated clauses: 18082
% Final active clauses: 468
% Final passive clauses: 1254
% Input formulas: 42
% Initial clauses: 183
% 
% Duplicate literals: 466
% Trivial inequalities: 77
% Fw subsumption resolutions: 1732
% Fw demodulations: 911
% Bw demodulations: 656
% Condensations: 574
% Global subsumptions: 14
% 
% Simple tautologies: 333
% Equational tautologies: 25
% Forward subsumptions: 9170
% Bw demodulations to eq. taut.: 19
% 
% Binary resolution: 3648
% Factoring: 175
% Forward superposition: 2924
% Backward superposition: 4116
% Self superposition: 49
% Equality factoring: 40
% Equality resolution: 9
% 
% Split clauses: 177
% Split components: 355
% SAT solver clauses: 2564
% SAT solver unit clauses: 198
% SAT solver binary clauses: 69
% SAT solver learnt clauses: 28
% SAT solver learnt literals: 12
% 
% Sat splits: 177
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 2508
% TWLsolver calls for satisfiability: 1942
% 
% Memory used [KB]: 12537
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.real_sigma
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.real_sigma
% SZS output start Proof for thm.real_sigma
fof(f10543,plain,(
  $false),
  inference(subsumption_resolution,[],[f10542,f244])).
fof(f244,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.real_sigma.REAL_SUM_IMAGE_INTER_ELIM',aHL_TRUTH)).
fof(f10542,plain,(
  ~p(s(bool,'T_0'))),
  inference(forward_demodulation,[],[f10136,f475])).
fof(f475,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK14,bool),sK15)))),
  inference(unit_resulting_resolution,[],[f241,f259])).
fof(f259,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f163])).
fof(f163,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X13] : (s(bool,'T_0') = s(bool,X13) <=> p(s(bool,X13)))),
  file('thm.real_sigma.REAL_SUM_IMAGE_INTER_ELIM','thm.bool.EQ_CLAUSES|split|0')).
fof(f241,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK14,bool),sK15))))),
  inference(cnf_transformation,[],[f160])).
fof(f160,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK14,bool),sK15)))) & (! [X4] : (p(s(bool,'const.bool.IN_2'(s(sK14,X4),s(fun(sK14,bool),sK17)))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),sK16),s(sK14,X4)))) & s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK14,'type.realax.real'),sK16),s(fun(sK14,bool),sK15))) != s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK14,'type.realax.real'),sK16),s(fun(sK14,bool),'const.pred_set.INTER_2'(s(fun(sK14,bool),sK15),s(fun(sK14,bool),sK17))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK14,sK15,sK16,sK17])],[f109])).
fof(f109,plain,(
  ? [X0,X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) & ? [X2,X3] : (! [X4] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X3)))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X4)))) & s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) != s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X3)))))))),
  inference(ennf_transformation,[],[f45])).
fof(f45,plain,(
  ~! [X0,X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) => ! [X2,X3] : (! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X3)))) => s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X4)))) => s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X3)))))))),
  inference(flattening,[],[f44])).
fof(f44,plain,(
  ~! [X0] : ! [X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) => ! [X2,X3] : (! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X3)))) => s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X4)))) => s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X3)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X7)))) => ! [X8,X9] : (! [X10] : (~p(s(bool,'const.bool.IN_2'(s(X6,X10),s(fun(X6,bool),X9)))) => s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(X6,'type.realax.real'),X8),s(X6,X10)))) => s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X6,'type.realax.real'),X8),s(fun(X6,bool),X7))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X6,'type.realax.real'),X8),s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X9)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X7)))) => ! [X8,X9] : (! [X10] : (~p(s(bool,'const.bool.IN_2'(s(X6,X10),s(fun(X6,bool),X9)))) => s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(X6,'type.realax.real'),X8),s(X6,X10)))) => s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X6,'type.realax.real'),X8),s(fun(X6,bool),X7))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X6,'type.realax.real'),X8),s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X9)))))))),
  file('thm.real_sigma.REAL_SUM_IMAGE_INTER_ELIM','thm.real_sigma.REAL_SUM_IMAGE_INTER_ELIM')).
fof(f10136,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK14,bool),sK15))))),
  inference(unit_resulting_resolution,[],[f243,f5183,f1224])).
fof(f1224,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK14,bool),X1)))) | s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK14,'type.realax.real'),sK16),s(fun(sK14,bool),X1))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK14,'type.realax.real'),sK16),s(fun(sK14,bool),'const.pred_set.INTER_2'(s(fun(sK14,bool),X1),s(fun(sK14,bool),X0))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK14,sK18(X0,sK16,X1,sK14)),s(fun(sK14,bool),sK17)))) )),
  inference(trivial_inequality_removal,[],[f1223])).
fof(f1223,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK14,'type.realax.real'),sK16),s(fun(sK14,bool),X1))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK14,'type.realax.real'),sK16),s(fun(sK14,bool),'const.pred_set.INTER_2'(s(fun(sK14,bool),X1),s(fun(sK14,bool),X0))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK14,bool),X1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK14,sK18(X0,sK16,X1,sK14)),s(fun(sK14,bool),sK17)))) )),
  inference(superposition,[],[f288,f770])).
fof(f770,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),sK16),s(sK14,X0))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK14,X0),s(fun(sK14,bool),sK17)))) )),
  inference(subsumption_resolution,[],[f769,f434])).
fof(f434,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f247,f244])).
fof(f247,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f161])).
fof(f161,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f46])).
fof(f46,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f18])).
fof(f18,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.real_sigma.REAL_SUM_IMAGE_INTER_ELIM','thm.bool.NOT_CLAUSES|split|2')).
fof(f769,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),sK16),s(sK14,X0))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK14,X0),s(fun(sK14,bool),sK17)))) )),
  inference(superposition,[],[f242,f251])).
fof(f251,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.real_sigma.REAL_SUM_IMAGE_INTER_ELIM',aHL_BOOL_CASES)).
fof(f242,plain,(
  ( ! [X4] : (p(s(bool,'const.bool.IN_2'(s(sK14,X4),s(fun(sK14,bool),sK17)))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),sK16),s(sK14,X4)))) )),
  inference(cnf_transformation,[],[f160])).
fof(f288,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,sK18(X3,X2,X1,X0)))) | s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X3))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f179])).
fof(f179,plain,(
  ! [X0,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) | ! [X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,sK18(X3,X2,X1,X0)),s(fun(X0,bool),X3)))) & s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,sK18(X3,X2,X1,X0))))) | s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X3)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK18])],[f120])).
fof(f120,plain,(
  ! [X0,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) | ! [X2,X3] : (? [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X3)))) & s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X4)))) | s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X3)))))))),
  inference(ennf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0,X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) => ! [X2,X3] : (! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X3)))) => s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X4)))) => s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X3)))))))),
  inference(flattening,[],[f75])).
fof(f75,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) => ! [X2,X3] : (! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X3)))) => s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X4)))) => s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X3)))))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X6] : ! [X7] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X7)))) => ! [X8,X9] : (! [X10] : (~p(s(bool,'const.bool.IN_2'(s(X6,X10),s(fun(X6,bool),X9)))) => s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(X6,'type.realax.real'),X8),s(X6,X10)))) => s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X6,'type.realax.real'),X8),s(fun(X6,bool),X7))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X6,'type.realax.real'),X8),s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X9)))))))),
  file('thm.real_sigma.REAL_SUM_IMAGE_INTER_ELIM','thm.real_sigma.real_sigma_25_prove_ltb')).
fof(f5183,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK14,sK18(sK17,sK16,sK15,sK14)),s(fun(sK14,bool),sK17)))),
  inference(unit_resulting_resolution,[],[f472,f472,f4728,f1396])).
fof(f1396,plain,(
  ( ! [X2,X3,X1] : (s(bool,X1) != s(bool,X2) | p(s(bool,X2)) | s(bool,X1) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f258,f705])).
fof(f705,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f251,f251])).
fof(f258,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f163])).
fof(f4728,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK14,sK18(sK17,sK16,sK15,sK14)),s(fun(sK14,bool),sK17))))),
  inference(unit_resulting_resolution,[],[f243,f1219])).
fof(f1219,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(sK14,sK18(X0,X1,sK15,sK14)),s(fun(sK14,bool),X0)))) | s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK14,'type.realax.real'),X1),s(fun(sK14,bool),sK15))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK14,'type.realax.real'),X1),s(fun(sK14,bool),'const.pred_set.INTER_2'(s(fun(sK14,bool),sK15),s(fun(sK14,bool),X0)))))) )),
  inference(subsumption_resolution,[],[f1217,f244])).
fof(f1217,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK14,sK18(X0,X1,sK15,sK14)),s(fun(sK14,bool),X0)))) | s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK14,'type.realax.real'),X1),s(fun(sK14,bool),sK15))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK14,'type.realax.real'),X1),s(fun(sK14,bool),'const.pred_set.INTER_2'(s(fun(sK14,bool),sK15),s(fun(sK14,bool),X0)))))) )),
  inference(superposition,[],[f287,f475])).
fof(f287,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,sK18(X3,X2,X1,X0)),s(fun(X0,bool),X3)))) | s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X3)))))) )),
  inference(cnf_transformation,[],[f179])).
fof(f472,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f434,f258])).
fof(f243,plain,(
  s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK14,'type.realax.real'),sK16),s(fun(sK14,bool),sK15))) != s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK14,'type.realax.real'),sK16),s(fun(sK14,bool),'const.pred_set.INTER_2'(s(fun(sK14,bool),sK15),s(fun(sK14,bool),sK17)))))),
  inference(cnf_transformation,[],[f160])).
% SZS output end Proof for thm.real_sigma
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 502
% Passive clauses: 1243
% Generated clauses: 10155
% Final active clauses: 479
% Final passive clauses: 610
% Input formulas: 42
% Initial clauses: 183
% 
% Duplicate literals: 813
% Trivial inequalities: 47
% Fw subsumption resolutions: 667
% Fw demodulations: 405
% Bw demodulations: 71
% Condensations: 1201
% Global subsumptions: 29
% 
% Simple tautologies: 92
% Equational tautologies: 279
% Forward subsumptions: 5028
% Bw demodulations to eq. taut.: 10
% 
% Unit resulting resolution: 2982
% Factoring: 1
% Forward superposition: 1276
% Backward superposition: 2192
% Self superposition: 106
% Equality factoring: 106
% Equality resolution: 3
% 
% Split clauses: 63
% Split components: 131
% SAT solver clauses: 1715
% SAT solver unit clauses: 549
% SAT solver binary clauses: 258
% 
% Sat splits: 63
% Sat splitting refutations: 315
% 
% Memory used [KB]: 6524
% Time elapsed: 0.724 s
% ------------------------------
% ------------------------------
% Success in time 2.043 s
