Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.pred_set
% (85281)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 96
% Passive clauses: 3347
% Generated clauses: 15881
% Final active clauses: 88
% Final passive clauses: 3078
% Input formulas: 9
% Initial clauses: 18
% 
% Duplicate literals: 851
% Trivial inequalities: 188
% Fw subsumption resolutions: 2343
% Fw demodulations: 1130
% Bw demodulations: 6
% 
% Simple tautologies: 266
% Equational tautologies: 89
% Forward subsumptions: 5458
% 
% Binary resolution: 400
% Factoring: 5
% Forward superposition: 5593
% Backward superposition: 4668
% Self superposition: 289
% Equality factoring: 227
% Equality resolution: 17
% 
% Split clauses: 585
% Split components: 1172
% SAT solver clauses: 1495
% SAT solver unit clauses: 25
% SAT solver binary clauses: 218
% 
% Sat splits: 585
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 1495
% TWLsolver calls for satisfiability: 188
% 
% Memory used [KB]: 7931
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pred_set
% (85302)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 110
% Passive clauses: 887
% Generated clauses: 11654
% Final active clauses: 100
% Final passive clauses: 736
% Input formulas: 9
% Initial clauses: 18
% 
% Duplicate literals: 491
% Trivial inequalities: 39
% Fw subsumption resolutions: 894
% Fw demodulations: 1240
% Bw demodulations: 41
% Condensations: 674
% 
% Simple tautologies: 80
% Equational tautologies: 48
% Forward subsumptions: 4428
% Bw demodulations to eq. taut.: 6
% 
% Binary resolution: 462
% Factoring: 9
% Forward superposition: 3104
% Backward superposition: 3995
% Self superposition: 215
% Equality factoring: 124
% Equality resolution: 4
% 
% Split clauses: 48
% Split components: 97
% SAT solver clauses: 1765
% SAT solver unit clauses: 16
% SAT solver binary clauses: 33
% SAT solver learnt clauses: 1
% SAT solver learnt literals: 4
% 
% Sat splits: 48
% Sat splitting refutations: 9
% 
% TWLsolver clauses: 1764
% TWLsolver calls for satisfiability: 199
% 
% Memory used [KB]: 12153
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.pred_set
% (85325)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 884
% Passive clauses: 7000
% Generated clauses: 182782
% Final active clauses: 50
% Final passive clauses: 267
% Input formulas: 9
% Initial clauses: 18
% 
% Duplicate literals: 9346
% Trivial inequalities: 1305
% Fw subsumption resolutions: 15512
% Fw demodulations: 28484
% Bw demodulations: 5106
% Condensations: 13968
% Global subsumptions: 790
% 
% Simple tautologies: 1402
% Equational tautologies: 1036
% Forward subsumptions: 102367
% Fw demodulations to eq. taut.: 70
% Bw demodulations to eq. taut.: 360
% 
% Unit resulting resolution: 5258
% Factoring: 1
% Forward superposition: 48431
% Backward superposition: 33306
% Self superposition: 854
% Equality factoring: 357
% Equality resolution: 74
% 
% Split clauses: 598
% Split components: 1197
% SAT solver clauses: 9830
% SAT solver unit clauses: 6
% SAT solver binary clauses: 384
% 
% Sat splits: 598
% Sat splitting refutations: 842
% 
% Memory used [KB]: 72280
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.pred_set
% (85752)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 420
% Passive clauses: 21296
% Generated clauses: 78035
% Final active clauses: 191
% Final passive clauses: 10967
% Input formulas: 9
% Initial clauses: 18
% 
% Selected by SInE selection: 9
% SInE iterations: 1
% 
% Duplicate literals: 2353
% Trivial inequalities: 1045
% Fw subsumption resolutions: 12975
% Fw demodulations: 6682
% Bw demodulations: 79
% 
% Simple tautologies: 1500
% Equational tautologies: 113
% Forward subsumptions: 22708
% Bw demodulations to eq. taut.: 3
% 
% Binary resolution: 8597
% Factoring: 6
% Forward superposition: 23773
% Backward superposition: 19638
% Self superposition: 385
% Equality factoring: 254
% Equality resolution: 73
% 
% Split clauses: 8592
% Split components: 17537
% SAT solver clauses: 8671
% SAT solver unit clauses: 11
% SAT solver binary clauses: 135
% 
% Sat splits: 8592
% Sat splitting refutations: 24
% 
% Memory used [KB]: 31982
% Time elapsed: 4.0000 s
% ------------------------------
% ------------------------------
% remaining time: 142 next slice time: 6
lrs+1002_2:3_br=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:sas=minisat:stl=30:sac=on:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none:sp=reverse_arity:urr=on_4 on thm.pred_set
% (85836)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 104
% Passive clauses: 1619
% Generated clauses: 12156
% Final active clauses: 79
% Final passive clauses: 1258
% Input formulas: 9
% Initial clauses: 18
% 
% Duplicate literals: 580
% Trivial inequalities: 296
% Fw subsumption resolutions: 2019
% Fw demodulations: 189
% Bw demodulations: 164
% Global subsumptions: 8
% 
% Simple tautologies: 642
% Equational tautologies: 70
% Forward subsumptions: 6500
% Bw demodulations to eq. taut.: 2
% 
% Unit resulting resolution: 219
% Factoring: 2
% Forward superposition: 4195
% Backward superposition: 3911
% Self superposition: 117
% Equality factoring: 59
% Equality resolution: 16
% 
% Split clauses: 16
% Split components: 32
% SAT solver clauses: 1663
% SAT solver unit clauses: 9
% SAT solver binary clauses: 71
% 
% Sat splits: 16
% Sat splitting refutations: 18
% 
% Memory used [KB]: 8827
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 135 next slice time: 15
lrs+1011_12_bs=on:bsr=unit_only:cond=on:gs=on:gsaa=from_current:gsssp=full:nwc=1.1:sas=minisat:stl=60:sos=all:sac=on:sdd=off:sser=off:sfr=on:ssfp=100000:ssfq=1.2:ssnc=none:sp=reverse_arity:updr=off_14 on thm.pred_set
% (85854)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 244
% Passive clauses: 2467
% Generated clauses: 18905
% Final active clauses: 99
% Final passive clauses: 1040
% Input formulas: 9
% Initial clauses: 18
% 
% Duplicate literals: 1111
% Trivial inequalities: 384
% Fw subsumption resolutions: 1388
% Bw subsumption resolutions: 247
% Fw demodulations: 1153
% Bw demodulations: 394
% Condensations: 1905
% Global subsumptions: 53
% 
% Simple tautologies: 708
% Equational tautologies: 72
% Forward subsumptions: 9484
% Backward subsumptions: 372
% 
% Binary resolution: 1446
% Factoring: 4
% Forward superposition: 4419
% Backward superposition: 2763
% Self superposition: 126
% Equality factoring: 136
% Equality resolution: 62
% 
% Split clauses: 6
% Split components: 12
% SAT solver clauses: 2645
% SAT solver unit clauses: 7
% SAT solver binary clauses: 99
% 
% Sat splits: 6
% Sat splitting refutations: 117
% 
% Memory used [KB]: 11001
% Time elapsed: 1.600 s
% ------------------------------
% ------------------------------
% remaining time: 119 next slice time: 9
lrs+10_8:1_bd=preordered:bs=on:ccuc=first:fsr=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:nicw=on:sas=minisat:stl=120:sos=on:sscc=on:sser=off:sfr=on:ssfp=4000:ssfq=1.0:ssnc=none:sp=reverse_arity:urr=on_7 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f1654,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f40,f60,f58,f39,f57,f65,f63,f71,f62,f73,f61,f75,f130,f128,f127,f132,f137,f138,f185,f183,f182,f181,f192,f193,f190,f194,f195,f202,f197,f206,f210,f208,f217,f216,f220,f222,f226,f228,f225,f230,f224,f232,f223,f234,f258,f256,f265,f289,f293,f291,f294,f298,f296,f301,f305,f308,f381,f380,f379,f378,f377,f376,f385,f383,f387,f388,f390,f400,f386,f430,f429,f428,f427,f426,f425,f434,f432,f436,f437,f439,f449,f435,f546,f545,f543,f542,f541,f540,f550,f539,f554,f537,f548,f552,f555,f557,f586,f585,f583,f582,f581,f580,f590,f579,f594,f577,f588,f592,f595,f597,f673,f674,f675,f676,f677,f678,f723,f724,f725,f726,f727,f728,f754,f755,f765,f766,f767,f768,f886,f985,f887,f986,f897,f989,f898,f990,f899,f988,f900,f987,f991,f643,f993,f992,f640,f994,f1019,f1027,f1001,f1034,f1032,f1031,f1030,f1029,f1047,f1037,f1050,f1045,f1089,f1087,f1086,f1085,f1084,f1103,f1093,f1108,f1101,f1118,f1116,f1115,f1114,f1113,f1132,f1122,f1137,f1130,f1152,f551,f1172,f1154,f1174,f1180,f302,f1189,f1188,f1187,f1186,f1185,f553,f1184,f1183,f1196,f1197,f1195,f303,f1200,f1232,f1235,f1255,f1173,f1246,f1249,f1253,f1293,f591,f1313,f1295,f1315,f1321,f309,f1330,f1329,f1328,f1327,f1326,f593,f1325,f1324,f1337,f1338,f1336,f310,f1339,f1342,f1362,f1314,f1353,f1356,f1360,f1426,f1453,f1427,f1454,f1540,f1538,f1554,f1555,f1556,f1557,f1533,f1553,f1558,f1561,f1562,f1560,f1563,f221,f1565,f1566,f1567,f1568,f1569,f1570,f1571,f1572,f1573,f1574,f1575,f1343,f312,f59,f1121,f1120,f1111,f1112,f72,f1589,f1653,f140,f1021])).
fof(f1021,plain,(
  ~p(s(bool,'T_0')) | $spl95),
  inference(cnf_transformation,[],[f1021_D])).
fof(f1021_D,plain,(
  ~p(s(bool,'T_0')) <=> ~$spl95),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl95])])).
fof(f140,plain,(
  p(s(bool,'F_0')) | $spl12),
  inference(cnf_transformation,[],[f140_D])).
fof(f140_D,plain,(
  p(s(bool,'F_0')) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f1653,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'T_0')) | ($spl2 | $spl10 | $spl14)),
  inference(forward_demodulation,[],[f1651,f137])).
fof(f1651,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2)))) | ($spl2 | $spl14)),
  inference(superposition,[],[f1593,f190])).
fof(f1593,plain,(
  ( ! [X1] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK1)))) | p(s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK2))))) ) | $spl2),
  inference(superposition,[],[f50,f59])).
fof(f50,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f37])).
fof(f37,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))))),
  inference(flattening,[],[f36])).
fof(f36,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))))),
  inference(nnf_transformation,[],[f19])).
fof(f19,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(flattening,[],[f18])).
fof(f18,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6] : ! [X7,X8,X11] : (p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X7)))) & p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X8))))))),
  file('thm.pred_set.SUBSET_INTER_ABSORPTION','thm.pred_set.IN_INTER')).
fof(f1589,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1)))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))),
  introduced(theory_axiom,[])).
fof(f72,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2)))) | $spl6),
  inference(cnf_transformation,[],[f72_D])).
fof(f72_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2)))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f1112,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2)))) | $spl98),
  inference(unit_resulting_resolution,[],[f1047,f50])).
fof(f1111,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2)))) | $spl98),
  inference(unit_resulting_resolution,[],[f1047,f49])).
fof(f49,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f37])).
fof(f1120,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2)))) | $spl98),
  inference(resolution,[],[f1047,f49])).
fof(f1121,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2)))) | $spl98),
  inference(resolution,[],[f1047,f50])).
fof(f59,plain,(
  s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | $spl2),
  inference(cnf_transformation,[],[f59_D])).
fof(f59_D,plain,(
  s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f312,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl47),
  inference(cnf_transformation,[],[f312_D])).
fof(f312_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl47),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl47])])).
fof(f1343,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | $spl74),
  inference(superposition,[],[f45,f596])).
fof(f596,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl74),
  inference(cnf_transformation,[],[f596_D])).
fof(f596_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f45,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.bool.IN_2'(s(X0,sK3(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK3(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) )),
  inference(cnf_transformation,[],[f32])).
fof(f32,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (s(bool,'const.bool.IN_2'(s(X0,sK3(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK3(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3])],[f31])).
fof(f31,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(rectify,[],[f30])).
fof(f30,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(nnf_transformation,[],[f15])).
fof(f15,plain,(
  ! [X0,X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(flattening,[],[f14])).
fof(f14,plain,(
  ! [X0] : ! [X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : ! [X7,X8] : (s(fun(X6,bool),X7) = s(fun(X6,bool),X8) <=> ! [X11] : s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X7))) = s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X8))))),
  file('thm.pred_set.SUBSET_INTER_ABSORPTION','thm.pred_set.EXTENSION')).
fof(f1575,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))))),
  introduced(theory_axiom,[])).
fof(f1574,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))) | s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))))),
  introduced(theory_axiom,[])).
fof(f1573,plain,(
  s(fun(sK0,bool),sK1) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),
  introduced(theory_axiom,[])).
fof(f1572,plain,(
  s(fun(sK0,bool),sK1) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),sK1))))),
  introduced(theory_axiom,[])).
fof(f1571,plain,(
  s(fun(sK0,bool),sK1) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))))),
  introduced(theory_axiom,[])).
fof(f1570,plain,(
  s(fun(sK0,bool),sK1) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),sK1))))),
  introduced(theory_axiom,[])).
fof(f1569,plain,(
  s(fun(sK0,bool),sK1) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))))),
  introduced(theory_axiom,[])).
fof(f1568,plain,(
  s(fun(sK0,bool),sK1) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))))),
  introduced(theory_axiom,[])).
fof(f1567,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))))),
  introduced(theory_axiom,[])).
fof(f1566,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | p(s(bool,'T_0')) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),
  introduced(theory_axiom,[])).
fof(f1565,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),
  introduced(theory_axiom,[])).
fof(f221,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl20),
  inference(cnf_transformation,[],[f221_D])).
fof(f221_D,plain,(
  s(bool,'T_0') = s(bool,'F_0') <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f1563,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl0 | $spl62 | $spl109)),
  inference(forward_demodulation,[],[f1527,f552])).
fof(f1527,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | ($spl0 | $spl109)),
  inference(unit_resulting_resolution,[],[f42,f1173,f286])).
fof(f286,plain,(
  ( ! [X2,X1] : (p(s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK2)))) | p(s(bool,X2)) | s(bool,X2) = s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK1)))) ) | $spl0),
  inference(resolution,[],[f212,f52])).
fof(f52,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f26])).
fof(f26,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f25])).
fof(f25,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f20])).
fof(f20,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X9,X10] : ((p(s(bool,X9)) => p(s(bool,X10))) => ((p(s(bool,X10)) => p(s(bool,X9))) => s(bool,X9) = s(bool,X10)))),
  file('thm.pred_set.SUBSET_INTER_ABSORPTION','thm.bool.IMP_ANTISYM_AX')).
fof(f212,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK1)))) | p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK2))))) ) | $spl0),
  inference(resolution,[],[f57,f46])).
fof(f46,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))))) & ((p(s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X2))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f34])).
fof(f34,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))))) & (? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f33])).
fof(f33,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & (? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(nnf_transformation,[],[f24])).
fof(f24,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(ennf_transformation,[],[f17])).
fof(f17,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) => p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(flattening,[],[f16])).
fof(f16,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) => p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6] : ! [X7,X8] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))) <=> ! [X11] : (p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X7)))) => p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X8))))))),
  file('thm.pred_set.SUBSET_INTER_ABSORPTION','thm.pred_set.SUBSET_DEF')).
fof(f42,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f13])).
fof(f13,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.pred_set.SUBSET_INTER_ABSORPTION',aHL_FALSITY)).
fof(f1560,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | $spl122),
  inference(cnf_transformation,[],[f1560_D])).
fof(f1560_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f1562,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | ($spl0 | $spl62 | $spl109)),
  inference(forward_demodulation,[],[f1528,f552])).
fof(f1528,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | ($spl0 | $spl109)),
  inference(unit_resulting_resolution,[],[f1173,f1173,f286])).
fof(f1561,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl62 | $spl109)),
  inference(forward_demodulation,[],[f1529,f552])).
fof(f1529,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) | ($spl0 | $spl109)),
  inference(unit_resulting_resolution,[],[f1173,f212])).
fof(f1558,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | ($spl0 | $spl62 | $spl109)),
  inference(forward_demodulation,[],[f1532,f552])).
fof(f1532,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | ($spl0 | $spl109)),
  inference(unit_resulting_resolution,[],[f1173,f1173,f286])).
fof(f1553,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | $spl120),
  inference(cnf_transformation,[],[f1553_D])).
fof(f1553_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f1533,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | ($spl0 | $spl18 | $spl21 | $spl109)),
  inference(unit_resulting_resolution,[],[f222,f1173,f245])).
fof(f245,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,X1)) ) | ($spl0 | $spl18)),
  inference(resolution,[],[f218,f52])).
fof(f218,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,'T_0') = s(bool,X1)) ) | ($spl0 | $spl18)),
  inference(backward_demodulation,[],[f216,f213])).
fof(f213,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) ) | $spl0),
  inference(resolution,[],[f57,f55])).
fof(f55,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f26])).
fof(f1557,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | ($spl0 | $spl18 | $spl41 | $spl66 | $spl109)),
  inference(forward_demodulation,[],[f1534,f556])).
fof(f556,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl66),
  inference(cnf_transformation,[],[f556_D])).
fof(f556_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f1534,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl0 | $spl18 | $spl41 | $spl109)),
  inference(unit_resulting_resolution,[],[f302,f1173,f245])).
fof(f1556,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | ($spl0 | $spl18 | $spl45 | $spl74 | $spl109)),
  inference(forward_demodulation,[],[f1535,f596])).
fof(f1535,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl0 | $spl18 | $spl45 | $spl109)),
  inference(unit_resulting_resolution,[],[f309,f1173,f245])).
fof(f1555,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | ($spl0 | $spl18 | $spl34 | $spl51 | $spl109)),
  inference(forward_demodulation,[],[f1536,f292])).
fof(f292,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK5(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl34),
  inference(cnf_transformation,[],[f292_D])).
fof(f292_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK5(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f1536,plain,(
  s(bool,f(s(sK0,sK5(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | ($spl0 | $spl18 | $spl51 | $spl109)),
  inference(unit_resulting_resolution,[],[f386,f1173,f245])).
fof(f1554,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | ($spl0 | $spl18 | $spl38 | $spl57 | $spl109)),
  inference(forward_demodulation,[],[f1537,f297])).
fof(f297,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | $spl38),
  inference(cnf_transformation,[],[f297_D])).
fof(f297_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f1537,plain,(
  s(bool,f(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | ($spl0 | $spl18 | $spl57 | $spl109)),
  inference(unit_resulting_resolution,[],[f435,f1173,f245])).
fof(f1538,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | $spl109),
  inference(unit_resulting_resolution,[],[f42,f1173,f52])).
fof(f1540,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | $spl109),
  inference(unit_resulting_resolution,[],[f42,f1173,f52])).
fof(f1454,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X1),s(fun(sK0,bool),sK2))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X1))))) ) | $spl118),
  inference(cnf_transformation,[],[f1454_D])).
fof(f1454_D,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X1),s(fun(sK0,bool),sK2))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X1))))) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f1427,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X1),s(fun(sK0,bool),sK2))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X1))))) ) | ($spl0 | $spl18 | $spl96)),
  inference(superposition,[],[f249,f1045])).
fof(f249,plain,(
  ( ! [X10,X8,X7,X9] : (~p(s(bool,'const.bool.IN_2'(s(X7,X8),s(fun(X7,bool),X10)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X7,X8),s(fun(X7,bool),'const.pred_set.INTER_2'(s(fun(X7,bool),X9),s(fun(X7,bool),X10))))) | ~p(s(bool,'const.bool.IN_2'(s(X7,X8),s(fun(X7,bool),X9))))) ) | ($spl0 | $spl18)),
  inference(resolution,[],[f218,f51])).
fof(f51,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f37])).
fof(f1453,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X0),s(fun(sK0,bool),sK1))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X0))))) ) | $spl116),
  inference(cnf_transformation,[],[f1453_D])).
fof(f1453_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X0),s(fun(sK0,bool),sK1))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X0))))) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f1426,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X0),s(fun(sK0,bool),sK1))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X0))))) ) | ($spl0 | $spl14 | $spl18)),
  inference(superposition,[],[f249,f190])).
fof(f1360,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) ) | $spl114),
  inference(cnf_transformation,[],[f1360_D])).
fof(f1360_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f1356,plain,(
  ( ! [X8] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X8) | s(bool,X8) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) ) | ($spl0 | $spl18 | $spl74)),
  inference(superposition,[],[f361,f596])).
fof(f361,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | ($spl0 | $spl18)),
  inference(resolution,[],[f245,f218])).
fof(f1353,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X5) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,X5)) ) | ($spl0 | $spl18 | $spl74)),
  inference(superposition,[],[f361,f596])).
fof(f1314,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),sK2)))) | $spl113),
  inference(cnf_transformation,[],[f1314_D])).
fof(f1314_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),sK2)))) <=> ~$spl113),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl113])])).
fof(f1362,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),sK2)))) | ($spl70 | $spl74)),
  inference(forward_demodulation,[],[f1346,f592])).
fof(f1346,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),sK1)))) | $spl74),
  inference(superposition,[],[f51,f596])).
fof(f1342,plain,(
  ( ! [X3] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X3) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,X3)) ) | ($spl0 | $spl18 | $spl74)),
  inference(superposition,[],[f596,f361])).
fof(f1339,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) ) | ($spl0 | $spl18 | $spl74)),
  inference(superposition,[],[f596,f361])).
fof(f310,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl44),
  inference(cnf_transformation,[],[f310_D])).
fof(f310_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f1336,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl75),
  inference(trivial_inequality_removal,[],[f1335])).
fof(f1335,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl75),
  inference(superposition,[],[f597,f43])).
fof(f43,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.pred_set.SUBSET_INTER_ABSORPTION',aHL_BOOL_CASES)).
fof(f1338,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl0 | $spl18 | $spl75)),
  inference(global_subsumption,[],[f1333,f40,f39,f597,f1330,f1327,f1324])).
fof(f1333,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,X2)) ) | ($spl0 | $spl18 | $spl75)),
  inference(superposition,[],[f597,f361])).
fof(f1337,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl0 | $spl18 | $spl75)),
  inference(global_subsumption,[],[f1332,f40,f39,f597,f1330,f1327,f1324])).
fof(f1332,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) ) | ($spl0 | $spl18 | $spl75)),
  inference(superposition,[],[f597,f361])).
fof(f1324,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl75),
  inference(unit_resulting_resolution,[],[f597,f43])).
fof(f1325,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl0 | $spl18 | $spl21 | $spl75)),
  inference(unit_resulting_resolution,[],[f222,f597,f361])).
fof(f593,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | $spl72),
  inference(cnf_transformation,[],[f593_D])).
fof(f593_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f1326,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | ($spl0 | $spl18 | $spl21 | $spl75)),
  inference(unit_resulting_resolution,[],[f222,f597,f245])).
fof(f1327,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | $spl75),
  inference(unit_resulting_resolution,[],[f42,f597,f52])).
fof(f1328,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl0 | $spl18 | $spl21 | $spl75)),
  inference(unit_resulting_resolution,[],[f222,f597,f361])).
fof(f1329,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl0 | $spl18 | $spl75)),
  inference(unit_resulting_resolution,[],[f42,f597,f245])).
fof(f1330,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | $spl75),
  inference(unit_resulting_resolution,[],[f42,f597,f52])).
fof(f309,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl45),
  inference(cnf_transformation,[],[f309_D])).
fof(f309_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl45),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl45])])).
fof(f1321,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl29 | $spl70)),
  inference(backward_demodulation,[],[f592,f234])).
fof(f1315,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),sK2)))) | $spl112),
  inference(cnf_transformation,[],[f1315_D])).
fof(f1315_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),sK2)))) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f1295,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),sK2)))) | $spl44),
  inference(superposition,[],[f50,f310])).
fof(f1313,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'T_0')) | ($spl44 | $spl46)),
  inference(forward_demodulation,[],[f1294,f311])).
fof(f311,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl46),
  inference(cnf_transformation,[],[f311_D])).
fof(f311_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f1294,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),sK1)))) | $spl44),
  inference(superposition,[],[f49,f310])).
fof(f591,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl71),
  inference(cnf_transformation,[],[f591_D])).
fof(f591_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl71),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl71])])).
fof(f1293,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | $spl44),
  inference(superposition,[],[f45,f310])).
fof(f1253,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) ) | $spl110),
  inference(cnf_transformation,[],[f1253_D])).
fof(f1253_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f1249,plain,(
  ( ! [X8] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X8) | s(bool,X8) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) ) | ($spl0 | $spl18 | $spl66)),
  inference(superposition,[],[f361,f556])).
fof(f1246,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X5) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,X5)) ) | ($spl0 | $spl18 | $spl66)),
  inference(superposition,[],[f361,f556])).
fof(f1173,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) | $spl109),
  inference(cnf_transformation,[],[f1173_D])).
fof(f1173_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) <=> ~$spl109),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl109])])).
fof(f1255,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) | ($spl62 | $spl66)),
  inference(forward_demodulation,[],[f1239,f552])).
fof(f1239,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) | $spl66),
  inference(superposition,[],[f51,f556])).
fof(f1235,plain,(
  ( ! [X3] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X3) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,X3)) ) | ($spl0 | $spl18 | $spl66)),
  inference(superposition,[],[f556,f361])).
fof(f1232,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) ) | ($spl0 | $spl18 | $spl66)),
  inference(superposition,[],[f556,f361])).
fof(f1200,plain,(
  $false | ($spl14 | $spl21 | $spl92)),
  inference(unit_resulting_resolution,[],[f42,f222,f997])).
fof(f997,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | ($spl14 | $spl92)),
  inference(forward_demodulation,[],[f994,f190])).
fof(f303,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl40),
  inference(cnf_transformation,[],[f303_D])).
fof(f303_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f1195,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl67),
  inference(trivial_inequality_removal,[],[f1194])).
fof(f1194,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl67),
  inference(superposition,[],[f557,f43])).
fof(f1197,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl0 | $spl18 | $spl67)),
  inference(global_subsumption,[],[f1192,f40,f39,f557,f1189,f1186,f1183])).
fof(f1192,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,X2)) ) | ($spl0 | $spl18 | $spl67)),
  inference(superposition,[],[f557,f361])).
fof(f1196,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl0 | $spl18 | $spl67)),
  inference(global_subsumption,[],[f1191,f40,f39,f557,f1189,f1186,f1183])).
fof(f1191,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) ) | ($spl0 | $spl18 | $spl67)),
  inference(superposition,[],[f557,f361])).
fof(f1183,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl67),
  inference(unit_resulting_resolution,[],[f557,f43])).
fof(f1184,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl0 | $spl18 | $spl21 | $spl67)),
  inference(unit_resulting_resolution,[],[f222,f557,f361])).
fof(f553,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | $spl64),
  inference(cnf_transformation,[],[f553_D])).
fof(f553_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f1185,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | ($spl0 | $spl18 | $spl21 | $spl67)),
  inference(unit_resulting_resolution,[],[f222,f557,f245])).
fof(f1186,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | $spl67),
  inference(unit_resulting_resolution,[],[f42,f557,f52])).
fof(f1187,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl0 | $spl18 | $spl21 | $spl67)),
  inference(unit_resulting_resolution,[],[f222,f557,f361])).
fof(f1188,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl0 | $spl18 | $spl67)),
  inference(unit_resulting_resolution,[],[f42,f557,f245])).
fof(f1189,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | $spl67),
  inference(unit_resulting_resolution,[],[f42,f557,f52])).
fof(f302,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl41),
  inference(cnf_transformation,[],[f302_D])).
fof(f302_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl41),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl41])])).
fof(f1180,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl25 | $spl62)),
  inference(backward_demodulation,[],[f552,f230])).
fof(f1174,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) | $spl108),
  inference(cnf_transformation,[],[f1174_D])).
fof(f1174_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f1154,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) | $spl40),
  inference(superposition,[],[f50,f303])).
fof(f1172,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'T_0')) | ($spl40 | $spl42)),
  inference(forward_demodulation,[],[f1153,f304])).
fof(f304,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | $spl42),
  inference(cnf_transformation,[],[f304_D])).
fof(f304_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f1153,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) | $spl40),
  inference(superposition,[],[f49,f303])).
fof(f551,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | $spl63),
  inference(cnf_transformation,[],[f551_D])).
fof(f551_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl63),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl63])])).
fof(f1152,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | $spl40),
  inference(superposition,[],[f45,f303])).
fof(f1130,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))) | $spl104),
  inference(cnf_transformation,[],[f1130_D])).
fof(f1130_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f1137,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))) | $spl98),
  inference(global_subsumption,[],[f40,f39,f1047,f1118,f1116,f1114,f1113,f1123,f1128])).
fof(f1128,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))) | $spl98),
  inference(superposition,[],[f1047,f43])).
fof(f1123,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))) ) | $spl98),
  inference(resolution,[],[f1047,f55])).
fof(f1122,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))) | ($spl0 | $spl18 | $spl98)),
  inference(resolution,[],[f1047,f218])).
fof(f1132,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | $spl106),
  inference(cnf_transformation,[],[f1132_D])).
fof(f1132_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f1113,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | $spl98),
  inference(unit_resulting_resolution,[],[f1047,f1047,f51])).
fof(f1114,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | $spl98),
  inference(unit_resulting_resolution,[],[f1047,f1047,f51])).
fof(f1115,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))) | ($spl0 | $spl18 | $spl98)),
  inference(unit_resulting_resolution,[],[f1047,f218])).
fof(f1116,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))) | $spl98),
  inference(unit_resulting_resolution,[],[f41,f1047,f55])).
fof(f41,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.SUBSET_INTER_ABSORPTION',aHL_TRUTH)).
fof(f1118,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))) | $spl98),
  inference(unit_resulting_resolution,[],[f41,f1047,f55])).
fof(f1101,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl100),
  inference(cnf_transformation,[],[f1101_D])).
fof(f1101_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f1108,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl16),
  inference(global_subsumption,[],[f40,f39,f192,f1089,f1087,f1085,f1084,f1094,f1099])).
fof(f1099,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl16),
  inference(superposition,[],[f192,f43])).
fof(f1094,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))) ) | $spl16),
  inference(resolution,[],[f192,f55])).
fof(f1093,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | ($spl0 | $spl16 | $spl18)),
  inference(resolution,[],[f192,f218])).
fof(f1103,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | $spl102),
  inference(cnf_transformation,[],[f1103_D])).
fof(f1103_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f1084,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | $spl16),
  inference(unit_resulting_resolution,[],[f192,f192,f51])).
fof(f1085,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | $spl16),
  inference(unit_resulting_resolution,[],[f192,f192,f51])).
fof(f1086,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | ($spl0 | $spl16 | $spl18)),
  inference(unit_resulting_resolution,[],[f192,f218])).
fof(f1087,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl16),
  inference(unit_resulting_resolution,[],[f41,f192,f55])).
fof(f1089,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl16),
  inference(unit_resulting_resolution,[],[f41,f192,f55])).
fof(f1045,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl96),
  inference(cnf_transformation,[],[f1045_D])).
fof(f1045_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f1050,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl6),
  inference(global_subsumption,[],[f40,f39,f72,f1034,f1032,f1030,f1029,f1038,f1043])).
fof(f1043,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl6),
  inference(superposition,[],[f72,f43])).
fof(f1038,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2)))) ) | $spl6),
  inference(resolution,[],[f72,f55])).
fof(f1037,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl0 | $spl6 | $spl18)),
  inference(resolution,[],[f72,f218])).
fof(f1047,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))) | $spl98),
  inference(cnf_transformation,[],[f1047_D])).
fof(f1047_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f1029,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))) | $spl6),
  inference(unit_resulting_resolution,[],[f72,f72,f51])).
fof(f1030,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))) | $spl6),
  inference(unit_resulting_resolution,[],[f72,f72,f51])).
fof(f1031,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl0 | $spl6 | $spl18)),
  inference(unit_resulting_resolution,[],[f72,f218])).
fof(f1032,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl6),
  inference(unit_resulting_resolution,[],[f41,f72,f55])).
fof(f1034,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl6),
  inference(unit_resulting_resolution,[],[f41,f72,f55])).
fof(f1001,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2)))) | ($spl0 | $spl14)),
  inference(superposition,[],[f212,f190])).
fof(f1027,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))))),
  introduced(theory_axiom,[])).
fof(f1019,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'T_0')) | ($spl0 | $spl10 | $spl14)),
  inference(forward_demodulation,[],[f1001,f137])).
fof(f994,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1)))) ) | $spl92),
  inference(cnf_transformation,[],[f994_D])).
fof(f994_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1)))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f640,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1)))) ) | ($spl0 | $spl10)),
  inference(superposition,[],[f286,f137])).
fof(f992,plain,(
  ( ! [X3] : (p(s(bool,'const.bool.IN_2'(s(sK0,X3),s(fun(sK0,bool),sK2)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X3),s(fun(sK0,bool),sK1)))) ) | $spl88),
  inference(cnf_transformation,[],[f992_D])).
fof(f992_D,plain,(
  ( ! [X3] : (p(s(bool,'const.bool.IN_2'(s(sK0,X3),s(fun(sK0,bool),sK2)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X3),s(fun(sK0,bool),sK1)))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f993,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,'T_0') = s(bool,X2)) ) | $spl90),
  inference(cnf_transformation,[],[f993_D])).
fof(f993_D,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f643,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | p(s(bool,'const.bool.IN_2'(s(sK0,X3),s(fun(sK0,bool),sK2)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X3),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X2)) ) | $spl0),
  inference(superposition,[],[f286,f43])).
fof(f991,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),
  introduced(theory_axiom,[])).
fof(f987,plain,(
  ( ! [X16] : (s(bool,X16) = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X16)) ) | $spl80),
  inference(cnf_transformation,[],[f987_D])).
fof(f987_D,plain,(
  ( ! [X16] : (s(bool,X16) = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X16)) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f900,plain,(
  ( ! [X72] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X72) | s(bool,X72) = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2)))) ) | ($spl0 | $spl10 | $spl18)),
  inference(superposition,[],[f137,f361])).
fof(f988,plain,(
  ( ! [X17] : (s(bool,X17) = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X17)) ) | $spl82),
  inference(cnf_transformation,[],[f988_D])).
fof(f988_D,plain,(
  ( ! [X17] : (s(bool,X17) = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X17)) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f899,plain,(
  ( ! [X71] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X71) | s(bool,X71) = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1)))) ) | ($spl0 | $spl18 | $spl30)),
  inference(superposition,[],[f265,f361])).
fof(f990,plain,(
  ( ! [X19] : (s(bool,X19) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X19)) ) | $spl86),
  inference(cnf_transformation,[],[f990_D])).
fof(f990_D,plain,(
  ( ! [X19] : (s(bool,X19) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X19)) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f898,plain,(
  ( ! [X70] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X70) | s(bool,X70) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),sK1)))) ) | ($spl0 | $spl18 | $spl46)),
  inference(superposition,[],[f311,f361])).
fof(f989,plain,(
  ( ! [X18] : (s(bool,X18) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X18)) ) | $spl84),
  inference(cnf_transformation,[],[f989_D])).
fof(f989_D,plain,(
  ( ! [X18] : (s(bool,X18) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X18)) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f897,plain,(
  ( ! [X69] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X69) | s(bool,X69) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) ) | ($spl0 | $spl18 | $spl42)),
  inference(superposition,[],[f304,f361])).
fof(f986,plain,(
  ( ! [X15] : (s(bool,X15) = s(bool,f(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,X15)) ) | $spl78),
  inference(cnf_transformation,[],[f986_D])).
fof(f986_D,plain,(
  ( ! [X15] : (s(bool,X15) = s(bool,f(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,X15)) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f887,plain,(
  ( ! [X30] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X30) | s(bool,X30) = s(bool,f(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,bool,sK0))))) ) | ($spl0 | $spl18 | $spl38)),
  inference(superposition,[],[f297,f361])).
fof(f985,plain,(
  ( ! [X14] : (s(bool,X14) = s(bool,f(s(sK0,sK5(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'T_0') = s(bool,X14)) ) | $spl76),
  inference(cnf_transformation,[],[f985_D])).
fof(f985_D,plain,(
  ( ! [X14] : (s(bool,X14) = s(bool,f(s(sK0,sK5(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'T_0') = s(bool,X14)) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f886,plain,(
  ( ! [X29] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X29) | s(bool,X29) = s(bool,f(s(sK0,sK5(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0))))) ) | ($spl0 | $spl18 | $spl34)),
  inference(superposition,[],[f292,f361])).
fof(f768,plain,(
  ( ! [X72] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X72) = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X72)) ) | ($spl0 | $spl10 | $spl18)),
  inference(superposition,[],[f137,f361])).
fof(f767,plain,(
  ( ! [X71] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X71) = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X71)) ) | ($spl0 | $spl18 | $spl30)),
  inference(superposition,[],[f265,f361])).
fof(f766,plain,(
  ( ! [X70] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X70) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X70)) ) | ($spl0 | $spl18 | $spl46)),
  inference(superposition,[],[f311,f361])).
fof(f765,plain,(
  ( ! [X69] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X69) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X69)) ) | ($spl0 | $spl18 | $spl42)),
  inference(superposition,[],[f304,f361])).
fof(f755,plain,(
  ( ! [X30] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X30) = s(bool,f(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,X30)) ) | ($spl0 | $spl18 | $spl38)),
  inference(superposition,[],[f297,f361])).
fof(f754,plain,(
  ( ! [X29] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X29) = s(bool,f(s(sK0,sK5(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'T_0') = s(bool,X29)) ) | ($spl0 | $spl18 | $spl34)),
  inference(superposition,[],[f292,f361])).
fof(f728,plain,(
  ( ! [X19] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X19) | s(bool,X19) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),sK1)))) ) | ($spl0 | $spl18 | $spl46)),
  inference(superposition,[],[f361,f311])).
fof(f727,plain,(
  ( ! [X18] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X18) | s(bool,X18) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) ) | ($spl0 | $spl18 | $spl42)),
  inference(superposition,[],[f361,f304])).
fof(f726,plain,(
  ( ! [X17] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X17) | s(bool,X17) = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1)))) ) | ($spl0 | $spl18 | $spl30)),
  inference(superposition,[],[f361,f265])).
fof(f725,plain,(
  ( ! [X16] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X16) | s(bool,X16) = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2)))) ) | ($spl0 | $spl10 | $spl18)),
  inference(superposition,[],[f361,f137])).
fof(f724,plain,(
  ( ! [X15] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X15) | s(bool,X15) = s(bool,f(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,bool,sK0))))) ) | ($spl0 | $spl18 | $spl38)),
  inference(superposition,[],[f361,f297])).
fof(f723,plain,(
  ( ! [X14] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X14) | s(bool,X14) = s(bool,f(s(sK0,sK5(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0))))) ) | ($spl0 | $spl18 | $spl34)),
  inference(superposition,[],[f361,f292])).
fof(f678,plain,(
  ( ! [X19] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X19) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X19)) ) | ($spl0 | $spl18 | $spl46)),
  inference(superposition,[],[f361,f311])).
fof(f677,plain,(
  ( ! [X18] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X18) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X18)) ) | ($spl0 | $spl18 | $spl42)),
  inference(superposition,[],[f361,f304])).
fof(f676,plain,(
  ( ! [X17] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X17) = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X17)) ) | ($spl0 | $spl18 | $spl30)),
  inference(superposition,[],[f361,f265])).
fof(f675,plain,(
  ( ! [X16] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X16) = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X16)) ) | ($spl0 | $spl10 | $spl18)),
  inference(superposition,[],[f361,f137])).
fof(f674,plain,(
  ( ! [X15] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X15) = s(bool,f(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,X15)) ) | ($spl0 | $spl18 | $spl38)),
  inference(superposition,[],[f361,f297])).
fof(f673,plain,(
  ( ! [X14] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X14) = s(bool,f(s(sK0,sK5(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'T_0') = s(bool,X14)) ) | ($spl0 | $spl18 | $spl34)),
  inference(superposition,[],[f361,f292])).
fof(f597,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl75),
  inference(cnf_transformation,[],[f597_D])).
fof(f597_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl75),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl75])])).
fof(f595,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl29 | $spl46)),
  inference(backward_demodulation,[],[f311,f234])).
fof(f592,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl70),
  inference(cnf_transformation,[],[f592_D])).
fof(f592_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f588,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl47),
  inference(trivial_inequality_removal,[],[f587])).
fof(f587,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl47),
  inference(superposition,[],[f312,f43])).
fof(f577,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl9 | $spl30 | $spl47)),
  inference(unit_resulting_resolution,[],[f41,f312,f316])).
fof(f316,plain,(
  ( ! [X2,X1] : (~p(s(bool,X2)) | s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,X2)) ) | ($spl9 | $spl30)),
  inference(resolution,[],[f268,f55])).
fof(f268,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) | ($spl9 | $spl30)),
  inference(backward_demodulation,[],[f265,f261])).
fof(f261,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1)))) ) | $spl9),
  inference(resolution,[],[f74,f52])).
fof(f74,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1)))) | $spl9),
  inference(cnf_transformation,[],[f74_D])).
fof(f74_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1)))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f594,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | $spl73),
  inference(cnf_transformation,[],[f594_D])).
fof(f594_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) <=> ~$spl73),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl73])])).
fof(f579,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | ($spl9 | $spl29 | $spl30 | $spl47)),
  inference(unit_resulting_resolution,[],[f234,f312,f316])).
fof(f590,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),sK1)))) | $spl68),
  inference(cnf_transformation,[],[f590_D])).
fof(f590_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),sK1)))) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f580,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),sK1)))) | ($spl9 | $spl30 | $spl47)),
  inference(unit_resulting_resolution,[],[f312,f268])).
fof(f581,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl47),
  inference(unit_resulting_resolution,[],[f312,f43])).
fof(f582,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),sK1)))) | ($spl0 | $spl18 | $spl21 | $spl47)),
  inference(unit_resulting_resolution,[],[f222,f312,f245])).
fof(f583,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),sK1)))) | $spl47),
  inference(unit_resulting_resolution,[],[f42,f312,f52])).
fof(f585,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl0 | $spl18 | $spl47)),
  inference(unit_resulting_resolution,[],[f42,f312,f245])).
fof(f586,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),sK1)))) | $spl47),
  inference(unit_resulting_resolution,[],[f42,f312,f52])).
fof(f557,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl67),
  inference(cnf_transformation,[],[f557_D])).
fof(f557_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl67),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl67])])).
fof(f555,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl25 | $spl42)),
  inference(backward_demodulation,[],[f304,f230])).
fof(f552,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | $spl62),
  inference(cnf_transformation,[],[f552_D])).
fof(f552_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f548,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | $spl43),
  inference(trivial_inequality_removal,[],[f547])).
fof(f547,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | $spl43),
  inference(superposition,[],[f305,f43])).
fof(f537,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | ($spl9 | $spl30 | $spl43)),
  inference(unit_resulting_resolution,[],[f41,f305,f316])).
fof(f554,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | $spl65),
  inference(cnf_transformation,[],[f554_D])).
fof(f554_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) <=> ~$spl65),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl65])])).
fof(f539,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | ($spl9 | $spl25 | $spl30 | $spl43)),
  inference(unit_resulting_resolution,[],[f230,f305,f316])).
fof(f550,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) | $spl60),
  inference(cnf_transformation,[],[f550_D])).
fof(f550_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f540,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) | ($spl9 | $spl30 | $spl43)),
  inference(unit_resulting_resolution,[],[f305,f268])).
fof(f541,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | $spl43),
  inference(unit_resulting_resolution,[],[f305,f43])).
fof(f542,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) | ($spl0 | $spl18 | $spl21 | $spl43)),
  inference(unit_resulting_resolution,[],[f222,f305,f245])).
fof(f543,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) | $spl43),
  inference(unit_resulting_resolution,[],[f42,f305,f52])).
fof(f545,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | ($spl0 | $spl18 | $spl43)),
  inference(unit_resulting_resolution,[],[f42,f305,f245])).
fof(f546,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) | $spl43),
  inference(unit_resulting_resolution,[],[f42,f305,f52])).
fof(f435,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | $spl57),
  inference(cnf_transformation,[],[f435_D])).
fof(f435_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) <=> ~$spl57),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl57])])).
fof(f449,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | $spl36),
  inference(superposition,[],[f56,f296])).
fof(f56,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK5(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK5(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK5(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK5(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5])],[f27])).
fof(f27,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f22])).
fof(f22,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f21])).
fof(f21,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.pred_set.SUBSET_INTER_ABSORPTION',aHL_EXT)).
fof(f439,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | $spl59),
  inference(cnf_transformation,[],[f439_D])).
fof(f439_D,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) <=> ~$spl59),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl59])])).
fof(f437,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | ($spl27 | $spl38)),
  inference(backward_demodulation,[],[f297,f232])).
fof(f436,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | $spl56),
  inference(cnf_transformation,[],[f436_D])).
fof(f436_D,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f432,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | $spl39),
  inference(trivial_inequality_removal,[],[f431])).
fof(f431,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | $spl39),
  inference(superposition,[],[f298,f43])).
fof(f434,plain,(
  p(s(bool,f(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,bool,sK0))))) | $spl54),
  inference(cnf_transformation,[],[f434_D])).
fof(f434_D,plain,(
  p(s(bool,f(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,bool,sK0))))) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f425,plain,(
  p(s(bool,f(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,bool,sK0))))) | ($spl9 | $spl30 | $spl39)),
  inference(unit_resulting_resolution,[],[f298,f268])).
fof(f426,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | $spl39),
  inference(unit_resulting_resolution,[],[f298,f43])).
fof(f427,plain,(
  p(s(bool,f(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,bool,sK0))))) | ($spl0 | $spl18 | $spl21 | $spl39)),
  inference(unit_resulting_resolution,[],[f222,f298,f245])).
fof(f428,plain,(
  p(s(bool,f(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,bool,sK0))))) | $spl39),
  inference(unit_resulting_resolution,[],[f42,f298,f52])).
fof(f429,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | ($spl0 | $spl18 | $spl39)),
  inference(unit_resulting_resolution,[],[f42,f298,f245])).
fof(f430,plain,(
  p(s(bool,f(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,bool,sK0))))) | $spl39),
  inference(unit_resulting_resolution,[],[f42,f298,f52])).
fof(f386,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK5(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl51),
  inference(cnf_transformation,[],[f386_D])).
fof(f386_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK5(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) <=> ~$spl51),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl51])])).
fof(f400,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK5(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | $spl32),
  inference(superposition,[],[f56,f291])).
fof(f390,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK5(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl53),
  inference(cnf_transformation,[],[f390_D])).
fof(f390_D,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK5(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) <=> ~$spl53),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl53])])).
fof(f388,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK5(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | ($spl23 | $spl34)),
  inference(backward_demodulation,[],[f292,f228])).
fof(f387,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK5(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl50),
  inference(cnf_transformation,[],[f387_D])).
fof(f387_D,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK5(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f383,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK5(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl35),
  inference(trivial_inequality_removal,[],[f382])).
fof(f382,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(sK0,sK5(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl35),
  inference(superposition,[],[f293,f43])).
fof(f385,plain,(
  p(s(bool,f(s(sK0,sK5(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0))))) | $spl48),
  inference(cnf_transformation,[],[f385_D])).
fof(f385_D,plain,(
  p(s(bool,f(s(sK0,sK5(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0))))) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f376,plain,(
  p(s(bool,f(s(sK0,sK5(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0))))) | ($spl9 | $spl30 | $spl35)),
  inference(unit_resulting_resolution,[],[f293,f268])).
fof(f377,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK5(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl35),
  inference(unit_resulting_resolution,[],[f293,f43])).
fof(f378,plain,(
  p(s(bool,f(s(sK0,sK5(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0))))) | ($spl0 | $spl18 | $spl21 | $spl35)),
  inference(unit_resulting_resolution,[],[f222,f293,f245])).
fof(f379,plain,(
  p(s(bool,f(s(sK0,sK5(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0))))) | $spl35),
  inference(unit_resulting_resolution,[],[f42,f293,f52])).
fof(f380,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK5(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | ($spl0 | $spl18 | $spl35)),
  inference(unit_resulting_resolution,[],[f42,f293,f245])).
fof(f381,plain,(
  p(s(bool,f(s(sK0,sK5(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0))))) | $spl35),
  inference(unit_resulting_resolution,[],[f42,f293,f52])).
fof(f308,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl29),
  inference(superposition,[],[f234,f43])).
fof(f305,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | $spl43),
  inference(cnf_transformation,[],[f305_D])).
fof(f305_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl43),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl43])])).
fof(f301,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl25),
  inference(superposition,[],[f230,f43])).
fof(f296,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | $spl36),
  inference(cnf_transformation,[],[f296_D])).
fof(f296_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f298,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | $spl39),
  inference(cnf_transformation,[],[f298_D])).
fof(f298_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) <=> ~$spl39),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl39])])).
fof(f294,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | $spl27),
  inference(superposition,[],[f232,f43])).
fof(f291,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK5(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl32),
  inference(cnf_transformation,[],[f291_D])).
fof(f291_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK5(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f293,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK5(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl35),
  inference(cnf_transformation,[],[f293_D])).
fof(f293_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK5(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) <=> ~$spl35),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl35])])).
fof(f289,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK5(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK5(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl23),
  inference(superposition,[],[f228,f43])).
fof(f265,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl30),
  inference(cnf_transformation,[],[f265_D])).
fof(f265_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f256,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl9),
  inference(unit_resulting_resolution,[],[f42,f74,f52])).
fof(f258,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl9),
  inference(unit_resulting_resolution,[],[f42,f74,f52])).
fof(f234,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),sK1))) != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl29),
  inference(cnf_transformation,[],[f234_D])).
fof(f234_D,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),sK1))) != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f223,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),sK1))) != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl3),
  inference(unit_resulting_resolution,[],[f60,f45])).
fof(f232,plain,(
  s(bool,f(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) != s(bool,g(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | $spl27),
  inference(cnf_transformation,[],[f232_D])).
fof(f232_D,plain,(
  s(bool,f(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) != s(bool,g(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f224,plain,(
  s(bool,f(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) != s(bool,g(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | $spl3),
  inference(unit_resulting_resolution,[],[f60,f56])).
fof(f230,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl25),
  inference(cnf_transformation,[],[f230_D])).
fof(f230_D,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f225,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl3),
  inference(unit_resulting_resolution,[],[f60,f45])).
fof(f228,plain,(
  s(bool,f(s(sK0,sK5(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) != s(bool,g(s(sK0,sK5(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl23),
  inference(cnf_transformation,[],[f228_D])).
fof(f228_D,plain,(
  s(bool,f(s(sK0,sK5(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) != s(bool,g(s(sK0,sK5(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f226,plain,(
  s(bool,f(s(sK0,sK5(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) != s(bool,g(s(sK0,sK5(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl3),
  inference(unit_resulting_resolution,[],[f60,f56])).
fof(f222,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl21),
  inference(cnf_transformation,[],[f222_D])).
fof(f222_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f220,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ($spl5 | $spl18)),
  inference(backward_demodulation,[],[f216,f70])).
fof(f70,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | $spl5),
  inference(cnf_transformation,[],[f70_D])).
fof(f70_D,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f216,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | $spl18),
  inference(cnf_transformation,[],[f216_D])).
fof(f216_D,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f217,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | $spl0),
  inference(global_subsumption,[],[f40,f39,f57,f210,f208,f212,f213,f214])).
fof(f214,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | $spl0),
  inference(superposition,[],[f57,f43])).
fof(f208,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | $spl0),
  inference(unit_resulting_resolution,[],[f41,f57,f55])).
fof(f210,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | $spl0),
  inference(unit_resulting_resolution,[],[f41,f57,f55])).
fof(f206,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))) | p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))))),
  introduced(theory_axiom,[])).
fof(f197,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f42,f140])).
fof(f202,plain,(
  $false | $spl12),
  inference(resolution,[],[f140,f42])).
fof(f195,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))) | ~p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))))),
  introduced(theory_axiom,[])).
fof(f194,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))) | s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))))),
  introduced(theory_axiom,[])).
fof(f190,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl14),
  inference(cnf_transformation,[],[f190_D])).
fof(f190_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f193,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl8),
  inference(global_subsumption,[],[f40,f39,f75,f185,f183,f182,f181,f187,f188])).
fof(f188,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl8),
  inference(superposition,[],[f75,f43])).
fof(f187,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1)))) ) | $spl8),
  inference(resolution,[],[f75,f55])).
fof(f192,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))) | $spl16),
  inference(cnf_transformation,[],[f192_D])).
fof(f192_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f181,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))) | $spl8),
  inference(unit_resulting_resolution,[],[f75,f75,f51])).
fof(f182,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))) | $spl8),
  inference(unit_resulting_resolution,[],[f75,f75,f51])).
fof(f183,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl8),
  inference(unit_resulting_resolution,[],[f41,f75,f55])).
fof(f185,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl8),
  inference(unit_resulting_resolution,[],[f41,f75,f55])).
fof(f138,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f57,f71])).
fof(f137,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl10),
  inference(cnf_transformation,[],[f137_D])).
fof(f137_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f132,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl1 | $spl4 | $spl7)),
  inference(resolution,[],[f73,f76])).
fof(f76,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) | ($spl1 | $spl4)),
  inference(backward_demodulation,[],[f71,f67])).
fof(f67,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,X0) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) ) | $spl1),
  inference(resolution,[],[f58,f52])).
fof(f127,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl1 | $spl4 | $spl7)),
  inference(unit_resulting_resolution,[],[f73,f76])).
fof(f128,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl7),
  inference(unit_resulting_resolution,[],[f42,f73,f52])).
fof(f130,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl7),
  inference(unit_resulting_resolution,[],[f42,f73,f52])).
fof(f75,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1)))) | $spl8),
  inference(cnf_transformation,[],[f75_D])).
fof(f75_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1)))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f61,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1)))) | $spl1),
  inference(unit_resulting_resolution,[],[f58,f47])).
fof(f47,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X1)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f35])).
fof(f73,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2)))) | $spl7),
  inference(cnf_transformation,[],[f73_D])).
fof(f73_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2)))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f62,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2)))) | $spl1),
  inference(unit_resulting_resolution,[],[f58,f48])).
fof(f48,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X2)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f35])).
fof(f71,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | $spl4),
  inference(cnf_transformation,[],[f71_D])).
fof(f71_D,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f63,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | $spl1),
  inference(unit_resulting_resolution,[],[f42,f58,f52])).
fof(f65,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | $spl1),
  inference(unit_resulting_resolution,[],[f42,f58,f52])).
fof(f57,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) | $spl0),
  inference(cnf_transformation,[],[f57_D])).
fof(f57_D,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f39,plain,(
  s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),
  inference(cnf_transformation,[],[f29])).
fof(f29,plain,(
  (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) | s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) & (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) | s(fun(sK0,bool),sK1) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f28])).
fof(f28,plain,(
  ? [X0,X1,X2] : ((p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) & (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | s(fun(X0,bool),X1) != s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))),
  inference(nnf_transformation,[],[f23])).
fof(f23,plain,(
  ? [X0,X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <~> s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))),
  inference(ennf_transformation,[],[f12])).
fof(f12,plain,(
  ~! [X0,X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))),
  inference(flattening,[],[f11])).
fof(f11,plain,(
  ~! [X0] : ! [X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))) <=> s(fun(X6,bool),X7) = s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))) <=> s(fun(X6,bool),X7) = s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8))))),
  file('thm.pred_set.SUBSET_INTER_ABSORPTION','thm.pred_set.SUBSET_INTER_ABSORPTION')).
fof(f58,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) | $spl1),
  inference(cnf_transformation,[],[f58_D])).
fof(f58_D,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f60,plain,(
  s(fun(sK0,bool),sK1) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | $spl3),
  inference(cnf_transformation,[],[f60_D])).
fof(f60_D,plain,(
  s(fun(sK0,bool),sK1) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f40,plain,(
  s(fun(sK0,bool),sK1) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),
  inference(cnf_transformation,[],[f29])).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 75
% Passive clauses: 219
% Generated clauses: 1590
% Final active clauses: 26
% Final passive clauses: 29
% Input formulas: 9
% Initial clauses: 18
% 
% Duplicate literals: 71
% Trivial inequalities: 16
% Fw demodulations: 265
% Bw demodulations: 33
% Global subsumptions: 9
% 
% Simple tautologies: 33
% Equational tautologies: 71
% Forward subsumptions: 710
% Fw demodulations to eq. taut.: 6
% 
% Binary resolution: 115
% Unit resulting resolution: 250
% Factoring: 1
% Forward superposition: 219
% Backward superposition: 430
% Self superposition: 20
% Equality factoring: 24
% 
% Split clauses: 56
% Split components: 114
% SAT solver clauses: 613
% SAT solver unit clauses: 2
% SAT solver binary clauses: 223
% 
% Sat splits: 56
% Sat splitting refutations: 3
% 
% Memory used [KB]: 2046
% Time elapsed: 0.060 s
% ------------------------------
% ------------------------------
% Success in time 18.107 s
