Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.relation
% (87986)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 761
% Passive clauses: 6718
% Generated clauses: 22590
% Final active clauses: 748
% Final passive clauses: 5596
% Input formulas: 36
% Initial clauses: 171
% 
% Duplicate literals: 933
% Trivial inequalities: 44
% Fw subsumption resolutions: 2097
% Fw demodulations: 651
% Bw demodulations: 10
% 
% Simple tautologies: 263
% Equational tautologies: 78
% Forward subsumptions: 7822
% 
% Binary resolution: 8032
% Factoring: 1390
% Forward superposition: 3174
% Backward superposition: 5623
% Self superposition: 110
% Equality factoring: 110
% Equality resolution: 8
% 
% Split clauses: 1412
% Split components: 3023
% SAT solver clauses: 1870
% SAT solver unit clauses: 24
% SAT solver binary clauses: 269
% 
% Sat splits: 1412
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1865
% TWLsolver calls for satisfiability: 467
% 
% Memory used [KB]: 11257
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.relation
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.relation
% SZS output start Proof for thm.relation
fof(f7661,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f7625,f221,f416,f532,f536,f219,f412,f235,f424,f533,f538,f218,f414,f1872,f1884,f1864,f1866,f417,f419,f265,f435,f552,f558,f220,f410,f239,f426,f1952,f1964,f1918,f1920,f264,f436,f2152,f2176,f261,f433,f267,f441])).
fof(f441,plain,(
  ( ! [X6,X0,X7,X5,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X7)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))),s(X0,X7)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X6)))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1))))) ) | $spl38),
  inference(cnf_transformation,[],[f441_D])).
fof(f441_D,plain,(
  ( ! [X6,X0,X7,X5,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X7)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))),s(X0,X7)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X6)))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1))))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f267,plain,(
  ( ! [X6,X0,X7,X5,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X7)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))),s(X0,X7)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X6)))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(cnf_transformation,[],[f166])).
fof(f166,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X5,X6,X7] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X6)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))),s(X0,X7)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X7)))))) & ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK17(X1,X0)))),s(X0,sK18(X1,X0))))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK18(X1,X0)))),s(X0,sK19(X1,X0))))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK17(X1,X0)))),s(X0,sK19(X1,X0)))))) | p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK17,sK18,sK19])],[f165])).
fof(f165,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X5,X6,X7] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X6)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))),s(X0,X7)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X7)))))) & (? [X2,X3,X4] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X4)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X4))))) | p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(rectify,[],[f164])).
fof(f164,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X2,X3,X4] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X4)))))) & (? [X2,X3,X4] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X4)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X4))))) | p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(nnf_transformation,[],[f106])).
fof(f106,plain,(
  ! [X0,X1] : (p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> ! [X2,X3,X4] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X4))))))),
  inference(flattening,[],[f105])).
fof(f105,plain,(
  ! [X0,X1] : (p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> ! [X2,X3,X4] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X4))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X4))))))),
  inference(ennf_transformation,[],[f68])).
fof(f68,plain,(
  ! [X0,X1] : (p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> ! [X2,X3,X4] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X4))))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X4))))))),
  inference(flattening,[],[f67])).
fof(f67,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> ! [X2,X3,X4] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X4))))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X4))))))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X6] : ! [X7] : (p(s(bool,'const.relation.transitive_1'(s(fun(X6,fun(X6,bool)),X7)))) <=> ! [X9,X10,X16] : ((p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X9))),s(X6,X10)))) & p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X10))),s(X6,X16))))) => p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X9))),s(X6,X16))))))),
  file('thm.relation.irrefl_trans_implies_antisym','thm.relation.transitive_def')).
fof(f433,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X3)))) | ~p(s(bool,'const.relation.irreflexive_1'(s(fun(X0,fun(X0,bool)),X1))))) ) | $spl22),
  inference(cnf_transformation,[],[f433_D])).
fof(f433_D,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X3)))) | ~p(s(bool,'const.relation.irreflexive_1'(s(fun(X0,fun(X0,bool)),X1))))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f261,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X3)))) | ~p(s(bool,'const.relation.irreflexive_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(cnf_transformation,[],[f160])).
fof(f160,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.irreflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X3))))) & (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK14(X1,X0)))),s(X0,sK14(X1,X0))))) | p(s(bool,'const.relation.irreflexive_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK14])],[f159])).
fof(f159,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.irreflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X3))))) & (? [X2] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X2)))) | p(s(bool,'const.relation.irreflexive_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(rectify,[],[f158])).
fof(f158,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.irreflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X2] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X2))))) & (? [X2] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X2)))) | p(s(bool,'const.relation.irreflexive_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(nnf_transformation,[],[f64])).
fof(f64,plain,(
  ! [X0,X1] : (p(s(bool,'const.relation.irreflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> ! [X2] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X2)))))),
  inference(flattening,[],[f63])).
fof(f63,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.relation.irreflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> ! [X2] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X2)))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X6] : ! [X7] : (p(s(bool,'const.relation.irreflexive_1'(s(fun(X6,fun(X6,bool)),X7)))) <=> ! [X9] : ~p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X9))),s(X6,X9)))))),
  file('thm.relation.irrefl_trans_implies_antisym','thm.relation.irreflexive_def')).
fof(f2176,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2)))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))) | ~p(s(bool,'const.relation.irreflexive_1'(s(fun(X0,fun(X0,bool)),X1))))) ) | $spl824),
  inference(cnf_transformation,[],[f2176_D])).
fof(f2176_D,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2)))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))) | ~p(s(bool,'const.relation.irreflexive_1'(s(fun(X0,fun(X0,bool)),X1))))) ) <=> ~$spl824),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl824])])).
fof(f2152,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2)))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))) | ~p(s(bool,'const.relation.irreflexive_1'(s(fun(X0,fun(X0,bool)),X1))))) ) | ($spl22 | $spl38)),
  inference(resolution,[],[f441,f433])).
fof(f436,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK15(X1,X0)))),s(X0,sK16(X1,X0)))))) ) | $spl28),
  inference(cnf_transformation,[],[f436_D])).
fof(f436_D,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK15(X1,X0)))),s(X0,sK16(X1,X0)))))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f264,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK15(X1,X0)))),s(X0,sK16(X1,X0)))))) )),
  inference(cnf_transformation,[],[f163])).
fof(f163,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X4,X5] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X4)))) | s(X0,X4) = s(X0,X5))) & ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK15(X1,X0)))),s(X0,sK16(X1,X0))))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK16(X1,X0)))),s(X0,sK15(X1,X0))))) & s(X0,sK15(X1,X0)) != s(X0,sK16(X1,X0))) | p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK15,sK16])],[f162])).
fof(f162,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X4,X5] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X4)))) | s(X0,X4) = s(X0,X5))) & (? [X2,X3] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2)))) & s(X0,X2) != s(X0,X3)) | p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(rectify,[],[f161])).
fof(f161,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X2,X3] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2)))) | s(X0,X2) = s(X0,X3))) & (? [X2,X3] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2)))) & s(X0,X2) != s(X0,X3)) | p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(nnf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0,X1] : (p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> ! [X2,X3] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2)))) | s(X0,X2) = s(X0,X3)))),
  inference(flattening,[],[f103])).
fof(f103,plain,(
  ! [X0,X1] : (p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> ! [X2,X3] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2))))) | s(X0,X2) = s(X0,X3)))),
  inference(ennf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0,X1] : (p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> ! [X2,X3] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2))))) => s(X0,X2) = s(X0,X3)))),
  inference(flattening,[],[f65])).
fof(f65,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> ! [X2,X3] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2))))) => s(X0,X2) = s(X0,X3)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X6] : ! [X7] : (p(s(bool,'const.relation.antisymmetric_1'(s(fun(X6,fun(X6,bool)),X7)))) <=> ! [X9,X10] : ((p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X9))),s(X6,X10)))) & p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X10))),s(X6,X9))))) => s(X6,X9) = s(X6,X10)))),
  file('thm.relation.irrefl_trans_implies_antisym','thm.relation.antisymmetric_def')).
fof(f1920,plain,(
  p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15(sK13,sK12)))),s(sK12,sK16(sK13,sK12))))) | $spl710),
  inference(cnf_transformation,[],[f1920_D])).
fof(f1920_D,plain,(
  p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15(sK13,sK12)))),s(sK12,sK16(sK13,sK12))))) <=> ~$spl710),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl710])])).
fof(f1918,plain,(
  p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15(sK13,sK12)))),s(sK12,sK16(sK13,sK12))))) | ($spl9 | $spl28 | $spl208)),
  inference(subsumption_resolution,[],[f1907,f419])).
fof(f1907,plain,(
  p(s(bool,'F_0')) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15(sK13,sK12)))),s(sK12,sK16(sK13,sK12))))) | ($spl28 | $spl208)),
  inference(superposition,[],[f436,f558])).
fof(f1964,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15(sK13,sK12)))),s(sK12,sK16(sK13,sK12)))) | $spl730),
  inference(cnf_transformation,[],[f1964_D])).
fof(f1964_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15(sK13,sK12)))),s(sK12,sK16(sK13,sK12)))) <=> ~$spl730),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl730])])).
fof(f1952,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15(sK13,sK12)))),s(sK12,sK16(sK13,sK12)))) | ($spl12 | $spl710)),
  inference(resolution,[],[f1920,f424])).
fof(f426,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) | $spl16),
  inference(cnf_transformation,[],[f426_D])).
fof(f426_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f239,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f144])).
fof(f144,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f49])).
fof(f49,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X8] : (s(bool,'F_0') = s(bool,X8) <=> ~p(s(bool,X8)))),
  file('thm.relation.irrefl_trans_implies_antisym','thm.bool.EQ_CLAUSES|split|3')).
fof(f410,plain,(
  ~p(s(bool,'const.relation.antisymmetric_1'(s(fun(sK12,fun(sK12,bool)),sK13)))) | $spl1),
  inference(cnf_transformation,[],[f410_D])).
fof(f410_D,plain,(
  ~p(s(bool,'const.relation.antisymmetric_1'(s(fun(sK12,fun(sK12,bool)),sK13)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f220,plain,(
  ~p(s(bool,'const.relation.antisymmetric_1'(s(fun(sK12,fun(sK12,bool)),sK13))))),
  inference(cnf_transformation,[],[f139])).
fof(f139,plain,(
  p(s(bool,'const.relation.irreflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)))) & p(s(bool,'const.relation.transitive_1'(s(fun(sK12,fun(sK12,bool)),sK13)))) & ~p(s(bool,'const.relation.antisymmetric_1'(s(fun(sK12,fun(sK12,bool)),sK13))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13])],[f93])).
fof(f93,plain,(
  ? [X0,X1] : (p(s(bool,'const.relation.irreflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) & p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))) & ~p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1)))))),
  inference(flattening,[],[f92])).
fof(f92,plain,(
  ? [X0,X1] : ((p(s(bool,'const.relation.irreflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) & p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1))))) & ~p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1)))))),
  inference(ennf_transformation,[],[f39])).
fof(f39,plain,(
  ~! [X0,X1] : ((p(s(bool,'const.relation.irreflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) & p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1))))) => p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1)))))),
  inference(flattening,[],[f38])).
fof(f38,plain,(
  ~! [X0] : ! [X1] : ((p(s(bool,'const.relation.irreflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) & p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1))))) => p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7] : ((p(s(bool,'const.relation.irreflexive_1'(s(fun(X6,fun(X6,bool)),X7)))) & p(s(bool,'const.relation.transitive_1'(s(fun(X6,fun(X6,bool)),X7))))) => p(s(bool,'const.relation.antisymmetric_1'(s(fun(X6,fun(X6,bool)),X7)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7] : ((p(s(bool,'const.relation.irreflexive_1'(s(fun(X6,fun(X6,bool)),X7)))) & p(s(bool,'const.relation.transitive_1'(s(fun(X6,fun(X6,bool)),X7))))) => p(s(bool,'const.relation.antisymmetric_1'(s(fun(X6,fun(X6,bool)),X7)))))),
  file('thm.relation.irrefl_trans_implies_antisym','thm.relation.irrefl_trans_implies_antisym')).
fof(f558,plain,(
  s(bool,'F_0') = s(bool,'const.relation.antisymmetric_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl208),
  inference(cnf_transformation,[],[f558_D])).
fof(f558_D,plain,(
  s(bool,'F_0') = s(bool,'const.relation.antisymmetric_1'(s(fun(sK12,fun(sK12,bool)),sK13))) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f552,plain,(
  s(bool,'F_0') = s(bool,'const.relation.antisymmetric_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | ($spl1 | $spl16)),
  inference(resolution,[],[f426,f410])).
fof(f435,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK16(X1,X0)))),s(X0,sK15(X1,X0)))))) ) | $spl26),
  inference(cnf_transformation,[],[f435_D])).
fof(f435_D,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK16(X1,X0)))),s(X0,sK15(X1,X0)))))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f265,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.relation.antisymmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK16(X1,X0)))),s(X0,sK15(X1,X0)))))) )),
  inference(cnf_transformation,[],[f163])).
fof(f419,plain,(
  ~p(s(bool,'F_0')) | $spl9),
  inference(cnf_transformation,[],[f419_D])).
fof(f419_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f417,plain,(
  ~p(s(bool,'F_0'))),
  inference(global_subsumption,[],[f224,f221])).
fof(f224,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f140])).
fof(f140,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f40])).
fof(f40,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f15])).
fof(f15,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.relation.irrefl_trans_implies_antisym','thm.bool.NOT_CLAUSES|split|2')).
fof(f1866,plain,(
  p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK16(sK13,sK12)))),s(sK12,sK15(sK13,sK12))))) | $spl672),
  inference(cnf_transformation,[],[f1866_D])).
fof(f1866_D,plain,(
  p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK16(sK13,sK12)))),s(sK12,sK15(sK13,sK12))))) <=> ~$spl672),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl672])])).
fof(f1864,plain,(
  p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK16(sK13,sK12)))),s(sK12,sK15(sK13,sK12))))) | ($spl9 | $spl26 | $spl208)),
  inference(subsumption_resolution,[],[f1853,f419])).
fof(f1853,plain,(
  p(s(bool,'F_0')) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK16(sK13,sK12)))),s(sK12,sK15(sK13,sK12))))) | ($spl26 | $spl208)),
  inference(superposition,[],[f435,f558])).
fof(f1884,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK16(sK13,sK12)))),s(sK12,sK15(sK13,sK12)))) | $spl674),
  inference(cnf_transformation,[],[f1884_D])).
fof(f1884_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK16(sK13,sK12)))),s(sK12,sK15(sK13,sK12)))) <=> ~$spl674),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl674])])).
fof(f1872,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK16(sK13,sK12)))),s(sK12,sK15(sK13,sK12)))) | ($spl12 | $spl672)),
  inference(resolution,[],[f1866,f424])).
fof(f414,plain,(
  p(s(bool,'const.relation.irreflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)))) | $spl4),
  inference(cnf_transformation,[],[f414_D])).
fof(f414_D,plain,(
  p(s(bool,'const.relation.irreflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f218,plain,(
  p(s(bool,'const.relation.irreflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))))),
  inference(cnf_transformation,[],[f139])).
fof(f538,plain,(
  s(bool,'T_0') = s(bool,'const.relation.irreflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl202),
  inference(cnf_transformation,[],[f538_D])).
fof(f538_D,plain,(
  s(bool,'T_0') = s(bool,'const.relation.irreflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f533,plain,(
  s(bool,'T_0') = s(bool,'const.relation.irreflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | ($spl4 | $spl12)),
  inference(resolution,[],[f424,f414])).
fof(f424,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl12),
  inference(cnf_transformation,[],[f424_D])).
fof(f424_D,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f235,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f142])).
fof(f142,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X8] : (s(bool,'T_0') = s(bool,X8) <=> p(s(bool,X8)))),
  file('thm.relation.irrefl_trans_implies_antisym','thm.bool.EQ_CLAUSES|split|1')).
fof(f412,plain,(
  p(s(bool,'const.relation.transitive_1'(s(fun(sK12,fun(sK12,bool)),sK13)))) | $spl2),
  inference(cnf_transformation,[],[f412_D])).
fof(f412_D,plain,(
  p(s(bool,'const.relation.transitive_1'(s(fun(sK12,fun(sK12,bool)),sK13)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f219,plain,(
  p(s(bool,'const.relation.transitive_1'(s(fun(sK12,fun(sK12,bool)),sK13))))),
  inference(cnf_transformation,[],[f139])).
fof(f536,plain,(
  s(bool,'T_0') = s(bool,'const.relation.transitive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | $spl200),
  inference(cnf_transformation,[],[f536_D])).
fof(f536_D,plain,(
  s(bool,'T_0') = s(bool,'const.relation.transitive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f532,plain,(
  s(bool,'T_0') = s(bool,'const.relation.transitive_1'(s(fun(sK12,fun(sK12,bool)),sK13))) | ($spl2 | $spl12)),
  inference(resolution,[],[f424,f412])).
fof(f416,plain,(
  p(s(bool,'T_0')) | $spl6),
  inference(cnf_transformation,[],[f416_D])).
fof(f416_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f221,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.relation.irrefl_trans_implies_antisym',aHL_TRUTH)).
fof(f7625,plain,(
  $false | ($spl6 | $spl200 | $spl202 | $spl674 | $spl730 | $spl824)),
  inference(subsumption_resolution,[],[f7624,f416])).
fof(f7624,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl200 | $spl202 | $spl674 | $spl730 | $spl824)),
  inference(forward_demodulation,[],[f7623,f538])).
fof(f7623,plain,(
  ~p(s(bool,'const.relation.irreflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)))) | ($spl6 | $spl200 | $spl674 | $spl730 | $spl824)),
  inference(subsumption_resolution,[],[f7622,f416])).
fof(f7622,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.relation.irreflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)))) | ($spl6 | $spl200 | $spl674 | $spl730 | $spl824)),
  inference(forward_demodulation,[],[f7621,f536])).
fof(f7621,plain,(
  ~p(s(bool,'const.relation.transitive_1'(s(fun(sK12,fun(sK12,bool)),sK13)))) | ~p(s(bool,'const.relation.irreflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)))) | ($spl6 | $spl674 | $spl730 | $spl824)),
  inference(subsumption_resolution,[],[f7620,f416])).
fof(f7620,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.relation.transitive_1'(s(fun(sK12,fun(sK12,bool)),sK13)))) | ~p(s(bool,'const.relation.irreflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)))) | ($spl6 | $spl674 | $spl730 | $spl824)),
  inference(forward_demodulation,[],[f7619,f1964])).
fof(f7619,plain,(
  ~p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15(sK13,sK12)))),s(sK12,sK16(sK13,sK12))))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(sK12,fun(sK12,bool)),sK13)))) | ~p(s(bool,'const.relation.irreflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)))) | ($spl6 | $spl674 | $spl824)),
  inference(subsumption_resolution,[],[f7510,f416])).
fof(f7510,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15(sK13,sK12)))),s(sK12,sK16(sK13,sK12))))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(sK12,fun(sK12,bool)),sK13)))) | ~p(s(bool,'const.relation.irreflexive_1'(s(fun(sK12,fun(sK12,bool)),sK13)))) | ($spl674 | $spl824)),
  inference(superposition,[],[f2176,f1884])).
% SZS output end Proof for thm.relation
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 362
% Passive clauses: 1490
% Generated clauses: 7311
% Final active clauses: 356
% Final passive clauses: 1128
% Input formulas: 36
% Initial clauses: 171
% 
% Duplicate literals: 175
% Trivial inequalities: 65
% Fw subsumption resolutions: 665
% Fw demodulations: 262
% Bw demodulations: 5
% Condensations: 189
% Global subsumptions: 13
% 
% Simple tautologies: 233
% Equational tautologies: 25
% Forward subsumptions: 2162
% 
% Binary resolution: 2399
% Factoring: 180
% Forward superposition: 1079
% Backward superposition: 556
% Self superposition: 22
% Equality factoring: 26
% Equality resolution: 11
% 
% Split clauses: 107
% Split components: 219
% SAT solver clauses: 2202
% SAT solver unit clauses: 182
% SAT solver binary clauses: 75
% SAT solver learnt clauses: 27
% SAT solver learnt literals: 10
% 
% Sat splits: 107
% Sat splitting refutations: 10
% 
% TWLsolver clauses: 2146
% TWLsolver calls for satisfiability: 1732
% 
% Memory used [KB]: 8187
% Time elapsed: 0.386 s
% ------------------------------
% ------------------------------
% Success in time 0.999 s
