Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.list
% (55731)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1125
% Passive clauses: 7666
% Generated clauses: 36501
% Final active clauses: 906
% Final passive clauses: 6221
% Input formulas: 260
% Initial clauses: 521
% 
% Duplicate literals: 207
% Trivial inequalities: 62
% Fw subsumption resolutions: 2163
% Fw demodulations: 4147
% Bw demodulations: 194
% 
% Simple tautologies: 291
% Equational tautologies: 24
% Forward subsumptions: 21721
% Fw demodulations to eq. taut.: 21
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 6694
% Factoring: 857
% Forward superposition: 9307
% Backward superposition: 12161
% Self superposition: 54
% Equality factoring: 2
% Equality resolution: 11
% 
% Split clauses: 271
% Split components: 544
% SAT solver clauses: 498
% SAT solver unit clauses: 12
% SAT solver binary clauses: 87
% 
% Sat splits: 271
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 498
% TWLsolver calls for satisfiability: 385
% 
% Memory used [KB]: 18293
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.list
% (55743)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 626
% Passive clauses: 3759
% Generated clauses: 27378
% Final active clauses: 572
% Final passive clauses: 2762
% Input formulas: 260
% Initial clauses: 521
% 
% Duplicate literals: 102
% Trivial inequalities: 35
% Fw subsumption resolutions: 1238
% Fw demodulations: 3835
% Bw demodulations: 395
% Condensations: 102
% Global subsumptions: 25
% 
% Simple tautologies: 293
% Equational tautologies: 12
% Forward subsumptions: 13459
% Fw demodulations to eq. taut.: 22
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 2732
% Factoring: 145
% Forward superposition: 7589
% Backward superposition: 6528
% Self superposition: 78
% Equality factoring: 2
% Equality resolution: 10
% 
% Split clauses: 15
% Split components: 30
% SAT solver clauses: 4648
% SAT solver unit clauses: 641
% SAT solver binary clauses: 184
% SAT solver learnt clauses: 54
% SAT solver learnt literals: 19
% 
% Sat splits: 15
% Sat splitting refutations: 66
% 
% TWLsolver clauses: 4544
% TWLsolver calls for satisfiability: 2991
% 
% Memory used [KB]: 22515
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.list
% (55756)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1230
% Passive clauses: 11128
% Generated clauses: 353605
% Final active clauses: 789
% Final passive clauses: 2556
% Input formulas: 260
% Initial clauses: 521
% Discarded non-redundant clauses: 15491
% 
% Duplicate literals: 9
% Trivial inequalities: 35
% Fw subsumption resolutions: 2765
% Fw demodulations: 149455
% Bw demodulations: 4376
% Condensations: 31
% Global subsumptions: 1
% 
% Simple tautologies: 126
% Equational tautologies: 57
% Forward subsumptions: 94462
% Fw demodulations to eq. taut.: 79906
% Bw demodulations to eq. taut.: 18
% 
% Unit resulting resolution: 149247
% Factoring: 3
% Forward superposition: 25254
% Backward superposition: 17780
% Self superposition: 215
% Equality factoring: 2
% Equality resolution: 16
% 
% Split clauses: 23
% Split components: 46
% SAT solver clauses: 13715
% SAT solver unit clauses: 5158
% SAT solver binary clauses: 6321
% 
% Sat splits: 23
% Sat splitting refutations: 2538
% 
% Memory used [KB]: 135605
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.list
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.list
% SZS output start Proof for thm.list
fof(f100521,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f100520])).
fof(f100520,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK14))))))))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK14)))))))))),
  inference(forward_demodulation,[],[f100511,f637])).
fof(f637,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))) )),
  inference(cnf_transformation,[],[f321])).
fof(f321,plain,(
  ! [X0,X1,X2] : s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))),
  inference(flattening,[],[f320])).
fof(f320,plain,(
  ! [X0] : ! [X1,X2] : s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))),
  inference(rectify,[],[f140])).
fof(f140,axiom,(
  ! [X6] : ! [X17,X18] : s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),X17),s(fun(X6,bool),X18))) = s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),X18),s(fun(X6,bool),X17)))),
  file('thm.list.list_39|49|_prove_ltb','thm.pred_set.INTER_COMM')).
fof(f100511,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK14))))))))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK14))),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))))))),
  inference(equality_resolution,[],[f2036])).
fof(f2036,plain,(
  ( ! [X1] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK14))),s(fun(sK12,bool),X1))))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK14))),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK14))),s(fun(sK12,bool),X1))))))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK14)))))))))) )),
  inference(subsumption_resolution,[],[f1976,f654])).
fof(f654,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1))))))) )),
  inference(cnf_transformation,[],[f346])).
fof(f346,plain,(
  ! [X0,X1] : p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1))))))),
  inference(flattening,[],[f345])).
fof(f345,plain,(
  ! [X0] : ! [X1] : p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1))))))),
  inference(rectify,[],[f126])).
fof(f126,axiom,(
  ! [X6] : ! [X33] : p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),X33))))))),
  file('thm.list.list_39|49|_prove_ltb','thm.list.FINITE_LIST_TO_SET')).
fof(f1976,plain,(
  ( ! [X1] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK14))),s(fun(sK12,bool),X1))))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK14))),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK14))),s(fun(sK12,bool),X1))))))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK14))))))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK14))))))) )),
  inference(superposition,[],[f854,f625])).
fof(f625,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f417])).
fof(f417,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))))))))),
  inference(ennf_transformation,[],[f305])).
fof(f305,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) => s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))))))))),
  inference(flattening,[],[f304])).
fof(f304,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) => s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))))))))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X6] : ! [X18,X17] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X17)))) => s('type.num.num','const.pred_set.CARD_1'(s(fun(X6,bool),'const.pred_set.DIFF_2'(s(fun(X6,bool),X17),s(fun(X6,bool),X18))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X6,bool),X17))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),X17),s(fun(X6,bool),X18))))))))),
  file('thm.list.list_39|49|_prove_ltb','thm.pred_set.CARD_DIFF_EQN')).
fof(f854,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK14))))),s('type.num.num',X0))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK14))),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK14)))))))))) )),
  inference(superposition,[],[f832,f805])).
fof(f805,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) = s('type.num.num',X1) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f556])).
fof(f556,plain,(
  ! [X0,X1] : ((s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) & (s('type.num.num',X0) != s('type.num.num',X1) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))),
  inference(nnf_transformation,[],[f395])).
fof(f395,plain,(
  ! [X0,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) <=> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(rectify,[],[f144])).
fof(f144,axiom,(
  ! [X47,X39] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X39))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X47))) <=> s('type.num.num',X39) = s('type.num.num',X47))),
  file('thm.list.list_39|49|_prove_ltb','thm.prim_rec.INV_SUC_EQ')).
fof(f832,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK14))),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK14))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK14)))))))))),
  inference(forward_demodulation,[],[f559,f607])).
fof(f607,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2))),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1))))) = s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(fun(X0,bool),bool)),'const.bool.IN_0'),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1))))))),s('type.list.list'(X0),X2)))))) )),
  inference(cnf_transformation,[],[f285])).
fof(f285,plain,(
  ! [X0,X1,X2] : s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2))),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1))))) = s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(fun(X0,bool),bool)),'const.bool.IN_0'),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1))))))),s('type.list.list'(X0),X2)))))),
  inference(flattening,[],[f284])).
fof(f284,plain,(
  ! [X0] : ! [X1,X2] : s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2))),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1))))) = s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(fun(X0,bool),bool)),'const.bool.IN_0'),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1))))))),s('type.list.list'(X0),X2)))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : ! [X7,X8] : s(fun(X6,bool),'const.pred_set.DIFF_2'(s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),X8))),s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),X7))))) = s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),'const.list.FILTER_2'(s(fun(X6,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(X6,bool),'_dst_x1x0_2'(s(fun(X6,fun(fun(X6,bool),bool)),'const.bool.IN_0'),s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),X7))))))),s('type.list.list'(X6),X8)))))),
  file('thm.list.list_39|49|_prove_ltb','thm.list.list_to_set_diff')).
fof(f559,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK14))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK14))))))))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.FILTER_2'(s(fun(sK12,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(sK12,bool),'_dst_x1x0_2'(s(fun(sK12,fun(fun(sK12,bool),bool)),'const.bool.IN_0'),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))))),s('type.list.list'(sK12),sK14)))))))),
  inference(cnf_transformation,[],[f465])).
fof(f465,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK14))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK14))))))))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.FILTER_2'(s(fun(sK12,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(sK12,bool),'_dst_x1x0_2'(s(fun(sK12,fun(fun(sK12,bool),bool)),'const.bool.IN_0'),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))))),s('type.list.list'(sK12),sK14)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13,sK14])],[f410])).
fof(f410,plain,(
  ? [X0,X1,X2] : s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1))),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2))))))))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(fun(X0,bool),bool)),'const.bool.IN_0'),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1))))))),s('type.list.list'(X0),X2)))))))),
  inference(ennf_transformation,[],[f263])).
fof(f263,plain,(
  ~! [X0,X1,X2] : s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1))),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2))))))))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(fun(X0,bool),bool)),'const.bool.IN_0'),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1))))))),s('type.list.list'(X0),X2)))))))),
  inference(flattening,[],[f262])).
fof(f262,plain,(
  ~! [X0] : ! [X1,X2] : s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1))),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2))))))))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(fun(X0,bool),bool)),'const.bool.IN_0'),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1))))))),s('type.list.list'(X0),X2)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),X8))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),X7))),s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),X8))))))))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),'const.list.FILTER_2'(s(fun(X6,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(X6,bool),'_dst_x1x0_2'(s(fun(X6,fun(fun(X6,bool),bool)),'const.bool.IN_0'),s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),X7))))))),s('type.list.list'(X6),X8)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),X8))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),X7))),s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),X8))))))))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),'const.list.FILTER_2'(s(fun(X6,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(X6,bool),'_dst_x1x0_2'(s(fun(X6,fun(fun(X6,bool),bool)),'const.bool.IN_0'),s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),X7))))))),s('type.list.list'(X6),X8)))))))),
  file('thm.list.list_39|49|_prove_ltb','thm.list.list_39|49|_prove_ltb')).
% SZS output end Proof for thm.list
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 2252
% Passive clauses: 42685
% Generated clauses: 99769
% Final active clauses: 1403
% Final passive clauses: 28768
% Input formulas: 260
% Initial clauses: 266
% 
% Selected by SInE selection: 96
% SInE iterations: 1
% 
% Duplicate literals: 1784
% Trivial inequalities: 309
% Fw subsumption resolutions: 3867
% Fw demodulations: 7669
% Bw demodulations: 5
% 
% Simple tautologies: 3556
% Equational tautologies: 184
% Forward subsumptions: 37749
% Fw demodulations to eq. taut.: 4
% 
% Binary resolution: 62874
% Factoring: 84
% Forward superposition: 21185
% Backward superposition: 1078
% Self superposition: 1
% Equality resolution: 90
% 
% Split clauses: 3406
% Split components: 7228
% SAT solver clauses: 3478
% SAT solver unit clauses: 15
% SAT solver binary clauses: 399
% 
% Sat splits: 3406
% Sat splitting refutations: 34
% 
% Memory used [KB]: 35692
% Time elapsed: 2.863 s
% ------------------------------
% ------------------------------
% Success in time 14.62 s
