Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.rat
% (82851)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1293
% Passive clauses: 8250
% Generated clauses: 33004
% Final active clauses: 1184
% Final passive clauses: 6305
% Input formulas: 517
% Initial clauses: 823
% 
% Duplicate literals: 257
% Trivial inequalities: 82
% Fw subsumption resolutions: 2417
% Fw demodulations: 2880
% Bw demodulations: 103
% 
% Simple tautologies: 312
% Equational tautologies: 34
% Forward subsumptions: 18709
% Fw demodulations to eq. taut.: 8
% 
% Binary resolution: 8519
% Factoring: 964
% Forward superposition: 8124
% Backward superposition: 8545
% Self superposition: 84
% Equality factoring: 2
% Equality resolution: 18
% 
% Split clauses: 386
% Split components: 774
% SAT solver clauses: 709
% SAT solver unit clauses: 15
% SAT solver binary clauses: 121
% 
% Sat splits: 386
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 709
% TWLsolver calls for satisfiability: 592
% 
% Memory used [KB]: 17782
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.rat
% (82883)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 740
% Passive clauses: 3518
% Generated clauses: 27094
% Final active clauses: 680
% Final passive clauses: 2371
% Input formulas: 517
% Initial clauses: 823
% 
% Duplicate literals: 126
% Trivial inequalities: 145
% Fw subsumption resolutions: 992
% Fw demodulations: 2781
% Bw demodulations: 459
% Condensations: 140
% Global subsumptions: 48
% 
% Simple tautologies: 273
% Equational tautologies: 20
% Forward subsumptions: 14793
% Fw demodulations to eq. taut.: 2
% Bw demodulations to eq. taut.: 5
% 
% Binary resolution: 2717
% Factoring: 326
% Forward superposition: 8554
% Backward superposition: 6271
% Self superposition: 108
% Equality factoring: 2
% Equality resolution: 26
% 
% Split clauses: 20
% Split components: 40
% SAT solver clauses: 5163
% SAT solver unit clauses: 1226
% SAT solver binary clauses: 249
% SAT solver learnt clauses: 96
% SAT solver learnt literals: 12
% 
% Sat splits: 20
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 4971
% TWLsolver calls for satisfiability: 4082
% 
% Memory used [KB]: 23411
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.rat
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.rat
% SZS output start Proof for thm.rat
fof(f15110,plain,(
  $false),
  inference(subsumption_resolution,[],[f15109,f14942])).
fof(f14942,plain,(
  ( ! [X0] : (s('type.rat.rat','const.rat.rat_|48|_0') = s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X0),s('type.rat.rat',X0)))) )),
  inference(forward_demodulation,[],[f14844,f1734])).
fof(f1734,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))) )),
  inference(cnf_transformation,[],[f752])).
fof(f752,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))),
  inference(flattening,[],[f751])).
fof(f751,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))),
  inference(rectify,[],[f230])).
fof(f230,axiom,(
  ! [X11] : ! [X12] : s(X11,X12) = s(X11,'const.combin.I_1'(s(X11,X12)))),
  file('thm.rat.RAT_CONS_TO_NUM|split|0','thm.combin.I_THM')).
fof(f14844,plain,(
  ( ! [X0] : (s('type.rat.rat','const.rat.rat_|48|_0') = s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat','const.combin.I_1'(s('type.rat.rat',X0))),s('type.rat.rat',X0)))) )),
  inference(unit_resulting_resolution,[],[f1734,f2391])).
fof(f2391,plain,(
  ( ! [X0,X1] : (s('type.rat.rat','const.rat.rat_|48|_0') = s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))) | s('type.rat.rat',X0) != s('type.rat.rat',X1)) )),
  inference(forward_demodulation,[],[f1969,f1431])).
fof(f1431,plain,(
  s('type.rat.rat','const.rat.rat_|48|_0') = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),
  inference(cnf_transformation,[],[f330])).
fof(f330,axiom,(
  s('type.rat.rat','const.rat.rat_|48|_0') = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),
  file('thm.rat.RAT_CONS_TO_NUM|split|0','thm.rat.rat_of_num_def|split|0')).
fof(f1969,plain,(
  ( ! [X0,X1] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))) | s('type.rat.rat',X0) != s('type.rat.rat',X1)) )),
  inference(cnf_transformation,[],[f1313])).
fof(f1313,plain,(
  ! [X0,X1] : ((s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))) | s('type.rat.rat',X0) = s('type.rat.rat',X1)) & (s('type.rat.rat',X0) != s('type.rat.rat',X1) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X0),s('type.rat.rat',X1)))))),
  inference(nnf_transformation,[],[f956])).
fof(f956,plain,(
  ! [X0,X1] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))) <=> s('type.rat.rat',X0) = s('type.rat.rat',X1))),
  inference(rectify,[],[f100])).
fof(f100,axiom,(
  ! [X19,X26] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X19),s('type.rat.rat',X26))) <=> s('type.rat.rat',X19) = s('type.rat.rat',X26))),
  file('thm.rat.RAT_CONS_TO_NUM|split|0','thm.rat.RAT_EQ_|48|SUB')).
fof(f15109,plain,(
  s('type.rat.rat','const.rat.rat_|48|_0') != s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num',sK13))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num',sK13)))))),
  inference(forward_demodulation,[],[f15108,f3047])).
fof(f3047,plain,(
  ( ! [X0] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num',X0))) = s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac','const.frac.frac_save_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.ZERO_0')))))) )),
  inference(backward_demodulation,[],[f3003,f1573])).
fof(f1573,plain,(
  ( ! [X0] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num',X0))) = s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac','const.frac.frac_save_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0')))))) )),
  inference(cnf_transformation,[],[f642])).
fof(f642,plain,(
  ! [X0] : s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num',X0))) = s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac','const.frac.frac_save_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0')))))),
  inference(rectify,[],[f92])).
fof(f92,axiom,(
  ! [X6] : s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num',X6))) = s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac','const.frac.frac_save_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6))),s('type.num.num','const.num.|48|_0')))))),
  file('thm.rat.RAT_CONS_TO_NUM|split|0','thm.rat.RAT_SAVE_NUM')).
fof(f3003,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.num.|48|_0')),
  inference(superposition,[],[f1472,f1458])).
fof(f1458,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f527])).
fof(f527,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f416])).
fof(f416,axiom,(
  ! [X31] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X31),s('type.num.num','const.num.|48|_0')))),
  file('thm.rat.RAT_CONS_TO_NUM|split|0','thm.numRing.num_rewrites|split|4')).
fof(f1472,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f541])).
fof(f541,plain,(
  ! [X0] : s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f371])).
fof(f371,axiom,(
  ! [X31] : s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X31)))),
  file('thm.rat.RAT_CONS_TO_NUM|split|0','thm.numRing.num_rewrites|split|68')).
fof(f15108,plain,(
  s('type.rat.rat','const.rat.rat_|48|_0') != s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac','const.frac.frac_save_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num',sK13)))))),
  inference(forward_demodulation,[],[f15044,f2366])).
fof(f2366,plain,(
  ( ! [X0,X1] : (s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac','const.frac.frac_save_2'(s('type.integer.int',X0),s('type.num.num',X1))))) = s('type.rat.rat','const.rat.rat_cons_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))))) )),
  inference(forward_demodulation,[],[f2365,f1461])).
fof(f1461,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f530])).
fof(f530,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f424])).
fof(f424,axiom,(
  ! [X31] : s('type.num.num',X31) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X31),s('type.num.num','const.num.|48|_0')))),
  file('thm.rat.RAT_CONS_TO_NUM|split|0','thm.numRing.num_rewrites|split|1')).
fof(f2365,plain,(
  ( ! [X0,X1] : (s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac','const.frac.frac_save_2'(s('type.integer.int',X0),s('type.num.num',X1))))) = s('type.rat.rat','const.rat.rat_cons_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0')))))))))) )),
  inference(forward_demodulation,[],[f2364,f1794])).
fof(f1794,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f810])).
fof(f810,plain,(
  ! [X0,X1] : s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f263])).
fof(f263,axiom,(
  ! [X6,X30] : s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X30),s('type.num.num',X6))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X30),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))),
  file('thm.rat.RAT_CONS_TO_NUM|split|0','thm.arithmetic.ADD_CLAUSES|split|3')).
fof(f2364,plain,(
  ( ! [X0,X1] : (s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac','const.frac.frac_save_2'(s('type.integer.int',X0),s('type.num.num',X1))))) = s('type.rat.rat','const.rat.rat_cons_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))))) )),
  inference(forward_demodulation,[],[f2363,f1435])).
fof(f1435,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(cnf_transformation,[],[f177])).
fof(f177,axiom,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  file('thm.rat.RAT_CONS_TO_NUM|split|0','thm.numeral.numeral_distrib|split|13')).
fof(f2363,plain,(
  ( ! [X0,X1] : (s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac','const.frac.frac_save_2'(s('type.integer.int',X0),s('type.num.num',X1))))) = s('type.rat.rat','const.rat.rat_cons_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) )),
  inference(forward_demodulation,[],[f1894,f1824])).
fof(f1824,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f840])).
fof(f840,plain,(
  ! [X0,X1] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f515])).
fof(f515,axiom,(
  ! [X6,X30] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X30))))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X30)))))),
  file('thm.rat.RAT_CONS_TO_NUM|split|0','thm.integer.INT_ADD_CALCULATE|split|2')).
fof(f1894,plain,(
  ( ! [X0,X1] : (s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac','const.frac.frac_save_2'(s('type.integer.int',X0),s('type.num.num',X1))))) = s('type.rat.rat','const.rat.rat_cons_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) )),
  inference(cnf_transformation,[],[f910])).
fof(f910,plain,(
  ! [X0,X1] : s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac','const.frac.frac_save_2'(s('type.integer.int',X0),s('type.num.num',X1))))) = s('type.rat.rat','const.rat.rat_cons_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X9,X10] : s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac','const.frac.frac_save_2'(s('type.integer.int',X9),s('type.num.num',X10))))) = s('type.rat.rat','const.rat.rat_cons_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X10))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))),
  file('thm.rat.RAT_CONS_TO_NUM|split|0','thm.rat.RAT_SAVE_TO_CONS')).
fof(f15044,plain,(
  s('type.rat.rat','const.rat.rat_|48|_0') != s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat','const.rat.rat_cons_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num',sK13)))))),
  inference(unit_resulting_resolution,[],[f3111,f2392])).
fof(f2392,plain,(
  ( ! [X0,X1] : (s('type.rat.rat','const.rat.rat_|48|_0') != s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))) | s('type.rat.rat',X0) = s('type.rat.rat',X1)) )),
  inference(forward_demodulation,[],[f1968,f1431])).
fof(f1968,plain,(
  ( ! [X0,X1] : (s('type.rat.rat',X0) = s('type.rat.rat',X1) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X0),s('type.rat.rat',X1)))) )),
  inference(cnf_transformation,[],[f1313])).
fof(f3111,plain,(
  s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num',sK13))) != s('type.rat.rat','const.rat.rat_cons_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(backward_demodulation,[],[f3003,f2280])).
fof(f2280,plain,(
  s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num',sK13))) != s('type.rat.rat','const.rat.rat_cons_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(backward_demodulation,[],[f1435,f1424])).
fof(f1424,plain,(
  s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num',sK13))) != s('type.rat.rat','const.rat.rat_cons_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),
  inference(cnf_transformation,[],[f1192])).
fof(f1192,plain,(
  s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num',sK13))) != s('type.rat.rat','const.rat.rat_cons_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13])],[f1094])).
fof(f1094,plain,(
  ? [X0] : s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num',X0))) != s('type.rat.rat','const.rat.rat_cons_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),
  inference(ennf_transformation,[],[f519])).
fof(f519,plain,(
  ~! [X0] : s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num',X0))) = s('type.rat.rat','const.rat.rat_cons_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num',X6))) = s('type.rat.rat','const.rat.rat_cons_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num',X6))) = s('type.rat.rat','const.rat.rat_cons_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),
  file('thm.rat.RAT_CONS_TO_NUM|split|0','thm.rat.RAT_CONS_TO_NUM|split|0')).
% SZS output end Proof for thm.rat
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 629
% Passive clauses: 2242
% Generated clauses: 13680
% Final active clauses: 560
% Final passive clauses: 1446
% Input formulas: 517
% Initial clauses: 823
% 
% Duplicate literals: 13
% Trivial inequalities: 33
% Fw subsumption resolutions: 113
% Fw demodulations: 2156
% Bw demodulations: 232
% Condensations: 19
% 
% Simple tautologies: 107
% Equational tautologies: 36
% Forward subsumptions: 8599
% Fw demodulations to eq. taut.: 294
% Bw demodulations to eq. taut.: 4
% 
% Unit resulting resolution: 3254
% Factoring: 1
% Forward superposition: 3906
% Backward superposition: 3022
% Self superposition: 88
% Equality factoring: 2
% Equality resolution: 16
% 
% Split clauses: 3
% Split components: 6
% SAT solver clauses: 2250
% SAT solver unit clauses: 1106
% SAT solver binary clauses: 898
% 
% Sat splits: 3
% 
% Memory used [KB]: 9850
% Time elapsed: 0.315 s
% ------------------------------
% ------------------------------
% Success in time 1.66 s
