Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.extreal
% (53546)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 114
% Passive clauses: 3970
% Generated clauses: 13374
% Final active clauses: 105
% Final passive clauses: 3713
% Input formulas: 47
% Initial clauses: 117
% 
% Duplicate literals: 501
% Trivial inequalities: 165
% Fw subsumption resolutions: 2045
% Fw demodulations: 820
% Bw demodulations: 8
% 
% Simple tautologies: 298
% Equational tautologies: 73
% Forward subsumptions: 4512
% 
% Binary resolution: 625
% Factoring: 2
% Forward superposition: 3595
% Backward superposition: 5059
% Self superposition: 238
% Equality factoring: 182
% Equality resolution: 16
% 
% Split clauses: 452
% Split components: 930
% SAT solver clauses: 1127
% SAT solver unit clauses: 40
% SAT solver binary clauses: 515
% 
% Sat splits: 452
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1126
% TWLsolver calls for satisfiability: 162
% 
% Memory used [KB]: 8827
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.extreal
% (53587)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 157
% Passive clauses: 1305
% Generated clauses: 12233
% Final active clauses: 130
% Final passive clauses: 1024
% Input formulas: 47
% Initial clauses: 117
% 
% Duplicate literals: 198
% Trivial inequalities: 172
% Fw subsumption resolutions: 1003
% Fw demodulations: 2477
% Bw demodulations: 133
% Condensations: 197
% Global subsumptions: 12
% 
% Simple tautologies: 162
% Equational tautologies: 15
% Forward subsumptions: 4834
% Bw demodulations to eq. taut.: 16
% 
% Binary resolution: 350
% Factoring: 1
% Forward superposition: 3313
% Backward superposition: 2811
% Self superposition: 66
% Equality factoring: 22
% Equality resolution: 25
% 
% Split clauses: 51
% Split components: 102
% SAT solver clauses: 1969
% SAT solver unit clauses: 130
% SAT solver binary clauses: 42
% SAT solver learnt clauses: 24
% 
% Sat splits: 51
% Sat splitting refutations: 7
% 
% TWLsolver clauses: 1921
% TWLsolver calls for satisfiability: 705
% 
% Memory used [KB]: 17014
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.extreal
% (53623)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 640
% Passive clauses: 11355
% Generated clauses: 142648
% Final active clauses: 179
% Final passive clauses: 1320
% Input formulas: 47
% Initial clauses: 117
% Discarded non-redundant clauses: 1683
% 
% Duplicate literals: 4399
% Trivial inequalities: 1057
% Fw subsumption resolutions: 13325
% Fw demodulations: 22930
% Bw demodulations: 7129
% Condensations: 9141
% Global subsumptions: 1485
% 
% Simple tautologies: 1367
% Equational tautologies: 408
% Forward subsumptions: 73971
% Fw demodulations to eq. taut.: 55
% Bw demodulations to eq. taut.: 366
% 
% Unit resulting resolution: 2348
% Factoring: 13
% Forward superposition: 36646
% Backward superposition: 26182
% Self superposition: 587
% Equality factoring: 283
% Equality resolution: 87
% 
% Split clauses: 134
% Split components: 269
% SAT solver clauses: 14838
% SAT solver unit clauses: 615
% SAT solver binary clauses: 2457
% 
% Sat splits: 134
% Sat splitting refutations: 1728
% 
% Memory used [KB]: 87120
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.extreal
% (53868)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 564
% Passive clauses: 15037
% Generated clauses: 96992
% Final active clauses: 225
% Final passive clauses: 7758
% Input formulas: 47
% Initial clauses: 86
% 
% Selected by SInE selection: 37
% SInE iterations: 1
% 
% Duplicate literals: 2544
% Trivial inequalities: 1726
% Fw subsumption resolutions: 18384
% Fw demodulations: 11026
% Bw demodulations: 134
% 
% Simple tautologies: 2256
% Equational tautologies: 83
% Forward subsumptions: 39870
% Bw demodulations to eq. taut.: 3
% 
% Binary resolution: 10037
% Factoring: 1
% Forward superposition: 26575
% Backward superposition: 24028
% Self superposition: 386
% Equality factoring: 265
% Equality resolution: 99
% 
% Split clauses: 7311
% Split components: 15696
% SAT solver clauses: 7497
% SAT solver unit clauses: 38
% SAT solver binary clauses: 892
% 
% Sat splits: 7311
% Sat splitting refutations: 67
% 
% Memory used [KB]: 32238
% Time elapsed: 4.0000 s
% ------------------------------
% ------------------------------
% remaining time: 142 next slice time: 6
lrs+1002_2:3_br=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:sas=minisat:stl=30:sac=on:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none:sp=reverse_arity:urr=on_4 on thm.extreal
% (53944)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 101
% Passive clauses: 2039
% Generated clauses: 14806
% Final active clauses: 66
% Final passive clauses: 1086
% Input formulas: 47
% Initial clauses: 117
% 
% Duplicate literals: 523
% Trivial inequalities: 262
% Fw subsumption resolutions: 1842
% Fw demodulations: 1263
% Bw demodulations: 837
% Global subsumptions: 5
% 
% Simple tautologies: 634
% Equational tautologies: 130
% Forward subsumptions: 7857
% Bw demodulations to eq. taut.: 4
% 
% Unit resulting resolution: 1132
% Forward superposition: 4809
% Backward superposition: 3041
% Self superposition: 83
% Equality factoring: 46
% Equality resolution: 9
% 
% Split clauses: 14
% Split components: 29
% SAT solver clauses: 2216
% SAT solver unit clauses: 136
% SAT solver binary clauses: 234
% 
% Sat splits: 14
% Sat splitting refutations: 156
% 
% Memory used [KB]: 8955
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 135 next slice time: 15
lrs+1011_12_bs=on:bsr=unit_only:cond=on:gs=on:gsaa=from_current:gsssp=full:nwc=1.1:sas=minisat:stl=60:sos=all:sac=on:sdd=off:sser=off:sfr=on:ssfp=100000:ssfq=1.2:ssnc=none:sp=reverse_arity:updr=off_14 on thm.extreal
% (53961)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 200
% Passive clauses: 3358
% Generated clauses: 20414
% Final active clauses: 189
% Final passive clauses: 3174
% Input formulas: 47
% Initial clauses: 117
% 
% Duplicate literals: 538
% Trivial inequalities: 440
% Fw subsumption resolutions: 1450
% Fw demodulations: 2949
% Bw demodulations: 38
% Condensations: 922
% 
% Simple tautologies: 511
% Equational tautologies: 17
% Forward subsumptions: 10290
% Backward subsumptions: 31
% 
% Binary resolution: 2207
% Factoring: 1
% Forward superposition: 6308
% Backward superposition: 5520
% Self superposition: 22
% Equality factoring: 36
% Equality resolution: 27
% 
% Split clauses: 1
% Split components: 2
% SAT solver clauses: 3361
% SAT solver unit clauses: 58
% SAT solver binary clauses: 613
% 
% Sat splits: 1
% 
% Memory used [KB]: 20724
% Time elapsed: 1.600 s
% ------------------------------
% ------------------------------
% remaining time: 119 next slice time: 9
lrs+10_8:1_bd=preordered:bs=on:ccuc=first:fsr=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:nicw=on:sas=minisat:stl=120:sos=on:sscc=on:sser=off:sfr=on:ssfp=4000:ssfq=1.0:ssnc=none:sp=reverse_arity:urr=on_7 on thm.extreal
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.extreal
% SZS output start Proof for thm.extreal
fof(f2344,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f193,f314,f409,f407,f406,f405,f403,f402,f401,f398,f397,f396,f394,f393,f388,f386,f385,f384,f382,f381,f380,f378,f377,f376,f374,f373,f372,f370,f369,f368,f366,f365,f364,f362,f361,f360,f425,f358,f357,f427,f356,f353,f352,f351,f348,f347,f346,f343,f342,f341,f338,f337,f336,f421,f334,f333,f329,f327,f320,f319,f318,f317,f423,f411,f412,f433,f619,f610,f600,f575,f574,f573,f553,f545,f536,f535,f527,f519,f510,f509,f641,f645,f926,f932,f934,f1017,f1018,f1023,f1022,f1068,f1149,f1069,f1152,f1071,f1154,f1072,f1155,f1073,f1156,f1253,f1252,f1248,f1247,f1245,f1244,f1243,f1242,f1235,f1234,f1231,f1230,f1225,f1224,f1223,f1222,f1218,f1216,f1212,f1211,f1210,f1209,f1207,f1206,f1205,f1204,f1200,f1198,f1194,f1193,f1192,f1191,f1187,f1185,f1179,f1178,f1177,f1176,f1271,f1257,f1258,f1294,f1369,f1368,f1367,f1363,f1360,f1359,f1358,f1357,f1356,f1355,f1354,f1385,f1388,f1391,f1390,f1392,f1387,f1393,f1396,f1395,f1346,f1345,f1343,f1342,f1336,f1335,f1381,f1334,f1333,f1372,f1373,f1399,f1398,f1400,f2229,f2322,f2323,f2325,f1151,f2343])).
fof(f2343,plain,(
  s(bool,'F_0') != s(bool,'const.util_prob.countable_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) | p(s(bool,'const.util_prob.countable_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')))))) | ~p(s(bool,'F_0'))),
  introduced(theory_axiom,[])).
fof(f1151,plain,(
  ~p(s(bool,'T_0')) | $spl21),
  inference(cnf_transformation,[],[f1151_D])).
fof(f1151_D,plain,(
  ~p(s(bool,'T_0')) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f2325,plain,(
  p(s(bool,'F_0')) | $spl46),
  inference(cnf_transformation,[],[f2325_D])).
fof(f2325_D,plain,(
  p(s(bool,'F_0')) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f2323,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'T_0')) | ($spl1 | $spl14)),
  inference(forward_demodulation,[],[f2230,f2077])).
fof(f2077,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.pred_set.SURJ_3'(s(fun('type.num.num','type.pair.prod'('type.num.num','type.num.num')),X0),s(fun('type.num.num',bool),X1),s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f200,f938,f257])).
fof(f257,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f127])).
fof(f127,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f126])).
fof(f126,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f88])).
fof(f88,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X7,X8] : ((p(s(bool,X7)) => p(s(bool,X8))) => ((p(s(bool,X8)) => p(s(bool,X7))) => s(bool,X7) = s(bool,X8)))),
  file('thm.extreal.CROSS_COUNTABLE_UNIV','thm.bool.IMP_ANTISYM_AX')).
fof(f938,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.pred_set.SURJ_3'(s(fun('type.num.num','type.pair.prod'('type.num.num','type.num.num')),X0),s(fun('type.num.num',bool),X1),s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f316,f315,f275])).
fof(f275,plain,(
  ( ! [X4,X2,X0,X8,X3,X1] : (~p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X8),s(fun(X1,bool),X4)))) | s(X1,X8) = s(X1,happ(s(fun(X0,X1),X2),s(X0,sK5(X8,X4,X3,X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f182])).
fof(f182,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | (! [X10] : (~p(s(bool,'const.bool.IN_2'(s(X0,X10),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X10))),s(fun(X1,bool),X4))))) & ! [X8] : (~p(s(bool,'const.bool.IN_2'(s(X1,X8),s(fun(X1,bool),X4)))) | (p(s(bool,'const.bool.IN_2'(s(X0,sK5(X8,X4,X3,X2,X1,X0)),s(fun(X0,bool),X3)))) & s(X1,X8) = s(X1,happ(s(fun(X0,X1),X2),s(X0,sK5(X8,X4,X3,X2,X1,X0)))))))) & ((p(s(bool,'const.bool.IN_2'(s(X0,sK6(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,sK6(X4,X3,X2,X1,X0)))),s(fun(X1,bool),X4))))) | (p(s(bool,'const.bool.IN_2'(s(X1,sK7(X4,X3,X2,X1,X0)),s(fun(X1,bool),X4)))) & ! [X6] : (~p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X3)))) | s(X1,sK7(X4,X3,X2,X1,X0)) != s(X1,happ(s(fun(X0,X1),X2),s(X0,X6))))) | p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5,sK6,sK7])],[f181])).
fof(f181,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | (! [X10] : (~p(s(bool,'const.bool.IN_2'(s(X0,X10),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X10))),s(fun(X1,bool),X4))))) & ! [X8] : (~p(s(bool,'const.bool.IN_2'(s(X1,X8),s(fun(X1,bool),X4)))) | ? [X9] : (p(s(bool,'const.bool.IN_2'(s(X0,X9),s(fun(X0,bool),X3)))) & s(X1,X8) = s(X1,happ(s(fun(X0,X1),X2),s(X0,X9))))))) & (? [X7] : (p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X7))),s(fun(X1,bool),X4))))) | ? [X5] : (p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X4)))) & ! [X6] : (~p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X3)))) | s(X1,X5) != s(X1,happ(s(fun(X0,X1),X2),s(X0,X6))))) | p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))),
  inference(rectify,[],[f180])).
fof(f180,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | (! [X7] : (~p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X7))),s(fun(X1,bool),X4))))) & ! [X5] : (~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X4)))) | ? [X6] : (p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X3)))) & s(X1,X5) = s(X1,happ(s(fun(X0,X1),X2),s(X0,X6))))))) & (? [X7] : (p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X7))),s(fun(X1,bool),X4))))) | ? [X5] : (p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X4)))) & ! [X6] : (~p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X3)))) | s(X1,X5) != s(X1,happ(s(fun(X0,X1),X2),s(X0,X6))))) | p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))),
  inference(flattening,[],[f179])).
fof(f179,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | (! [X7] : (~p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X7))),s(fun(X1,bool),X4))))) & ! [X5] : (~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X4)))) | ? [X6] : (p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X3)))) & s(X1,X5) = s(X1,happ(s(fun(X0,X1),X2),s(X0,X6))))))) & ((? [X7] : (p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X7))),s(fun(X1,bool),X4))))) | ? [X5] : (p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X4)))) & ! [X6] : (~p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X3)))) | s(X1,X5) != s(X1,happ(s(fun(X0,X1),X2),s(X0,X6)))))) | p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))),
  inference(nnf_transformation,[],[f133])).
fof(f133,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) <=> (! [X7] : (~p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X7))),s(fun(X1,bool),X4))))) & ! [X5] : (~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X4)))) | ? [X6] : (p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X3)))) & s(X1,X5) = s(X1,happ(s(fun(X0,X1),X2),s(X0,X6)))))))),
  inference(ennf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) <=> (! [X7] : (p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),X3)))) => p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X7))),s(fun(X1,bool),X4))))) & ! [X5] : (p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X4)))) => ? [X6] : (p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X3)))) & s(X1,X5) = s(X1,happ(s(fun(X0,X1),X2),s(X0,X6)))))))),
  inference(flattening,[],[f103])).
fof(f103,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) <=> (! [X7] : (p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),X3)))) => p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X7))),s(fun(X1,bool),X4))))) & ! [X5] : (p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X4)))) => ? [X6] : (p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X3)))) & s(X1,X5) = s(X1,happ(s(fun(X0,X1),X2),s(X0,X6)))))))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X6,X15] : ! [X16,X17,X9] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(X6,X15),X16),s(fun(X6,bool),X17),s(fun(X15,bool),X9)))) <=> (! [X10] : (p(s(bool,'const.bool.IN_2'(s(X6,X10),s(fun(X6,bool),X17)))) => p(s(bool,'const.bool.IN_2'(s(X15,happ(s(fun(X6,X15),X16),s(X6,X10))),s(fun(X15,bool),X9))))) & ! [X10] : (p(s(bool,'const.bool.IN_2'(s(X15,X10),s(fun(X15,bool),X9)))) => ? [X11] : (p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X17)))) & s(X15,X10) = s(X15,happ(s(fun(X6,X15),X16),s(X6,X11)))))))),
  file('thm.extreal.CROSS_COUNTABLE_UNIV','thm.pred_set.SURJ_DEF')).
fof(f315,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK3(X0,'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'type.pair.prod'('type.num.num','type.num.num'))),s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f314,f252])).
fof(f252,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK3(X2,X1,X0)),s(fun(X0,bool),X1)))) | p(s(bool,'const.util_prob.countable_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f172])).
fof(f172,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.util_prob.countable_1'(s(fun(X0,bool),X1)))) | ! [X6] : (~p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X1)))) | s(X0,X6) = s(X0,happ(s(fun('type.num.num',X0),sK1(X1,X0)),s('type.num.num',sK2(X6,X1,X0)))))) & (! [X2] : (p(s(bool,'const.bool.IN_2'(s(X0,sK3(X2,X1,X0)),s(fun(X0,bool),X1)))) & ! [X4] : s(X0,sK3(X2,X1,X0)) != s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X4)))) | p(s(bool,'const.util_prob.countable_1'(s(fun(X0,bool),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1,sK2,sK3])],[f171])).
fof(f171,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.util_prob.countable_1'(s(fun(X0,bool),X1)))) | ? [X5] : ! [X6] : (~p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X1)))) | ? [X7] : s(X0,X6) = s(X0,happ(s(fun('type.num.num',X0),X5),s('type.num.num',X7))))) & (! [X2] : ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ! [X4] : s(X0,X3) != s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X4)))) | p(s(bool,'const.util_prob.countable_1'(s(fun(X0,bool),X1))))))),
  inference(rectify,[],[f170])).
fof(f170,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.util_prob.countable_1'(s(fun(X0,bool),X1)))) | ? [X2] : ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ? [X4] : s(X0,X3) = s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X4))))) & (! [X2] : ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ! [X4] : s(X0,X3) != s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X4)))) | p(s(bool,'const.util_prob.countable_1'(s(fun(X0,bool),X1))))))),
  inference(nnf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0,X1] : (p(s(bool,'const.util_prob.countable_1'(s(fun(X0,bool),X1)))) <=> ? [X2] : ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ? [X4] : s(X0,X3) = s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X4)))))),
  inference(ennf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0,X1] : (p(s(bool,'const.util_prob.countable_1'(s(fun(X0,bool),X1)))) <=> ? [X2] : ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) => ? [X4] : s(X0,X3) = s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X4)))))),
  inference(flattening,[],[f82])).
fof(f82,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.util_prob.countable_1'(s(fun(X0,bool),X1)))) <=> ? [X2] : ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) => ? [X4] : s(X0,X3) = s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X4)))))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X6] : ! [X17] : (p(s(bool,'const.util_prob.countable_1'(s(fun(X6,bool),X17)))) <=> ? [X16] : ! [X10] : (p(s(bool,'const.bool.IN_2'(s(X6,X10),s(fun(X6,bool),X17)))) => ? [X22] : s(X6,X10) = s(X6,happ(s(fun('type.num.num',X6),X16),s('type.num.num',X22)))))),
  file('thm.extreal.CROSS_COUNTABLE_UNIV','thm.util_prob.countable_def')).
fof(f316,plain,(
  ( ! [X0,X1] : (s('type.pair.prod'('type.num.num','type.num.num'),happ(s(fun('type.num.num','type.pair.prod'('type.num.num','type.num.num')),X0),s('type.num.num',X1))) != s('type.pair.prod'('type.num.num','type.num.num'),sK3(X0,'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'type.pair.prod'('type.num.num','type.num.num')))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f314,f253])).
fof(f253,plain,(
  ( ! [X4,X2,X0,X1] : (s(X0,sK3(X2,X1,X0)) != s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X4))) | p(s(bool,'const.util_prob.countable_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f172])).
fof(f200,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f52])).
fof(f52,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.extreal.CROSS_COUNTABLE_UNIV',aHL_FALSITY)).
fof(f2230,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.pred_set.SURJ_3'(s(fun('type.num.num','type.pair.prod'('type.num.num','type.num.num')),sK11),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')))))) | $spl14),
  inference(superposition,[],[f271,f934])).
fof(f271,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))) )),
  inference(cnf_transformation,[],[f178])).
fof(f178,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | (p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) & p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))))) & (~p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | ~p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))),
  inference(flattening,[],[f177])).
fof(f177,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | (p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) & p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))))) & ((~p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | ~p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))) | p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))),
  inference(nnf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) <=> (p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) & p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))),
  inference(flattening,[],[f101])).
fof(f101,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) <=> (p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) & p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X6,X15] : ! [X16,X17,X9] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(X6,X15),X16),s(fun(X6,bool),X17),s(fun(X15,bool),X9)))) <=> (p(s(bool,'const.pred_set.INJ_3'(s(fun(X6,X15),X16),s(fun(X6,bool),X17),s(fun(X15,bool),X9)))) & p(s(bool,'const.pred_set.SURJ_3'(s(fun(X6,X15),X16),s(fun(X6,bool),X17),s(fun(X15,bool),X9))))))),
  file('thm.extreal.CROSS_COUNTABLE_UNIV','thm.pred_set.BIJ_DEF')).
fof(f2322,plain,(
  p(s(bool,'const.pred_set.INJ_3'(s(fun('type.num.num','type.pair.prod'('type.num.num','type.num.num')),sK11),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')))))) | $spl44),
  inference(cnf_transformation,[],[f2322_D])).
fof(f2322_D,plain,(
  p(s(bool,'const.pred_set.INJ_3'(s(fun('type.num.num','type.pair.prod'('type.num.num','type.num.num')),sK11),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')))))) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f2229,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.pred_set.INJ_3'(s(fun('type.num.num','type.pair.prod'('type.num.num','type.num.num')),sK11),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')))))) | $spl14),
  inference(superposition,[],[f270,f934])).
fof(f270,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))) )),
  inference(cnf_transformation,[],[f178])).
fof(f1400,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,'T_0'))) | p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,'T_0'))))),
  introduced(theory_axiom,[])).
fof(f1398,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,'T_0'))) | $spl42),
  inference(cnf_transformation,[],[f1398_D])).
fof(f1398_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,'T_0'))) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f1399,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,'T_0'))) | $spl16),
  inference(global_subsumption,[],[f193,f1022,f1369,f1363,f1360,f1359,f1358,f1385,f1388,f1391,f1392,f1393,f1396,f1346,f1345,f1343,f1342,f1336,f1335,f1334,f1333,f1372,f1373,f1378,f1379])).
fof(f1379,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,'T_0'))) | $spl16),
  inference(superposition,[],[f1022,f201])).
fof(f201,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.extreal.CROSS_COUNTABLE_UNIV',aHL_BOOL_CASES)).
fof(f1378,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,'T_0')))) ) | $spl16),
  inference(resolution,[],[f1022,f260])).
fof(f260,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f127])).
fof(f1373,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,'T_0'))) | $spl16),
  inference(resolution,[],[f1022,f210])).
fof(f210,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X9] : (s(bool,'T_0') = s(bool,X9) <=> p(s(bool,X9)))),
  file('thm.extreal.CROSS_COUNTABLE_UNIV','thm.bool.EQ_CLAUSES|split|1')).
fof(f1372,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,'T_0'))) | $spl16),
  inference(resolution,[],[f1022,f208])).
fof(f208,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f144])).
fof(f144,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X9] : (s(bool,'T_0') = s(bool,X9) <=> p(s(bool,X9)))),
  file('thm.extreal.CROSS_COUNTABLE_UNIV','thm.bool.EQ_CLAUSES|split|0')).
fof(f1333,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,'T_0'))) | $spl16),
  inference(unit_resulting_resolution,[],[f1022,f208])).
fof(f1334,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,'T_0'))) | $spl16),
  inference(unit_resulting_resolution,[],[f1022,f210])).
fof(f1381,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,'T_0'))) | $spl35),
  inference(cnf_transformation,[],[f1381_D])).
fof(f1381_D,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,'T_0'))) <=> ~$spl35),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl35])])).
fof(f1335,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,'T_0'))) | $spl16),
  inference(unit_resulting_resolution,[],[f1022,f211])).
fof(f211,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f59])).
fof(f59,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X9] : (s(bool,'F_0') = s(bool,X9) <=> ~p(s(bool,X9)))),
  file('thm.extreal.CROSS_COUNTABLE_UNIV','thm.bool.EQ_CLAUSES|split|3')).
fof(f1336,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,'T_0'))) | $spl16),
  inference(unit_resulting_resolution,[],[f1022,f213])).
fof(f213,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f62])).
fof(f62,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f61])).
fof(f61,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X9] : (s(bool,'F_0') = s(bool,X9) <=> ~p(s(bool,X9)))),
  file('thm.extreal.CROSS_COUNTABLE_UNIV','thm.bool.EQ_CLAUSES|split|2')).
fof(f1342,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,'T_0'))) | $spl16),
  inference(unit_resulting_resolution,[],[f195,f1022,f260])).
fof(f195,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f8])).
fof(f8,axiom,(
  p(s(bool,'T_0'))),
  file('thm.extreal.CROSS_COUNTABLE_UNIV','thm.bool.TRUTH')).
fof(f1343,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,'T_0'))) | $spl16),
  inference(unit_resulting_resolution,[],[f194,f1022,f260])).
fof(f194,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.extreal.CROSS_COUNTABLE_UNIV',aHL_TRUTH)).
fof(f1345,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,'T_0'))) | $spl16),
  inference(unit_resulting_resolution,[],[f195,f1022,f260])).
fof(f1346,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,'T_0'))) | $spl16),
  inference(unit_resulting_resolution,[],[f194,f1022,f260])).
fof(f1395,plain,(
  p(s(bool,happ(s(fun(bool,bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | $spl40),
  inference(cnf_transformation,[],[f1395_D])).
fof(f1395_D,plain,(
  p(s(bool,happ(s(fun(bool,bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'T_0')))) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f1396,plain,(
  p(s(bool,happ(s(fun(bool,bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | $spl16),
  inference(forward_demodulation,[],[f1348,f307])).
fof(f307,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun(X2,X0),'_dst_x0x1x0_3'(s(fun(X1,fun(X3,X0)),X4),s(fun(X2,X1),X5),s(X3,X6))),s(X2,X7))) = s(X0,happ(s(fun(X3,X0),happ(s(fun(X1,fun(X3,X0)),X4),s(X1,happ(s(fun(X2,X1),X5),s(X2,X7))))),s(X3,X6)))) )),
  inference(cnf_transformation,[],[f112])).
fof(f112,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun(X2,X0),'_dst_x0x1x0_3'(s(fun(X1,fun(X3,X0)),X4),s(fun(X2,X1),X5),s(X3,X6))),s(X2,X7))) = s(X0,happ(s(fun(X3,X0),happ(s(fun(X1,fun(X3,X0)),X4),s(X1,happ(s(fun(X2,X1),X5),s(X2,X7))))),s(X3,X6)))),
  inference(flattening,[],[f111])).
fof(f111,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6,X7] : s(X0,happ(s(fun(X2,X0),'_dst_x0x1x0_3'(s(fun(X1,fun(X3,X0)),X4),s(fun(X2,X1),X5),s(X3,X6))),s(X2,X7))) = s(X0,happ(s(fun(X3,X0),happ(s(fun(X1,fun(X3,X0)),X4),s(X1,happ(s(fun(X2,X1),X5),s(X2,X7))))),s(X3,X6)))),
  inference(rectify,[],[f48])).
fof(f48,axiom,(
  ! [X27,X28,X29,X33] : ! [X30,X31,X34,X32] : s(X27,happ(s(fun(X29,X27),'_dst_x0x1x0_3'(s(fun(X28,fun(X33,X27)),X30),s(fun(X29,X28),X31),s(X33,X34))),s(X29,X32))) = s(X27,happ(s(fun(X33,X27),happ(s(fun(X28,fun(X33,X27)),X30),s(X28,happ(s(fun(X29,X28),X31),s(X29,X32))))),s(X33,X34)))),
  file('thm.extreal.CROSS_COUNTABLE_UNIV','_dst_x0x1x0_ax')).
fof(f1348,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,'T_0'))))),s(bool,'T_0')))) | $spl16),
  inference(unit_resulting_resolution,[],[f195,f1022,f265])).
fof(f265,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X0))),s(bool,X1)))) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f175])).
fof(f175,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X0))),s(bool,X1)))) | (p(s(bool,X0)) & p(s(bool,X1)))) & (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X0))),s(bool,X1))))))),
  inference(flattening,[],[f174])).
fof(f174,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X0))),s(bool,X1)))) | (p(s(bool,X0)) & p(s(bool,X1)))) & ((~p(s(bool,X0)) | ~p(s(bool,X1))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X0))),s(bool,X1))))))),
  inference(nnf_transformation,[],[f92])).
fof(f92,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X0))),s(bool,X1)))) <=> (p(s(bool,X0)) & p(s(bool,X1))))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X25,X26] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X25))),s(bool,X26)))) <=> (p(s(bool,X25)) & p(s(bool,X26))))),
  file('thm.extreal.CROSS_COUNTABLE_UNIV','|slash||bslash|_HL_0')).
fof(f1393,plain,(
  p(s(bool,happ(s(fun(bool,bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | $spl16),
  inference(forward_demodulation,[],[f1349,f307])).
fof(f1349,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,'T_0'))))),s(bool,'T_0')))) | $spl16),
  inference(unit_resulting_resolution,[],[f194,f1022,f265])).
fof(f1387,plain,(
  p(s(bool,happ(s(fun(bool,bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,'T_0'))))),s(bool,'T_0')))) | $spl36),
  inference(cnf_transformation,[],[f1387_D])).
fof(f1387_D,plain,(
  p(s(bool,happ(s(fun(bool,bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,'T_0'))))),s(bool,'T_0')))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f1392,plain,(
  p(s(bool,happ(s(fun(bool,bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,'T_0'))))),s(bool,'T_0')))) | $spl16),
  inference(forward_demodulation,[],[f1350,f307])).
fof(f1350,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,'T_0'))))),s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,'T_0')))))) | $spl16),
  inference(unit_resulting_resolution,[],[f1022,f1022,f265])).
fof(f1390,plain,(
  p(s(bool,happ(s(fun(bool,bool),'_dst_x1x1_2'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))))),s(bool,'T_0')))) | $spl38),
  inference(cnf_transformation,[],[f1390_D])).
fof(f1390_D,plain,(
  p(s(bool,happ(s(fun(bool,bool),'_dst_x1x1_2'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))))),s(bool,'T_0')))) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f1391,plain,(
  p(s(bool,happ(s(fun(bool,bool),'_dst_x1x1_2'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))))),s(bool,'T_0')))) | $spl16),
  inference(forward_demodulation,[],[f1351,f298])).
fof(f298,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X2,X0),'_dst_x1x1_2'(s(fun(X2,fun(X1,X0)),X3),s(fun(X2,X1),X4))),s(X2,X5))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(X2,fun(X1,X0)),X3),s(X2,X5))),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5)))))) )),
  inference(cnf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X2,X0),'_dst_x1x1_2'(s(fun(X2,fun(X1,X0)),X3),s(fun(X2,X1),X4))),s(X2,X5))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(X2,fun(X1,X0)),X3),s(X2,X5))),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5)))))),
  inference(flattening,[],[f108])).
fof(f108,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X2,X0),'_dst_x1x1_2'(s(fun(X2,fun(X1,X0)),X3),s(fun(X2,X1),X4))),s(X2,X5))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(X2,fun(X1,X0)),X3),s(X2,X5))),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5)))))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X27,X28,X29] : ! [X30,X31,X32] : s(X27,happ(s(fun(X29,X27),'_dst_x1x1_2'(s(fun(X29,fun(X28,X27)),X30),s(fun(X29,X28),X31))),s(X29,X32))) = s(X27,happ(s(fun(X28,X27),happ(s(fun(X29,fun(X28,X27)),X30),s(X29,X32))),s(X28,happ(s(fun(X29,X28),X31),s(X29,X32)))))),
  file('thm.extreal.CROSS_COUNTABLE_UNIV','_dst_x1x1_ax')).
fof(f1351,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,'T_0')))))) | $spl16),
  inference(unit_resulting_resolution,[],[f195,f1022,f265])).
fof(f1388,plain,(
  p(s(bool,happ(s(fun(bool,bool),'_dst_x1x1_2'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))))),s(bool,'T_0')))) | $spl16),
  inference(forward_demodulation,[],[f1352,f298])).
fof(f1352,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,'T_0')))))) | $spl16),
  inference(unit_resulting_resolution,[],[f194,f1022,f265])).
fof(f1385,plain,(
  p(s(bool,happ(s(fun(bool,bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,'T_0'))))),s(bool,'T_0')))) | $spl16),
  inference(forward_demodulation,[],[f1353,f307])).
fof(f1353,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,'T_0'))))),s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,'T_0')))))) | $spl16),
  inference(unit_resulting_resolution,[],[f1022,f1022,f265])).
fof(f1354,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,'T_0'))) | ($spl10 | $spl16)),
  inference(unit_resulting_resolution,[],[f200,f200,f433,f1022,f300])).
fof(f300,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X2) != s(bool,X3) | ~p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,X3)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f191])).
fof(f191,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f190])).
fof(f190,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f139])).
fof(f139,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f138])).
fof(f138,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f110])).
fof(f110,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X10,X13,X11,X14] : ((s(bool,X10) = s(bool,X13) & (p(s(bool,X13)) => s(bool,X11) = s(bool,X14))) => ((p(s(bool,X10)) => p(s(bool,X11))) <=> (p(s(bool,X13)) => p(s(bool,X14)))))),
  file('thm.extreal.CROSS_COUNTABLE_UNIV','thm.bool.IMP_CONG')).
fof(f1355,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,'T_0'))) | ($spl1 | $spl16)),
  inference(unit_resulting_resolution,[],[f200,f195,f454,f1022,f302])).
fof(f302,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X2) != s(bool,X3) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f191])).
fof(f454,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f404,f201])).
fof(f404,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f249,f200,f314,f312])).
fof(f312,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(duplicate_literal_removal,[],[f303])).
fof(f303,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f191])).
fof(f249,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0'))))) )),
  inference(cnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0,X1] : p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0'))))),
  inference(flattening,[],[f78])).
fof(f78,plain,(
  ! [X0] : ! [X1] : p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0'))))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X6] : ! [X10] : p(s(bool,'const.bool.IN_2'(s(X6,X10),s(fun(X6,bool),'const.pred_set.UNIV_0'))))),
  file('thm.extreal.CROSS_COUNTABLE_UNIV','thm.pred_set.IN_UNIV')).
fof(f1356,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,'T_0'))) | ($spl1 | $spl16)),
  inference(unit_resulting_resolution,[],[f200,f194,f454,f1022,f302])).
fof(f1357,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,'T_0'))) | ($spl10 | $spl16)),
  inference(unit_resulting_resolution,[],[f200,f200,f433,f1022,f304])).
fof(f304,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X2) != s(bool,X3) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f191])).
fof(f1358,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,'T_0'))) | $spl16),
  inference(unit_resulting_resolution,[],[f200,f200,f195,f1022,f305])).
fof(f305,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f191])).
fof(f1359,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,'T_0'))) | $spl16),
  inference(unit_resulting_resolution,[],[f200,f200,f194,f1022,f305])).
fof(f1360,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,'T_0'))) | $spl16),
  inference(unit_resulting_resolution,[],[f200,f200,f1022,f1022,f305])).
fof(f1363,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,'T_0'))) | $spl16),
  inference(unit_resulting_resolution,[],[f200,f200,f1022,f1022,f305])).
fof(f1367,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,'T_0'))) | ($spl1 | $spl16)),
  inference(unit_resulting_resolution,[],[f200,f195,f454,f1022,f306])).
fof(f306,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X2) != s(bool,X3) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X2)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f191])).
fof(f1368,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,'T_0'))) | ($spl1 | $spl16)),
  inference(unit_resulting_resolution,[],[f200,f194,f454,f1022,f306])).
fof(f1369,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,'T_0'))) | $spl16),
  inference(unit_resulting_resolution,[],[f200,f200,f1022,f312])).
fof(f1294,plain,(
  s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0'))))) | $spl32),
  inference(cnf_transformation,[],[f1294_D])).
fof(f1294_D,plain,(
  s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0'))))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f1258,plain,(
  s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0'))))) | $spl13),
  inference(resolution,[],[f645,f214])).
fof(f214,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f147])).
fof(f1257,plain,(
  s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0'))))) | $spl13),
  inference(resolution,[],[f645,f212])).
fof(f212,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f146])).
fof(f1271,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0'))))) | $spl31),
  inference(cnf_transformation,[],[f1271_D])).
fof(f1271_D,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0'))))) <=> ~$spl31),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl31])])).
fof(f1176,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0'))))) | $spl13),
  inference(unit_resulting_resolution,[],[f645,f207])).
fof(f207,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f144])).
fof(f1177,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0'))))) | $spl13),
  inference(unit_resulting_resolution,[],[f645,f209])).
fof(f209,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f145])).
fof(f1178,plain,(
  s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0'))))) | $spl13),
  inference(unit_resulting_resolution,[],[f645,f212])).
fof(f1179,plain,(
  s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0'))))) | $spl13),
  inference(unit_resulting_resolution,[],[f645,f214])).
fof(f1185,plain,(
  s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0'))))) | $spl13),
  inference(unit_resulting_resolution,[],[f200,f645,f257])).
fof(f1187,plain,(
  s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0'))))) | $spl13),
  inference(unit_resulting_resolution,[],[f200,f645,f257])).
fof(f1191,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0'))))) | ($spl10 | $spl13)),
  inference(unit_resulting_resolution,[],[f195,f200,f433,f645,f300])).
fof(f1192,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0'))))) | ($spl10 | $spl13)),
  inference(unit_resulting_resolution,[],[f194,f200,f433,f645,f300])).
fof(f1193,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0'))))) | $spl13),
  inference(unit_resulting_resolution,[],[f195,f645,f250,f645,f300])).
fof(f250,plain,(
  ( ! [X0,X1] : (s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) )),
  inference(cnf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0,X1] : s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f80])).
fof(f80,plain,(
  ! [X0] : ! [X1] : s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6] : ! [X5] : s(bool,'|dollar|exists_1'(s(fun(X6,bool),X5))) = s(bool,happ(s(fun(X6,bool),X5),s(X6,'const.min.@_1'(s(fun(X6,bool),X5)))))),
  file('thm.extreal.CROSS_COUNTABLE_UNIV','thm.bool.EXISTS_DEF')).
fof(f1194,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0'))))) | $spl13),
  inference(unit_resulting_resolution,[],[f194,f645,f250,f645,f300])).
fof(f1198,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0'))))) | $spl13),
  inference(unit_resulting_resolution,[],[f645,f195,f250,f645,f300])).
fof(f1200,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0'))))) | $spl13),
  inference(unit_resulting_resolution,[],[f645,f194,f250,f645,f300])).
fof(f1204,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0'))))) | ($spl1 | $spl13)),
  inference(unit_resulting_resolution,[],[f195,f195,f454,f645,f302])).
fof(f1205,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0'))))) | ($spl1 | $spl13)),
  inference(unit_resulting_resolution,[],[f194,f195,f454,f645,f302])).
fof(f1206,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0'))))) | ($spl1 | $spl13)),
  inference(unit_resulting_resolution,[],[f195,f194,f454,f645,f302])).
fof(f1207,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0'))))) | ($spl1 | $spl13)),
  inference(unit_resulting_resolution,[],[f194,f194,f454,f645,f302])).
fof(f1209,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0'))))) | ($spl10 | $spl13)),
  inference(unit_resulting_resolution,[],[f195,f200,f433,f645,f304])).
fof(f1210,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0'))))) | ($spl10 | $spl13)),
  inference(unit_resulting_resolution,[],[f194,f200,f433,f645,f304])).
fof(f1211,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0'))))) | $spl13),
  inference(unit_resulting_resolution,[],[f195,f645,f250,f645,f304])).
fof(f1212,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0'))))) | $spl13),
  inference(unit_resulting_resolution,[],[f194,f645,f250,f645,f304])).
fof(f1216,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0'))))) | $spl13),
  inference(unit_resulting_resolution,[],[f645,f195,f250,f645,f304])).
fof(f1218,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0'))))) | $spl13),
  inference(unit_resulting_resolution,[],[f645,f194,f250,f645,f304])).
fof(f1222,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0'))))) | $spl13),
  inference(unit_resulting_resolution,[],[f200,f195,f195,f645,f305])).
fof(f1223,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0'))))) | $spl13),
  inference(unit_resulting_resolution,[],[f200,f194,f195,f645,f305])).
fof(f1224,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0'))))) | $spl13),
  inference(unit_resulting_resolution,[],[f200,f195,f194,f645,f305])).
fof(f1225,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0'))))) | $spl13),
  inference(unit_resulting_resolution,[],[f200,f194,f194,f645,f305])).
fof(f1230,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0'))))) | $spl13),
  inference(unit_resulting_resolution,[],[f195,f645,f195,f645,f305])).
fof(f1231,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0'))))) | $spl13),
  inference(unit_resulting_resolution,[],[f194,f645,f195,f645,f305])).
fof(f1234,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0'))))) | $spl13),
  inference(unit_resulting_resolution,[],[f195,f645,f194,f645,f305])).
fof(f1235,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0'))))) | $spl13),
  inference(unit_resulting_resolution,[],[f194,f645,f194,f645,f305])).
fof(f1242,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0'))))) | ($spl1 | $spl13)),
  inference(unit_resulting_resolution,[],[f195,f195,f454,f645,f306])).
fof(f1243,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0'))))) | ($spl1 | $spl13)),
  inference(unit_resulting_resolution,[],[f194,f195,f454,f645,f306])).
fof(f1244,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0'))))) | ($spl1 | $spl13)),
  inference(unit_resulting_resolution,[],[f195,f194,f454,f645,f306])).
fof(f1245,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0'))))) | ($spl1 | $spl13)),
  inference(unit_resulting_resolution,[],[f194,f194,f454,f645,f306])).
fof(f1247,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0'))))) | $spl13),
  inference(unit_resulting_resolution,[],[f200,f195,f645,f312])).
fof(f1248,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0'))))) | $spl13),
  inference(unit_resulting_resolution,[],[f200,f194,f645,f312])).
fof(f1252,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0'))))) | $spl13),
  inference(unit_resulting_resolution,[],[f195,f645,f645,f312])).
fof(f1253,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0'))))) | $spl13),
  inference(unit_resulting_resolution,[],[f194,f645,f645,f312])).
fof(f1156,plain,(
  ( ! [X21,X19,X22,X20] : (s(X19,happ(s(fun(X20,X19),X21),s(X20,sK9('const.pred_set.UNIV_0',X22,X21,X19,X20)))) = s(X19,happ(s(fun(X20,X19),X21),s(X20,sK10('const.pred_set.UNIV_0',X22,X21,X19,X20)))) | sP0(X20,X19,X21,X22,'const.pred_set.UNIV_0')) ) | $spl28),
  inference(cnf_transformation,[],[f1156_D])).
fof(f1156_D,plain,(
  ( ! [X21,X19,X22,X20] : (s(X19,happ(s(fun(X20,X19),X21),s(X20,sK9('const.pred_set.UNIV_0',X22,X21,X19,X20)))) = s(X19,happ(s(fun(X20,X19),X21),s(X20,sK10('const.pred_set.UNIV_0',X22,X21,X19,X20)))) | sP0(X20,X19,X21,X22,'const.pred_set.UNIV_0')) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f1073,plain,(
  ( ! [X21,X19,X22,X20] : (~p(s(bool,'T_0')) | s(X19,happ(s(fun(X20,X19),X21),s(X20,sK9('const.pred_set.UNIV_0',X22,X21,X19,X20)))) = s(X19,happ(s(fun(X20,X19),X21),s(X20,sK10('const.pred_set.UNIV_0',X22,X21,X19,X20)))) | sP0(X20,X19,X21,X22,'const.pred_set.UNIV_0')) ) | $spl1),
  inference(superposition,[],[f287,f454])).
fof(f287,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,sK8(X4,X3,X2,X1,X0)))),s(fun(X1,bool),X4)))) | s(X1,happ(s(fun(X0,X1),X2),s(X0,sK9(X4,X3,X2,X1,X0)))) = s(X1,happ(s(fun(X0,X1),X2),s(X0,sK10(X4,X3,X2,X1,X0)))) | sP0(X0,X1,X2,X3,X4)) )),
  inference(cnf_transformation,[],[f186])).
fof(f186,plain,(
  ! [X0,X1,X2,X3,X4] : ((~sP0(X0,X1,X2,X3,X4) | (! [X10] : (~p(s(bool,'const.bool.IN_2'(s(X0,X10),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X10))),s(fun(X1,bool),X4))))) & ! [X8,X9] : (~p(s(bool,'const.bool.IN_2'(s(X0,X8),s(fun(X0,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X9),s(fun(X0,bool),X3)))) | s(X1,happ(s(fun(X0,X1),X2),s(X0,X8))) != s(X1,happ(s(fun(X0,X1),X2),s(X0,X9))) | s(X0,X8) = s(X0,X9)))) & ((p(s(bool,'const.bool.IN_2'(s(X0,sK8(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,sK8(X4,X3,X2,X1,X0)))),s(fun(X1,bool),X4))))) | (p(s(bool,'const.bool.IN_2'(s(X0,sK9(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3)))) & p(s(bool,'const.bool.IN_2'(s(X0,sK10(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3)))) & s(X1,happ(s(fun(X0,X1),X2),s(X0,sK9(X4,X3,X2,X1,X0)))) = s(X1,happ(s(fun(X0,X1),X2),s(X0,sK10(X4,X3,X2,X1,X0)))) & s(X0,sK9(X4,X3,X2,X1,X0)) != s(X0,sK10(X4,X3,X2,X1,X0))) | sP0(X0,X1,X2,X3,X4)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK8,sK9,sK10])],[f185])).
fof(f185,plain,(
  ! [X0,X1,X2,X3,X4] : ((~sP0(X0,X1,X2,X3,X4) | (! [X10] : (~p(s(bool,'const.bool.IN_2'(s(X0,X10),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X10))),s(fun(X1,bool),X4))))) & ! [X8,X9] : (~p(s(bool,'const.bool.IN_2'(s(X0,X8),s(fun(X0,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X9),s(fun(X0,bool),X3)))) | s(X1,happ(s(fun(X0,X1),X2),s(X0,X8))) != s(X1,happ(s(fun(X0,X1),X2),s(X0,X9))) | s(X0,X8) = s(X0,X9)))) & (? [X7] : (p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X7))),s(fun(X1,bool),X4))))) | ? [X5,X6] : (p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),X3)))) & p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X3)))) & s(X1,happ(s(fun(X0,X1),X2),s(X0,X5))) = s(X1,happ(s(fun(X0,X1),X2),s(X0,X6))) & s(X0,X5) != s(X0,X6)) | sP0(X0,X1,X2,X3,X4)))),
  inference(rectify,[],[f184])).
fof(f184,plain,(
  ! [X1,X0,X2,X3,X4] : ((~sP0(X1,X0,X2,X3,X4) | (! [X7] : (~p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))),s(fun(X0,bool),X4))))) & ! [X5,X6] : (~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3)))) | s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) | s(X1,X5) = s(X1,X6)))) & (? [X7] : (p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))),s(fun(X0,bool),X4))))) | ? [X5,X6] : (p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) & p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3)))) & s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) & s(X1,X5) != s(X1,X6)) | sP0(X1,X0,X2,X3,X4)))),
  inference(flattening,[],[f183])).
fof(f183,plain,(
  ! [X1,X0,X2,X3,X4] : ((~sP0(X1,X0,X2,X3,X4) | (! [X7] : (~p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))),s(fun(X0,bool),X4))))) & ! [X5,X6] : (~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3)))) | s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) | s(X1,X5) = s(X1,X6)))) & ((? [X7] : (p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))),s(fun(X0,bool),X4))))) | ? [X5,X6] : (p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) & p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3)))) & s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) & s(X1,X5) != s(X1,X6))) | sP0(X1,X0,X2,X3,X4)))),
  inference(nnf_transformation,[],[f140])).
fof(f140,plain,(
  ! [X1,X0,X2,X3,X4] : (sP0(X1,X0,X2,X3,X4) <=> (! [X7] : (~p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))),s(fun(X0,bool),X4))))) & ! [X5,X6] : (~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3)))) | s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) | s(X1,X5) = s(X1,X6))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f1155,plain,(
  ( ! [X17,X15,X18,X16] : (p(s(bool,'const.bool.IN_2'(s(X16,sK10('const.pred_set.UNIV_0',X18,X17,X15,X16)),s(fun(X16,bool),X18)))) | sP0(X16,X15,X17,X18,'const.pred_set.UNIV_0')) ) | $spl26),
  inference(cnf_transformation,[],[f1155_D])).
fof(f1155_D,plain,(
  ( ! [X17,X15,X18,X16] : (p(s(bool,'const.bool.IN_2'(s(X16,sK10('const.pred_set.UNIV_0',X18,X17,X15,X16)),s(fun(X16,bool),X18)))) | sP0(X16,X15,X17,X18,'const.pred_set.UNIV_0')) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f1072,plain,(
  ( ! [X17,X15,X18,X16] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(X16,sK10('const.pred_set.UNIV_0',X18,X17,X15,X16)),s(fun(X16,bool),X18)))) | sP0(X16,X15,X17,X18,'const.pred_set.UNIV_0')) ) | $spl1),
  inference(superposition,[],[f285,f454])).
fof(f285,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,sK8(X4,X3,X2,X1,X0)))),s(fun(X1,bool),X4)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK10(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3)))) | sP0(X0,X1,X2,X3,X4)) )),
  inference(cnf_transformation,[],[f186])).
fof(f1154,plain,(
  ( ! [X14,X12,X13,X11] : (p(s(bool,'const.bool.IN_2'(s(X12,sK9('const.pred_set.UNIV_0',X14,X13,X11,X12)),s(fun(X12,bool),X14)))) | sP0(X12,X11,X13,X14,'const.pred_set.UNIV_0')) ) | $spl24),
  inference(cnf_transformation,[],[f1154_D])).
fof(f1154_D,plain,(
  ( ! [X14,X12,X13,X11] : (p(s(bool,'const.bool.IN_2'(s(X12,sK9('const.pred_set.UNIV_0',X14,X13,X11,X12)),s(fun(X12,bool),X14)))) | sP0(X12,X11,X13,X14,'const.pred_set.UNIV_0')) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f1071,plain,(
  ( ! [X14,X12,X13,X11] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(X12,sK9('const.pred_set.UNIV_0',X14,X13,X11,X12)),s(fun(X12,bool),X14)))) | sP0(X12,X11,X13,X14,'const.pred_set.UNIV_0')) ) | $spl1),
  inference(superposition,[],[f283,f454])).
fof(f283,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,sK8(X4,X3,X2,X1,X0)))),s(fun(X1,bool),X4)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK9(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3)))) | sP0(X0,X1,X2,X3,X4)) )),
  inference(cnf_transformation,[],[f186])).
fof(f1152,plain,(
  ( ! [X6,X5] : (s(X5,X6) = s(X5,happ(s(fun('type.num.num',X5),sK1('const.pred_set.UNIV_0',X5)),s('type.num.num',sK2(X6,'const.pred_set.UNIV_0',X5)))) | ~p(s(bool,'const.util_prob.countable_1'(s(fun(X5,bool),'const.pred_set.UNIV_0'))))) ) | $spl22),
  inference(cnf_transformation,[],[f1152_D])).
fof(f1152_D,plain,(
  ( ! [X6,X5] : (s(X5,X6) = s(X5,happ(s(fun('type.num.num',X5),sK1('const.pred_set.UNIV_0',X5)),s('type.num.num',sK2(X6,'const.pred_set.UNIV_0',X5)))) | ~p(s(bool,'const.util_prob.countable_1'(s(fun(X5,bool),'const.pred_set.UNIV_0'))))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f1069,plain,(
  ( ! [X6,X5] : (~p(s(bool,'T_0')) | s(X5,X6) = s(X5,happ(s(fun('type.num.num',X5),sK1('const.pred_set.UNIV_0',X5)),s('type.num.num',sK2(X6,'const.pred_set.UNIV_0',X5)))) | ~p(s(bool,'const.util_prob.countable_1'(s(fun(X5,bool),'const.pred_set.UNIV_0'))))) ) | $spl1),
  inference(superposition,[],[f251,f454])).
fof(f251,plain,(
  ( ! [X6,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X1)))) | s(X0,X6) = s(X0,happ(s(fun('type.num.num',X0),sK1(X1,X0)),s('type.num.num',sK2(X6,X1,X0)))) | ~p(s(bool,'const.util_prob.countable_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f172])).
fof(f1149,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X2,happ(s(fun(X0,X2),X3),s(X0,X1))),s(fun(X2,bool),X4)))) | ~sP0(X0,X2,X3,'const.pred_set.UNIV_0',X4)) ) | $spl18),
  inference(cnf_transformation,[],[f1149_D])).
fof(f1149_D,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X2,happ(s(fun(X0,X2),X3),s(X0,X1))),s(fun(X2,bool),X4)))) | ~sP0(X0,X2,X3,'const.pred_set.UNIV_0',X4)) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f1068,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(X2,happ(s(fun(X0,X2),X3),s(X0,X1))),s(fun(X2,bool),X4)))) | ~sP0(X0,X2,X3,'const.pred_set.UNIV_0',X4)) ) | $spl1),
  inference(superposition,[],[f280,f454])).
fof(f280,plain,(
  ( ! [X4,X2,X0,X10,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X10),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X10))),s(fun(X1,bool),X4)))) | ~sP0(X0,X1,X2,X3,X4)) )),
  inference(cnf_transformation,[],[f186])).
fof(f1022,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,'T_0')))) | $spl16),
  inference(cnf_transformation,[],[f1022_D])).
fof(f1022_D,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,'T_0')))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f1023,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,'T_0')))) | $spl1),
  inference(forward_demodulation,[],[f1019,f952])).
fof(f952,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK3(X0,'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'type.pair.prod'('type.num.num','type.num.num'))),s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f194,f315,f260])).
fof(f1019,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK3(X1,'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'type.pair.prod'('type.num.num','type.num.num'))),s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))))))) ) | $spl1),
  inference(backward_demodulation,[],[f952,f959])).
fof(f959,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK3(X0,'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'type.pair.prod'('type.num.num','type.num.num'))),s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))))),s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK3(X1,'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'type.pair.prod'('type.num.num','type.num.num'))),s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f315,f315,f265])).
fof(f1018,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,'T_0')))) | $spl1),
  inference(backward_demodulation,[],[f952,f958])).
fof(f958,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK3(X0,'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'type.pair.prod'('type.num.num','type.num.num'))),s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f194,f315,f265])).
fof(f1017,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,'T_0')))) | $spl1),
  inference(backward_demodulation,[],[f952,f955])).
fof(f955,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK3(X0,'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'type.pair.prod'('type.num.num','type.num.num'))),s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))))),s(bool,'T_0'))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f194,f315,f265])).
fof(f934,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num','type.pair.prod'('type.num.num','type.num.num')),sK11),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) | $spl14),
  inference(cnf_transformation,[],[f934_D])).
fof(f934_D,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num','type.pair.prod'('type.num.num','type.num.num')),sK11),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f932,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num','type.pair.prod'('type.num.num','type.num.num')),sK11),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) | $spl7),
  inference(trivial_inequality_removal,[],[f931])).
fof(f931,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num','type.pair.prod'('type.num.num','type.num.num')),sK11),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) | $spl7),
  inference(superposition,[],[f425,f201])).
fof(f926,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num','type.pair.prod'('type.num.num','type.num.num')),sK11),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) | $spl7),
  inference(unit_resulting_resolution,[],[f425,f201])).
fof(f645,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0')))))) | $spl13),
  inference(cnf_transformation,[],[f645_D])).
fof(f645_D,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0')))))) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f641,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0')))))) | ($spl1 | $spl10)),
  inference(superposition,[],[f435,f250])).
fof(f435,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0'))),s(bool,X0))))) ) | ($spl1 | $spl10)),
  inference(backward_demodulation,[],[f433,f331])).
fof(f331,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'const.util_prob.countable_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))))),s(bool,X0))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f314,f263])).
fof(f263,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X0))),s(bool,X1)))) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f175])).
fof(f509,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0')))))) | ($spl1 | $spl10)),
  inference(unit_resulting_resolution,[],[f435,f250,f250,f435,f300])).
fof(f510,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0')))))) | ($spl1 | $spl10)),
  inference(unit_resulting_resolution,[],[f435,f298,f250,f435,f300])).
fof(f519,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0')))))) | ($spl1 | $spl10)),
  inference(unit_resulting_resolution,[],[f435,f250,f250,f435,f300])).
fof(f527,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0')))))) | ($spl1 | $spl10)),
  inference(unit_resulting_resolution,[],[f435,f250,f298,f435,f300])).
fof(f535,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0')))))) | ($spl1 | $spl10)),
  inference(unit_resulting_resolution,[],[f435,f250,f250,f435,f304])).
fof(f536,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0')))))) | ($spl1 | $spl10)),
  inference(unit_resulting_resolution,[],[f435,f298,f250,f435,f304])).
fof(f545,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0')))))) | ($spl1 | $spl10)),
  inference(unit_resulting_resolution,[],[f435,f250,f250,f435,f304])).
fof(f553,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0')))))) | ($spl1 | $spl10)),
  inference(unit_resulting_resolution,[],[f435,f250,f298,f435,f304])).
fof(f573,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0')))))) | ($spl1 | $spl10)),
  inference(unit_resulting_resolution,[],[f200,f195,f250,f435,f305])).
fof(f574,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0')))))) | ($spl1 | $spl10)),
  inference(unit_resulting_resolution,[],[f200,f194,f250,f435,f305])).
fof(f575,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0')))))) | ($spl1 | $spl10)),
  inference(unit_resulting_resolution,[],[f200,f309,f250,f435,f305])).
fof(f309,plain,(
  p(s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num','type.pair.prod'('type.num.num','type.num.num')),sK11),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))))),
  inference(cnf_transformation,[],[f192])).
fof(f192,plain,(
  p(s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num','type.pair.prod'('type.num.num','type.num.num')),sK11),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK11])],[f115])).
fof(f115,plain,(
  ? [X0] : p(s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num','type.pair.prod'('type.num.num','type.num.num')),X0),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ? [X16] : p(s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num','type.pair.prod'('type.num.num','type.num.num')),X16),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))))),
  file('thm.extreal.CROSS_COUNTABLE_UNIV','thm.util_prob.NUM_2D_BIJ_INV')).
fof(f600,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0')))))) | ($spl1 | $spl10)),
  inference(unit_resulting_resolution,[],[f195,f435,f250,f435,f305])).
fof(f610,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0')))))) | ($spl1 | $spl10)),
  inference(unit_resulting_resolution,[],[f200,f250,f435,f312])).
fof(f619,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0')))))) | ($spl1 | $spl10)),
  inference(unit_resulting_resolution,[],[f435,f250,f435,f312])).
fof(f433,plain,(
  s(bool,'F_0') = s(bool,'const.util_prob.countable_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) | $spl10),
  inference(cnf_transformation,[],[f433_D])).
fof(f433_D,plain,(
  s(bool,'F_0') = s(bool,'const.util_prob.countable_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f412,plain,(
  s(bool,'F_0') = s(bool,'const.util_prob.countable_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) | $spl1),
  inference(resolution,[],[f314,f214])).
fof(f411,plain,(
  s(bool,'F_0') = s(bool,'const.util_prob.countable_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) | $spl1),
  inference(resolution,[],[f314,f212])).
fof(f423,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) | $spl5),
  inference(cnf_transformation,[],[f423_D])).
fof(f423_D,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f317,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) | $spl1),
  inference(unit_resulting_resolution,[],[f314,f207])).
fof(f318,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) | $spl1),
  inference(unit_resulting_resolution,[],[f314,f209])).
fof(f319,plain,(
  s(bool,'F_0') = s(bool,'const.util_prob.countable_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) | $spl1),
  inference(unit_resulting_resolution,[],[f314,f212])).
fof(f320,plain,(
  s(bool,'F_0') = s(bool,'const.util_prob.countable_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) | $spl1),
  inference(unit_resulting_resolution,[],[f314,f214])).
fof(f327,plain,(
  s(bool,'F_0') = s(bool,'const.util_prob.countable_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) | $spl1),
  inference(unit_resulting_resolution,[],[f200,f314,f257])).
fof(f329,plain,(
  s(bool,'F_0') = s(bool,'const.util_prob.countable_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) | $spl1),
  inference(unit_resulting_resolution,[],[f200,f314,f257])).
fof(f333,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) | $spl1),
  inference(unit_resulting_resolution,[],[f195,f249,f254,f314,f302])).
fof(f254,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2))) = s(bool,happ(s(fun(fun(X0,bool),bool),happ(s(fun(X0,fun(fun(X0,bool),bool)),'const.bool.IN_0'),s(X0,X1))),s(fun(X0,bool),X2)))) )),
  inference(cnf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0,X1,X2] : s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2))) = s(bool,happ(s(fun(fun(X0,bool),bool),happ(s(fun(X0,fun(fun(X0,bool),bool)),'const.bool.IN_0'),s(X0,X1))),s(fun(X0,bool),X2)))),
  inference(flattening,[],[f84])).
fof(f84,plain,(
  ! [X0] : ! [X1,X2] : s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2))) = s(bool,happ(s(fun(fun(X0,bool),bool),happ(s(fun(X0,fun(fun(X0,bool),bool)),'const.bool.IN_0'),s(X0,X1))),s(fun(X0,bool),X2)))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X6] : ! [X23,X24] : s(bool,'const.bool.IN_2'(s(X6,X23),s(fun(X6,bool),X24))) = s(bool,happ(s(fun(fun(X6,bool),bool),happ(s(fun(X6,fun(fun(X6,bool),bool)),'const.bool.IN_0'),s(X6,X23))),s(fun(X6,bool),X24)))),
  file('thm.extreal.CROSS_COUNTABLE_UNIV','ARITY_0_2_const.bool.IN')).
fof(f334,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) | $spl1),
  inference(unit_resulting_resolution,[],[f194,f249,f254,f314,f302])).
fof(f421,plain,(
  s(bool,'const.util_prob.countable_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) != s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num','type.pair.prod'('type.num.num','type.num.num')),sK11),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) | $spl3),
  inference(cnf_transformation,[],[f421_D])).
fof(f421_D,plain,(
  s(bool,'const.util_prob.countable_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) != s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num','type.pair.prod'('type.num.num','type.num.num')),sK11),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f336,plain,(
  s(bool,'const.util_prob.countable_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) != s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num','type.pair.prod'('type.num.num','type.num.num')),sK11),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) | $spl1),
  inference(unit_resulting_resolution,[],[f309,f249,f254,f314,f302])).
fof(f337,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) | $spl1),
  inference(unit_resulting_resolution,[],[f200,f195,f195,f314,f305])).
fof(f338,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) | $spl1),
  inference(unit_resulting_resolution,[],[f200,f194,f195,f314,f305])).
fof(f341,plain,(
  s(bool,'const.util_prob.countable_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) != s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num','type.pair.prod'('type.num.num','type.num.num')),sK11),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) | $spl1),
  inference(unit_resulting_resolution,[],[f200,f309,f195,f314,f305])).
fof(f342,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) | $spl1),
  inference(unit_resulting_resolution,[],[f200,f195,f194,f314,f305])).
fof(f343,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) | $spl1),
  inference(unit_resulting_resolution,[],[f200,f194,f194,f314,f305])).
fof(f346,plain,(
  s(bool,'const.util_prob.countable_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) != s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num','type.pair.prod'('type.num.num','type.num.num')),sK11),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) | $spl1),
  inference(unit_resulting_resolution,[],[f200,f309,f194,f314,f305])).
fof(f347,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) | $spl1),
  inference(unit_resulting_resolution,[],[f200,f195,f249,f314,f305])).
fof(f348,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) | $spl1),
  inference(unit_resulting_resolution,[],[f200,f194,f249,f314,f305])).
fof(f351,plain,(
  s(bool,'const.util_prob.countable_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) != s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num','type.pair.prod'('type.num.num','type.num.num')),sK11),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) | $spl1),
  inference(unit_resulting_resolution,[],[f200,f309,f249,f314,f305])).
fof(f352,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) | $spl1),
  inference(unit_resulting_resolution,[],[f200,f195,f309,f314,f305])).
fof(f353,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) | $spl1),
  inference(unit_resulting_resolution,[],[f200,f194,f309,f314,f305])).
fof(f356,plain,(
  s(bool,'const.util_prob.countable_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) != s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num','type.pair.prod'('type.num.num','type.num.num')),sK11),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) | $spl1),
  inference(unit_resulting_resolution,[],[f200,f309,f309,f314,f305])).
fof(f427,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl9),
  inference(cnf_transformation,[],[f427_D])).
fof(f427_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f357,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f195,f200,f195,f314,f305])).
fof(f358,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f194,f200,f195,f314,f305])).
fof(f425,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num','type.pair.prod'('type.num.num','type.num.num')),sK11),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) | $spl7),
  inference(cnf_transformation,[],[f425_D])).
fof(f425_D,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num','type.pair.prod'('type.num.num','type.num.num')),sK11),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f360,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num','type.pair.prod'('type.num.num','type.num.num')),sK11),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) | $spl1),
  inference(unit_resulting_resolution,[],[f309,f200,f195,f314,f305])).
fof(f361,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) | $spl1),
  inference(unit_resulting_resolution,[],[f195,f314,f195,f314,f305])).
fof(f362,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) | $spl1),
  inference(unit_resulting_resolution,[],[f194,f314,f195,f314,f305])).
fof(f364,plain,(
  s(bool,'const.util_prob.countable_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) != s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num','type.pair.prod'('type.num.num','type.num.num')),sK11),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) | $spl1),
  inference(unit_resulting_resolution,[],[f309,f314,f195,f314,f305])).
fof(f365,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f195,f200,f194,f314,f305])).
fof(f366,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f194,f200,f194,f314,f305])).
fof(f368,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num','type.pair.prod'('type.num.num','type.num.num')),sK11),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) | $spl1),
  inference(unit_resulting_resolution,[],[f309,f200,f194,f314,f305])).
fof(f369,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) | $spl1),
  inference(unit_resulting_resolution,[],[f195,f314,f194,f314,f305])).
fof(f370,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) | $spl1),
  inference(unit_resulting_resolution,[],[f194,f314,f194,f314,f305])).
fof(f372,plain,(
  s(bool,'const.util_prob.countable_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) != s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num','type.pair.prod'('type.num.num','type.num.num')),sK11),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) | $spl1),
  inference(unit_resulting_resolution,[],[f309,f314,f194,f314,f305])).
fof(f373,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f195,f200,f249,f314,f305])).
fof(f374,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f194,f200,f249,f314,f305])).
fof(f376,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num','type.pair.prod'('type.num.num','type.num.num')),sK11),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) | $spl1),
  inference(unit_resulting_resolution,[],[f309,f200,f249,f314,f305])).
fof(f377,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) | $spl1),
  inference(unit_resulting_resolution,[],[f195,f314,f249,f314,f305])).
fof(f378,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) | $spl1),
  inference(unit_resulting_resolution,[],[f194,f314,f249,f314,f305])).
fof(f380,plain,(
  s(bool,'const.util_prob.countable_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) != s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num','type.pair.prod'('type.num.num','type.num.num')),sK11),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) | $spl1),
  inference(unit_resulting_resolution,[],[f309,f314,f249,f314,f305])).
fof(f381,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f195,f200,f309,f314,f305])).
fof(f382,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f194,f200,f309,f314,f305])).
fof(f384,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num','type.pair.prod'('type.num.num','type.num.num')),sK11),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) | $spl1),
  inference(unit_resulting_resolution,[],[f309,f200,f309,f314,f305])).
fof(f385,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) | $spl1),
  inference(unit_resulting_resolution,[],[f195,f314,f309,f314,f305])).
fof(f386,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) | $spl1),
  inference(unit_resulting_resolution,[],[f194,f314,f309,f314,f305])).
fof(f388,plain,(
  s(bool,'const.util_prob.countable_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) != s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num','type.pair.prod'('type.num.num','type.num.num')),sK11),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) | $spl1),
  inference(unit_resulting_resolution,[],[f309,f314,f309,f314,f305])).
fof(f393,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) | $spl1),
  inference(unit_resulting_resolution,[],[f195,f249,f254,f314,f306])).
fof(f394,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) | $spl1),
  inference(unit_resulting_resolution,[],[f194,f249,f254,f314,f306])).
fof(f396,plain,(
  s(bool,'const.util_prob.countable_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) != s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num','type.pair.prod'('type.num.num','type.num.num')),sK11),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) | $spl1),
  inference(unit_resulting_resolution,[],[f309,f249,f254,f314,f306])).
fof(f397,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) | $spl1),
  inference(unit_resulting_resolution,[],[f200,f195,f314,f312])).
fof(f398,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) | $spl1),
  inference(unit_resulting_resolution,[],[f200,f194,f314,f312])).
fof(f401,plain,(
  s(bool,'const.util_prob.countable_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) != s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num','type.pair.prod'('type.num.num','type.num.num')),sK11),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) | $spl1),
  inference(unit_resulting_resolution,[],[f200,f309,f314,f312])).
fof(f402,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f195,f200,f314,f312])).
fof(f403,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f194,f200,f314,f312])).
fof(f405,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num','type.pair.prod'('type.num.num','type.num.num')),sK11),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) | $spl1),
  inference(unit_resulting_resolution,[],[f309,f200,f314,f312])).
fof(f406,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) | $spl1),
  inference(unit_resulting_resolution,[],[f195,f314,f314,f312])).
fof(f407,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) | $spl1),
  inference(unit_resulting_resolution,[],[f194,f314,f314,f312])).
fof(f409,plain,(
  s(bool,'const.util_prob.countable_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) != s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num','type.pair.prod'('type.num.num','type.num.num')),sK11),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) | $spl1),
  inference(unit_resulting_resolution,[],[f309,f314,f314,f312])).
fof(f314,plain,(
  ~p(s(bool,'const.util_prob.countable_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')))))) | $spl1),
  inference(cnf_transformation,[],[f314_D])).
fof(f314_D,plain,(
  ~p(s(bool,'const.util_prob.countable_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f193,plain,(
  ~p(s(bool,'const.util_prob.countable_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))))),
  inference(cnf_transformation,[],[f49])).
fof(f49,plain,(
  ~p(s(bool,'const.util_prob.countable_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))))),
  inference(flattening,[],[f6])).
fof(f6,negated_conjecture,(
  ~p(s(bool,'const.util_prob.countable_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  p(s(bool,'const.util_prob.countable_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.pred_set.CROSS_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))))),
  file('thm.extreal.CROSS_COUNTABLE_UNIV','thm.extreal.CROSS_COUNTABLE_UNIV')).
% SZS output end Proof for thm.extreal
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 99
% Passive clauses: 132
% Generated clauses: 2013
% Final active clauses: 84
% Final passive clauses: 82
% Input formulas: 47
% Initial clauses: 117
% 
% Duplicate literals: 12
% Trivial inequalities: 13
% Fw demodulations: 437
% Bw demodulations: 37
% Global subsumptions: 1
% 
% Simple tautologies: 59
% Equational tautologies: 9
% Forward subsumptions: 1239
% Backward subsumptions: 3
% 
% Binary resolution: 79
% Unit resulting resolution: 849
% Forward superposition: 260
% Backward superposition: 298
% Self superposition: 3
% 
% Split clauses: 7
% Split components: 14
% SAT solver clauses: 447
% SAT solver unit clauses: 2
% SAT solver binary clauses: 273
% 
% Sat splits: 7
% 
% Memory used [KB]: 3070
% Time elapsed: 0.132 s
% ------------------------------
% ------------------------------
% Success in time 18.162 s
