Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.option
% (55650)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 124
% Passive clauses: 3786
% Generated clauses: 12877
% Final active clauses: 106
% Final passive clauses: 3157
% Input formulas: 31
% Initial clauses: 85
% 
% Duplicate literals: 635
% Trivial inequalities: 304
% Fw subsumption resolutions: 2230
% Fw demodulations: 189
% Bw demodulations: 10
% 
% Simple tautologies: 367
% Equational tautologies: 98
% Forward subsumptions: 3633
% Fw demodulations to eq. taut.: 4
% 
% Binary resolution: 487
% Factoring: 2
% Forward superposition: 3095
% Backward superposition: 5187
% Self superposition: 267
% Equality factoring: 222
% Equality resolution: 36
% 
% Split clauses: 580
% Split components: 1208
% SAT solver clauses: 1360
% SAT solver unit clauses: 47
% SAT solver binary clauses: 419
% 
% Sat splits: 580
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1325
% TWLsolver calls for satisfiability: 203
% 
% Memory used [KB]: 6268
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.option
% (55653)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 204
% Passive clauses: 1403
% Generated clauses: 16148
% Final active clauses: 98
% Final passive clauses: 565
% Input formulas: 31
% Initial clauses: 85
% 
% Duplicate literals: 450
% Trivial inequalities: 284
% Fw subsumption resolutions: 1689
% Fw demodulations: 1661
% Bw demodulations: 669
% Condensations: 804
% Global subsumptions: 9
% 
% Simple tautologies: 241
% Equational tautologies: 26
% Forward subsumptions: 7628
% Fw demodulations to eq. taut.: 23
% Bw demodulations to eq. taut.: 32
% 
% Binary resolution: 505
% Factoring: 1
% Forward superposition: 4668
% Backward superposition: 3403
% Self superposition: 79
% Equality factoring: 38
% Equality resolution: 47
% 
% Split clauses: 224
% Split components: 450
% SAT solver clauses: 1947
% SAT solver unit clauses: 90
% SAT solver binary clauses: 59
% SAT solver learnt clauses: 19
% SAT solver learnt literals: 3
% 
% Sat splits: 224
% Sat splitting refutations: 36
% 
% TWLsolver clauses: 1910
% TWLsolver calls for satisfiability: 545
% 
% Memory used [KB]: 10490
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.option
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.option
% SZS output start Proof for thm.option
fof(f20797,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f999,f1002,f1000,f1005,f1011,f1015,f1017,f1019,f1021,f1025,f1027,f1029,f1031,f1033,f1035,f1091,f1092,f1093,f1094,f1096,f1101,f1102,f1103,f1104,f1105,f1106,f1107,f1108,f1833,f1826,f1861,f1860,f1819,f1810,f1858,f1763,f1945,f1711,f2007,f2006,f2099,f2100,f2101,f2102,f2104,f2105,f2107,f2111,f2113,f2115,f2117,f2138,f2140,f2142,f2148,f2150,f3593,f3620,f3823,f3831,f3832,f3827,f3825,f3976,f3978,f3977,f4202,f4204,f4203,f4205,f4208,f4212,f4214,f4218,f4220,f4225,f4233,f4372,f4373,f6108,f6110,f6114,f6118,f6122,f6124,f6144,f6145,f6147,f6149,f6152,f6153,f6417,f6419,f6418,f6505,f6509,f6507,f6592,f6824,f6825,f6826,f6827,f6829,f6846,f6849,f6856,f6858,f6862,f6864,f6866,f6868,f6870,f6872,f6874,f6876,f6878,f6880,f6882,f6884,f6886,f6888,f6894,f8554,f8556,f8555,f8773,f8775,f8774,f8980,f8982,f8981,f11112,f11145,f12025,f12043,f12551,f12559,f12546,f12561,f12573,f12574,f12828,f12900,f13324,f13326,f13333,f13334,f13364,f13365,f13366,f13367,f13368,f13369,f13370,f13371,f13372,f13373,f13397,f13390,f13374,f13398,f13396,f13375,f13392,f13399,f13376,f13389,f13400,f13377,f13393,f13394,f13378,f13401,f13387,f13386,f13565,f13568,f13647,f13657,f14053,f14055,f15316,f15300,f15285,f15366,f15490,f1001,f15491,f15206,f15501,f15204,f15502,f16052,f16128,f16516,f16594,f17601,f17589,f17572,f17668,f17754,f17755,f17762,f17763,f17790,f17791,f17793,f17794,f17795,f17796,f17797,f17798,f17799,f17800,f17500,f17837,f17948,f17949,f18224,f18225,f18733,f18734,f18937,f18947,f19048,f19049,f19128,f19138,f19230,f19240,f19337,f19338,f19526,f19536,f13325,f19633,f19634,f19748,f19749,f20693,f20716,f1859,f20728,f3828,f20729,f20784])).
fof(f20784,plain,(
  $false | ($spl18 | $spl21)),
  inference(subsumption_resolution,[],[f20783,f128])).
fof(f128,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f35])).
fof(f35,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.option.OPTREL_MONO',aHL_FALSITY)).
fof(f20783,plain,(
  p(s(bool,'F_0')) | ($spl18 | $spl21)),
  inference(forward_demodulation,[],[f20732,f232])).
fof(f232,plain,(
  s(bool,'F_0') = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)))),
  inference(unit_resulting_resolution,[],[f125,f136])).
fof(f136,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f40])).
fof(f40,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X16] : (s(bool,'F_0') = s(bool,X16) <=> ~p(s(bool,X16)))),
  file('thm.option.OPTREL_MONO','thm.bool.EQ_CLAUSES|split|3')).
fof(f125,plain,(
  ~p(s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))))),
  inference(cnf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X6,X7] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK5),s(sK0,X6))),s(sK1,X7)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,X6))),s(sK1,X7))))) & p(s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)))) & ~p(s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5])],[f72])).
fof(f72,plain,(
  ? [X0,X1,X2,X3,X4,X5] : (! [X6,X7] : (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X5),s(X0,X6))),s(X1,X7)))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X4),s(X0,X6))),s(X1,X7))))) & p(s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(X1,bool)),X5),s('type.option.option'(X0),X3),s('type.option.option'(X1),X2)))) & ~p(s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(X1,bool)),X4),s('type.option.option'(X0),X3),s('type.option.option'(X1),X2)))))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ? [X0,X1,X2,X3,X4,X5] : (! [X6,X7] : (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X5),s(X0,X6))),s(X1,X7)))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X4),s(X0,X6))),s(X1,X7))))) & (p(s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(X1,bool)),X5),s('type.option.option'(X0),X3),s('type.option.option'(X1),X2)))) & ~p(s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(X1,bool)),X4),s('type.option.option'(X0),X3),s('type.option.option'(X1),X2))))))),
  inference(ennf_transformation,[],[f34])).
fof(f34,plain,(
  ~! [X0,X1,X2,X3,X4,X5] : (! [X6,X7] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X5),s(X0,X6))),s(X1,X7)))) => p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X4),s(X0,X6))),s(X1,X7))))) => (p(s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(X1,bool)),X5),s('type.option.option'(X0),X3),s('type.option.option'(X1),X2)))) => p(s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(X1,bool)),X4),s('type.option.option'(X0),X3),s('type.option.option'(X1),X2))))))),
  inference(flattening,[],[f33])).
fof(f33,plain,(
  ~! [X0,X1] : ! [X2,X3,X4,X5] : (! [X6,X7] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X5),s(X0,X6))),s(X1,X7)))) => p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X4),s(X0,X6))),s(X1,X7))))) => (p(s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(X1,bool)),X5),s('type.option.option'(X0),X3),s('type.option.option'(X1),X2)))) => p(s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(X1,bool)),X4),s('type.option.option'(X0),X3),s('type.option.option'(X1),X2))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10,X11] : (! [X12,X13] : (p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X6,fun(X7,bool)),X11),s(X6,X12))),s(X7,X13)))) => p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X6,fun(X7,bool)),X10),s(X6,X12))),s(X7,X13))))) => (p(s(bool,'const.option.OPTREL_3'(s(fun(X6,fun(X7,bool)),X11),s('type.option.option'(X6),X9),s('type.option.option'(X7),X8)))) => p(s(bool,'const.option.OPTREL_3'(s(fun(X6,fun(X7,bool)),X10),s('type.option.option'(X6),X9),s('type.option.option'(X7),X8))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10,X11] : (! [X12,X13] : (p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X6,fun(X7,bool)),X11),s(X6,X12))),s(X7,X13)))) => p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X6,fun(X7,bool)),X10),s(X6,X12))),s(X7,X13))))) => (p(s(bool,'const.option.OPTREL_3'(s(fun(X6,fun(X7,bool)),X11),s('type.option.option'(X6),X9),s('type.option.option'(X7),X8)))) => p(s(bool,'const.option.OPTREL_3'(s(fun(X6,fun(X7,bool)),X10),s('type.option.option'(X6),X9),s('type.option.option'(X7),X8))))))),
  file('thm.option.OPTREL_MONO','thm.option.OPTREL_MONO')).
fof(f20732,plain,(
  p(s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)))) | ($spl18 | $spl21)),
  inference(unit_resulting_resolution,[],[f3825,f6511,f20545,f193])).
fof(f193,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.option.option'(X0),X3) != s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X5))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6)))) | s('type.option.option'(X1),X4) != s('type.option.option'(X1),'const.option.SOME_1'(s(X1,X6))) | p(s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.option.option'(X0),X3),s('type.option.option'(X1),X4))))) )),
  inference(cnf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.option.option'(X0),X3),s('type.option.option'(X1),X4)))) | (s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X3) & s('type.option.option'(X1),'const.option.NONE_0') = s('type.option.option'(X1),X4)) | (s('type.option.option'(X0),X3) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,sK8(X4,X3,X2,X1,X0)))) & s('type.option.option'(X1),X4) = s('type.option.option'(X1),'const.option.SOME_1'(s(X1,sK9(X4,X3,X2,X1,X0)))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,sK8(X4,X3,X2,X1,X0)))),s(X1,sK9(X4,X3,X2,X1,X0))))))) & (((s('type.option.option'(X0),'const.option.NONE_0') != s('type.option.option'(X0),X3) | s('type.option.option'(X1),'const.option.NONE_0') != s('type.option.option'(X1),X4)) & ! [X5,X6] : (s('type.option.option'(X0),X3) != s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X5))) | s('type.option.option'(X1),X4) != s('type.option.option'(X1),'const.option.SOME_1'(s(X1,X6))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6)))))) | p(s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.option.option'(X0),X3),s('type.option.option'(X1),X4))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK8,sK9])],[f117])).
fof(f117,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.option.option'(X0),X3),s('type.option.option'(X1),X4)))) | (s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X3) & s('type.option.option'(X1),'const.option.NONE_0') = s('type.option.option'(X1),X4)) | ? [X7,X8] : (s('type.option.option'(X0),X3) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X7))) & s('type.option.option'(X1),X4) = s('type.option.option'(X1),'const.option.SOME_1'(s(X1,X8))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X7))),s(X1,X8)))))) & (((s('type.option.option'(X0),'const.option.NONE_0') != s('type.option.option'(X0),X3) | s('type.option.option'(X1),'const.option.NONE_0') != s('type.option.option'(X1),X4)) & ! [X5,X6] : (s('type.option.option'(X0),X3) != s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X5))) | s('type.option.option'(X1),X4) != s('type.option.option'(X1),'const.option.SOME_1'(s(X1,X6))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6)))))) | p(s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.option.option'(X0),X3),s('type.option.option'(X1),X4))))))),
  inference(rectify,[],[f116])).
fof(f116,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.option.option'(X0),X3),s('type.option.option'(X1),X4)))) | (s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X3) & s('type.option.option'(X1),'const.option.NONE_0') = s('type.option.option'(X1),X4)) | ? [X5,X6] : (s('type.option.option'(X0),X3) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X5))) & s('type.option.option'(X1),X4) = s('type.option.option'(X1),'const.option.SOME_1'(s(X1,X6))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6)))))) & (((s('type.option.option'(X0),'const.option.NONE_0') != s('type.option.option'(X0),X3) | s('type.option.option'(X1),'const.option.NONE_0') != s('type.option.option'(X1),X4)) & ! [X5,X6] : (s('type.option.option'(X0),X3) != s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X5))) | s('type.option.option'(X1),X4) != s('type.option.option'(X1),'const.option.SOME_1'(s(X1,X6))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6)))))) | p(s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.option.option'(X0),X3),s('type.option.option'(X1),X4))))))),
  inference(flattening,[],[f115])).
fof(f115,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.option.option'(X0),X3),s('type.option.option'(X1),X4)))) | ((s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X3) & s('type.option.option'(X1),'const.option.NONE_0') = s('type.option.option'(X1),X4)) | ? [X5,X6] : (s('type.option.option'(X0),X3) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X5))) & s('type.option.option'(X1),X4) = s('type.option.option'(X1),'const.option.SOME_1'(s(X1,X6))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6))))))) & (((s('type.option.option'(X0),'const.option.NONE_0') != s('type.option.option'(X0),X3) | s('type.option.option'(X1),'const.option.NONE_0') != s('type.option.option'(X1),X4)) & ! [X5,X6] : (s('type.option.option'(X0),X3) != s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X5))) | s('type.option.option'(X1),X4) != s('type.option.option'(X1),'const.option.SOME_1'(s(X1,X6))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6)))))) | p(s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.option.option'(X0),X3),s('type.option.option'(X1),X4))))))),
  inference(nnf_transformation,[],[f68])).
fof(f68,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.option.option'(X0),X3),s('type.option.option'(X1),X4)))) <=> ((s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X3) & s('type.option.option'(X1),'const.option.NONE_0') = s('type.option.option'(X1),X4)) | ? [X5,X6] : (s('type.option.option'(X0),X3) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X5))) & s('type.option.option'(X1),X4) = s('type.option.option'(X1),'const.option.SOME_1'(s(X1,X6))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6)))))))),
  inference(flattening,[],[f67])).
fof(f67,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.option.option'(X0),X3),s('type.option.option'(X1),X4)))) <=> ((s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X3) & s('type.option.option'(X1),'const.option.NONE_0') = s('type.option.option'(X1),X4)) | ? [X5,X6] : (s('type.option.option'(X0),X3) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X5))) & s('type.option.option'(X1),X4) = s('type.option.option'(X1),'const.option.SOME_1'(s(X1,X6))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6)))))))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X6,X7] : ! [X21,X9,X8] : (p(s(bool,'const.option.OPTREL_3'(s(fun(X6,fun(X7,bool)),X21),s('type.option.option'(X6),X9),s('type.option.option'(X7),X8)))) <=> ((s('type.option.option'(X6),'const.option.NONE_0') = s('type.option.option'(X6),X9) & s('type.option.option'(X7),'const.option.NONE_0') = s('type.option.option'(X7),X8)) | ? [X12,X13] : (s('type.option.option'(X6),X9) = s('type.option.option'(X6),'const.option.SOME_1'(s(X6,X12))) & s('type.option.option'(X7),X8) = s('type.option.option'(X7),'const.option.SOME_1'(s(X7,X13))) & p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X6,fun(X7,bool)),X21),s(X6,X12))),s(X7,X13)))))))),
  file('thm.option.OPTREL_MONO','thm.option.OPTREL_def')).
fof(f20545,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,sK8(sK2,sK3,sK5,sK1,sK0)))),s(sK1,sK9(sK2,sK3,sK5,sK1,sK0))))) | $spl21),
  inference(unit_resulting_resolution,[],[f8552,f123])).
fof(f123,plain,(
  ( ! [X6,X7] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK5),s(sK0,X6))),s(sK1,X7)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,X6))),s(sK1,X7))))) )),
  inference(cnf_transformation,[],[f81])).
fof(f8552,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK5),s(sK0,sK8(sK2,sK3,sK5,sK1,sK0)))),s(sK1,sK9(sK2,sK3,sK5,sK1,sK0))))) | $spl21),
  inference(subsumption_resolution,[],[f8551,f3826])).
fof(f3826,plain,(
  s('type.option.option'(sK0),'const.option.NONE_0') != s('type.option.option'(sK0),sK3) | $spl21),
  inference(cnf_transformation,[],[f3826_D])).
fof(f3826_D,plain,(
  s('type.option.option'(sK0),'const.option.NONE_0') != s('type.option.option'(sK0),sK3) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f8551,plain,(
  s('type.option.option'(sK0),'const.option.NONE_0') = s('type.option.option'(sK0),sK3) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK5),s(sK0,sK8(sK2,sK3,sK5,sK1,sK0)))),s(sK1,sK9(sK2,sK3,sK5,sK1,sK0)))))),
  inference(subsumption_resolution,[],[f8528,f126])).
fof(f126,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.option.OPTREL_MONO',aHL_TRUTH)).
fof(f8528,plain,(
  ~p(s(bool,'T_0')) | s('type.option.option'(sK0),'const.option.NONE_0') = s('type.option.option'(sK0),sK3) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK5),s(sK0,sK8(sK2,sK3,sK5,sK1,sK0)))),s(sK1,sK9(sK2,sK3,sK5,sK1,sK0)))))),
  inference(superposition,[],[f190,f227])).
fof(f227,plain,(
  s(bool,'T_0') = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)))),
  inference(unit_resulting_resolution,[],[f124,f132])).
fof(f132,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f37])).
fof(f37,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X16] : (s(bool,'T_0') = s(bool,X16) <=> p(s(bool,X16)))),
  file('thm.option.OPTREL_MONO','thm.bool.EQ_CLAUSES|split|0')).
fof(f124,plain,(
  p(s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))))),
  inference(cnf_transformation,[],[f81])).
fof(f190,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.option.option'(X0),X3),s('type.option.option'(X1),X4)))) | s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X3) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,sK8(X4,X3,X2,X1,X0)))),s(X1,sK9(X4,X3,X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f118])).
fof(f6511,plain,(
  s('type.option.option'(sK1),sK2) = s('type.option.option'(sK1),'const.option.SOME_1'(s(sK1,sK9(sK2,sK3,sK5,sK1,sK0)))) | $spl21),
  inference(subsumption_resolution,[],[f6510,f3826])).
fof(f6510,plain,(
  s('type.option.option'(sK0),'const.option.NONE_0') = s('type.option.option'(sK0),sK3) | s('type.option.option'(sK1),sK2) = s('type.option.option'(sK1),'const.option.SOME_1'(s(sK1,sK9(sK2,sK3,sK5,sK1,sK0))))),
  inference(subsumption_resolution,[],[f6445,f126])).
fof(f6445,plain,(
  ~p(s(bool,'T_0')) | s('type.option.option'(sK0),'const.option.NONE_0') = s('type.option.option'(sK0),sK3) | s('type.option.option'(sK1),sK2) = s('type.option.option'(sK1),'const.option.SOME_1'(s(sK1,sK9(sK2,sK3,sK5,sK1,sK0))))),
  inference(superposition,[],[f188,f227])).
fof(f188,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.option.option'(X0),X3),s('type.option.option'(X1),X4)))) | s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X3) | s('type.option.option'(X1),X4) = s('type.option.option'(X1),'const.option.SOME_1'(s(X1,sK9(X4,X3,X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f118])).
fof(f20729,plain,(
  ( ! [X74,X75] : (s(bool,'F_0') = s(bool,sK6('F_0',X74,bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(bool,bool),X74),s(bool,X75))))) | s(bool,'T_0') = s(bool,X75)) ) | $spl142),
  inference(cnf_transformation,[],[f20729_D])).
fof(f20729_D,plain,(
  ( ! [X74,X75] : (s(bool,'F_0') = s(bool,sK6('F_0',X74,bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(bool,bool),X74),s(bool,X75))))) | s(bool,'T_0') = s(bool,X75)) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f3828,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl22),
  inference(cnf_transformation,[],[f3828_D])).
fof(f3828_D,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f20728,plain,(
  ( ! [X76,X74,X75] : (~p(s(bool,X76)) | s(bool,'F_0') = s(bool,sK6('F_0',X74,bool)) | s(bool,'T_0') = s(bool,X75) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(bool,bool),X74),s(bool,X75)))))) ) | $spl42),
  inference(subsumption_resolution,[],[f20687,f135])).
fof(f135,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f84])).
fof(f20687,plain,(
  ( ! [X76,X74,X75] : (~p(s(bool,X76)) | s(bool,'F_0') = s(bool,sK6('F_0',X74,bool)) | s(bool,'T_0') = s(bool,X75) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(bool,bool),X74),s(bool,X75))))) | s(bool,'F_0') = s(bool,X76)) ) | $spl42),
  inference(superposition,[],[f926,f12043])).
fof(f926,plain,(
  ( ! [X23,X22] : (~p(s(bool,happ(s(fun(bool,bool),X23),s(bool,X22)))) | s(bool,'F_0') = s(bool,sK6('F_0',X23,bool)) | s(bool,'T_0') = s(bool,X22)) )),
  inference(superposition,[],[f179,f129])).
fof(f129,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.option.OPTREL_MONO',aHL_BOOL_CASES)).
fof(f179,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | s(X0,X2) = s(X0,sK6(X2,X1,X0))) )),
  inference(cnf_transformation,[],[f113])).
fof(f113,plain,(
  ! [X0,X1,X2] : ((! [X4] : (s(X0,X2) != s(X0,X4) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | (s(X0,X2) = s(X0,sK6(X2,X1,X0)) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK6(X2,X1,X0))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK6])],[f112])).
fof(f112,plain,(
  ! [X0,X1,X2] : ((! [X4] : (s(X0,X2) != s(X0,X4) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ? [X3] : (s(X0,X2) = s(X0,X3) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))))),
  inference(rectify,[],[f111])).
fof(f111,plain,(
  ! [X0,X1,X2] : ((! [X3] : (s(X0,X2) != s(X0,X3) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ? [X3] : (s(X0,X2) = s(X0,X3) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f63])).
fof(f63,plain,(
  ! [X0,X1,X2] : (? [X3] : (s(X0,X2) = s(X0,X3) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) <=> p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(flattening,[],[f62])).
fof(f62,plain,(
  ! [X0] : ! [X1,X2] : (? [X3] : (s(X0,X2) = s(X0,X3) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) <=> p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X6] : ! [X11,X20] : (? [X9] : (s(X6,X9) = s(X6,X20) & p(s(bool,happ(s(fun(X6,bool),X11),s(X6,X9))))) <=> p(s(bool,happ(s(fun(X6,bool),X11),s(X6,X20)))))),
  file('thm.option.OPTREL_MONO','thm.bool.UNWIND_THM2')).
fof(f1859,plain,(
  ( ! [X26] : (s(bool,'T_0') = s(bool,X26)) ) | $spl6),
  inference(cnf_transformation,[],[f1859_D])).
fof(f1859_D,plain,(
  ( ! [X26] : (s(bool,'T_0') = s(bool,X26)) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f20716,plain,(
  ( ! [X87,X86] : (~p(s(bool,happ(s(fun(bool,bool),X87),s(bool,X86)))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X86))) | s(bool,'F_0') = s(bool,sK6('F_0',X87,bool))) ) | $spl140),
  inference(cnf_transformation,[],[f20716_D])).
fof(f20716_D,plain,(
  ( ! [X87,X86] : (~p(s(bool,happ(s(fun(bool,bool),X87),s(bool,X86)))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X86))) | s(bool,'F_0') = s(bool,sK6('F_0',X87,bool))) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f20693,plain,(
  ( ! [X87,X85,X86] : (~p(s(bool,happ(s(fun(bool,bool),X87),s(bool,X86)))) | s(bool,'F_0') = s(bool,sK6('F_0',X87,bool)) | s(bool,'T_0') = s(bool,X85) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X86)))) ) | $spl56),
  inference(duplicate_literal_removal,[],[f20637])).
fof(f20637,plain,(
  ( ! [X87,X85,X86] : (~p(s(bool,happ(s(fun(bool,bool),X87),s(bool,X86)))) | s(bool,'F_0') = s(bool,sK6('F_0',X87,bool)) | s(bool,'T_0') = s(bool,X85) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X86))) | s(bool,'T_0') = s(bool,X85)) ) | $spl56),
  inference(superposition,[],[f926,f13334])).
fof(f19749,plain,(
  ( ! [X94,X95,X96] : (s(bool,'T_0') = s(bool,happ(s(fun(X94,bool),X95),s(X94,sK6(X96,X95,X94)))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(X94,bool),X95),s(X94,X96)))))) ) | $spl138),
  inference(cnf_transformation,[],[f19749_D])).
fof(f19749_D,plain,(
  ( ! [X94,X95,X96] : (s(bool,'T_0') = s(bool,happ(s(fun(X94,bool),X95),s(X94,sK6(X96,X95,X94)))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(X94,bool),X95),s(X94,X96)))))) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f19748,plain,(
  ( ! [X94,X97,X95,X96] : (~p(s(bool,X97)) | s(bool,'T_0') = s(bool,happ(s(fun(X94,bool),X95),s(X94,sK6(X96,X95,X94)))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(X94,bool),X95),s(X94,X96)))))) ) | $spl42),
  inference(subsumption_resolution,[],[f19722,f135])).
fof(f19722,plain,(
  ( ! [X94,X97,X95,X96] : (~p(s(bool,X97)) | s(bool,'T_0') = s(bool,happ(s(fun(X94,bool),X95),s(X94,sK6(X96,X95,X94)))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(X94,bool),X95),s(X94,X96))))) | s(bool,'F_0') = s(bool,X97)) ) | $spl42),
  inference(superposition,[],[f1004,f12043])).
fof(f1004,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,happ(s(fun(X16,bool),X17),s(X16,X18)))) | s(bool,'T_0') = s(bool,happ(s(fun(X16,bool),X17),s(X16,sK6(X18,X17,X16))))) )),
  inference(subsumption_resolution,[],[f1003,f128])).
fof(f1003,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,'F_0')) | ~p(s(bool,happ(s(fun(X16,bool),X17),s(X16,X18)))) | s(bool,'T_0') = s(bool,happ(s(fun(X16,bool),X17),s(X16,sK6(X18,X17,X16))))) )),
  inference(forward_demodulation,[],[f996,f232])).
fof(f996,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)))) | ~p(s(bool,happ(s(fun(X16,bool),X17),s(X16,X18)))) | s(bool,'T_0') = s(bool,happ(s(fun(X16,bool),X17),s(X16,sK6(X18,X17,X16))))) )),
  inference(superposition,[],[f180,f335])).
fof(f335,plain,(
  ( ! [X1] : (s(bool,X1) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(subsumption_resolution,[],[f332,f128])).
fof(f332,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | s(bool,X1) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f296,f232])).
fof(f296,plain,(
  ( ! [X4,X5] : (p(s(bool,X5)) | s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f289,f128])).
fof(f289,plain,(
  ( ! [X4,X5] : (p(s(bool,'F_0')) | p(s(bool,X5)) | s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f181,f129])).
fof(f181,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f74])).
fof(f74,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f64])).
fof(f64,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X14,X15] : ((p(s(bool,X14)) => p(s(bool,X15))) => ((p(s(bool,X15)) => p(s(bool,X14))) => s(bool,X14) = s(bool,X15)))),
  file('thm.option.OPTREL_MONO','thm.bool.IMP_ANTISYM_AX')).
fof(f180,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK6(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) )),
  inference(cnf_transformation,[],[f113])).
fof(f19634,plain,(
  ( ! [X167,X165,X163,X166,X164] : (s('type.option.option'(X164),'const.option.NONE_0') = s('type.option.option'(X164),X167) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.OPTREL_3'(s(fun(X163,fun(X164,bool)),X165),s('type.option.option'(X163),X166),s('type.option.option'(X164),X167))))) | p(s(bool,happ(s(fun(X164,bool),happ(s(fun(X163,fun(X164,bool)),X165),s(X163,sK8(X167,X166,X165,X164,X163)))),s(X164,sK9(X167,X166,X165,X164,X163)))))) ) | $spl136),
  inference(cnf_transformation,[],[f19634_D])).
fof(f19634_D,plain,(
  ( ! [X167,X165,X163,X166,X164] : (s('type.option.option'(X164),'const.option.NONE_0') = s('type.option.option'(X164),X167) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.OPTREL_3'(s(fun(X163,fun(X164,bool)),X165),s('type.option.option'(X163),X166),s('type.option.option'(X164),X167))))) | p(s(bool,happ(s(fun(X164,bool),happ(s(fun(X163,fun(X164,bool)),X165),s(X163,sK8(X167,X166,X165,X164,X163)))),s(X164,sK9(X167,X166,X165,X164,X163)))))) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f19633,plain,(
  ( ! [X167,X165,X163,X166,X164,X168] : (~p(s(bool,X168)) | s('type.option.option'(X164),'const.option.NONE_0') = s('type.option.option'(X164),X167) | p(s(bool,happ(s(fun(X164,bool),happ(s(fun(X163,fun(X164,bool)),X165),s(X163,sK8(X167,X166,X165,X164,X163)))),s(X164,sK9(X167,X166,X165,X164,X163))))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.OPTREL_3'(s(fun(X163,fun(X164,bool)),X165),s('type.option.option'(X163),X166),s('type.option.option'(X164),X167)))))) ) | $spl42),
  inference(subsumption_resolution,[],[f19614,f135])).
fof(f19614,plain,(
  ( ! [X167,X165,X163,X166,X164,X168] : (~p(s(bool,X168)) | s('type.option.option'(X164),'const.option.NONE_0') = s('type.option.option'(X164),X167) | p(s(bool,happ(s(fun(X164,bool),happ(s(fun(X163,fun(X164,bool)),X165),s(X163,sK8(X167,X166,X165,X164,X163)))),s(X164,sK9(X167,X166,X165,X164,X163))))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.OPTREL_3'(s(fun(X163,fun(X164,bool)),X165),s('type.option.option'(X163),X166),s('type.option.option'(X164),X167))))) | s(bool,'F_0') = s(bool,X168)) ) | $spl42),
  inference(superposition,[],[f191,f12043])).
fof(f191,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.option.option'(X0),X3),s('type.option.option'(X1),X4)))) | s('type.option.option'(X1),'const.option.NONE_0') = s('type.option.option'(X1),X4) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,sK8(X4,X3,X2,X1,X0)))),s(X1,sK9(X4,X3,X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f118])).
fof(f13325,plain,(
  ( ! [X16] : (s(bool,'F_0') = s(bool,X16)) ) | $spl52),
  inference(cnf_transformation,[],[f13325_D])).
fof(f13325_D,plain,(
  ( ! [X16] : (s(bool,'F_0') = s(bool,X16)) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f19536,plain,(
  ( ! [X47,X50,X48,X49] : (s(bool,X47) != s(bool,'const.option.OPTREL_3'(s(fun(X48,fun(sK1,bool)),X49),s('type.option.option'(X48),'const.option.NONE_0'),s('type.option.option'(sK1),sK2))) | s(bool,'T_0') = s(bool,X50) | s(bool,X47) = s(bool,X50)) ) | $spl134),
  inference(cnf_transformation,[],[f19536_D])).
fof(f19536_D,plain,(
  ( ! [X47,X50,X48,X49] : (s(bool,X47) != s(bool,'const.option.OPTREL_3'(s(fun(X48,fun(sK1,bool)),X49),s('type.option.option'(X48),'const.option.NONE_0'),s('type.option.option'(sK1),sK2))) | s(bool,'T_0') = s(bool,X50) | s(bool,X47) = s(bool,X50)) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f19526,plain,(
  ( ! [X47,X50,X48,X46,X49] : (s(bool,X47) != s(bool,'const.option.OPTREL_3'(s(fun(X48,fun(sK1,bool)),X49),s('type.option.option'(X48),'const.option.NONE_0'),s('type.option.option'(sK1),sK2))) | s(bool,'F_0') = s(bool,X46) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) ) | ($spl8 | $spl21 | $spl33)),
  inference(duplicate_literal_removal,[],[f19463])).
fof(f19463,plain,(
  ( ! [X47,X50,X48,X46,X49] : (s(bool,X47) != s(bool,'const.option.OPTREL_3'(s(fun(X48,fun(sK1,bool)),X49),s('type.option.option'(X48),'const.option.NONE_0'),s('type.option.option'(sK1),sK2))) | s(bool,'F_0') = s(bool,X46) | s(bool,'F_0') = s(bool,X46) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) ) | ($spl8 | $spl21 | $spl33)),
  inference(superposition,[],[f8519,f2533])).
fof(f2533,plain,(
  ( ! [X4,X2,X3] : (s(bool,X2) = s(bool,X3) | s(bool,'F_0') = s(bool,X3) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) ) | $spl8),
  inference(superposition,[],[f1860,f982])).
fof(f982,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) )),
  inference(trivial_inequality_removal,[],[f962])).
fof(f962,plain,(
  ( ! [X0,X1] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | s(bool,X0) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f774,f259])).
fof(f259,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f129,f129])).
fof(f774,plain,(
  ( ! [X24] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X24))) | s(bool,X24) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)))) )),
  inference(superposition,[],[f176,f227])).
fof(f176,plain,(
  ( ! [X2,X0,X1] : (s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1))) != s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2))) | s(X0,X1) = s(X0,X2)) )),
  inference(cnf_transformation,[],[f110])).
fof(f110,plain,(
  ! [X0,X1,X2] : ((s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1))) != s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1))) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2)))))),
  inference(nnf_transformation,[],[f61])).
fof(f61,plain,(
  ! [X0,X1,X2] : (s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1))) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2))) <=> s(X0,X1) = s(X0,X2))),
  inference(flattening,[],[f60])).
fof(f60,plain,(
  ! [X0] : ! [X1,X2] : (s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1))) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X6] : ! [X9,X8] : (s('type.option.option'(X6),'const.option.SOME_1'(s(X6,X8))) = s('type.option.option'(X6),'const.option.SOME_1'(s(X6,X9))) <=> s(X6,X8) = s(X6,X9))),
  file('thm.option.OPTREL_MONO','thm.option.SOME_|49||49|')).
fof(f8519,plain,(
  ( ! [X94,X95,X93] : (s(bool,X95) != s(bool,'const.option.OPTREL_3'(s(fun(X93,fun(sK1,bool)),X94),s('type.option.option'(X93),'const.option.NONE_0'),s('type.option.option'(sK1),sK2))) | s(bool,'F_0') = s(bool,X95)) ) | ($spl8 | $spl21 | $spl33)),
  inference(subsumption_resolution,[],[f8481,f128])).
fof(f8481,plain,(
  ( ! [X94,X95,X93] : (p(s(bool,'F_0')) | s(bool,X95) != s(bool,'const.option.OPTREL_3'(s(fun(X93,fun(sK1,bool)),X94),s('type.option.option'(X93),'const.option.NONE_0'),s('type.option.option'(sK1),sK2))) | s(bool,'F_0') = s(bool,X95)) ) | ($spl8 | $spl21 | $spl33)),
  inference(superposition,[],[f2646,f7985])).
fof(f7985,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') = s(bool,'const.option.OPTREL_3'(s(fun(X2,fun(sK1,bool)),X3),s('type.option.option'(X2),'const.option.NONE_0'),s('type.option.option'(sK1),sK2)))) ) | ($spl8 | $spl21 | $spl33)),
  inference(forward_demodulation,[],[f7968,f6243])).
fof(f6243,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(sK0,bool)),X1),s('type.option.option'(X0),'const.option.NONE_0'),s('type.option.option'(sK0),sK3)))) ) | ($spl8 | $spl21)),
  inference(unit_resulting_resolution,[],[f223,f223,f6206,f3500])).
fof(f3500,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X10)) | s(bool,X9) = s(bool,X10) | s(bool,X9) = s(bool,X11) | s(bool,'F_0') = s(bool,X11)) ) | $spl8),
  inference(subsumption_resolution,[],[f3147,f128])).
fof(f3147,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,'F_0')) | p(s(bool,X10)) | s(bool,X9) = s(bool,X10) | s(bool,X9) = s(bool,X11) | s(bool,'F_0') = s(bool,X11)) ) | $spl8),
  inference(superposition,[],[f181,f2532])).
fof(f2532,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | $spl8),
  inference(superposition,[],[f1860,f1860])).
fof(f6206,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(sK0,bool)),X1),s('type.option.option'(X0),'const.option.NONE_0'),s('type.option.option'(sK0),sK3))))) ) | $spl21),
  inference(unit_resulting_resolution,[],[f3826,f173,f187])).
fof(f187,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.option.option'(X0),X3),s('type.option.option'(X1),X4)))) | s('type.option.option'(X1),'const.option.NONE_0') = s('type.option.option'(X1),X4) | s('type.option.option'(X0),X3) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,sK8(X4,X3,X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f118])).
fof(f173,plain,(
  ( ! [X0,X1] : (s('type.option.option'(X0),'const.option.NONE_0') != s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1)))) )),
  inference(cnf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0,X1] : s('type.option.option'(X0),'const.option.NONE_0') != s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1)))),
  inference(flattening,[],[f56])).
fof(f56,plain,(
  ! [X0] : ! [X1] : ~s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1)))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X6] : ! [X9] : ~s('type.option.option'(X6),'const.option.NONE_0') = s('type.option.option'(X6),'const.option.SOME_1'(s(X6,X9)))),
  file('thm.option.OPTREL_MONO','thm.option.NOT_NONE_SOME')).
fof(f223,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f128,f131])).
fof(f131,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f82])).
fof(f7968,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(sK0,bool)),X1),s('type.option.option'(X0),'const.option.NONE_0'),s('type.option.option'(sK0),sK3))) = s(bool,'const.option.OPTREL_3'(s(fun(X2,fun(sK1,bool)),X3),s('type.option.option'(X2),'const.option.NONE_0'),s('type.option.option'(sK1),sK2)))) ) | ($spl8 | $spl21 | $spl33)),
  inference(unit_resulting_resolution,[],[f6242,f7927,f296])).
fof(f7927,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(sK1,bool)),X1),s('type.option.option'(X0),'const.option.NONE_0'),s('type.option.option'(sK1),sK2))))) ) | $spl33),
  inference(unit_resulting_resolution,[],[f6508,f173,f187])).
fof(f6508,plain,(
  s('type.option.option'(sK1),'const.option.NONE_0') != s('type.option.option'(sK1),sK2) | $spl33),
  inference(cnf_transformation,[],[f6508_D])).
fof(f6508_D,plain,(
  s('type.option.option'(sK1),'const.option.NONE_0') != s('type.option.option'(sK1),sK2) <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f6242,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(sK0,bool)),X1),s('type.option.option'(X0),'const.option.NONE_0'),s('type.option.option'(sK0),sK3)))) ) | ($spl8 | $spl21)),
  inference(unit_resulting_resolution,[],[f223,f6206,f2646])).
fof(f2646,plain,(
  ( ! [X12,X11] : (p(s(bool,X12)) | s(bool,X11) != s(bool,X12) | s(bool,'F_0') = s(bool,X11)) ) | $spl8),
  inference(subsumption_resolution,[],[f2645,f126])).
fof(f2645,plain,(
  ( ! [X12,X11] : (~p(s(bool,'T_0')) | s(bool,X11) != s(bool,X12) | p(s(bool,X12)) | s(bool,'F_0') = s(bool,X11)) ) | $spl8),
  inference(forward_demodulation,[],[f2554,f227])).
fof(f2554,plain,(
  ( ! [X12,X11] : (~p(s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)))) | s(bool,X11) != s(bool,X12) | p(s(bool,X12)) | s(bool,'F_0') = s(bool,X11)) ) | $spl8),
  inference(superposition,[],[f219,f1860])).
fof(f219,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,X0) != s(bool,X1) | p(s(bool,X1))) )),
  inference(condensation,[],[f218])).
fof(f218,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2)) )),
  inference(condensation,[],[f207])).
fof(f207,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f121])).
fof(f121,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f79])).
fof(f79,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f70])).
fof(f70,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X9,X18,X8,X19] : ((s(bool,X9) = s(bool,X18) & (p(s(bool,X18)) => s(bool,X8) = s(bool,X19))) => ((p(s(bool,X9)) => p(s(bool,X8))) <=> (p(s(bool,X18)) => p(s(bool,X19)))))),
  file('thm.option.OPTREL_MONO','thm.bool.IMP_CONG')).
fof(f19338,plain,(
  ( ! [X167,X165,X163,X166,X164] : (s('type.option.option'(X163),'const.option.NONE_0') = s('type.option.option'(X163),X166) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.OPTREL_3'(s(fun(X163,fun(X164,bool)),X165),s('type.option.option'(X163),X166),s('type.option.option'(X164),X167))))) | p(s(bool,happ(s(fun(X164,bool),happ(s(fun(X163,fun(X164,bool)),X165),s(X163,sK8(X167,X166,X165,X164,X163)))),s(X164,sK9(X167,X166,X165,X164,X163)))))) ) | $spl132),
  inference(cnf_transformation,[],[f19338_D])).
fof(f19338_D,plain,(
  ( ! [X167,X165,X163,X166,X164] : (s('type.option.option'(X163),'const.option.NONE_0') = s('type.option.option'(X163),X166) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.OPTREL_3'(s(fun(X163,fun(X164,bool)),X165),s('type.option.option'(X163),X166),s('type.option.option'(X164),X167))))) | p(s(bool,happ(s(fun(X164,bool),happ(s(fun(X163,fun(X164,bool)),X165),s(X163,sK8(X167,X166,X165,X164,X163)))),s(X164,sK9(X167,X166,X165,X164,X163)))))) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f19337,plain,(
  ( ! [X167,X165,X163,X166,X164,X168] : (~p(s(bool,X168)) | s('type.option.option'(X163),'const.option.NONE_0') = s('type.option.option'(X163),X166) | p(s(bool,happ(s(fun(X164,bool),happ(s(fun(X163,fun(X164,bool)),X165),s(X163,sK8(X167,X166,X165,X164,X163)))),s(X164,sK9(X167,X166,X165,X164,X163))))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.OPTREL_3'(s(fun(X163,fun(X164,bool)),X165),s('type.option.option'(X163),X166),s('type.option.option'(X164),X167)))))) ) | $spl42),
  inference(subsumption_resolution,[],[f19318,f135])).
fof(f19318,plain,(
  ( ! [X167,X165,X163,X166,X164,X168] : (~p(s(bool,X168)) | s('type.option.option'(X163),'const.option.NONE_0') = s('type.option.option'(X163),X166) | p(s(bool,happ(s(fun(X164,bool),happ(s(fun(X163,fun(X164,bool)),X165),s(X163,sK8(X167,X166,X165,X164,X163)))),s(X164,sK9(X167,X166,X165,X164,X163))))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.OPTREL_3'(s(fun(X163,fun(X164,bool)),X165),s('type.option.option'(X163),X166),s('type.option.option'(X164),X167))))) | s(bool,'F_0') = s(bool,X168)) ) | $spl42),
  inference(superposition,[],[f190,f12043])).
fof(f19240,plain,(
  ( ! [X47,X50,X48,X49] : (s(bool,X47) != s(bool,'const.option.OPTREL_3'(s(fun(sK1,fun(X48,bool)),X49),s('type.option.option'(sK1),sK2),s('type.option.option'(X48),'const.option.NONE_0'))) | s(bool,'T_0') = s(bool,X50) | s(bool,X47) = s(bool,X50)) ) | $spl130),
  inference(cnf_transformation,[],[f19240_D])).
fof(f19240_D,plain,(
  ( ! [X47,X50,X48,X49] : (s(bool,X47) != s(bool,'const.option.OPTREL_3'(s(fun(sK1,fun(X48,bool)),X49),s('type.option.option'(sK1),sK2),s('type.option.option'(X48),'const.option.NONE_0'))) | s(bool,'T_0') = s(bool,X50) | s(bool,X47) = s(bool,X50)) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f19230,plain,(
  ( ! [X47,X50,X48,X46,X49] : (s(bool,X47) != s(bool,'const.option.OPTREL_3'(s(fun(sK1,fun(X48,bool)),X49),s('type.option.option'(sK1),sK2),s('type.option.option'(X48),'const.option.NONE_0'))) | s(bool,'F_0') = s(bool,X46) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) ) | ($spl8 | $spl21 | $spl33)),
  inference(duplicate_literal_removal,[],[f19171])).
fof(f19171,plain,(
  ( ! [X47,X50,X48,X46,X49] : (s(bool,X47) != s(bool,'const.option.OPTREL_3'(s(fun(sK1,fun(X48,bool)),X49),s('type.option.option'(sK1),sK2),s('type.option.option'(X48),'const.option.NONE_0'))) | s(bool,'F_0') = s(bool,X46) | s(bool,'F_0') = s(bool,X46) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) ) | ($spl8 | $spl21 | $spl33)),
  inference(superposition,[],[f8421,f2533])).
fof(f8421,plain,(
  ( ! [X94,X95,X93] : (s(bool,X95) != s(bool,'const.option.OPTREL_3'(s(fun(sK1,fun(X93,bool)),X94),s('type.option.option'(sK1),sK2),s('type.option.option'(X93),'const.option.NONE_0'))) | s(bool,'F_0') = s(bool,X95)) ) | ($spl8 | $spl21 | $spl33)),
  inference(subsumption_resolution,[],[f8383,f128])).
fof(f8383,plain,(
  ( ! [X94,X95,X93] : (p(s(bool,'F_0')) | s(bool,X95) != s(bool,'const.option.OPTREL_3'(s(fun(sK1,fun(X93,bool)),X94),s('type.option.option'(sK1),sK2),s('type.option.option'(X93),'const.option.NONE_0'))) | s(bool,'F_0') = s(bool,X95)) ) | ($spl8 | $spl21 | $spl33)),
  inference(superposition,[],[f2646,f8166])).
fof(f8166,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') = s(bool,'const.option.OPTREL_3'(s(fun(sK1,fun(X2,bool)),X3),s('type.option.option'(sK1),sK2),s('type.option.option'(X2),'const.option.NONE_0')))) ) | ($spl8 | $spl21 | $spl33)),
  inference(forward_demodulation,[],[f8149,f7985])).
fof(f8149,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(sK1,bool)),X1),s('type.option.option'(X0),'const.option.NONE_0'),s('type.option.option'(sK1),sK2))) = s(bool,'const.option.OPTREL_3'(s(fun(sK1,fun(X2,bool)),X3),s('type.option.option'(sK1),sK2),s('type.option.option'(X2),'const.option.NONE_0')))) ) | ($spl8 | $spl33)),
  inference(unit_resulting_resolution,[],[f7969,f8107,f296])).
fof(f8107,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.option.OPTREL_3'(s(fun(sK1,fun(X0,bool)),X1),s('type.option.option'(sK1),sK2),s('type.option.option'(X0),'const.option.NONE_0'))))) ) | $spl33),
  inference(unit_resulting_resolution,[],[f6508,f173,f188])).
fof(f7969,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(sK1,bool)),X1),s('type.option.option'(X0),'const.option.NONE_0'),s('type.option.option'(sK1),sK2)))) ) | ($spl8 | $spl33)),
  inference(unit_resulting_resolution,[],[f223,f7927,f2646])).
fof(f19138,plain,(
  ( ! [X47,X50,X48,X49] : (s(bool,X47) != s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(X48,bool)),X49),s('type.option.option'(sK0),sK3),s('type.option.option'(X48),'const.option.NONE_0'))) | s(bool,'T_0') = s(bool,X50) | s(bool,X47) = s(bool,X50)) ) | $spl128),
  inference(cnf_transformation,[],[f19138_D])).
fof(f19138_D,plain,(
  ( ! [X47,X50,X48,X49] : (s(bool,X47) != s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(X48,bool)),X49),s('type.option.option'(sK0),sK3),s('type.option.option'(X48),'const.option.NONE_0'))) | s(bool,'T_0') = s(bool,X50) | s(bool,X47) = s(bool,X50)) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f19128,plain,(
  ( ! [X47,X50,X48,X46,X49] : (s(bool,X47) != s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(X48,bool)),X49),s('type.option.option'(sK0),sK3),s('type.option.option'(X48),'const.option.NONE_0'))) | s(bool,'F_0') = s(bool,X46) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) ) | ($spl8 | $spl21)),
  inference(duplicate_literal_removal,[],[f19065])).
fof(f19065,plain,(
  ( ! [X47,X50,X48,X46,X49] : (s(bool,X47) != s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(X48,bool)),X49),s('type.option.option'(sK0),sK3),s('type.option.option'(X48),'const.option.NONE_0'))) | s(bool,'F_0') = s(bool,X46) | s(bool,'F_0') = s(bool,X46) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) ) | ($spl8 | $spl21)),
  inference(superposition,[],[f7857,f2533])).
fof(f7857,plain,(
  ( ! [X87,X85,X86] : (s(bool,X87) != s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(X85,bool)),X86),s('type.option.option'(sK0),sK3),s('type.option.option'(X85),'const.option.NONE_0'))) | s(bool,'F_0') = s(bool,X87)) ) | ($spl8 | $spl21)),
  inference(subsumption_resolution,[],[f7821,f128])).
fof(f7821,plain,(
  ( ! [X87,X85,X86] : (p(s(bool,'F_0')) | s(bool,X87) != s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(X85,bool)),X86),s('type.option.option'(sK0),sK3),s('type.option.option'(X85),'const.option.NONE_0'))) | s(bool,'F_0') = s(bool,X87)) ) | ($spl8 | $spl21)),
  inference(superposition,[],[f2646,f6301])).
fof(f6301,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(X0,bool)),X1),s('type.option.option'(sK0),sK3),s('type.option.option'(X0),'const.option.NONE_0')))) ) | ($spl8 | $spl21)),
  inference(unit_resulting_resolution,[],[f223,f223,f6268,f3500])).
fof(f6268,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(X0,bool)),X1),s('type.option.option'(sK0),sK3),s('type.option.option'(X0),'const.option.NONE_0'))))) ) | $spl21),
  inference(unit_resulting_resolution,[],[f3826,f173,f188])).
fof(f19049,plain,(
  ( ! [X167,X165,X163,X166,X164] : (s('type.option.option'(X164),'const.option.NONE_0') = s('type.option.option'(X164),X167) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.OPTREL_3'(s(fun(X163,fun(X164,bool)),X165),s('type.option.option'(X163),X166),s('type.option.option'(X164),X167))))) | s('type.option.option'(X164),X167) = s('type.option.option'(X164),'const.option.SOME_1'(s(X164,sK9(X167,X166,X165,X164,X163))))) ) | $spl126),
  inference(cnf_transformation,[],[f19049_D])).
fof(f19049_D,plain,(
  ( ! [X167,X165,X163,X166,X164] : (s('type.option.option'(X164),'const.option.NONE_0') = s('type.option.option'(X164),X167) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.OPTREL_3'(s(fun(X163,fun(X164,bool)),X165),s('type.option.option'(X163),X166),s('type.option.option'(X164),X167))))) | s('type.option.option'(X164),X167) = s('type.option.option'(X164),'const.option.SOME_1'(s(X164,sK9(X167,X166,X165,X164,X163))))) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f19048,plain,(
  ( ! [X167,X165,X163,X166,X164,X168] : (~p(s(bool,X168)) | s('type.option.option'(X164),'const.option.NONE_0') = s('type.option.option'(X164),X167) | s('type.option.option'(X164),X167) = s('type.option.option'(X164),'const.option.SOME_1'(s(X164,sK9(X167,X166,X165,X164,X163)))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.OPTREL_3'(s(fun(X163,fun(X164,bool)),X165),s('type.option.option'(X163),X166),s('type.option.option'(X164),X167)))))) ) | $spl42),
  inference(subsumption_resolution,[],[f19025,f135])).
fof(f19025,plain,(
  ( ! [X167,X165,X163,X166,X164,X168] : (~p(s(bool,X168)) | s('type.option.option'(X164),'const.option.NONE_0') = s('type.option.option'(X164),X167) | s('type.option.option'(X164),X167) = s('type.option.option'(X164),'const.option.SOME_1'(s(X164,sK9(X167,X166,X165,X164,X163)))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.OPTREL_3'(s(fun(X163,fun(X164,bool)),X165),s('type.option.option'(X163),X166),s('type.option.option'(X164),X167))))) | s(bool,'F_0') = s(bool,X168)) ) | $spl42),
  inference(superposition,[],[f189,f12043])).
fof(f189,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.option.option'(X0),X3),s('type.option.option'(X1),X4)))) | s('type.option.option'(X1),'const.option.NONE_0') = s('type.option.option'(X1),X4) | s('type.option.option'(X1),X4) = s('type.option.option'(X1),'const.option.SOME_1'(s(X1,sK9(X4,X3,X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f118])).
fof(f18947,plain,(
  ( ! [X47,X50,X48,X49] : (s(bool,X47) != s(bool,'const.option.OPTREL_3'(s(fun(X48,fun(sK0,bool)),X49),s('type.option.option'(X48),'const.option.NONE_0'),s('type.option.option'(sK0),sK3))) | s(bool,'T_0') = s(bool,X50) | s(bool,X47) = s(bool,X50)) ) | $spl124),
  inference(cnf_transformation,[],[f18947_D])).
fof(f18947_D,plain,(
  ( ! [X47,X50,X48,X49] : (s(bool,X47) != s(bool,'const.option.OPTREL_3'(s(fun(X48,fun(sK0,bool)),X49),s('type.option.option'(X48),'const.option.NONE_0'),s('type.option.option'(sK0),sK3))) | s(bool,'T_0') = s(bool,X50) | s(bool,X47) = s(bool,X50)) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f18937,plain,(
  ( ! [X47,X50,X48,X46,X49] : (s(bool,X47) != s(bool,'const.option.OPTREL_3'(s(fun(X48,fun(sK0,bool)),X49),s('type.option.option'(X48),'const.option.NONE_0'),s('type.option.option'(sK0),sK3))) | s(bool,'F_0') = s(bool,X46) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) ) | ($spl8 | $spl21)),
  inference(duplicate_literal_removal,[],[f18874])).
fof(f18874,plain,(
  ( ! [X47,X50,X48,X46,X49] : (s(bool,X47) != s(bool,'const.option.OPTREL_3'(s(fun(X48,fun(sK0,bool)),X49),s('type.option.option'(X48),'const.option.NONE_0'),s('type.option.option'(sK0),sK3))) | s(bool,'F_0') = s(bool,X46) | s(bool,'F_0') = s(bool,X46) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) ) | ($spl8 | $spl21)),
  inference(superposition,[],[f7668,f2533])).
fof(f7668,plain,(
  ( ! [X87,X85,X86] : (s(bool,X87) != s(bool,'const.option.OPTREL_3'(s(fun(X85,fun(sK0,bool)),X86),s('type.option.option'(X85),'const.option.NONE_0'),s('type.option.option'(sK0),sK3))) | s(bool,'F_0') = s(bool,X87)) ) | ($spl8 | $spl21)),
  inference(subsumption_resolution,[],[f7633,f128])).
fof(f7633,plain,(
  ( ! [X87,X85,X86] : (p(s(bool,'F_0')) | s(bool,X87) != s(bool,'const.option.OPTREL_3'(s(fun(X85,fun(sK0,bool)),X86),s('type.option.option'(X85),'const.option.NONE_0'),s('type.option.option'(sK0),sK3))) | s(bool,'F_0') = s(bool,X87)) ) | ($spl8 | $spl21)),
  inference(superposition,[],[f2646,f6243])).
fof(f18734,plain,(
  ( ! [X167,X165,X163,X166,X164] : (s('type.option.option'(X163),'const.option.NONE_0') = s('type.option.option'(X163),X166) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.OPTREL_3'(s(fun(X163,fun(X164,bool)),X165),s('type.option.option'(X163),X166),s('type.option.option'(X164),X167))))) | s('type.option.option'(X164),X167) = s('type.option.option'(X164),'const.option.SOME_1'(s(X164,sK9(X167,X166,X165,X164,X163))))) ) | $spl122),
  inference(cnf_transformation,[],[f18734_D])).
fof(f18734_D,plain,(
  ( ! [X167,X165,X163,X166,X164] : (s('type.option.option'(X163),'const.option.NONE_0') = s('type.option.option'(X163),X166) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.OPTREL_3'(s(fun(X163,fun(X164,bool)),X165),s('type.option.option'(X163),X166),s('type.option.option'(X164),X167))))) | s('type.option.option'(X164),X167) = s('type.option.option'(X164),'const.option.SOME_1'(s(X164,sK9(X167,X166,X165,X164,X163))))) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f18733,plain,(
  ( ! [X167,X165,X163,X166,X164,X168] : (~p(s(bool,X168)) | s('type.option.option'(X163),'const.option.NONE_0') = s('type.option.option'(X163),X166) | s('type.option.option'(X164),X167) = s('type.option.option'(X164),'const.option.SOME_1'(s(X164,sK9(X167,X166,X165,X164,X163)))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.OPTREL_3'(s(fun(X163,fun(X164,bool)),X165),s('type.option.option'(X163),X166),s('type.option.option'(X164),X167)))))) ) | $spl42),
  inference(subsumption_resolution,[],[f18709,f135])).
fof(f18709,plain,(
  ( ! [X167,X165,X163,X166,X164,X168] : (~p(s(bool,X168)) | s('type.option.option'(X163),'const.option.NONE_0') = s('type.option.option'(X163),X166) | s('type.option.option'(X164),X167) = s('type.option.option'(X164),'const.option.SOME_1'(s(X164,sK9(X167,X166,X165,X164,X163)))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.OPTREL_3'(s(fun(X163,fun(X164,bool)),X165),s('type.option.option'(X163),X166),s('type.option.option'(X164),X167))))) | s(bool,'F_0') = s(bool,X168)) ) | $spl42),
  inference(superposition,[],[f188,f12043])).
fof(f18225,plain,(
  ( ! [X167,X165,X163,X166,X164] : (s('type.option.option'(X164),'const.option.NONE_0') = s('type.option.option'(X164),X167) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.OPTREL_3'(s(fun(X163,fun(X164,bool)),X165),s('type.option.option'(X163),X166),s('type.option.option'(X164),X167))))) | s('type.option.option'(X163),X166) = s('type.option.option'(X163),'const.option.SOME_1'(s(X163,sK8(X167,X166,X165,X164,X163))))) ) | $spl120),
  inference(cnf_transformation,[],[f18225_D])).
fof(f18225_D,plain,(
  ( ! [X167,X165,X163,X166,X164] : (s('type.option.option'(X164),'const.option.NONE_0') = s('type.option.option'(X164),X167) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.OPTREL_3'(s(fun(X163,fun(X164,bool)),X165),s('type.option.option'(X163),X166),s('type.option.option'(X164),X167))))) | s('type.option.option'(X163),X166) = s('type.option.option'(X163),'const.option.SOME_1'(s(X163,sK8(X167,X166,X165,X164,X163))))) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f18224,plain,(
  ( ! [X167,X165,X163,X166,X164,X168] : (~p(s(bool,X168)) | s('type.option.option'(X164),'const.option.NONE_0') = s('type.option.option'(X164),X167) | s('type.option.option'(X163),X166) = s('type.option.option'(X163),'const.option.SOME_1'(s(X163,sK8(X167,X166,X165,X164,X163)))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.OPTREL_3'(s(fun(X163,fun(X164,bool)),X165),s('type.option.option'(X163),X166),s('type.option.option'(X164),X167)))))) ) | $spl42),
  inference(subsumption_resolution,[],[f18201,f135])).
fof(f18201,plain,(
  ( ! [X167,X165,X163,X166,X164,X168] : (~p(s(bool,X168)) | s('type.option.option'(X164),'const.option.NONE_0') = s('type.option.option'(X164),X167) | s('type.option.option'(X163),X166) = s('type.option.option'(X163),'const.option.SOME_1'(s(X163,sK8(X167,X166,X165,X164,X163)))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.OPTREL_3'(s(fun(X163,fun(X164,bool)),X165),s('type.option.option'(X163),X166),s('type.option.option'(X164),X167))))) | s(bool,'F_0') = s(bool,X168)) ) | $spl42),
  inference(superposition,[],[f187,f12043])).
fof(f17949,plain,(
  ( ! [X167,X165,X163,X166,X164] : (s('type.option.option'(X163),'const.option.NONE_0') = s('type.option.option'(X163),X166) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.OPTREL_3'(s(fun(X163,fun(X164,bool)),X165),s('type.option.option'(X163),X166),s('type.option.option'(X164),X167))))) | s('type.option.option'(X163),X166) = s('type.option.option'(X163),'const.option.SOME_1'(s(X163,sK8(X167,X166,X165,X164,X163))))) ) | $spl118),
  inference(cnf_transformation,[],[f17949_D])).
fof(f17949_D,plain,(
  ( ! [X167,X165,X163,X166,X164] : (s('type.option.option'(X163),'const.option.NONE_0') = s('type.option.option'(X163),X166) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.OPTREL_3'(s(fun(X163,fun(X164,bool)),X165),s('type.option.option'(X163),X166),s('type.option.option'(X164),X167))))) | s('type.option.option'(X163),X166) = s('type.option.option'(X163),'const.option.SOME_1'(s(X163,sK8(X167,X166,X165,X164,X163))))) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f17948,plain,(
  ( ! [X167,X165,X163,X166,X164,X168] : (~p(s(bool,X168)) | s('type.option.option'(X163),'const.option.NONE_0') = s('type.option.option'(X163),X166) | s('type.option.option'(X163),X166) = s('type.option.option'(X163),'const.option.SOME_1'(s(X163,sK8(X167,X166,X165,X164,X163)))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.OPTREL_3'(s(fun(X163,fun(X164,bool)),X165),s('type.option.option'(X163),X166),s('type.option.option'(X164),X167)))))) ) | $spl42),
  inference(subsumption_resolution,[],[f17926,f135])).
fof(f17926,plain,(
  ( ! [X167,X165,X163,X166,X164,X168] : (~p(s(bool,X168)) | s('type.option.option'(X163),'const.option.NONE_0') = s('type.option.option'(X163),X166) | s('type.option.option'(X163),X166) = s('type.option.option'(X163),'const.option.SOME_1'(s(X163,sK8(X167,X166,X165,X164,X163)))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.OPTREL_3'(s(fun(X163,fun(X164,bool)),X165),s('type.option.option'(X163),X166),s('type.option.option'(X164),X167))))) | s(bool,'F_0') = s(bool,X168)) ) | $spl42),
  inference(superposition,[],[f186,f12043])).
fof(f186,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.option.option'(X0),X3),s('type.option.option'(X1),X4)))) | s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X3) | s('type.option.option'(X0),X3) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,sK8(X4,X3,X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f118])).
fof(f17837,plain,(
  ( ! [X47,X46] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X46))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X47))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X46))) | s(bool,'T_0') = s(bool,X47)) ) | $spl116),
  inference(cnf_transformation,[],[f17837_D])).
fof(f17837_D,plain,(
  ( ! [X47,X46] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X46))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X47))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X46))) | s(bool,'T_0') = s(bool,X47)) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f17500,plain,(
  ( ! [X47,X45,X46] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X46))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X47))) | s(bool,'F_0') = s(bool,X45) | s(bool,'T_0') = s(bool,X47) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X46)))) ) | $spl42),
  inference(duplicate_literal_removal,[],[f17344])).
fof(f17344,plain,(
  ( ! [X47,X45,X46] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X46))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X47))) | s(bool,'F_0') = s(bool,X45) | s(bool,'T_0') = s(bool,X47) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X46))) | s(bool,'F_0') = s(bool,X45)) ) | $spl42),
  inference(superposition,[],[f773,f12043])).
fof(f773,plain,(
  ( ! [X23,X22] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X22))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X23))) | s(bool,'F_0') = s(bool,X23) | s(bool,'T_0') = s(bool,X22)) )),
  inference(superposition,[],[f176,f129])).
fof(f17800,plain,(
  ( ! [X321,X323] : (s(bool,X323) = s(bool,sK6(X323,X321,bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(bool,bool),X321),s(bool,'T_0'))))) | s(bool,'F_0') = s(bool,X323)) ) | $spl114),
  inference(cnf_transformation,[],[f17800_D])).
fof(f17800_D,plain,(
  ( ! [X321,X323] : (s(bool,X323) = s(bool,sK6(X323,X321,bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(bool,bool),X321),s(bool,'T_0'))))) | s(bool,'F_0') = s(bool,X323)) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f17799,plain,(
  ( ! [X321,X323,X322] : (~p(s(bool,X322)) | s(bool,X323) = s(bool,sK6(X323,X321,bool)) | s(bool,'F_0') = s(bool,X323) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(bool,bool),X321),s(bool,'T_0')))))) ) | ($spl8 | $spl42)),
  inference(subsumption_resolution,[],[f17285,f135])).
fof(f17285,plain,(
  ( ! [X321,X323,X322] : (~p(s(bool,X322)) | s(bool,X323) = s(bool,sK6(X323,X321,bool)) | s(bool,'F_0') = s(bool,X323) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(bool,bool),X321),s(bool,'T_0'))))) | s(bool,'F_0') = s(bool,X322)) ) | ($spl8 | $spl42)),
  inference(superposition,[],[f2673,f12043])).
fof(f2673,plain,(
  ( ! [X61,X60] : (~p(s(bool,happ(s(fun(bool,bool),X61),s(bool,'T_0')))) | s(bool,X60) = s(bool,sK6(X60,X61,bool)) | s(bool,'F_0') = s(bool,X60)) ) | $spl8),
  inference(forward_demodulation,[],[f2581,f227])).
fof(f2581,plain,(
  ( ! [X61,X60] : (~p(s(bool,happ(s(fun(bool,bool),X61),s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)))))) | s(bool,X60) = s(bool,sK6(X60,X61,bool)) | s(bool,'F_0') = s(bool,X60)) ) | $spl8),
  inference(superposition,[],[f179,f1860])).
fof(f17798,plain,(
  ( ! [X318,X320] : (s(bool,X320) = s(bool,sK6(X320,X318,bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(bool,bool),X318),s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X320)) ) | $spl112),
  inference(cnf_transformation,[],[f17798_D])).
fof(f17798_D,plain,(
  ( ! [X318,X320] : (s(bool,X320) = s(bool,sK6(X320,X318,bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(bool,bool),X318),s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X320)) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f17797,plain,(
  ( ! [X318,X320,X319] : (~p(s(bool,X319)) | s(bool,X320) = s(bool,sK6(X320,X318,bool)) | s(bool,'T_0') = s(bool,X320) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(bool,bool),X318),s(bool,'F_0')))))) ) | $spl42),
  inference(subsumption_resolution,[],[f17284,f135])).
fof(f17284,plain,(
  ( ! [X318,X320,X319] : (~p(s(bool,X319)) | s(bool,X320) = s(bool,sK6(X320,X318,bool)) | s(bool,'T_0') = s(bool,X320) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(bool,bool),X318),s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X319)) ) | $spl42),
  inference(superposition,[],[f936,f12043])).
fof(f936,plain,(
  ( ! [X12,X13] : (~p(s(bool,happ(s(fun(bool,bool),X13),s(bool,'F_0')))) | s(bool,X12) = s(bool,sK6(X12,X13,bool)) | s(bool,'T_0') = s(bool,X12)) )),
  inference(forward_demodulation,[],[f922,f232])).
fof(f922,plain,(
  ( ! [X12,X13] : (~p(s(bool,happ(s(fun(bool,bool),X13),s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)))))) | s(bool,X12) = s(bool,sK6(X12,X13,bool)) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f179,f335])).
fof(f17796,plain,(
  ( ! [X315,X316] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,X315))),s(sK1,X316)))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK5),s(sK0,X315))),s(sK1,X316)))))) ) | $spl110),
  inference(cnf_transformation,[],[f17796_D])).
fof(f17796_D,plain,(
  ( ! [X315,X316] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,X315))),s(sK1,X316)))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK5),s(sK0,X315))),s(sK1,X316)))))) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f17795,plain,(
  ( ! [X315,X316,X317] : (~p(s(bool,X317)) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,X315))),s(sK1,X316)))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK5),s(sK0,X315))),s(sK1,X316)))))) ) | $spl42),
  inference(subsumption_resolution,[],[f17283,f135])).
fof(f17283,plain,(
  ( ! [X315,X316,X317] : (~p(s(bool,X317)) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,X315))),s(sK1,X316)))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK5),s(sK0,X315))),s(sK1,X316))))) | s(bool,'F_0') = s(bool,X317)) ) | $spl42),
  inference(superposition,[],[f123,f12043])).
fof(f17794,plain,(
  ( ! [X304,X303,X305] : (s(X303,sK6(X305,X304,X303)) = s(X303,sK6(sK6(X305,X304,X303),X304,X303)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(X303,bool),X304),s(X303,X305)))))) ) | $spl108),
  inference(cnf_transformation,[],[f17794_D])).
fof(f17794_D,plain,(
  ( ! [X304,X303,X305] : (s(X303,sK6(X305,X304,X303)) = s(X303,sK6(sK6(X305,X304,X303),X304,X303)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(X303,bool),X304),s(X303,X305)))))) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f17793,plain,(
  ( ! [X304,X306,X303,X305] : (~p(s(bool,X306)) | s(X303,sK6(X305,X304,X303)) = s(X303,sK6(sK6(X305,X304,X303),X304,X303)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(X303,bool),X304),s(X303,X305)))))) ) | ($spl8 | $spl42)),
  inference(subsumption_resolution,[],[f17280,f135])).
fof(f17280,plain,(
  ( ! [X304,X306,X303,X305] : (~p(s(bool,X306)) | s(X303,sK6(X305,X304,X303)) = s(X303,sK6(sK6(X305,X304,X303),X304,X303)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(X303,bool),X304),s(X303,X305))))) | s(bool,'F_0') = s(bool,X306)) ) | ($spl8 | $spl42)),
  inference(superposition,[],[f3798,f12043])).
fof(f3798,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,happ(s(fun(X9,bool),X10),s(X9,X11)))) | s(X9,sK6(X11,X10,X9)) = s(X9,sK6(sK6(X11,X10,X9),X10,X9))) ) | $spl8),
  inference(subsumption_resolution,[],[f3739,f128])).
fof(f3739,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,'F_0')) | ~p(s(bool,happ(s(fun(X9,bool),X10),s(X9,X11)))) | s(X9,sK6(X11,X10,X9)) = s(X9,sK6(sK6(X11,X10,X9),X10,X9))) ) | $spl8),
  inference(superposition,[],[f180,f2667])).
fof(f2667,plain,(
  ( ! [X41,X42,X40] : (s(bool,'F_0') = s(bool,happ(s(fun(X40,bool),X41),s(X40,X42))) | s(X40,X42) = s(X40,sK6(X42,X41,X40))) ) | $spl8),
  inference(subsumption_resolution,[],[f2666,f126])).
fof(f2666,plain,(
  ( ! [X41,X42,X40] : (~p(s(bool,'T_0')) | s(X40,X42) = s(X40,sK6(X42,X41,X40)) | s(bool,'F_0') = s(bool,happ(s(fun(X40,bool),X41),s(X40,X42)))) ) | $spl8),
  inference(forward_demodulation,[],[f2570,f227])).
fof(f2570,plain,(
  ( ! [X41,X42,X40] : (~p(s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)))) | s(X40,X42) = s(X40,sK6(X42,X41,X40)) | s(bool,'F_0') = s(bool,happ(s(fun(X40,bool),X41),s(X40,X42)))) ) | $spl8),
  inference(superposition,[],[f179,f1860])).
fof(f17791,plain,(
  ( ! [X294,X296,X298,X295] : (p(s(bool,happ(s(fun(X294,bool),X295),s(X294,X298)))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(X294,bool),X295),s(X294,X296))))) | s(X294,X296) != s(X294,X298)) ) | $spl106),
  inference(cnf_transformation,[],[f17791_D])).
fof(f17791_D,plain,(
  ( ! [X294,X296,X298,X295] : (p(s(bool,happ(s(fun(X294,bool),X295),s(X294,X298)))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(X294,bool),X295),s(X294,X296))))) | s(X294,X296) != s(X294,X298)) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f17790,plain,(
  ( ! [X294,X296,X298,X295,X297] : (~p(s(bool,X297)) | p(s(bool,happ(s(fun(X294,bool),X295),s(X294,X298)))) | s(X294,X296) != s(X294,X298) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(X294,bool),X295),s(X294,X296)))))) ) | $spl42),
  inference(subsumption_resolution,[],[f17278,f135])).
fof(f17278,plain,(
  ( ! [X294,X296,X298,X295,X297] : (~p(s(bool,X297)) | p(s(bool,happ(s(fun(X294,bool),X295),s(X294,X298)))) | s(X294,X296) != s(X294,X298) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(X294,bool),X295),s(X294,X296))))) | s(bool,'F_0') = s(bool,X297)) ) | $spl42),
  inference(superposition,[],[f178,f12043])).
fof(f178,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | s(X0,X2) != s(X0,X4)) )),
  inference(cnf_transformation,[],[f113])).
fof(f17763,plain,(
  ( ! [X21,X19] : (s(bool,X19) != s(bool,X21) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X19))) | p(s(bool,X21))) ) | $spl104),
  inference(cnf_transformation,[],[f17763_D])).
fof(f17763_D,plain,(
  ( ! [X21,X19] : (s(bool,X19) != s(bool,X21) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X19))) | p(s(bool,X21))) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f17762,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,X20)) | s(bool,X19) != s(bool,X21) | p(s(bool,X21)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X19)))) ) | $spl42),
  inference(subsumption_resolution,[],[f17194,f135])).
fof(f17194,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,X20)) | s(bool,X19) != s(bool,X21) | p(s(bool,X21)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X19))) | s(bool,'F_0') = s(bool,X20)) ) | $spl42),
  inference(superposition,[],[f219,f12043])).
fof(f17755,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X2)))) ) | $spl102),
  inference(cnf_transformation,[],[f17755_D])).
fof(f17755_D,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X2)))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f17754,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X2))) | s(bool,'F_0') = s(bool,X3)) ) | $spl42),
  inference(subsumption_resolution,[],[f17187,f129])).
fof(f17187,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X3) | p(s(bool,X2)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X2))) | s(bool,'F_0') = s(bool,X3)) ) | $spl42),
  inference(superposition,[],[f131,f12043])).
fof(f17668,plain,(
  ( ! [X64,X62] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X64))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X62))) | s(bool,X62) = s(bool,X64)) ) | $spl100),
  inference(cnf_transformation,[],[f17668_D])).
fof(f17668_D,plain,(
  ( ! [X64,X62] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X64))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X62))) | s(bool,X62) = s(bool,X64)) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f17572,plain,(
  ( ! [X74,X75,X73] : (s(bool,X74) = s(bool,X75) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X75))) | s(bool,'F_0') = s(bool,X73) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X74)))) ) | $spl42),
  inference(duplicate_literal_removal,[],[f16896])).
fof(f16896,plain,(
  ( ! [X74,X75,X73] : (s(bool,X74) = s(bool,X75) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X75))) | s(bool,'F_0') = s(bool,X73) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X74))) | s(bool,'F_0') = s(bool,X73)) ) | $spl42),
  inference(superposition,[],[f12043,f12043])).
fof(f17589,plain,(
  ( ! [X78,X76,X77] : (s(bool,'F_0') = s(bool,X77) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X76))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X78))) | s(bool,X76) = s(bool,X78)) ) | $spl42),
  inference(duplicate_literal_removal,[],[f16846])).
fof(f16846,plain,(
  ( ! [X78,X76,X77] : (s(bool,'F_0') = s(bool,X77) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X76))) | s(bool,'F_0') = s(bool,X77) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X78))) | s(bool,X76) = s(bool,X78)) ) | $spl42),
  inference(superposition,[],[f12043,f12043])).
fof(f17601,plain,(
  ( ! [X64,X62,X63] : (s(bool,'F_0') = s(bool,X63) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X64))) | s(bool,X62) = s(bool,X64) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X62)))) ) | $spl42),
  inference(duplicate_literal_removal,[],[f16811])).
fof(f16811,plain,(
  ( ! [X64,X62,X63] : (s(bool,'F_0') = s(bool,X63) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X64))) | s(bool,X62) = s(bool,X64) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X62))) | s(bool,'F_0') = s(bool,X63)) ) | $spl42),
  inference(superposition,[],[f12043,f12043])).
fof(f16594,plain,(
  ( ! [X80,X78,X79] : (s(bool,X80) != s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(X78,bool)),X79),s('type.option.option'(sK0),sK3),s('type.option.option'(X78),'const.option.NONE_0'))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X80)))) ) | $spl98),
  inference(cnf_transformation,[],[f16594_D])).
fof(f16594_D,plain,(
  ( ! [X80,X78,X79] : (s(bool,X80) != s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(X78,bool)),X79),s('type.option.option'(sK0),sK3),s('type.option.option'(X78),'const.option.NONE_0'))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X80)))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f16516,plain,(
  ( ! [X80,X78,X79,X77] : (s(bool,X80) != s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(X78,bool)),X79),s('type.option.option'(sK0),sK3),s('type.option.option'(X78),'const.option.NONE_0'))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X80))) | s(bool,'T_0') = s(bool,X77)) ) | $spl21),
  inference(duplicate_literal_removal,[],[f16325])).
fof(f16325,plain,(
  ( ! [X80,X78,X79,X77] : (s(bool,X80) != s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(X78,bool)),X79),s('type.option.option'(sK0),sK3),s('type.option.option'(X78),'const.option.NONE_0'))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X80))) | s(bool,'T_0') = s(bool,X77) | s(bool,'T_0') = s(bool,X77)) ) | $spl21),
  inference(superposition,[],[f830,f6320])).
fof(f6320,plain,(
  ( ! [X4,X2,X3] : (s(bool,X4) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(X2,bool)),X3),s('type.option.option'(sK0),sK3),s('type.option.option'(X2),'const.option.NONE_0'))) | s(bool,'T_0') = s(bool,X4)) ) | $spl21),
  inference(subsumption_resolution,[],[f6303,f126])).
fof(f6303,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | s(bool,X4) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(X2,bool)),X3),s('type.option.option'(sK0),sK3),s('type.option.option'(X2),'const.option.NONE_0'))) | s(bool,'T_0') = s(bool,X4)) ) | $spl21),
  inference(superposition,[],[f6268,f259])).
fof(f830,plain,(
  ( ! [X23,X22] : (s(bool,X22) != s(bool,X23) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X23))) | s(bool,'T_0') = s(bool,X22)) )),
  inference(superposition,[],[f177,f129])).
fof(f177,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) != s(X0,X2) | s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1))) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2)))) )),
  inference(cnf_transformation,[],[f110])).
fof(f16128,plain,(
  ( ! [X80,X78,X79] : (s(bool,X80) != s(bool,'const.option.OPTREL_3'(s(fun(X78,fun(sK0,bool)),X79),s('type.option.option'(X78),'const.option.NONE_0'),s('type.option.option'(sK0),sK3))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X80)))) ) | $spl96),
  inference(cnf_transformation,[],[f16128_D])).
fof(f16128_D,plain,(
  ( ! [X80,X78,X79] : (s(bool,X80) != s(bool,'const.option.OPTREL_3'(s(fun(X78,fun(sK0,bool)),X79),s('type.option.option'(X78),'const.option.NONE_0'),s('type.option.option'(sK0),sK3))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X80)))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f16052,plain,(
  ( ! [X80,X78,X79,X77] : (s(bool,X80) != s(bool,'const.option.OPTREL_3'(s(fun(X78,fun(sK0,bool)),X79),s('type.option.option'(X78),'const.option.NONE_0'),s('type.option.option'(sK0),sK3))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X80))) | s(bool,'T_0') = s(bool,X77)) ) | $spl21),
  inference(duplicate_literal_removal,[],[f15865])).
fof(f15865,plain,(
  ( ! [X80,X78,X79,X77] : (s(bool,X80) != s(bool,'const.option.OPTREL_3'(s(fun(X78,fun(sK0,bool)),X79),s('type.option.option'(X78),'const.option.NONE_0'),s('type.option.option'(sK0),sK3))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X80))) | s(bool,'T_0') = s(bool,X77) | s(bool,'T_0') = s(bool,X77)) ) | $spl21),
  inference(superposition,[],[f830,f6262])).
fof(f6262,plain,(
  ( ! [X4,X2,X3] : (s(bool,X4) = s(bool,'const.option.OPTREL_3'(s(fun(X2,fun(sK0,bool)),X3),s('type.option.option'(X2),'const.option.NONE_0'),s('type.option.option'(sK0),sK3))) | s(bool,'T_0') = s(bool,X4)) ) | $spl21),
  inference(subsumption_resolution,[],[f6245,f126])).
fof(f6245,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | s(bool,X4) = s(bool,'const.option.OPTREL_3'(s(fun(X2,fun(sK0,bool)),X3),s('type.option.option'(X2),'const.option.NONE_0'),s('type.option.option'(sK0),sK3))) | s(bool,'T_0') = s(bool,X4)) ) | $spl21),
  inference(superposition,[],[f6206,f259])).
fof(f15502,plain,(
  ( ! [X57,X58] : (s(bool,X57) != s(bool,X58) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X57))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X58)))) ) | $spl94),
  inference(cnf_transformation,[],[f15502_D])).
fof(f15502_D,plain,(
  ( ! [X57,X58] : (s(bool,X57) != s(bool,X58) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X57))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X58)))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f15204,plain,(
  ( ! [X57,X58,X56] : (s(bool,X57) != s(bool,X58) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X58))) | s(bool,'T_0') = s(bool,X56) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X57)))) ) | $spl56),
  inference(duplicate_literal_removal,[],[f15077])).
fof(f15077,plain,(
  ( ! [X57,X58,X56] : (s(bool,X57) != s(bool,X58) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X58))) | s(bool,'T_0') = s(bool,X56) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X57))) | s(bool,'T_0') = s(bool,X56)) ) | $spl56),
  inference(superposition,[],[f830,f13334])).
fof(f15501,plain,(
  ( ! [X43,X44] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X43))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X44))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X43))) | s(bool,'F_0') = s(bool,X44)) ) | $spl92),
  inference(cnf_transformation,[],[f15501_D])).
fof(f15501_D,plain,(
  ( ! [X43,X44] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X43))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X44))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X43))) | s(bool,'F_0') = s(bool,X44)) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f15206,plain,(
  ( ! [X43,X44,X42] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X43))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X44))) | s(bool,'F_0') = s(bool,X44) | s(bool,'T_0') = s(bool,X42) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X43)))) ) | $spl56),
  inference(duplicate_literal_removal,[],[f15072])).
fof(f15072,plain,(
  ( ! [X43,X44,X42] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X43))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X44))) | s(bool,'F_0') = s(bool,X44) | s(bool,'T_0') = s(bool,X42) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X43))) | s(bool,'T_0') = s(bool,X42)) ) | $spl56),
  inference(superposition,[],[f773,f13334])).
fof(f15491,plain,(
  ( ! [X277,X275,X276] : (~p(s(bool,happ(s(fun(X275,bool),X276),s(X275,X277)))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(X275,bool),X276),s(X275,sK6(X277,X276,X275))))))) ) | $spl90),
  inference(cnf_transformation,[],[f15491_D])).
fof(f15491_D,plain,(
  ( ! [X277,X275,X276] : (~p(s(bool,happ(s(fun(X275,bool),X276),s(X275,X277)))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(X275,bool),X276),s(X275,sK6(X277,X276,X275))))))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f1001,plain,(
  ( ! [X11] : (p(s(bool,X11))) ) | $spl2),
  inference(cnf_transformation,[],[f1001_D])).
fof(f1001_D,plain,(
  ( ! [X11] : (p(s(bool,X11))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f15490,plain,(
  ( ! [X277,X275,X276,X278] : (p(s(bool,X278)) | ~p(s(bool,happ(s(fun(X275,bool),X276),s(X275,X277)))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(X275,bool),X276),s(X275,sK6(X277,X276,X275))))))) ) | $spl56),
  inference(subsumption_resolution,[],[f15014,f131])).
fof(f15014,plain,(
  ( ! [X277,X275,X276,X278] : (p(s(bool,X278)) | ~p(s(bool,happ(s(fun(X275,bool),X276),s(X275,X277)))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(X275,bool),X276),s(X275,sK6(X277,X276,X275)))))) | s(bool,'T_0') = s(bool,X278)) ) | $spl56),
  inference(superposition,[],[f180,f13334])).
fof(f15366,plain,(
  ( ! [X58,X56] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X58))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X56))) | s(bool,X56) = s(bool,X58)) ) | $spl88),
  inference(cnf_transformation,[],[f15366_D])).
fof(f15366_D,plain,(
  ( ! [X58,X56] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X58))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X56))) | s(bool,X56) = s(bool,X58)) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f15285,plain,(
  ( ! [X70,X68,X69] : (s(bool,X69) = s(bool,X70) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X70))) | s(bool,'T_0') = s(bool,X68) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X69)))) ) | $spl56),
  inference(duplicate_literal_removal,[],[f14662])).
fof(f14662,plain,(
  ( ! [X70,X68,X69] : (s(bool,X69) = s(bool,X70) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X70))) | s(bool,'T_0') = s(bool,X68) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X69))) | s(bool,'T_0') = s(bool,X68)) ) | $spl56),
  inference(superposition,[],[f13334,f13334])).
fof(f15300,plain,(
  ( ! [X64,X62,X63] : (s(bool,'T_0') = s(bool,X63) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X62))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X64))) | s(bool,X62) = s(bool,X64)) ) | $spl56),
  inference(duplicate_literal_removal,[],[f14617])).
fof(f14617,plain,(
  ( ! [X64,X62,X63] : (s(bool,'T_0') = s(bool,X63) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X62))) | s(bool,'T_0') = s(bool,X63) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X64))) | s(bool,X62) = s(bool,X64)) ) | $spl56),
  inference(superposition,[],[f13334,f13334])).
fof(f15316,plain,(
  ( ! [X57,X58,X56] : (s(bool,'T_0') = s(bool,X57) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X58))) | s(bool,X56) = s(bool,X58) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X56)))) ) | $spl56),
  inference(duplicate_literal_removal,[],[f14580])).
fof(f14580,plain,(
  ( ! [X57,X58,X56] : (s(bool,'T_0') = s(bool,X57) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X58))) | s(bool,X56) = s(bool,X58) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X56))) | s(bool,'T_0') = s(bool,X57)) ) | $spl56),
  inference(superposition,[],[f13334,f13334])).
fof(f14055,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f13927,f128])).
fof(f13927,plain,(
  p(s(bool,'F_0')) | $spl52),
  inference(backward_demodulation,[],[f13325,f11012])).
fof(f11012,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.option.option'(X0),'const.option.NONE_0'),s('type.option.option'(X1),'const.option.NONE_0'))))) )),
  inference(equality_resolution,[],[f3625])).
fof(f3625,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.option.option'(X0),'const.option.NONE_0') != s('type.option.option'(X0),X1) | p(s(bool,'const.option.OPTREL_3'(s(fun(X2,fun(X0,bool)),X3),s('type.option.option'(X2),'const.option.NONE_0'),s('type.option.option'(X0),X1))))) )),
  inference(equality_resolution,[],[f192])).
fof(f192,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.option.option'(X0),'const.option.NONE_0') != s('type.option.option'(X0),X3) | s('type.option.option'(X1),'const.option.NONE_0') != s('type.option.option'(X1),X4) | p(s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.option.option'(X0),X3),s('type.option.option'(X1),X4))))) )),
  inference(cnf_transformation,[],[f118])).
fof(f14053,plain,(
  $false | ($spl21 | $spl52)),
  inference(subsumption_resolution,[],[f13884,f128])).
fof(f13884,plain,(
  p(s(bool,'F_0')) | ($spl21 | $spl52)),
  inference(backward_demodulation,[],[f13325,f8552])).
fof(f13657,plain,(
  ( ! [X24,X23] : (s(bool,X23) != s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,'T_0') = s(bool,X24) | s(bool,X23) = s(bool,X24)) ) | $spl86),
  inference(cnf_transformation,[],[f13657_D])).
fof(f13657_D,plain,(
  ( ! [X24,X23] : (s(bool,X23) != s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,'T_0') = s(bool,X24) | s(bool,X23) = s(bool,X24)) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f13647,plain,(
  ( ! [X24,X23,X22] : (s(bool,X23) != s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,'F_0') = s(bool,X22) | s(bool,X23) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) ) | $spl8),
  inference(duplicate_literal_removal,[],[f13603])).
fof(f13603,plain,(
  ( ! [X24,X23,X22] : (s(bool,X23) != s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,'F_0') = s(bool,X22) | s(bool,'F_0') = s(bool,X22) | s(bool,X23) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) ) | $spl8),
  inference(superposition,[],[f2827,f2533])).
fof(f2827,plain,(
  ( ! [X47] : (s(bool,X47) != s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,'F_0') = s(bool,X47)) ) | $spl8),
  inference(subsumption_resolution,[],[f2810,f128])).
fof(f2810,plain,(
  ( ! [X47] : (p(s(bool,'F_0')) | s(bool,X47) != s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,'F_0') = s(bool,X47)) ) | $spl8),
  inference(superposition,[],[f2646,f232])).
fof(f13568,plain,(
  ( ! [X23,X22] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X22))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X23))) | s(bool,'F_0') = s(bool,X22) | s(bool,'F_0') = s(bool,X23)) ) | $spl84),
  inference(cnf_transformation,[],[f13568_D])).
fof(f13568_D,plain,(
  ( ! [X23,X22] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X22))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X23))) | s(bool,'F_0') = s(bool,X22) | s(bool,'F_0') = s(bool,X23)) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f13565,plain,(
  ( ! [X24,X23,X22] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X22))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X23))) | s(bool,'F_0') = s(bool,X23) | s(bool,'F_0') = s(bool,X22) | s(bool,'T_0') = s(bool,X24)) ) | $spl8),
  inference(duplicate_literal_removal,[],[f13413])).
fof(f13413,plain,(
  ( ! [X24,X23,X22] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X22))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X23))) | s(bool,'F_0') = s(bool,X23) | s(bool,'F_0') = s(bool,X22) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) ) | $spl8),
  inference(superposition,[],[f13331,f2533])).
fof(f13331,plain,(
  ( ! [X0] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X0))) | s(bool,'F_0') = s(bool,X0)) ) | $spl8),
  inference(condensation,[],[f13330])).
fof(f13330,plain,(
  ( ! [X17,X16] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X16))) | s(bool,'F_0') = s(bool,X17) | s(bool,'F_0') = s(bool,X16)) ) | $spl8),
  inference(subsumption_resolution,[],[f13283,f1860])).
fof(f13283,plain,(
  ( ! [X17,X16] : (s(bool,X17) != s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X16))) | s(bool,'F_0') = s(bool,X17) | s(bool,'F_0') = s(bool,X16)) ) | $spl8),
  inference(superposition,[],[f2690,f1860])).
fof(f2690,plain,(
  ( ! [X50,X49] : (s(bool,X49) != s(bool,X50) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X50))) | s(bool,'F_0') = s(bool,X49)) ) | $spl8),
  inference(forward_demodulation,[],[f2612,f227])).
fof(f2612,plain,(
  ( ! [X50,X49] : (s(bool,X49) != s(bool,X50) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X50))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))))) | s(bool,'F_0') = s(bool,X49)) ) | $spl8),
  inference(superposition,[],[f177,f1860])).
fof(f13386,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl58),
  inference(cnf_transformation,[],[f13386_D])).
fof(f13386_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f13387,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK7(X0,X1,bool,bool))) ) | $spl60),
  inference(cnf_transformation,[],[f13387_D])).
fof(f13387_D,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK7(X0,X1,bool,bool))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f13401,plain,(
  ( ! [X46,X49] : (s(bool,X46) != s(bool,f(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X49) | s(bool,X46) = s(bool,X49)) ) | $spl82),
  inference(cnf_transformation,[],[f13401_D])).
fof(f13401_D,plain,(
  ( ! [X46,X49] : (s(bool,X46) != s(bool,f(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X49) | s(bool,X46) = s(bool,X49)) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f13378,plain,(
  ( ! [X47,X48,X46,X49] : (s(bool,X46) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X47) = s(fun(bool,bool),X48) | s(bool,'T_0') = s(bool,sK7(X47,X48,bool,bool)) | s(bool,X46) = s(bool,X49) | s(bool,'F_0') = s(bool,X49) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) ) | $spl8),
  inference(superposition,[],[f2829,f2533])).
fof(f2829,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK7(X0,X1,X2,bool))) )),
  inference(superposition,[],[f185,f129])).
fof(f185,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK7(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK7(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f114])).
fof(f114,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK7(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK7(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7])],[f76])).
fof(f76,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f65])).
fof(f65,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.option.OPTREL_MONO',aHL_EXT)).
fof(f13394,plain,(
  ( ! [X16] : (s(bool,X16) = s(bool,g(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X16)) ) | $spl70),
  inference(cnf_transformation,[],[f13394_D])).
fof(f13394_D,plain,(
  ( ! [X16] : (s(bool,X16) = s(bool,g(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X16)) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f13393,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) ) | $spl68),
  inference(cnf_transformation,[],[f13393_D])).
fof(f13393_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f13377,plain,(
  ( ! [X45,X43,X44,X42] : (s(bool,X42) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X43) = s(fun(bool,bool),X44) | s(bool,'T_0') = s(bool,sK7(X43,X44,bool,bool)) | s(bool,X45) = s(bool,g(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X45) | s(bool,'T_0') = s(bool,X42)) ) | $spl8),
  inference(superposition,[],[f2829,f2533])).
fof(f13400,plain,(
  ( ! [X41,X40] : (s(bool,X40) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X41) | s(bool,X40) = s(bool,X41)) ) | $spl80),
  inference(cnf_transformation,[],[f13400_D])).
fof(f13400_D,plain,(
  ( ! [X41,X40] : (s(bool,X40) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X41) | s(bool,X40) = s(bool,X41)) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f13389,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl63),
  inference(cnf_transformation,[],[f13389_D])).
fof(f13389_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl63),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl63])])).
fof(f13376,plain,(
  ( ! [X39,X41,X38,X40] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X38) = s(fun(bool,bool),X39) | s(bool,'T_0') = s(bool,sK7(X38,X39,bool,bool)) | s(bool,X40) = s(bool,g(s(bool,'F_0'))) | s(bool,X40) = s(bool,X41) | s(bool,'T_0') = s(bool,X41)) ) | $spl8),
  inference(superposition,[],[f2829,f2533])).
fof(f13399,plain,(
  ( ! [X37,X36] : (s(bool,X36) = s(bool,X37) | s(bool,X36) = s(bool,g(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X37)) ) | $spl78),
  inference(cnf_transformation,[],[f13399_D])).
fof(f13399_D,plain,(
  ( ! [X37,X36] : (s(bool,X36) = s(bool,X37) | s(bool,X36) = s(bool,g(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X37)) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f13392,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl67),
  inference(cnf_transformation,[],[f13392_D])).
fof(f13392_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl67),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl67])])).
fof(f13375,plain,(
  ( ! [X37,X35,X36,X34] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X34) = s(fun(bool,bool),X35) | s(bool,'T_0') = s(bool,sK7(X34,X35,bool,bool)) | s(bool,X36) = s(bool,X37) | s(bool,'F_0') = s(bool,X37) | s(bool,X36) = s(bool,g(s(bool,'F_0')))) ) | $spl8),
  inference(superposition,[],[f2829,f2533])).
fof(f13396,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | $spl72),
  inference(cnf_transformation,[],[f13396_D])).
fof(f13396_D,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f13398,plain,(
  ( ! [X30,X33] : (s(bool,X30) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X33) | s(bool,X30) = s(bool,X33)) ) | $spl76),
  inference(cnf_transformation,[],[f13398_D])).
fof(f13398_D,plain,(
  ( ! [X30,X33] : (s(bool,X30) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X33) | s(bool,X30) = s(bool,X33)) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f13374,plain,(
  ( ! [X30,X33,X31,X32] : (s(bool,X30) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X31) = s(fun(bool,bool),X32) | s(bool,'T_0') = s(bool,sK7(X31,X32,bool,bool)) | s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | s(bool,X30) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) ) | $spl8),
  inference(superposition,[],[f2829,f2533])).
fof(f13390,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) ) | $spl64),
  inference(cnf_transformation,[],[f13390_D])).
fof(f13390_D,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f13397,plain,(
  ( ! [X17] : (s(bool,X17) != s(bool,f(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X17)) ) | $spl74),
  inference(cnf_transformation,[],[f13397_D])).
fof(f13397_D,plain,(
  ( ! [X17] : (s(bool,X17) != s(bool,f(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X17)) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f13373,plain,(
  ( ! [X28,X26,X29,X27] : (s(bool,X26) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X27) = s(fun(bool,bool),X28) | s(bool,'T_0') = s(bool,sK7(X27,X28,bool,bool)) | s(bool,'F_0') = s(bool,X26) | s(bool,X29) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X29)) ) | $spl8),
  inference(superposition,[],[f2829,f2533])).
fof(f13372,plain,(
  ( ! [X24,X23,X25] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X23) = s(fun(bool,bool),X24) | s(bool,'T_0') = s(bool,sK7(X23,X24,bool,bool)) | s(bool,X25) = s(bool,g(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X25)) ) | $spl8),
  inference(superposition,[],[f2829,f2532])).
fof(f13371,plain,(
  ( ! [X21,X22,X20] : (s(bool,X20) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X21) = s(fun(bool,bool),X22) | s(bool,'T_0') = s(bool,sK7(X21,X22,bool,bool)) | s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X20)) ) | $spl8),
  inference(superposition,[],[f2829,f2532])).
fof(f13370,plain,(
  ( ! [X19,X17,X18] : (s(bool,X17) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X18) = s(fun(bool,bool),X19) | s(bool,'T_0') = s(bool,sK7(X18,X19,bool,bool)) | s(bool,'F_0') = s(bool,X17) | s(bool,'F_0') = s(bool,g(s(bool,'F_0')))) ) | $spl8),
  inference(superposition,[],[f2829,f2532])).
fof(f13369,plain,(
  ( ! [X14,X15,X16] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X14) = s(fun(bool,bool),X15) | s(bool,'T_0') = s(bool,sK7(X14,X15,bool,bool)) | s(bool,X16) = s(bool,g(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X16)) ) | $spl8),
  inference(superposition,[],[f2829,f2532])).
fof(f13368,plain,(
  ( ! [X12,X13,X11] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X11) = s(fun(bool,bool),X12) | s(bool,'T_0') = s(bool,sK7(X11,X12,bool,bool)) | s(bool,X13) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f2829,f259])).
fof(f13367,plain,(
  ( ! [X10,X8,X9] : (s(bool,X8) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X9) = s(fun(bool,bool),X10) | s(bool,'T_0') = s(bool,sK7(X9,X10,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f2829,f259])).
fof(f13366,plain,(
  ( ! [X6,X7,X5] : (s(bool,X5) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X6) = s(fun(bool,bool),X7) | s(bool,'T_0') = s(bool,sK7(X6,X7,bool,bool)) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f2829,f259])).
fof(f13365,plain,(
  ( ! [X4,X2,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,sK7(X2,X3,bool,bool)) | s(bool,X4) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f2829,f259])).
fof(f13364,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK7(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f2829,f129])).
fof(f13334,plain,(
  ( ! [X2,X1] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X1))) | s(bool,'T_0') = s(bool,X2) | s(bool,X1) = s(bool,X2)) ) | $spl56),
  inference(cnf_transformation,[],[f13334_D])).
fof(f13334_D,plain,(
  ( ! [X2,X1] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X1))) | s(bool,'T_0') = s(bool,X2) | s(bool,X1) = s(bool,X2)) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f13333,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X0) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X1))) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl8),
  inference(condensation,[],[f13332])).
fof(f13332,plain,(
  ( ! [X30,X33,X31,X32] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X30))) | s(bool,'F_0') = s(bool,X32) | s(bool,'F_0') = s(bool,X31) | s(bool,X30) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) ) | $spl8),
  inference(subsumption_resolution,[],[f13288,f2532])).
fof(f13288,plain,(
  ( ! [X30,X33,X31,X32] : (s(bool,X31) != s(bool,X32) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X30))) | s(bool,'F_0') = s(bool,X32) | s(bool,'F_0') = s(bool,X31) | s(bool,X30) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) ) | $spl8),
  inference(superposition,[],[f2690,f2533])).
fof(f13326,plain,(
  ( ! [X17] : (s(bool,X17) != s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X17)))) ) | $spl54),
  inference(cnf_transformation,[],[f13326_D])).
fof(f13326_D,plain,(
  ( ! [X17] : (s(bool,X17) != s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X17)))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f13324,plain,(
  ( ! [X17,X16] : (s(bool,X17) != s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X17))) | s(bool,'F_0') = s(bool,X16)) ) | $spl8),
  inference(duplicate_literal_removal,[],[f13244])).
fof(f13244,plain,(
  ( ! [X17,X16] : (s(bool,X17) != s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X17))) | s(bool,'F_0') = s(bool,X16) | s(bool,'F_0') = s(bool,X16)) ) | $spl8),
  inference(superposition,[],[f2690,f1860])).
fof(f12900,plain,(
  ( ! [X80,X78,X79] : (s(bool,X80) != s(bool,'const.option.OPTREL_3'(s(fun(sK1,fun(X78,bool)),X79),s('type.option.option'(sK1),sK2),s('type.option.option'(X78),'const.option.NONE_0'))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X80)))) ) | $spl50),
  inference(cnf_transformation,[],[f12900_D])).
fof(f12900_D,plain,(
  ( ! [X80,X78,X79] : (s(bool,X80) != s(bool,'const.option.OPTREL_3'(s(fun(sK1,fun(X78,bool)),X79),s('type.option.option'(sK1),sK2),s('type.option.option'(X78),'const.option.NONE_0'))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X80)))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f12828,plain,(
  ( ! [X80,X78,X79,X77] : (s(bool,X80) != s(bool,'const.option.OPTREL_3'(s(fun(sK1,fun(X78,bool)),X79),s('type.option.option'(sK1),sK2),s('type.option.option'(X78),'const.option.NONE_0'))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X80))) | s(bool,'T_0') = s(bool,X77)) ) | $spl33),
  inference(duplicate_literal_removal,[],[f12665])).
fof(f12665,plain,(
  ( ! [X80,X78,X79,X77] : (s(bool,X80) != s(bool,'const.option.OPTREL_3'(s(fun(sK1,fun(X78,bool)),X79),s('type.option.option'(sK1),sK2),s('type.option.option'(X78),'const.option.NONE_0'))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X80))) | s(bool,'T_0') = s(bool,X77) | s(bool,'T_0') = s(bool,X77)) ) | $spl33),
  inference(superposition,[],[f830,f8173])).
fof(f8173,plain,(
  ( ! [X4,X2,X3] : (s(bool,X4) = s(bool,'const.option.OPTREL_3'(s(fun(sK1,fun(X2,bool)),X3),s('type.option.option'(sK1),sK2),s('type.option.option'(X2),'const.option.NONE_0'))) | s(bool,'T_0') = s(bool,X4)) ) | $spl33),
  inference(subsumption_resolution,[],[f8152,f126])).
fof(f8152,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | s(bool,X4) = s(bool,'const.option.OPTREL_3'(s(fun(sK1,fun(X2,bool)),X3),s('type.option.option'(sK1),sK2),s('type.option.option'(X2),'const.option.NONE_0'))) | s(bool,'T_0') = s(bool,X4)) ) | $spl33),
  inference(superposition,[],[f8107,f259])).
fof(f12574,plain,(
  ( ! [X78,X76,X77] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(X76,bool),X77),s(X76,X78))))) | s(X76,X78) = s(X76,sK6(X78,X77,X76))) ) | $spl48),
  inference(cnf_transformation,[],[f12574_D])).
fof(f12574_D,plain,(
  ( ! [X78,X76,X77] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(X76,bool),X77),s(X76,X78))))) | s(X76,X78) = s(X76,sK6(X78,X77,X76))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f12573,plain,(
  ( ! [X78,X76,X79,X77] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(X76,bool),X77),s(X76,X78))))) | s(bool,'T_0') = s(bool,X79) | s(X76,X78) = s(X76,sK6(X78,X77,X76))) ) | $spl8),
  inference(subsumption_resolution,[],[f12533,f129])).
fof(f12533,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,'F_0') != s(bool,X79) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(X76,bool),X77),s(X76,X78))))) | s(bool,'T_0') = s(bool,X79) | s(X76,X78) = s(X76,sK6(X78,X77,X76))) ) | $spl8),
  inference(superposition,[],[f830,f2667])).
fof(f12561,plain,(
  ( ! [X57,X56,X55] : (s(bool,X57) != s(bool,'const.option.OPTREL_3'(s(fun(X55,fun(sK1,bool)),X56),s('type.option.option'(X55),'const.option.NONE_0'),s('type.option.option'(sK1),sK2))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X57)))) ) | $spl46),
  inference(cnf_transformation,[],[f12561_D])).
fof(f12561_D,plain,(
  ( ! [X57,X56,X55] : (s(bool,X57) != s(bool,'const.option.OPTREL_3'(s(fun(X55,fun(sK1,bool)),X56),s('type.option.option'(X55),'const.option.NONE_0'),s('type.option.option'(sK1),sK2))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X57)))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f12546,plain,(
  ( ! [X57,X54,X56,X55] : (s(bool,X57) != s(bool,'const.option.OPTREL_3'(s(fun(X55,fun(sK1,bool)),X56),s('type.option.option'(X55),'const.option.NONE_0'),s('type.option.option'(sK1),sK2))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X57))) | s(bool,'T_0') = s(bool,X54)) ) | $spl33),
  inference(duplicate_literal_removal,[],[f12488])).
fof(f12488,plain,(
  ( ! [X57,X54,X56,X55] : (s(bool,X57) != s(bool,'const.option.OPTREL_3'(s(fun(X55,fun(sK1,bool)),X56),s('type.option.option'(X55),'const.option.NONE_0'),s('type.option.option'(sK1),sK2))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X57))) | s(bool,'T_0') = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) ) | $spl33),
  inference(superposition,[],[f830,f7991])).
fof(f7991,plain,(
  ( ! [X4,X2,X3] : (s(bool,X4) = s(bool,'const.option.OPTREL_3'(s(fun(X2,fun(sK1,bool)),X3),s('type.option.option'(X2),'const.option.NONE_0'),s('type.option.option'(sK1),sK2))) | s(bool,'T_0') = s(bool,X4)) ) | $spl33),
  inference(subsumption_resolution,[],[f7971,f126])).
fof(f7971,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | s(bool,X4) = s(bool,'const.option.OPTREL_3'(s(fun(X2,fun(sK1,bool)),X3),s('type.option.option'(X2),'const.option.NONE_0'),s('type.option.option'(sK1),sK2))) | s(bool,'T_0') = s(bool,X4)) ) | $spl33),
  inference(superposition,[],[f7927,f259])).
fof(f12559,plain,(
  ( ! [X15] : (s(bool,X15) != s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X15)))) ) | $spl44),
  inference(cnf_transformation,[],[f12559_D])).
fof(f12559_D,plain,(
  ( ! [X15] : (s(bool,X15) != s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X15)))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f12551,plain,(
  ( ! [X14,X15] : (s(bool,X15) != s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X15))) | s(bool,'T_0') = s(bool,X14)) )),
  inference(duplicate_literal_removal,[],[f12476])).
fof(f12476,plain,(
  ( ! [X14,X15] : (s(bool,X15) != s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X15))) | s(bool,'T_0') = s(bool,X14) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f830,f335])).
fof(f12043,plain,(
  ( ! [X31,X32] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X31))) | s(bool,'F_0') = s(bool,X32) | s(bool,X31) = s(bool,X32)) ) | $spl42),
  inference(cnf_transformation,[],[f12043_D])).
fof(f12043_D,plain,(
  ( ! [X31,X32] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X31))) | s(bool,'F_0') = s(bool,X32) | s(bool,X31) = s(bool,X32)) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f12025,plain,(
  ( ! [X30,X31,X32] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X31))) | s(bool,'T_0') = s(bool,X30) | s(bool,X31) = s(bool,X32) | s(bool,'F_0') = s(bool,X32)) ) | $spl8),
  inference(duplicate_literal_removal,[],[f11924])).
fof(f11924,plain,(
  ( ! [X30,X31,X32] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X31))) | s(bool,'T_0') = s(bool,X30) | s(bool,X31) = s(bool,X32) | s(bool,'F_0') = s(bool,X32) | s(bool,'T_0') = s(bool,X30)) ) | $spl8),
  inference(superposition,[],[f11813,f2533])).
fof(f11813,plain,(
  ( ! [X0] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl8),
  inference(condensation,[],[f11812])).
fof(f11812,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,X2) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X2))) | s(bool,'T_0') = s(bool,X3)) ) | $spl8),
  inference(duplicate_literal_removal,[],[f11811])).
fof(f11811,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,X2) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X2))) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) ) | $spl8),
  inference(condensation,[],[f11810])).
fof(f11810,plain,(
  ( ! [X59,X57,X60,X58] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X58))) | s(bool,'T_0') = s(bool,X58) | s(bool,X57) = s(bool,X59) | s(bool,X59) = s(bool,X60) | s(bool,'T_0') = s(bool,X60)) ) | $spl8),
  inference(trivial_inequality_removal,[],[f11780])).
fof(f11780,plain,(
  ( ! [X59,X57,X60,X58] : (s(bool,X57) != s(bool,X57) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X58))) | s(bool,'T_0') = s(bool,X58) | s(bool,X57) = s(bool,X59) | s(bool,X59) = s(bool,X60) | s(bool,'T_0') = s(bool,X60)) ) | $spl8),
  inference(superposition,[],[f827,f2533])).
fof(f827,plain,(
  ( ! [X14,X15] : (s(bool,'F_0') != s(bool,X15) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X14))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X15))) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f177,f129])).
fof(f11145,plain,(
  ( ! [X35,X34] : (s(bool,X34) != s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,'F_0') = s(bool,X35) | s(bool,X34) = s(bool,X35)) ) | $spl40),
  inference(cnf_transformation,[],[f11145_D])).
fof(f11145_D,plain,(
  ( ! [X35,X34] : (s(bool,X34) != s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,'F_0') = s(bool,X35) | s(bool,X34) = s(bool,X35)) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f11112,plain,(
  ( ! [X35,X33,X34] : (s(bool,X34) != s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,'T_0') = s(bool,X33) | s(bool,X34) = s(bool,X35) | s(bool,'F_0') = s(bool,X35)) ) | $spl8),
  inference(duplicate_literal_removal,[],[f11084])).
fof(f11084,plain,(
  ( ! [X35,X33,X34] : (s(bool,X34) != s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,'T_0') = s(bool,X33) | s(bool,X34) = s(bool,X35) | s(bool,'F_0') = s(bool,X35) | s(bool,'T_0') = s(bool,X33)) ) | $spl8),
  inference(superposition,[],[f1739,f2533])).
fof(f1739,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(condensation,[],[f1717])).
fof(f1717,plain,(
  ( ! [X0,X1] : (s(bool,X1) != s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) )),
  inference(equality_factoring,[],[f982])).
fof(f8981,plain,(
  ( ! [X28,X26,X27] : (s(X26,sK6(X28,X27,X26)) = s(X26,sK6(sK6(X28,X27,X26),X27,X26)) | s(bool,'F_0') = s(bool,happ(s(fun(X26,bool),X27),s(X26,X28)))) ) | $spl38),
  inference(cnf_transformation,[],[f8981_D])).
fof(f8981_D,plain,(
  ( ! [X28,X26,X27] : (s(X26,sK6(X28,X27,X26)) = s(X26,sK6(sK6(X28,X27,X26),X27,X26)) | s(bool,'F_0') = s(bool,happ(s(fun(X26,bool),X27),s(X26,X28)))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f8982,plain,(
  ( ! [X30,X33,X31,X32] : (~p(s(bool,X33)) | s(X30,sK6(X32,X31,X30)) = s(X30,sK6(sK6(X32,X31,X30),X31,X30)) | s(bool,'F_0') = s(bool,happ(s(fun(X30,bool),X31),s(X30,X32)))) ) | $spl8),
  inference(subsumption_resolution,[],[f8971,f135])).
fof(f8971,plain,(
  ( ! [X30,X33,X31,X32] : (~p(s(bool,X33)) | s(X30,sK6(X32,X31,X30)) = s(X30,sK6(sK6(X32,X31,X30),X31,X30)) | s(bool,'F_0') = s(bool,happ(s(fun(X30,bool),X31),s(X30,X32))) | s(bool,'F_0') = s(bool,X33)) ) | $spl8),
  inference(superposition,[],[f3798,f2532])).
fof(f8980,plain,(
  ( ! [X28,X26,X29,X27] : (~p(s(bool,X29)) | s(X26,sK6(X28,X27,X26)) = s(X26,sK6(sK6(X28,X27,X26),X27,X26)) | s(bool,'F_0') = s(bool,happ(s(fun(X26,bool),X27),s(X26,X28)))) ) | $spl8),
  inference(subsumption_resolution,[],[f8970,f135])).
fof(f8970,plain,(
  ( ! [X28,X26,X29,X27] : (~p(s(bool,X29)) | s(X26,sK6(X28,X27,X26)) = s(X26,sK6(sK6(X28,X27,X26),X27,X26)) | s(bool,'F_0') = s(bool,X29) | s(bool,'F_0') = s(bool,happ(s(fun(X26,bool),X27),s(X26,X28)))) ) | $spl8),
  inference(superposition,[],[f3798,f2532])).
fof(f8774,plain,(
  ( ! [X47,X45,X43,X46,X44] : (s('type.option.option'(X44),'const.option.NONE_0') = s('type.option.option'(X44),X47) | s(bool,'F_0') = s(bool,'const.option.OPTREL_3'(s(fun(X43,fun(X44,bool)),X45),s('type.option.option'(X43),X46),s('type.option.option'(X44),X47))) | p(s(bool,happ(s(fun(X44,bool),happ(s(fun(X43,fun(X44,bool)),X45),s(X43,sK8(X47,X46,X45,X44,X43)))),s(X44,sK9(X47,X46,X45,X44,X43)))))) ) | $spl36),
  inference(cnf_transformation,[],[f8774_D])).
fof(f8774_D,plain,(
  ( ! [X47,X45,X43,X46,X44] : (s('type.option.option'(X44),'const.option.NONE_0') = s('type.option.option'(X44),X47) | s(bool,'F_0') = s(bool,'const.option.OPTREL_3'(s(fun(X43,fun(X44,bool)),X45),s('type.option.option'(X43),X46),s('type.option.option'(X44),X47))) | p(s(bool,happ(s(fun(X44,bool),happ(s(fun(X43,fun(X44,bool)),X45),s(X43,sK8(X47,X46,X45,X44,X43)))),s(X44,sK9(X47,X46,X45,X44,X43)))))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f8775,plain,(
  ( ! [X54,X52,X50,X53,X51,X49] : (~p(s(bool,X54)) | s('type.option.option'(X50),'const.option.NONE_0') = s('type.option.option'(X50),X53) | p(s(bool,happ(s(fun(X50,bool),happ(s(fun(X49,fun(X50,bool)),X51),s(X49,sK8(X53,X52,X51,X50,X49)))),s(X50,sK9(X53,X52,X51,X50,X49))))) | s(bool,'F_0') = s(bool,'const.option.OPTREL_3'(s(fun(X49,fun(X50,bool)),X51),s('type.option.option'(X49),X52),s('type.option.option'(X50),X53)))) ) | $spl8),
  inference(subsumption_resolution,[],[f8762,f135])).
fof(f8762,plain,(
  ( ! [X54,X52,X50,X53,X51,X49] : (~p(s(bool,X54)) | s('type.option.option'(X50),'const.option.NONE_0') = s('type.option.option'(X50),X53) | p(s(bool,happ(s(fun(X50,bool),happ(s(fun(X49,fun(X50,bool)),X51),s(X49,sK8(X53,X52,X51,X50,X49)))),s(X50,sK9(X53,X52,X51,X50,X49))))) | s(bool,'F_0') = s(bool,'const.option.OPTREL_3'(s(fun(X49,fun(X50,bool)),X51),s('type.option.option'(X49),X52),s('type.option.option'(X50),X53))) | s(bool,'F_0') = s(bool,X54)) ) | $spl8),
  inference(superposition,[],[f191,f2532])).
fof(f8773,plain,(
  ( ! [X47,X45,X43,X48,X46,X44] : (~p(s(bool,X48)) | s('type.option.option'(X44),'const.option.NONE_0') = s('type.option.option'(X44),X47) | p(s(bool,happ(s(fun(X44,bool),happ(s(fun(X43,fun(X44,bool)),X45),s(X43,sK8(X47,X46,X45,X44,X43)))),s(X44,sK9(X47,X46,X45,X44,X43))))) | s(bool,'F_0') = s(bool,'const.option.OPTREL_3'(s(fun(X43,fun(X44,bool)),X45),s('type.option.option'(X43),X46),s('type.option.option'(X44),X47)))) ) | $spl8),
  inference(subsumption_resolution,[],[f8761,f135])).
fof(f8761,plain,(
  ( ! [X47,X45,X43,X48,X46,X44] : (~p(s(bool,X48)) | s('type.option.option'(X44),'const.option.NONE_0') = s('type.option.option'(X44),X47) | p(s(bool,happ(s(fun(X44,bool),happ(s(fun(X43,fun(X44,bool)),X45),s(X43,sK8(X47,X46,X45,X44,X43)))),s(X44,sK9(X47,X46,X45,X44,X43))))) | s(bool,'F_0') = s(bool,X48) | s(bool,'F_0') = s(bool,'const.option.OPTREL_3'(s(fun(X43,fun(X44,bool)),X45),s('type.option.option'(X43),X46),s('type.option.option'(X44),X47)))) ) | $spl8),
  inference(superposition,[],[f191,f2532])).
fof(f8555,plain,(
  ( ! [X47,X45,X43,X46,X44] : (s('type.option.option'(X43),'const.option.NONE_0') = s('type.option.option'(X43),X46) | s(bool,'F_0') = s(bool,'const.option.OPTREL_3'(s(fun(X43,fun(X44,bool)),X45),s('type.option.option'(X43),X46),s('type.option.option'(X44),X47))) | p(s(bool,happ(s(fun(X44,bool),happ(s(fun(X43,fun(X44,bool)),X45),s(X43,sK8(X47,X46,X45,X44,X43)))),s(X44,sK9(X47,X46,X45,X44,X43)))))) ) | $spl34),
  inference(cnf_transformation,[],[f8555_D])).
fof(f8555_D,plain,(
  ( ! [X47,X45,X43,X46,X44] : (s('type.option.option'(X43),'const.option.NONE_0') = s('type.option.option'(X43),X46) | s(bool,'F_0') = s(bool,'const.option.OPTREL_3'(s(fun(X43,fun(X44,bool)),X45),s('type.option.option'(X43),X46),s('type.option.option'(X44),X47))) | p(s(bool,happ(s(fun(X44,bool),happ(s(fun(X43,fun(X44,bool)),X45),s(X43,sK8(X47,X46,X45,X44,X43)))),s(X44,sK9(X47,X46,X45,X44,X43)))))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f8556,plain,(
  ( ! [X54,X52,X50,X53,X51,X49] : (~p(s(bool,X54)) | s('type.option.option'(X49),'const.option.NONE_0') = s('type.option.option'(X49),X52) | p(s(bool,happ(s(fun(X50,bool),happ(s(fun(X49,fun(X50,bool)),X51),s(X49,sK8(X53,X52,X51,X50,X49)))),s(X50,sK9(X53,X52,X51,X50,X49))))) | s(bool,'F_0') = s(bool,'const.option.OPTREL_3'(s(fun(X49,fun(X50,bool)),X51),s('type.option.option'(X49),X52),s('type.option.option'(X50),X53)))) ) | $spl8),
  inference(subsumption_resolution,[],[f8541,f135])).
fof(f8541,plain,(
  ( ! [X54,X52,X50,X53,X51,X49] : (~p(s(bool,X54)) | s('type.option.option'(X49),'const.option.NONE_0') = s('type.option.option'(X49),X52) | p(s(bool,happ(s(fun(X50,bool),happ(s(fun(X49,fun(X50,bool)),X51),s(X49,sK8(X53,X52,X51,X50,X49)))),s(X50,sK9(X53,X52,X51,X50,X49))))) | s(bool,'F_0') = s(bool,'const.option.OPTREL_3'(s(fun(X49,fun(X50,bool)),X51),s('type.option.option'(X49),X52),s('type.option.option'(X50),X53))) | s(bool,'F_0') = s(bool,X54)) ) | $spl8),
  inference(superposition,[],[f190,f2532])).
fof(f8554,plain,(
  ( ! [X47,X45,X43,X48,X46,X44] : (~p(s(bool,X48)) | s('type.option.option'(X43),'const.option.NONE_0') = s('type.option.option'(X43),X46) | p(s(bool,happ(s(fun(X44,bool),happ(s(fun(X43,fun(X44,bool)),X45),s(X43,sK8(X47,X46,X45,X44,X43)))),s(X44,sK9(X47,X46,X45,X44,X43))))) | s(bool,'F_0') = s(bool,'const.option.OPTREL_3'(s(fun(X43,fun(X44,bool)),X45),s('type.option.option'(X43),X46),s('type.option.option'(X44),X47)))) ) | $spl8),
  inference(subsumption_resolution,[],[f8540,f135])).
fof(f8540,plain,(
  ( ! [X47,X45,X43,X48,X46,X44] : (~p(s(bool,X48)) | s('type.option.option'(X43),'const.option.NONE_0') = s('type.option.option'(X43),X46) | p(s(bool,happ(s(fun(X44,bool),happ(s(fun(X43,fun(X44,bool)),X45),s(X43,sK8(X47,X46,X45,X44,X43)))),s(X44,sK9(X47,X46,X45,X44,X43))))) | s(bool,'F_0') = s(bool,X48) | s(bool,'F_0') = s(bool,'const.option.OPTREL_3'(s(fun(X43,fun(X44,bool)),X45),s('type.option.option'(X43),X46),s('type.option.option'(X44),X47)))) ) | $spl8),
  inference(superposition,[],[f190,f2532])).
fof(f6894,plain,(
  $false | ($spl21 | $spl32)),
  inference(subsumption_resolution,[],[f6591,f173])).
fof(f6591,plain,(
  s('type.option.option'(sK1),'const.option.NONE_0') = s('type.option.option'(sK1),'const.option.SOME_1'(s(sK1,sK9(sK2,sK3,sK5,sK1,sK0)))) | ($spl21 | $spl32)),
  inference(backward_demodulation,[],[f6509,f6511])).
fof(f6888,plain,(
  $false | ($spl8 | $spl21 | $spl32)),
  inference(global_subsumption,[],[f6887,f125,f124,f126,f129,f132,f131,f173,f177,f176,f184,f181,f185,f219,f223,f259,f296,f428,f498,f714,f715,f726,f764,f765,f767,f806,f822,f824,f828,f884,f999,f1002,f1833,f1826,f1861,f1819,f1862,f1810,f1910,f1763,f1711,f128,f135,f136,f773,f790,f792,f795,f805,f827,f830,f954,f940,f1515,f2829,f2830,f2831,f3830,f3832,f3833,f3975,f4201,f4187,f3961,f3625,f192,f189,f188,f187,f186,f123,f716,f1330,f936,f1333,f1312,f918,f1372,f1362,f1360,f1356,f1354,f1352,f1341,f1335,f1320,f1318,f1314,f998,f988,f986,f926,f924,f920,f191,f190,f193,f1376,f1331,f1004,f180,f179,f178,f2842,f2841,f2840,f3836,f2849,f2848,f2838,f2837,f2151,f1995,f1994,f1993,f1981,f1951,f1950,f1949,f1751,f1739,f1705,f1703,f1655,f1637,f1636,f1626,f1572,f1570,f1523,f1489,f1438,f1437,f1344,f1323,f1322,f1321,f982,f956,f939,f938,f937,f826,f776,f775,f774,f335,f232,f227,f224,f3826,f6181,f6206,f6233,f6262,f6268,f6320,f6416,f6505,f6511,f6509,f6532,f6533,f6534,f6535,f6536,f6537,f6538,f6539,f6540,f6541,f6542,f6543,f6544,f6625,f6828,f6770,f6857,f6877,f6879,f6881,f6883])).
fof(f6883,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.option.option'(X1),'const.option.NONE_0') = s('type.option.option'(X1),X4) | s('type.option.option'(X1),X4) = s('type.option.option'(X1),'const.option.SOME_1'(s(X1,sK9(X4,X3,X2,X1,X0))))) ) | ($spl21 | $spl32)),
  inference(subsumption_resolution,[],[f6804,f126])).
fof(f6804,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | s('type.option.option'(X1),'const.option.NONE_0') = s('type.option.option'(X1),X4) | s('type.option.option'(X1),X4) = s('type.option.option'(X1),'const.option.SOME_1'(s(X1,sK9(X4,X3,X2,X1,X0))))) ) | ($spl21 | $spl32)),
  inference(backward_demodulation,[],[f6625,f189])).
fof(f6881,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X3) | s('type.option.option'(X1),X4) = s('type.option.option'(X1),'const.option.SOME_1'(s(X1,sK9(X4,X3,X2,X1,X0))))) ) | ($spl21 | $spl32)),
  inference(subsumption_resolution,[],[f6803,f126])).
fof(f6803,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X3) | s('type.option.option'(X1),X4) = s('type.option.option'(X1),'const.option.SOME_1'(s(X1,sK9(X4,X3,X2,X1,X0))))) ) | ($spl21 | $spl32)),
  inference(backward_demodulation,[],[f6625,f188])).
fof(f6879,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.option.option'(X1),'const.option.NONE_0') = s('type.option.option'(X1),X4) | s('type.option.option'(X0),X3) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,sK8(X4,X3,X2,X1,X0))))) ) | ($spl21 | $spl32)),
  inference(subsumption_resolution,[],[f6802,f126])).
fof(f6802,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | s('type.option.option'(X1),'const.option.NONE_0') = s('type.option.option'(X1),X4) | s('type.option.option'(X0),X3) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,sK8(X4,X3,X2,X1,X0))))) ) | ($spl21 | $spl32)),
  inference(backward_demodulation,[],[f6625,f187])).
fof(f6877,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X3) | s('type.option.option'(X0),X3) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,sK8(X4,X3,X2,X1,X0))))) ) | ($spl21 | $spl32)),
  inference(subsumption_resolution,[],[f6801,f126])).
fof(f6801,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X3) | s('type.option.option'(X0),X3) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,sK8(X4,X3,X2,X1,X0))))) ) | ($spl21 | $spl32)),
  inference(backward_demodulation,[],[f6625,f186])).
fof(f6857,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,sK6(X2,X1,X0))) ) | ($spl21 | $spl32)),
  inference(subsumption_resolution,[],[f6774,f126])).
fof(f6774,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s(X0,X2) = s(X0,sK6(X2,X1,X0))) ) | ($spl21 | $spl32)),
  inference(backward_demodulation,[],[f6625,f179])).
fof(f6770,plain,(
  ( ! [X27] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X27))) | s(bool,X27) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),'const.option.NONE_0')))) ) | ($spl21 | $spl32)),
  inference(backward_demodulation,[],[f6625,f6534])).
fof(f6828,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl21 | $spl32)),
  inference(subsumption_resolution,[],[f6627,f6625])).
fof(f6627,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | ($spl21 | $spl32)),
  inference(backward_demodulation,[],[f6625,f135])).
fof(f6625,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl21 | $spl32)),
  inference(subsumption_resolution,[],[f6557,f6320])).
fof(f6557,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),'const.option.NONE_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl32),
  inference(backward_demodulation,[],[f6509,f1739])).
fof(f6544,plain,(
  ( ! [X26,X27] : (s(bool,'T_0') = s(bool,sK6('const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),'const.option.NONE_0')),X27,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X27),s(bool,X26)))) | p(s(bool,happ(s(fun(bool,bool),X27),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X26)) ) | $spl32),
  inference(backward_demodulation,[],[f6509,f1344])).
fof(f6543,plain,(
  ( ! [X26,X27] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X27),s(bool,sK6('const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),'const.option.NONE_0')),X27,bool)))) | ~p(s(bool,happ(s(fun(bool,bool),X27),s(bool,X26)))) | s(bool,'T_0') = s(bool,X26)) ) | $spl32),
  inference(backward_demodulation,[],[f6509,f1323])).
fof(f6542,plain,(
  ( ! [X25] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X25),s(bool,sK6('const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),'const.option.NONE_0')),X25,bool)))) | ~p(s(bool,happ(s(fun(bool,bool),X25),s(bool,'F_0'))))) ) | $spl32),
  inference(backward_demodulation,[],[f6509,f1322])).
fof(f6541,plain,(
  ( ! [X24] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X24),s(bool,sK6('const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),'const.option.NONE_0')),X24,bool)))) | ~p(s(bool,happ(s(fun(bool,bool),X24),s(bool,'T_0'))))) ) | $spl32),
  inference(backward_demodulation,[],[f6509,f1321])).
fof(f6540,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),'const.option.NONE_0'))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | $spl32),
  inference(backward_demodulation,[],[f6509,f982])).
fof(f6539,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),'const.option.NONE_0'))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X0))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X1))) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl32),
  inference(backward_demodulation,[],[f6509,f956])).
fof(f6538,plain,(
  ( ! [X27] : (s(bool,'F_0') = s(bool,sK6('const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),'const.option.NONE_0')),X27,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X27),s(bool,'F_0'))))) ) | $spl32),
  inference(backward_demodulation,[],[f6509,f939])).
fof(f6537,plain,(
  ( ! [X26,X25] : (s(bool,'F_0') = s(bool,sK6('const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),'const.option.NONE_0')),X26,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X26),s(bool,X25)))) | s(bool,'T_0') = s(bool,X25)) ) | $spl32),
  inference(backward_demodulation,[],[f6509,f938])).
fof(f6536,plain,(
  ( ! [X24] : (s(bool,'T_0') = s(bool,sK6('const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),'const.option.NONE_0')),X24,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X24),s(bool,'T_0'))))) ) | $spl32),
  inference(backward_demodulation,[],[f6509,f937])).
fof(f6535,plain,(
  ( ! [X12,X13] : (s(bool,X13) != s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),'const.option.NONE_0'))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X12))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X13))) | s(bool,'T_0') = s(bool,X12)) ) | $spl32),
  inference(backward_demodulation,[],[f6509,f826])).
fof(f6534,plain,(
  ( ! [X27] : (s(bool,X27) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),'const.option.NONE_0'))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X27)))) ) | $spl32),
  inference(backward_demodulation,[],[f6509,f776])).
fof(f6533,plain,(
  ( ! [X26,X25] : (s(bool,X26) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),'const.option.NONE_0'))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X25))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X26))) | s(bool,'T_0') = s(bool,X25)) ) | $spl32),
  inference(backward_demodulation,[],[f6509,f775])).
fof(f6532,plain,(
  ( ! [X24] : (s(bool,X24) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),'const.option.NONE_0'))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X24)))) ) | $spl32),
  inference(backward_demodulation,[],[f6509,f774])).
fof(f6416,plain,(
  ( ! [X6,X10,X8,X7,X5,X9] : (s('type.option.option'(X6),'const.option.NONE_0') = s('type.option.option'(X6),X9) | s('type.option.option'(X6),X9) = s('type.option.option'(X6),'const.option.SOME_1'(s(X6,sK9(X9,X8,X7,X6,X5)))) | s(bool,X10) = s(bool,'const.option.OPTREL_3'(s(fun(X5,fun(X6,bool)),X7),s('type.option.option'(X5),X8),s('type.option.option'(X6),X9))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(subsumption_resolution,[],[f6400,f126])).
fof(f6400,plain,(
  ( ! [X6,X10,X8,X7,X5,X9] : (~p(s(bool,'T_0')) | s('type.option.option'(X6),'const.option.NONE_0') = s('type.option.option'(X6),X9) | s('type.option.option'(X6),X9) = s('type.option.option'(X6),'const.option.SOME_1'(s(X6,sK9(X9,X8,X7,X6,X5)))) | s(bool,X10) = s(bool,'const.option.OPTREL_3'(s(fun(X5,fun(X6,bool)),X7),s('type.option.option'(X5),X8),s('type.option.option'(X6),X9))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f189,f259])).
fof(f6233,plain,(
  s('type.option.option'('type.option.option'('type.option.option'(sK0))),'const.option.SOME_1'(s('type.option.option'('type.option.option'(sK0)),'const.option.SOME_1'(s('type.option.option'(sK0),'const.option.NONE_0'))))) != s('type.option.option'('type.option.option'('type.option.option'(sK0))),'const.option.SOME_1'(s('type.option.option'('type.option.option'(sK0)),'const.option.SOME_1'(s('type.option.option'(sK0),sK3))))) | $spl21),
  inference(unit_resulting_resolution,[],[f6181,f176])).
fof(f6181,plain,(
  s('type.option.option'('type.option.option'(sK0)),'const.option.SOME_1'(s('type.option.option'(sK0),'const.option.NONE_0'))) != s('type.option.option'('type.option.option'(sK0)),'const.option.SOME_1'(s('type.option.option'(sK0),sK3))) | $spl21),
  inference(unit_resulting_resolution,[],[f3826,f176])).
fof(f224,plain,(
  s(bool,'T_0') != s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)))),
  inference(unit_resulting_resolution,[],[f125,f131])).
fof(f775,plain,(
  ( ! [X26,X25] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X25))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X26))) | s(bool,X26) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f176,f335])).
fof(f776,plain,(
  ( ! [X27] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X27))) | s(bool,X27) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)))) )),
  inference(superposition,[],[f176,f232])).
fof(f826,plain,(
  ( ! [X12,X13] : (s(bool,X13) != s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X12))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X13))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f177,f335])).
fof(f937,plain,(
  ( ! [X24] : (s(bool,'T_0') = s(bool,sK6('const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)),X24,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X24),s(bool,'T_0'))))) )),
  inference(forward_demodulation,[],[f927,f227])).
fof(f927,plain,(
  ( ! [X24] : (~p(s(bool,happ(s(fun(bool,bool),X24),s(bool,'T_0')))) | s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) = s(bool,sK6('const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)),X24,bool))) )),
  inference(superposition,[],[f179,f227])).
fof(f938,plain,(
  ( ! [X26,X25] : (s(bool,'F_0') = s(bool,sK6('const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)),X26,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X26),s(bool,X25)))) | s(bool,'T_0') = s(bool,X25)) )),
  inference(forward_demodulation,[],[f928,f232])).
fof(f928,plain,(
  ( ! [X26,X25] : (~p(s(bool,happ(s(fun(bool,bool),X26),s(bool,X25)))) | s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) = s(bool,sK6('const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)),X26,bool)) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f179,f335])).
fof(f939,plain,(
  ( ! [X27] : (s(bool,'F_0') = s(bool,sK6('const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)),X27,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X27),s(bool,'F_0'))))) )),
  inference(forward_demodulation,[],[f929,f232])).
fof(f929,plain,(
  ( ! [X27] : (~p(s(bool,happ(s(fun(bool,bool),X27),s(bool,'F_0')))) | s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) = s(bool,sK6('const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)),X27,bool))) )),
  inference(superposition,[],[f179,f232])).
fof(f956,plain,(
  ( ! [X2,X0,X1] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X0))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X1))) | s(bool,X1) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f774,f259])).
fof(f1321,plain,(
  ( ! [X24] : (~p(s(bool,happ(s(fun(bool,bool),X24),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X24),s(bool,sK6('const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)),X24,bool))))) )),
  inference(superposition,[],[f1004,f227])).
fof(f1322,plain,(
  ( ! [X25] : (~p(s(bool,happ(s(fun(bool,bool),X25),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X25),s(bool,sK6('const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)),X25,bool))))) )),
  inference(superposition,[],[f1004,f232])).
fof(f1323,plain,(
  ( ! [X26,X27] : (~p(s(bool,happ(s(fun(bool,bool),X27),s(bool,X26)))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X27),s(bool,sK6('const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)),X27,bool)))) | s(bool,'T_0') = s(bool,X26)) )),
  inference(superposition,[],[f1004,f335])).
fof(f1344,plain,(
  ( ! [X26,X27] : (~p(s(bool,happ(s(fun(bool,bool),X27),s(bool,X26)))) | p(s(bool,happ(s(fun(bool,bool),X27),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK6('const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)),X27,bool)) | s(bool,'T_0') = s(bool,X26)) )),
  inference(superposition,[],[f998,f335])).
fof(f1437,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X2) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,X1) = s(bool,X2) | s(bool,X3) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,X0) = s(bool,X3)) )),
  inference(superposition,[],[f982,f982])).
fof(f1438,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X4) = s(bool,X5) | s(bool,X6) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,X5) = s(bool,X6) | s(bool,X4) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f982,f259])).
fof(f1489,plain,(
  ( ! [X88,X87,X89] : (s(bool,X88) != s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X87))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X88))) | s(bool,X87) = s(bool,X89) | s(bool,'T_0') = s(bool,X89)) )),
  inference(superposition,[],[f177,f982])).
fof(f1523,plain,(
  ( ! [X76,X74,X75] : (~p(s(bool,happ(s(fun(bool,bool),X75),s(bool,X74)))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X75),s(bool,sK6('const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)),X75,bool)))) | s(bool,X74) = s(bool,X76) | s(bool,'T_0') = s(bool,X76)) )),
  inference(superposition,[],[f1004,f982])).
fof(f1570,plain,(
  ( ! [X118,X119,X117] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X118))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X119))) | s(bool,X117) = s(bool,X119) | s(bool,X117) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,'T_0') = s(bool,X118)) )),
  inference(superposition,[],[f176,f982])).
fof(f1572,plain,(
  ( ! [X125,X123,X124] : (s(bool,X124) != s(bool,X125) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X123))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X125))) | s(bool,X123) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,'T_0') = s(bool,X124)) )),
  inference(superposition,[],[f177,f982])).
fof(f1626,plain,(
  ( ! [X125,X123,X124] : (s(bool,X124) != s(bool,X125) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X123))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X125))) | s(bool,X124) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,'T_0') = s(bool,X123)) )),
  inference(superposition,[],[f177,f982])).
fof(f1636,plain,(
  ( ! [X17,X15,X16] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X15))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X16))) | s(bool,X16) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,X17) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,X15) = s(bool,X17)) )),
  inference(superposition,[],[f774,f982])).
fof(f1637,plain,(
  ( ! [X21,X22,X20] : (s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X20))))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X21))))) | s(bool,'T_0') = s(bool,X21) | s(bool,X22) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,X20) = s(bool,X22)) )),
  inference(superposition,[],[f806,f982])).
fof(f1655,plain,(
  ( ! [X70,X71,X69] : (s(bool,X69) != s(bool,X70) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X70))) | s(bool,X71) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,X69) = s(bool,X71)) )),
  inference(superposition,[],[f177,f982])).
fof(f1703,plain,(
  ( ! [X109,X110,X108] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X109))) | s(bool,X108) = s(bool,X109) | s(bool,X110) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,X108) = s(bool,X110)) )),
  inference(superposition,[],[f176,f982])).
fof(f1705,plain,(
  ( ! [X116,X114,X115] : (s(bool,'T_0') != s(bool,X115) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X114))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X115))) | s(bool,X116) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,X114) = s(bool,X116)) )),
  inference(superposition,[],[f177,f982])).
fof(f1751,plain,(
  ( ! [X39,X38] : (s(bool,X38) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,X39) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,X38) = s(bool,X39)) )),
  inference(trivial_inequality_removal,[],[f1674])).
fof(f1674,plain,(
  ( ! [X39,X38] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | s(bool,X38) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,X39) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,X38) = s(bool,X39)) )),
  inference(superposition,[],[f774,f982])).
fof(f1949,plain,(
  ( ! [X72] : (s(bool,'T_0') != s(bool,X72) | s(bool,X72) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)))) )),
  inference(subsumption_resolution,[],[f1606,f223])).
fof(f1606,plain,(
  ( ! [X72] : (s(bool,'T_0') != s(bool,X72) | s(bool,X72) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,'T_0') = s(bool,'F_0')) )),
  inference(superposition,[],[f223,f982])).
fof(f1950,plain,(
  ( ! [X74,X75,X73] : (s(bool,X73) != s(bool,X74) | s(bool,X74) = s(bool,X75) | s(bool,'T_0') = s(bool,X75) | s(bool,X73) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)))) )),
  inference(subsumption_resolution,[],[f1607,f223])).
fof(f1607,plain,(
  ( ! [X74,X75,X73] : (s(bool,X73) != s(bool,X74) | s(bool,X74) = s(bool,X75) | s(bool,'T_0') = s(bool,X75) | s(bool,X73) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,'T_0') = s(bool,'F_0')) )),
  inference(superposition,[],[f714,f982])).
fof(f1951,plain,(
  ( ! [X80] : (s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X80))))) | s(bool,X80) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)))) )),
  inference(subsumption_resolution,[],[f1610,f223])).
fof(f1610,plain,(
  ( ! [X80] : (s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X80))))) | s(bool,X80) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,'T_0') = s(bool,'F_0')) )),
  inference(superposition,[],[f795,f982])).
fof(f1981,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X7) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,X6) = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f1663,f126])).
fof(f1663,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X6) | s(bool,X7) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,X6) = s(bool,X7)) )),
  inference(superposition,[],[f135,f982])).
fof(f1993,plain,(
  ( ! [X87,X85,X86,X84] : (s(bool,'T_0') = s(bool,happ(s(fun(X84,bool),X85),s(X84,sK6(X86,X85,X84)))) | s(bool,X87) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,X87) = s(bool,happ(s(fun(X84,bool),X85),s(X84,X86)))) )),
  inference(subsumption_resolution,[],[f1694,f126])).
fof(f1694,plain,(
  ( ! [X87,X85,X86,X84] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,happ(s(fun(X84,bool),X85),s(X84,sK6(X86,X85,X84)))) | s(bool,X87) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,X87) = s(bool,happ(s(fun(X84,bool),X85),s(X84,X86)))) )),
  inference(superposition,[],[f1004,f982])).
fof(f1994,plain,(
  ( ! [X90,X88,X91,X89] : (s(X88,X90) = s(X88,sK6(X90,X89,X88)) | s(bool,X91) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,X91) = s(bool,happ(s(fun(X88,bool),X89),s(X88,X90)))) )),
  inference(subsumption_resolution,[],[f1695,f126])).
fof(f1695,plain,(
  ( ! [X90,X88,X91,X89] : (~p(s(bool,'T_0')) | s(X88,X90) = s(X88,sK6(X90,X89,X88)) | s(bool,X91) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,X91) = s(bool,happ(s(fun(X88,bool),X89),s(X88,X90)))) )),
  inference(superposition,[],[f179,f982])).
fof(f1995,plain,(
  ( ! [X97,X98,X96] : (s(bool,X97) = s(bool,sK6(X97,X96,bool)) | s(bool,'T_0') = s(bool,X97) | s(bool,X98) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,X98) = s(bool,happ(s(fun(bool,bool),X96),s(bool,'F_0')))) )),
  inference(subsumption_resolution,[],[f1697,f126])).
fof(f1697,plain,(
  ( ! [X97,X98,X96] : (~p(s(bool,'T_0')) | s(bool,X97) = s(bool,sK6(X97,X96,bool)) | s(bool,'T_0') = s(bool,X97) | s(bool,X98) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,X98) = s(bool,happ(s(fun(bool,bool),X96),s(bool,'F_0')))) )),
  inference(superposition,[],[f936,f982])).
fof(f2151,plain,(
  ( ! [X52,X50,X51] : (s(bool,'T_0') = s(bool,X52) | s(bool,'T_0') = s(bool,sK6('const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)),X51,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X51),s(bool,X50)))) | s(bool,X50) = s(bool,X52)) )),
  inference(global_subsumption,[],[f125,f124,f126,f129,f132,f131,f173,f177,f176,f184,f181,f185,f219,f224,f223,f227,f259,f296,f335,f428,f498,f714,f715,f726,f764,f765,f767,f774,f775,f806,f822,f824,f826,f828,f884,f956,f982,f999,f1002,f191,f190,f193,f123,f716,f1321,f1322,f1323,f1833,f1826,f1861,f1819,f1862,f1437,f1438,f1810,f1489,f1910,f1523,f1570,f1572,f1763,f1949,f1950,f1951,f1626,f1636,f1637,f1655,f1751,f1703,f1705,f1711,f1739,f128,f135,f136,f232,f773,f776,f790,f792,f795,f805,f827,f830,f926,f936,f954,f998,f1330,f1333,f1335,f1341,f1352,f1362,f1372,f1981,f1995,f178,f1994,f940,f179,f180,f1004,f1331,f1376,f1993,f920,f924,f986,f988,f1314,f1318,f1320,f1354,f1356,f1360,f918,f192,f189,f188,f187,f186,f1312,f938,f939,f1344,f937,f1515])).
fof(f2837,plain,(
  ( ! [X30,X33,X31,X32] : (s(X32,f(s(bool,X33))) != s(X32,g(s(bool,X33))) | s(fun(bool,X32),X30) = s(fun(bool,X32),X31) | s(bool,X33) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,'T_0') = s(bool,sK7(X30,X31,X32,bool))) )),
  inference(superposition,[],[f185,f982])).
fof(f2838,plain,(
  ( ! [X37,X35,X36,X34] : (s(X36,f(s(bool,'T_0'))) != s(X36,g(s(bool,'T_0'))) | s(fun(bool,X36),X34) = s(fun(bool,X36),X35) | s(bool,X37) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,X37) = s(bool,sK7(X34,X35,X36,bool))) )),
  inference(superposition,[],[f185,f982])).
fof(f2848,plain,(
  ( ! [X30,X33,X31,X32] : (s(bool,X33) != s(bool,f(s(X30,sK7(X31,X32,bool,X30)))) | s(fun(X30,bool),X31) = s(fun(X30,bool),X32) | s(bool,X33) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,'T_0') = s(bool,g(s(X30,sK7(X31,X32,bool,X30))))) )),
  inference(superposition,[],[f185,f982])).
fof(f2849,plain,(
  ( ! [X37,X35,X36,X34] : (s(bool,'T_0') != s(bool,f(s(X34,sK7(X35,X36,bool,X34)))) | s(fun(X34,bool),X35) = s(fun(X34,bool),X36) | s(bool,X37) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,X37) = s(bool,g(s(X34,sK7(X35,X36,bool,X34))))) )),
  inference(superposition,[],[f185,f982])).
fof(f3836,plain,(
  ( ! [X61,X59,X57,X60,X58,X56] : (s('type.option.option'(X56),'const.option.NONE_0') = s('type.option.option'(X56),X59) | s('type.option.option'(X56),X59) = s('type.option.option'(X56),'const.option.SOME_1'(s(X56,sK8(X60,X59,X58,X57,X56)))) | s(bool,X61) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,X61) = s(bool,'const.option.OPTREL_3'(s(fun(X56,fun(X57,bool)),X58),s('type.option.option'(X56),X59),s('type.option.option'(X57),X60)))) )),
  inference(subsumption_resolution,[],[f3817,f126])).
fof(f3817,plain,(
  ( ! [X61,X59,X57,X60,X58,X56] : (~p(s(bool,'T_0')) | s('type.option.option'(X56),'const.option.NONE_0') = s('type.option.option'(X56),X59) | s('type.option.option'(X56),X59) = s('type.option.option'(X56),'const.option.SOME_1'(s(X56,sK8(X60,X59,X58,X57,X56)))) | s(bool,X61) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,X61) = s(bool,'const.option.OPTREL_3'(s(fun(X56,fun(X57,bool)),X58),s('type.option.option'(X56),X59),s('type.option.option'(X57),X60)))) )),
  inference(superposition,[],[f186,f982])).
fof(f2840,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK7(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK7(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f185,f129])).
fof(f2841,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(X3,sK7(X4,X5,bool,X3)))) | s(fun(X3,bool),X4) = s(fun(X3,bool),X5) | s(bool,X6) = s(bool,g(s(X3,sK7(X4,X5,bool,X3)))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f185,f259])).
fof(f2842,plain,(
  ( ! [X10,X8,X7,X9] : (s(bool,X10) != s(bool,f(s(X7,sK7(X8,X9,bool,X7)))) | s(fun(X7,bool),X8) = s(fun(X7,bool),X9) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,g(s(X7,sK7(X8,X9,bool,X7))))) )),
  inference(superposition,[],[f185,f259])).
fof(f1331,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK6(X2,X1,X0)))) | s(bool,X3) = s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f1324,f126])).
fof(f1324,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK6(X2,X1,X0)))) | s(bool,X3) = s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1004,f259])).
fof(f1376,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | s(X0,X2) != s(X0,X3) | s(bool,X4) = s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f1366,f126])).
fof(f1366,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | s(X0,X2) != s(X0,X3) | s(bool,X4) = s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f178,f259])).
fof(f920,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,happ(s(fun(bool,bool),X8),s(bool,X7)))) | s(bool,X6) = s(bool,sK6(X6,X8,bool)) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f179,f259])).
fof(f924,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,happ(s(fun(bool,bool),X17),s(bool,X16)))) | s(bool,'T_0') = s(bool,sK6('T_0',X17,bool)) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f179,f259])).
fof(f986,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'T_0')))) | ~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,X0)))) | s(bool,X2) = s(bool,sK6(X0,X1,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f180,f259])).
fof(f988,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,happ(s(fun(bool,bool),X7),s(bool,X8)))) | ~p(s(bool,happ(s(fun(bool,bool),X7),s(bool,X6)))) | s(bool,'T_0') = s(bool,sK6(X6,X7,bool)) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f180,f259])).
fof(f998,plain,(
  ( ! [X12,X13] : (~p(s(bool,happ(s(fun(bool,bool),X13),s(bool,X12)))) | p(s(bool,happ(s(fun(bool,bool),X13),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK6(X12,X13,bool))) )),
  inference(forward_demodulation,[],[f990,f232])).
fof(f990,plain,(
  ( ! [X12,X13] : (p(s(bool,happ(s(fun(bool,bool),X13),s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)))))) | ~p(s(bool,happ(s(fun(bool,bool),X13),s(bool,X12)))) | s(bool,'T_0') = s(bool,sK6(X12,X13,bool))) )),
  inference(superposition,[],[f180,f335])).
fof(f1314,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,happ(s(fun(bool,bool),X8),s(bool,X7)))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X8),s(bool,sK6(X6,X8,bool)))) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f1004,f259])).
fof(f1318,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,happ(s(fun(bool,bool),X17),s(bool,X16)))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X17),s(bool,sK6('T_0',X17,bool)))) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f1004,f259])).
fof(f1320,plain,(
  ( ! [X23,X22] : (~p(s(bool,happ(s(fun(bool,bool),X23),s(bool,X22)))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X23),s(bool,sK6('F_0',X23,bool)))) | s(bool,'T_0') = s(bool,X22)) )),
  inference(superposition,[],[f1004,f129])).
fof(f1335,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,happ(s(fun(bool,bool),X8),s(bool,X7)))) | p(s(bool,happ(s(fun(bool,bool),X8),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK6(X6,X8,bool)) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f998,f259])).
fof(f1341,plain,(
  ( ! [X23,X22] : (~p(s(bool,happ(s(fun(bool,bool),X23),s(bool,X22)))) | p(s(bool,happ(s(fun(bool,bool),X23),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK6('F_0',X23,bool)) | s(bool,'T_0') = s(bool,X22)) )),
  inference(superposition,[],[f998,f129])).
fof(f1352,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK6(X1,X0,bool)) | s(bool,X2) = s(bool,happ(s(fun(bool,bool),X0),s(bool,X1))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f1345,f126])).
fof(f1345,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK6(X1,X0,bool)) | s(bool,X2) = s(bool,happ(s(fun(bool,bool),X0),s(bool,X1))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f998,f259])).
fof(f1354,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X1),s(bool,X2)))) | s(bool,X0) != s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f178,f259])).
fof(f1356,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,happ(s(fun(bool,bool),X10),s(bool,X9)))) | p(s(bool,happ(s(fun(bool,bool),X10),s(bool,X11)))) | s(bool,X8) != s(bool,X11) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f178,f259])).
fof(f1360,plain,(
  ( ! [X24,X23,X25,X22] : (~p(s(bool,happ(s(fun(bool,bool),X23),s(bool,X22)))) | p(s(bool,happ(s(fun(bool,bool),X23),s(bool,X24)))) | s(bool,'T_0') != s(bool,X24) | s(bool,X22) = s(bool,X25) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f178,f259])).
fof(f1362,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,happ(s(fun(bool,bool),X31),s(bool,X30)))) | p(s(bool,happ(s(fun(bool,bool),X31),s(bool,X32)))) | s(bool,'F_0') != s(bool,X32) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f178,f129])).
fof(f1372,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,happ(s(fun(bool,bool),X17),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X17),s(bool,X18)))) | s(bool,X16) != s(bool,X18) | s(bool,'T_0') = s(bool,X16)) )),
  inference(forward_demodulation,[],[f1358,f232])).
fof(f1358,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,happ(s(fun(bool,bool),X17),s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)))))) | p(s(bool,happ(s(fun(bool,bool),X17),s(bool,X18)))) | s(bool,X16) != s(bool,X18) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f178,f335])).
fof(f918,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'T_0')))) | s(bool,X0) = s(bool,sK6(X0,X1,bool)) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f179,f259])).
fof(f1312,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X1),s(bool,sK6(X0,X1,bool)))) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1004,f259])).
fof(f1333,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK6(X0,X1,bool)) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f998,f259])).
fof(f1330,plain,(
  ( ! [X12,X13] : (~p(s(bool,happ(s(fun(bool,bool),X13),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X13),s(bool,sK6(X12,X13,bool)))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(forward_demodulation,[],[f1316,f232])).
fof(f1316,plain,(
  ( ! [X12,X13] : (~p(s(bool,happ(s(fun(bool,bool),X13),s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)))))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X13),s(bool,sK6(X12,X13,bool)))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f1004,f335])).
fof(f716,plain,(
  ( ! [X35,X36,X34] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,X34))),s(sK1,X35)))) | s(bool,X36) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK5),s(sK0,X34))),s(sK1,X35))) | s(bool,'T_0') = s(bool,X36)) )),
  inference(subsumption_resolution,[],[f526,f126])).
fof(f526,plain,(
  ( ! [X35,X36,X34] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,X34))),s(sK1,X35)))) | s(bool,X36) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK5),s(sK0,X34))),s(sK1,X35))) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f123,f259])).
fof(f3961,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.option.option'(X0),'const.option.NONE_0'),s('type.option.option'(X1),'const.option.SOME_1'(s(X1,X3))))))) )),
  inference(unit_resulting_resolution,[],[f173,f173,f187])).
fof(f4187,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X3))),s('type.option.option'(X1),'const.option.NONE_0'))))) )),
  inference(unit_resulting_resolution,[],[f173,f173,f188])).
fof(f4201,plain,(
  ( ! [X6,X10,X8,X7,X5,X9] : (s('type.option.option'(X5),'const.option.NONE_0') = s('type.option.option'(X5),X8) | s('type.option.option'(X6),X9) = s('type.option.option'(X6),'const.option.SOME_1'(s(X6,sK9(X9,X8,X7,X6,X5)))) | s(bool,X10) = s(bool,'const.option.OPTREL_3'(s(fun(X5,fun(X6,bool)),X7),s('type.option.option'(X5),X8),s('type.option.option'(X6),X9))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(subsumption_resolution,[],[f4191,f126])).
fof(f4191,plain,(
  ( ! [X6,X10,X8,X7,X5,X9] : (~p(s(bool,'T_0')) | s('type.option.option'(X5),'const.option.NONE_0') = s('type.option.option'(X5),X8) | s('type.option.option'(X6),X9) = s('type.option.option'(X6),'const.option.SOME_1'(s(X6,sK9(X9,X8,X7,X6,X5)))) | s(bool,X10) = s(bool,'const.option.OPTREL_3'(s(fun(X5,fun(X6,bool)),X7),s('type.option.option'(X5),X8),s('type.option.option'(X6),X9))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f188,f259])).
fof(f3975,plain,(
  ( ! [X6,X10,X8,X7,X5,X9] : (s('type.option.option'(X6),'const.option.NONE_0') = s('type.option.option'(X6),X9) | s('type.option.option'(X5),X8) = s('type.option.option'(X5),'const.option.SOME_1'(s(X5,sK8(X9,X8,X7,X6,X5)))) | s(bool,X10) = s(bool,'const.option.OPTREL_3'(s(fun(X5,fun(X6,bool)),X7),s('type.option.option'(X5),X8),s('type.option.option'(X6),X9))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(subsumption_resolution,[],[f3965,f126])).
fof(f3965,plain,(
  ( ! [X6,X10,X8,X7,X5,X9] : (~p(s(bool,'T_0')) | s('type.option.option'(X6),'const.option.NONE_0') = s('type.option.option'(X6),X9) | s('type.option.option'(X5),X8) = s('type.option.option'(X5),'const.option.SOME_1'(s(X5,sK8(X9,X8,X7,X6,X5)))) | s(bool,X10) = s(bool,'const.option.OPTREL_3'(s(fun(X5,fun(X6,bool)),X7),s('type.option.option'(X5),X8),s('type.option.option'(X6),X9))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f187,f259])).
fof(f3833,plain,(
  ( ! [X14,X12,X10,X13,X11,X9] : (s('type.option.option'(X9),'const.option.NONE_0') = s('type.option.option'(X9),X12) | s('type.option.option'(X9),X12) = s('type.option.option'(X9),'const.option.SOME_1'(s(X9,sK8(X13,X12,X11,X10,X9)))) | s(bool,X14) = s(bool,'const.option.OPTREL_3'(s(fun(X9,fun(X10,bool)),X11),s('type.option.option'(X9),X12),s('type.option.option'(X10),X13))) | s(bool,'T_0') = s(bool,X14)) )),
  inference(subsumption_resolution,[],[f3809,f126])).
fof(f3809,plain,(
  ( ! [X14,X12,X10,X13,X11,X9] : (~p(s(bool,'T_0')) | s('type.option.option'(X9),'const.option.NONE_0') = s('type.option.option'(X9),X12) | s('type.option.option'(X9),X12) = s('type.option.option'(X9),'const.option.SOME_1'(s(X9,sK8(X13,X12,X11,X10,X9)))) | s(bool,X14) = s(bool,'const.option.OPTREL_3'(s(fun(X9,fun(X10,bool)),X11),s('type.option.option'(X9),X12),s('type.option.option'(X10),X13))) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f186,f259])).
fof(f3830,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s('type.option.option'(sK0),'const.option.NONE_0') = s('type.option.option'(sK0),sK3) | s('type.option.option'(sK0),sK3) = s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK8(sK2,sK3,sK5,sK1,sK0))))) )),
  inference(condensation,[],[f3829])).
fof(f3829,plain,(
  ( ! [X2,X1] : (s('type.option.option'(sK0),'const.option.NONE_0') = s('type.option.option'(sK0),sK3) | s('type.option.option'(sK0),sK3) = s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK8(sK2,sK3,sK5,sK1,sK0)))) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f3804,f296])).
fof(f3804,plain,(
  ( ! [X2,X1] : (~p(s(bool,X1)) | s('type.option.option'(sK0),'const.option.NONE_0') = s('type.option.option'(sK0),sK3) | s('type.option.option'(sK0),sK3) = s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK8(sK2,sK3,sK5,sK1,sK0)))) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f186,f982])).
fof(f2831,plain,(
  ( ! [X10,X8,X7,X9] : (s(X9,f(s(bool,X10))) != s(X9,g(s(bool,X10))) | s(fun(bool,X9),X7) = s(fun(bool,X9),X8) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,sK7(X7,X8,X9,bool))) )),
  inference(superposition,[],[f185,f259])).
fof(f2830,plain,(
  ( ! [X6,X4,X5,X3] : (s(X5,f(s(bool,'T_0'))) != s(X5,g(s(bool,'T_0'))) | s(fun(bool,X5),X3) = s(fun(bool,X5),X4) | s(bool,X6) = s(bool,sK7(X3,X4,X5,bool)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f185,f259])).
fof(f1515,plain,(
  ( ! [X52,X50,X51] : (~p(s(bool,happ(s(fun(bool,bool),X51),s(bool,X50)))) | p(s(bool,happ(s(fun(bool,bool),X51),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK6('const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)),X51,bool)) | s(bool,X50) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(superposition,[],[f998,f982])).
fof(f940,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X2) = s(X0,sK6(X2,X1,X0)) | s(bool,X3) = s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f930,f126])).
fof(f930,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | s(X0,X2) = s(X0,sK6(X2,X1,X0)) | s(bool,X3) = s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f179,f259])).
fof(f954,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) = s(bool,sK6(X1,X0,bool)) | s(bool,'T_0') = s(bool,X1) | s(bool,X2) = s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f948,f126])).
fof(f948,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s(bool,X1) = s(bool,sK6(X1,X0,bool)) | s(bool,'T_0') = s(bool,X1) | s(bool,X2) = s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f936,f259])).
fof(f805,plain,(
  s('type.option.option'('type.option.option'('type.option.option'(bool))),'const.option.SOME_1'(s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))))) != s('type.option.option'('type.option.option'('type.option.option'(bool))),'const.option.SOME_1'(s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0')))))))),
  inference(unit_resulting_resolution,[],[f795,f176])).
fof(f795,plain,(
  s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0')))))),
  inference(unit_resulting_resolution,[],[f790,f176])).
fof(f792,plain,(
  ( ! [X12,X13] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X13))) | s(bool,X12) = s(bool,X13) | s(bool,'T_0') = s(bool,X12)) )),
  inference(forward_demodulation,[],[f769,f232])).
fof(f769,plain,(
  ( ! [X12,X13] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X13))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))))) | s(bool,X12) = s(bool,X13) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f176,f335])).
fof(f790,plain,(
  s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0')))),
  inference(forward_demodulation,[],[f763,f232])).
fof(f763,plain,(
  s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)))))),
  inference(unit_resulting_resolution,[],[f224,f176])).
fof(f1910,plain,(
  ( ! [X72,X71,X73] : (s(bool,'T_0') = s(bool,sK6('const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)),X72,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X72),s(bool,X71)))) | s(bool,X71) = s(bool,X73) | s(bool,'T_0') = s(bool,X73)) )),
  inference(forward_demodulation,[],[f1522,f227])).
fof(f1522,plain,(
  ( ! [X72,X71,X73] : (~p(s(bool,happ(s(fun(bool,bool),X72),s(bool,X71)))) | s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) = s(bool,sK6('const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)),X72,bool)) | s(bool,X71) = s(bool,X73) | s(bool,'T_0') = s(bool,X73)) )),
  inference(superposition,[],[f179,f982])).
fof(f1862,plain,(
  ( ! [X40] : (s(bool,'F_0') = s(bool,X40) | s(bool,X40) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)))) )),
  inference(subsumption_resolution,[],[f1435,f224])).
fof(f1435,plain,(
  ( ! [X40] : (s(bool,'F_0') = s(bool,X40) | s(bool,X40) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,'T_0') = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)))) )),
  inference(superposition,[],[f982,f232])).
fof(f884,plain,(
  ( ! [X2,X0,X1] : (s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X0))))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X1))))) | s(bool,'T_0') = s(bool,X1) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f806,f259])).
fof(f828,plain,(
  ( ! [X17,X18,X16] : (s(bool,X16) != s(bool,X17) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X17))) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f177,f259])).
fof(f824,plain,(
  ( ! [X6,X8,X7] : (s(bool,X7) != s(bool,X8) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X6))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X8))) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f177,f259])).
fof(f822,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,X1) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X0))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X1))) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f177,f259])).
fof(f806,plain,(
  ( ! [X0] : (s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X0))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f795,f129])).
fof(f767,plain,(
  ( ! [X6,X8,X7] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X7))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X8))) | s(bool,X6) = s(bool,X8) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f176,f259])).
fof(f765,plain,(
  ( ! [X2,X0,X1] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X1))) | s(bool,X0) = s(bool,X1) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f176,f259])).
fof(f764,plain,(
  ( ! [X0,X1] : (s('type.option.option'('type.option.option'(X0)),'const.option.SOME_1'(s('type.option.option'(X0),'const.option.NONE_0'))) != s('type.option.option'('type.option.option'(X0)),'const.option.SOME_1'(s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1)))))) )),
  inference(unit_resulting_resolution,[],[f173,f176])).
fof(f726,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f714,f129])).
fof(f715,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f517,f126])).
fof(f517,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f184,f259])).
fof(f714,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f513,f126])).
fof(f513,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f135,f259])).
fof(f498,plain,(
  ( ! [X4,X2,X3] : (s(bool,X2) != s(bool,X3) | p(s(bool,X3)) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f131,f259])).
fof(f428,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f259,f259])).
fof(f184,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f75])).
fof(f6887,plain,(
  ( ! [X4,X1] : (s('type.option.option'(X1),'const.option.NONE_0') = s('type.option.option'(X1),X4)) ) | ($spl8 | $spl21 | $spl32)),
  inference(subsumption_resolution,[],[f6807,f126])).
fof(f6807,plain,(
  ( ! [X4,X1] : (~p(s(bool,'T_0')) | s('type.option.option'(X1),'const.option.NONE_0') = s('type.option.option'(X1),X4)) ) | ($spl8 | $spl21 | $spl32)),
  inference(backward_demodulation,[],[f6625,f6619])).
fof(f6619,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.option.option'(X1),'const.option.NONE_0') = s('type.option.option'(X1),X4) | ~p(s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.option.option'(X0),X3),s('type.option.option'(X1),X4))))) ) | ($spl8 | $spl21 | $spl32)),
  inference(subsumption_resolution,[],[f6618,f128])).
fof(f6618,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'F_0')) | s('type.option.option'(X1),'const.option.NONE_0') = s('type.option.option'(X1),X4) | ~p(s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.option.option'(X0),X3),s('type.option.option'(X1),X4))))) ) | ($spl8 | $spl21 | $spl32)),
  inference(forward_demodulation,[],[f6601,f6301])).
fof(f6601,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),'const.option.NONE_0')))) | s('type.option.option'(X1),'const.option.NONE_0') = s('type.option.option'(X1),X4) | ~p(s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.option.option'(X0),X3),s('type.option.option'(X1),X4))))) ) | ($spl8 | $spl21 | $spl32)),
  inference(backward_demodulation,[],[f6598,f191])).
fof(f6598,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),'const.option.NONE_0')))) ) | ($spl8 | $spl21 | $spl32)),
  inference(condensation,[],[f6597])).
fof(f6597,plain,(
  ( ! [X2,X3] : (s(bool,X2) = s(bool,X3) | s(bool,X2) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),'const.option.NONE_0')))) ) | ($spl8 | $spl21 | $spl32)),
  inference(duplicate_literal_removal,[],[f6596])).
fof(f6596,plain,(
  ( ! [X2,X3] : (s(bool,X2) = s(bool,X3) | s(bool,X2) = s(bool,X3) | s(bool,X2) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),'const.option.NONE_0')))) ) | ($spl8 | $spl21 | $spl32)),
  inference(condensation,[],[f6595])).
fof(f6595,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,X0) = s(bool,X1) | s(bool,X3) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),'const.option.NONE_0')))) ) | ($spl8 | $spl21 | $spl32)),
  inference(global_subsumption,[],[f125,f124,f126,f129,f132,f131,f173,f177,f176,f184,f181,f185,f219,f223,f259,f296,f428,f498,f714,f715,f726,f764,f765,f767,f806,f822,f824,f828,f884,f999,f1002,f1833,f1826,f1861,f1819,f1862,f1810,f1910,f1763,f1711,f128,f135,f136,f773,f790,f792,f795,f805,f827,f830,f954,f940,f1515,f2532,f2533,f2644,f2646,f2667,f2670,f2671,f2583,f2690,f2829,f2830,f2831,f2854,f2916,f3033,f3035,f3041,f3108,f3109,f3111,f3112,f3139,f3500,f3507,f3189,f3191,f3197,f3593,f3698,f3699,f3792,f3708,f3709,f3802,f3773,f3823,f3830,f3831,f3832,f3833,f3838,f3975,f3976,f3978,f4201,f4202,f4204,f3715,f3717,f3713,f4187,f3961,f3625,f192,f189,f188,f187,f186,f123,f716,f2665,f3194,f1330,f936,f2673,f1333,f1312,f918,f3196,f3193,f3142,f3141,f3113,f3040,f3038,f3037,f2732,f2730,f2674,f1372,f1362,f1360,f1356,f1354,f1352,f1341,f1335,f1320,f1318,f1314,f998,f988,f986,f926,f924,f920,f191,f190,f193,f3799,f3798,f3778,f3777,f3775,f3718,f3508,f2742,f1376,f1331,f1004,f180,f179,f178,f3195,f3039,f2858,f2842,f2841,f2840,f4744,f5869,f4873,f4957,f5655,f5043,f5045,f5047,f5048,f5049,f5050,f5051,f5052,f5109,f5111,f5112,f5113,f5114,f5115,f5116,f5119,f5120,f5122,f5123,f5153,f5155,f5156,f5157,f5972,f5973,f5974,f5201,f5203,f5205,f5206,f5207,f5208,f5209,f5210,f5265,f5267,f5269,f5270,f5271,f5272,f5273,f5274,f5331,f5333,f5335,f5336,f5337,f5338,f5340,f5343,f5344,f5368,f5370,f5371,f5372,f6039,f6040,f6045,f6046,f5423,f5425,f5427,f5428,f5429,f5430,f5431,f5432,f3836,f3797,f3796,f3795,f3794,f3793,f3712,f3706,f3705,f3704,f3629,f3617,f2849,f2848,f2838,f2837,f2827,f2781,f2692,f2610,f2579,f2151,f1995,f1994,f1993,f1981,f1951,f1950,f1949,f1860,f1751,f1739,f1705,f1703,f1655,f1637,f1636,f1626,f1572,f1570,f1523,f1489,f1438,f1437,f1344,f1323,f1322,f1321,f982,f956,f939,f938,f937,f826,f776,f775,f774,f335,f232,f227,f224,f3826,f6179,f6181,f6206,f6229,f6233,f6243,f6242,f6262,f6266,f6268,f6291,f6301,f6300,f6320,f6324,f6386,f6416,f6417,f6419,f6422,f6423,f6424,f6505,f6511,f6509,f6532,f6533,f6534,f6535,f6536,f6537,f6538,f6539,f6540,f6541,f6542,f6543,f6544,f6594])).
fof(f6594,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X3) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),'const.option.NONE_0'))) | s(bool,'F_0') = s(bool,X2) | s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,X3)) ) | ($spl8 | $spl21 | $spl32)),
  inference(forward_demodulation,[],[f6593,f6509])).
fof(f6593,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,X2) | s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,X3) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,X0) = s(bool,X3)) ) | ($spl8 | $spl21 | $spl32)),
  inference(forward_demodulation,[],[f6545,f6301])).
fof(f6545,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X2) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),'const.option.NONE_0'))) | s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,X3) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,X0) = s(bool,X3)) ) | $spl32),
  inference(backward_demodulation,[],[f6509,f1437])).
fof(f6424,plain,(
  ( ! [X70,X68,X72,X71,X69,X67,X73] : (s('type.option.option'(X68),'const.option.NONE_0') = s('type.option.option'(X68),X71) | s('type.option.option'(X68),X71) = s('type.option.option'(X68),'const.option.SOME_1'(s(X68,sK9(X71,X70,X69,X68,X67)))) | s(bool,X72) = s(bool,X73) | s(bool,'F_0') = s(bool,X73) | s(bool,X72) = s(bool,'const.option.OPTREL_3'(s(fun(X67,fun(X68,bool)),X69),s('type.option.option'(X67),X70),s('type.option.option'(X68),X71)))) ) | $spl8),
  inference(subsumption_resolution,[],[f6410,f126])).
fof(f6410,plain,(
  ( ! [X70,X68,X72,X71,X69,X67,X73] : (~p(s(bool,'T_0')) | s('type.option.option'(X68),'const.option.NONE_0') = s('type.option.option'(X68),X71) | s('type.option.option'(X68),X71) = s('type.option.option'(X68),'const.option.SOME_1'(s(X68,sK9(X71,X70,X69,X68,X67)))) | s(bool,X72) = s(bool,X73) | s(bool,'F_0') = s(bool,X73) | s(bool,X72) = s(bool,'const.option.OPTREL_3'(s(fun(X67,fun(X68,bool)),X69),s('type.option.option'(X67),X70),s('type.option.option'(X68),X71)))) ) | $spl8),
  inference(superposition,[],[f189,f2533])).
fof(f6423,plain,(
  ( ! [X4,X2,X5,X3,X1] : (s('type.option.option'(X1),'const.option.NONE_0') = s('type.option.option'(X1),X2) | s(bool,'F_0') = s(bool,'const.option.OPTREL_3'(s(fun(X5,fun(X1,bool)),X4),s('type.option.option'(X5),X3),s('type.option.option'(X1),X2))) | s('type.option.option'(X1),X2) = s('type.option.option'(X1),'const.option.SOME_1'(s(X1,sK9(X2,X3,X4,X1,X5))))) ) | $spl8),
  inference(global_subsumption,[],[f125,f124,f126,f129,f132,f131,f173,f177,f176,f184,f181,f185,f219,f223,f259,f296,f428,f498,f714,f715,f726,f764,f765,f767,f806,f822,f824,f828,f884,f999,f1002,f1833,f1826,f1861,f1819,f1862,f1810,f1910,f1763,f1711,f128,f135,f136,f773,f790,f792,f795,f805,f827,f830,f954,f940,f1515,f2532,f2533,f2644,f2646,f2667,f2670,f2671,f2583,f2690,f2829,f2830,f2831,f2854,f2916,f3033,f3035,f3041,f3108,f3109,f3111,f3112,f3139,f3500,f3507,f3189,f3191,f3197,f3593,f3698,f3699,f3792,f3708,f3709,f3802,f3773,f3823,f3830,f3831,f3832,f3833,f3838,f3975,f3976,f3978,f4201,f4202,f4204,f3715,f3717,f3713,f4187,f3961,f3625,f192,f189,f188,f187,f186,f123,f716,f2665,f3194,f1330,f936,f2673,f1333,f1312,f918,f3196,f3193,f3142,f3141,f3113,f3040,f3038,f3037,f2732,f2730,f2674,f1372,f1362,f1360,f1356,f1354,f1352,f1341,f1335,f1320,f1318,f1314,f998,f988,f986,f926,f924,f920,f191,f190,f193,f3799,f3798,f3778,f3777,f3775,f3718,f3508,f2742,f1376,f1331,f1004,f180,f179,f178,f3195,f3039,f2858,f2842,f2841,f2840,f4744,f5869,f4873,f4957,f5655,f5043,f5045,f5047,f5048,f5049,f5050,f5051,f5052,f5109,f5111,f5112,f5113,f5114,f5115,f5116,f5119,f5120,f5122,f5123,f5153,f5155,f5156,f5157,f5972,f5973,f5974,f5201,f5203,f5205,f5206,f5207,f5208,f5209,f5210,f5265,f5267,f5269,f5270,f5271,f5272,f5273,f5274,f5331,f5333,f5335,f5336,f5337,f5338,f5340,f5343,f5344,f5368,f5370,f5371,f5372,f6039,f6040,f6045,f6046,f5423,f5425,f5427,f5428,f5429,f5430,f5431,f5432,f3836,f3797,f3796,f3795,f3794,f3793,f3712,f3706,f3705,f3704,f3629,f3617,f2849,f2848,f2838,f2837,f2827,f2781,f2692,f2610,f2579,f2151,f1995,f1994,f1993,f1981,f1951,f1950,f1949,f1860,f1751,f1739,f1705,f1703,f1655,f1637,f1636,f1626,f1572,f1570,f1523,f1489,f1438,f1437,f1344,f1323,f1322,f1321,f982,f956,f939,f938,f937,f826,f776,f775,f774,f335,f232,f227,f224,f6179,f6229,f6291,f6386,f6416,f6417,f6419,f6422])).
fof(f6422,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s('type.option.option'(X1),'const.option.NONE_0') = s('type.option.option'(X1),X2) | s('type.option.option'(X1),X2) = s('type.option.option'(X1),'const.option.SOME_1'(s(X1,sK9(X2,X3,X4,X1,X5)))) | s(bool,'F_0') = s(bool,'const.option.OPTREL_3'(s(fun(X5,fun(X1,bool)),X4),s('type.option.option'(X5),X3),s('type.option.option'(X1),X2)))) ) | $spl8),
  inference(condensation,[],[f6421])).
fof(f6421,plain,(
  ( ! [X61,X66,X64,X62,X60,X65,X63] : (s('type.option.option'(X61),'const.option.NONE_0') = s('type.option.option'(X61),X64) | s('type.option.option'(X61),X64) = s('type.option.option'(X61),'const.option.SOME_1'(s(X61,sK9(X64,X63,X62,X61,X60)))) | s(bool,'F_0') = s(bool,'const.option.OPTREL_3'(s(fun(X60,fun(X61,bool)),X62),s('type.option.option'(X60),X63),s('type.option.option'(X61),X64))) | s(bool,X65) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) ) | $spl8),
  inference(subsumption_resolution,[],[f6409,f296])).
fof(f6409,plain,(
  ( ! [X61,X66,X64,X62,X60,X65,X63] : (~p(s(bool,X65)) | s('type.option.option'(X61),'const.option.NONE_0') = s('type.option.option'(X61),X64) | s('type.option.option'(X61),X64) = s('type.option.option'(X61),'const.option.SOME_1'(s(X61,sK9(X64,X63,X62,X61,X60)))) | s(bool,'F_0') = s(bool,'const.option.OPTREL_3'(s(fun(X60,fun(X61,bool)),X62),s('type.option.option'(X60),X63),s('type.option.option'(X61),X64))) | s(bool,X65) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) ) | $spl8),
  inference(superposition,[],[f189,f2533])).
fof(f6386,plain,(
  ( ! [X12,X13] : (s(bool,X12) = s(bool,X13) | s(bool,'F_0') = s(bool,X13) | s(bool,X12) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)))) ) | $spl8),
  inference(trivial_inequality_removal,[],[f6382])).
fof(f6382,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X12) = s(bool,X13) | s(bool,'F_0') = s(bool,X13) | s(bool,X12) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)))) ) | $spl8),
  inference(superposition,[],[f224,f2533])).
fof(f6324,plain,(
  ( ! [X37,X35,X36,X34] : (s(bool,X36) = s(bool,X37) | s(bool,'F_0') = s(bool,X37) | s(bool,X36) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(X34,bool)),X35),s('type.option.option'(sK0),sK3),s('type.option.option'(X34),'const.option.NONE_0')))) ) | ($spl8 | $spl21)),
  inference(subsumption_resolution,[],[f6313,f126])).
fof(f6313,plain,(
  ( ! [X37,X35,X36,X34] : (~p(s(bool,'T_0')) | s(bool,X36) = s(bool,X37) | s(bool,'F_0') = s(bool,X37) | s(bool,X36) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(X34,bool)),X35),s('type.option.option'(sK0),sK3),s('type.option.option'(X34),'const.option.NONE_0')))) ) | ($spl8 | $spl21)),
  inference(superposition,[],[f6268,f2533])).
fof(f6300,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(X0,bool)),X1),s('type.option.option'(sK0),sK3),s('type.option.option'(X0),'const.option.NONE_0')))) ) | ($spl8 | $spl21)),
  inference(unit_resulting_resolution,[],[f223,f6268,f2646])).
fof(f6291,plain,(
  ( ! [X70,X68,X72,X71,X69,X67,X73] : (s('type.option.option'(X67),'const.option.NONE_0') = s('type.option.option'(X67),X70) | s('type.option.option'(X68),X71) = s('type.option.option'(X68),'const.option.SOME_1'(s(X68,sK9(X71,X70,X69,X68,X67)))) | s(bool,X72) = s(bool,X73) | s(bool,'F_0') = s(bool,X73) | s(bool,X72) = s(bool,'const.option.OPTREL_3'(s(fun(X67,fun(X68,bool)),X69),s('type.option.option'(X67),X70),s('type.option.option'(X68),X71)))) ) | $spl8),
  inference(subsumption_resolution,[],[f6282,f126])).
fof(f6282,plain,(
  ( ! [X70,X68,X72,X71,X69,X67,X73] : (~p(s(bool,'T_0')) | s('type.option.option'(X67),'const.option.NONE_0') = s('type.option.option'(X67),X70) | s('type.option.option'(X68),X71) = s('type.option.option'(X68),'const.option.SOME_1'(s(X68,sK9(X71,X70,X69,X68,X67)))) | s(bool,X72) = s(bool,X73) | s(bool,'F_0') = s(bool,X73) | s(bool,X72) = s(bool,'const.option.OPTREL_3'(s(fun(X67,fun(X68,bool)),X69),s('type.option.option'(X67),X70),s('type.option.option'(X68),X71)))) ) | $spl8),
  inference(superposition,[],[f188,f2533])).
fof(f6266,plain,(
  ( ! [X37,X35,X36,X34] : (s(bool,X36) = s(bool,X37) | s(bool,'F_0') = s(bool,X37) | s(bool,X36) = s(bool,'const.option.OPTREL_3'(s(fun(X34,fun(sK0,bool)),X35),s('type.option.option'(X34),'const.option.NONE_0'),s('type.option.option'(sK0),sK3)))) ) | ($spl8 | $spl21)),
  inference(subsumption_resolution,[],[f6255,f126])).
fof(f6255,plain,(
  ( ! [X37,X35,X36,X34] : (~p(s(bool,'T_0')) | s(bool,X36) = s(bool,X37) | s(bool,'F_0') = s(bool,X37) | s(bool,X36) = s(bool,'const.option.OPTREL_3'(s(fun(X34,fun(sK0,bool)),X35),s('type.option.option'(X34),'const.option.NONE_0'),s('type.option.option'(sK0),sK3)))) ) | ($spl8 | $spl21)),
  inference(superposition,[],[f6206,f2533])).
fof(f6229,plain,(
  ( ! [X70,X68,X72,X71,X69,X67,X73] : (s('type.option.option'(X68),'const.option.NONE_0') = s('type.option.option'(X68),X71) | s('type.option.option'(X67),X70) = s('type.option.option'(X67),'const.option.SOME_1'(s(X67,sK8(X71,X70,X69,X68,X67)))) | s(bool,X72) = s(bool,X73) | s(bool,'F_0') = s(bool,X73) | s(bool,X72) = s(bool,'const.option.OPTREL_3'(s(fun(X67,fun(X68,bool)),X69),s('type.option.option'(X67),X70),s('type.option.option'(X68),X71)))) ) | $spl8),
  inference(subsumption_resolution,[],[f6220,f126])).
fof(f6220,plain,(
  ( ! [X70,X68,X72,X71,X69,X67,X73] : (~p(s(bool,'T_0')) | s('type.option.option'(X68),'const.option.NONE_0') = s('type.option.option'(X68),X71) | s('type.option.option'(X67),X70) = s('type.option.option'(X67),'const.option.SOME_1'(s(X67,sK8(X71,X70,X69,X68,X67)))) | s(bool,X72) = s(bool,X73) | s(bool,'F_0') = s(bool,X73) | s(bool,X72) = s(bool,'const.option.OPTREL_3'(s(fun(X67,fun(X68,bool)),X69),s('type.option.option'(X67),X70),s('type.option.option'(X68),X71)))) ) | $spl8),
  inference(superposition,[],[f187,f2533])).
fof(f6179,plain,(
  ( ! [X70,X68,X72,X71,X69,X67,X73] : (s('type.option.option'(X67),'const.option.NONE_0') = s('type.option.option'(X67),X70) | s('type.option.option'(X67),X70) = s('type.option.option'(X67),'const.option.SOME_1'(s(X67,sK8(X71,X70,X69,X68,X67)))) | s(bool,X72) = s(bool,X73) | s(bool,'F_0') = s(bool,X73) | s(bool,X72) = s(bool,'const.option.OPTREL_3'(s(fun(X67,fun(X68,bool)),X69),s('type.option.option'(X67),X70),s('type.option.option'(X68),X71)))) ) | $spl8),
  inference(subsumption_resolution,[],[f6172,f126])).
fof(f6172,plain,(
  ( ! [X70,X68,X72,X71,X69,X67,X73] : (~p(s(bool,'T_0')) | s('type.option.option'(X67),'const.option.NONE_0') = s('type.option.option'(X67),X70) | s('type.option.option'(X67),X70) = s('type.option.option'(X67),'const.option.SOME_1'(s(X67,sK8(X71,X70,X69,X68,X67)))) | s(bool,X72) = s(bool,X73) | s(bool,'F_0') = s(bool,X73) | s(bool,X72) = s(bool,'const.option.OPTREL_3'(s(fun(X67,fun(X68,bool)),X69),s('type.option.option'(X67),X70),s('type.option.option'(X68),X71)))) ) | $spl8),
  inference(superposition,[],[f186,f2533])).
fof(f2579,plain,(
  ( ! [X57,X56] : (s(bool,X57) != s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X56))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X57))) | s(bool,'F_0') = s(bool,X56)) ) | $spl8),
  inference(superposition,[],[f177,f1860])).
fof(f2610,plain,(
  ( ! [X45,X46] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X45))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X46))) | s(bool,X46) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,'F_0') = s(bool,X45)) ) | $spl8),
  inference(superposition,[],[f176,f1860])).
fof(f2692,plain,(
  ( ! [X54,X53] : (s(bool,'T_0') = s(bool,sK6('const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)),X54,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X54),s(bool,X53)))) | s(bool,'F_0') = s(bool,X53)) ) | $spl8),
  inference(forward_demodulation,[],[f2614,f227])).
fof(f2614,plain,(
  ( ! [X54,X53] : (~p(s(bool,happ(s(fun(bool,bool),X54),s(bool,X53)))) | s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) = s(bool,sK6('const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)),X54,bool)) | s(bool,'F_0') = s(bool,X53)) ) | $spl8),
  inference(superposition,[],[f179,f1860])).
fof(f2781,plain,(
  ( ! [X26,X27,X25] : (s(bool,X25) = s(bool,X26) | s(bool,'F_0') = s(bool,X26) | s(bool,X27) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,X25) = s(bool,X27)) ) | $spl8),
  inference(subsumption_resolution,[],[f2757,f126])).
fof(f2757,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'T_0')) | s(bool,X25) = s(bool,X26) | s(bool,'F_0') = s(bool,X26) | s(bool,X27) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,X25) = s(bool,X27)) ) | $spl8),
  inference(superposition,[],[f2644,f982])).
fof(f3617,plain,(
  ( ! [X43,X44,X42] : (s(bool,X42) != s(bool,X43) | s(bool,'T_0') = s(bool,X43) | s(bool,X44) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,X42) = s(bool,X44)) ) | $spl8),
  inference(subsumption_resolution,[],[f3555,f223])).
fof(f3555,plain,(
  ( ! [X43,X44,X42] : (s(bool,X42) != s(bool,X43) | s(bool,'T_0') = s(bool,X43) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X44) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,X42) = s(bool,X44)) ) | $spl8),
  inference(superposition,[],[f2583,f982])).
fof(f3629,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,X3) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,X0) = s(bool,X3)) ) | $spl8),
  inference(superposition,[],[f3108,f982])).
fof(f3704,plain,(
  ( ! [X21,X19,X20] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X21),s(bool,X20))) | s(bool,X19) = s(bool,sK6(X19,X21,bool)) | s(bool,X19) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,'T_0') = s(bool,X20)) ) | $spl8),
  inference(superposition,[],[f2667,f982])).
fof(f3705,plain,(
  ( ! [X24,X23,X22] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X24),s(bool,X23))) | s(bool,X22) = s(bool,sK6(X22,X24,bool)) | s(bool,X23) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,'T_0') = s(bool,X22)) ) | $spl8),
  inference(superposition,[],[f2667,f982])).
fof(f3706,plain,(
  ( ! [X26,X27,X25] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X26),s(bool,'T_0'))) | s(bool,X25) = s(bool,sK6(X25,X26,bool)) | s(bool,X27) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,X25) = s(bool,X27)) ) | $spl8),
  inference(superposition,[],[f2667,f982])).
fof(f3712,plain,(
  ( ! [X43,X44,X42] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X43),s(bool,X42))) | s(bool,'T_0') = s(bool,sK6('T_0',X43,bool)) | s(bool,X44) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,X42) = s(bool,X44)) ) | $spl8),
  inference(superposition,[],[f2667,f982])).
fof(f3793,plain,(
  ( ! [X64,X63] : (s(bool,'T_0') = s(bool,sK6('const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)),X64,bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X64),s(bool,X63))) | s(bool,'F_0') = s(bool,X63)) ) | $spl8),
  inference(forward_demodulation,[],[f3719,f227])).
fof(f3719,plain,(
  ( ! [X64,X63] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X64),s(bool,X63))) | s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) = s(bool,sK6('const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)),X64,bool)) | s(bool,'F_0') = s(bool,X63)) ) | $spl8),
  inference(superposition,[],[f2667,f1860])).
fof(f3794,plain,(
  ( ! [X66,X67,X65] : (s(bool,'T_0') = s(bool,sK6('const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)),X66,bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X66),s(bool,X65))) | s(bool,X65) = s(bool,X67) | s(bool,'T_0') = s(bool,X67)) ) | $spl8),
  inference(forward_demodulation,[],[f3720,f227])).
fof(f3720,plain,(
  ( ! [X66,X67,X65] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X66),s(bool,X65))) | s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) = s(bool,sK6('const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)),X66,bool)) | s(bool,X65) = s(bool,X67) | s(bool,'T_0') = s(bool,X67)) ) | $spl8),
  inference(superposition,[],[f2667,f982])).
fof(f3795,plain,(
  ( ! [X68] : (s(bool,'T_0') = s(bool,sK6('const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)),X68,bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X68),s(bool,'T_0')))) ) | $spl8),
  inference(forward_demodulation,[],[f3721,f227])).
fof(f3721,plain,(
  ( ! [X68] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X68),s(bool,'T_0'))) | s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) = s(bool,sK6('const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)),X68,bool))) ) | $spl8),
  inference(superposition,[],[f2667,f227])).
fof(f3796,plain,(
  ( ! [X69] : (s(bool,'F_0') = s(bool,sK6('const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)),X69,bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X69),s(bool,'F_0')))) ) | $spl8),
  inference(forward_demodulation,[],[f3722,f232])).
fof(f3722,plain,(
  ( ! [X69] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X69),s(bool,'F_0'))) | s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) = s(bool,sK6('const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)),X69,bool))) ) | $spl8),
  inference(superposition,[],[f2667,f232])).
fof(f3797,plain,(
  ( ! [X70,X71] : (s(bool,'F_0') = s(bool,sK6('const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)),X71,bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X71),s(bool,X70))) | s(bool,'T_0') = s(bool,X70)) ) | $spl8),
  inference(forward_demodulation,[],[f3723,f232])).
fof(f3723,plain,(
  ( ! [X70,X71] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X71),s(bool,X70))) | s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) = s(bool,sK6('const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)),X71,bool)) | s(bool,'T_0') = s(bool,X70)) ) | $spl8),
  inference(superposition,[],[f2667,f335])).
fof(f5432,plain,(
  ( ! [X227,X225,X228,X226,X224] : (s(X226,f(s(bool,'T_0'))) != s(X226,g(s(bool,'T_0'))) | s(fun(bool,X226),X224) = s(fun(bool,X226),X225) | s(bool,X227) = s(bool,X228) | s(bool,'F_0') = s(bool,X228) | s(bool,X227) = s(bool,sK7(X224,X225,X226,bool))) ) | $spl8),
  inference(superposition,[],[f185,f2533])).
fof(f5431,plain,(
  ( ! [X222,X220,X223,X221] : (p(s(bool,happ(s(fun(bool,bool),X221),s(bool,'T_0')))) | ~p(s(bool,happ(s(fun(bool,bool),X221),s(bool,X220)))) | s(bool,X222) = s(bool,X223) | s(bool,'F_0') = s(bool,X223) | s(bool,X222) = s(bool,sK6(X220,X221,bool))) ) | $spl8),
  inference(superposition,[],[f180,f2533])).
fof(f5430,plain,(
  ( ! [X218,X216,X215,X219,X217] : (s(bool,'T_0') != s(bool,f(s(X215,sK7(X216,X217,bool,X215)))) | s(fun(X215,bool),X216) = s(fun(X215,bool),X217) | s(bool,X218) = s(bool,X219) | s(bool,'F_0') = s(bool,X219) | s(bool,X218) = s(bool,g(s(X215,sK7(X216,X217,bool,X215))))) ) | $spl8),
  inference(superposition,[],[f185,f2533])).
fof(f5429,plain,(
  ( ! [X210,X208,X211,X209] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X209),s(bool,'T_0'))) | s(bool,X208) = s(bool,sK6(X208,X209,bool)) | s(bool,X210) = s(bool,X211) | s(bool,'F_0') = s(bool,X211) | s(bool,X208) = s(bool,X210)) ) | $spl8),
  inference(superposition,[],[f2667,f2533])).
fof(f5428,plain,(
  ( ! [X206,X204,X207,X205] : (~p(s(bool,happ(s(fun(bool,bool),X205),s(bool,'T_0')))) | s(bool,X204) = s(bool,sK6(X204,X205,bool)) | s(bool,X206) = s(bool,X207) | s(bool,'F_0') = s(bool,X207) | s(bool,X204) = s(bool,X206)) ) | $spl8),
  inference(superposition,[],[f179,f2533])).
fof(f5427,plain,(
  ( ! [X202,X200,X199,X203,X201] : (~p(s(bool,happ(s(fun(bool,bool),X200),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X200),s(bool,X201)))) | s(bool,X199) != s(bool,X201) | s(bool,X202) = s(bool,X203) | s(bool,'F_0') = s(bool,X203) | s(bool,X199) = s(bool,X202)) ) | $spl8),
  inference(superposition,[],[f178,f2533])).
fof(f5425,plain,(
  ( ! [X191,X194,X192,X193] : (s(bool,'T_0') != s(bool,X192) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X191))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X192))) | s(bool,X193) = s(bool,X194) | s(bool,'F_0') = s(bool,X194) | s(bool,X191) = s(bool,X193)) ) | $spl8),
  inference(superposition,[],[f177,f2533])).
fof(f5423,plain,(
  ( ! [X185,X186,X184,X183] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X184))) | s(bool,X183) = s(bool,X184) | s(bool,X185) = s(bool,X186) | s(bool,'F_0') = s(bool,X186) | s(bool,X183) = s(bool,X185)) ) | $spl8),
  inference(superposition,[],[f176,f2533])).
fof(f6046,plain,(
  ( ! [X167,X165,X166,X168] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,X165))),s(sK1,X166)))) | s(bool,X167) = s(bool,X168) | s(bool,'F_0') = s(bool,X168) | s(bool,X167) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK5),s(sK0,X165))),s(sK1,X166)))) ) | $spl8),
  inference(subsumption_resolution,[],[f5417,f126])).
fof(f5417,plain,(
  ( ! [X167,X165,X166,X168] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,X165))),s(sK1,X166)))) | s(bool,X167) = s(bool,X168) | s(bool,'F_0') = s(bool,X168) | s(bool,X167) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK5),s(sK0,X165))),s(sK1,X166)))) ) | $spl8),
  inference(superposition,[],[f123,f2533])).
fof(f6045,plain,(
  ( ! [X146,X144,X149,X147,X145,X148] : (p(s(bool,happ(s(fun(X144,bool),X145),s(X144,X147)))) | s(X144,X146) != s(X144,X147) | s(bool,X148) = s(bool,X149) | s(bool,'F_0') = s(bool,X149) | s(bool,X148) = s(bool,happ(s(fun(X144,bool),X145),s(X144,X146)))) ) | $spl8),
  inference(subsumption_resolution,[],[f5413,f126])).
fof(f5413,plain,(
  ( ! [X146,X144,X149,X147,X145,X148] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X144,bool),X145),s(X144,X147)))) | s(X144,X146) != s(X144,X147) | s(bool,X148) = s(bool,X149) | s(bool,'F_0') = s(bool,X149) | s(bool,X148) = s(bool,happ(s(fun(X144,bool),X145),s(X144,X146)))) ) | $spl8),
  inference(superposition,[],[f178,f2533])).
fof(f6040,plain,(
  ( ! [X26,X24,X23,X25] : (~p(s(bool,X24)) | s(bool,X23) = s(bool,X24) | s(bool,X25) = s(bool,X26) | s(bool,'F_0') = s(bool,X26) | s(bool,X23) = s(bool,X25)) ) | $spl8),
  inference(subsumption_resolution,[],[f5380,f126])).
fof(f5380,plain,(
  ( ! [X26,X24,X23,X25] : (~p(s(bool,'T_0')) | ~p(s(bool,X24)) | s(bool,X23) = s(bool,X24) | s(bool,X25) = s(bool,X26) | s(bool,'F_0') = s(bool,X26) | s(bool,X23) = s(bool,X25)) ) | $spl8),
  inference(superposition,[],[f184,f2533])).
fof(f6039,plain,(
  ( ! [X10,X11,X9] : (s(bool,'F_0') != s(bool,X9) | s(bool,X10) = s(bool,X11) | s(bool,'F_0') = s(bool,X11) | s(bool,X9) = s(bool,X10)) ) | $spl8),
  inference(subsumption_resolution,[],[f5376,f126])).
fof(f5376,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X9) | s(bool,X10) = s(bool,X11) | s(bool,'F_0') = s(bool,X11) | s(bool,X9) = s(bool,X10)) ) | $spl8),
  inference(superposition,[],[f135,f2533])).
fof(f5372,plain,(
  ( ! [X134,X132,X133,X131] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X132),s(bool,X131))) | s(bool,'T_0') = s(bool,sK6('T_0',X132,bool)) | s(bool,X133) = s(bool,X134) | s(bool,'F_0') = s(bool,X134) | s(bool,X131) = s(bool,X133)) ) | $spl8),
  inference(superposition,[],[f2667,f2533])).
fof(f5371,plain,(
  ( ! [X127,X130,X128,X129] : (~p(s(bool,happ(s(fun(bool,bool),X128),s(bool,X127)))) | s(bool,'T_0') = s(bool,sK6('T_0',X128,bool)) | s(bool,X129) = s(bool,X130) | s(bool,'F_0') = s(bool,X130) | s(bool,X127) = s(bool,X129)) ) | $spl8),
  inference(superposition,[],[f179,f2533])).
fof(f5370,plain,(
  ( ! [X125,X123,X126,X124,X122] : (~p(s(bool,happ(s(fun(bool,bool),X123),s(bool,X122)))) | p(s(bool,happ(s(fun(bool,bool),X123),s(bool,X124)))) | s(bool,'T_0') != s(bool,X124) | s(bool,X125) = s(bool,X126) | s(bool,'F_0') = s(bool,X126) | s(bool,X122) = s(bool,X125)) ) | $spl8),
  inference(superposition,[],[f178,f2533])).
fof(f5368,plain,(
  ( ! [X116,X114,X117,X115] : (s(bool,X114) != s(bool,X115) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X115))) | s(bool,X116) = s(bool,X117) | s(bool,'F_0') = s(bool,X117) | s(bool,X114) = s(bool,X116)) ) | $spl8),
  inference(superposition,[],[f177,f2533])).
fof(f5344,plain,(
  ( ! [X28,X26,X24,X27,X25] : (s(bool,X24) != s(bool,X25) | s(bool,X25) = s(bool,X26) | s(bool,'F_0') = s(bool,X26) | s(bool,X27) = s(bool,X28) | s(bool,'F_0') = s(bool,X28) | s(bool,X24) = s(bool,X27)) ) | $spl8),
  inference(superposition,[],[f3108,f2533])).
fof(f5343,plain,(
  ( ! [X23,X21,X22,X20] : (s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X20))))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X21))))) | s(bool,'T_0') = s(bool,X21) | s(bool,X22) = s(bool,X23) | s(bool,'F_0') = s(bool,X23) | s(bool,X20) = s(bool,X22)) ) | $spl8),
  inference(superposition,[],[f806,f2533])).
fof(f5340,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,X3) != s(bool,X4) | p(s(bool,X4)) | s(bool,X5) = s(bool,X6) | s(bool,'F_0') = s(bool,X6) | s(bool,X3) = s(bool,X5)) ) | $spl8),
  inference(superposition,[],[f131,f2533])).
fof(f5338,plain,(
  ( ! [X243,X241,X239,X242,X240] : (s(X241,f(s(bool,X242))) != s(X241,g(s(bool,X242))) | s(fun(bool,X241),X239) = s(fun(bool,X241),X240) | s(bool,X242) = s(bool,X243) | s(bool,'F_0') = s(bool,X243) | s(bool,'T_0') = s(bool,sK7(X239,X240,X241,bool))) ) | $spl8),
  inference(superposition,[],[f185,f2533])).
fof(f5337,plain,(
  ( ! [X237,X235,X238,X236] : (p(s(bool,happ(s(fun(bool,bool),X236),s(bool,X237)))) | ~p(s(bool,happ(s(fun(bool,bool),X236),s(bool,X235)))) | s(bool,X237) = s(bool,X238) | s(bool,'F_0') = s(bool,X238) | s(bool,'T_0') = s(bool,sK6(X235,X236,bool))) ) | $spl8),
  inference(superposition,[],[f180,f2533])).
fof(f5336,plain,(
  ( ! [X231,X233,X230,X234,X232] : (s(bool,X233) != s(bool,f(s(X230,sK7(X231,X232,bool,X230)))) | s(fun(X230,bool),X231) = s(fun(X230,bool),X232) | s(bool,X233) = s(bool,X234) | s(bool,'F_0') = s(bool,X234) | s(bool,'T_0') = s(bool,g(s(X230,sK7(X231,X232,bool,X230))))) ) | $spl8),
  inference(superposition,[],[f185,f2533])).
fof(f5335,plain,(
  ( ! [X222,X225,X223,X224] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X224),s(bool,X223))) | s(bool,X222) = s(bool,sK6(X222,X224,bool)) | s(bool,X223) = s(bool,X225) | s(bool,'F_0') = s(bool,X225) | s(bool,'T_0') = s(bool,X222)) ) | $spl8),
  inference(superposition,[],[f2667,f2533])).
fof(f5333,plain,(
  ( ! [X216,X215,X213,X217,X214] : (~p(s(bool,happ(s(fun(bool,bool),X215),s(bool,X214)))) | p(s(bool,happ(s(fun(bool,bool),X215),s(bool,X216)))) | s(bool,X213) != s(bool,X216) | s(bool,X214) = s(bool,X217) | s(bool,'F_0') = s(bool,X217) | s(bool,'T_0') = s(bool,X213)) ) | $spl8),
  inference(superposition,[],[f178,f2533])).
fof(f5331,plain,(
  ( ! [X208,X206,X207,X205] : (s(bool,X206) != s(bool,X207) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X205))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X207))) | s(bool,X206) = s(bool,X208) | s(bool,'F_0') = s(bool,X208) | s(bool,'T_0') = s(bool,X205)) ) | $spl8),
  inference(superposition,[],[f177,f2533])).
fof(f5274,plain,(
  ( ! [X243,X241,X239,X242,X240] : (s(X241,f(s(bool,X242))) != s(X241,g(s(bool,X242))) | s(fun(bool,X241),X239) = s(fun(bool,X241),X240) | s(bool,X243) = s(bool,sK7(X239,X240,X241,bool)) | s(bool,'F_0') = s(bool,X243) | s(bool,'T_0') = s(bool,X242)) ) | $spl8),
  inference(superposition,[],[f185,f2533])).
fof(f5273,plain,(
  ( ! [X237,X235,X238,X236] : (p(s(bool,happ(s(fun(bool,bool),X236),s(bool,X237)))) | ~p(s(bool,happ(s(fun(bool,bool),X236),s(bool,X235)))) | s(bool,X238) = s(bool,sK6(X235,X236,bool)) | s(bool,'F_0') = s(bool,X238) | s(bool,'T_0') = s(bool,X237)) ) | $spl8),
  inference(superposition,[],[f180,f2533])).
fof(f5272,plain,(
  ( ! [X231,X233,X230,X234,X232] : (s(bool,X233) != s(bool,f(s(X230,sK7(X231,X232,bool,X230)))) | s(fun(X230,bool),X231) = s(fun(X230,bool),X232) | s(bool,X234) = s(bool,g(s(X230,sK7(X231,X232,bool,X230)))) | s(bool,'F_0') = s(bool,X234) | s(bool,'T_0') = s(bool,X233)) ) | $spl8),
  inference(superposition,[],[f185,f2533])).
fof(f5271,plain,(
  ( ! [X222,X225,X223,X224] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X224),s(bool,X223))) | s(bool,X222) = s(bool,sK6(X222,X224,bool)) | s(bool,X222) = s(bool,X225) | s(bool,'F_0') = s(bool,X225) | s(bool,'T_0') = s(bool,X223)) ) | $spl8),
  inference(superposition,[],[f2667,f2533])).
fof(f5270,plain,(
  ( ! [X218,X220,X219,X221] : (~p(s(bool,happ(s(fun(bool,bool),X220),s(bool,X219)))) | s(bool,X218) = s(bool,sK6(X218,X220,bool)) | s(bool,X218) = s(bool,X221) | s(bool,'F_0') = s(bool,X221) | s(bool,'T_0') = s(bool,X219)) ) | $spl8),
  inference(superposition,[],[f179,f2533])).
fof(f5269,plain,(
  ( ! [X216,X215,X213,X217,X214] : (~p(s(bool,happ(s(fun(bool,bool),X215),s(bool,X214)))) | p(s(bool,happ(s(fun(bool,bool),X215),s(bool,X216)))) | s(bool,X213) != s(bool,X216) | s(bool,X213) = s(bool,X217) | s(bool,'F_0') = s(bool,X217) | s(bool,'T_0') = s(bool,X214)) ) | $spl8),
  inference(superposition,[],[f178,f2533])).
fof(f5267,plain,(
  ( ! [X208,X206,X207,X205] : (s(bool,X206) != s(bool,X207) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X205))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X207))) | s(bool,X205) = s(bool,X208) | s(bool,'F_0') = s(bool,X208) | s(bool,'T_0') = s(bool,X206)) ) | $spl8),
  inference(superposition,[],[f177,f2533])).
fof(f5265,plain,(
  ( ! [X198,X200,X199,X197] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X198))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X199))) | s(bool,X197) = s(bool,X199) | s(bool,X197) = s(bool,X200) | s(bool,'F_0') = s(bool,X200) | s(bool,'T_0') = s(bool,X198)) ) | $spl8),
  inference(superposition,[],[f176,f2533])).
fof(f5210,plain,(
  ( ! [X206,X204,X202,X205,X203] : (s(X204,f(s(bool,'F_0'))) != s(X204,g(s(bool,'F_0'))) | s(fun(bool,X204),X202) = s(fun(bool,X204),X203) | s(bool,X205) = s(bool,sK7(X202,X203,X204,bool)) | s(bool,X205) = s(bool,X206) | s(bool,'T_0') = s(bool,X206)) ) | $spl8),
  inference(superposition,[],[f185,f2533])).
fof(f5209,plain,(
  ( ! [X198,X200,X199,X201] : (p(s(bool,happ(s(fun(bool,bool),X199),s(bool,'F_0')))) | ~p(s(bool,happ(s(fun(bool,bool),X199),s(bool,X198)))) | s(bool,X200) = s(bool,sK6(X198,X199,bool)) | s(bool,X200) = s(bool,X201) | s(bool,'T_0') = s(bool,X201)) ) | $spl8),
  inference(superposition,[],[f180,f2533])).
fof(f5208,plain,(
  ( ! [X196,X194,X197,X195,X193] : (s(bool,'F_0') != s(bool,f(s(X193,sK7(X194,X195,bool,X193)))) | s(fun(X193,bool),X194) = s(fun(X193,bool),X195) | s(bool,X196) = s(bool,g(s(X193,sK7(X194,X195,bool,X193)))) | s(bool,X196) = s(bool,X197) | s(bool,'T_0') = s(bool,X197)) ) | $spl8),
  inference(superposition,[],[f185,f2533])).
fof(f5207,plain,(
  ( ! [X189,X187,X188,X186] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X187),s(bool,'F_0'))) | s(bool,X186) = s(bool,sK6(X186,X187,bool)) | s(bool,X186) = s(bool,X188) | s(bool,X188) = s(bool,X189) | s(bool,'T_0') = s(bool,X189)) ) | $spl8),
  inference(superposition,[],[f2667,f2533])).
fof(f5206,plain,(
  ( ! [X185,X182,X184,X183] : (~p(s(bool,happ(s(fun(bool,bool),X183),s(bool,'F_0')))) | s(bool,X182) = s(bool,sK6(X182,X183,bool)) | s(bool,X182) = s(bool,X184) | s(bool,X184) = s(bool,X185) | s(bool,'T_0') = s(bool,X185)) ) | $spl8),
  inference(superposition,[],[f179,f2533])).
fof(f5205,plain,(
  ( ! [X177,X180,X178,X181,X179] : (~p(s(bool,happ(s(fun(bool,bool),X178),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X178),s(bool,X179)))) | s(bool,X177) != s(bool,X179) | s(bool,X177) = s(bool,X180) | s(bool,X180) = s(bool,X181) | s(bool,'T_0') = s(bool,X181)) ) | $spl8),
  inference(superposition,[],[f178,f2533])).
fof(f5203,plain,(
  ( ! [X171,X169,X172,X170] : (s(bool,'F_0') != s(bool,X170) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X169))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X170))) | s(bool,X169) = s(bool,X171) | s(bool,X171) = s(bool,X172) | s(bool,'T_0') = s(bool,X172)) ) | $spl8),
  inference(superposition,[],[f177,f2533])).
fof(f5201,plain,(
  ( ! [X163,X161,X164,X162] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X162))) | s(bool,X161) = s(bool,X162) | s(bool,X161) = s(bool,X163) | s(bool,X163) = s(bool,X164) | s(bool,'T_0') = s(bool,X164)) ) | $spl8),
  inference(superposition,[],[f176,f2533])).
fof(f5974,plain,(
  ( ! [X140,X138,X136,X139,X137] : (~p(s(bool,happ(s(fun(X136,bool),X137),s(X136,X138)))) | s(bool,X139) = s(bool,happ(s(fun(X136,bool),X137),s(X136,sK6(X138,X137,X136)))) | s(bool,X139) = s(bool,X140) | s(bool,'T_0') = s(bool,X140)) ) | $spl8),
  inference(subsumption_resolution,[],[f5193,f128])).
fof(f5193,plain,(
  ( ! [X140,X138,X136,X139,X137] : (p(s(bool,'F_0')) | ~p(s(bool,happ(s(fun(X136,bool),X137),s(X136,X138)))) | s(bool,X139) = s(bool,happ(s(fun(X136,bool),X137),s(X136,sK6(X138,X137,X136)))) | s(bool,X139) = s(bool,X140) | s(bool,'T_0') = s(bool,X140)) ) | $spl8),
  inference(superposition,[],[f180,f2533])).
fof(f5973,plain,(
  ( ! [X85,X83,X84,X82] : (s(bool,X82) != s(bool,X83) | s(bool,'F_0') = s(bool,X83) | s(bool,X82) = s(bool,X84) | s(bool,X84) = s(bool,X85) | s(bool,'T_0') = s(bool,X85)) ) | $spl8),
  inference(subsumption_resolution,[],[f5179,f128])).
fof(f5179,plain,(
  ( ! [X85,X83,X84,X82] : (p(s(bool,'F_0')) | s(bool,X82) != s(bool,X83) | s(bool,'F_0') = s(bool,X83) | s(bool,X82) = s(bool,X84) | s(bool,X84) = s(bool,X85) | s(bool,'T_0') = s(bool,X85)) ) | $spl8),
  inference(superposition,[],[f2646,f2533])).
fof(f5972,plain,(
  ( ! [X14,X15,X13,X16] : (p(s(bool,X14)) | s(bool,X13) = s(bool,X14) | s(bool,X13) = s(bool,X15) | s(bool,X15) = s(bool,X16) | s(bool,'T_0') = s(bool,X16)) ) | $spl8),
  inference(subsumption_resolution,[],[f5162,f128])).
fof(f5162,plain,(
  ( ! [X14,X15,X13,X16] : (p(s(bool,'F_0')) | p(s(bool,X14)) | s(bool,X13) = s(bool,X14) | s(bool,X13) = s(bool,X15) | s(bool,X15) = s(bool,X16) | s(bool,'T_0') = s(bool,X16)) ) | $spl8),
  inference(superposition,[],[f181,f2533])).
fof(f5157,plain,(
  ( ! [X154,X152,X151,X153] : (s(bool,X151) = s(bool,happ(s(fun(bool,bool),X152),s(bool,X151))) | s(bool,'F_0') = s(bool,sK6('F_0',X152,bool)) | s(bool,X151) = s(bool,X153) | s(bool,X153) = s(bool,X154) | s(bool,'T_0') = s(bool,X154)) ) | $spl8),
  inference(superposition,[],[f2667,f2533])).
fof(f5156,plain,(
  ( ! [X149,X147,X150,X148] : (~p(s(bool,happ(s(fun(bool,bool),X148),s(bool,X147)))) | s(bool,'F_0') = s(bool,sK6('F_0',X148,bool)) | s(bool,X147) = s(bool,X149) | s(bool,X149) = s(bool,X150) | s(bool,'T_0') = s(bool,X150)) ) | $spl8),
  inference(superposition,[],[f179,f2533])).
fof(f5155,plain,(
  ( ! [X146,X144,X142,X145,X143] : (~p(s(bool,happ(s(fun(bool,bool),X143),s(bool,X142)))) | p(s(bool,happ(s(fun(bool,bool),X143),s(bool,X144)))) | s(bool,'F_0') != s(bool,X144) | s(bool,X142) = s(bool,X145) | s(bool,X145) = s(bool,X146) | s(bool,'T_0') = s(bool,X146)) ) | $spl8),
  inference(superposition,[],[f178,f2533])).
fof(f5153,plain,(
  ( ! [X134,X136,X135,X137] : (s(bool,X134) != s(bool,X135) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X135))) | s(bool,X134) = s(bool,X136) | s(bool,X136) = s(bool,X137) | s(bool,'T_0') = s(bool,X137)) ) | $spl8),
  inference(superposition,[],[f177,f2533])).
fof(f5123,plain,(
  ( ! [X24,X23,X25] : (s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X23))))) | s(bool,X23) = s(bool,X24) | s(bool,X24) = s(bool,X25) | s(bool,'T_0') = s(bool,X25)) ) | $spl8),
  inference(superposition,[],[f795,f2533])).
fof(f5122,plain,(
  ( ! [X21,X19,X22,X20,X18] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X18))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X19))) | s(bool,X19) = s(bool,X20) | s(bool,'T_0') = s(bool,X20) | s(bool,X18) = s(bool,X21) | s(bool,X21) = s(bool,X22) | s(bool,'T_0') = s(bool,X22)) ) | $spl8),
  inference(superposition,[],[f792,f2533])).
fof(f5120,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s(bool,X10) != s(bool,X11) | s(bool,X11) = s(bool,X12) | s(bool,'T_0') = s(bool,X12) | s(bool,X10) = s(bool,X13) | s(bool,X13) = s(bool,X14) | s(bool,'T_0') = s(bool,X14)) ) | $spl8),
  inference(superposition,[],[f714,f2533])).
fof(f5119,plain,(
  ( ! [X8,X7,X9] : (s(bool,'T_0') != s(bool,X7) | s(bool,X7) = s(bool,X8) | s(bool,X8) = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) ) | $spl8),
  inference(superposition,[],[f223,f2533])).
fof(f5116,plain,(
  ( ! [X231,X233,X230,X234,X232] : (s(X232,f(s(bool,X233))) != s(X232,g(s(bool,X233))) | s(fun(bool,X232),X230) = s(fun(bool,X232),X231) | s(bool,'F_0') = s(bool,sK7(X230,X231,X232,bool)) | s(bool,X233) = s(bool,X234) | s(bool,'T_0') = s(bool,X234)) ) | $spl8),
  inference(superposition,[],[f185,f2533])).
fof(f5115,plain,(
  ( ! [X229,X227,X228,X226] : (p(s(bool,happ(s(fun(bool,bool),X227),s(bool,X228)))) | ~p(s(bool,happ(s(fun(bool,bool),X227),s(bool,X226)))) | s(bool,'F_0') = s(bool,sK6(X226,X227,bool)) | s(bool,X228) = s(bool,X229) | s(bool,'T_0') = s(bool,X229)) ) | $spl8),
  inference(superposition,[],[f180,f2533])).
fof(f5114,plain,(
  ( ! [X222,X225,X223,X221,X224] : (s(bool,X224) != s(bool,f(s(X221,sK7(X222,X223,bool,X221)))) | s(fun(X221,bool),X222) = s(fun(X221,bool),X223) | s(bool,'F_0') = s(bool,g(s(X221,sK7(X222,X223,bool,X221)))) | s(bool,X224) = s(bool,X225) | s(bool,'T_0') = s(bool,X225)) ) | $spl8),
  inference(superposition,[],[f185,f2533])).
fof(f5113,plain,(
  ( ! [X216,X215,X213,X214] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X215),s(bool,X214))) | s(bool,X213) = s(bool,sK6(X213,X215,bool)) | s(bool,'F_0') = s(bool,X213) | s(bool,X214) = s(bool,X216) | s(bool,'T_0') = s(bool,X216)) ) | $spl8),
  inference(superposition,[],[f2667,f2533])).
fof(f5112,plain,(
  ( ! [X212,X210,X211,X209] : (~p(s(bool,happ(s(fun(bool,bool),X211),s(bool,X210)))) | s(bool,X209) = s(bool,sK6(X209,X211,bool)) | s(bool,'F_0') = s(bool,X209) | s(bool,X210) = s(bool,X212) | s(bool,'T_0') = s(bool,X212)) ) | $spl8),
  inference(superposition,[],[f179,f2533])).
fof(f5111,plain,(
  ( ! [X208,X206,X204,X207,X205] : (~p(s(bool,happ(s(fun(bool,bool),X206),s(bool,X205)))) | p(s(bool,happ(s(fun(bool,bool),X206),s(bool,X207)))) | s(bool,X204) != s(bool,X207) | s(bool,'F_0') = s(bool,X204) | s(bool,X205) = s(bool,X208) | s(bool,'T_0') = s(bool,X208)) ) | $spl8),
  inference(superposition,[],[f178,f2533])).
fof(f5109,plain,(
  ( ! [X198,X196,X199,X197] : (s(bool,X197) != s(bool,X198) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X196))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X198))) | s(bool,'F_0') = s(bool,X196) | s(bool,X197) = s(bool,X199) | s(bool,'T_0') = s(bool,X199)) ) | $spl8),
  inference(superposition,[],[f177,f2533])).
fof(f5052,plain,(
  ( ! [X243,X241,X239,X242,X240] : (s(X241,f(s(bool,X242))) != s(X241,g(s(bool,X242))) | s(fun(bool,X241),X239) = s(fun(bool,X241),X240) | s(bool,'F_0') = s(bool,X242) | s(bool,X243) = s(bool,sK7(X239,X240,X241,bool)) | s(bool,'T_0') = s(bool,X243)) ) | $spl8),
  inference(superposition,[],[f185,f2533])).
fof(f5051,plain,(
  ( ! [X237,X235,X238,X236] : (p(s(bool,happ(s(fun(bool,bool),X236),s(bool,X237)))) | ~p(s(bool,happ(s(fun(bool,bool),X236),s(bool,X235)))) | s(bool,'F_0') = s(bool,X237) | s(bool,X238) = s(bool,sK6(X235,X236,bool)) | s(bool,'T_0') = s(bool,X238)) ) | $spl8),
  inference(superposition,[],[f180,f2533])).
fof(f5050,plain,(
  ( ! [X231,X233,X230,X234,X232] : (s(bool,X233) != s(bool,f(s(X230,sK7(X231,X232,bool,X230)))) | s(fun(X230,bool),X231) = s(fun(X230,bool),X232) | s(bool,'F_0') = s(bool,X233) | s(bool,X234) = s(bool,g(s(X230,sK7(X231,X232,bool,X230)))) | s(bool,'T_0') = s(bool,X234)) ) | $spl8),
  inference(superposition,[],[f185,f2533])).
fof(f5049,plain,(
  ( ! [X222,X225,X223,X224] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X224),s(bool,X223))) | s(bool,X222) = s(bool,sK6(X222,X224,bool)) | s(bool,'F_0') = s(bool,X223) | s(bool,X222) = s(bool,X225) | s(bool,'T_0') = s(bool,X225)) ) | $spl8),
  inference(superposition,[],[f2667,f2533])).
fof(f5048,plain,(
  ( ! [X218,X220,X219,X221] : (~p(s(bool,happ(s(fun(bool,bool),X220),s(bool,X219)))) | s(bool,X218) = s(bool,sK6(X218,X220,bool)) | s(bool,'F_0') = s(bool,X219) | s(bool,X218) = s(bool,X221) | s(bool,'T_0') = s(bool,X221)) ) | $spl8),
  inference(superposition,[],[f179,f2533])).
fof(f5047,plain,(
  ( ! [X216,X215,X213,X217,X214] : (~p(s(bool,happ(s(fun(bool,bool),X215),s(bool,X214)))) | p(s(bool,happ(s(fun(bool,bool),X215),s(bool,X216)))) | s(bool,X213) != s(bool,X216) | s(bool,'F_0') = s(bool,X214) | s(bool,X213) = s(bool,X217) | s(bool,'T_0') = s(bool,X217)) ) | $spl8),
  inference(superposition,[],[f178,f2533])).
fof(f5045,plain,(
  ( ! [X208,X206,X207,X205] : (s(bool,X206) != s(bool,X207) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X205))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X207))) | s(bool,'F_0') = s(bool,X206) | s(bool,X205) = s(bool,X208) | s(bool,'T_0') = s(bool,X208)) ) | $spl8),
  inference(superposition,[],[f177,f2533])).
fof(f5043,plain,(
  ( ! [X198,X200,X199,X197] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X198))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X199))) | s(bool,X197) = s(bool,X199) | s(bool,'F_0') = s(bool,X198) | s(bool,X197) = s(bool,X200) | s(bool,'T_0') = s(bool,X200)) ) | $spl8),
  inference(superposition,[],[f176,f2533])).
fof(f5655,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X0))))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X1))))) | s(bool,'F_0') = s(bool,X1)) ) | $spl8),
  inference(condensation,[],[f5654])).
fof(f5654,plain,(
  ( ! [X121,X120,X119] : (s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X119))))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X120))))) | s(bool,'T_0') = s(bool,X120) | s(bool,'F_0') = s(bool,X119) | s(bool,'T_0') = s(bool,X121)) ) | $spl8),
  inference(duplicate_literal_removal,[],[f5020])).
fof(f5020,plain,(
  ( ! [X121,X120,X119] : (s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X119))))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X120))))) | s(bool,'T_0') = s(bool,X120) | s(bool,'F_0') = s(bool,X119) | s(bool,'T_0') = s(bool,X121) | s(bool,'T_0') = s(bool,X121)) ) | $spl8),
  inference(superposition,[],[f806,f2533])).
fof(f4957,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X2) = s(bool,X3) | s(bool,'F_0') = s(bool,X3) | s(bool,X1) = s(bool,X2) | s(bool,X4) = s(bool,X5) | s(bool,'F_0') = s(bool,X5) | s(bool,X0) = s(bool,X4)) ) | $spl8),
  inference(superposition,[],[f2533,f2533])).
fof(f4873,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X3) | s(bool,X0) = s(bool,X4) | s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) ) | $spl8),
  inference(superposition,[],[f2533,f2533])).
fof(f5869,plain,(
  ( ! [X68,X66,X69,X67,X65] : (s(bool,X65) = s(bool,X66) | s(bool,X66) = s(bool,X67) | s(bool,'T_0') = s(bool,X67) | s(bool,X68) = s(bool,X69) | s(bool,'F_0') = s(bool,X69) | s(bool,X65) = s(bool,X68)) ) | $spl8),
  inference(subsumption_resolution,[],[f4862,f223])).
fof(f4862,plain,(
  ( ! [X68,X66,X69,X67,X65] : (s(bool,X65) = s(bool,X66) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X66) = s(bool,X67) | s(bool,'T_0') = s(bool,X67) | s(bool,X68) = s(bool,X69) | s(bool,'F_0') = s(bool,X69) | s(bool,X65) = s(bool,X68)) ) | $spl8),
  inference(superposition,[],[f2533,f2533])).
fof(f4744,plain,(
  ( ! [X57,X54,X58,X56,X55] : (s(bool,X57) = s(bool,happ(s(fun(X54,bool),X55),s(X54,X56))) | s(bool,X57) = s(bool,X58) | s(bool,'F_0') = s(bool,X58) | s(X54,X56) = s(X54,sK6(X56,X55,X54))) ) | $spl8),
  inference(subsumption_resolution,[],[f4729,f128])).
fof(f4729,plain,(
  ( ! [X57,X54,X58,X56,X55] : (p(s(bool,'F_0')) | s(bool,X57) = s(bool,happ(s(fun(X54,bool),X55),s(X54,X56))) | s(bool,X57) = s(bool,X58) | s(bool,'F_0') = s(bool,X58) | s(X54,X56) = s(X54,sK6(X56,X55,X54))) ) | $spl8),
  inference(superposition,[],[f3500,f2667])).
fof(f2858,plain,(
  ( ! [X39,X38,X40] : (s(bool,'T_0') != s(bool,f(s(X38,sK7(X39,X40,bool,X38)))) | s(fun(X38,bool),X39) = s(fun(X38,bool),X40) | s(bool,'F_0') = s(bool,g(s(X38,sK7(X39,X40,bool,X38))))) ) | $spl8),
  inference(forward_demodulation,[],[f2850,f227])).
fof(f2850,plain,(
  ( ! [X39,X38,X40] : (s(bool,f(s(X38,sK7(X39,X40,bool,X38)))) != s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(fun(X38,bool),X39) = s(fun(X38,bool),X40) | s(bool,'F_0') = s(bool,g(s(X38,sK7(X39,X40,bool,X38))))) ) | $spl8),
  inference(superposition,[],[f185,f1860])).
fof(f3039,plain,(
  ( ! [X165,X163,X164,X162] : (s(bool,X165) != s(bool,f(s(X162,sK7(X163,X164,bool,X162)))) | s(fun(X162,bool),X163) = s(fun(X162,bool),X164) | s(bool,'F_0') = s(bool,X165) | s(bool,'F_0') = s(bool,g(s(X162,sK7(X163,X164,bool,X162))))) ) | $spl8),
  inference(superposition,[],[f185,f2532])).
fof(f3195,plain,(
  ( ! [X142,X140,X141,X139] : (s(bool,'F_0') != s(bool,f(s(X139,sK7(X140,X141,bool,X139)))) | s(fun(X139,bool),X140) = s(fun(X139,bool),X141) | s(bool,X142) = s(bool,g(s(X139,sK7(X140,X141,bool,X139)))) | s(bool,'F_0') = s(bool,X142)) ) | $spl8),
  inference(superposition,[],[f185,f2532])).
fof(f2742,plain,(
  ( ! [X50,X48,X51,X49] : (p(s(bool,happ(s(fun(X48,bool),X49),s(X48,X51)))) | s(X48,X50) != s(X48,X51) | s(bool,'F_0') = s(bool,happ(s(fun(X48,bool),X49),s(X48,X50)))) ) | $spl8),
  inference(subsumption_resolution,[],[f2741,f126])).
fof(f2741,plain,(
  ( ! [X50,X48,X51,X49] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X48,bool),X49),s(X48,X51)))) | s(X48,X50) != s(X48,X51) | s(bool,'F_0') = s(bool,happ(s(fun(X48,bool),X49),s(X48,X50)))) ) | $spl8),
  inference(forward_demodulation,[],[f2724,f227])).
fof(f2724,plain,(
  ( ! [X50,X48,X51,X49] : (~p(s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)))) | p(s(bool,happ(s(fun(X48,bool),X49),s(X48,X51)))) | s(X48,X50) != s(X48,X51) | s(bool,'F_0') = s(bool,happ(s(fun(X48,bool),X49),s(X48,X50)))) ) | $spl8),
  inference(superposition,[],[f178,f1860])).
fof(f3508,plain,(
  ( ! [X99,X97,X100,X98] : (~p(s(bool,happ(s(fun(X97,bool),X98),s(X97,X99)))) | s(bool,X100) = s(bool,happ(s(fun(X97,bool),X98),s(X97,sK6(X99,X98,X97)))) | s(bool,'F_0') = s(bool,X100)) ) | $spl8),
  inference(subsumption_resolution,[],[f3180,f128])).
fof(f3180,plain,(
  ( ! [X99,X97,X100,X98] : (p(s(bool,'F_0')) | ~p(s(bool,happ(s(fun(X97,bool),X98),s(X97,X99)))) | s(bool,X100) = s(bool,happ(s(fun(X97,bool),X98),s(X97,sK6(X99,X98,X97)))) | s(bool,'F_0') = s(bool,X100)) ) | $spl8),
  inference(superposition,[],[f180,f2532])).
fof(f3718,plain,(
  ( ! [X61,X59,X62,X60] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X62),s(bool,'F_0'))) | s(bool,happ(s(fun(X59,bool),X60),s(X59,X61))) = s(bool,sK6(happ(s(fun(X59,bool),X60),s(X59,X61)),X62,bool)) | s(X59,X61) = s(X59,sK6(X61,X60,X59))) ) | $spl8),
  inference(superposition,[],[f2667,f2667])).
fof(f3775,plain,(
  ( ! [X152,X151,X149,X150] : (s(bool,'F_0') != s(bool,X152) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X152))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(X149,bool),X150),s(X149,X151))))) | s(X149,X151) = s(X149,sK6(X151,X150,X149))) ) | $spl8),
  inference(superposition,[],[f177,f2667])).
fof(f3777,plain,(
  ( ! [X158,X161,X159,X157,X160] : (~p(s(bool,happ(s(fun(bool,bool),X160),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X160),s(bool,X161)))) | s(bool,X161) != s(bool,happ(s(fun(X157,bool),X158),s(X157,X159))) | s(X157,X159) = s(X157,sK6(X159,X158,X157))) ) | $spl8),
  inference(superposition,[],[f178,f2667])).
fof(f3778,plain,(
  ( ! [X165,X163,X164,X162] : (~p(s(bool,happ(s(fun(bool,bool),X165),s(bool,'F_0')))) | s(bool,happ(s(fun(X162,bool),X163),s(X162,X164))) = s(bool,sK6(happ(s(fun(X162,bool),X163),s(X162,X164)),X165,bool)) | s(X162,X164) = s(X162,sK6(X164,X163,X162))) ) | $spl8),
  inference(superposition,[],[f179,f2667])).
fof(f3799,plain,(
  ( ! [X30,X31,X29,X32] : (p(s(bool,X32)) | s(bool,X32) = s(bool,happ(s(fun(X29,bool),X30),s(X29,X31))) | s(X29,X31) = s(X29,sK6(X31,X30,X29))) ) | $spl8),
  inference(subsumption_resolution,[],[f3745,f128])).
fof(f3745,plain,(
  ( ! [X30,X31,X29,X32] : (p(s(bool,'F_0')) | p(s(bool,X32)) | s(bool,X32) = s(bool,happ(s(fun(X29,bool),X30),s(X29,X31))) | s(X29,X31) = s(X29,sK6(X31,X30,X29))) ) | $spl8),
  inference(superposition,[],[f181,f2667])).
fof(f2674,plain,(
  ( ! [X62,X63] : (p(s(bool,happ(s(fun(bool,bool),X63),s(bool,'T_0')))) | ~p(s(bool,happ(s(fun(bool,bool),X63),s(bool,X62)))) | s(bool,'F_0') = s(bool,sK6(X62,X63,bool))) ) | $spl8),
  inference(forward_demodulation,[],[f2582,f227])).
fof(f2582,plain,(
  ( ! [X62,X63] : (p(s(bool,happ(s(fun(bool,bool),X63),s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)))))) | ~p(s(bool,happ(s(fun(bool,bool),X63),s(bool,X62)))) | s(bool,'F_0') = s(bool,sK6(X62,X63,bool))) ) | $spl8),
  inference(superposition,[],[f180,f1860])).
fof(f2730,plain,(
  ( ! [X39,X38,X40] : (~p(s(bool,happ(s(fun(bool,bool),X39),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X39),s(bool,X40)))) | s(bool,X38) != s(bool,X40) | s(bool,'F_0') = s(bool,X38)) ) | $spl8),
  inference(forward_demodulation,[],[f2704,f227])).
fof(f2704,plain,(
  ( ! [X39,X38,X40] : (~p(s(bool,happ(s(fun(bool,bool),X39),s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)))))) | p(s(bool,happ(s(fun(bool,bool),X39),s(bool,X40)))) | s(bool,X38) != s(bool,X40) | s(bool,'F_0') = s(bool,X38)) ) | $spl8),
  inference(superposition,[],[f178,f1860])).
fof(f2732,plain,(
  ( ! [X57,X58,X56] : (s(bool,'T_0') != s(bool,X58) | ~p(s(bool,happ(s(fun(bool,bool),X57),s(bool,X56)))) | p(s(bool,happ(s(fun(bool,bool),X57),s(bool,X58)))) | s(bool,'F_0') = s(bool,X56)) ) | $spl8),
  inference(forward_demodulation,[],[f2709,f227])).
fof(f2709,plain,(
  ( ! [X57,X58,X56] : (~p(s(bool,happ(s(fun(bool,bool),X57),s(bool,X56)))) | p(s(bool,happ(s(fun(bool,bool),X57),s(bool,X58)))) | s(bool,X58) != s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,'F_0') = s(bool,X56)) ) | $spl8),
  inference(superposition,[],[f178,f1860])).
fof(f3037,plain,(
  ( ! [X154,X152,X151,X153] : (~p(s(bool,happ(s(fun(bool,bool),X153),s(bool,X152)))) | p(s(bool,happ(s(fun(bool,bool),X153),s(bool,X154)))) | s(bool,X151) != s(bool,X154) | s(bool,'F_0') = s(bool,X152) | s(bool,'F_0') = s(bool,X151)) ) | $spl8),
  inference(superposition,[],[f178,f2532])).
fof(f3038,plain,(
  ( ! [X156,X157,X155] : (~p(s(bool,happ(s(fun(bool,bool),X157),s(bool,X156)))) | s(bool,X155) = s(bool,sK6(X155,X157,bool)) | s(bool,'F_0') = s(bool,X156) | s(bool,'F_0') = s(bool,X155)) ) | $spl8),
  inference(superposition,[],[f179,f2532])).
fof(f3040,plain,(
  ( ! [X167,X166,X168] : (p(s(bool,happ(s(fun(bool,bool),X167),s(bool,X168)))) | ~p(s(bool,happ(s(fun(bool,bool),X167),s(bool,X166)))) | s(bool,'F_0') = s(bool,X168) | s(bool,'F_0') = s(bool,sK6(X166,X167,bool))) ) | $spl8),
  inference(superposition,[],[f180,f2532])).
fof(f3113,plain,(
  ( ! [X21,X19,X22,X20] : (~p(s(bool,happ(s(fun(bool,bool),X20),s(bool,X19)))) | s(bool,X21) = s(bool,sK6(X21,X20,bool)) | s(bool,'T_0') = s(bool,X21) | s(bool,X19) = s(bool,X22) | s(bool,'F_0') = s(bool,X22)) ) | $spl8),
  inference(superposition,[],[f936,f2532])).
fof(f3141,plain,(
  ( ! [X103,X101,X102,X104] : (~p(s(bool,happ(s(fun(bool,bool),X102),s(bool,X101)))) | p(s(bool,happ(s(fun(bool,bool),X102),s(bool,X103)))) | s(bool,'F_0') != s(bool,X103) | s(bool,X101) = s(bool,X104) | s(bool,'F_0') = s(bool,X104)) ) | $spl8),
  inference(superposition,[],[f178,f2532])).
fof(f3142,plain,(
  ( ! [X107,X105,X106] : (~p(s(bool,happ(s(fun(bool,bool),X106),s(bool,X105)))) | s(bool,'F_0') = s(bool,sK6('F_0',X106,bool)) | s(bool,X105) = s(bool,X107) | s(bool,'F_0') = s(bool,X107)) ) | $spl8),
  inference(superposition,[],[f179,f2532])).
fof(f3193,plain,(
  ( ! [X132,X130,X131,X129] : (~p(s(bool,happ(s(fun(bool,bool),X130),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X130),s(bool,X131)))) | s(bool,X129) != s(bool,X131) | s(bool,X129) = s(bool,X132) | s(bool,'F_0') = s(bool,X132)) ) | $spl8),
  inference(superposition,[],[f178,f2532])).
fof(f3196,plain,(
  ( ! [X144,X145,X143] : (p(s(bool,happ(s(fun(bool,bool),X144),s(bool,'F_0')))) | ~p(s(bool,happ(s(fun(bool,bool),X144),s(bool,X143)))) | s(bool,X145) = s(bool,sK6(X143,X144,bool)) | s(bool,'F_0') = s(bool,X145)) ) | $spl8),
  inference(superposition,[],[f180,f2532])).
fof(f3194,plain,(
  ( ! [X134,X135,X133] : (~p(s(bool,happ(s(fun(bool,bool),X134),s(bool,'F_0')))) | s(bool,X133) = s(bool,sK6(X133,X134,bool)) | s(bool,X133) = s(bool,X135) | s(bool,'F_0') = s(bool,X135)) ) | $spl8),
  inference(superposition,[],[f179,f2532])).
fof(f2665,plain,(
  ( ! [X39,X38] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,X38))),s(sK1,X39)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK5),s(sK0,X38))),s(sK1,X39)))) ) | $spl8),
  inference(subsumption_resolution,[],[f2664,f126])).
fof(f2664,plain,(
  ( ! [X39,X38] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,X38))),s(sK1,X39)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK5),s(sK0,X38))),s(sK1,X39)))) ) | $spl8),
  inference(forward_demodulation,[],[f2569,f227])).
fof(f2569,plain,(
  ( ! [X39,X38] : (~p(s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,X38))),s(sK1,X39)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK5),s(sK0,X38))),s(sK1,X39)))) ) | $spl8),
  inference(superposition,[],[f123,f1860])).
fof(f3713,plain,(
  ( ! [X47,X45,X46] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X46),s(bool,X45))) | s(bool,'T_0') = s(bool,sK6('T_0',X46,bool)) | s(bool,X45) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) ) | $spl8),
  inference(superposition,[],[f2667,f259])).
fof(f3717,plain,(
  ( ! [X57,X58] : (s(bool,X57) = s(bool,happ(s(fun(bool,bool),X58),s(bool,X57))) | s(bool,'F_0') = s(bool,sK6('F_0',X58,bool)) | s(bool,'T_0') = s(bool,X57)) ) | $spl8),
  inference(superposition,[],[f2667,f129])).
fof(f3715,plain,(
  ( ! [X52,X53,X51] : (s(bool,X51) = s(bool,happ(s(fun(bool,bool),X52),s(bool,X51))) | s(bool,'F_0') = s(bool,sK6('F_0',X52,bool)) | s(bool,X51) = s(bool,X53) | s(bool,'F_0') = s(bool,X53)) ) | $spl8),
  inference(superposition,[],[f2667,f2532])).
fof(f3838,plain,(
  ( ! [X66,X64,X62,X65,X63] : (s('type.option.option'(X62),'const.option.NONE_0') = s('type.option.option'(X62),X65) | s('type.option.option'(X62),X65) = s('type.option.option'(X62),'const.option.SOME_1'(s(X62,sK8(X66,X65,X64,X63,X62)))) | s(bool,'F_0') = s(bool,'const.option.OPTREL_3'(s(fun(X62,fun(X63,bool)),X64),s('type.option.option'(X62),X65),s('type.option.option'(X63),X66)))) ) | $spl8),
  inference(subsumption_resolution,[],[f3837,f126])).
fof(f3837,plain,(
  ( ! [X66,X64,X62,X65,X63] : (~p(s(bool,'T_0')) | s('type.option.option'(X62),'const.option.NONE_0') = s('type.option.option'(X62),X65) | s('type.option.option'(X62),X65) = s('type.option.option'(X62),'const.option.SOME_1'(s(X62,sK8(X66,X65,X64,X63,X62)))) | s(bool,'F_0') = s(bool,'const.option.OPTREL_3'(s(fun(X62,fun(X63,bool)),X64),s('type.option.option'(X62),X65),s('type.option.option'(X63),X66)))) ) | $spl8),
  inference(forward_demodulation,[],[f3818,f227])).
fof(f3818,plain,(
  ( ! [X66,X64,X62,X65,X63] : (~p(s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)))) | s('type.option.option'(X62),'const.option.NONE_0') = s('type.option.option'(X62),X65) | s('type.option.option'(X62),X65) = s('type.option.option'(X62),'const.option.SOME_1'(s(X62,sK8(X66,X65,X64,X63,X62)))) | s(bool,'F_0') = s(bool,'const.option.OPTREL_3'(s(fun(X62,fun(X63,bool)),X64),s('type.option.option'(X62),X65),s('type.option.option'(X63),X66)))) ) | $spl8),
  inference(superposition,[],[f186,f1860])).
fof(f3773,plain,(
  ( ! [X144,X142,X143,X141] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X144))) | s(bool,X144) = s(bool,happ(s(fun(X141,bool),X142),s(X141,X143))) | s(X141,X143) = s(X141,sK6(X143,X142,X141))) ) | $spl8),
  inference(superposition,[],[f176,f2667])).
fof(f3802,plain,(
  ( ! [X132,X130,X133,X131] : (s(bool,X133) != s(bool,happ(s(fun(X130,bool),X131),s(X130,X132))) | s(bool,'F_0') = s(bool,X133) | s(X130,X132) = s(X130,sK6(X132,X131,X130))) ) | $spl8),
  inference(subsumption_resolution,[],[f3770,f128])).
fof(f3770,plain,(
  ( ! [X132,X130,X133,X131] : (p(s(bool,'F_0')) | s(bool,X133) != s(bool,happ(s(fun(X130,bool),X131),s(X130,X132))) | s(bool,'F_0') = s(bool,X133) | s(X130,X132) = s(X130,sK6(X132,X131,X130))) ) | $spl8),
  inference(superposition,[],[f2646,f2667])).
fof(f3709,plain,(
  ( ! [X35,X33,X34] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X35),s(bool,X34))) | s(bool,X33) = s(bool,sK6(X33,X35,bool)) | s(bool,'F_0') = s(bool,X34) | s(bool,'F_0') = s(bool,X33)) ) | $spl8),
  inference(superposition,[],[f2667,f2532])).
fof(f3708,plain,(
  ( ! [X30,X31,X32] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X31),s(bool,'F_0'))) | s(bool,X30) = s(bool,sK6(X30,X31,bool)) | s(bool,X30) = s(bool,X32) | s(bool,'F_0') = s(bool,X32)) ) | $spl8),
  inference(superposition,[],[f2667,f2532])).
fof(f3792,plain,(
  ( ! [X28,X29] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X29),s(bool,'T_0'))) | s(bool,X28) = s(bool,sK6(X28,X29,bool)) | s(bool,'F_0') = s(bool,X28)) ) | $spl8),
  inference(forward_demodulation,[],[f3707,f227])).
fof(f3707,plain,(
  ( ! [X28,X29] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X29),s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))))) | s(bool,X28) = s(bool,sK6(X28,X29,bool)) | s(bool,'F_0') = s(bool,X28)) ) | $spl8),
  inference(superposition,[],[f2667,f1860])).
fof(f3699,plain,(
  ( ! [X6,X7,X5] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X7),s(bool,X6))) | s(bool,X5) = s(bool,sK6(X5,X7,bool)) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,X5)) ) | $spl8),
  inference(superposition,[],[f2667,f259])).
fof(f3698,plain,(
  ( ! [X4,X2,X3] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X3),s(bool,'T_0'))) | s(bool,X2) = s(bool,sK6(X2,X3,bool)) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) ) | $spl8),
  inference(superposition,[],[f2667,f259])).
fof(f3197,plain,(
  ( ! [X146,X149,X147,X148] : (s(X148,f(s(bool,'F_0'))) != s(X148,g(s(bool,'F_0'))) | s(fun(bool,X148),X146) = s(fun(bool,X148),X147) | s(bool,X149) = s(bool,sK7(X146,X147,X148,bool)) | s(bool,'F_0') = s(bool,X149)) ) | $spl8),
  inference(superposition,[],[f185,f2532])).
fof(f3191,plain,(
  ( ! [X125,X123,X124] : (s(bool,'F_0') != s(bool,X124) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X123))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X124))) | s(bool,X123) = s(bool,X125) | s(bool,'F_0') = s(bool,X125)) ) | $spl8),
  inference(superposition,[],[f177,f2532])).
fof(f3189,plain,(
  ( ! [X118,X119,X117] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X118))) | s(bool,X117) = s(bool,X118) | s(bool,X117) = s(bool,X119) | s(bool,'F_0') = s(bool,X119)) ) | $spl8),
  inference(superposition,[],[f176,f2532])).
fof(f3507,plain,(
  ( ! [X68,X69,X67] : (s(bool,X67) != s(bool,X68) | s(bool,'F_0') = s(bool,X68) | s(bool,X67) = s(bool,X69) | s(bool,'F_0') = s(bool,X69)) ) | $spl8),
  inference(subsumption_resolution,[],[f3168,f128])).
fof(f3168,plain,(
  ( ! [X68,X69,X67] : (p(s(bool,'F_0')) | s(bool,X67) != s(bool,X68) | s(bool,'F_0') = s(bool,X68) | s(bool,X67) = s(bool,X69) | s(bool,'F_0') = s(bool,X69)) ) | $spl8),
  inference(superposition,[],[f2646,f2532])).
fof(f3139,plain,(
  ( ! [X97,X95,X96] : (s(bool,X95) != s(bool,X96) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X96))) | s(bool,X95) = s(bool,X97) | s(bool,'F_0') = s(bool,X97)) ) | $spl8),
  inference(superposition,[],[f177,f2532])).
fof(f3112,plain,(
  ( ! [X17,X18] : (s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X17))))) | s(bool,X17) = s(bool,X18) | s(bool,'F_0') = s(bool,X18)) ) | $spl8),
  inference(superposition,[],[f795,f2532])).
fof(f3111,plain,(
  ( ! [X14,X15,X13,X16] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X13))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X14))) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X15) | s(bool,X13) = s(bool,X16) | s(bool,'F_0') = s(bool,X16)) ) | $spl8),
  inference(superposition,[],[f792,f2532])).
fof(f3109,plain,(
  ( ! [X10,X8,X7,X9] : (s(bool,X7) != s(bool,X8) | s(bool,X8) = s(bool,X9) | s(bool,'T_0') = s(bool,X9) | s(bool,X7) = s(bool,X10) | s(bool,'F_0') = s(bool,X10)) ) | $spl8),
  inference(superposition,[],[f714,f2532])).
fof(f3108,plain,(
  ( ! [X6,X5] : (s(bool,'T_0') != s(bool,X5) | s(bool,X5) = s(bool,X6) | s(bool,'F_0') = s(bool,X6)) ) | $spl8),
  inference(superposition,[],[f223,f2532])).
fof(f3041,plain,(
  ( ! [X171,X169,X172,X170] : (s(X171,f(s(bool,X172))) != s(X171,g(s(bool,X172))) | s(fun(bool,X171),X169) = s(fun(bool,X171),X170) | s(bool,'F_0') = s(bool,X172) | s(bool,'F_0') = s(bool,sK7(X169,X170,X171,bool))) ) | $spl8),
  inference(superposition,[],[f185,f2532])).
fof(f3035,plain,(
  ( ! [X146,X147,X145] : (s(bool,X146) != s(bool,X147) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X145))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X147))) | s(bool,'F_0') = s(bool,X146) | s(bool,'F_0') = s(bool,X145)) ) | $spl8),
  inference(superposition,[],[f177,f2532])).
fof(f3033,plain,(
  ( ! [X140,X141,X139] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X140))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X141))) | s(bool,X139) = s(bool,X141) | s(bool,'F_0') = s(bool,X140) | s(bool,'F_0') = s(bool,X139)) ) | $spl8),
  inference(superposition,[],[f176,f2532])).
fof(f2916,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'F_0') = s(bool,X3)) ) | $spl8),
  inference(superposition,[],[f2532,f2532])).
fof(f2854,plain,(
  ( ! [X39,X38,X40] : (s(X40,f(s(bool,'T_0'))) != s(X40,g(s(bool,'T_0'))) | s(fun(bool,X40),X38) = s(fun(bool,X40),X39) | s(bool,'F_0') = s(bool,sK7(X38,X39,X40,bool))) ) | $spl8),
  inference(forward_demodulation,[],[f2839,f227])).
fof(f2839,plain,(
  ( ! [X39,X38,X40] : (s(X40,f(s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))))) != s(X40,g(s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))))) | s(fun(bool,X40),X38) = s(fun(bool,X40),X39) | s(bool,'F_0') = s(bool,sK7(X38,X39,X40,bool))) ) | $spl8),
  inference(superposition,[],[f185,f1860])).
fof(f2583,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | $spl8),
  inference(superposition,[],[f1739,f1860])).
fof(f2671,plain,(
  ( ! [X52,X53] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X53))) | s(bool,X52) = s(bool,X53) | s(bool,'F_0') = s(bool,X52)) ) | $spl8),
  inference(forward_demodulation,[],[f2577,f227])).
fof(f2577,plain,(
  ( ! [X52,X53] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X53))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))))) | s(bool,X52) = s(bool,X53) | s(bool,'F_0') = s(bool,X52)) ) | $spl8),
  inference(superposition,[],[f176,f1860])).
fof(f2670,plain,(
  ( ! [X47,X46] : (s(bool,X47) = s(bool,sK6(X47,X46,bool)) | s(bool,'T_0') = s(bool,X47) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X46),s(bool,'F_0')))) ) | $spl8),
  inference(subsumption_resolution,[],[f2669,f126])).
fof(f2669,plain,(
  ( ! [X47,X46] : (~p(s(bool,'T_0')) | s(bool,X47) = s(bool,sK6(X47,X46,bool)) | s(bool,'T_0') = s(bool,X47) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X46),s(bool,'F_0')))) ) | $spl8),
  inference(forward_demodulation,[],[f2572,f227])).
fof(f2572,plain,(
  ( ! [X47,X46] : (~p(s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)))) | s(bool,X47) = s(bool,sK6(X47,X46,bool)) | s(bool,'T_0') = s(bool,X47) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X46),s(bool,'F_0')))) ) | $spl8),
  inference(superposition,[],[f936,f1860])).
fof(f2644,plain,(
  ( ! [X10,X9] : (~p(s(bool,X10)) | s(bool,X9) = s(bool,X10) | s(bool,'F_0') = s(bool,X9)) ) | $spl8),
  inference(subsumption_resolution,[],[f2643,f126])).
fof(f2643,plain,(
  ( ! [X10,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,X10)) | s(bool,X9) = s(bool,X10) | s(bool,'F_0') = s(bool,X9)) ) | $spl8),
  inference(forward_demodulation,[],[f2553,f227])).
fof(f2553,plain,(
  ( ! [X10,X9] : (~p(s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)))) | ~p(s(bool,X10)) | s(bool,X9) = s(bool,X10) | s(bool,'F_0') = s(bool,X9)) ) | $spl8),
  inference(superposition,[],[f184,f1860])).
fof(f6886,plain,(
  $false | ($spl8 | $spl21 | $spl32)),
  inference(global_subsumption,[],[f6885,f125,f124,f126,f129,f132,f131,f173,f177,f176,f184,f181,f185,f219,f223,f259,f296,f428,f498,f714,f715,f726,f764,f765,f767,f806,f822,f824,f828,f884,f999,f1002,f1833,f1826,f1861,f1819,f1862,f1810,f1910,f1763,f1711,f128,f135,f136,f773,f790,f792,f795,f805,f827,f830,f954,f940,f1515,f2829,f2830,f2831,f3830,f3832,f3833,f3975,f4201,f4187,f3961,f3625,f192,f189,f188,f187,f186,f123,f716,f1330,f936,f1333,f1312,f918,f1372,f1362,f1360,f1356,f1354,f1352,f1341,f1335,f1320,f1318,f1314,f998,f988,f986,f926,f924,f920,f191,f190,f193,f1376,f1331,f1004,f180,f179,f178,f2842,f2841,f2840,f3836,f2849,f2848,f2838,f2837,f2151,f1995,f1994,f1993,f1981,f1951,f1950,f1949,f1751,f1739,f1705,f1703,f1655,f1637,f1636,f1626,f1572,f1570,f1523,f1489,f1438,f1437,f1344,f1323,f1322,f1321,f982,f956,f939,f938,f937,f826,f776,f775,f774,f335,f232,f227,f224,f3826,f6181,f6206,f6233,f6262,f6268,f6320,f6416,f6505,f6511,f6509,f6532,f6533,f6534,f6535,f6536,f6537,f6538,f6539,f6540,f6541,f6542,f6543,f6544,f6625,f6828,f6770,f6857,f6877,f6879,f6881,f6883])).
fof(f6885,plain,(
  ( ! [X0,X3] : (s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X3)) ) | ($spl8 | $spl21 | $spl32)),
  inference(subsumption_resolution,[],[f6806,f126])).
fof(f6806,plain,(
  ( ! [X0,X3] : (~p(s(bool,'T_0')) | s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X3)) ) | ($spl8 | $spl21 | $spl32)),
  inference(backward_demodulation,[],[f6625,f6617])).
fof(f6617,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X3) | ~p(s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.option.option'(X0),X3),s('type.option.option'(X1),X4))))) ) | ($spl8 | $spl21 | $spl32)),
  inference(subsumption_resolution,[],[f6616,f128])).
fof(f6616,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'F_0')) | s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X3) | ~p(s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.option.option'(X0),X3),s('type.option.option'(X1),X4))))) ) | ($spl8 | $spl21 | $spl32)),
  inference(forward_demodulation,[],[f6600,f6301])).
fof(f6600,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),'const.option.NONE_0')))) | s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X3) | ~p(s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.option.option'(X0),X3),s('type.option.option'(X1),X4))))) ) | ($spl8 | $spl21 | $spl32)),
  inference(backward_demodulation,[],[f6598,f190])).
fof(f6884,plain,(
  $false | ($spl21 | $spl32)),
  inference(global_subsumption,[],[f125,f124,f126,f129,f132,f131,f173,f177,f176,f184,f181,f185,f219,f223,f259,f296,f428,f498,f714,f715,f726,f764,f765,f767,f806,f822,f824,f828,f884,f999,f1002,f1833,f1826,f1861,f1819,f1862,f1810,f1910,f1763,f1711,f128,f135,f136,f773,f790,f792,f795,f805,f827,f830,f954,f940,f1515,f2829,f2830,f2831,f3830,f3832,f3833,f3975,f4201,f4187,f3961,f3625,f192,f189,f188,f187,f186,f123,f716,f1330,f936,f1333,f1312,f918,f1372,f1362,f1360,f1356,f1354,f1352,f1341,f1335,f1320,f1318,f1314,f998,f988,f986,f926,f924,f920,f191,f190,f193,f1376,f1331,f1004,f180,f179,f178,f2842,f2841,f2840,f3836,f2849,f2848,f2838,f2837,f2151,f1995,f1994,f1993,f1981,f1951,f1950,f1949,f1751,f1739,f1705,f1703,f1655,f1637,f1636,f1626,f1572,f1570,f1523,f1489,f1438,f1437,f1344,f1323,f1322,f1321,f982,f956,f939,f938,f937,f826,f776,f775,f774,f335,f232,f227,f224,f3826,f6181,f6206,f6233,f6262,f6268,f6320,f6416,f6505,f6511,f6509,f6532,f6533,f6534,f6535,f6536,f6537,f6538,f6539,f6540,f6541,f6542,f6543,f6544,f6625,f6828,f6770,f6857,f6877,f6879,f6881,f6883])).
fof(f6882,plain,(
  $false | ($spl21 | $spl32)),
  inference(global_subsumption,[],[f125,f124,f126,f129,f132,f131,f173,f177,f176,f184,f181,f185,f219,f223,f259,f296,f428,f498,f714,f715,f726,f764,f765,f767,f806,f822,f824,f828,f884,f999,f1002,f1833,f1826,f1861,f1819,f1862,f1810,f1910,f1763,f1711,f128,f135,f136,f773,f790,f792,f795,f805,f827,f830,f954,f940,f1515,f2829,f2830,f2831,f3830,f3832,f3833,f3975,f4201,f4187,f3961,f3625,f192,f189,f188,f187,f186,f123,f716,f1330,f936,f1333,f1312,f918,f1372,f1362,f1360,f1356,f1354,f1352,f1341,f1335,f1320,f1318,f1314,f998,f988,f986,f926,f924,f920,f191,f190,f193,f1376,f1331,f1004,f180,f179,f178,f2842,f2841,f2840,f3836,f2849,f2848,f2838,f2837,f2151,f1995,f1994,f1993,f1981,f1951,f1950,f1949,f1751,f1739,f1705,f1703,f1655,f1637,f1636,f1626,f1572,f1570,f1523,f1489,f1438,f1437,f1344,f1323,f1322,f1321,f982,f956,f939,f938,f937,f826,f776,f775,f774,f335,f232,f227,f224,f3826,f6181,f6206,f6233,f6262,f6268,f6320,f6416,f6505,f6511,f6509,f6532,f6533,f6534,f6535,f6536,f6537,f6538,f6539,f6540,f6541,f6542,f6543,f6544,f6625,f6828,f6770,f6857,f6877,f6879,f6881])).
fof(f6880,plain,(
  $false | ($spl21 | $spl32)),
  inference(global_subsumption,[],[f125,f124,f126,f129,f132,f131,f173,f177,f176,f184,f181,f185,f219,f223,f259,f296,f428,f498,f714,f715,f726,f764,f765,f767,f806,f822,f824,f828,f884,f999,f1002,f1833,f1826,f1861,f1819,f1862,f1810,f1910,f1763,f1711,f128,f135,f136,f773,f790,f792,f795,f805,f827,f830,f954,f940,f1515,f2829,f2830,f2831,f3830,f3832,f3833,f3975,f4201,f4187,f3961,f3625,f192,f189,f188,f187,f186,f123,f716,f1330,f936,f1333,f1312,f918,f1372,f1362,f1360,f1356,f1354,f1352,f1341,f1335,f1320,f1318,f1314,f998,f988,f986,f926,f924,f920,f191,f190,f193,f1376,f1331,f1004,f180,f179,f178,f2842,f2841,f2840,f3836,f2849,f2848,f2838,f2837,f2151,f1995,f1994,f1993,f1981,f1951,f1950,f1949,f1751,f1739,f1705,f1703,f1655,f1637,f1636,f1626,f1572,f1570,f1523,f1489,f1438,f1437,f1344,f1323,f1322,f1321,f982,f956,f939,f938,f937,f826,f776,f775,f774,f335,f232,f227,f224,f3826,f6181,f6206,f6233,f6262,f6268,f6320,f6416,f6505,f6511,f6509,f6532,f6533,f6534,f6535,f6536,f6537,f6538,f6539,f6540,f6541,f6542,f6543,f6544,f6625,f6828,f6770,f6857,f6877,f6879])).
fof(f6878,plain,(
  $false | ($spl21 | $spl32)),
  inference(global_subsumption,[],[f125,f124,f126,f129,f132,f131,f173,f177,f176,f184,f181,f185,f219,f223,f259,f296,f428,f498,f714,f715,f726,f764,f765,f767,f806,f822,f824,f828,f884,f999,f1002,f1833,f1826,f1861,f1819,f1862,f1810,f1910,f1763,f1711,f128,f135,f136,f773,f790,f792,f795,f805,f827,f830,f954,f940,f1515,f2829,f2830,f2831,f3830,f3832,f3833,f3975,f4201,f4187,f3961,f3625,f192,f189,f188,f187,f186,f123,f716,f1330,f936,f1333,f1312,f918,f1372,f1362,f1360,f1356,f1354,f1352,f1341,f1335,f1320,f1318,f1314,f998,f988,f986,f926,f924,f920,f191,f190,f193,f1376,f1331,f1004,f180,f179,f178,f2842,f2841,f2840,f3836,f2849,f2848,f2838,f2837,f2151,f1995,f1994,f1993,f1981,f1951,f1950,f1949,f1751,f1739,f1705,f1703,f1655,f1637,f1636,f1626,f1572,f1570,f1523,f1489,f1438,f1437,f1344,f1323,f1322,f1321,f982,f956,f939,f938,f937,f826,f776,f775,f774,f335,f232,f227,f224,f3826,f6181,f6206,f6233,f6262,f6268,f6320,f6416,f6505,f6511,f6509,f6532,f6533,f6534,f6535,f6536,f6537,f6538,f6539,f6540,f6541,f6542,f6543,f6544,f6625,f6828,f6770,f6857,f6877])).
fof(f6876,plain,(
  $false | ($spl18 | $spl21 | $spl32)),
  inference(global_subsumption,[],[f6875,f125,f124,f126,f129,f132,f131,f173,f177,f176,f184,f181,f185,f219,f223,f259,f296,f428,f498,f714,f715,f726,f764,f765,f767,f806,f822,f824,f828,f884,f999,f1002,f1833,f1826,f1861,f1819,f1862,f1810,f1910,f1763,f1711,f128,f135,f136,f773,f790,f792,f795,f805,f827,f830,f954,f940,f1515,f2829,f2830,f2831,f3830,f3832,f3833,f3975,f4201,f4187,f3961,f3625,f192,f189,f188,f187,f186,f123,f716,f1330,f936,f1333,f1312,f918,f1372,f1362,f1360,f1356,f1354,f1352,f1341,f1335,f1320,f1318,f1314,f998,f988,f986,f926,f924,f920,f191,f190,f193,f1376,f1331,f1004,f180,f179,f178,f2842,f2841,f2840,f3836,f2849,f2848,f2838,f2837,f2151,f1995,f1994,f1993,f1981,f1951,f1950,f1949,f1751,f1739,f1705,f1703,f1655,f1637,f1636,f1626,f1572,f1570,f1523,f1489,f1438,f1437,f1344,f1323,f1322,f1321,f982,f956,f939,f938,f937,f826,f776,f775,f774,f335,f232,f227,f224,f3826,f6181,f6206,f6233,f6262,f6268,f6320,f6416,f6505,f6511,f6509,f6532,f6533,f6534,f6535,f6536,f6537,f6538,f6539,f6540,f6541,f6542,f6543,f6544,f6625,f6828,f6770,f6857])).
fof(f6875,plain,(
  ( ! [X6,X7,X5] : (s('type.option.option'(sK0),sK3) = s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK8(X7,'const.option.SOME_1'(s(sK0,sK8(sK2,sK3,sK5,sK1,sK0))),X6,X5,sK0))))) ) | ($spl18 | $spl21 | $spl32)),
  inference(subsumption_resolution,[],[f6800,f126])).
fof(f6800,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,'T_0')) | s('type.option.option'(sK0),sK3) = s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK8(X7,'const.option.SOME_1'(s(sK0,sK8(sK2,sK3,sK5,sK1,sK0))),X6,X5,sK0))))) ) | ($spl18 | $spl21 | $spl32)),
  inference(backward_demodulation,[],[f6625,f6199])).
fof(f6199,plain,(
  ( ! [X6,X7,X5] : (s('type.option.option'(sK0),sK3) = s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK8(X7,'const.option.SOME_1'(s(sK0,sK8(sK2,sK3,sK5,sK1,sK0))),X6,X5,sK0)))) | ~p(s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(X5,bool)),X6),s('type.option.option'(sK0),sK3),s('type.option.option'(X5),X7))))) ) | $spl18),
  inference(forward_demodulation,[],[f6198,f3825])).
fof(f6198,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(X5,bool)),X6),s('type.option.option'(sK0),sK3),s('type.option.option'(X5),X7)))) | s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK8(sK2,sK3,sK5,sK1,sK0)))) = s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK8(X7,'const.option.SOME_1'(s(sK0,sK8(sK2,sK3,sK5,sK1,sK0))),X6,X5,sK0))))) ) | $spl18),
  inference(subsumption_resolution,[],[f6188,f173])).
fof(f6188,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(X5,bool)),X6),s('type.option.option'(sK0),sK3),s('type.option.option'(X5),X7)))) | s('type.option.option'(sK0),'const.option.NONE_0') = s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK8(sK2,sK3,sK5,sK1,sK0)))) | s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK8(sK2,sK3,sK5,sK1,sK0)))) = s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK8(X7,'const.option.SOME_1'(s(sK0,sK8(sK2,sK3,sK5,sK1,sK0))),X6,X5,sK0))))) ) | $spl18),
  inference(superposition,[],[f186,f3825])).
fof(f6874,plain,(
  $false | ($spl18 | $spl21 | $spl32)),
  inference(global_subsumption,[],[f6873,f125,f124,f126,f129,f132,f131,f173,f177,f176,f184,f181,f185,f219,f223,f259,f296,f428,f498,f714,f715,f726,f764,f765,f767,f806,f822,f824,f828,f884,f999,f1002,f1833,f1826,f1861,f1819,f1862,f1810,f1910,f1763,f1711,f128,f135,f136,f773,f790,f792,f795,f805,f827,f830,f954,f940,f1515,f2829,f2830,f2831,f3830,f3832,f3833,f3975,f4201,f4187,f3961,f3625,f192,f189,f188,f187,f186,f123,f716,f1330,f936,f1333,f1312,f918,f1372,f1362,f1360,f1356,f1354,f1352,f1341,f1335,f1320,f1318,f1314,f998,f988,f986,f926,f924,f920,f191,f190,f193,f1376,f1331,f1004,f180,f179,f178,f2842,f2841,f2840,f3836,f2849,f2848,f2838,f2837,f2151,f1995,f1994,f1993,f1981,f1951,f1950,f1949,f1751,f1739,f1705,f1703,f1655,f1637,f1636,f1626,f1572,f1570,f1523,f1489,f1438,f1437,f1344,f1323,f1322,f1321,f982,f956,f939,f938,f937,f826,f776,f775,f774,f335,f232,f227,f224,f3826,f6181,f6206,f6233,f6262,f6268,f6320,f6416,f6505,f6511,f6509,f6532,f6533,f6534,f6535,f6536,f6537,f6538,f6539,f6540,f6541,f6542,f6543,f6544,f6625,f6828,f6770,f6857])).
fof(f6873,plain,(
  ( ! [X2,X0,X1] : (s('type.option.option'(X0),X2) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,sK9(X2,'const.option.SOME_1'(s(sK0,sK8(sK2,sK3,sK5,sK1,sK0))),X1,X0,sK0))))) ) | ($spl18 | $spl21 | $spl32)),
  inference(subsumption_resolution,[],[f6799,f126])).
fof(f6799,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s('type.option.option'(X0),X2) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,sK9(X2,'const.option.SOME_1'(s(sK0,sK8(sK2,sK3,sK5,sK1,sK0))),X1,X0,sK0))))) ) | ($spl18 | $spl21 | $spl32)),
  inference(backward_demodulation,[],[f6625,f6286])).
fof(f6286,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(X0,bool)),X1),s('type.option.option'(sK0),sK3),s('type.option.option'(X0),X2)))) | s('type.option.option'(X0),X2) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,sK9(X2,'const.option.SOME_1'(s(sK0,sK8(sK2,sK3,sK5,sK1,sK0))),X1,X0,sK0))))) ) | $spl18),
  inference(subsumption_resolution,[],[f6269,f173])).
fof(f6269,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(X0,bool)),X1),s('type.option.option'(sK0),sK3),s('type.option.option'(X0),X2)))) | s('type.option.option'(sK0),'const.option.NONE_0') = s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK8(sK2,sK3,sK5,sK1,sK0)))) | s('type.option.option'(X0),X2) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,sK9(X2,'const.option.SOME_1'(s(sK0,sK8(sK2,sK3,sK5,sK1,sK0))),X1,X0,sK0))))) ) | $spl18),
  inference(superposition,[],[f188,f3825])).
fof(f6872,plain,(
  $false | ($spl18 | $spl21 | $spl32)),
  inference(global_subsumption,[],[f6871,f125,f124,f126,f129,f132,f131,f173,f177,f176,f184,f181,f185,f219,f223,f259,f296,f428,f498,f714,f715,f726,f764,f765,f767,f806,f822,f824,f828,f884,f999,f1002,f1833,f1826,f1861,f1819,f1862,f1810,f1910,f1763,f1711,f128,f135,f136,f773,f790,f792,f795,f805,f827,f830,f954,f940,f1515,f2829,f2830,f2831,f3830,f3832,f3833,f3975,f4201,f4187,f3961,f3625,f192,f189,f188,f187,f186,f123,f716,f1330,f936,f1333,f1312,f918,f1372,f1362,f1360,f1356,f1354,f1352,f1341,f1335,f1320,f1318,f1314,f998,f988,f986,f926,f924,f920,f191,f190,f193,f1376,f1331,f1004,f180,f179,f178,f2842,f2841,f2840,f3836,f2849,f2848,f2838,f2837,f2151,f1995,f1994,f1993,f1981,f1951,f1950,f1949,f1751,f1739,f1705,f1703,f1655,f1637,f1636,f1626,f1572,f1570,f1523,f1489,f1438,f1437,f1344,f1323,f1322,f1321,f982,f956,f939,f938,f937,f826,f776,f775,f774,f335,f232,f227,f224,f3826,f6181,f6206,f6233,f6262,f6268,f6320,f6416,f6505,f6511,f6509,f6532,f6533,f6534,f6535,f6536,f6537,f6538,f6539,f6540,f6541,f6542,f6543,f6544,f6625,f6828,f6770,f6857])).
fof(f6871,plain,(
  ( ! [X4,X2,X3] : (s('type.option.option'(X2),'const.option.NONE_0') = s('type.option.option'(X2),X4) | s('type.option.option'(X2),X4) = s('type.option.option'(X2),'const.option.SOME_1'(s(X2,sK8('const.option.SOME_1'(s(sK0,sK8(sK2,sK3,sK5,sK1,sK0))),X4,X3,sK0,X2))))) ) | ($spl18 | $spl21 | $spl32)),
  inference(subsumption_resolution,[],[f6798,f126])).
fof(f6798,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | s('type.option.option'(X2),'const.option.NONE_0') = s('type.option.option'(X2),X4) | s('type.option.option'(X2),X4) = s('type.option.option'(X2),'const.option.SOME_1'(s(X2,sK8('const.option.SOME_1'(s(sK0,sK8(sK2,sK3,sK5,sK1,sK0))),X4,X3,sK0,X2))))) ) | ($spl18 | $spl21 | $spl32)),
  inference(backward_demodulation,[],[f6625,f6187])).
fof(f6187,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'const.option.OPTREL_3'(s(fun(X2,fun(sK0,bool)),X3),s('type.option.option'(X2),X4),s('type.option.option'(sK0),sK3)))) | s('type.option.option'(X2),'const.option.NONE_0') = s('type.option.option'(X2),X4) | s('type.option.option'(X2),X4) = s('type.option.option'(X2),'const.option.SOME_1'(s(X2,sK8('const.option.SOME_1'(s(sK0,sK8(sK2,sK3,sK5,sK1,sK0))),X4,X3,sK0,X2))))) ) | $spl18),
  inference(superposition,[],[f186,f3825])).
fof(f6870,plain,(
  $false | ($spl18 | $spl21 | $spl32)),
  inference(global_subsumption,[],[f6869,f125,f124,f126,f129,f132,f131,f173,f177,f176,f184,f181,f185,f219,f223,f259,f296,f428,f498,f714,f715,f726,f764,f765,f767,f806,f822,f824,f828,f884,f999,f1002,f1833,f1826,f1861,f1819,f1862,f1810,f1910,f1763,f1711,f128,f135,f136,f773,f790,f792,f795,f805,f827,f830,f954,f940,f1515,f2829,f2830,f2831,f3830,f3832,f3833,f3975,f4201,f4187,f3961,f3625,f192,f189,f188,f187,f186,f123,f716,f1330,f936,f1333,f1312,f918,f1372,f1362,f1360,f1356,f1354,f1352,f1341,f1335,f1320,f1318,f1314,f998,f988,f986,f926,f924,f920,f191,f190,f193,f1376,f1331,f1004,f180,f179,f178,f2842,f2841,f2840,f3836,f2849,f2848,f2838,f2837,f2151,f1995,f1994,f1993,f1981,f1951,f1950,f1949,f1751,f1739,f1705,f1703,f1655,f1637,f1636,f1626,f1572,f1570,f1523,f1489,f1438,f1437,f1344,f1323,f1322,f1321,f982,f956,f939,f938,f937,f826,f776,f775,f774,f335,f232,f227,f224,f3826,f6181,f6206,f6233,f6262,f6268,f6320,f6416,f6505,f6511,f6509,f6532,f6533,f6534,f6535,f6536,f6537,f6538,f6539,f6540,f6541,f6542,f6543,f6544,f6625,f6828,f6770,f6857])).
fof(f6869,plain,(
  ( ! [X2,X0,X1] : (s('type.option.option'(X0),X2) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,sK8('const.option.SOME_1'(s(sK0,sK8(sK2,sK3,sK5,sK1,sK0))),X2,X1,sK0,X0))))) ) | ($spl18 | $spl21 | $spl32)),
  inference(subsumption_resolution,[],[f6797,f126])).
fof(f6797,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s('type.option.option'(X0),X2) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,sK8('const.option.SOME_1'(s(sK0,sK8(sK2,sK3,sK5,sK1,sK0))),X2,X1,sK0,X0))))) ) | ($spl18 | $spl21 | $spl32)),
  inference(backward_demodulation,[],[f6625,f6225])).
fof(f6225,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(sK0,bool)),X1),s('type.option.option'(X0),X2),s('type.option.option'(sK0),sK3)))) | s('type.option.option'(X0),X2) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,sK8('const.option.SOME_1'(s(sK0,sK8(sK2,sK3,sK5,sK1,sK0))),X2,X1,sK0,X0))))) ) | $spl18),
  inference(subsumption_resolution,[],[f6208,f173])).
fof(f6208,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(sK0,bool)),X1),s('type.option.option'(X0),X2),s('type.option.option'(sK0),sK3)))) | s('type.option.option'(sK0),'const.option.NONE_0') = s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK8(sK2,sK3,sK5,sK1,sK0)))) | s('type.option.option'(X0),X2) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,sK8('const.option.SOME_1'(s(sK0,sK8(sK2,sK3,sK5,sK1,sK0))),X2,X1,sK0,X0))))) ) | $spl18),
  inference(superposition,[],[f187,f3825])).
fof(f6868,plain,(
  $false | ($spl18 | $spl21 | $spl32)),
  inference(global_subsumption,[],[f6867,f125,f124,f126,f129,f132,f131,f173,f177,f176,f184,f181,f185,f219,f223,f259,f296,f428,f498,f714,f715,f726,f764,f765,f767,f806,f822,f824,f828,f884,f999,f1002,f1833,f1826,f1861,f1819,f1862,f1810,f1910,f1763,f1711,f128,f135,f136,f773,f790,f792,f795,f805,f827,f830,f954,f940,f1515,f2829,f2830,f2831,f3830,f3832,f3833,f3975,f4201,f4187,f3961,f3625,f192,f189,f188,f187,f186,f123,f716,f1330,f936,f1333,f1312,f918,f1372,f1362,f1360,f1356,f1354,f1352,f1341,f1335,f1320,f1318,f1314,f998,f988,f986,f926,f924,f920,f191,f190,f193,f1376,f1331,f1004,f180,f179,f178,f2842,f2841,f2840,f3836,f2849,f2848,f2838,f2837,f2151,f1995,f1994,f1993,f1981,f1951,f1950,f1949,f1751,f1739,f1705,f1703,f1655,f1637,f1636,f1626,f1572,f1570,f1523,f1489,f1438,f1437,f1344,f1323,f1322,f1321,f982,f956,f939,f938,f937,f826,f776,f775,f774,f335,f232,f227,f224,f3826,f6181,f6206,f6233,f6262,f6268,f6320,f6416,f6505,f6511,f6509,f6532,f6533,f6534,f6535,f6536,f6537,f6538,f6539,f6540,f6541,f6542,f6543,f6544,f6625,f6828,f6770,f6857])).
fof(f6867,plain,(
  ( ! [X2,X0,X1] : (s('type.option.option'(sK0),sK3) = s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK9('const.option.SOME_1'(s(sK0,sK8(sK2,sK3,sK5,sK1,sK0))),X2,X1,sK0,X0)))) | s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X2)) ) | ($spl18 | $spl21 | $spl32)),
  inference(subsumption_resolution,[],[f6796,f126])).
fof(f6796,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s('type.option.option'(sK0),sK3) = s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK9('const.option.SOME_1'(s(sK0,sK8(sK2,sK3,sK5,sK1,sK0))),X2,X1,sK0,X0)))) | s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X2)) ) | ($spl18 | $spl21 | $spl32)),
  inference(backward_demodulation,[],[f6625,f6287])).
fof(f6287,plain,(
  ( ! [X2,X0,X1] : (s('type.option.option'(sK0),sK3) = s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK9('const.option.SOME_1'(s(sK0,sK8(sK2,sK3,sK5,sK1,sK0))),X2,X1,sK0,X0)))) | ~p(s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(sK0,bool)),X1),s('type.option.option'(X0),X2),s('type.option.option'(sK0),sK3)))) | s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X2)) ) | $spl18),
  inference(forward_demodulation,[],[f6270,f3825])).
fof(f6270,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(sK0,bool)),X1),s('type.option.option'(X0),X2),s('type.option.option'(sK0),sK3)))) | s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X2) | s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK8(sK2,sK3,sK5,sK1,sK0)))) = s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK9('const.option.SOME_1'(s(sK0,sK8(sK2,sK3,sK5,sK1,sK0))),X2,X1,sK0,X0))))) ) | $spl18),
  inference(superposition,[],[f188,f3825])).
fof(f6866,plain,(
  $false | ($spl18 | $spl21 | $spl32)),
  inference(global_subsumption,[],[f6865,f125,f124,f126,f129,f132,f131,f173,f177,f176,f184,f181,f185,f219,f223,f259,f296,f428,f498,f714,f715,f726,f764,f765,f767,f806,f822,f824,f828,f884,f999,f1002,f1833,f1826,f1861,f1819,f1862,f1810,f1910,f1763,f1711,f128,f135,f136,f773,f790,f792,f795,f805,f827,f830,f954,f940,f1515,f2829,f2830,f2831,f3830,f3832,f3833,f3975,f4201,f4187,f3961,f3625,f192,f189,f188,f187,f186,f123,f716,f1330,f936,f1333,f1312,f918,f1372,f1362,f1360,f1356,f1354,f1352,f1341,f1335,f1320,f1318,f1314,f998,f988,f986,f926,f924,f920,f191,f190,f193,f1376,f1331,f1004,f180,f179,f178,f2842,f2841,f2840,f3836,f2849,f2848,f2838,f2837,f2151,f1995,f1994,f1993,f1981,f1951,f1950,f1949,f1751,f1739,f1705,f1703,f1655,f1637,f1636,f1626,f1572,f1570,f1523,f1489,f1438,f1437,f1344,f1323,f1322,f1321,f982,f956,f939,f938,f937,f826,f776,f775,f774,f335,f232,f227,f224,f3826,f6181,f6206,f6233,f6262,f6268,f6320,f6416,f6505,f6511,f6509,f6532,f6533,f6534,f6535,f6536,f6537,f6538,f6539,f6540,f6541,f6542,f6543,f6544,f6625,f6828,f6770,f6857])).
fof(f6865,plain,(
  ( ! [X2,X0,X1] : (s('type.option.option'(sK0),sK3) = s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK9('const.option.SOME_1'(s(sK0,sK8(sK2,sK3,sK5,sK1,sK0))),X2,X1,sK0,X0))))) ) | ($spl18 | $spl21 | $spl32)),
  inference(subsumption_resolution,[],[f6795,f126])).
fof(f6795,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s('type.option.option'(sK0),sK3) = s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK9('const.option.SOME_1'(s(sK0,sK8(sK2,sK3,sK5,sK1,sK0))),X2,X1,sK0,X0))))) ) | ($spl18 | $spl21 | $spl32)),
  inference(backward_demodulation,[],[f6625,f6415])).
fof(f6415,plain,(
  ( ! [X2,X0,X1] : (s('type.option.option'(sK0),sK3) = s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK9('const.option.SOME_1'(s(sK0,sK8(sK2,sK3,sK5,sK1,sK0))),X2,X1,sK0,X0)))) | ~p(s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(sK0,bool)),X1),s('type.option.option'(X0),X2),s('type.option.option'(sK0),sK3))))) ) | $spl18),
  inference(forward_demodulation,[],[f6414,f3825])).
fof(f6414,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(sK0,bool)),X1),s('type.option.option'(X0),X2),s('type.option.option'(sK0),sK3)))) | s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK8(sK2,sK3,sK5,sK1,sK0)))) = s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK9('const.option.SOME_1'(s(sK0,sK8(sK2,sK3,sK5,sK1,sK0))),X2,X1,sK0,X0))))) ) | $spl18),
  inference(subsumption_resolution,[],[f6398,f173])).
fof(f6398,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(sK0,bool)),X1),s('type.option.option'(X0),X2),s('type.option.option'(sK0),sK3)))) | s('type.option.option'(sK0),'const.option.NONE_0') = s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK8(sK2,sK3,sK5,sK1,sK0)))) | s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK8(sK2,sK3,sK5,sK1,sK0)))) = s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK9('const.option.SOME_1'(s(sK0,sK8(sK2,sK3,sK5,sK1,sK0))),X2,X1,sK0,X0))))) ) | $spl18),
  inference(superposition,[],[f189,f3825])).
fof(f6864,plain,(
  $false | ($spl18 | $spl21 | $spl32)),
  inference(global_subsumption,[],[f6863,f125,f124,f126,f129,f132,f131,f173,f177,f176,f184,f181,f185,f219,f223,f259,f296,f428,f498,f714,f715,f726,f764,f765,f767,f806,f822,f824,f828,f884,f999,f1002,f1833,f1826,f1861,f1819,f1862,f1810,f1910,f1763,f1711,f128,f135,f136,f773,f790,f792,f795,f805,f827,f830,f954,f940,f1515,f2829,f2830,f2831,f3830,f3832,f3833,f3975,f4201,f4187,f3961,f3625,f192,f189,f188,f187,f186,f123,f716,f1330,f936,f1333,f1312,f918,f1372,f1362,f1360,f1356,f1354,f1352,f1341,f1335,f1320,f1318,f1314,f998,f988,f986,f926,f924,f920,f191,f190,f193,f1376,f1331,f1004,f180,f179,f178,f2842,f2841,f2840,f3836,f2849,f2848,f2838,f2837,f2151,f1995,f1994,f1993,f1981,f1951,f1950,f1949,f1751,f1739,f1705,f1703,f1655,f1637,f1636,f1626,f1572,f1570,f1523,f1489,f1438,f1437,f1344,f1323,f1322,f1321,f982,f956,f939,f938,f937,f826,f776,f775,f774,f335,f232,f227,f224,f3826,f6181,f6206,f6233,f6262,f6268,f6320,f6416,f6505,f6511,f6509,f6532,f6533,f6534,f6535,f6536,f6537,f6538,f6539,f6540,f6541,f6542,f6543,f6544,f6625,f6828,f6770,f6857])).
fof(f6863,plain,(
  ( ! [X17] : (s('type.option.option'(sK0),sK3) = s('type.option.option'(sK0),sK6('const.option.SOME_1'(s(sK0,sK8(sK2,sK3,sK5,sK1,sK0))),X17,'type.option.option'(sK0)))) ) | ($spl18 | $spl21 | $spl32)),
  inference(subsumption_resolution,[],[f6794,f126])).
fof(f6794,plain,(
  ( ! [X17] : (~p(s(bool,'T_0')) | s('type.option.option'(sK0),sK3) = s('type.option.option'(sK0),sK6('const.option.SOME_1'(s(sK0,sK8(sK2,sK3,sK5,sK1,sK0))),X17,'type.option.option'(sK0)))) ) | ($spl18 | $spl21 | $spl32)),
  inference(backward_demodulation,[],[f6625,f6203])).
fof(f6203,plain,(
  ( ! [X17] : (s('type.option.option'(sK0),sK3) = s('type.option.option'(sK0),sK6('const.option.SOME_1'(s(sK0,sK8(sK2,sK3,sK5,sK1,sK0))),X17,'type.option.option'(sK0))) | ~p(s(bool,happ(s(fun('type.option.option'(sK0),bool),X17),s('type.option.option'(sK0),sK3))))) ) | $spl18),
  inference(forward_demodulation,[],[f6196,f3825])).
fof(f6196,plain,(
  ( ! [X17] : (~p(s(bool,happ(s(fun('type.option.option'(sK0),bool),X17),s('type.option.option'(sK0),sK3)))) | s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK8(sK2,sK3,sK5,sK1,sK0)))) = s('type.option.option'(sK0),sK6('const.option.SOME_1'(s(sK0,sK8(sK2,sK3,sK5,sK1,sK0))),X17,'type.option.option'(sK0)))) ) | $spl18),
  inference(superposition,[],[f179,f3825])).
fof(f6862,plain,(
  $false | ($spl8 | $spl21 | $spl32)),
  inference(global_subsumption,[],[f6861,f125,f124,f126,f129,f132,f131,f173,f177,f176,f184,f181,f185,f219,f223,f259,f296,f428,f498,f714,f715,f726,f764,f765,f767,f806,f822,f824,f828,f884,f999,f1002,f1833,f1826,f1861,f1819,f1862,f1810,f1910,f1763,f1711,f128,f135,f136,f773,f790,f792,f795,f805,f827,f830,f954,f940,f1515,f2829,f2830,f2831,f3830,f3832,f3833,f3975,f4201,f4187,f3961,f3625,f192,f189,f188,f187,f186,f123,f716,f1330,f936,f1333,f1312,f918,f1372,f1362,f1360,f1356,f1354,f1352,f1341,f1335,f1320,f1318,f1314,f998,f988,f986,f926,f924,f920,f191,f190,f193,f1376,f1331,f1004,f180,f179,f178,f2842,f2841,f2840,f3836,f2849,f2848,f2838,f2837,f2151,f1995,f1994,f1993,f1981,f1951,f1950,f1949,f1751,f1739,f1705,f1703,f1655,f1637,f1636,f1626,f1572,f1570,f1523,f1489,f1438,f1437,f1344,f1323,f1322,f1321,f982,f956,f939,f938,f937,f826,f776,f775,f774,f335,f232,f227,f224,f3826,f6181,f6206,f6233,f6262,f6268,f6320,f6416,f6505,f6511,f6509,f6532,f6533,f6534,f6535,f6536,f6537,f6538,f6539,f6540,f6541,f6542,f6543,f6544,f6625,f6828,f6770,f6857])).
fof(f6861,plain,(
  ( ! [X10,X11,X9] : (s(X9,sK6(X11,X10,X9)) = s(X9,sK6(sK6(X11,X10,X9),X10,X9))) ) | ($spl8 | $spl21 | $spl32)),
  inference(subsumption_resolution,[],[f6779,f126])).
fof(f6779,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,'T_0')) | s(X9,sK6(X11,X10,X9)) = s(X9,sK6(sK6(X11,X10,X9),X10,X9))) ) | ($spl8 | $spl21 | $spl32)),
  inference(backward_demodulation,[],[f6625,f3798])).
fof(f6858,plain,(
  $false | ($spl21 | $spl32)),
  inference(global_subsumption,[],[f125,f124,f126,f129,f132,f131,f173,f177,f176,f184,f181,f185,f219,f223,f259,f296,f428,f498,f714,f715,f726,f764,f765,f767,f806,f822,f824,f828,f884,f999,f1002,f1833,f1826,f1861,f1819,f1862,f1810,f1910,f1763,f1711,f128,f135,f136,f773,f790,f792,f795,f805,f827,f830,f954,f940,f1515,f2829,f2830,f2831,f3830,f3832,f3833,f3975,f4201,f4187,f3961,f3625,f192,f189,f188,f187,f186,f123,f716,f1330,f936,f1333,f1312,f918,f1372,f1362,f1360,f1356,f1354,f1352,f1341,f1335,f1320,f1318,f1314,f998,f988,f986,f926,f924,f920,f191,f190,f193,f1376,f1331,f1004,f180,f179,f178,f2842,f2841,f2840,f3836,f2849,f2848,f2838,f2837,f2151,f1995,f1994,f1993,f1981,f1951,f1950,f1949,f1751,f1739,f1705,f1703,f1655,f1637,f1636,f1626,f1572,f1570,f1523,f1489,f1438,f1437,f1344,f1323,f1322,f1321,f982,f956,f939,f938,f937,f826,f776,f775,f774,f335,f232,f227,f224,f3826,f6181,f6206,f6233,f6262,f6268,f6320,f6416,f6505,f6511,f6509,f6532,f6533,f6534,f6535,f6536,f6537,f6538,f6539,f6540,f6541,f6542,f6543,f6544,f6625,f6828,f6770,f6857])).
fof(f6856,plain,(
  $false | ($spl21 | $spl32)),
  inference(global_subsumption,[],[f125,f124,f126,f129,f132,f131,f173,f177,f176,f184,f181,f185,f219,f223,f259,f296,f428,f498,f714,f715,f726,f764,f765,f767,f806,f822,f824,f828,f884,f999,f1002,f1833,f1826,f1861,f1819,f1862,f1810,f1910,f1763,f1711,f128,f135,f136,f773,f790,f792,f795,f805,f827,f830,f954,f940,f1515,f2829,f2830,f2831,f3830,f3832,f3833,f3975,f4201,f4187,f3961,f3625,f192,f189,f188,f187,f186,f123,f716,f1330,f936,f1333,f1312,f918,f1372,f1362,f1360,f1356,f1354,f1352,f1341,f1335,f1320,f1318,f1314,f998,f988,f986,f926,f924,f920,f191,f190,f193,f1376,f1331,f1004,f180,f179,f178,f2842,f2841,f2840,f3836,f2849,f2848,f2838,f2837,f2151,f1995,f1994,f1993,f1981,f1951,f1950,f1949,f1751,f1739,f1705,f1703,f1655,f1637,f1636,f1626,f1572,f1570,f1523,f1489,f1438,f1437,f1344,f1323,f1322,f1321,f982,f956,f939,f938,f937,f826,f776,f775,f774,f335,f232,f227,f224,f3826,f6181,f6206,f6233,f6262,f6268,f6320,f6416,f6505,f6511,f6509,f6532,f6533,f6534,f6535,f6536,f6537,f6538,f6539,f6540,f6541,f6542,f6543,f6544,f6625,f6828,f6770])).
fof(f6849,plain,(
  $false | ($spl8 | $spl21 | $spl32)),
  inference(global_subsumption,[],[f6848,f125,f124,f126,f129,f132,f131,f173,f177,f176,f184,f181,f185,f219,f223,f259,f296,f428,f498,f714,f715,f726,f764,f765,f767,f806,f822,f824,f828,f884,f999,f1002,f1833,f1826,f1861,f1819,f1862,f1810,f1910,f1763,f1711,f128,f135,f136,f773,f790,f792,f795,f805,f827,f830,f954,f940,f1515,f2829,f2830,f2831,f3830,f3832,f3833,f3975,f4201,f4187,f3961,f3625,f192,f189,f188,f187,f186,f123,f716,f1330,f936,f1333,f1312,f918,f1372,f1362,f1360,f1356,f1354,f1352,f1341,f1335,f1320,f1318,f1314,f998,f988,f986,f926,f924,f920,f191,f190,f193,f1376,f1331,f1004,f180,f179,f178,f2842,f2841,f2840,f3836,f2849,f2848,f2838,f2837,f2151,f1995,f1994,f1993,f1981,f1951,f1950,f1949,f1751,f1739,f1705,f1703,f1655,f1637,f1636,f1626,f1572,f1570,f1523,f1489,f1438,f1437,f1344,f1323,f1322,f1321,f982,f956,f939,f938,f937,f826,f776,f775,f774,f335,f232,f227,f224,f3826,f6181,f6206,f6233,f6262,f6268,f6320,f6416,f6505,f6511,f6509,f6532,f6533,f6534,f6535,f6536,f6537,f6538,f6539,f6540,f6541,f6542,f6543,f6544,f6625,f6828])).
fof(f6848,plain,(
  ( ! [X152,X151,X149,X150] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X152))) | s(X149,X151) = s(X149,sK6(X151,X150,X149))) ) | ($spl8 | $spl21 | $spl32)),
  inference(forward_demodulation,[],[f6847,f6625])).
fof(f6847,plain,(
  ( ! [X152,X151,X149,X150] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X152))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(X149,bool),X150),s(X149,X151))))) | s(X149,X151) = s(X149,sK6(X151,X150,X149))) ) | ($spl8 | $spl21 | $spl32)),
  inference(subsumption_resolution,[],[f6693,f6625])).
fof(f6693,plain,(
  ( ! [X152,X151,X149,X150] : (s(bool,'T_0') != s(bool,X152) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X152))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(X149,bool),X150),s(X149,X151))))) | s(X149,X151) = s(X149,sK6(X151,X150,X149))) ) | ($spl8 | $spl21 | $spl32)),
  inference(backward_demodulation,[],[f6625,f3775])).
fof(f6846,plain,(
  $false | ($spl8 | $spl21 | $spl32)),
  inference(global_subsumption,[],[f6692,f125,f124,f126,f129,f132,f131,f173,f177,f176,f184,f181,f185,f219,f223,f259,f296,f428,f498,f714,f715,f726,f764,f765,f767,f806,f822,f824,f828,f884,f999,f1002,f1833,f1826,f1861,f1819,f1862,f1810,f1910,f1763,f1711,f128,f135,f136,f773,f790,f792,f795,f805,f827,f830,f954,f940,f1515,f2829,f2830,f2831,f3830,f3832,f3833,f3975,f4201,f4187,f3961,f3625,f192,f189,f188,f187,f186,f123,f716,f1330,f936,f1333,f1312,f918,f1372,f1362,f1360,f1356,f1354,f1352,f1341,f1335,f1320,f1318,f1314,f998,f988,f986,f926,f924,f920,f191,f190,f193,f1376,f1331,f1004,f180,f179,f178,f2842,f2841,f2840,f3836,f2849,f2848,f2838,f2837,f2151,f1995,f1994,f1993,f1981,f1951,f1950,f1949,f1751,f1739,f1705,f1703,f1655,f1637,f1636,f1626,f1572,f1570,f1523,f1489,f1438,f1437,f1344,f1323,f1322,f1321,f982,f956,f939,f938,f937,f826,f776,f775,f774,f335,f232,f227,f224,f3826,f6181,f6206,f6233,f6262,f6268,f6320,f6416,f6505,f6511,f6509,f6532,f6533,f6534,f6535,f6536,f6537,f6538,f6539,f6540,f6541,f6542,f6543,f6544,f6625,f6828])).
fof(f6692,plain,(
  ( ! [X144,X142,X143,X141] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X144))) | s(bool,X144) = s(bool,happ(s(fun(X141,bool),X142),s(X141,X143))) | s(X141,X143) = s(X141,sK6(X143,X142,X141))) ) | ($spl8 | $spl21 | $spl32)),
  inference(backward_demodulation,[],[f6625,f3773])).
fof(f6829,plain,(
  $false | ($spl21 | $spl32)),
  inference(global_subsumption,[],[f125,f124,f126,f129,f132,f131,f173,f177,f176,f184,f181,f185,f219,f223,f259,f296,f428,f498,f714,f715,f726,f764,f765,f767,f806,f822,f824,f828,f884,f999,f1002,f1833,f1826,f1861,f1819,f1862,f1810,f1910,f1763,f1711,f128,f135,f136,f773,f790,f792,f795,f805,f827,f830,f954,f940,f1515,f2829,f2830,f2831,f3830,f3832,f3833,f3975,f4201,f4187,f3961,f3625,f192,f189,f188,f187,f186,f123,f716,f1330,f936,f1333,f1312,f918,f1372,f1362,f1360,f1356,f1354,f1352,f1341,f1335,f1320,f1318,f1314,f998,f988,f986,f926,f924,f920,f191,f190,f193,f1376,f1331,f1004,f180,f179,f178,f2842,f2841,f2840,f3836,f2849,f2848,f2838,f2837,f2151,f1995,f1994,f1993,f1981,f1951,f1950,f1949,f1751,f1739,f1705,f1703,f1655,f1637,f1636,f1626,f1572,f1570,f1523,f1489,f1438,f1437,f1344,f1323,f1322,f1321,f982,f956,f939,f938,f937,f826,f776,f775,f774,f335,f232,f227,f224,f3826,f6181,f6206,f6233,f6262,f6268,f6320,f6416,f6505,f6511,f6509,f6532,f6533,f6534,f6535,f6536,f6537,f6538,f6539,f6540,f6541,f6542,f6543,f6544,f6625,f6828])).
fof(f6827,plain,(
  $false | ($spl21 | $spl32)),
  inference(subsumption_resolution,[],[f6626,f126])).
fof(f6626,plain,(
  ~p(s(bool,'T_0')) | ($spl21 | $spl32)),
  inference(backward_demodulation,[],[f6625,f128])).
fof(f6826,plain,(
  $false | ($spl21 | $spl32)),
  inference(trivial_inequality_removal,[],[f6630])).
fof(f6630,plain,(
  s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | ($spl21 | $spl32)),
  inference(backward_demodulation,[],[f6625,f790])).
fof(f6825,plain,(
  $false | ($spl21 | $spl32)),
  inference(trivial_inequality_removal,[],[f6632])).
fof(f6632,plain,(
  s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))) | ($spl21 | $spl32)),
  inference(backward_demodulation,[],[f6625,f795])).
fof(f6824,plain,(
  $false | ($spl21 | $spl32)),
  inference(trivial_inequality_removal,[],[f6633])).
fof(f6633,plain,(
  s('type.option.option'('type.option.option'('type.option.option'(bool))),'const.option.SOME_1'(s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))))) != s('type.option.option'('type.option.option'('type.option.option'(bool))),'const.option.SOME_1'(s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))))) | ($spl21 | $spl32)),
  inference(backward_demodulation,[],[f6625,f805])).
fof(f6592,plain,(
  $false | ($spl8 | $spl21 | $spl32)),
  inference(subsumption_resolution,[],[f6529,f6300])).
fof(f6529,plain,(
  s(bool,'T_0') = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),'const.option.NONE_0'))) | $spl32),
  inference(backward_demodulation,[],[f6509,f227])).
fof(f6507,plain,(
  s('type.option.option'(sK1),sK2) = s('type.option.option'(sK1),'const.option.SOME_1'(s(sK1,sK9(sK2,sK3,sK5,sK1,sK0)))) | $spl30),
  inference(cnf_transformation,[],[f6507_D])).
fof(f6507_D,plain,(
  s('type.option.option'(sK1),sK2) = s('type.option.option'(sK1),'const.option.SOME_1'(s(sK1,sK9(sK2,sK3,sK5,sK1,sK0)))) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f6509,plain,(
  s('type.option.option'(sK1),'const.option.NONE_0') = s('type.option.option'(sK1),sK2) | $spl32),
  inference(cnf_transformation,[],[f6509_D])).
fof(f6509_D,plain,(
  s('type.option.option'(sK1),'const.option.NONE_0') = s('type.option.option'(sK1),sK2) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f6505,plain,(
  s('type.option.option'(sK1),'const.option.NONE_0') = s('type.option.option'(sK1),sK2) | s('type.option.option'(sK1),sK2) = s('type.option.option'(sK1),'const.option.SOME_1'(s(sK1,sK9(sK2,sK3,sK5,sK1,sK0))))),
  inference(subsumption_resolution,[],[f6444,f126])).
fof(f6444,plain,(
  ~p(s(bool,'T_0')) | s('type.option.option'(sK1),'const.option.NONE_0') = s('type.option.option'(sK1),sK2) | s('type.option.option'(sK1),sK2) = s('type.option.option'(sK1),'const.option.SOME_1'(s(sK1,sK9(sK2,sK3,sK5,sK1,sK0))))),
  inference(superposition,[],[f189,f227])).
fof(f6418,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s('type.option.option'(X36),'const.option.NONE_0') = s('type.option.option'(X36),X39) | s(bool,'F_0') = s(bool,'const.option.OPTREL_3'(s(fun(X35,fun(X36,bool)),X37),s('type.option.option'(X35),X38),s('type.option.option'(X36),X39))) | s('type.option.option'(X36),X39) = s('type.option.option'(X36),'const.option.SOME_1'(s(X36,sK9(X39,X38,X37,X36,X35))))) ) | $spl28),
  inference(cnf_transformation,[],[f6418_D])).
fof(f6418_D,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s('type.option.option'(X36),'const.option.NONE_0') = s('type.option.option'(X36),X39) | s(bool,'F_0') = s(bool,'const.option.OPTREL_3'(s(fun(X35,fun(X36,bool)),X37),s('type.option.option'(X35),X38),s('type.option.option'(X36),X39))) | s('type.option.option'(X36),X39) = s('type.option.option'(X36),'const.option.SOME_1'(s(X36,sK9(X39,X38,X37,X36,X35))))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f6419,plain,(
  ( ! [X45,X43,X41,X46,X44,X42] : (~p(s(bool,X46)) | s('type.option.option'(X42),'const.option.NONE_0') = s('type.option.option'(X42),X45) | s('type.option.option'(X42),X45) = s('type.option.option'(X42),'const.option.SOME_1'(s(X42,sK9(X45,X44,X43,X42,X41)))) | s(bool,'F_0') = s(bool,'const.option.OPTREL_3'(s(fun(X41,fun(X42,bool)),X43),s('type.option.option'(X41),X44),s('type.option.option'(X42),X45)))) ) | $spl8),
  inference(subsumption_resolution,[],[f6406,f135])).
fof(f6406,plain,(
  ( ! [X45,X43,X41,X46,X44,X42] : (~p(s(bool,X46)) | s('type.option.option'(X42),'const.option.NONE_0') = s('type.option.option'(X42),X45) | s('type.option.option'(X42),X45) = s('type.option.option'(X42),'const.option.SOME_1'(s(X42,sK9(X45,X44,X43,X42,X41)))) | s(bool,'F_0') = s(bool,'const.option.OPTREL_3'(s(fun(X41,fun(X42,bool)),X43),s('type.option.option'(X41),X44),s('type.option.option'(X42),X45))) | s(bool,'F_0') = s(bool,X46)) ) | $spl8),
  inference(superposition,[],[f189,f2532])).
fof(f6417,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X40)) | s('type.option.option'(X36),'const.option.NONE_0') = s('type.option.option'(X36),X39) | s('type.option.option'(X36),X39) = s('type.option.option'(X36),'const.option.SOME_1'(s(X36,sK9(X39,X38,X37,X36,X35)))) | s(bool,'F_0') = s(bool,'const.option.OPTREL_3'(s(fun(X35,fun(X36,bool)),X37),s('type.option.option'(X35),X38),s('type.option.option'(X36),X39)))) ) | $spl8),
  inference(subsumption_resolution,[],[f6405,f135])).
fof(f6405,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X40)) | s('type.option.option'(X36),'const.option.NONE_0') = s('type.option.option'(X36),X39) | s('type.option.option'(X36),X39) = s('type.option.option'(X36),'const.option.SOME_1'(s(X36,sK9(X39,X38,X37,X36,X35)))) | s(bool,'F_0') = s(bool,X40) | s(bool,'F_0') = s(bool,'const.option.OPTREL_3'(s(fun(X35,fun(X36,bool)),X37),s('type.option.option'(X35),X38),s('type.option.option'(X36),X39)))) ) | $spl8),
  inference(superposition,[],[f189,f2532])).
fof(f6153,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f6126,f128])).
fof(f6126,plain,(
  p(s(bool,'F_0')) | $spl20),
  inference(superposition,[],[f6064,f4091])).
fof(f4091,plain,(
  s(bool,'F_0') = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),'const.option.NONE_0'),s('type.option.option'(sK1),'const.option.NONE_0'))) | $spl20),
  inference(backward_demodulation,[],[f4088,f3843])).
fof(f3843,plain,(
  s(bool,'F_0') = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),'const.option.NONE_0'),s('type.option.option'(sK1),sK2))) | $spl20),
  inference(backward_demodulation,[],[f3827,f232])).
fof(f4088,plain,(
  s('type.option.option'(sK1),'const.option.NONE_0') = s('type.option.option'(sK1),sK2) | $spl20),
  inference(subsumption_resolution,[],[f4087,f173])).
fof(f4087,plain,(
  s('type.option.option'(sK1),'const.option.NONE_0') = s('type.option.option'(sK1),sK2) | s('type.option.option'(sK0),'const.option.NONE_0') = s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK8(sK2,'const.option.NONE_0',sK5,sK1,sK0)))) | $spl20),
  inference(subsumption_resolution,[],[f4039,f126])).
fof(f4039,plain,(
  ~p(s(bool,'T_0')) | s('type.option.option'(sK1),'const.option.NONE_0') = s('type.option.option'(sK1),sK2) | s('type.option.option'(sK0),'const.option.NONE_0') = s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK8(sK2,'const.option.NONE_0',sK5,sK1,sK0)))) | $spl20),
  inference(superposition,[],[f187,f3842])).
fof(f3842,plain,(
  s(bool,'T_0') = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),'const.option.NONE_0'),s('type.option.option'(sK1),sK2))) | $spl20),
  inference(backward_demodulation,[],[f3827,f227])).
fof(f6064,plain,(
  ( ! [X0] : (p(s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),X0),s('type.option.option'(sK0),'const.option.NONE_0'),s('type.option.option'(sK1),'const.option.NONE_0'))))) ) | $spl20),
  inference(forward_demodulation,[],[f6063,f3827])).
fof(f6063,plain,(
  ( ! [X0] : (p(s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),X0),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),'const.option.NONE_0'))))) ) | $spl20),
  inference(forward_demodulation,[],[f6055,f4088])).
fof(f6055,plain,(
  ( ! [X0] : (p(s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),X0),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f3827,f4088,f192])).
fof(f6152,plain,(
  $false | ($spl8 | $spl20)),
  inference(subsumption_resolution,[],[f6107,f223])).
fof(f6107,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl8 | $spl20)),
  inference(unit_resulting_resolution,[],[f4089,f6064,f2644])).
fof(f4089,plain,(
  s(bool,'T_0') != s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),'const.option.NONE_0'),s('type.option.option'(sK1),'const.option.NONE_0'))) | $spl20),
  inference(backward_demodulation,[],[f4088,f3841])).
fof(f3841,plain,(
  s(bool,'T_0') != s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),'const.option.NONE_0'),s('type.option.option'(sK1),sK2))) | $spl20),
  inference(backward_demodulation,[],[f3827,f224])).
fof(f6149,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f6113,f128])).
fof(f6113,plain,(
  p(s(bool,'F_0')) | $spl20),
  inference(unit_resulting_resolution,[],[f4091,f6064,f219])).
fof(f6147,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f6117,f126])).
fof(f6117,plain,(
  ~p(s(bool,'T_0')) | $spl20),
  inference(unit_resulting_resolution,[],[f4089,f6064,f184])).
fof(f6145,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f6121,f126])).
fof(f6121,plain,(
  ~p(s(bool,'T_0')) | $spl20),
  inference(unit_resulting_resolution,[],[f4089,f6064,f184])).
fof(f6144,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f6143,f223])).
fof(f6143,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl20),
  inference(backward_demodulation,[],[f6123,f4091])).
fof(f6123,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),X0),s('type.option.option'(sK0),'const.option.NONE_0'),s('type.option.option'(sK1),'const.option.NONE_0')))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f6064,f132])).
fof(f6124,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f4089,f6064,f132])).
fof(f6122,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f126,f4089,f6064,f184])).
fof(f6118,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f126,f4089,f6064,f184])).
fof(f6114,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f128,f4091,f6064,f219])).
fof(f6110,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f4091,f6064,f135])).
fof(f6108,plain,(
  $false | ($spl8 | $spl20)),
  inference(unit_resulting_resolution,[],[f223,f4089,f6064,f2644])).
fof(f4373,plain,(
  $false | $spl22),
  inference(subsumption_resolution,[],[f4213,f3828])).
fof(f4213,plain,(
  p(s(bool,'F_0')) | $spl22),
  inference(unit_resulting_resolution,[],[f223,f3828,f181])).
fof(f4372,plain,(
  $false | $spl22),
  inference(subsumption_resolution,[],[f4219,f3828])).
fof(f4219,plain,(
  p(s(bool,'F_0')) | $spl22),
  inference(unit_resulting_resolution,[],[f223,f3828,f181])).
fof(f4233,plain,(
  $false | ($spl20 | $spl22)),
  inference(unit_resulting_resolution,[],[f3827,f3827,f3828,f192])).
fof(f4225,plain,(
  $false | $spl22),
  inference(unit_resulting_resolution,[],[f223,f223,f3828,f296])).
fof(f4220,plain,(
  $false | $spl22),
  inference(unit_resulting_resolution,[],[f128,f223,f3828,f181])).
fof(f4218,plain,(
  $false | $spl22),
  inference(unit_resulting_resolution,[],[f3828,f223,f3828,f181])).
fof(f4214,plain,(
  $false | $spl22),
  inference(unit_resulting_resolution,[],[f128,f223,f3828,f181])).
fof(f4212,plain,(
  $false | $spl22),
  inference(unit_resulting_resolution,[],[f3828,f223,f3828,f181])).
fof(f4208,plain,(
  $false | $spl22),
  inference(unit_resulting_resolution,[],[f223,f3828,f136])).
fof(f4205,plain,(
  $false | $spl22),
  inference(unit_resulting_resolution,[],[f126,f3828])).
fof(f4203,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s('type.option.option'(X35),'const.option.NONE_0') = s('type.option.option'(X35),X38) | s(bool,'F_0') = s(bool,'const.option.OPTREL_3'(s(fun(X35,fun(X36,bool)),X37),s('type.option.option'(X35),X38),s('type.option.option'(X36),X39))) | s('type.option.option'(X36),X39) = s('type.option.option'(X36),'const.option.SOME_1'(s(X36,sK9(X39,X38,X37,X36,X35))))) ) | $spl26),
  inference(cnf_transformation,[],[f4203_D])).
fof(f4203_D,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s('type.option.option'(X35),'const.option.NONE_0') = s('type.option.option'(X35),X38) | s(bool,'F_0') = s(bool,'const.option.OPTREL_3'(s(fun(X35,fun(X36,bool)),X37),s('type.option.option'(X35),X38),s('type.option.option'(X36),X39))) | s('type.option.option'(X36),X39) = s('type.option.option'(X36),'const.option.SOME_1'(s(X36,sK9(X39,X38,X37,X36,X35))))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f4204,plain,(
  ( ! [X45,X43,X41,X46,X44,X42] : (~p(s(bool,X46)) | s('type.option.option'(X41),'const.option.NONE_0') = s('type.option.option'(X41),X44) | s('type.option.option'(X42),X45) = s('type.option.option'(X42),'const.option.SOME_1'(s(X42,sK9(X45,X44,X43,X42,X41)))) | s(bool,'F_0') = s(bool,'const.option.OPTREL_3'(s(fun(X41,fun(X42,bool)),X43),s('type.option.option'(X41),X44),s('type.option.option'(X42),X45)))) ) | $spl8),
  inference(subsumption_resolution,[],[f4197,f135])).
fof(f4197,plain,(
  ( ! [X45,X43,X41,X46,X44,X42] : (~p(s(bool,X46)) | s('type.option.option'(X41),'const.option.NONE_0') = s('type.option.option'(X41),X44) | s('type.option.option'(X42),X45) = s('type.option.option'(X42),'const.option.SOME_1'(s(X42,sK9(X45,X44,X43,X42,X41)))) | s(bool,'F_0') = s(bool,'const.option.OPTREL_3'(s(fun(X41,fun(X42,bool)),X43),s('type.option.option'(X41),X44),s('type.option.option'(X42),X45))) | s(bool,'F_0') = s(bool,X46)) ) | $spl8),
  inference(superposition,[],[f188,f2532])).
fof(f4202,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X40)) | s('type.option.option'(X35),'const.option.NONE_0') = s('type.option.option'(X35),X38) | s('type.option.option'(X36),X39) = s('type.option.option'(X36),'const.option.SOME_1'(s(X36,sK9(X39,X38,X37,X36,X35)))) | s(bool,'F_0') = s(bool,'const.option.OPTREL_3'(s(fun(X35,fun(X36,bool)),X37),s('type.option.option'(X35),X38),s('type.option.option'(X36),X39)))) ) | $spl8),
  inference(subsumption_resolution,[],[f4196,f135])).
fof(f4196,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X40)) | s('type.option.option'(X35),'const.option.NONE_0') = s('type.option.option'(X35),X38) | s('type.option.option'(X36),X39) = s('type.option.option'(X36),'const.option.SOME_1'(s(X36,sK9(X39,X38,X37,X36,X35)))) | s(bool,'F_0') = s(bool,X40) | s(bool,'F_0') = s(bool,'const.option.OPTREL_3'(s(fun(X35,fun(X36,bool)),X37),s('type.option.option'(X35),X38),s('type.option.option'(X36),X39)))) ) | $spl8),
  inference(superposition,[],[f188,f2532])).
fof(f3977,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s('type.option.option'(X36),'const.option.NONE_0') = s('type.option.option'(X36),X39) | s(bool,'F_0') = s(bool,'const.option.OPTREL_3'(s(fun(X35,fun(X36,bool)),X37),s('type.option.option'(X35),X38),s('type.option.option'(X36),X39))) | s('type.option.option'(X35),X38) = s('type.option.option'(X35),'const.option.SOME_1'(s(X35,sK8(X39,X38,X37,X36,X35))))) ) | $spl24),
  inference(cnf_transformation,[],[f3977_D])).
fof(f3977_D,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s('type.option.option'(X36),'const.option.NONE_0') = s('type.option.option'(X36),X39) | s(bool,'F_0') = s(bool,'const.option.OPTREL_3'(s(fun(X35,fun(X36,bool)),X37),s('type.option.option'(X35),X38),s('type.option.option'(X36),X39))) | s('type.option.option'(X35),X38) = s('type.option.option'(X35),'const.option.SOME_1'(s(X35,sK8(X39,X38,X37,X36,X35))))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f3978,plain,(
  ( ! [X45,X43,X41,X46,X44,X42] : (~p(s(bool,X46)) | s('type.option.option'(X42),'const.option.NONE_0') = s('type.option.option'(X42),X45) | s('type.option.option'(X41),X44) = s('type.option.option'(X41),'const.option.SOME_1'(s(X41,sK8(X45,X44,X43,X42,X41)))) | s(bool,'F_0') = s(bool,'const.option.OPTREL_3'(s(fun(X41,fun(X42,bool)),X43),s('type.option.option'(X41),X44),s('type.option.option'(X42),X45)))) ) | $spl8),
  inference(subsumption_resolution,[],[f3971,f135])).
fof(f3971,plain,(
  ( ! [X45,X43,X41,X46,X44,X42] : (~p(s(bool,X46)) | s('type.option.option'(X42),'const.option.NONE_0') = s('type.option.option'(X42),X45) | s('type.option.option'(X41),X44) = s('type.option.option'(X41),'const.option.SOME_1'(s(X41,sK8(X45,X44,X43,X42,X41)))) | s(bool,'F_0') = s(bool,'const.option.OPTREL_3'(s(fun(X41,fun(X42,bool)),X43),s('type.option.option'(X41),X44),s('type.option.option'(X42),X45))) | s(bool,'F_0') = s(bool,X46)) ) | $spl8),
  inference(superposition,[],[f187,f2532])).
fof(f3976,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X40)) | s('type.option.option'(X36),'const.option.NONE_0') = s('type.option.option'(X36),X39) | s('type.option.option'(X35),X38) = s('type.option.option'(X35),'const.option.SOME_1'(s(X35,sK8(X39,X38,X37,X36,X35)))) | s(bool,'F_0') = s(bool,'const.option.OPTREL_3'(s(fun(X35,fun(X36,bool)),X37),s('type.option.option'(X35),X38),s('type.option.option'(X36),X39)))) ) | $spl8),
  inference(subsumption_resolution,[],[f3970,f135])).
fof(f3970,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X40)) | s('type.option.option'(X36),'const.option.NONE_0') = s('type.option.option'(X36),X39) | s('type.option.option'(X35),X38) = s('type.option.option'(X35),'const.option.SOME_1'(s(X35,sK8(X39,X38,X37,X36,X35)))) | s(bool,'F_0') = s(bool,X40) | s(bool,'F_0') = s(bool,'const.option.OPTREL_3'(s(fun(X35,fun(X36,bool)),X37),s('type.option.option'(X35),X38),s('type.option.option'(X36),X39)))) ) | $spl8),
  inference(superposition,[],[f187,f2532])).
fof(f3825,plain,(
  s('type.option.option'(sK0),sK3) = s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK8(sK2,sK3,sK5,sK1,sK0)))) | $spl18),
  inference(cnf_transformation,[],[f3825_D])).
fof(f3825_D,plain,(
  s('type.option.option'(sK0),sK3) = s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK8(sK2,sK3,sK5,sK1,sK0)))) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f3827,plain,(
  s('type.option.option'(sK0),'const.option.NONE_0') = s('type.option.option'(sK0),sK3) | $spl20),
  inference(cnf_transformation,[],[f3827_D])).
fof(f3827_D,plain,(
  s('type.option.option'(sK0),'const.option.NONE_0') = s('type.option.option'(sK0),sK3) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f3832,plain,(
  s('type.option.option'(sK0),'const.option.NONE_0') = s('type.option.option'(sK0),sK3) | s('type.option.option'(sK0),sK3) = s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK8(sK2,sK3,sK5,sK1,sK0))))),
  inference(subsumption_resolution,[],[f3805,f126])).
fof(f3805,plain,(
  ~p(s(bool,'T_0')) | s('type.option.option'(sK0),'const.option.NONE_0') = s('type.option.option'(sK0),sK3) | s('type.option.option'(sK0),sK3) = s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK8(sK2,sK3,sK5,sK1,sK0))))),
  inference(superposition,[],[f186,f227])).
fof(f3831,plain,(
  s('type.option.option'(sK0),'const.option.NONE_0') = s('type.option.option'(sK0),sK3) | s('type.option.option'(sK0),sK3) = s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK8(sK2,sK3,sK5,sK1,sK0)))) | $spl8),
  inference(global_subsumption,[],[f125,f124,f126,f129,f132,f131,f173,f177,f176,f184,f181,f185,f219,f224,f223,f227,f259,f296,f335,f428,f498,f714,f715,f726,f764,f765,f767,f774,f775,f806,f822,f824,f826,f828,f884,f956,f982,f999,f1002,f1833,f1826,f1861,f1819,f1862,f1437,f1438,f1810,f1489,f1910,f1570,f1572,f1763,f1949,f1950,f1951,f1626,f1636,f1637,f1655,f1751,f1703,f1705,f1711,f1739,f1860,f128,f135,f136,f232,f773,f776,f790,f792,f795,f805,f827,f830,f926,f936,f954,f998,f1330,f1333,f1335,f1341,f1352,f1362,f1372,f1981,f1995,f178,f1994,f940,f179,f180,f1004,f1331,f1376,f1993,f920,f924,f986,f988,f1314,f1318,f1320,f1354,f1356,f1360,f918,f192,f189,f188,f187,f186,f1312,f938,f939,f1344,f937,f1515,f2151,f1322,f1323,f1321,f1523,f123,f716,f193,f191,f190,f2532,f2533,f2644,f2646,f2665,f2667,f2670,f2671,f2579,f2673,f2674,f2583,f2610,f2690,f2692,f2730,f2732,f2742,f2781,f2827,f2829,f2830,f2831,f2837,f2838,f2854,f2840,f2841,f2842,f2848,f2849,f2858,f2916,f3033,f3035,f3037,f3038,f3039,f3040,f3041,f3108,f3109,f3111,f3112,f3113,f3139,f3141,f3142,f3500,f3507,f3508,f3189,f3191,f3193,f3194,f3195,f3196,f3197,f3617,f3593,f3625,f3629,f3698,f3699,f3704,f3705,f3706,f3792,f3708,f3709,f3712,f3713,f3715,f3717,f3718,f3793,f3794,f3795,f3796,f3797,f3798,f3799,f3802,f3773,f3775,f3777,f3778,f3823,f3830])).
fof(f3823,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s('type.option.option'(sK0),'const.option.NONE_0') = s('type.option.option'(sK0),sK3) | s('type.option.option'(sK0),sK3) = s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK8(sK2,sK3,sK5,sK1,sK0))))) ) | $spl8),
  inference(subsumption_resolution,[],[f3803,f135])).
fof(f3803,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s('type.option.option'(sK0),'const.option.NONE_0') = s('type.option.option'(sK0),sK3) | s('type.option.option'(sK0),sK3) = s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK8(sK2,sK3,sK5,sK1,sK0)))) | s(bool,'F_0') = s(bool,X0)) ) | $spl8),
  inference(superposition,[],[f186,f1860])).
fof(f3620,plain,(
  ( ! [X24,X23] : (s(bool,X23) != s(bool,X24) | s(bool,X23) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,'F_0') = s(bool,X24)) ) | $spl16),
  inference(cnf_transformation,[],[f3620_D])).
fof(f3620_D,plain,(
  ( ! [X24,X23] : (s(bool,X23) != s(bool,X24) | s(bool,X23) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,'F_0') = s(bool,X24)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f3593,plain,(
  ( ! [X24,X23,X22] : (s(bool,X23) != s(bool,X24) | s(bool,'T_0') = s(bool,X22) | s(bool,'F_0') = s(bool,X24) | s(bool,X23) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f3574])).
fof(f3574,plain,(
  ( ! [X24,X23,X22] : (s(bool,X23) != s(bool,X24) | s(bool,'T_0') = s(bool,X22) | s(bool,'F_0') = s(bool,X24) | s(bool,X23) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,'T_0') = s(bool,X22)) ) | $spl8),
  inference(superposition,[],[f2583,f982])).
fof(f2150,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f2149,f126])).
fof(f2149,plain,(
  ~p(s(bool,'T_0')) | $spl12),
  inference(forward_demodulation,[],[f128,f2118])).
fof(f2118,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,X2)) ) | $spl12),
  inference(global_subsumption,[],[f125,f124,f126,f129,f132,f131,f173,f177,f176,f184,f181,f185,f219,f224,f223,f227,f259,f296,f335,f428,f498,f714,f715,f726,f764,f765,f767,f774,f775,f806,f822,f824,f826,f828,f884,f956,f982,f999,f1002,f178,f191,f190,f193,f123,f716,f954,f830,f827,f805,f795,f792,f790,f776,f773,f232,f136,f135,f128,f937,f924,f998,f988,f986,f936,f920,f918,f180,f179,f940,f1004,f926,f938,f939,f1314,f1330,f1318,f1320,f1321,f1322,f1323,f1331,f1333,f1335,f1341,f1344,f1352,f1354,f1356,f1372,f1360,f1362,f1376,f1833,f1826,f1861,f1819,f1862,f1437,f1438,f1810,f1489,f1515,f1910,f1523,f1570,f1572,f1763,f1949,f1950,f1951,f1626,f1636,f1637,f1655,f1981,f1751,f1993,f1994,f1995,f1703,f1705,f1711,f1739,f2006,f2028,f192,f189,f188,f187,f186,f1312])).
fof(f2028,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.option.option'(X0),X3),s('type.option.option'(X1),X4)))) | s('type.option.option'(X1),X4) != s('type.option.option'(X1),'const.option.SOME_1'(s(X1,X6))) | s('type.option.option'(X0),X3) != s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X5)))) ) | $spl12),
  inference(subsumption_resolution,[],[f2027,f126])).
fof(f2027,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.option.option'(X0),X3),s('type.option.option'(X1),X4)))) | s('type.option.option'(X1),X4) != s('type.option.option'(X1),'const.option.SOME_1'(s(X1,X6))) | s('type.option.option'(X0),X3) != s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X5)))) ) | $spl12),
  inference(forward_demodulation,[],[f2012,f227])).
fof(f2012,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)))) | p(s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.option.option'(X0),X3),s('type.option.option'(X1),X4)))) | s('type.option.option'(X1),X4) != s('type.option.option'(X1),'const.option.SOME_1'(s(X1,X6))) | s('type.option.option'(X0),X3) != s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X5)))) ) | $spl12),
  inference(backward_demodulation,[],[f2006,f193])).
fof(f2148,plain,(
  $false | $spl12),
  inference(global_subsumption,[],[f125,f124,f126,f129,f132,f131,f173,f177,f176,f184,f181,f185,f219,f224,f223,f227,f259,f296,f335,f428,f498,f714,f715,f726,f764,f765,f767,f774,f775,f806,f822,f824,f826,f828,f884,f956,f982,f999,f1002,f178,f191,f190,f193,f123,f716,f954,f830,f827,f805,f795,f792,f790,f776,f773,f232,f136,f135,f128,f937,f924,f998,f988,f986,f936,f920,f918,f180,f179,f940,f1004,f926,f938,f939,f1314,f1330,f1318,f1320,f1321,f1322,f1323,f1331,f1333,f1335,f1341,f1344,f1352,f1354,f1356,f1372,f1360,f1362,f1376,f1833,f1826,f1861,f1819,f1862,f1437,f1438,f1810,f1489,f1515,f1910,f1523,f1570,f1572,f1763,f1949,f1950,f1951,f1626,f1636,f1637,f1655,f1981,f1751,f1993,f1994,f1995,f1703,f1705,f1711,f1739,f2006,f2028,f192,f189,f188,f187,f186,f1312,f2118,f2127,f2128,f2129,f2130,f2133,f2143,f2147])).
fof(f2147,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl12),
  inference(subsumption_resolution,[],[f2146,f2118])).
fof(f2146,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl12),
  inference(forward_demodulation,[],[f135,f2118])).
fof(f2143,plain,(
  ( ! [X27] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X27))) | s(bool,X27) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)))) ) | $spl12),
  inference(forward_demodulation,[],[f776,f2118])).
fof(f2133,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,sK6(X2,X1,X0))) ) | $spl12),
  inference(subsumption_resolution,[],[f2132,f126])).
fof(f2132,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s(X0,X2) = s(X0,sK6(X2,X1,X0))) ) | $spl12),
  inference(forward_demodulation,[],[f179,f2118])).
fof(f2130,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.option.option'(X1),X4) = s('type.option.option'(X1),'const.option.SOME_1'(s(X1,sK9(X4,X3,X2,X1,X0)))) | s('type.option.option'(X1),'const.option.NONE_0') = s('type.option.option'(X1),X4)) ) | $spl12),
  inference(subsumption_resolution,[],[f2122,f126])).
fof(f2122,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | s('type.option.option'(X1),X4) = s('type.option.option'(X1),'const.option.SOME_1'(s(X1,sK9(X4,X3,X2,X1,X0)))) | s('type.option.option'(X1),'const.option.NONE_0') = s('type.option.option'(X1),X4)) ) | $spl12),
  inference(backward_demodulation,[],[f2118,f189])).
fof(f2129,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.option.option'(X1),X4) = s('type.option.option'(X1),'const.option.SOME_1'(s(X1,sK9(X4,X3,X2,X1,X0)))) | s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X3)) ) | $spl12),
  inference(subsumption_resolution,[],[f2121,f126])).
fof(f2121,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | s('type.option.option'(X1),X4) = s('type.option.option'(X1),'const.option.SOME_1'(s(X1,sK9(X4,X3,X2,X1,X0)))) | s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X3)) ) | $spl12),
  inference(backward_demodulation,[],[f2118,f188])).
fof(f2128,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.option.option'(X0),X3) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,sK8(X4,X3,X2,X1,X0)))) | s('type.option.option'(X1),'const.option.NONE_0') = s('type.option.option'(X1),X4)) ) | $spl12),
  inference(subsumption_resolution,[],[f2120,f126])).
fof(f2120,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | s('type.option.option'(X0),X3) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,sK8(X4,X3,X2,X1,X0)))) | s('type.option.option'(X1),'const.option.NONE_0') = s('type.option.option'(X1),X4)) ) | $spl12),
  inference(backward_demodulation,[],[f2118,f187])).
fof(f2127,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.option.option'(X0),X3) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,sK8(X4,X3,X2,X1,X0)))) | s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X3)) ) | $spl12),
  inference(subsumption_resolution,[],[f2119,f126])).
fof(f2119,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | s('type.option.option'(X0),X3) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,sK8(X4,X3,X2,X1,X0)))) | s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X3)) ) | $spl12),
  inference(backward_demodulation,[],[f2118,f186])).
fof(f2142,plain,(
  $false | $spl12),
  inference(trivial_inequality_removal,[],[f2141])).
fof(f2141,plain,(
  s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | $spl12),
  inference(forward_demodulation,[],[f790,f2118])).
fof(f2140,plain,(
  $false | $spl12),
  inference(trivial_inequality_removal,[],[f2139])).
fof(f2139,plain,(
  s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))) | $spl12),
  inference(forward_demodulation,[],[f795,f2118])).
fof(f2138,plain,(
  $false | $spl12),
  inference(trivial_inequality_removal,[],[f2137])).
fof(f2137,plain,(
  s('type.option.option'('type.option.option'('type.option.option'(bool))),'const.option.SOME_1'(s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))))) != s('type.option.option'('type.option.option'('type.option.option'(bool))),'const.option.SOME_1'(s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))))) | $spl12),
  inference(forward_demodulation,[],[f805,f2118])).
fof(f2117,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f125,f124,f126,f129,f132,f131,f173,f177,f176,f184,f181,f185,f219,f224,f223,f227,f259,f296,f335,f428,f498,f714,f715,f726,f764,f765,f767,f774,f775,f806,f822,f824,f826,f828,f884,f956,f982,f999,f1002,f192,f189,f188,f187,f186,f178,f191,f190,f193,f123,f716,f954,f830,f827,f805,f795,f792,f790,f776,f773,f232,f136,f135,f128,f937,f924,f998,f988,f986,f936,f920,f918,f180,f179,f940,f1004,f926,f938,f939,f1312,f1314,f1330,f1318,f1320,f1321,f1322,f1323,f1331,f1333,f1335,f1341,f1344,f1352,f1354,f1356,f1372,f1360,f1362,f1376,f1833,f1826,f1861,f1819,f1862,f1437,f1438,f1810,f1489,f1515,f1910,f1523,f1570,f1572,f1763,f1949,f1950,f1951,f1626,f1636,f1637,f1655,f1981,f1751,f1993,f1994,f1995,f1703,f1705,f1711,f1739,f1859,f2103,f2038,f2106,f2110,f2112,f2114,f2116])).
fof(f2116,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.option.option'(X1),X4) = s('type.option.option'(X1),'const.option.SOME_1'(s(X1,sK9(X4,X3,X2,X1,X0)))) | s('type.option.option'(X1),'const.option.NONE_0') = s('type.option.option'(X1),X4)) ) | $spl6),
  inference(subsumption_resolution,[],[f2081,f126])).
fof(f2081,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | s('type.option.option'(X1),X4) = s('type.option.option'(X1),'const.option.SOME_1'(s(X1,sK9(X4,X3,X2,X1,X0)))) | s('type.option.option'(X1),'const.option.NONE_0') = s('type.option.option'(X1),X4)) ) | $spl6),
  inference(backward_demodulation,[],[f1859,f189])).
fof(f2114,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.option.option'(X1),X4) = s('type.option.option'(X1),'const.option.SOME_1'(s(X1,sK9(X4,X3,X2,X1,X0)))) | s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X3)) ) | $spl6),
  inference(subsumption_resolution,[],[f2080,f126])).
fof(f2080,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | s('type.option.option'(X1),X4) = s('type.option.option'(X1),'const.option.SOME_1'(s(X1,sK9(X4,X3,X2,X1,X0)))) | s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X3)) ) | $spl6),
  inference(backward_demodulation,[],[f1859,f188])).
fof(f2112,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.option.option'(X0),X3) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,sK8(X4,X3,X2,X1,X0)))) | s('type.option.option'(X1),'const.option.NONE_0') = s('type.option.option'(X1),X4)) ) | $spl6),
  inference(subsumption_resolution,[],[f2079,f126])).
fof(f2079,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | s('type.option.option'(X0),X3) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,sK8(X4,X3,X2,X1,X0)))) | s('type.option.option'(X1),'const.option.NONE_0') = s('type.option.option'(X1),X4)) ) | $spl6),
  inference(backward_demodulation,[],[f1859,f187])).
fof(f2110,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.option.option'(X0),X3) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,sK8(X4,X3,X2,X1,X0)))) | s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X3)) ) | $spl6),
  inference(subsumption_resolution,[],[f2078,f126])).
fof(f2078,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | s('type.option.option'(X0),X3) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,sK8(X4,X3,X2,X1,X0)))) | s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X3)) ) | $spl6),
  inference(backward_demodulation,[],[f1859,f186])).
fof(f2106,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,sK6(X2,X1,X0))) ) | $spl6),
  inference(subsumption_resolution,[],[f2060,f126])).
fof(f2060,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s(X0,X2) = s(X0,sK6(X2,X1,X0))) ) | $spl6),
  inference(backward_demodulation,[],[f1859,f179])).
fof(f2038,plain,(
  ( ! [X27] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X27))) | s(bool,X27) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)))) ) | $spl6),
  inference(backward_demodulation,[],[f1859,f776])).
fof(f2103,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl6),
  inference(subsumption_resolution,[],[f2034,f1859])).
fof(f2034,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl6),
  inference(backward_demodulation,[],[f1859,f135])).
fof(f2115,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f125,f124,f126,f129,f132,f131,f173,f177,f176,f184,f181,f185,f219,f224,f223,f227,f259,f296,f335,f428,f498,f714,f715,f726,f764,f765,f767,f774,f775,f806,f822,f824,f826,f828,f884,f956,f982,f999,f1002,f192,f189,f188,f187,f186,f178,f191,f190,f193,f123,f716,f954,f830,f827,f805,f795,f792,f790,f776,f773,f232,f136,f135,f128,f937,f924,f998,f988,f986,f936,f920,f918,f180,f179,f940,f1004,f926,f938,f939,f1312,f1314,f1330,f1318,f1320,f1321,f1322,f1323,f1331,f1333,f1335,f1341,f1344,f1352,f1354,f1356,f1372,f1360,f1362,f1376,f1833,f1826,f1861,f1819,f1862,f1437,f1438,f1810,f1489,f1515,f1910,f1523,f1570,f1572,f1763,f1949,f1950,f1951,f1626,f1636,f1637,f1655,f1981,f1751,f1993,f1994,f1995,f1703,f1705,f1711,f1739,f1859,f2103,f2038,f2106,f2110,f2112,f2114])).
fof(f2113,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f125,f124,f126,f129,f132,f131,f173,f177,f176,f184,f181,f185,f219,f224,f223,f227,f259,f296,f335,f428,f498,f714,f715,f726,f764,f765,f767,f774,f775,f806,f822,f824,f826,f828,f884,f956,f982,f999,f1002,f192,f189,f188,f187,f186,f178,f191,f190,f193,f123,f716,f954,f830,f827,f805,f795,f792,f790,f776,f773,f232,f136,f135,f128,f937,f924,f998,f988,f986,f936,f920,f918,f180,f179,f940,f1004,f926,f938,f939,f1312,f1314,f1330,f1318,f1320,f1321,f1322,f1323,f1331,f1333,f1335,f1341,f1344,f1352,f1354,f1356,f1372,f1360,f1362,f1376,f1833,f1826,f1861,f1819,f1862,f1437,f1438,f1810,f1489,f1515,f1910,f1523,f1570,f1572,f1763,f1949,f1950,f1951,f1626,f1636,f1637,f1655,f1981,f1751,f1993,f1994,f1995,f1703,f1705,f1711,f1739,f1859,f2103,f2038,f2106,f2110,f2112])).
fof(f2111,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f125,f124,f126,f129,f132,f131,f173,f177,f176,f184,f181,f185,f219,f224,f223,f227,f259,f296,f335,f428,f498,f714,f715,f726,f764,f765,f767,f774,f775,f806,f822,f824,f826,f828,f884,f956,f982,f999,f1002,f192,f189,f188,f187,f186,f178,f191,f190,f193,f123,f716,f954,f830,f827,f805,f795,f792,f790,f776,f773,f232,f136,f135,f128,f937,f924,f998,f988,f986,f936,f920,f918,f180,f179,f940,f1004,f926,f938,f939,f1312,f1314,f1330,f1318,f1320,f1321,f1322,f1323,f1331,f1333,f1335,f1341,f1344,f1352,f1354,f1356,f1372,f1360,f1362,f1376,f1833,f1826,f1861,f1819,f1862,f1437,f1438,f1810,f1489,f1515,f1910,f1523,f1570,f1572,f1763,f1949,f1950,f1951,f1626,f1636,f1637,f1655,f1981,f1751,f1993,f1994,f1995,f1703,f1705,f1711,f1739,f1859,f2103,f2038,f2106,f2110])).
fof(f2107,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f125,f124,f126,f129,f132,f131,f173,f177,f176,f184,f181,f185,f219,f224,f223,f227,f259,f296,f335,f428,f498,f714,f715,f726,f764,f765,f767,f774,f775,f806,f822,f824,f826,f828,f884,f956,f982,f999,f1002,f192,f189,f188,f187,f186,f178,f191,f190,f193,f123,f716,f954,f830,f827,f805,f795,f792,f790,f776,f773,f232,f136,f135,f128,f937,f924,f998,f988,f986,f936,f920,f918,f180,f179,f940,f1004,f926,f938,f939,f1312,f1314,f1330,f1318,f1320,f1321,f1322,f1323,f1331,f1333,f1335,f1341,f1344,f1352,f1354,f1356,f1372,f1360,f1362,f1376,f1833,f1826,f1861,f1819,f1862,f1437,f1438,f1810,f1489,f1515,f1910,f1523,f1570,f1572,f1763,f1949,f1950,f1951,f1626,f1636,f1637,f1655,f1981,f1751,f1993,f1994,f1995,f1703,f1705,f1711,f1739,f1859,f2103,f2038,f2106])).
fof(f2105,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f125,f124,f126,f129,f132,f131,f173,f177,f176,f184,f181,f185,f219,f224,f223,f227,f259,f296,f335,f428,f498,f714,f715,f726,f764,f765,f767,f774,f775,f806,f822,f824,f826,f828,f884,f956,f982,f999,f1002,f192,f189,f188,f187,f186,f178,f191,f190,f193,f123,f716,f954,f830,f827,f805,f795,f792,f790,f776,f773,f232,f136,f135,f128,f937,f924,f998,f988,f986,f936,f920,f918,f180,f179,f940,f1004,f926,f938,f939,f1312,f1314,f1330,f1318,f1320,f1321,f1322,f1323,f1331,f1333,f1335,f1341,f1344,f1352,f1354,f1356,f1372,f1360,f1362,f1376,f1833,f1826,f1861,f1819,f1862,f1437,f1438,f1810,f1489,f1515,f1910,f1523,f1570,f1572,f1763,f1949,f1950,f1951,f1626,f1636,f1637,f1655,f1981,f1751,f1993,f1994,f1995,f1703,f1705,f1711,f1739,f1859,f2103,f2038])).
fof(f2104,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f125,f124,f126,f129,f132,f131,f173,f177,f176,f184,f181,f185,f219,f224,f223,f227,f259,f296,f335,f428,f498,f714,f715,f726,f764,f765,f767,f774,f775,f806,f822,f824,f826,f828,f884,f956,f982,f999,f1002,f192,f189,f188,f187,f186,f178,f191,f190,f193,f123,f716,f954,f830,f827,f805,f795,f792,f790,f776,f773,f232,f136,f135,f128,f937,f924,f998,f988,f986,f936,f920,f918,f180,f179,f940,f1004,f926,f938,f939,f1312,f1314,f1330,f1318,f1320,f1321,f1322,f1323,f1331,f1333,f1335,f1341,f1344,f1352,f1354,f1356,f1372,f1360,f1362,f1376,f1833,f1826,f1861,f1819,f1862,f1437,f1438,f1810,f1489,f1515,f1910,f1523,f1570,f1572,f1763,f1949,f1950,f1951,f1626,f1636,f1637,f1655,f1981,f1751,f1993,f1994,f1995,f1703,f1705,f1711,f1739,f1859,f2103])).
fof(f2102,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f2033,f126])).
fof(f2033,plain,(
  ~p(s(bool,'T_0')) | $spl6),
  inference(backward_demodulation,[],[f1859,f128])).
fof(f2101,plain,(
  $false | $spl6),
  inference(trivial_inequality_removal,[],[f2039])).
fof(f2039,plain,(
  s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | $spl6),
  inference(backward_demodulation,[],[f1859,f790])).
fof(f2100,plain,(
  $false | $spl6),
  inference(trivial_inequality_removal,[],[f2041])).
fof(f2041,plain,(
  s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))) | $spl6),
  inference(backward_demodulation,[],[f1859,f795])).
fof(f2099,plain,(
  $false | $spl6),
  inference(trivial_inequality_removal,[],[f2042])).
fof(f2042,plain,(
  s('type.option.option'('type.option.option'('type.option.option'(bool))),'const.option.SOME_1'(s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))))) != s('type.option.option'('type.option.option'('type.option.option'(bool))),'const.option.SOME_1'(s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))))) | $spl6),
  inference(backward_demodulation,[],[f1859,f805])).
fof(f2006,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)))) ) | $spl12),
  inference(cnf_transformation,[],[f2006_D])).
fof(f2006_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f2007,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,'T_0') = s(bool,X1)) ) | $spl14),
  inference(cnf_transformation,[],[f2007_D])).
fof(f2007_D,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,'T_0') = s(bool,X1)) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f1711,plain,(
  ( ! [X0,X1] : (s(bool,X1) != s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,X0) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(equality_factoring,[],[f982])).
fof(f1945,plain,(
  ( ! [X48] : (s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X48))))) | s(bool,X48) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)))) ) | $spl10),
  inference(cnf_transformation,[],[f1945_D])).
fof(f1945_D,plain,(
  ( ! [X48] : (s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X48))))) | s(bool,X48) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f1763,plain,(
  ( ! [X47,X48] : (s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X48))))) | s(bool,'T_0') = s(bool,X47) | s(bool,X48) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)))) )),
  inference(duplicate_literal_removal,[],[f1596])).
fof(f1596,plain,(
  ( ! [X47,X48] : (s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X48))))) | s(bool,'T_0') = s(bool,X47) | s(bool,X48) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f806,f982])).
fof(f1858,plain,(
  ( ! [X27,X25] : (s(bool,X25) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,X25) = s(bool,X27) | s(bool,X27) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)))) ) | $spl4),
  inference(cnf_transformation,[],[f1858_D])).
fof(f1858_D,plain,(
  ( ! [X27,X25] : (s(bool,X25) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,X25) = s(bool,X27) | s(bool,X27) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f1810,plain,(
  ( ! [X17,X15,X16] : (s(bool,'T_0') = s(bool,X16) | s(bool,X17) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,X15) = s(bool,X17) | s(bool,X15) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)))) )),
  inference(duplicate_literal_removal,[],[f1452])).
fof(f1452,plain,(
  ( ! [X17,X15,X16] : (s(bool,'T_0') = s(bool,X16) | s(bool,X17) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,X15) = s(bool,X17) | s(bool,X15) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f982,f982])).
fof(f1819,plain,(
  ( ! [X24,X23,X22] : (s(bool,X23) = s(bool,X24) | s(bool,X24) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,'T_0') = s(bool,X22) | s(bool,X23) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)))) )),
  inference(duplicate_literal_removal,[],[f1427])).
fof(f1427,plain,(
  ( ! [X24,X23,X22] : (s(bool,X23) = s(bool,X24) | s(bool,X24) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,'T_0') = s(bool,X22) | s(bool,X23) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,'T_0') = s(bool,X22)) )),
  inference(superposition,[],[f982,f982])).
fof(f1860,plain,(
  ( ! [X1] : (s(bool,X1) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,'F_0') = s(bool,X1)) ) | $spl8),
  inference(cnf_transformation,[],[f1860_D])).
fof(f1860_D,plain,(
  ( ! [X1] : (s(bool,X1) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,'F_0') = s(bool,X1)) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f1861,plain,(
  ( ! [X14,X15] : (s(bool,'F_0') = s(bool,X15) | s(bool,X15) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,'T_0') = s(bool,X14)) )),
  inference(forward_demodulation,[],[f1821,f232])).
fof(f1821,plain,(
  ( ! [X14,X15] : (s(bool,X15) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,X15) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,'T_0') = s(bool,X14)) )),
  inference(duplicate_literal_removal,[],[f1424])).
fof(f1424,plain,(
  ( ! [X14,X15] : (s(bool,X15) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,X15) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,'T_0') = s(bool,X14) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f982,f335])).
fof(f1826,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(duplicate_literal_removal,[],[f1419])).
fof(f1419,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f982,f129])).
fof(f1833,plain,(
  ( ! [X26,X27,X25] : (s(bool,'T_0') = s(bool,X26) | s(bool,X25) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,X27) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,X25) = s(bool,X27)) )),
  inference(duplicate_literal_removal,[],[f1410])).
fof(f1410,plain,(
  ( ! [X26,X27,X25] : (s(bool,'T_0') = s(bool,X26) | s(bool,X25) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,'T_0') = s(bool,X26) | s(bool,X27) = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2))) | s(bool,X25) = s(bool,X27)) )),
  inference(superposition,[],[f982,f982])).
fof(f1108,plain,(
  $false | $spl2),
  inference(subsumption_resolution,[],[f1014,f126])).
fof(f1014,plain,(
  ~p(s(bool,'T_0')) | $spl2),
  inference(unit_resulting_resolution,[],[f223,f1001,f184])).
fof(f1107,plain,(
  $false | $spl2),
  inference(subsumption_resolution,[],[f1016,f126])).
fof(f1016,plain,(
  ~p(s(bool,'T_0')) | $spl2),
  inference(unit_resulting_resolution,[],[f224,f1001,f184])).
fof(f1106,plain,(
  $false | $spl2),
  inference(subsumption_resolution,[],[f1018,f1001])).
fof(f1018,plain,(
  ~p(s(bool,'F_0')) | $spl2),
  inference(unit_resulting_resolution,[],[f223,f1001,f184])).
fof(f1105,plain,(
  $false | $spl2),
  inference(subsumption_resolution,[],[f1020,f1001])).
fof(f1020,plain,(
  ~p(s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)))) | $spl2),
  inference(unit_resulting_resolution,[],[f224,f1001,f184])).
fof(f1104,plain,(
  $false | $spl2),
  inference(subsumption_resolution,[],[f1024,f126])).
fof(f1024,plain,(
  ~p(s(bool,'T_0')) | $spl2),
  inference(unit_resulting_resolution,[],[f223,f1001,f184])).
fof(f1103,plain,(
  $false | $spl2),
  inference(subsumption_resolution,[],[f1026,f126])).
fof(f1026,plain,(
  ~p(s(bool,'T_0')) | $spl2),
  inference(unit_resulting_resolution,[],[f224,f1001,f184])).
fof(f1102,plain,(
  $false | $spl2),
  inference(subsumption_resolution,[],[f1028,f1001])).
fof(f1028,plain,(
  ~p(s(bool,'F_0')) | $spl2),
  inference(unit_resulting_resolution,[],[f223,f1001,f184])).
fof(f1101,plain,(
  $false | $spl2),
  inference(subsumption_resolution,[],[f1030,f1001])).
fof(f1030,plain,(
  ~p(s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.option.option'(sK0),sK3),s('type.option.option'(sK1),sK2)))) | $spl2),
  inference(unit_resulting_resolution,[],[f224,f1001,f184])).
fof(f1096,plain,(
  $false | $spl2),
  inference(subsumption_resolution,[],[f1095,f1001])).
fof(f1095,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl2),
  inference(subsumption_resolution,[],[f1065,f1034])).
fof(f1034,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f1001,f132])).
fof(f1065,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl2),
  inference(backward_demodulation,[],[f1034,f135])).
fof(f1094,plain,(
  $false | $spl2),
  inference(subsumption_resolution,[],[f1064,f126])).
fof(f1064,plain,(
  ~p(s(bool,'T_0')) | $spl2),
  inference(backward_demodulation,[],[f1034,f128])).
fof(f1093,plain,(
  $false | $spl2),
  inference(trivial_inequality_removal,[],[f1070])).
fof(f1070,plain,(
  s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | $spl2),
  inference(backward_demodulation,[],[f1034,f790])).
fof(f1092,plain,(
  $false | $spl2),
  inference(trivial_inequality_removal,[],[f1072])).
fof(f1072,plain,(
  s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))) | $spl2),
  inference(backward_demodulation,[],[f1034,f795])).
fof(f1091,plain,(
  $false | $spl2),
  inference(trivial_inequality_removal,[],[f1073])).
fof(f1073,plain,(
  s('type.option.option'('type.option.option'('type.option.option'(bool))),'const.option.SOME_1'(s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))))) != s('type.option.option'('type.option.option'('type.option.option'(bool))),'const.option.SOME_1'(s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))))) | $spl2),
  inference(backward_demodulation,[],[f1034,f805])).
fof(f1035,plain,(
  $false | $spl2),
  inference(unit_resulting_resolution,[],[f223,f1001,f132])).
fof(f1033,plain,(
  $false | $spl2),
  inference(unit_resulting_resolution,[],[f232,f1001,f135])).
fof(f1031,plain,(
  $false | $spl2),
  inference(unit_resulting_resolution,[],[f1001,f224,f1001,f184])).
fof(f1029,plain,(
  $false | $spl2),
  inference(unit_resulting_resolution,[],[f1001,f223,f1001,f184])).
fof(f1027,plain,(
  $false | $spl2),
  inference(unit_resulting_resolution,[],[f126,f224,f1001,f184])).
fof(f1025,plain,(
  $false | $spl2),
  inference(unit_resulting_resolution,[],[f126,f223,f1001,f184])).
fof(f1021,plain,(
  $false | $spl2),
  inference(unit_resulting_resolution,[],[f1001,f224,f1001,f184])).
fof(f1019,plain,(
  $false | $spl2),
  inference(unit_resulting_resolution,[],[f1001,f223,f1001,f184])).
fof(f1017,plain,(
  $false | $spl2),
  inference(unit_resulting_resolution,[],[f126,f224,f1001,f184])).
fof(f1015,plain,(
  $false | $spl2),
  inference(unit_resulting_resolution,[],[f126,f223,f1001,f184])).
fof(f1011,plain,(
  $false | $spl2),
  inference(unit_resulting_resolution,[],[f128,f232,f1001,f219])).
fof(f1005,plain,(
  $false | $spl2),
  inference(unit_resulting_resolution,[],[f128,f1001])).
fof(f1000,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | s(bool,'T_0') = s(bool,happ(s(fun(X8,bool),X9),s(X8,sK6(X10,X9,X8))))) ) | $spl0),
  inference(cnf_transformation,[],[f1000_D])).
fof(f1000_D,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | s(bool,'T_0') = s(bool,happ(s(fun(X8,bool),X9),s(X8,sK6(X10,X9,X8))))) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f1002,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X15)) | ~p(s(bool,happ(s(fun(X12,bool),X13),s(X12,X14)))) | s(bool,'T_0') = s(bool,happ(s(fun(X12,bool),X13),s(X12,sK6(X14,X13,X12))))) )),
  inference(subsumption_resolution,[],[f995,f131])).
fof(f995,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X15)) | ~p(s(bool,happ(s(fun(X12,bool),X13),s(X12,X14)))) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,happ(s(fun(X12,bool),X13),s(X12,sK6(X14,X13,X12))))) )),
  inference(superposition,[],[f180,f259])).
fof(f999,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X11)) | ~p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | s(bool,'T_0') = s(bool,happ(s(fun(X8,bool),X9),s(X8,sK6(X10,X9,X8))))) )),
  inference(subsumption_resolution,[],[f994,f131])).
fof(f994,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X11)) | ~p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | s(bool,'T_0') = s(bool,happ(s(fun(X8,bool),X9),s(X8,sK6(X10,X9,X8)))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f180,f259])).
% SZS output end Proof for thm.option
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 216
% Passive clauses: 2111
% Generated clauses: 21052
% Final active clauses: 111
% Final passive clauses: 990
% Input formulas: 31
% Initial clauses: 85
% 
% Duplicate literals: 1246
% Trivial inequalities: 339
% Fw subsumption resolutions: 1762
% Fw demodulations: 1578
% Bw demodulations: 850
% Condensations: 1874
% Global subsumptions: 48
% 
% Simple tautologies: 344
% Equational tautologies: 326
% Forward subsumptions: 11063
% Fw demodulations to eq. taut.: 225
% Bw demodulations to eq. taut.: 36
% 
% Unit resulting resolution: 1717
% Factoring: 1
% Forward superposition: 5943
% Backward superposition: 4363
% Self superposition: 191
% Equality factoring: 148
% Equality resolution: 38
% 
% Split clauses: 83
% Split components: 182
% SAT solver clauses: 2418
% SAT solver unit clauses: 103
% SAT solver binary clauses: 176
% 
% Sat splits: 83
% Sat splitting refutations: 91
% 
% Memory used [KB]: 12920
% Time elapsed: 1.276 s
% ------------------------------
% ------------------------------
% Success in time 2.6 s
