Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.ringNorm
% (82289)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 66
% Passive clauses: 1915
% Generated clauses: 11571
% Final active clauses: 64
% Final passive clauses: 1796
% Input formulas: 21
% Initial clauses: 22
% 
% Duplicate literals: 738
% Trivial inequalities: 107
% Fw subsumption resolutions: 948
% Fw demodulations: 463
% 
% Simple tautologies: 231
% Equational tautologies: 95
% Forward subsumptions: 5682
% 
% Binary resolution: 25
% Forward superposition: 4644
% Backward superposition: 3939
% Self superposition: 345
% Equality factoring: 262
% Equality resolution: 6
% 
% Split clauses: 401
% Split components: 812
% SAT solver clauses: 1265
% SAT solver unit clauses: 428
% SAT solver binary clauses: 401
% 
% Sat splits: 401
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1261
% TWLsolver calls for satisfiability: 44
% 
% Memory used [KB]: 10618
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.ringNorm
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.ringNorm
% SZS output start Proof for thm.ringNorm
fof(f14356,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f14355,f65,f88,f3789,f3814,f71,f96,f1986,f1987,f69,f94,f78,f101,f79,f102])).
fof(f102,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_varlist_insert_3'(s('type.ring.ring'(X0),X1),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_varlist_insert_1'(s('type.ring.ring'(X0),X1))),s('type.list.list'('type.quote.index'),X2))),s('type.canonical.canonical_sum'(X0),X3)))) ) | $spl24),
  inference(cnf_transformation,[],[f102_D])).
fof(f102_D,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_varlist_insert_3'(s('type.ring.ring'(X0),X1),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_varlist_insert_1'(s('type.ring.ring'(X0),X1))),s('type.list.list'('type.quote.index'),X2))),s('type.canonical.canonical_sum'(X0),X3)))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f79,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_varlist_insert_3'(s('type.ring.ring'(X0),X1),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_varlist_insert_1'(s('type.ring.ring'(X0),X1))),s('type.list.list'('type.quote.index'),X2))),s('type.canonical.canonical_sum'(X0),X3)))) )),
  inference(cnf_transformation,[],[f45])).
fof(f45,plain,(
  ! [X0,X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_varlist_insert_3'(s('type.ring.ring'(X0),X1),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_varlist_insert_1'(s('type.ring.ring'(X0),X1))),s('type.list.list'('type.quote.index'),X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(flattening,[],[f44])).
fof(f44,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_varlist_insert_3'(s('type.ring.ring'(X0),X1),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_varlist_insert_1'(s('type.ring.ring'(X0),X1))),s('type.list.list'('type.quote.index'),X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X6] : ! [X19,X20,X21] : s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_varlist_insert_3'(s('type.ring.ring'(X6),X19),s('type.list.list'('type.quote.index'),X20),s('type.canonical.canonical_sum'(X6),X21))) = s('type.canonical.canonical_sum'(X6),happ(s(fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6)),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6))),'const.ringNorm.r_varlist_insert_1'(s('type.ring.ring'(X6),X19))),s('type.list.list'('type.quote.index'),X20))),s('type.canonical.canonical_sum'(X6),X21)))),
  file('thm.ringNorm.varlist_insert_def|split|2','ARITY_1_3_const.ringNorm.r_varlist_insert')).
fof(f101,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.varlist_insert_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.varlist_insert_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.list.list'('type.quote.index'),X2))),s('type.canonical.canonical_sum'(X0),X3)))) ) | $spl22),
  inference(cnf_transformation,[],[f101_D])).
fof(f101_D,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.varlist_insert_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.varlist_insert_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.list.list'('type.quote.index'),X2))),s('type.canonical.canonical_sum'(X0),X3)))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f78,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.varlist_insert_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.varlist_insert_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.list.list'('type.quote.index'),X2))),s('type.canonical.canonical_sum'(X0),X3)))) )),
  inference(cnf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0,X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.canonical.varlist_insert_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.varlist_insert_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.list.list'('type.quote.index'),X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(flattening,[],[f42])).
fof(f42,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.canonical.varlist_insert_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.varlist_insert_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.list.list'('type.quote.index'),X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X6] : ! [X19,X20,X21] : s('type.canonical.canonical_sum'(X6),'const.canonical.varlist_insert_3'(s('type.semi_ring.semi_ring'(X6),X19),s('type.list.list'('type.quote.index'),X20),s('type.canonical.canonical_sum'(X6),X21))) = s('type.canonical.canonical_sum'(X6),happ(s(fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6)),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6))),'const.canonical.varlist_insert_1'(s('type.semi_ring.semi_ring'(X6),X19))),s('type.list.list'('type.quote.index'),X20))),s('type.canonical.canonical_sum'(X6),X21)))),
  file('thm.ringNorm.varlist_insert_def|split|2','ARITY_1_3_const.canonical.varlist_insert')).
fof(f94,plain,(
  ( ! [X0,X1] : (s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_varlist_insert_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.varlist_insert_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))) ) | $spl8),
  inference(cnf_transformation,[],[f94_D])).
fof(f94_D,plain,(
  ( ! [X0,X1] : (s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_varlist_insert_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.varlist_insert_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f69,plain,(
  ( ! [X0,X1] : (s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_varlist_insert_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.varlist_insert_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f27])).
fof(f27,plain,(
  ! [X0,X1] : s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_varlist_insert_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.varlist_insert_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f26])).
fof(f26,plain,(
  ! [X0] : ! [X1] : s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_varlist_insert_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.varlist_insert_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6] : ! [X7] : s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6))),'const.ringNorm.r_varlist_insert_1'(s('type.ring.ring'(X6),X7))) = s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6))),'const.canonical.varlist_insert_1'(s('type.semi_ring.semi_ring'(X6),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X6),X7)))))),
  file('thm.ringNorm.varlist_insert_def|split|2','thm.ringNorm.r_varlist_insert_def')).
fof(f1987,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_varlist_insert_3'(s('type.ring.ring'(X0),X1),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),'const.canonical.varlist_insert_3'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X3)))) ) | $spl298),
  inference(cnf_transformation,[],[f1987_D])).
fof(f1987_D,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_varlist_insert_3'(s('type.ring.ring'(X0),X1),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),'const.canonical.varlist_insert_3'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X3)))) ) <=> ~$spl298),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl298])])).
fof(f1986,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_varlist_insert_3'(s('type.ring.ring'(X0),X1),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),'const.canonical.varlist_insert_3'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X3)))) ) | ($spl8 | $spl22 | $spl24)),
  inference(forward_demodulation,[],[f1955,f102])).
fof(f1955,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.varlist_insert_3'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_varlist_insert_1'(s('type.ring.ring'(X0),X1))),s('type.list.list'('type.quote.index'),X2))),s('type.canonical.canonical_sum'(X0),X3)))) ) | ($spl8 | $spl22)),
  inference(superposition,[],[f101,f94])).
fof(f96,plain,(
  ( ! [X2,X0,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0'))) = s('type.canonical.canonical_sum'(X0),'const.canonical.varlist_insert_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))) ) | $spl12),
  inference(cnf_transformation,[],[f96_D])).
fof(f96_D,plain,(
  ( ! [X2,X0,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0'))) = s('type.canonical.canonical_sum'(X0),'const.canonical.varlist_insert_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f71,plain,(
  ( ! [X2,X0,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0'))) = s('type.canonical.canonical_sum'(X0),'const.canonical.varlist_insert_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))) )),
  inference(cnf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0,X1,X2] : s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0'))) = s('type.canonical.canonical_sum'(X0),'const.canonical.varlist_insert_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))),
  inference(flattening,[],[f30])).
fof(f30,plain,(
  ! [X0] : ! [X1,X2] : s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0'))) = s('type.canonical.canonical_sum'(X0),'const.canonical.varlist_insert_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6] : ! [X12,X8] : s('type.canonical.canonical_sum'(X6),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X8),s('type.canonical.canonical_sum'(X6),'const.canonical.Nil_monom_0'))) = s('type.canonical.canonical_sum'(X6),'const.canonical.varlist_insert_3'(s('type.semi_ring.semi_ring'(X6),X12),s('type.list.list'('type.quote.index'),X8),s('type.canonical.canonical_sum'(X6),'const.canonical.Nil_monom_0')))),
  file('thm.ringNorm.varlist_insert_def|split|2','thm.canonical.varlist_insert_def|split|2')).
fof(f3814,plain,(
  ( ! [X6,X4,X5] : (s('type.canonical.canonical_sum'(X4),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X6),s('type.canonical.canonical_sum'(X4),'const.canonical.Nil_monom_0'))) = s('type.canonical.canonical_sum'(X4),'const.ringNorm.r_varlist_insert_3'(s('type.ring.ring'(X4),X5),s('type.list.list'('type.quote.index'),X6),s('type.canonical.canonical_sum'(X4),'const.canonical.Nil_monom_0')))) ) | $spl516),
  inference(cnf_transformation,[],[f3814_D])).
fof(f3814_D,plain,(
  ( ! [X6,X4,X5] : (s('type.canonical.canonical_sum'(X4),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X6),s('type.canonical.canonical_sum'(X4),'const.canonical.Nil_monom_0'))) = s('type.canonical.canonical_sum'(X4),'const.ringNorm.r_varlist_insert_3'(s('type.ring.ring'(X4),X5),s('type.list.list'('type.quote.index'),X6),s('type.canonical.canonical_sum'(X4),'const.canonical.Nil_monom_0')))) ) <=> ~$spl516),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl516])])).
fof(f3789,plain,(
  ( ! [X6,X4,X5] : (s('type.canonical.canonical_sum'(X4),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X6),s('type.canonical.canonical_sum'(X4),'const.canonical.Nil_monom_0'))) = s('type.canonical.canonical_sum'(X4),'const.ringNorm.r_varlist_insert_3'(s('type.ring.ring'(X4),X5),s('type.list.list'('type.quote.index'),X6),s('type.canonical.canonical_sum'(X4),'const.canonical.Nil_monom_0')))) ) | ($spl12 | $spl298)),
  inference(superposition,[],[f1987,f96])).
fof(f88,plain,(
  s('type.canonical.canonical_sum'(sK0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),sK2),s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0'))) != s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_varlist_insert_3'(s('type.ring.ring'(sK0),sK1),s('type.list.list'('type.quote.index'),sK2),s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0'))) | $spl1),
  inference(cnf_transformation,[],[f88_D])).
fof(f88_D,plain,(
  s('type.canonical.canonical_sum'(sK0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),sK2),s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0'))) != s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_varlist_insert_3'(s('type.ring.ring'(sK0),sK1),s('type.list.list'('type.quote.index'),sK2),s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0'))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f65,plain,(
  s('type.canonical.canonical_sum'(sK0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),sK2),s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0'))) != s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_varlist_insert_3'(s('type.ring.ring'(sK0),sK1),s('type.list.list'('type.quote.index'),sK2),s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0')))),
  inference(cnf_transformation,[],[f62])).
fof(f62,plain,(
  s('type.canonical.canonical_sum'(sK0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),sK2),s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0'))) != s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_varlist_insert_3'(s('type.ring.ring'(sK0),sK1),s('type.list.list'('type.quote.index'),sK2),s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0')))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f60])).
fof(f60,plain,(
  ? [X0,X1,X2] : s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0'))) != s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_varlist_insert_3'(s('type.ring.ring'(X0),X1),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))),
  inference(ennf_transformation,[],[f24])).
fof(f24,plain,(
  ~! [X0,X1,X2] : s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0'))) = s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_varlist_insert_3'(s('type.ring.ring'(X0),X1),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))),
  inference(flattening,[],[f23])).
fof(f23,plain,(
  ~! [X0] : ! [X1,X2] : s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0'))) = s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_varlist_insert_3'(s('type.ring.ring'(X0),X1),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : s('type.canonical.canonical_sum'(X6),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X8),s('type.canonical.canonical_sum'(X6),'const.canonical.Nil_monom_0'))) = s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_varlist_insert_3'(s('type.ring.ring'(X6),X7),s('type.list.list'('type.quote.index'),X8),s('type.canonical.canonical_sum'(X6),'const.canonical.Nil_monom_0')))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : s('type.canonical.canonical_sum'(X6),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X8),s('type.canonical.canonical_sum'(X6),'const.canonical.Nil_monom_0'))) = s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_varlist_insert_3'(s('type.ring.ring'(X6),X7),s('type.list.list'('type.quote.index'),X8),s('type.canonical.canonical_sum'(X6),'const.canonical.Nil_monom_0')))),
  file('thm.ringNorm.varlist_insert_def|split|2','thm.ringNorm.varlist_insert_def|split|2')).
fof(f14355,plain,(
  $false | ($spl1 | $spl516)),
  inference(trivial_inequality_removal,[],[f14319])).
fof(f14319,plain,(
  s('type.canonical.canonical_sum'(sK0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),sK2),s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0'))) != s('type.canonical.canonical_sum'(sK0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),sK2),s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0'))) | ($spl1 | $spl516)),
  inference(superposition,[],[f88,f3814])).
% SZS output end Proof for thm.ringNorm
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 102
% Passive clauses: 884
% Generated clauses: 14183
% Final active clauses: 98
% Final passive clauses: 771
% Input formulas: 21
% Initial clauses: 22
% 
% Duplicate literals: 771
% Trivial inequalities: 24
% Fw subsumption resolutions: 662
% Fw demodulations: 1881
% Bw demodulations: 11
% Condensations: 999
% 
% Simple tautologies: 124
% Equational tautologies: 17
% Forward subsumptions: 7653
% 
% Binary resolution: 9
% Forward superposition: 5301
% Backward superposition: 3447
% Self superposition: 133
% Equality factoring: 32
% Equality resolution: 3
% 
% Split clauses: 12
% Split components: 29
% SAT solver clauses: 1188
% SAT solver unit clauses: 38
% SAT solver binary clauses: 17
% SAT solver learnt clauses: 1
% 
% Sat splits: 12
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 1187
% TWLsolver calls for satisfiability: 119
% 
% Memory used [KB]: 18933
% Time elapsed: 0.673 s
% ------------------------------
% ------------------------------
% Success in time 1.308 s
