Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f5205,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f5197,f2656,f81,f129,f127,f587,f585,f173,f340,f124,f132,f84,f134,f3370,f135,f3372,f80,f128,f126,f2737])).
fof(f2737,plain,(
  $false | ($spl9 | $spl10)),
  inference(resolution,[],[f135,f134])).
fof(f126,plain,(
  p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK0,sK1),sK2),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK1,bool),sK3)))) | $spl0),
  inference(cnf_transformation,[],[f126_D])).
fof(f126_D,plain,(
  p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK0,sK1),sK2),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK1,bool),sK3)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f128,plain,(
  s(fun(sK1,bool),'const.pred_set.EMPTY_0') = s(fun(sK1,bool),sK3) | $spl2),
  inference(cnf_transformation,[],[f128_D])).
fof(f128_D,plain,(
  s(fun(sK1,bool),'const.pred_set.EMPTY_0') = s(fun(sK1,bool),sK3) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f80,plain,(
  s(fun(sK1,bool),'const.pred_set.EMPTY_0') = s(fun(sK1,bool),sK3) | p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK0,sK1),sK2),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK1,bool),sK3))))),
  inference(cnf_transformation,[],[f56])).
fof(f56,plain,(
  (p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK0,sK1),sK2),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK1,bool),sK3)))) | s(fun(sK1,bool),'const.pred_set.EMPTY_0') = s(fun(sK1,bool),sK3)) & (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK0,sK1),sK2),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK1,bool),sK3)))) | s(fun(sK1,bool),'const.pred_set.EMPTY_0') != s(fun(sK1,bool),sK3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f55])).
fof(f55,plain,(
  ? [X0,X1,X2,X3] : ((p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X1,bool),X3)))) | s(fun(X1,bool),'const.pred_set.EMPTY_0') = s(fun(X1,bool),X3)) & (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X1,bool),X3)))) | s(fun(X1,bool),'const.pred_set.EMPTY_0') != s(fun(X1,bool),X3)))),
  inference(nnf_transformation,[],[f53])).
fof(f53,plain,(
  ? [X0,X1,X2,X3] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X1,bool),X3)))) <~> s(fun(X1,bool),'const.pred_set.EMPTY_0') = s(fun(X1,bool),X3))),
  inference(ennf_transformation,[],[f25])).
fof(f25,plain,(
  ~! [X0,X1,X2,X3] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X1,bool),X3)))) <=> s(fun(X1,bool),'const.pred_set.EMPTY_0') = s(fun(X1,bool),X3))),
  inference(flattening,[],[f24])).
fof(f24,plain,(
  ~! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X1,bool),X3)))) <=> s(fun(X1,bool),'const.pred_set.EMPTY_0') = s(fun(X1,bool),X3))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(X6,X7),X8),s(fun(X6,bool),'const.pred_set.EMPTY_0'),s(fun(X7,bool),X9)))) <=> s(fun(X7,bool),'const.pred_set.EMPTY_0') = s(fun(X7,bool),X9))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(X6,X7),X8),s(fun(X6,bool),'const.pred_set.EMPTY_0'),s(fun(X7,bool),X9)))) <=> s(fun(X7,bool),'const.pred_set.EMPTY_0') = s(fun(X7,bool),X9))),
  file('thm.pred_set.BIJ_EMPTY|split|0','thm.pred_set.BIJ_EMPTY|split|0')).
fof(f3372,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.BIJ_3'(s(fun(sK0,sK1),sK2),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK1,bool),sK3))) | $spl134),
  inference(cnf_transformation,[],[f3372_D])).
fof(f3372_D,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.BIJ_3'(s(fun(sK0,sK1),sK2),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK1,bool),sK3))) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f135,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(cnf_transformation,[],[f135_D])).
fof(f135_D,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f3370,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'T_0') = s(bool,'const.pred_set.BIJ_3'(s(fun(sK0,sK1),sK2),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK1,bool),sK3)))) ) | $spl0),
  inference(subsumption_resolution,[],[f3366,f86])).
fof(f86,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f27])).
fof(f27,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X10] : (s(bool,'T_0') = s(bool,X10) <=> p(s(bool,X10)))),
  file('thm.pred_set.BIJ_EMPTY|split|0','thm.bool.EQ_CLAUSES|split|0')).
fof(f3366,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,'const.pred_set.BIJ_3'(s(fun(sK0,sK1),sK2),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK1,bool),sK3)))) ) | $spl0),
  inference(superposition,[],[f126,f142])).
fof(f142,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f85,f85])).
fof(f85,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.pred_set.BIJ_EMPTY|split|0',aHL_BOOL_CASES)).
fof(f134,plain,(
  ~p(s(bool,'F_0')) | $spl9),
  inference(cnf_transformation,[],[f134_D])).
fof(f134_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f84,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f26])).
fof(f26,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.pred_set.BIJ_EMPTY|split|0',aHL_FALSITY)).
fof(f132,plain,(
  p(s(bool,'T_0')) | $spl6),
  inference(cnf_transformation,[],[f132_D])).
fof(f132_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f124,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f109])).
fof(f109,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f39])).
fof(f39,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f38])).
fof(f38,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X6] : ! [X11] : (s(X6,X11) = s(X6,X11) <=> p(s(bool,'T_0')))),
  file('thm.pred_set.BIJ_EMPTY|split|0','thm.bool.REFL_CLAUSE')).
fof(f340,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(X22,X23),X24),s(fun(X22,bool),'const.pred_set.EMPTY_0'),s(fun(X23,bool),X25)))) ) | $spl12),
  inference(cnf_transformation,[],[f340_D])).
fof(f340_D,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(X22,X23),X24),s(fun(X22,bool),'const.pred_set.EMPTY_0'),s(fun(X23,bool),X25)))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f173,plain,(
  ( ! [X6,X8,X7,X5] : (s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(X5,X6),X7),s(fun(X5,bool),'const.pred_set.EMPTY_0'),s(fun(X6,bool),X8)))) )),
  inference(resolution,[],[f113,f87])).
fof(f87,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f57])).
fof(f113,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X1,bool),X3))))) )),
  inference(cnf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0,X1,X2,X3] : p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X1,bool),X3))))),
  inference(flattening,[],[f41])).
fof(f41,plain,(
  ! [X0,X1] : ! [X2,X3] : p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X1,bool),X3))))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X6,X7] : ! [X8,X9] : p(s(bool,'const.pred_set.INJ_3'(s(fun(X6,X7),X8),s(fun(X6,bool),'const.pred_set.EMPTY_0'),s(fun(X7,bool),X9))))),
  file('thm.pred_set.BIJ_EMPTY|split|0','thm.pred_set.INJ_EMPTY|split|0')).
fof(f585,plain,(
  ( ! [X10,X8,X9] : (s(bool,'T_0') = s(bool,'const.pred_set.SURJ_3'(s(fun(X8,X9),X10),s(fun(X8,bool),'const.pred_set.EMPTY_0'),s(fun(X9,bool),'const.pred_set.EMPTY_0')))) ) | $spl22),
  inference(cnf_transformation,[],[f585_D])).
fof(f585_D,plain,(
  ( ! [X10,X8,X9] : (s(bool,'T_0') = s(bool,'const.pred_set.SURJ_3'(s(fun(X8,X9),X10),s(fun(X8,bool),'const.pred_set.EMPTY_0'),s(fun(X9,bool),'const.pred_set.EMPTY_0')))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f587,plain,(
  ( ! [X17,X18,X16] : (s(bool,'T_0') = s(bool,'const.pred_set.SURJ_3'(s(fun(X16,X17),X18),s(fun(X16,bool),'const.pred_set.EMPTY_0'),s(fun(X17,bool),'const.pred_set.EMPTY_0')))) ) | $spl9),
  inference(subsumption_resolution,[],[f583,f134])).
fof(f583,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.pred_set.SURJ_3'(s(fun(X16,X17),X18),s(fun(X16,bool),'const.pred_set.EMPTY_0'),s(fun(X17,bool),'const.pred_set.EMPTY_0')))) )),
  inference(superposition,[],[f559,f85])).
fof(f559,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X1,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(equality_resolution,[],[f118])).
fof(f118,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun(X1,bool),'const.pred_set.EMPTY_0') != s(fun(X1,bool),X3) | p(s(bool,'const.pred_set.SURJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.pred_set.SURJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) | s(fun(X1,bool),'const.pred_set.EMPTY_0') = s(fun(X1,bool),X3)) & (s(fun(X1,bool),'const.pred_set.EMPTY_0') != s(fun(X1,bool),X3) | p(s(bool,'const.pred_set.SURJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))))),
  inference(nnf_transformation,[],[f48])).
fof(f48,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) <=> s(fun(X1,bool),'const.pred_set.EMPTY_0') = s(fun(X1,bool),X3))),
  inference(flattening,[],[f47])).
fof(f47,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) <=> s(fun(X1,bool),'const.pred_set.EMPTY_0') = s(fun(X1,bool),X3))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X7,X6] : ! [X8,X9] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(X6,X7),X8),s(fun(X6,bool),X9),s(fun(X7,bool),'const.pred_set.EMPTY_0')))) <=> s(fun(X6,bool),'const.pred_set.EMPTY_0') = s(fun(X6,bool),X9))),
  file('thm.pred_set.BIJ_EMPTY|split|0','thm.pred_set.SURJ_EMPTY|split|1')).
fof(f127,plain,(
  ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK0,sK1),sK2),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK1,bool),sK3)))) | $spl1),
  inference(cnf_transformation,[],[f127_D])).
fof(f127_D,plain,(
  ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK0,sK1),sK2),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK1,bool),sK3)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f129,plain,(
  s(fun(sK1,bool),'const.pred_set.EMPTY_0') != s(fun(sK1,bool),sK3) | $spl3),
  inference(cnf_transformation,[],[f129_D])).
fof(f129_D,plain,(
  s(fun(sK1,bool),'const.pred_set.EMPTY_0') != s(fun(sK1,bool),sK3) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f81,plain,(
  s(fun(sK1,bool),'const.pred_set.EMPTY_0') != s(fun(sK1,bool),sK3) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK0,sK1),sK2),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK1,bool),sK3))))),
  inference(cnf_transformation,[],[f56])).
fof(f2656,plain,(
  $false | ($spl1 | $spl2 | $spl6 | $spl9 | $spl12 | $spl22)),
  inference(subsumption_resolution,[],[f2645,f134])).
fof(f2645,plain,(
  p(s(bool,'F_0')) | ($spl1 | $spl2 | $spl6 | $spl9 | $spl12 | $spl22)),
  inference(superposition,[],[f2410,f169])).
fof(f169,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.BIJ_3'(s(fun(sK0,sK1),sK2),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0'))) | ($spl1 | $spl2)),
  inference(resolution,[],[f163,f91])).
fof(f91,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f30])).
fof(f30,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f29])).
fof(f29,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X10] : (s(bool,'F_0') = s(bool,X10) <=> ~p(s(bool,X10)))),
  file('thm.pred_set.BIJ_EMPTY|split|0','thm.bool.EQ_CLAUSES|split|3')).
fof(f163,plain,(
  ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK0,sK1),sK2),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0')))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f127,f128])).
fof(f2410,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(X2,X3),X4),s(fun(X2,bool),'const.pred_set.EMPTY_0'),s(fun(X3,bool),'const.pred_set.EMPTY_0'))))) ) | ($spl6 | $spl9 | $spl12 | $spl22)),
  inference(subsumption_resolution,[],[f2409,f136])).
fof(f136,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl9),
  inference(resolution,[],[f86,f134])).
fof(f2409,plain,(
  ( ! [X4,X2,X3] : (s(bool,'T_0') = s(bool,'F_0') | p(s(bool,'const.pred_set.BIJ_3'(s(fun(X2,X3),X4),s(fun(X2,bool),'const.pred_set.EMPTY_0'),s(fun(X3,bool),'const.pred_set.EMPTY_0'))))) ) | ($spl6 | $spl12 | $spl22)),
  inference(forward_demodulation,[],[f2408,f340])).
fof(f2408,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(X2,X3),X4),s(fun(X2,bool),'const.pred_set.EMPTY_0'),s(fun(X3,bool),'const.pred_set.EMPTY_0')))) | s(bool,'F_0') = s(bool,'const.pred_set.INJ_3'(s(fun(X2,X3),X4),s(fun(X2,bool),'const.pred_set.EMPTY_0'),s(fun(X3,bool),'const.pred_set.EMPTY_0')))) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f2390,f132])).
fof(f2390,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.pred_set.BIJ_3'(s(fun(X2,X3),X4),s(fun(X2,bool),'const.pred_set.EMPTY_0'),s(fun(X3,bool),'const.pred_set.EMPTY_0')))) | s(bool,'F_0') = s(bool,'const.pred_set.INJ_3'(s(fun(X2,X3),X4),s(fun(X2,bool),'const.pred_set.EMPTY_0'),s(fun(X3,bool),'const.pred_set.EMPTY_0')))) ) | $spl22),
  inference(superposition,[],[f1270,f585])).
fof(f1270,plain,(
  ( ! [X14,X12,X15,X13,X11] : (~p(s(bool,'const.pred_set.SURJ_3'(s(fun(X11,X12),X13),s(fun(X11,bool),X14),s(fun(X12,bool),X15)))) | p(s(bool,'const.pred_set.BIJ_3'(s(fun(X11,X12),X13),s(fun(X11,bool),X14),s(fun(X12,bool),X15)))) | s(bool,'F_0') = s(bool,'const.pred_set.INJ_3'(s(fun(X11,X12),X13),s(fun(X11,bool),X14),s(fun(X12,bool),X15)))) )),
  inference(resolution,[],[f123,f91])).
fof(f123,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | ~p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))) )),
  inference(cnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | (p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) & p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))))) & (~p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | ~p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))),
  inference(flattening,[],[f78])).
fof(f78,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | (p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) & p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))))) & ((~p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | ~p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))) | p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))),
  inference(nnf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) <=> (p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) & p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))),
  inference(flattening,[],[f51])).
fof(f51,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) <=> (p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) & p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X6,X7] : ! [X8,X9,X10] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(X6,X7),X8),s(fun(X6,bool),X9),s(fun(X7,bool),X10)))) <=> (p(s(bool,'const.pred_set.INJ_3'(s(fun(X6,X7),X8),s(fun(X6,bool),X9),s(fun(X7,bool),X10)))) & p(s(bool,'const.pred_set.SURJ_3'(s(fun(X6,X7),X8),s(fun(X6,bool),X9),s(fun(X7,bool),X10))))))),
  file('thm.pred_set.BIJ_EMPTY|split|0','thm.pred_set.BIJ_DEF')).
fof(f5197,plain,(
  $false | ($spl3 | $spl6 | $spl134)),
  inference(subsumption_resolution,[],[f5188,f129])).
fof(f5188,plain,(
  s(fun(sK1,bool),'const.pred_set.EMPTY_0') = s(fun(sK1,bool),sK3) | ($spl6 | $spl134)),
  inference(resolution,[],[f5178,f119])).
fof(f119,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X1,bool),X3)))) | s(fun(X1,bool),'const.pred_set.EMPTY_0') = s(fun(X1,bool),X3)) )),
  inference(cnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X1,bool),X3)))) | s(fun(X1,bool),'const.pred_set.EMPTY_0') = s(fun(X1,bool),X3)) & (s(fun(X1,bool),'const.pred_set.EMPTY_0') != s(fun(X1,bool),X3) | p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X1,bool),X3))))))),
  inference(nnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X1,bool),X3)))) <=> s(fun(X1,bool),'const.pred_set.EMPTY_0') = s(fun(X1,bool),X3))),
  inference(flattening,[],[f49])).
fof(f49,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X1,bool),X3)))) <=> s(fun(X1,bool),'const.pred_set.EMPTY_0') = s(fun(X1,bool),X3))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X6,X7] : ! [X8,X9] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(X6,X7),X8),s(fun(X6,bool),'const.pred_set.EMPTY_0'),s(fun(X7,bool),X9)))) <=> s(fun(X7,bool),'const.pred_set.EMPTY_0') = s(fun(X7,bool),X9))),
  file('thm.pred_set.BIJ_EMPTY|split|0','thm.pred_set.SURJ_EMPTY|split|0')).
fof(f5178,plain,(
  p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK0,sK1),sK2),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK1,bool),sK3)))) | ($spl6 | $spl134)),
  inference(subsumption_resolution,[],[f5168,f132])).
fof(f5168,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK0,sK1),sK2),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK1,bool),sK3)))) | $spl134),
  inference(superposition,[],[f122,f3372])).
fof(f122,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))) )),
  inference(cnf_transformation,[],[f79])).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 127
% Passive clauses: 625
% Generated clauses: 5053
% Final active clauses: 69
% Final passive clauses: 339
% Input formulas: 22
% Initial clauses: 44
% 
% Duplicate literals: 313
% Trivial inequalities: 97
% Fw subsumption resolutions: 722
% Fw demodulations: 445
% Bw demodulations: 45
% 
% Simple tautologies: 190
% Equational tautologies: 26
% Forward subsumptions: 2311
% 
% Binary resolution: 266
% Forward superposition: 1752
% Backward superposition: 1100
% Self superposition: 44
% Equality factoring: 32
% Equality resolution: 22
% 
% Split clauses: 195
% Split components: 401
% SAT solver clauses: 438
% SAT solver unit clauses: 17
% SAT solver binary clauses: 79
% SAT solver learnt clauses: 2
% SAT solver learnt literals: 1
% 
% Sat splits: 195
% Sat splitting refutations: 9
% 
% TWLsolver clauses: 428
% TWLsolver calls for satisfiability: 125
% 
% Memory used [KB]: 2814
% Time elapsed: 0.117 s
% ------------------------------
% ------------------------------
% Success in time 0.119 s
