Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.list
% (110681)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 698
% Passive clauses: 9962
% Generated clauses: 21125
% Final active clauses: 600
% Final passive clauses: 8746
% Input formulas: 69
% Initial clauses: 225
% 
% Duplicate literals: 280
% Trivial inequalities: 160
% Fw subsumption resolutions: 1806
% Fw demodulations: 750
% Bw demodulations: 96
% 
% Simple tautologies: 223
% Equational tautologies: 66
% Forward subsumptions: 7759
% 
% Binary resolution: 10269
% Factoring: 1012
% Forward superposition: 2889
% Backward superposition: 3408
% Self superposition: 34
% Equality factoring: 26
% Equality resolution: 33
% 
% Split clauses: 278
% Split components: 557
% SAT solver clauses: 546
% SAT solver unit clauses: 20
% SAT solver binary clauses: 133
% 
% Sat splits: 278
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 544
% TWLsolver calls for satisfiability: 316
% 
% Memory used [KB]: 10746
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.list
% (110708)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 544
% Passive clauses: 3731
% Generated clauses: 17557
% Final active clauses: 449
% Final passive clauses: 2112
% Input formulas: 69
% Initial clauses: 225
% 
% Duplicate literals: 192
% Trivial inequalities: 61
% Fw subsumption resolutions: 1674
% Fw demodulations: 1438
% Bw demodulations: 1105
% Condensations: 257
% Global subsumptions: 16
% 
% Simple tautologies: 270
% Equational tautologies: 19
% Forward subsumptions: 6556
% Bw demodulations to eq. taut.: 51
% 
% Binary resolution: 4071
% Factoring: 116
% Forward superposition: 2094
% Backward superposition: 2451
% Self superposition: 31
% Equality factoring: 26
% Equality resolution: 4
% 
% Split clauses: 44
% Split components: 88
% SAT solver clauses: 3202
% SAT solver unit clauses: 233
% SAT solver binary clauses: 80
% SAT solver learnt clauses: 32
% SAT solver learnt literals: 10
% 
% Sat splits: 44
% Sat splitting refutations: 37
% 
% TWLsolver clauses: 3138
% TWLsolver calls for satisfiability: 1919
% 
% Memory used [KB]: 14200
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.list
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.list
% SZS output start Proof for thm.list
fof(f58052,plain,(
  $false),
  inference(subsumption_resolution,[],[f58051,f12261])).
fof(f12261,plain,(
  ~p(s(bool,'F_0'))),
  inference(global_subsumption,[],[f349,f348,f350,f357,f358,f360,f374,f373,f407,f408,f409,f418,f419,f420,f421,f422,f454,f451,f465,f464,f587,f483,f480,f479,f478,f497,f504,f503,f502,f501,f500,f499,f498,f508,f507,f506,f505,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f574,f524,f523,f522,f521,f531,f530,f529,f528,f537,f534,f533,f532,f544,f543,f542,f573,f540,f539,f538,f548,f547,f546,f545,f552,f551,f550,f549,f561,f560,f559,f572,f557,f556,f555,f565,f564,f563,f562,f569,f568,f567,f566,f589,f588,f591,f590,f593,f592,f595,f594,f597,f596,f599,f598,f601,f600,f603,f602,f605,f604,f607,f606,f609,f608,f611,f610,f613,f612,f615,f614,f617,f616,f619,f618,f621,f620,f623,f622,f625,f624,f627,f626,f636,f635,f634,f633,f632,f631,f630,f629,f628,f644,f643,f642,f647,f683,f685,f688,f691,f694,f697,f706,f711,f714,f741,f740,f737,f736,f732,f731,f730,f729,f728,f727,f726,f725,f724,f723,f722,f721,f720,f719,f717,f715,f756,f760,f768,f779,f788,f794,f803,f806,f810,f813,f816,f825,f826,f827,f835,f838,f842,f845,f848,f869,f867,f866,f872,f887,f885,f890,f902,f901,f899,f898,f966,f965,f964,f963,f962,f961,f960,f959,f958,f957,f956,f955,f951,f949,f945,f944,f940,f939,f935,f934,f933,f929,f925,f924,f923,f922,f921,f920,f919,f918,f917,f916,f915,f914,f913,f911,f910,f909,f908,f907,f974,f983,f1059,f1099,f1098,f1097,f1096,f1109,f1108,f1107,f1106,f1122,f1121,f1118,f1117,f1135,f1134,f1133,f1132,f1150,f1149,f1148,f1147,f1185,f1186,f1187,f1188,f1189,f1190,f1191,f1192,f1193,f1194,f1216,f1217,f1230,f1241,f1242,f1257,f1268,f1269,f1273,f1279,f1288,f1293,f1292,f1307,f1318,f1319,f1323,f1353,f1354,f1358,f1365,f1380,f1388,f1387,f1386,f1400,f1410,f1411,f1415,f1428,f1454,f1455,f1461,f1476,f1492,f1503,f1502,f1501,f1516,f1527,f1528,f1532,f1546,f1570,f1571,f1582,f1581,f1580,f1579,f1646,f1658,f1665,f1664,f1663,f1662,f1693,f1694,f1695,f1696,f1697,f1698,f1699,f1700,f1701,f1702,f1703,f1704,f1705,f1706,f1707,f1708,f1709,f1710,f1711,f1746,f1774,f1840,f1839,f1835,f1834,f1828,f1824,f1823,f1822,f1821,f1820,f1819,f1818,f1912,f1911,f1910,f1902,f1899,f1898,f1897,f1893,f1888,f1995,f1994,f1984,f1983,f1982,f1973,f1958,f2058,f2057,f2056,f2055,f2225,f2226,f2227,f2228,f2229,f2230,f2232,f2505,f2502,f2501,f2498,f2495,f2492,f2491,f2490,f2489,f2488,f2487,f2486,f2485,f2484,f2483,f2482,f2481,f2480,f2479,f2478,f2477,f2473,f2472,f2468,f2465,f2464,f2458,f2457,f2450,f2449,f2448,f2426,f2425,f2424,f2420,f2419,f2415,f2414,f2413,f2408,f2400,f2396,f2392,f2391,f2386,f2385,f2384,f2383,f2382,f2379,f2377,f2372,f2371,f2365,f2364,f2357,f2355,f2354,f2349,f2346,f2345,f2344,f2343,f2342,f2341,f2340,f2339,f2338,f2337,f2336,f2335,f2334,f2333,f2332,f2331,f2330,f2329,f2328,f2327,f2326,f2325,f2319,f2318,f2311,f2310,f2304,f2303,f2296,f2295,f2294,f2293,f2292,f2287,f2286,f2283,f2282,f2281,f2280,f2278,f2277,f2276,f2275,f2274,f2559,f2558,f2556,f2555,f2570,f2571,f2663,f2669,f2683,f2689,f2710,f2716,f2810,f2821,f2841,f2850,f2878,f2887,f2907,f2916,f2919,f2965,f2970,f3000,f3041,f3063,f3079,f3091,f3130,f3139,f3165,f3177,f3246,f3254,f3271,f3279,f3307,f3317,f3346,f3355,f3388,f3402,f3429,f3437,f3463,f3473,f3505,f3514,f3554,f3562,f3580,f3589,f3599,f3597,f3595,f3616,f3624,f3634,f3633,f3632,f3630,f3654,f3662,f3674,f3673,f3696,f3706,f3718,f3717,f3716,f3715,f3713,f3739,f3753,f3759,f3780,f3788,f3794,f3815,f3824,f3831,f3849,f3872,f3921,f4103,f4182,f4220,f4227,f4279,f4274,f4296,f4339,f4337,f4336,f4437,f4462,f4461,f4460,f4479,f4501,f4500,f4498,f4544,f4541,f4538,f4535,f4585,f4584,f4583,f4582,f4624,f4623,f4622,f4621,f4620,f4619,f4618,f4765,f4846,f4865,f4872,f4897,f4976,f4984,f4991,f5014,f5009,f5040,f5101,f5099,f5098,f5110,f5220,f5273,f5313,f5372,f5371,f5369,f5368,f5367,f5387,f5431,f5457,f5458,f5435,f5436,f5438,f5445,f5477,f5481,f5482,f5484,f5491,f5507,f5546,f5547,f5522,f5530,f5531,f5533,f5558,f5579,f5584,f5585,f5587,f5603,f5615,f5628,f5627,f5623,f5740,f5741,f5742,f5743,f5744,f5745,f5746,f5747,f5748,f5749,f5750,f5751,f5752,f5753,f5754,f5755,f5756,f5665,f5666,f5667,f5668,f5774,f5781,f5799,f5807,f5826,f5833,f5847,f5857,f5863,f5872,f5868,f5881,f5923,f5925,f5926,f6345,f6391,f6543,f6958,f6957,f6956,f6955,f6954,f6950,f6945,f6944,f6939,f6934,f6929,f6928,f6927,f6926,f6925,f6924,f6923,f6922,f6921,f6920,f6919,f6918,f6917,f6916,f6915,f6914,f6913,f6912,f6911,f6910,f6909,f6908,f6907,f6906,f6905,f6904,f6898,f6897,f6893,f6888,f6887,f6876,f6875,f6863,f6862,f6861,f6860,f6817,f6816,f6815,f6809,f6808,f6802,f6801,f6800,f6791,f6772,f6761,f6756,f6755,f6749,f6748,f6747,f6746,f6745,f6742,f6740,f6735,f6734,f6723,f6722,f6710,f6706,f6705,f6704,f6703,f6702,f6701,f6694,f6692,f6690,f6687,f6686,f6685,f6684,f6683,f6682,f6681,f6680,f6679,f6678,f6677,f6676,f6675,f6674,f6673,f6672,f6671,f6670,f6669,f6668,f6667,f6666,f6665,f6664,f6663,f6662,f6661,f6660,f6659,f6658,f6657,f6656,f6655,f6654,f6653,f6652,f6641,f6640,f6628,f6627,f6616,f6615,f6603,f6602,f6601,f6600,f6599,f6598,f6597,f6596,f6595,f6594,f6593,f6592,f6591,f6590,f6589,f6588,f6587,f6569,f6568,f6565,f6564,f6563,f6562,f6560,f6559,f6558,f6557,f6556,f7063,f7490,f7489,f7488,f7487,f7486,f7482,f7477,f7476,f7471,f7466,f7461,f7460,f7459,f7458,f7457,f7456,f7455,f7454,f7453,f7452,f7451,f7450,f7449,f7448,f7447,f7446,f7445,f7444,f7443,f7442,f7441,f7440,f7439,f7438,f7437,f7436,f7430,f7429,f7425,f7420,f7419,f7408,f7407,f7395,f7394,f7393,f7392,f7349,f7348,f7347,f7341,f7340,f7334,f7333,f7332,f7323,f7304,f7293,f7288,f7287,f7281,f7280,f7279,f7278,f7277,f7274,f7272,f7267,f7266,f7255,f7254,f7242,f7238,f7237,f7236,f7235,f7234,f7233,f7226,f7224,f7222,f7219,f7218,f7217,f7216,f7215,f7214,f7213,f7212,f7211,f7210,f7209,f7208,f7207,f7206,f7205,f7204,f7203,f7202,f7201,f7200,f7199,f7198,f7197,f7196,f7195,f7194,f7193,f7192,f7191,f7190,f7189,f7188,f7187,f7186,f7185,f7184,f7173,f7172,f7160,f7159,f7148,f7147,f7135,f7134,f7133,f7132,f7131,f7130,f7129,f7128,f7127,f7126,f7125,f7124,f7123,f7122,f7121,f7120,f7119,f7101,f7100,f7097,f7096,f7095,f7094,f7092,f7091,f7090,f7089,f7088,f8107,f8106,f8105,f8104,f8103,f8099,f8094,f8093,f8088,f8083,f8078,f8077,f8076,f8075,f8074,f8073,f8072,f8071,f8070,f8069,f8068,f8067,f8066,f8065,f8064,f8063,f8061,f8060,f8059,f8058,f8057,f8056,f8055,f8054,f8050,f8049,f8041,f8039,f8034,f8027,f8026,f8015,f8014,f8002,f8001,f8000,f7999,f7951,f7948,f7946,f7939,f7937,f7930,f7929,f7928,f7919,f7896,f7883,f7877,f7876,f7869,f7866,f7865,f7864,f7863,f7858,f7856,f7849,f7848,f7837,f7836,f7824,f7817,f7816,f7815,f7814,f7804,f7802,f7798,f7796,f7793,f7792,f7791,f7790,f7789,f7788,f7787,f7786,f7783,f7782,f7781,f7780,f7778,f7777,f7776,f7775,f7774,f7773,f7772,f7771,f7770,f7769,f7767,f7766,f7765,f7764,f7763,f7762,f7761,f7760,f7759,f7758,f7747,f7746,f7734,f7733,f7722,f7721,f7709,f7703,f7702,f7701,f7699,f7698,f7697,f7694,f7693,f7692,f7632,f7631,f7625,f7623,f7621,f7620,f7617,f7616,f7615,f7614,f7613,f8349,f8371,f8389,f8439,f8440,f8441,f8442,f8639,f8872,f8873,f8874,f8875,f8876,f8877,f8878,f8879,f9970,f9974,f9981,f9984,f9987,f9990,f9993,f9996,f9999,f10002,f10005,f10008,f10011,f10014,f10017,f10020,f10023,f10026,f10029,f8909,f8910,f8911,f8912,f10031,f10032,f8931,f10033,f10036,f10037,f10038,f10039,f10040,f10041,f10042,f10043,f10044,f10045,f10046,f10047,f10048,f10049,f10050,f10051,f10052,f10053,f10054,f10055,f9008,f9009,f9010,f9011,f9204,f10079,f10081,f10087,f10089,f10091,f10093,f10095,f10097,f10099,f10101,f10103,f10105,f10107,f10109,f10111,f10113,f10115,f10117,f10119,f8387,f8347,f5592,f5538,f5489,f5443,f1619,f1618,f1617,f1616,f1615,f1614,f1613,f1612,f1611,f1610,f1609,f1608,f1607,f1606,f1605,f1604,f1603,f1602,f1601,f6551,f8928,f7083,f8929,f9146,f9145,f8581,f8580,f9144,f9143,f8930,f8927,f8582,f8579,f8578,f8926,f9147,f7608,f447,f448,f446,f10076,f10077,f8388,f8348,f5593,f5539,f5490,f5444,f2613,f2612,f2611,f2610,f2609,f2608,f2607,f2606,f2605,f2604,f2603,f2602,f2601,f2600,f2599,f2598,f2597,f2596,f2595,f8384,f8342,f5588,f5534,f5485,f5439,f2094,f2093,f2092,f2091,f2090,f2089,f2088,f2087,f2086,f2085,f2084,f2083,f2082,f2081,f2080,f2079,f2078,f2077,f2076,f449,f9202,f8637,f2702,f2899,f2699,f2703,f2900,f10071,f2902,f2901,f2898,f2708,f2704,f2701,f444,f429,f426,f3542,f3302,f3499,f347,f10070,f2660,f3266,f2126,f3264,f584,f2127,f10069,f365,f3268,f10068,f2726,f3574,f2725,f3572,f576,f2727,f10067,f411,f3576,f10066,f8924,f9198,f9203,f9201,f9200,f9141,f8638,f8636,f8635,f8633,f8576,f2680,f3155,f2801,f2764,f2659,f2658,f2679,f3154,f2800,f2763,f2166,f2144,f3459,f3458,f3424,f3423,f3339,f3338,f3377,f3376,f2681,f430,f10065,f3151,f425,f10073,f3158,f3159,f3157,f3156,f3160,f427,f10062,f10075,f10074,f10072,f10064,f10063,f10061,f10060,f10059,f10058,f10057,f10056,f10035,f10034,f5737,f5736,f3500,f3460,f3425,f3378,f3340,f3303,f2865,f2864,f2863,f2862,f2861,f2860,f2836,f2835,f2834,f2833,f2832,f2831,f2802,f2765,f2167,f2145,f445,f443,f442,f441,f440,f439,f438,f436,f434,f432,f428,f415,f414,f8630,f9195,f7779,f7768,f7708,f7707,f7706,f7705,f7704,f7700,f7695,f7672,f7671,f7670,f7668,f7667,f7666,f7664,f7663,f7662,f5823,f5591,f5590,f5589,f5586,f5583,f5550,f5549,f5548,f5545,f5537,f5536,f5535,f5532,f5529,f5526,f5488,f5487,f5486,f5483,f5480,f5456,f5452,f5451,f5450,f5442,f5441,f5440,f5437,f5434,f5427,f4503,f4205,f4204,f4203,f4202,f4201,f4200,f4199,f4198,f4197,f4196,f4195,f4194,f4193,f4192,f4191,f4190,f4189,f4186,f4185,f3714,f3675,f3631,f3198,f3197,f3186,f3185,f3152,f3069,f3068,f3039,f3038,f3037,f3036,f3035,f3009,f3008,f2987,f2986,f2985,f2959,f2931,f2930,f2913,f2906,f2892,f2886,f2875,f2874,f2849,f2677,f2572,f2216,f2215,f2205,f2204,f2164,f2142,f2067,f1676,f1584,f716,f668,f667,f585,f378,f377,f12257,f12260])).
fof(f12260,plain,(
  ( ! [X9] : (~p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,X9)) )),
  inference(subsumption_resolution,[],[f11820,f360])).
fof(f11820,plain,(
  ( ! [X9] : (~p(s(bool,'F_0')) | s(bool,'F_0') != s(bool,X9) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f377,f360])).
fof(f12257,plain,(
  ( ! [X22] : (~p(s(bool,X22)) | s(bool,'F_0') != s(bool,X22)) )),
  inference(forward_demodulation,[],[f11830,f419])).
fof(f11830,plain,(
  ( ! [X23,X22] : (~p(s(bool,X22)) | s(bool,'F_0') != s(bool,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X22),s('type.list.list'(bool),X23)))))) )),
  inference(superposition,[],[f377,f419])).
fof(f377,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f245])).
fof(f245,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f94])).
fof(f94,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f93])).
fof(f93,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X15] : (s(bool,'F_0') = s(bool,X15) <=> ~p(s(bool,X15)))),
  file('thm.list.MONO_EVERY','thm.bool.EQ_CLAUSES|split|3')).
fof(f378,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f245])).
fof(f585,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f353,f350])).
fof(f353,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f239])).
fof(f239,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f73])).
fof(f73,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f67])).
fof(f67,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.list.MONO_EVERY','thm.bool.NOT_CLAUSES|split|1')).
fof(f667,plain,(
  s(bool,'F_0') = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)))),
  inference(unit_resulting_resolution,[],[f349,f378])).
fof(f668,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f407,f378])).
fof(f716,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f418,f378])).
fof(f1584,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(bool),X0))) | s(bool,'F_0') = s(bool,'const.list.HD_1'(s('type.list.list'(bool),X0)))) )),
  inference(superposition,[],[f409,f360])).
fof(f1676,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X0),s('type.list.list'(bool),X1))))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f419,f360])).
fof(f2067,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))) | s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num',X3)))) )),
  inference(superposition,[],[f449,f360])).
fof(f2142,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X1)))) | s(bool,'T_0') = s(bool,sK19(X0,bool))) )),
  inference(superposition,[],[f414,f360])).
fof(f2164,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X1)))) | s(bool,'T_0') = s(bool,sK20(X0,bool))) )),
  inference(superposition,[],[f415,f360])).
fof(f2204,plain,(
  ( ! [X10,X11,X9] : (s(bool,'F_0') != s(bool,X10) | p(s(bool,'const.bool.IN_2'(s(bool,X9),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X10),s(fun(bool,bool),X11)))))) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f447,f360])).
fof(f2205,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X13),s(fun(bool,bool),X14)))))) | s(bool,X12) != s(bool,X13) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f447,f360])).
fof(f2215,plain,(
  ( ! [X10,X11,X9] : (s(bool,'F_0') != s(bool,X10) | p(s(bool,'const.bool.IN_2'(s(bool,X10),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X9),s(fun(bool,bool),X11)))))) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f447,f360])).
fof(f2216,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,'const.bool.IN_2'(s(bool,X13),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X14)))))) | s(bool,X12) != s(bool,X13) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f447,f360])).
fof(f2572,plain,(
  ( ! [X6,X5] : (s(bool,'T_0') = s(bool,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(bool),'const.list.TL_1'(s('type.list.list'(bool),X6))))) | s(bool,'F_0') = s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))),s('type.list.list'(bool),X6)))) )),
  inference(superposition,[],[f422,f360])).
fof(f2677,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0')))) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X1),s('type.list.list'(bool),X0)))) | s(bool,'T_0') = s(bool,sK22(X0,X1,bool))) )),
  inference(superposition,[],[f430,f360])).
fof(f2849,plain,(
  ( ! [X0] : (~sP5('T_0','F_0','const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))))) )),
  inference(unit_resulting_resolution,[],[f2303,f500])).
fof(f2874,plain,(
  ( ! [X0] : (sP5('T_0','const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f598,f608,f2310,f502])).
fof(f2875,plain,(
  ( ! [X0] : (sP5('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f599,f608,f2310,f502])).
fof(f2886,plain,(
  ( ! [X0] : (~sP5('F_0','const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f2311,f499])).
fof(f2892,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X0)))))) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X1),s('type.list.list'(bool),X0)))) | s(bool,'T_0') = s(bool,sK22(X0,X1,bool))) )),
  inference(superposition,[],[f429,f360])).
fof(f2906,plain,(
  ( ! [X0] : (~sP5('F_0','T_0','const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))))) )),
  inference(unit_resulting_resolution,[],[f2318,f499])).
fof(f2913,plain,(
  ( ! [X0] : (sP5('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'T_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f608,f592,f2325,f502])).
fof(f2930,plain,(
  ( ! [X14,X15,X16] : (~p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X15),s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'F_0'),s('type.list.list'(bool),X16)))))) | p(s(bool,happ(s(fun(bool,bool),X15),s(bool,X14)))) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f443,f360])).
fof(f2931,plain,(
  ( ! [X19,X17,X18] : (~p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X18),s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X17),s('type.list.list'(bool),X19)))))) | p(s(bool,happ(s(fun(bool,bool),X18),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X17)) )),
  inference(superposition,[],[f443,f360])).
fof(f2959,plain,(
  ( ! [X0] : (sP5('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f602,f594,f2326,f503])).
fof(f2985,plain,(
  ( ! [X0] : (sP5('F_0','const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f600,f596,f2327,f504])).
fof(f2986,plain,(
  ( ! [X0] : (sP5('F_0','const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)))) )),
  inference(unit_resulting_resolution,[],[f585,f601,f596,f2327,f504])).
fof(f2987,plain,(
  ( ! [X0,X1] : (sP5('F_0','const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f585,f630,f596,f2327,f504])).
fof(f3008,plain,(
  ( ! [X14,X15,X16] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X15)))) | p(s(bool,'const.bool.IN_2'(s(bool,X14),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X16),s(fun(bool,bool),X15)))))) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f448,f360])).
fof(f3009,plain,(
  ( ! [X19,X17,X18] : (~p(s(bool,'const.bool.IN_2'(s(bool,X17),s(fun(bool,bool),X18)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X19),s(fun(bool,bool),X18)))))) | s(bool,'T_0') = s(bool,X17)) )),
  inference(superposition,[],[f448,f360])).
fof(f3035,plain,(
  ( ! [X0] : (sP5('F_0','const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)),'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))))) )),
  inference(unit_resulting_resolution,[],[f585,f603,f596,f2328,f504])).
fof(f3036,plain,(
  ( ! [X0,X1] : (sP5('F_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f585,f631,f596,f2328,f504])).
fof(f3037,plain,(
  ( ! [X0] : (sP5('F_0','F_0','const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))))) )),
  inference(unit_resulting_resolution,[],[f585,f602,f596,f2328,f503])).
fof(f3038,plain,(
  ( ! [X0] : (sP5('const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)),'F_0','const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))))) )),
  inference(unit_resulting_resolution,[],[f585,f602,f597,f2328,f503])).
fof(f3039,plain,(
  ( ! [X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'F_0','const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f585,f602,f629,f2328,f503])).
fof(f3068,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK30(X0,X1,X2,bool))) )),
  inference(superposition,[],[f497,f360])).
fof(f3069,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK30(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK30(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f497,f360])).
fof(f3152,plain,(
  ( ! [X4,X3] : (~p(s(bool,happ(s(fun(bool,bool),X3),s(bool,'F_0')))) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X3),s('type.list.list'(bool),X4)))) | s(bool,'T_0') = s(bool,'const.list.EL_2'(s('type.num.num',sK21(X3,X4,bool)),s('type.list.list'(bool),X4)))) )),
  inference(superposition,[],[f427,f360])).
fof(f3185,plain,(
  ( ! [X14,X15,X16] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X15),s(fun(bool,bool),X16)))))) | s(bool,X14) = s(bool,X15) | p(s(bool,'const.bool.IN_2'(s(bool,X14),s(fun(bool,bool),X16)))) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f446,f360])).
fof(f3186,plain,(
  ( ! [X19,X17,X18] : (~p(s(bool,'const.bool.IN_2'(s(bool,X17),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X18),s(fun(bool,bool),X19)))))) | s(bool,'F_0') = s(bool,X18) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X19)))) | s(bool,'T_0') = s(bool,X17)) )),
  inference(superposition,[],[f446,f360])).
fof(f3197,plain,(
  ( ! [X14,X15,X16] : (~p(s(bool,'const.bool.IN_2'(s(bool,X15),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X16)))))) | s(bool,X14) = s(bool,X15) | p(s(bool,'const.bool.IN_2'(s(bool,X15),s(fun(bool,bool),X16)))) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f446,f360])).
fof(f3198,plain,(
  ( ! [X19,X17,X18] : (~p(s(bool,'const.bool.IN_2'(s(bool,X18),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X17),s(fun(bool,bool),X19)))))) | s(bool,'F_0') = s(bool,X18) | p(s(bool,'const.bool.IN_2'(s(bool,X18),s(fun(bool,bool),X19)))) | s(bool,'T_0') = s(bool,X17)) )),
  inference(superposition,[],[f446,f360])).
fof(f3631,plain,(
  ( ! [X2,X0,X1] : (sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0','const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f585,f602,f2328,f724,f503])).
fof(f3675,plain,(
  ( ! [X2,X0,X1] : (sP5('F_0','const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f585,f2327,f596,f725,f504])).
fof(f3714,plain,(
  ( ! [X2,X0,X1] : (sP5('F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f585,f2328,f596,f726,f504])).
fof(f4185,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,happ(s(fun(X9,bool),X10),s(X9,X11)))) | p(s(bool,happ(s(fun(X9,bool),X10),s(X9,sK25(X10,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X8)))),X9)))))) )),
  inference(subsumption_resolution,[],[f4116,f350])).
fof(f4116,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X9,bool),X10),s(X9,X11)))) | p(s(bool,happ(s(fun(X9,bool),X10),s(X9,sK25(X10,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X8)))),X9)))))) )),
  inference(superposition,[],[f436,f2274])).
fof(f4186,plain,(
  ( ! [X14,X12,X15,X13] : (~p(s(bool,happ(s(fun(X13,bool),X14),s(X13,X15)))) | p(s(bool,happ(s(fun(X13,bool),X14),s(X13,sK26(X14,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X12)))),X13)))))) )),
  inference(subsumption_resolution,[],[f4117,f350])).
fof(f4117,plain,(
  ( ! [X14,X12,X15,X13] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X13,bool),X14),s(X13,X15)))) | p(s(bool,happ(s(fun(X13,bool),X14),s(X13,sK26(X14,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X12)))),X13)))))) )),
  inference(superposition,[],[f438,f2274])).
fof(f4189,plain,(
  ( ! [X26,X27] : (sP1('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X26)))),X27) | p(s(bool,X27))) )),
  inference(subsumption_resolution,[],[f4123,f350])).
fof(f4123,plain,(
  ( ! [X26,X27] : (~p(s(bool,'T_0')) | p(s(bool,X27)) | sP1('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X26)))),X27)) )),
  inference(superposition,[],[f480,f2274])).
fof(f4190,plain,(
  ( ! [X28,X29] : (sP1(X29,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X28))))) | p(s(bool,X29))) )),
  inference(subsumption_resolution,[],[f4124,f350])).
fof(f4124,plain,(
  ( ! [X28,X29] : (~p(s(bool,'T_0')) | p(s(bool,X29)) | sP1(X29,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X28)))))) )),
  inference(superposition,[],[f483,f2274])).
fof(f4191,plain,(
  ( ! [X30,X31,X32] : (~sP5(X31,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X30)))),X32) | p(s(bool,X32)) | ~p(s(bool,X31))) )),
  inference(subsumption_resolution,[],[f4125,f350])).
fof(f4125,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X31)) | p(s(bool,X32)) | ~sP5(X31,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X30)))),X32)) )),
  inference(superposition,[],[f501,f2274])).
fof(f4192,plain,(
  ( ! [X35,X33,X34] : (~sP4(X34,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X33)))),X35) | p(s(bool,X35)) | ~p(s(bool,X34))) )),
  inference(subsumption_resolution,[],[f4126,f350])).
fof(f4126,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,'T_0')) | ~p(s(bool,X34)) | p(s(bool,X35)) | ~sP4(X34,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X33)))),X35)) )),
  inference(superposition,[],[f505,f2274])).
fof(f4193,plain,(
  ( ! [X37,X38,X36] : (~sP3(X37,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X36)))),X38) | p(s(bool,X38)) | ~p(s(bool,X37))) )),
  inference(subsumption_resolution,[],[f4127,f350])).
fof(f4127,plain,(
  ( ! [X37,X38,X36] : (~p(s(bool,'T_0')) | ~p(s(bool,X37)) | p(s(bool,X38)) | ~sP3(X37,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X36)))),X38)) )),
  inference(superposition,[],[f509,f2274])).
fof(f4194,plain,(
  ( ! [X39,X41,X40] : (~sP6(X40,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X39)))),X41) | p(s(bool,X41)) | ~p(s(bool,X40))) )),
  inference(subsumption_resolution,[],[f4128,f350])).
fof(f4128,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X40)) | p(s(bool,X41)) | ~sP6(X40,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X39)))),X41)) )),
  inference(superposition,[],[f528,f2274])).
fof(f4195,plain,(
  ( ! [X43,X44,X42] : (~sP7(X44,X43,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X42))))) | p(s(bool,X44)) | ~p(s(bool,X43))) )),
  inference(subsumption_resolution,[],[f4129,f350])).
fof(f4129,plain,(
  ( ! [X43,X44,X42] : (~p(s(bool,'T_0')) | ~p(s(bool,X43)) | p(s(bool,X44)) | ~sP7(X44,X43,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X42)))))) )),
  inference(superposition,[],[f537,f2274])).
fof(f4196,plain,(
  ( ! [X47,X45,X46] : (~sP8('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X45)))),X46,X47) | p(s(bool,X47)) | p(s(bool,X46))) )),
  inference(subsumption_resolution,[],[f4130,f350])).
fof(f4130,plain,(
  ( ! [X47,X45,X46] : (~p(s(bool,'T_0')) | p(s(bool,X46)) | p(s(bool,X47)) | ~sP8('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X45)))),X46,X47)) )),
  inference(superposition,[],[f545,f2274])).
fof(f4197,plain,(
  ( ! [X50,X48,X49] : (sP9('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X48)))),X50,X49) | ~p(s(bool,X49))) )),
  inference(subsumption_resolution,[],[f4131,f350])).
fof(f4131,plain,(
  ( ! [X50,X48,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,X49)) | sP9('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X48)))),X50,X49)) )),
  inference(superposition,[],[f550,f2274])).
fof(f4198,plain,(
  ( ! [X52,X53,X51] : (sP9('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X51)))),X52,X53) | ~p(s(bool,X52))) )),
  inference(subsumption_resolution,[],[f4132,f350])).
fof(f4132,plain,(
  ( ! [X52,X53,X51] : (~p(s(bool,'T_0')) | ~p(s(bool,X52)) | sP9('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X51)))),X52,X53)) )),
  inference(superposition,[],[f551,f2274])).
fof(f4199,plain,(
  ( ! [X54,X56,X55] : (~sP9('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X54)))),X56,X55) | p(s(bool,X56)) | p(s(bool,X55))) )),
  inference(subsumption_resolution,[],[f4133,f350])).
fof(f4133,plain,(
  ( ! [X54,X56,X55] : (~p(s(bool,'T_0')) | p(s(bool,X55)) | p(s(bool,X56)) | ~sP9('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X54)))),X56,X55)) )),
  inference(superposition,[],[f552,f2274])).
fof(f4200,plain,(
  ( ! [X59,X57,X58] : (~sP10(X59,X58,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X57))))) | ~p(s(bool,X59)) | p(s(bool,X58))) )),
  inference(subsumption_resolution,[],[f4134,f350])).
fof(f4134,plain,(
  ( ! [X59,X57,X58] : (~p(s(bool,'T_0')) | p(s(bool,X58)) | ~p(s(bool,X59)) | ~sP10(X59,X58,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X57)))))) )),
  inference(superposition,[],[f562,f2274])).
fof(f4201,plain,(
  ( ! [X61,X62,X60] : (sP11('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X60)))),X62,X61) | p(s(bool,X61))) )),
  inference(subsumption_resolution,[],[f4135,f350])).
fof(f4135,plain,(
  ( ! [X61,X62,X60] : (~p(s(bool,'T_0')) | p(s(bool,X61)) | sP11('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X60)))),X62,X61)) )),
  inference(superposition,[],[f567,f2274])).
fof(f4202,plain,(
  ( ! [X64,X65,X63] : (sP11('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X63)))),X64,X65) | ~p(s(bool,X64))) )),
  inference(subsumption_resolution,[],[f4136,f350])).
fof(f4136,plain,(
  ( ! [X64,X65,X63] : (~p(s(bool,'T_0')) | ~p(s(bool,X64)) | sP11('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X63)))),X64,X65)) )),
  inference(superposition,[],[f568,f2274])).
fof(f4203,plain,(
  ( ! [X68,X66,X67] : (~sP11('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X66)))),X68,X67) | p(s(bool,X68)) | ~p(s(bool,X67))) )),
  inference(subsumption_resolution,[],[f4137,f350])).
fof(f4137,plain,(
  ( ! [X68,X66,X67] : (~p(s(bool,'T_0')) | ~p(s(bool,X67)) | p(s(bool,X68)) | ~sP11('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X66)))),X68,X67)) )),
  inference(superposition,[],[f569,f2274])).
fof(f4204,plain,(
  ( ! [X70,X71,X69] : (~sP10('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X69)))),X70,X71) | sP11('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X69)))),X70,X71)) )),
  inference(subsumption_resolution,[],[f4138,f350])).
fof(f4138,plain,(
  ( ! [X70,X71,X69] : (~p(s(bool,'T_0')) | ~sP10('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X69)))),X70,X71) | sP11('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X69)))),X70,X71)) )),
  inference(superposition,[],[f572,f2274])).
fof(f4205,plain,(
  ( ! [X74,X72,X73] : (~sP8('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X72)))),X73,X74) | sP9('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X72)))),X73,X74)) )),
  inference(subsumption_resolution,[],[f4139,f350])).
fof(f4139,plain,(
  ( ! [X74,X72,X73] : (~p(s(bool,'T_0')) | ~sP8('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X72)))),X73,X74) | sP9('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X72)))),X73,X74)) )),
  inference(superposition,[],[f573,f2274])).
fof(f4503,plain,(
  ( ! [X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f608,f2310,f918,f502])).
fof(f5427,plain,(
  ( ! [X0] : (~sP2('F_0','const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'F_0'),s('type.list.list'(bool),X0)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f600,f594,f1662,f504])).
fof(f5434,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'F_0'),s('type.list.list'(bool),X1)))),X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1662,f360])).
fof(f5437,plain,(
  ( ! [X5] : (~sP5('F_0','const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'F_0'),s('type.list.list'(bool),X5)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)))) )),
  inference(superposition,[],[f1662,f667])).
fof(f5440,plain,(
  ( ! [X12,X13] : (~sP5('F_0','const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'F_0'),s('type.list.list'(bool),X13)))),'const.prim_rec.<_2'(s('type.num.num',X12),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1662,f668])).
fof(f5441,plain,(
  ( ! [X14,X15] : (~sP5('F_0','const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'T_0'),s('type.list.list'(bool),X15)))),'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X14)))))) )),
  inference(superposition,[],[f1662,f2274])).
fof(f5442,plain,(
  ( ! [X17,X18,X16] : (~sP5('F_0','const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'F_0'),s('type.list.list'(bool),X18)))),'const.bool.IN_2'(s(X16,X17),s(fun(X16,bool),'const.pred_set.EMPTY_0')))) )),
  inference(superposition,[],[f1662,f716])).
fof(f5450,plain,(
  ( ! [X2,X0,X1] : (~sP2('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'F_0'),s('type.list.list'(bool),X2)))),'F_0')) )),
  inference(forward_demodulation,[],[f5430,f716])).
fof(f5430,plain,(
  ( ! [X2,X0,X1] : (~sP2('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))),s('type.list.list'(bool),X2)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f725,f723,f1662,f504])).
fof(f5451,plain,(
  ( ! [X0,X1] : (~sP2('const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'F_0'),s('type.list.list'(bool),X1)))),'F_0')) )),
  inference(forward_demodulation,[],[f5429,f668])).
fof(f5429,plain,(
  ( ! [X0,X1] : (~sP2('const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))),s('type.list.list'(bool),X1)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f630,f628,f1662,f504])).
fof(f5452,plain,(
  ( ! [X0] : (~sP2('const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)),'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'F_0'),s('type.list.list'(bool),X0)))),'F_0')) )),
  inference(forward_demodulation,[],[f5428,f667])).
fof(f5428,plain,(
  ( ! [X0] : (~sP2('const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)),'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13))),s('type.list.list'(bool),X0)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f601,f595,f1662,f504])).
fof(f5456,plain,(
  ( ! [X0,X1] : (~sP3('const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'T_0'),s('type.list.list'(bool),X1)))),'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'F_0')) )),
  inference(forward_demodulation,[],[f5422,f2274])).
fof(f5422,plain,(
  ( ! [X0,X1] : (~sP3('const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0))))),s('type.list.list'(bool),X1)))),'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f2328,f596,f1662,f504])).
fof(f5480,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'F_0'),s('type.list.list'(bool),X1)))),X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1663,f360])).
fof(f5483,plain,(
  ( ! [X5] : (sP5('T_0','const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'F_0'),s('type.list.list'(bool),X5)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)))) )),
  inference(superposition,[],[f1663,f667])).
fof(f5486,plain,(
  ( ! [X12,X13] : (sP5('T_0','const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'F_0'),s('type.list.list'(bool),X13)))),'const.prim_rec.<_2'(s('type.num.num',X12),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1663,f668])).
fof(f5487,plain,(
  ( ! [X14,X15] : (sP5('T_0','const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'T_0'),s('type.list.list'(bool),X15)))),'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X14)))))) )),
  inference(superposition,[],[f1663,f2274])).
fof(f5488,plain,(
  ( ! [X17,X18,X16] : (sP5('T_0','const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'F_0'),s('type.list.list'(bool),X18)))),'const.bool.IN_2'(s(X16,X17),s(fun(X16,bool),'const.pred_set.EMPTY_0')))) )),
  inference(superposition,[],[f1663,f716])).
fof(f5526,plain,(
  ( ! [X0] : (~sP2('const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'F_0'),s('type.list.list'(bool),X0)))),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f602,f596,f1664,f503])).
fof(f5529,plain,(
  ( ! [X0,X1] : (~sP5('F_0',X0,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'F_0'),s('type.list.list'(bool),X1))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1664,f360])).
fof(f5532,plain,(
  ( ! [X5] : (~sP5('F_0','const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)),'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'F_0'),s('type.list.list'(bool),X5)))))) )),
  inference(superposition,[],[f1664,f667])).
fof(f5535,plain,(
  ( ! [X12,X13] : (~sP5('F_0','const.prim_rec.<_2'(s('type.num.num',X12),s('type.num.num','const.num.|48|_0')),'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'F_0'),s('type.list.list'(bool),X13)))))) )),
  inference(superposition,[],[f1664,f668])).
fof(f5536,plain,(
  ( ! [X14,X15] : (~sP5('F_0','const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X14)))),'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'T_0'),s('type.list.list'(bool),X15)))))) )),
  inference(superposition,[],[f1664,f2274])).
fof(f5537,plain,(
  ( ! [X17,X18,X16] : (~sP5('F_0','const.bool.IN_2'(s(X16,X17),s(fun(X16,bool),'const.pred_set.EMPTY_0')),'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'F_0'),s('type.list.list'(bool),X18)))))) )),
  inference(superposition,[],[f1664,f716])).
fof(f5545,plain,(
  ( ! [X0,X1] : (~sP3('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'T_0'),s('type.list.list'(bool),X1)))),'F_0')) )),
  inference(forward_demodulation,[],[f5525,f2274])).
fof(f5525,plain,(
  ( ! [X0,X1] : (~sP3('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0))))),s('type.list.list'(bool),X1)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f2327,f2310,f1664,f504])).
fof(f5548,plain,(
  ( ! [X2,X0,X1] : (~sP2('const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'F_0'),s('type.list.list'(bool),X2)))),'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0')) )),
  inference(forward_demodulation,[],[f5521,f716])).
fof(f5521,plain,(
  ( ! [X2,X0,X1] : (~sP2('const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))),s('type.list.list'(bool),X2)))),'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f726,f596,f1664,f504])).
fof(f5549,plain,(
  ( ! [X0,X1] : (~sP2('const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'F_0'),s('type.list.list'(bool),X1)))),'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'F_0')) )),
  inference(forward_demodulation,[],[f5520,f668])).
fof(f5520,plain,(
  ( ! [X0,X1] : (~sP2('const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))),s('type.list.list'(bool),X1)))),'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f631,f596,f1664,f504])).
fof(f5550,plain,(
  ( ! [X0] : (~sP2('const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'F_0'),s('type.list.list'(bool),X0)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)),'F_0')) )),
  inference(forward_demodulation,[],[f5519,f667])).
fof(f5519,plain,(
  ( ! [X0] : (~sP2('const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13))),s('type.list.list'(bool),X0)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f603,f596,f1664,f504])).
fof(f5583,plain,(
  ( ! [X0,X1] : (sP5('T_0',X0,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'F_0'),s('type.list.list'(bool),X1))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1665,f360])).
fof(f5586,plain,(
  ( ! [X5] : (sP5('T_0','const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)),'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'F_0'),s('type.list.list'(bool),X5)))))) )),
  inference(superposition,[],[f1665,f667])).
fof(f5589,plain,(
  ( ! [X12,X13] : (sP5('T_0','const.prim_rec.<_2'(s('type.num.num',X12),s('type.num.num','const.num.|48|_0')),'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'F_0'),s('type.list.list'(bool),X13)))))) )),
  inference(superposition,[],[f1665,f668])).
fof(f5590,plain,(
  ( ! [X14,X15] : (sP5('T_0','const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X14)))),'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'T_0'),s('type.list.list'(bool),X15)))))) )),
  inference(superposition,[],[f1665,f2274])).
fof(f5591,plain,(
  ( ! [X17,X18,X16] : (sP5('T_0','const.bool.IN_2'(s(X16,X17),s(fun(X16,bool),'const.pred_set.EMPTY_0')),'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'F_0'),s('type.list.list'(bool),X18)))))) )),
  inference(superposition,[],[f1665,f716])).
fof(f5823,plain,(
  ( ! [X0] : (sP4('T_0','const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'F_0'),s('type.list.list'(bool),X0)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)))) )),
  inference(unit_resulting_resolution,[],[f5586,f500])).
fof(f7662,plain,(
  ( ! [X0] : (sP5('const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)),'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)),'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))))) )),
  inference(unit_resulting_resolution,[],[f603,f597,f2328,f2232,f502])).
fof(f7663,plain,(
  ( ! [X0,X1] : (sP5('const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)),'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f631,f597,f2328,f2232,f502])).
fof(f7664,plain,(
  ( ! [X2,X0,X1] : (sP5('const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)),'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f726,f597,f2328,f2232,f502])).
fof(f7666,plain,(
  ( ! [X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)),'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f603,f629,f2328,f2232,f502])).
fof(f7667,plain,(
  ( ! [X2,X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0')),'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f631,f629,f2328,f2232,f502])).
fof(f7668,plain,(
  ( ! [X2,X0,X3,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')),'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X3)))))) )),
  inference(unit_resulting_resolution,[],[f726,f629,f2328,f2232,f502])).
fof(f7670,plain,(
  ( ! [X2,X0,X1] : (sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)),'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f603,f724,f2328,f2232,f502])).
fof(f7671,plain,(
  ( ! [X2,X0,X3,X1] : (sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num','const.num.|48|_0')),'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X3)))))) )),
  inference(unit_resulting_resolution,[],[f631,f724,f2328,f2232,f502])).
fof(f7672,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.pred_set.EMPTY_0')),'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X4)))))) )),
  inference(unit_resulting_resolution,[],[f726,f724,f2328,f2232,f502])).
fof(f7695,plain,(
  ( ! [X2,X0,X3,X1] : (sP5('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'T_0','const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.INSERT_2'(s(X1,X2),s(fun(X1,bool),X3)))))) )),
  inference(unit_resulting_resolution,[],[f2325,f592,f606,f2232,f502])).
fof(f7700,plain,(
  ( ! [X2,X0,X3,X1] : (sP5('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.INSERT_2'(s(X1,X2),s(fun(X1,bool),X3)))))) )),
  inference(unit_resulting_resolution,[],[f2325,f593,f607,f2232,f502])).
fof(f7704,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP5('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))),'const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.pred_set.INSERT_2'(s(X2,X3),s(fun(X2,bool),X4)))))) )),
  inference(unit_resulting_resolution,[],[f2325,f917,f920,f2232,f502])).
fof(f7705,plain,(
  ( ! [X2,X0,X3,X1] : (sP5('T_0','const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.INSERT_2'(s(X1,X2),s(fun(X1,bool),X3)))))) )),
  inference(unit_resulting_resolution,[],[f598,f2310,f2327,f2232,f502])).
fof(f7706,plain,(
  ( ! [X2,X0,X3,X1] : (sP5('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.INSERT_2'(s(X1,X2),s(fun(X1,bool),X3)))))) )),
  inference(unit_resulting_resolution,[],[f599,f2310,f2327,f2232,f502])).
fof(f7707,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))),'const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.pred_set.INSERT_2'(s(X2,X3),s(fun(X2,bool),X4)))))) )),
  inference(unit_resulting_resolution,[],[f918,f2310,f2327,f2232,f502])).
fof(f7708,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP5('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))),'const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.pred_set.INSERT_2'(s(X2,X3),s(fun(X2,bool),X4)))))) )),
  inference(unit_resulting_resolution,[],[f2325,f2310,f2327,f2232,f502])).
fof(f7768,plain,(
  ( ! [X2,X0,X3,X1] : (sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X3)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f2274,f2232,f518])).
fof(f7779,plain,(
  ( ! [X2,X0,X3,X1] : (sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'T_0','const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X3)))))) )),
  inference(unit_resulting_resolution,[],[f2274,f2232,f518])).
fof(f9195,plain,(
  ( ! [X656,X654,X653,X655] : (s(bool,'T_0') != s(bool,f(s(X653,sK30(X654,X655,bool,X653)))) | s(fun(X653,bool),X654) = s(fun(X653,bool),X655) | s(bool,X656) = s(bool,g(s(X653,sK30(X654,X655,bool,X653)))) | s(bool,'T_0') = s(bool,X656)) )),
  inference(superposition,[],[f497,f983])).
fof(f8630,plain,(
  ( ! [X662,X664,X663,X665] : (s(bool,X665) != s(bool,f(s(X662,sK30(X663,X664,bool,X662)))) | s(fun(X662,bool),X663) = s(fun(X662,bool),X664) | s(bool,'T_0') = s(bool,X665) | s(bool,'T_0') = s(bool,g(s(X662,sK30(X663,X664,bool,X662))))) )),
  inference(superposition,[],[f497,f983])).
fof(f414,plain,(
  ( ! [X0,X5,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK19(X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))) )),
  inference(cnf_transformation,[],[f269])).
fof(f269,plain,(
  ! [X0,X1] : ((! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK19(X1,X0)))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK20(X1,X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK19,sK20])],[f268])).
fof(f268,plain,(
  ! [X0,X1] : ((! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(rectify,[],[f267])).
fof(f267,plain,(
  ! [X0,X1] : ((! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(nnf_transformation,[],[f198])).
fof(f198,plain,(
  ! [X0,X1] : (? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(ennf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0,X1] : (~! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(flattening,[],[f118])).
fof(f118,plain,(
  ! [X0] : ! [X1] : (~! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(rectify,[],[f62])).
fof(f62,axiom,(
  ! [X6] : ! [X9] : (~! [X10] : p(s(bool,happ(s(fun(X6,bool),X9),s(X6,X10)))) <=> ? [X10] : ~p(s(bool,happ(s(fun(X6,bool),X9),s(X6,X10)))))),
  file('thm.list.MONO_EVERY','thm.bool.NOT_FORALL_THM')).
fof(f415,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK20(X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f269])).
fof(f428,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) )),
  inference(cnf_transformation,[],[f278])).
fof(f278,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))))) & ((p(s(bool,'const.bool.IN_2'(s(X0,sK22(X2,X1,X0)),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK22(X2,X1,X0)))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK22])],[f277])).
fof(f277,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))))) & (? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))))))),
  inference(rectify,[],[f276])).
fof(f276,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) | ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))) & (? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))))))),
  inference(nnf_transformation,[],[f200])).
fof(f200,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) <=> ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))))),
  inference(ennf_transformation,[],[f136])).
fof(f136,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))))),
  inference(flattening,[],[f135])).
fof(f135,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : ! [X9,X7] : (p(s(bool,'const.list.EVERY_2'(s(fun(X6,bool),X9),s('type.list.list'(X6),X7)))) <=> ! [X12] : (p(s(bool,'const.bool.IN_2'(s(X6,X12),s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),X7)))))) => p(s(bool,happ(s(fun(X6,bool),X9),s(X6,X12))))))),
  file('thm.list.MONO_EVERY','thm.list.EVERY_MEM')).
fof(f432,plain,(
  ( ! [X6,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK23(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f282])).
fof(f282,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK23(X2,X1,X0))))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK24(X2,X1,X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK23,sK24])],[f281])).
fof(f281,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | ? [X5] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(rectify,[],[f280])).
fof(f280,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(flattening,[],[f279])).
fof(f279,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f138])).
fof(f138,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f137])).
fof(f137,plain,(
  ! [X0] : ! [X1,X2] : ((p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f60])).
fof(f60,axiom,(
  ! [X6] : ! [X9,X8] : ((p(s(bool,X9)) | ? [X10] : p(s(bool,happ(s(fun(X6,bool),X8),s(X6,X10))))) <=> ? [X10] : (p(s(bool,X9)) | p(s(bool,happ(s(fun(X6,bool),X8),s(X6,X10))))))),
  file('thm.list.MONO_EVERY','thm.bool.RIGHT_OR_EXISTS_THM')).
fof(f434,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK24(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f282])).
fof(f436,plain,(
  ( ! [X6,X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK25(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f286])).
fof(f286,plain,(
  ! [X0,X1,X2] : ((! [X6] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK25(X2,X1,X0))))))) & (~p(s(bool,X1)) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK26(X2,X1,X0))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK25,sK26])],[f285])).
fof(f285,plain,(
  ! [X0,X1,X2] : ((! [X6] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | (p(s(bool,X1)) & ? [X5] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))))) & (~p(s(bool,X1)) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X3] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(rectify,[],[f284])).
fof(f284,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (p(s(bool,X1)) & ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (~p(s(bool,X1)) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X3] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(flattening,[],[f283])).
fof(f283,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (p(s(bool,X1)) & ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & ((~p(s(bool,X1)) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ? [X3] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f140])).
fof(f140,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (p(s(bool,X1)) & ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f139])).
fof(f139,plain,(
  ! [X0] : ! [X1,X2] : (? [X3] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (p(s(bool,X1)) & ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f64])).
fof(f64,axiom,(
  ! [X6] : ! [X9,X8] : (? [X10] : (p(s(bool,X9)) & p(s(bool,happ(s(fun(X6,bool),X8),s(X6,X10))))) <=> (p(s(bool,X9)) & ? [X10] : p(s(bool,happ(s(fun(X6,bool),X8),s(X6,X10))))))),
  file('thm.list.MONO_EVERY','thm.bool.RIGHT_EXISTS_AND_THM')).
fof(f438,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK26(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f286])).
fof(f439,plain,(
  ( ! [X2,X0,X8,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X8)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK27(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK28(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f290])).
fof(f290,plain,(
  ! [X0,X1,X2] : ((! [X8] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X8)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X8))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK27(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK28(X2,X1,X0)))))) & ((! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK29(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK29(X2,X1,X0))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK27,sK28,sK29])],[f289])).
fof(f289,plain,(
  ! [X0,X1,X2] : ((! [X8] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X8)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X8))))) | ? [X7] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X7)))) | ? [X6] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) & ((! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(rectify,[],[f288])).
fof(f288,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) & ((! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(flattening,[],[f287])).
fof(f287,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & ((! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f142])).
fof(f142,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f141])).
fof(f141,plain,(
  ! [X0] : ! [X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f61])).
fof(f61,axiom,(
  ! [X6] : ! [X9,X8] : (? [X10] : (p(s(bool,happ(s(fun(X6,bool),X9),s(X6,X10)))) | p(s(bool,happ(s(fun(X6,bool),X8),s(X6,X10))))) <=> (? [X10] : p(s(bool,happ(s(fun(X6,bool),X9),s(X6,X10)))) | ? [X10] : p(s(bool,happ(s(fun(X6,bool),X8),s(X6,X10))))))),
  file('thm.list.MONO_EVERY','thm.bool.EXISTS_OR_THM')).
fof(f440,plain,(
  ( ! [X2,X0,X8,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X8)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK27(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK28(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f290])).
fof(f441,plain,(
  ( ! [X2,X0,X5,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK29(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK29(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f290])).
fof(f442,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK29(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK29(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f290])).
fof(f443,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3)))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) )),
  inference(cnf_transformation,[],[f292])).
fof(f292,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3)))))) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) & p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X3)))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X3)))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3))))))))),
  inference(flattening,[],[f291])).
fof(f291,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3)))))) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) & p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X3)))))) & ((~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X3))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3))))))))),
  inference(nnf_transformation,[],[f144])).
fof(f144,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3)))))) <=> (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) & p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X3))))))),
  inference(flattening,[],[f143])).
fof(f143,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3)))))) <=> (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) & p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X3))))))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X6] : ! [X9,X17,X15] : (p(s(bool,'const.list.EVERY_2'(s(fun(X6,bool),X9),s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X17),s('type.list.list'(X6),X15)))))) <=> (p(s(bool,happ(s(fun(X6,bool),X9),s(X6,X17)))) & p(s(bool,'const.list.EVERY_2'(s(fun(X6,bool),X9),s('type.list.list'(X6),X15))))))),
  file('thm.list.MONO_EVERY','thm.list.EVERY_DEF|split|1')).
fof(f445,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X3)))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3)))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) )),
  inference(cnf_transformation,[],[f292])).
fof(f2145,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,happ(s(fun(X3,bool),X4),s(X3,sK19('const.list.HD_1'(s('type.list.list'(fun(X3,bool)),'const.list.CONS_2'(s(fun(X3,bool),X4),s('type.list.list'(fun(X3,bool)),X5)))),X3))))) | p(s(bool,happ(s(fun(X3,bool),X4),s(X3,X6))))) )),
  inference(forward_demodulation,[],[f2141,f419])).
fof(f2141,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,happ(s(fun(X3,bool),X4),s(X3,sK19('const.list.HD_1'(s('type.list.list'(fun(X3,bool)),'const.list.CONS_2'(s(fun(X3,bool),X4),s('type.list.list'(fun(X3,bool)),X5)))),X3))))) | p(s(bool,happ(s(fun(X3,bool),'const.list.HD_1'(s('type.list.list'(fun(X3,bool)),'const.list.CONS_2'(s(fun(X3,bool),X4),s('type.list.list'(fun(X3,bool)),X5))))),s(X3,X6))))) )),
  inference(superposition,[],[f414,f419])).
fof(f2167,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,happ(s(fun(X3,bool),X4),s(X3,sK20('const.list.HD_1'(s('type.list.list'(fun(X3,bool)),'const.list.CONS_2'(s(fun(X3,bool),X4),s('type.list.list'(fun(X3,bool)),X5)))),X3))))) | p(s(bool,happ(s(fun(X3,bool),X4),s(X3,X6))))) )),
  inference(forward_demodulation,[],[f2163,f419])).
fof(f2163,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,happ(s(fun(X3,bool),X4),s(X3,sK20('const.list.HD_1'(s('type.list.list'(fun(X3,bool)),'const.list.CONS_2'(s(fun(X3,bool),X4),s('type.list.list'(fun(X3,bool)),X5)))),X3))))) | p(s(bool,happ(s(fun(X3,bool),'const.list.HD_1'(s('type.list.list'(fun(X3,bool)),'const.list.CONS_2'(s(fun(X3,bool),X4),s('type.list.list'(fun(X3,bool)),X5))))),s(X3,X6))))) )),
  inference(superposition,[],[f415,f419])).
fof(f2765,plain,(
  ( ! [X12,X10,X13,X11,X9] : (~p(s(bool,happ(s(fun(X9,bool),X10),s(X9,X12)))) | p(s(bool,happ(s(fun(X9,bool),X10),s(X9,sK23('const.list.HD_1'(s('type.list.list'(fun(X9,bool)),'const.list.CONS_2'(s(fun(X9,bool),X10),s('type.list.list'(fun(X9,bool)),X11)))),X13,X9))))) | p(s(bool,X13))) )),
  inference(forward_demodulation,[],[f2749,f419])).
fof(f2749,plain,(
  ( ! [X12,X10,X13,X11,X9] : (~p(s(bool,happ(s(fun(X9,bool),X10),s(X9,X12)))) | p(s(bool,X13)) | p(s(bool,happ(s(fun(X9,bool),'const.list.HD_1'(s('type.list.list'(fun(X9,bool)),'const.list.CONS_2'(s(fun(X9,bool),X10),s('type.list.list'(fun(X9,bool)),X11))))),s(X9,sK23('const.list.HD_1'(s('type.list.list'(fun(X9,bool)),'const.list.CONS_2'(s(fun(X9,bool),X10),s('type.list.list'(fun(X9,bool)),X11)))),X13,X9)))))) )),
  inference(superposition,[],[f432,f419])).
fof(f2802,plain,(
  ( ! [X12,X10,X13,X11,X9] : (~p(s(bool,happ(s(fun(X9,bool),X10),s(X9,X12)))) | p(s(bool,happ(s(fun(X9,bool),X10),s(X9,sK24('const.list.HD_1'(s('type.list.list'(fun(X9,bool)),'const.list.CONS_2'(s(fun(X9,bool),X10),s('type.list.list'(fun(X9,bool)),X11)))),X13,X9))))) | p(s(bool,X13))) )),
  inference(forward_demodulation,[],[f2786,f419])).
fof(f2786,plain,(
  ( ! [X12,X10,X13,X11,X9] : (~p(s(bool,happ(s(fun(X9,bool),X10),s(X9,X12)))) | p(s(bool,X13)) | p(s(bool,happ(s(fun(X9,bool),'const.list.HD_1'(s('type.list.list'(fun(X9,bool)),'const.list.CONS_2'(s(fun(X9,bool),X10),s('type.list.list'(fun(X9,bool)),X11))))),s(X9,sK24('const.list.HD_1'(s('type.list.list'(fun(X9,bool)),'const.list.CONS_2'(s(fun(X9,bool),X10),s('type.list.list'(fun(X9,bool)),X11)))),X13,X9)))))) )),
  inference(superposition,[],[f434,f419])).
fof(f2831,plain,(
  ( ! [X23,X21,X19,X22,X20] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X19),s('type.num.num',X20)))) | ~p(s(bool,happ(s(fun(X21,bool),X22),s(X21,X23)))) | p(s(bool,happ(s(fun(X21,bool),X22),s(X21,sK25(X22,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X19))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X20)))),X21)))))) )),
  inference(superposition,[],[f436,f449])).
fof(f2832,plain,(
  ( ! [X26,X24,X27,X25] : (~p(s(bool,'const.list.HD_1'(s('type.list.list'(bool),X24)))) | ~p(s(bool,happ(s(fun(X25,bool),X26),s(X25,X27)))) | p(s(bool,happ(s(fun(X25,bool),X26),s(X25,sK25(X26,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(bool),X24)),X25)))))) )),
  inference(superposition,[],[f436,f409])).
fof(f2833,plain,(
  ( ! [X30,X28,X31,X29,X32] : (~p(s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X28))),s('type.list.list'(bool),X29)))) | ~p(s(bool,happ(s(fun(X30,bool),X31),s(X30,X32)))) | p(s(bool,happ(s(fun(X30,bool),X31),s(X30,sK25(X31,'const.list.EL_2'(s('type.num.num',X28),s('type.list.list'(bool),'const.list.TL_1'(s('type.list.list'(bool),X29)))),X30)))))) )),
  inference(superposition,[],[f436,f422])).
fof(f2834,plain,(
  ( ! [X37,X35,X33,X36,X34] : (~p(s(bool,X33)) | ~p(s(bool,happ(s(fun(X35,bool),X36),s(X35,X37)))) | p(s(bool,happ(s(fun(X35,bool),X36),s(X35,sK25(X36,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X33),s('type.list.list'(bool),X34)))),X35)))))) )),
  inference(superposition,[],[f436,f419])).
fof(f2835,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | p(s(bool,happ(s(fun(X8,bool),X9),s(X8,sK25(X9,'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X8)))))) )),
  inference(subsumption_resolution,[],[f2828,f350])).
fof(f2828,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | p(s(bool,happ(s(fun(X8,bool),X9),s(X8,sK25(X9,'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X8)))))) )),
  inference(superposition,[],[f436,f647])).
fof(f2836,plain,(
  ( ! [X14,X17,X15,X18,X16] : (~p(s(bool,happ(s(fun(X16,bool),X17),s(X16,X18)))) | p(s(bool,happ(s(fun(X16,bool),X17),s(X16,sK25(X17,'const.list.EVERY_2'(s(fun(X14,bool),X15),s('type.list.list'(X14),'const.list.NIL_0')),X16)))))) )),
  inference(subsumption_resolution,[],[f2830,f350])).
fof(f2830,plain,(
  ( ! [X14,X17,X15,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X16,bool),X17),s(X16,X18)))) | p(s(bool,happ(s(fun(X16,bool),X17),s(X16,sK25(X17,'const.list.EVERY_2'(s(fun(X14,bool),X15),s('type.list.list'(X14),'const.list.NIL_0')),X16)))))) )),
  inference(superposition,[],[f436,f408])).
fof(f2860,plain,(
  ( ! [X23,X21,X19,X22,X20] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X19),s('type.num.num',X20)))) | ~p(s(bool,happ(s(fun(X21,bool),X22),s(X21,X23)))) | p(s(bool,happ(s(fun(X21,bool),X22),s(X21,sK26(X22,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X19))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X20)))),X21)))))) )),
  inference(superposition,[],[f438,f449])).
fof(f2861,plain,(
  ( ! [X26,X24,X27,X25] : (~p(s(bool,'const.list.HD_1'(s('type.list.list'(bool),X24)))) | ~p(s(bool,happ(s(fun(X25,bool),X26),s(X25,X27)))) | p(s(bool,happ(s(fun(X25,bool),X26),s(X25,sK26(X26,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(bool),X24)),X25)))))) )),
  inference(superposition,[],[f438,f409])).
fof(f2862,plain,(
  ( ! [X30,X28,X31,X29,X32] : (~p(s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X28))),s('type.list.list'(bool),X29)))) | ~p(s(bool,happ(s(fun(X30,bool),X31),s(X30,X32)))) | p(s(bool,happ(s(fun(X30,bool),X31),s(X30,sK26(X31,'const.list.EL_2'(s('type.num.num',X28),s('type.list.list'(bool),'const.list.TL_1'(s('type.list.list'(bool),X29)))),X30)))))) )),
  inference(superposition,[],[f438,f422])).
fof(f2863,plain,(
  ( ! [X37,X35,X33,X36,X34] : (~p(s(bool,X33)) | ~p(s(bool,happ(s(fun(X35,bool),X36),s(X35,X37)))) | p(s(bool,happ(s(fun(X35,bool),X36),s(X35,sK26(X36,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X33),s('type.list.list'(bool),X34)))),X35)))))) )),
  inference(superposition,[],[f438,f419])).
fof(f2864,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | p(s(bool,happ(s(fun(X8,bool),X9),s(X8,sK26(X9,'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X8)))))) )),
  inference(subsumption_resolution,[],[f2857,f350])).
fof(f2857,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | p(s(bool,happ(s(fun(X8,bool),X9),s(X8,sK26(X9,'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X8)))))) )),
  inference(superposition,[],[f438,f647])).
fof(f2865,plain,(
  ( ! [X14,X17,X15,X18,X16] : (~p(s(bool,happ(s(fun(X16,bool),X17),s(X16,X18)))) | p(s(bool,happ(s(fun(X16,bool),X17),s(X16,sK26(X17,'const.list.EVERY_2'(s(fun(X14,bool),X15),s('type.list.list'(X14),'const.list.NIL_0')),X16)))))) )),
  inference(subsumption_resolution,[],[f2859,f350])).
fof(f2859,plain,(
  ( ! [X14,X17,X15,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X16,bool),X17),s(X16,X18)))) | p(s(bool,happ(s(fun(X16,bool),X17),s(X16,sK26(X17,'const.list.EVERY_2'(s(fun(X14,bool),X15),s('type.list.list'(X14),'const.list.NIL_0')),X16)))))) )),
  inference(superposition,[],[f438,f408])).
fof(f3303,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'const.bool.IN_2'(s(X2,X4),s(fun(X2,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X2),'const.list.NIL_0')))))) | p(s(bool,happ(s(fun(X2,bool),X3),s(X2,X4))))) )),
  inference(subsumption_resolution,[],[f3293,f350])).
fof(f3293,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(X2,X4),s(fun(X2,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X2),'const.list.NIL_0')))))) | p(s(bool,happ(s(fun(X2,bool),X3),s(X2,X4))))) )),
  inference(superposition,[],[f428,f408])).
fof(f3340,plain,(
  ( ! [X12,X10,X13,X11,X9] : (~p(s(bool,happ(s(fun(X9,bool),X10),s(X9,X12)))) | p(s(bool,happ(s(fun(X9,bool),X10),s(X9,sK27(X13,'const.list.HD_1'(s('type.list.list'(fun(X9,bool)),'const.list.CONS_2'(s(fun(X9,bool),X10),s('type.list.list'(fun(X9,bool)),X11)))),X9))))) | p(s(bool,happ(s(fun(X9,bool),X13),s(X9,sK28(X13,'const.list.HD_1'(s('type.list.list'(fun(X9,bool)),'const.list.CONS_2'(s(fun(X9,bool),X10),s('type.list.list'(fun(X9,bool)),X11)))),X9)))))) )),
  inference(forward_demodulation,[],[f3324,f419])).
fof(f3324,plain,(
  ( ! [X12,X10,X13,X11,X9] : (~p(s(bool,happ(s(fun(X9,bool),X10),s(X9,X12)))) | p(s(bool,happ(s(fun(X9,bool),'const.list.HD_1'(s('type.list.list'(fun(X9,bool)),'const.list.CONS_2'(s(fun(X9,bool),X10),s('type.list.list'(fun(X9,bool)),X11))))),s(X9,sK27(X13,'const.list.HD_1'(s('type.list.list'(fun(X9,bool)),'const.list.CONS_2'(s(fun(X9,bool),X10),s('type.list.list'(fun(X9,bool)),X11)))),X9))))) | p(s(bool,happ(s(fun(X9,bool),X13),s(X9,sK28(X13,'const.list.HD_1'(s('type.list.list'(fun(X9,bool)),'const.list.CONS_2'(s(fun(X9,bool),X10),s('type.list.list'(fun(X9,bool)),X11)))),X9)))))) )),
  inference(superposition,[],[f439,f419])).
fof(f3378,plain,(
  ( ! [X12,X10,X13,X11,X9] : (~p(s(bool,happ(s(fun(X9,bool),X10),s(X9,X12)))) | p(s(bool,happ(s(fun(X9,bool),X10),s(X9,sK28('const.list.HD_1'(s('type.list.list'(fun(X9,bool)),'const.list.CONS_2'(s(fun(X9,bool),X10),s('type.list.list'(fun(X9,bool)),X11)))),X13,X9))))) | p(s(bool,happ(s(fun(X9,bool),X13),s(X9,sK27('const.list.HD_1'(s('type.list.list'(fun(X9,bool)),'const.list.CONS_2'(s(fun(X9,bool),X10),s('type.list.list'(fun(X9,bool)),X11)))),X13,X9)))))) )),
  inference(forward_demodulation,[],[f3362,f419])).
fof(f3362,plain,(
  ( ! [X12,X10,X13,X11,X9] : (~p(s(bool,happ(s(fun(X9,bool),X10),s(X9,X12)))) | p(s(bool,happ(s(fun(X9,bool),X13),s(X9,sK27('const.list.HD_1'(s('type.list.list'(fun(X9,bool)),'const.list.CONS_2'(s(fun(X9,bool),X10),s('type.list.list'(fun(X9,bool)),X11)))),X13,X9))))) | p(s(bool,happ(s(fun(X9,bool),'const.list.HD_1'(s('type.list.list'(fun(X9,bool)),'const.list.CONS_2'(s(fun(X9,bool),X10),s('type.list.list'(fun(X9,bool)),X11))))),s(X9,sK28('const.list.HD_1'(s('type.list.list'(fun(X9,bool)),'const.list.CONS_2'(s(fun(X9,bool),X10),s('type.list.list'(fun(X9,bool)),X11)))),X13,X9)))))) )),
  inference(superposition,[],[f440,f419])).
fof(f3425,plain,(
  ( ! [X12,X10,X13,X11,X9] : (~p(s(bool,happ(s(fun(X9,bool),X10),s(X9,X12)))) | p(s(bool,happ(s(fun(X9,bool),X10),s(X9,sK29(X13,'const.list.HD_1'(s('type.list.list'(fun(X9,bool)),'const.list.CONS_2'(s(fun(X9,bool),X10),s('type.list.list'(fun(X9,bool)),X11)))),X9))))) | p(s(bool,happ(s(fun(X9,bool),X13),s(X9,sK29(X13,'const.list.HD_1'(s('type.list.list'(fun(X9,bool)),'const.list.CONS_2'(s(fun(X9,bool),X10),s('type.list.list'(fun(X9,bool)),X11)))),X9)))))) )),
  inference(forward_demodulation,[],[f3409,f419])).
fof(f3409,plain,(
  ( ! [X12,X10,X13,X11,X9] : (~p(s(bool,happ(s(fun(X9,bool),X10),s(X9,X12)))) | p(s(bool,happ(s(fun(X9,bool),'const.list.HD_1'(s('type.list.list'(fun(X9,bool)),'const.list.CONS_2'(s(fun(X9,bool),X10),s('type.list.list'(fun(X9,bool)),X11))))),s(X9,sK29(X13,'const.list.HD_1'(s('type.list.list'(fun(X9,bool)),'const.list.CONS_2'(s(fun(X9,bool),X10),s('type.list.list'(fun(X9,bool)),X11)))),X9))))) | p(s(bool,happ(s(fun(X9,bool),X13),s(X9,sK29(X13,'const.list.HD_1'(s('type.list.list'(fun(X9,bool)),'const.list.CONS_2'(s(fun(X9,bool),X10),s('type.list.list'(fun(X9,bool)),X11)))),X9)))))) )),
  inference(superposition,[],[f441,f419])).
fof(f3460,plain,(
  ( ! [X12,X10,X13,X11,X9] : (~p(s(bool,happ(s(fun(X9,bool),X10),s(X9,X12)))) | p(s(bool,happ(s(fun(X9,bool),X10),s(X9,sK29('const.list.HD_1'(s('type.list.list'(fun(X9,bool)),'const.list.CONS_2'(s(fun(X9,bool),X10),s('type.list.list'(fun(X9,bool)),X11)))),X13,X9))))) | p(s(bool,happ(s(fun(X9,bool),X13),s(X9,sK29('const.list.HD_1'(s('type.list.list'(fun(X9,bool)),'const.list.CONS_2'(s(fun(X9,bool),X10),s('type.list.list'(fun(X9,bool)),X11)))),X13,X9)))))) )),
  inference(forward_demodulation,[],[f3444,f419])).
fof(f3444,plain,(
  ( ! [X12,X10,X13,X11,X9] : (~p(s(bool,happ(s(fun(X9,bool),X10),s(X9,X12)))) | p(s(bool,happ(s(fun(X9,bool),X13),s(X9,sK29('const.list.HD_1'(s('type.list.list'(fun(X9,bool)),'const.list.CONS_2'(s(fun(X9,bool),X10),s('type.list.list'(fun(X9,bool)),X11)))),X13,X9))))) | p(s(bool,happ(s(fun(X9,bool),'const.list.HD_1'(s('type.list.list'(fun(X9,bool)),'const.list.CONS_2'(s(fun(X9,bool),X10),s('type.list.list'(fun(X9,bool)),X11))))),s(X9,sK29('const.list.HD_1'(s('type.list.list'(fun(X9,bool)),'const.list.CONS_2'(s(fun(X9,bool),X10),s('type.list.list'(fun(X9,bool)),X11)))),X13,X9)))))) )),
  inference(superposition,[],[f442,f419])).
fof(f3500,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,happ(s(fun(X2,bool),X3),s(X2,X4)))) | p(s(bool,'const.list.EVERY_2'(s(fun(X2,bool),X3),s('type.list.list'(X2),'const.list.CONS_2'(s(X2,X4),s('type.list.list'(X2),'const.list.NIL_0'))))))) )),
  inference(subsumption_resolution,[],[f3487,f350])).
fof(f3487,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.list.EVERY_2'(s(fun(X2,bool),X3),s('type.list.list'(X2),'const.list.CONS_2'(s(X2,X4),s('type.list.list'(X2),'const.list.NIL_0')))))) | ~p(s(bool,happ(s(fun(X2,bool),X3),s(X2,X4))))) )),
  inference(superposition,[],[f445,f408])).
fof(f5736,plain,(
  ( ! [X6,X8,X7,X5] : (~p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X8)))) | p(s(bool,happ(s(fun(X6,bool),X7),s(X6,sK25(X7,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X5)))),X6)))))) )),
  inference(subsumption_resolution,[],[f5641,f350])).
fof(f5641,plain,(
  ( ! [X6,X8,X7,X5] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X8)))) | p(s(bool,happ(s(fun(X6,bool),X7),s(X6,sK25(X7,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X5)))),X6)))))) )),
  inference(superposition,[],[f436,f907])).
fof(f5737,plain,(
  ( ! [X12,X10,X11,X9] : (~p(s(bool,happ(s(fun(X10,bool),X11),s(X10,X12)))) | p(s(bool,happ(s(fun(X10,bool),X11),s(X10,sK26(X11,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X9)))),X10)))))) )),
  inference(subsumption_resolution,[],[f5642,f350])).
fof(f5642,plain,(
  ( ! [X12,X10,X11,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X10,bool),X11),s(X10,X12)))) | p(s(bool,happ(s(fun(X10,bool),X11),s(X10,sK26(X11,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X9)))),X10)))))) )),
  inference(superposition,[],[f438,f907])).
fof(f10034,plain,(
  ( ! [X14,X12,X10,X13,X11] : (~p(s(bool,happ(s(fun(X11,bool),X12),s(X11,X13)))) | p(s(bool,happ(s(fun(X11,bool),X12),s(X11,sK25(X12,X10,X11))))) | s(bool,X10) = s(bool,X14) | s(bool,'T_0') = s(bool,X14)) )),
  inference(subsumption_resolution,[],[f8983,f350])).
fof(f8983,plain,(
  ( ! [X14,X12,X10,X13,X11] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X11,bool),X12),s(X11,X13)))) | p(s(bool,happ(s(fun(X11,bool),X12),s(X11,sK25(X12,X10,X11))))) | s(bool,X10) = s(bool,X14) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f436,f983])).
fof(f10035,plain,(
  ( ! [X19,X17,X15,X18,X16] : (~p(s(bool,happ(s(fun(X16,bool),X17),s(X16,X18)))) | p(s(bool,happ(s(fun(X16,bool),X17),s(X16,sK26(X17,X15,X16))))) | s(bool,X15) = s(bool,X19) | s(bool,'T_0') = s(bool,X19)) )),
  inference(subsumption_resolution,[],[f8984,f350])).
fof(f8984,plain,(
  ( ! [X19,X17,X15,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X16,bool),X17),s(X16,X18)))) | p(s(bool,happ(s(fun(X16,bool),X17),s(X16,sK26(X17,X15,X16))))) | s(bool,X15) = s(bool,X19) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f438,f983])).
fof(f10056,plain,(
  ( ! [X344,X341,X343,X345,X342] : (p(s(bool,happ(s(fun(X341,bool),X344),s(X341,sK29(X342,X344,X341))))) | p(s(bool,happ(s(fun(X341,bool),X342),s(X341,sK29(X342,X344,X341))))) | s(bool,X345) = s(bool,happ(s(fun(X341,bool),X342),s(X341,X343))) | s(bool,'T_0') = s(bool,X345)) )),
  inference(subsumption_resolution,[],[f9100,f350])).
fof(f9100,plain,(
  ( ! [X344,X341,X343,X345,X342] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X341,bool),X344),s(X341,sK29(X342,X344,X341))))) | p(s(bool,happ(s(fun(X341,bool),X342),s(X341,sK29(X342,X344,X341))))) | s(bool,X345) = s(bool,happ(s(fun(X341,bool),X342),s(X341,X343))) | s(bool,'T_0') = s(bool,X345)) )),
  inference(superposition,[],[f442,f983])).
fof(f10057,plain,(
  ( ! [X350,X346,X349,X347,X348] : (p(s(bool,happ(s(fun(X346,bool),X347),s(X346,sK29(X349,X347,X346))))) | p(s(bool,happ(s(fun(X346,bool),X349),s(X346,sK29(X349,X347,X346))))) | s(bool,X350) = s(bool,happ(s(fun(X346,bool),X347),s(X346,X348))) | s(bool,'T_0') = s(bool,X350)) )),
  inference(subsumption_resolution,[],[f9101,f350])).
fof(f9101,plain,(
  ( ! [X350,X346,X349,X347,X348] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X346,bool),X347),s(X346,sK29(X349,X347,X346))))) | p(s(bool,happ(s(fun(X346,bool),X349),s(X346,sK29(X349,X347,X346))))) | s(bool,X350) = s(bool,happ(s(fun(X346,bool),X347),s(X346,X348))) | s(bool,'T_0') = s(bool,X350)) )),
  inference(superposition,[],[f441,f983])).
fof(f10058,plain,(
  ( ! [X352,X354,X351,X353,X355] : (p(s(bool,happ(s(fun(X351,bool),X354),s(X351,sK27(X352,X354,X351))))) | p(s(bool,happ(s(fun(X351,bool),X352),s(X351,sK28(X352,X354,X351))))) | s(bool,X355) = s(bool,happ(s(fun(X351,bool),X352),s(X351,X353))) | s(bool,'T_0') = s(bool,X355)) )),
  inference(subsumption_resolution,[],[f9102,f350])).
fof(f9102,plain,(
  ( ! [X352,X354,X351,X353,X355] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X351,bool),X354),s(X351,sK27(X352,X354,X351))))) | p(s(bool,happ(s(fun(X351,bool),X352),s(X351,sK28(X352,X354,X351))))) | s(bool,X355) = s(bool,happ(s(fun(X351,bool),X352),s(X351,X353))) | s(bool,'T_0') = s(bool,X355)) )),
  inference(superposition,[],[f440,f983])).
fof(f10059,plain,(
  ( ! [X356,X358,X360,X357,X359] : (p(s(bool,happ(s(fun(X356,bool),X357),s(X356,sK27(X359,X357,X356))))) | p(s(bool,happ(s(fun(X356,bool),X359),s(X356,sK28(X359,X357,X356))))) | s(bool,X360) = s(bool,happ(s(fun(X356,bool),X357),s(X356,X358))) | s(bool,'T_0') = s(bool,X360)) )),
  inference(subsumption_resolution,[],[f9103,f350])).
fof(f9103,plain,(
  ( ! [X356,X358,X360,X357,X359] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X356,bool),X357),s(X356,sK27(X359,X357,X356))))) | p(s(bool,happ(s(fun(X356,bool),X359),s(X356,sK28(X359,X357,X356))))) | s(bool,X360) = s(bool,happ(s(fun(X356,bool),X357),s(X356,X358))) | s(bool,'T_0') = s(bool,X360)) )),
  inference(superposition,[],[f439,f983])).
fof(f10060,plain,(
  ( ! [X364,X362,X365,X361,X363] : (p(s(bool,X364)) | p(s(bool,happ(s(fun(X361,bool),X362),s(X361,sK24(X362,X364,X361))))) | s(bool,X365) = s(bool,happ(s(fun(X361,bool),X362),s(X361,X363))) | s(bool,'T_0') = s(bool,X365)) )),
  inference(subsumption_resolution,[],[f9104,f350])).
fof(f9104,plain,(
  ( ! [X364,X362,X365,X361,X363] : (~p(s(bool,'T_0')) | p(s(bool,X364)) | p(s(bool,happ(s(fun(X361,bool),X362),s(X361,sK24(X362,X364,X361))))) | s(bool,X365) = s(bool,happ(s(fun(X361,bool),X362),s(X361,X363))) | s(bool,'T_0') = s(bool,X365)) )),
  inference(superposition,[],[f434,f983])).
fof(f10061,plain,(
  ( ! [X366,X368,X370,X367,X369] : (p(s(bool,X369)) | p(s(bool,happ(s(fun(X366,bool),X367),s(X366,sK23(X367,X369,X366))))) | s(bool,X370) = s(bool,happ(s(fun(X366,bool),X367),s(X366,X368))) | s(bool,'T_0') = s(bool,X370)) )),
  inference(subsumption_resolution,[],[f9105,f350])).
fof(f9105,plain,(
  ( ! [X366,X368,X370,X367,X369] : (~p(s(bool,'T_0')) | p(s(bool,X369)) | p(s(bool,happ(s(fun(X366,bool),X367),s(X366,sK23(X367,X369,X366))))) | s(bool,X370) = s(bool,happ(s(fun(X366,bool),X367),s(X366,X368))) | s(bool,'T_0') = s(bool,X370)) )),
  inference(superposition,[],[f432,f983])).
fof(f10063,plain,(
  ( ! [X377,X376,X378,X375] : (p(s(bool,happ(s(fun(X375,bool),X376),s(X375,X377)))) | s(bool,X378) = s(bool,happ(s(fun(X375,bool),X376),s(X375,sK19(X376,X375)))) | s(bool,'T_0') = s(bool,X378)) )),
  inference(subsumption_resolution,[],[f9107,f350])).
fof(f9107,plain,(
  ( ! [X377,X376,X378,X375] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X375,bool),X376),s(X375,X377)))) | s(bool,X378) = s(bool,happ(s(fun(X375,bool),X376),s(X375,sK19(X376,X375)))) | s(bool,'T_0') = s(bool,X378)) )),
  inference(superposition,[],[f414,f983])).
fof(f10064,plain,(
  ( ! [X379,X380,X382,X381] : (p(s(bool,happ(s(fun(X379,bool),X380),s(X379,X381)))) | s(bool,X382) = s(bool,happ(s(fun(X379,bool),X380),s(X379,sK20(X380,X379)))) | s(bool,'T_0') = s(bool,X382)) )),
  inference(subsumption_resolution,[],[f9108,f350])).
fof(f9108,plain,(
  ( ! [X379,X380,X382,X381] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X379,bool),X380),s(X379,X381)))) | s(bool,X382) = s(bool,happ(s(fun(X379,bool),X380),s(X379,sK20(X380,X379)))) | s(bool,'T_0') = s(bool,X382)) )),
  inference(superposition,[],[f415,f983])).
fof(f10072,plain,(
  ( ! [X412,X414,X416,X413,X415] : (p(s(bool,happ(s(fun(X412,bool),X413),s(X412,X414)))) | s(bool,X416) = s(bool,'const.list.EVERY_2'(s(fun(X412,bool),X413),s('type.list.list'(X412),'const.list.CONS_2'(s(X412,X414),s('type.list.list'(X412),X415))))) | s(bool,'T_0') = s(bool,X416)) )),
  inference(subsumption_resolution,[],[f9119,f350])).
fof(f9119,plain,(
  ( ! [X412,X414,X416,X413,X415] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X412,bool),X413),s(X412,X414)))) | s(bool,X416) = s(bool,'const.list.EVERY_2'(s(fun(X412,bool),X413),s('type.list.list'(X412),'const.list.CONS_2'(s(X412,X414),s('type.list.list'(X412),X415))))) | s(bool,'T_0') = s(bool,X416)) )),
  inference(superposition,[],[f443,f983])).
fof(f10074,plain,(
  ( ! [X422,X424,X426,X423,X425] : (p(s(bool,'const.list.EVERY_2'(s(fun(X422,bool),X423),s('type.list.list'(X422),'const.list.CONS_2'(s(X422,X425),s('type.list.list'(X422),X424)))))) | ~p(s(bool,happ(s(fun(X422,bool),X423),s(X422,X425)))) | s(bool,X426) = s(bool,'const.list.EVERY_2'(s(fun(X422,bool),X423),s('type.list.list'(X422),X424))) | s(bool,'T_0') = s(bool,X426)) )),
  inference(subsumption_resolution,[],[f9121,f350])).
fof(f9121,plain,(
  ( ! [X422,X424,X426,X423,X425] : (~p(s(bool,'T_0')) | p(s(bool,'const.list.EVERY_2'(s(fun(X422,bool),X423),s('type.list.list'(X422),'const.list.CONS_2'(s(X422,X425),s('type.list.list'(X422),X424)))))) | ~p(s(bool,happ(s(fun(X422,bool),X423),s(X422,X425)))) | s(bool,X426) = s(bool,'const.list.EVERY_2'(s(fun(X422,bool),X423),s('type.list.list'(X422),X424))) | s(bool,'T_0') = s(bool,X426)) )),
  inference(superposition,[],[f445,f983])).
fof(f10075,plain,(
  ( ! [X428,X430,X429,X431,X427] : (~p(s(bool,'const.bool.IN_2'(s(X427,X430),s(fun(X427,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X427),X429)))))) | p(s(bool,happ(s(fun(X427,bool),X428),s(X427,X430)))) | s(bool,X431) = s(bool,'const.list.EVERY_2'(s(fun(X427,bool),X428),s('type.list.list'(X427),X429))) | s(bool,'T_0') = s(bool,X431)) )),
  inference(subsumption_resolution,[],[f9122,f350])).
fof(f9122,plain,(
  ( ! [X428,X430,X429,X431,X427] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(X427,X430),s(fun(X427,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X427),X429)))))) | p(s(bool,happ(s(fun(X427,bool),X428),s(X427,X430)))) | s(bool,X431) = s(bool,'const.list.EVERY_2'(s(fun(X427,bool),X428),s('type.list.list'(X427),X429))) | s(bool,'T_0') = s(bool,X431)) )),
  inference(superposition,[],[f428,f983])).
fof(f10062,plain,(
  ( ! [X372,X374,X373,X371] : (p(s(bool,'const.list.EVERY_2'(s(fun(X371,bool),X372),s('type.list.list'(X371),X373)))) | s(bool,X374) = s(bool,happ(s(fun(X371,bool),X372),s(X371,'const.list.EL_2'(s('type.num.num',sK21(X372,X373,X371)),s('type.list.list'(X371),X373))))) | s(bool,'T_0') = s(bool,X374)) )),
  inference(subsumption_resolution,[],[f9106,f350])).
fof(f9106,plain,(
  ( ! [X372,X374,X373,X371] : (~p(s(bool,'T_0')) | p(s(bool,'const.list.EVERY_2'(s(fun(X371,bool),X372),s('type.list.list'(X371),X373)))) | s(bool,X374) = s(bool,happ(s(fun(X371,bool),X372),s(X371,'const.list.EL_2'(s('type.num.num',sK21(X372,X373,X371)),s('type.list.list'(X371),X373))))) | s(bool,'T_0') = s(bool,X374)) )),
  inference(superposition,[],[f427,f983])).
fof(f427,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,'const.list.EL_2'(s('type.num.num',sK21(X2,X1,X0)),s('type.list.list'(X0),X1)))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X2),s('type.list.list'(X0),X1))))) )),
  inference(cnf_transformation,[],[f275])).
fof(f275,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X2),s('type.list.list'(X0),X1)))) | ! [X4] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1)))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,'const.list.EL_2'(s('type.num.num',X4),s('type.list.list'(X0),X1)))))))) & ((p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK21(X2,X1,X0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1)))))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,'const.list.EL_2'(s('type.num.num',sK21(X2,X1,X0)),s('type.list.list'(X0),X1))))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X2),s('type.list.list'(X0),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK21])],[f274])).
fof(f274,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X2),s('type.list.list'(X0),X1)))) | ! [X4] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1)))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,'const.list.EL_2'(s('type.num.num',X4),s('type.list.list'(X0),X1)))))))) & (? [X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1)))))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,'const.list.EL_2'(s('type.num.num',X3),s('type.list.list'(X0),X1))))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X2),s('type.list.list'(X0),X1))))))),
  inference(rectify,[],[f273])).
fof(f273,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X2),s('type.list.list'(X0),X1)))) | ! [X3] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1)))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,'const.list.EL_2'(s('type.num.num',X3),s('type.list.list'(X0),X1)))))))) & (? [X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1)))))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,'const.list.EL_2'(s('type.num.num',X3),s('type.list.list'(X0),X1))))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X2),s('type.list.list'(X0),X1))))))),
  inference(nnf_transformation,[],[f199])).
fof(f199,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X2),s('type.list.list'(X0),X1)))) <=> ! [X3] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1)))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,'const.list.EL_2'(s('type.num.num',X3),s('type.list.list'(X0),X1))))))))),
  inference(ennf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X2),s('type.list.list'(X0),X1)))) <=> ! [X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1)))))) => p(s(bool,happ(s(fun(X0,bool),X2),s(X0,'const.list.EL_2'(s('type.num.num',X3),s('type.list.list'(X0),X1))))))))),
  inference(flattening,[],[f133])).
fof(f133,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X2),s('type.list.list'(X0),X1)))) <=> ! [X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1)))))) => p(s(bool,happ(s(fun(X0,bool),X2),s(X0,'const.list.EL_2'(s('type.num.num',X3),s('type.list.list'(X0),X1))))))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6] : ! [X7,X9] : (p(s(bool,'const.list.EVERY_2'(s(fun(X6,bool),X9),s('type.list.list'(X6),X7)))) <=> ! [X11] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X7)))))) => p(s(bool,happ(s(fun(X6,bool),X9),s(X6,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(X6),X7))))))))),
  file('thm.list.MONO_EVERY','thm.list.EVERY_EL')).
fof(f3160,plain,(
  ( ! [X14,X12,X13,X11] : (~p(s(bool,happ(s(fun(X11,bool),X14),s(X11,'const.list.EL_2'(s('type.num.num',sK21(X14,'const.list.HD_1'(s('type.list.list'('type.list.list'(X11)),'const.list.CONS_2'(s('type.list.list'(X11),X12),s('type.list.list'('type.list.list'(X11)),X13)))),X11)),s('type.list.list'(X11),X12)))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X11,bool),X14),s('type.list.list'(X11),X12))))) )),
  inference(forward_demodulation,[],[f3150,f419])).
fof(f3150,plain,(
  ( ! [X14,X12,X13,X11] : (~p(s(bool,happ(s(fun(X11,bool),X14),s(X11,'const.list.EL_2'(s('type.num.num',sK21(X14,'const.list.HD_1'(s('type.list.list'('type.list.list'(X11)),'const.list.CONS_2'(s('type.list.list'(X11),X12),s('type.list.list'('type.list.list'(X11)),X13)))),X11)),s('type.list.list'(X11),X12)))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X11,bool),X14),s('type.list.list'(X11),'const.list.HD_1'(s('type.list.list'('type.list.list'(X11)),'const.list.CONS_2'(s('type.list.list'(X11),X12),s('type.list.list'('type.list.list'(X11)),X13))))))))) )),
  inference(superposition,[],[f427,f419])).
fof(f3156,plain,(
  ( ! [X10,X8,X7,X9] : (~p(s(bool,happ(s(fun(X7,bool),X8),s(X7,'const.list.EL_2'(s('type.num.num',sK21('const.list.HD_1'(s('type.list.list'(fun(X7,bool)),'const.list.CONS_2'(s(fun(X7,bool),X8),s('type.list.list'(fun(X7,bool)),X9)))),X10,X7)),s('type.list.list'(X7),X10)))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X7,bool),X8),s('type.list.list'(X7),X10))))) )),
  inference(forward_demodulation,[],[f3146,f419])).
fof(f3146,plain,(
  ( ! [X10,X8,X7,X9] : (~p(s(bool,happ(s(fun(X7,bool),X8),s(X7,'const.list.EL_2'(s('type.num.num',sK21('const.list.HD_1'(s('type.list.list'(fun(X7,bool)),'const.list.CONS_2'(s(fun(X7,bool),X8),s('type.list.list'(fun(X7,bool)),X9)))),X10,X7)),s('type.list.list'(X7),X10)))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X7,bool),'const.list.HD_1'(s('type.list.list'(fun(X7,bool)),'const.list.CONS_2'(s(fun(X7,bool),X8),s('type.list.list'(fun(X7,bool)),X9))))),s('type.list.list'(X7),X10))))) )),
  inference(superposition,[],[f427,f419])).
fof(f3157,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,'const.list.EL_2'(s('type.num.num',sK21(X3,'const.list.TL_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))),X0)),s('type.list.list'(X0),X2)))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X3),s('type.list.list'(X0),X2))))) )),
  inference(forward_demodulation,[],[f3147,f420])).
fof(f3147,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,'const.list.EL_2'(s('type.num.num',sK21(X3,'const.list.TL_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))),X0)),s('type.list.list'(X0),X2)))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X3),s('type.list.list'(X0),'const.list.TL_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2))))))))) )),
  inference(superposition,[],[f427,f420])).
fof(f3159,plain,(
  ( ! [X10,X8,X7,X9] : (~p(s(bool,happ(s(fun(X7,bool),X10),s(X7,'const.list.EL_2'(s('type.num.num',sK21(X10,'const.list.EL_2'(s('type.num.num',X8),s('type.list.list'('type.list.list'(X7)),'const.list.TL_1'(s('type.list.list'('type.list.list'(X7)),X9)))),X7)),s('type.list.list'(X7),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.list.list'('type.list.list'(X7)),X9)))))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X7,bool),X10),s('type.list.list'(X7),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.list.list'('type.list.list'(X7)),X9))))))) )),
  inference(forward_demodulation,[],[f3149,f422])).
fof(f3149,plain,(
  ( ! [X10,X8,X7,X9] : (~p(s(bool,happ(s(fun(X7,bool),X10),s(X7,'const.list.EL_2'(s('type.num.num',sK21(X10,'const.list.EL_2'(s('type.num.num',X8),s('type.list.list'('type.list.list'(X7)),'const.list.TL_1'(s('type.list.list'('type.list.list'(X7)),X9)))),X7)),s('type.list.list'(X7),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.list.list'('type.list.list'(X7)),X9)))))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X7,bool),X10),s('type.list.list'(X7),'const.list.EL_2'(s('type.num.num',X8),s('type.list.list'('type.list.list'(X7)),'const.list.TL_1'(s('type.list.list'('type.list.list'(X7)),X9))))))))) )),
  inference(superposition,[],[f427,f422])).
fof(f3158,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,happ(s(fun(X4,bool),X6),s(X4,'const.list.EL_2'(s('type.num.num',sK21(X6,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'(X4)),X5)),X4)),s('type.list.list'(X4),'const.list.HD_1'(s('type.list.list'('type.list.list'(X4)),X5)))))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X4,bool),X6),s('type.list.list'(X4),'const.list.HD_1'(s('type.list.list'('type.list.list'(X4)),X5))))))) )),
  inference(forward_demodulation,[],[f3148,f409])).
fof(f3148,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,happ(s(fun(X4,bool),X6),s(X4,'const.list.EL_2'(s('type.num.num',sK21(X6,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'(X4)),X5)),X4)),s('type.list.list'(X4),'const.list.HD_1'(s('type.list.list'('type.list.list'(X4)),X5)))))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X4,bool),X6),s('type.list.list'(X4),'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'(X4)),X5))))))) )),
  inference(superposition,[],[f427,f409])).
fof(f10073,plain,(
  ( ! [X420,X418,X421,X417,X419] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X420),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X417),X419)))))) | p(s(bool,happ(s(fun(X417,bool),X418),s(X417,'const.list.EL_2'(s('type.num.num',X420),s('type.list.list'(X417),X419)))))) | s(bool,X421) = s(bool,'const.list.EVERY_2'(s(fun(X417,bool),X418),s('type.list.list'(X417),X419))) | s(bool,'T_0') = s(bool,X421)) )),
  inference(subsumption_resolution,[],[f9120,f350])).
fof(f9120,plain,(
  ( ! [X420,X418,X421,X417,X419] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X420),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X417),X419)))))) | p(s(bool,happ(s(fun(X417,bool),X418),s(X417,'const.list.EL_2'(s('type.num.num',X420),s('type.list.list'(X417),X419)))))) | s(bool,X421) = s(bool,'const.list.EVERY_2'(s(fun(X417,bool),X418),s('type.list.list'(X417),X419))) | s(bool,'T_0') = s(bool,X421)) )),
  inference(superposition,[],[f425,f983])).
fof(f425,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X2),s('type.list.list'(X0),X1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1)))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,'const.list.EL_2'(s('type.num.num',X4),s('type.list.list'(X0),X1))))))) )),
  inference(cnf_transformation,[],[f275])).
fof(f3151,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK21(X1,'const.list.TL_1'(s('type.list.list'(X0),X2)),X0)))),s('type.list.list'(X0),X2)))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.TL_1'(s('type.list.list'(X0),X2))))))) )),
  inference(superposition,[],[f427,f422])).
fof(f10065,plain,(
  ( ! [X383,X385,X384,X386] : (p(s(bool,'const.list.EVERY_2'(s(fun(X383,bool),X384),s('type.list.list'(X383),X385)))) | s(bool,X386) = s(bool,happ(s(fun(X383,bool),X384),s(X383,sK22(X385,X384,X383)))) | s(bool,'T_0') = s(bool,X386)) )),
  inference(subsumption_resolution,[],[f9109,f350])).
fof(f9109,plain,(
  ( ! [X383,X385,X384,X386] : (~p(s(bool,'T_0')) | p(s(bool,'const.list.EVERY_2'(s(fun(X383,bool),X384),s('type.list.list'(X383),X385)))) | s(bool,X386) = s(bool,happ(s(fun(X383,bool),X384),s(X383,sK22(X385,X384,X383)))) | s(bool,'T_0') = s(bool,X386)) )),
  inference(superposition,[],[f430,f983])).
fof(f430,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK22(X2,X1,X0))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))))) )),
  inference(cnf_transformation,[],[f278])).
fof(f2681,plain,(
  ( ! [X10,X8,X7,X9] : (~p(s(bool,happ(s(fun(X7,bool),X8),s(X7,sK22(X10,'const.list.HD_1'(s('type.list.list'(fun(X7,bool)),'const.list.CONS_2'(s(fun(X7,bool),X8),s('type.list.list'(fun(X7,bool)),X9)))),X7))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X7,bool),X8),s('type.list.list'(X7),X10))))) )),
  inference(forward_demodulation,[],[f2676,f419])).
fof(f2676,plain,(
  ( ! [X10,X8,X7,X9] : (~p(s(bool,happ(s(fun(X7,bool),X8),s(X7,sK22(X10,'const.list.HD_1'(s('type.list.list'(fun(X7,bool)),'const.list.CONS_2'(s(fun(X7,bool),X8),s('type.list.list'(fun(X7,bool)),X9)))),X7))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X7,bool),'const.list.HD_1'(s('type.list.list'(fun(X7,bool)),'const.list.CONS_2'(s(fun(X7,bool),X8),s('type.list.list'(fun(X7,bool)),X9))))),s('type.list.list'(X7),X10))))) )),
  inference(superposition,[],[f430,f419])).
fof(f3376,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),'const.list.HD_1'(s('type.list.list'(fun(X0,bool)),X1))),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),'const.list.HD_1'(s('type.list.list'(fun(X0,bool)),X1))),s(X0,sK28('const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(fun(X0,bool)),X1)),X3,X0))))) | p(s(bool,happ(s(fun(X0,bool),X3),s(X0,sK27('const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(fun(X0,bool)),X1)),X3,X0)))))) )),
  inference(forward_demodulation,[],[f3360,f409])).
fof(f3360,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),'const.list.HD_1'(s('type.list.list'(fun(X0,bool)),X1))),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X3),s(X0,sK27('const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(fun(X0,bool)),X1)),X3,X0))))) | p(s(bool,happ(s(fun(X0,bool),'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(fun(X0,bool)),X1))),s(X0,sK28('const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(fun(X0,bool)),X1)),X3,X0)))))) )),
  inference(superposition,[],[f440,f409])).
fof(f3377,plain,(
  ( ! [X6,X4,X8,X7,X5] : (~p(s(bool,happ(s(fun(X4,bool),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))),s('type.list.list'(fun(X4,bool)),X6))),s(X4,X7)))) | p(s(bool,happ(s(fun(X4,bool),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))),s('type.list.list'(fun(X4,bool)),X6))),s(X4,sK28('const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(fun(X4,bool)),'const.list.TL_1'(s('type.list.list'(fun(X4,bool)),X6)))),X8,X4))))) | p(s(bool,happ(s(fun(X4,bool),X8),s(X4,sK27('const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(fun(X4,bool)),'const.list.TL_1'(s('type.list.list'(fun(X4,bool)),X6)))),X8,X4)))))) )),
  inference(forward_demodulation,[],[f3361,f422])).
fof(f3361,plain,(
  ( ! [X6,X4,X8,X7,X5] : (~p(s(bool,happ(s(fun(X4,bool),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))),s('type.list.list'(fun(X4,bool)),X6))),s(X4,X7)))) | p(s(bool,happ(s(fun(X4,bool),X8),s(X4,sK27('const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(fun(X4,bool)),'const.list.TL_1'(s('type.list.list'(fun(X4,bool)),X6)))),X8,X4))))) | p(s(bool,happ(s(fun(X4,bool),'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(fun(X4,bool)),'const.list.TL_1'(s('type.list.list'(fun(X4,bool)),X6))))),s(X4,sK28('const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(fun(X4,bool)),'const.list.TL_1'(s('type.list.list'(fun(X4,bool)),X6)))),X8,X4)))))) )),
  inference(superposition,[],[f440,f422])).
fof(f3338,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),'const.list.HD_1'(s('type.list.list'(fun(X0,bool)),X1))),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),'const.list.HD_1'(s('type.list.list'(fun(X0,bool)),X1))),s(X0,sK27(X3,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(fun(X0,bool)),X1)),X0))))) | p(s(bool,happ(s(fun(X0,bool),X3),s(X0,sK28(X3,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(fun(X0,bool)),X1)),X0)))))) )),
  inference(forward_demodulation,[],[f3322,f409])).
fof(f3322,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),'const.list.HD_1'(s('type.list.list'(fun(X0,bool)),X1))),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(fun(X0,bool)),X1))),s(X0,sK27(X3,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(fun(X0,bool)),X1)),X0))))) | p(s(bool,happ(s(fun(X0,bool),X3),s(X0,sK28(X3,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(fun(X0,bool)),X1)),X0)))))) )),
  inference(superposition,[],[f439,f409])).
fof(f3339,plain,(
  ( ! [X6,X4,X8,X7,X5] : (~p(s(bool,happ(s(fun(X4,bool),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))),s('type.list.list'(fun(X4,bool)),X6))),s(X4,X7)))) | p(s(bool,happ(s(fun(X4,bool),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))),s('type.list.list'(fun(X4,bool)),X6))),s(X4,sK27(X8,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(fun(X4,bool)),'const.list.TL_1'(s('type.list.list'(fun(X4,bool)),X6)))),X4))))) | p(s(bool,happ(s(fun(X4,bool),X8),s(X4,sK28(X8,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(fun(X4,bool)),'const.list.TL_1'(s('type.list.list'(fun(X4,bool)),X6)))),X4)))))) )),
  inference(forward_demodulation,[],[f3323,f422])).
fof(f3323,plain,(
  ( ! [X6,X4,X8,X7,X5] : (~p(s(bool,happ(s(fun(X4,bool),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))),s('type.list.list'(fun(X4,bool)),X6))),s(X4,X7)))) | p(s(bool,happ(s(fun(X4,bool),'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(fun(X4,bool)),'const.list.TL_1'(s('type.list.list'(fun(X4,bool)),X6))))),s(X4,sK27(X8,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(fun(X4,bool)),'const.list.TL_1'(s('type.list.list'(fun(X4,bool)),X6)))),X4))))) | p(s(bool,happ(s(fun(X4,bool),X8),s(X4,sK28(X8,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(fun(X4,bool)),'const.list.TL_1'(s('type.list.list'(fun(X4,bool)),X6)))),X4)))))) )),
  inference(superposition,[],[f439,f422])).
fof(f3423,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),'const.list.HD_1'(s('type.list.list'(fun(X0,bool)),X1))),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),'const.list.HD_1'(s('type.list.list'(fun(X0,bool)),X1))),s(X0,sK29(X3,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(fun(X0,bool)),X1)),X0))))) | p(s(bool,happ(s(fun(X0,bool),X3),s(X0,sK29(X3,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(fun(X0,bool)),X1)),X0)))))) )),
  inference(forward_demodulation,[],[f3407,f409])).
fof(f3407,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),'const.list.HD_1'(s('type.list.list'(fun(X0,bool)),X1))),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(fun(X0,bool)),X1))),s(X0,sK29(X3,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(fun(X0,bool)),X1)),X0))))) | p(s(bool,happ(s(fun(X0,bool),X3),s(X0,sK29(X3,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(fun(X0,bool)),X1)),X0)))))) )),
  inference(superposition,[],[f441,f409])).
fof(f3424,plain,(
  ( ! [X6,X4,X8,X7,X5] : (~p(s(bool,happ(s(fun(X4,bool),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))),s('type.list.list'(fun(X4,bool)),X6))),s(X4,X7)))) | p(s(bool,happ(s(fun(X4,bool),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))),s('type.list.list'(fun(X4,bool)),X6))),s(X4,sK29(X8,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(fun(X4,bool)),'const.list.TL_1'(s('type.list.list'(fun(X4,bool)),X6)))),X4))))) | p(s(bool,happ(s(fun(X4,bool),X8),s(X4,sK29(X8,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(fun(X4,bool)),'const.list.TL_1'(s('type.list.list'(fun(X4,bool)),X6)))),X4)))))) )),
  inference(forward_demodulation,[],[f3408,f422])).
fof(f3408,plain,(
  ( ! [X6,X4,X8,X7,X5] : (~p(s(bool,happ(s(fun(X4,bool),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))),s('type.list.list'(fun(X4,bool)),X6))),s(X4,X7)))) | p(s(bool,happ(s(fun(X4,bool),'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(fun(X4,bool)),'const.list.TL_1'(s('type.list.list'(fun(X4,bool)),X6))))),s(X4,sK29(X8,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(fun(X4,bool)),'const.list.TL_1'(s('type.list.list'(fun(X4,bool)),X6)))),X4))))) | p(s(bool,happ(s(fun(X4,bool),X8),s(X4,sK29(X8,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(fun(X4,bool)),'const.list.TL_1'(s('type.list.list'(fun(X4,bool)),X6)))),X4)))))) )),
  inference(superposition,[],[f441,f422])).
fof(f3458,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),'const.list.HD_1'(s('type.list.list'(fun(X0,bool)),X1))),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),'const.list.HD_1'(s('type.list.list'(fun(X0,bool)),X1))),s(X0,sK29('const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(fun(X0,bool)),X1)),X3,X0))))) | p(s(bool,happ(s(fun(X0,bool),X3),s(X0,sK29('const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(fun(X0,bool)),X1)),X3,X0)))))) )),
  inference(forward_demodulation,[],[f3442,f409])).
fof(f3442,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),'const.list.HD_1'(s('type.list.list'(fun(X0,bool)),X1))),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X3),s(X0,sK29('const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(fun(X0,bool)),X1)),X3,X0))))) | p(s(bool,happ(s(fun(X0,bool),'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(fun(X0,bool)),X1))),s(X0,sK29('const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(fun(X0,bool)),X1)),X3,X0)))))) )),
  inference(superposition,[],[f442,f409])).
fof(f3459,plain,(
  ( ! [X6,X4,X8,X7,X5] : (~p(s(bool,happ(s(fun(X4,bool),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))),s('type.list.list'(fun(X4,bool)),X6))),s(X4,X7)))) | p(s(bool,happ(s(fun(X4,bool),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))),s('type.list.list'(fun(X4,bool)),X6))),s(X4,sK29('const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(fun(X4,bool)),'const.list.TL_1'(s('type.list.list'(fun(X4,bool)),X6)))),X8,X4))))) | p(s(bool,happ(s(fun(X4,bool),X8),s(X4,sK29('const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(fun(X4,bool)),'const.list.TL_1'(s('type.list.list'(fun(X4,bool)),X6)))),X8,X4)))))) )),
  inference(forward_demodulation,[],[f3443,f422])).
fof(f3443,plain,(
  ( ! [X6,X4,X8,X7,X5] : (~p(s(bool,happ(s(fun(X4,bool),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))),s('type.list.list'(fun(X4,bool)),X6))),s(X4,X7)))) | p(s(bool,happ(s(fun(X4,bool),X8),s(X4,sK29('const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(fun(X4,bool)),'const.list.TL_1'(s('type.list.list'(fun(X4,bool)),X6)))),X8,X4))))) | p(s(bool,happ(s(fun(X4,bool),'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(fun(X4,bool)),'const.list.TL_1'(s('type.list.list'(fun(X4,bool)),X6))))),s(X4,sK29('const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(fun(X4,bool)),'const.list.TL_1'(s('type.list.list'(fun(X4,bool)),X6)))),X8,X4)))))) )),
  inference(superposition,[],[f442,f422])).
fof(f2144,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),'const.list.HD_1'(s('type.list.list'(fun(X0,bool)),X1))),s(X0,sK19('const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(fun(X0,bool)),X1)),X0))))) | p(s(bool,happ(s(fun(X0,bool),'const.list.HD_1'(s('type.list.list'(fun(X0,bool)),X1))),s(X0,X2))))) )),
  inference(forward_demodulation,[],[f2140,f409])).
fof(f2140,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),'const.list.HD_1'(s('type.list.list'(fun(X0,bool)),X1))),s(X0,sK19('const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(fun(X0,bool)),X1)),X0))))) | p(s(bool,happ(s(fun(X0,bool),'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(fun(X0,bool)),X1))),s(X0,X2))))) )),
  inference(superposition,[],[f414,f409])).
fof(f2166,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),'const.list.HD_1'(s('type.list.list'(fun(X0,bool)),X1))),s(X0,sK20('const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(fun(X0,bool)),X1)),X0))))) | p(s(bool,happ(s(fun(X0,bool),'const.list.HD_1'(s('type.list.list'(fun(X0,bool)),X1))),s(X0,X2))))) )),
  inference(forward_demodulation,[],[f2162,f409])).
fof(f2162,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),'const.list.HD_1'(s('type.list.list'(fun(X0,bool)),X1))),s(X0,sK20('const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(fun(X0,bool)),X1)),X0))))) | p(s(bool,happ(s(fun(X0,bool),'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(fun(X0,bool)),X1))),s(X0,X2))))) )),
  inference(superposition,[],[f415,f409])).
fof(f2763,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),'const.list.HD_1'(s('type.list.list'(fun(X0,bool)),X1))),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),'const.list.HD_1'(s('type.list.list'(fun(X0,bool)),X1))),s(X0,sK23('const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(fun(X0,bool)),X1)),X3,X0))))) | p(s(bool,X3))) )),
  inference(forward_demodulation,[],[f2747,f409])).
fof(f2747,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),'const.list.HD_1'(s('type.list.list'(fun(X0,bool)),X1))),s(X0,X2)))) | p(s(bool,X3)) | p(s(bool,happ(s(fun(X0,bool),'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(fun(X0,bool)),X1))),s(X0,sK23('const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(fun(X0,bool)),X1)),X3,X0)))))) )),
  inference(superposition,[],[f432,f409])).
fof(f2800,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),'const.list.HD_1'(s('type.list.list'(fun(X0,bool)),X1))),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),'const.list.HD_1'(s('type.list.list'(fun(X0,bool)),X1))),s(X0,sK24('const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(fun(X0,bool)),X1)),X3,X0))))) | p(s(bool,X3))) )),
  inference(forward_demodulation,[],[f2784,f409])).
fof(f2784,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),'const.list.HD_1'(s('type.list.list'(fun(X0,bool)),X1))),s(X0,X2)))) | p(s(bool,X3)) | p(s(bool,happ(s(fun(X0,bool),'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(fun(X0,bool)),X1))),s(X0,sK24('const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(fun(X0,bool)),X1)),X3,X0)))))) )),
  inference(superposition,[],[f434,f409])).
fof(f3154,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),'const.list.HD_1'(s('type.list.list'(fun(X0,bool)),X1))),s(X0,'const.list.EL_2'(s('type.num.num',sK21('const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(fun(X0,bool)),X1)),X2,X0)),s('type.list.list'(X0),X2)))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),'const.list.HD_1'(s('type.list.list'(fun(X0,bool)),X1))),s('type.list.list'(X0),X2))))) )),
  inference(forward_demodulation,[],[f3144,f409])).
fof(f3144,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),'const.list.HD_1'(s('type.list.list'(fun(X0,bool)),X1))),s(X0,'const.list.EL_2'(s('type.num.num',sK21('const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(fun(X0,bool)),X1)),X2,X0)),s('type.list.list'(X0),X2)))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(fun(X0,bool)),X1))),s('type.list.list'(X0),X2))))) )),
  inference(superposition,[],[f427,f409])).
fof(f2679,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),'const.list.HD_1'(s('type.list.list'(fun(X0,bool)),X1))),s(X0,sK22(X2,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(fun(X0,bool)),X1)),X0))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),'const.list.HD_1'(s('type.list.list'(fun(X0,bool)),X1))),s('type.list.list'(X0),X2))))) )),
  inference(forward_demodulation,[],[f2674,f409])).
fof(f2674,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),'const.list.HD_1'(s('type.list.list'(fun(X0,bool)),X1))),s(X0,sK22(X2,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(fun(X0,bool)),X1)),X0))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(fun(X0,bool)),X1))),s('type.list.list'(X0),X2))))) )),
  inference(superposition,[],[f430,f409])).
fof(f2658,plain,(
  ( ! [X142,X140,X141,X139] : (~p(s(bool,happ(s(fun(X139,bool),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X140))),s('type.list.list'(fun(X139,bool)),X141))),s(X139,sK20('const.list.EL_2'(s('type.num.num',X140),s('type.list.list'(fun(X139,bool)),'const.list.TL_1'(s('type.list.list'(fun(X139,bool)),X141)))),X139))))) | p(s(bool,happ(s(fun(X139,bool),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X140))),s('type.list.list'(fun(X139,bool)),X141))),s(X139,X142))))) )),
  inference(forward_demodulation,[],[f2614,f422])).
fof(f2614,plain,(
  ( ! [X142,X140,X141,X139] : (~p(s(bool,happ(s(fun(X139,bool),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X140))),s('type.list.list'(fun(X139,bool)),X141))),s(X139,sK20('const.list.EL_2'(s('type.num.num',X140),s('type.list.list'(fun(X139,bool)),'const.list.TL_1'(s('type.list.list'(fun(X139,bool)),X141)))),X139))))) | p(s(bool,happ(s(fun(X139,bool),'const.list.EL_2'(s('type.num.num',X140),s('type.list.list'(fun(X139,bool)),'const.list.TL_1'(s('type.list.list'(fun(X139,bool)),X141))))),s(X139,X142))))) )),
  inference(superposition,[],[f415,f422])).
fof(f2659,plain,(
  ( ! [X146,X144,X145,X143] : (~p(s(bool,happ(s(fun(X143,bool),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X144))),s('type.list.list'(fun(X143,bool)),X145))),s(X143,sK19('const.list.EL_2'(s('type.num.num',X144),s('type.list.list'(fun(X143,bool)),'const.list.TL_1'(s('type.list.list'(fun(X143,bool)),X145)))),X143))))) | p(s(bool,happ(s(fun(X143,bool),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X144))),s('type.list.list'(fun(X143,bool)),X145))),s(X143,X146))))) )),
  inference(forward_demodulation,[],[f2615,f422])).
fof(f2615,plain,(
  ( ! [X146,X144,X145,X143] : (~p(s(bool,happ(s(fun(X143,bool),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X144))),s('type.list.list'(fun(X143,bool)),X145))),s(X143,sK19('const.list.EL_2'(s('type.num.num',X144),s('type.list.list'(fun(X143,bool)),'const.list.TL_1'(s('type.list.list'(fun(X143,bool)),X145)))),X143))))) | p(s(bool,happ(s(fun(X143,bool),'const.list.EL_2'(s('type.num.num',X144),s('type.list.list'(fun(X143,bool)),'const.list.TL_1'(s('type.list.list'(fun(X143,bool)),X145))))),s(X143,X146))))) )),
  inference(superposition,[],[f414,f422])).
fof(f2764,plain,(
  ( ! [X6,X4,X8,X7,X5] : (~p(s(bool,happ(s(fun(X4,bool),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))),s('type.list.list'(fun(X4,bool)),X6))),s(X4,X7)))) | p(s(bool,happ(s(fun(X4,bool),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))),s('type.list.list'(fun(X4,bool)),X6))),s(X4,sK23('const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(fun(X4,bool)),'const.list.TL_1'(s('type.list.list'(fun(X4,bool)),X6)))),X8,X4))))) | p(s(bool,X8))) )),
  inference(forward_demodulation,[],[f2748,f422])).
fof(f2748,plain,(
  ( ! [X6,X4,X8,X7,X5] : (~p(s(bool,happ(s(fun(X4,bool),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))),s('type.list.list'(fun(X4,bool)),X6))),s(X4,X7)))) | p(s(bool,X8)) | p(s(bool,happ(s(fun(X4,bool),'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(fun(X4,bool)),'const.list.TL_1'(s('type.list.list'(fun(X4,bool)),X6))))),s(X4,sK23('const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(fun(X4,bool)),'const.list.TL_1'(s('type.list.list'(fun(X4,bool)),X6)))),X8,X4)))))) )),
  inference(superposition,[],[f432,f422])).
fof(f2801,plain,(
  ( ! [X6,X4,X8,X7,X5] : (~p(s(bool,happ(s(fun(X4,bool),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))),s('type.list.list'(fun(X4,bool)),X6))),s(X4,X7)))) | p(s(bool,happ(s(fun(X4,bool),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))),s('type.list.list'(fun(X4,bool)),X6))),s(X4,sK24('const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(fun(X4,bool)),'const.list.TL_1'(s('type.list.list'(fun(X4,bool)),X6)))),X8,X4))))) | p(s(bool,X8))) )),
  inference(forward_demodulation,[],[f2785,f422])).
fof(f2785,plain,(
  ( ! [X6,X4,X8,X7,X5] : (~p(s(bool,happ(s(fun(X4,bool),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))),s('type.list.list'(fun(X4,bool)),X6))),s(X4,X7)))) | p(s(bool,X8)) | p(s(bool,happ(s(fun(X4,bool),'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(fun(X4,bool)),'const.list.TL_1'(s('type.list.list'(fun(X4,bool)),X6))))),s(X4,sK24('const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(fun(X4,bool)),'const.list.TL_1'(s('type.list.list'(fun(X4,bool)),X6)))),X8,X4)))))) )),
  inference(superposition,[],[f434,f422])).
fof(f3155,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,happ(s(fun(X3,bool),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))),s('type.list.list'(fun(X3,bool)),X5))),s(X3,'const.list.EL_2'(s('type.num.num',sK21('const.list.EL_2'(s('type.num.num',X4),s('type.list.list'(fun(X3,bool)),'const.list.TL_1'(s('type.list.list'(fun(X3,bool)),X5)))),X6,X3)),s('type.list.list'(X3),X6)))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X3,bool),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))),s('type.list.list'(fun(X3,bool)),X5))),s('type.list.list'(X3),X6))))) )),
  inference(forward_demodulation,[],[f3145,f422])).
fof(f3145,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,happ(s(fun(X3,bool),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))),s('type.list.list'(fun(X3,bool)),X5))),s(X3,'const.list.EL_2'(s('type.num.num',sK21('const.list.EL_2'(s('type.num.num',X4),s('type.list.list'(fun(X3,bool)),'const.list.TL_1'(s('type.list.list'(fun(X3,bool)),X5)))),X6,X3)),s('type.list.list'(X3),X6)))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X3,bool),'const.list.EL_2'(s('type.num.num',X4),s('type.list.list'(fun(X3,bool)),'const.list.TL_1'(s('type.list.list'(fun(X3,bool)),X5))))),s('type.list.list'(X3),X6))))) )),
  inference(superposition,[],[f427,f422])).
fof(f2680,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,happ(s(fun(X3,bool),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))),s('type.list.list'(fun(X3,bool)),X5))),s(X3,sK22(X6,'const.list.EL_2'(s('type.num.num',X4),s('type.list.list'(fun(X3,bool)),'const.list.TL_1'(s('type.list.list'(fun(X3,bool)),X5)))),X3))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X3,bool),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))),s('type.list.list'(fun(X3,bool)),X5))),s('type.list.list'(X3),X6))))) )),
  inference(forward_demodulation,[],[f2675,f422])).
fof(f2675,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,happ(s(fun(X3,bool),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))),s('type.list.list'(fun(X3,bool)),X5))),s(X3,sK22(X6,'const.list.EL_2'(s('type.num.num',X4),s('type.list.list'(fun(X3,bool)),'const.list.TL_1'(s('type.list.list'(fun(X3,bool)),X5)))),X3))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X3,bool),'const.list.EL_2'(s('type.num.num',X4),s('type.list.list'(fun(X3,bool)),'const.list.TL_1'(s('type.list.list'(fun(X3,bool)),X5))))),s('type.list.list'(X3),X6))))) )),
  inference(superposition,[],[f430,f422])).
fof(f8576,plain,(
  ( ! [X509,X507,X508,X506] : (~p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X508),s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X507),s('type.list.list'(bool),X509)))))) | p(s(bool,happ(s(fun(bool,bool),X508),s(bool,X506)))) | s(bool,'T_0') = s(bool,X507) | s(bool,'T_0') = s(bool,X506)) )),
  inference(superposition,[],[f443,f983])).
fof(f8633,plain,(
  ( ! [X673,X671,X672] : (~p(s(bool,happ(s(fun(bool,bool),X671),s(bool,X673)))) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X671),s('type.list.list'(bool),X672)))) | s(bool,'T_0') = s(bool,X673) | s(bool,'T_0') = s(bool,'const.list.EL_2'(s('type.num.num',sK21(X671,X672,bool)),s('type.list.list'(bool),X672)))) )),
  inference(superposition,[],[f427,f983])).
fof(f8635,plain,(
  ( ! [X677,X679,X678] : (~p(s(bool,happ(s(fun(bool,bool),X677),s(bool,X678)))) | p(s(bool,happ(s(fun(bool,bool),X677),s(bool,X679)))) | s(bool,'T_0') = s(bool,X678) | s(bool,'T_0') = s(bool,sK19(X677,bool))) )),
  inference(superposition,[],[f414,f983])).
fof(f8636,plain,(
  ( ! [X681,X680,X682] : (~p(s(bool,happ(s(fun(bool,bool),X680),s(bool,X681)))) | p(s(bool,happ(s(fun(bool,bool),X680),s(bool,X682)))) | s(bool,'T_0') = s(bool,X681) | s(bool,'T_0') = s(bool,sK20(X680,bool))) )),
  inference(superposition,[],[f415,f983])).
fof(f8638,plain,(
  ( ! [X687,X688,X686] : (~p(s(bool,happ(s(fun(bool,bool),X687),s(bool,X688)))) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X687),s('type.list.list'(bool),X686)))) | s(bool,'T_0') = s(bool,X688) | s(bool,'T_0') = s(bool,sK22(X686,X687,bool))) )),
  inference(superposition,[],[f430,f983])).
fof(f9141,plain,(
  ( ! [X499,X500,X498,X501] : (~p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X499),s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'T_0'),s('type.list.list'(bool),X500)))))) | p(s(bool,happ(s(fun(bool,bool),X499),s(bool,X498)))) | s(bool,X498) = s(bool,X501) | s(bool,'T_0') = s(bool,X501)) )),
  inference(superposition,[],[f443,f983])).
fof(f9200,plain,(
  ( ! [X668,X670,X669] : (~p(s(bool,happ(s(fun(bool,bool),X668),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X668),s(bool,X669)))) | s(bool,X670) = s(bool,sK19(X668,bool)) | s(bool,'T_0') = s(bool,X670)) )),
  inference(superposition,[],[f414,f983])).
fof(f9201,plain,(
  ( ! [X673,X671,X672] : (~p(s(bool,happ(s(fun(bool,bool),X671),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X671),s(bool,X672)))) | s(bool,X673) = s(bool,sK20(X671,bool)) | s(bool,'T_0') = s(bool,X673)) )),
  inference(superposition,[],[f415,f983])).
fof(f9203,plain,(
  ( ! [X677,X679,X678] : (~p(s(bool,happ(s(fun(bool,bool),X678),s(bool,'T_0')))) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X678),s('type.list.list'(bool),X677)))) | s(bool,X679) = s(bool,sK22(X677,X678,bool)) | s(bool,'T_0') = s(bool,X679)) )),
  inference(superposition,[],[f430,f983])).
fof(f9198,plain,(
  ( ! [X662,X664,X663] : (~p(s(bool,happ(s(fun(bool,bool),X662),s(bool,'T_0')))) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X662),s('type.list.list'(bool),X663)))) | s(bool,X664) = s(bool,'const.list.EL_2'(s('type.num.num',sK21(X662,X663,bool)),s('type.list.list'(bool),X663))) | s(bool,'T_0') = s(bool,X664)) )),
  inference(superposition,[],[f427,f983])).
fof(f8924,plain,(
  ( ! [X191,X189,X192,X190] : (~p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X190),s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X189),s('type.list.list'(bool),X191)))))) | p(s(bool,happ(s(fun(bool,bool),X190),s(bool,'T_0')))) | s(bool,X189) = s(bool,X192) | s(bool,'T_0') = s(bool,X192)) )),
  inference(superposition,[],[f443,f983])).
fof(f10066,plain,(
  ( ! [X389,X387,X388,X390] : (~p(s(bool,happ(s(fun('type.list.list'(X387),bool),X388),s('type.list.list'(X387),'const.list.CONS_2'(s(X387,sK18(X388,X387)),s('type.list.list'(X387),sK17(X388,X387))))))) | p(s(bool,happ(s(fun('type.list.list'(X387),bool),X388),s('type.list.list'(X387),X389)))) | s(bool,X390) = s(bool,happ(s(fun('type.list.list'(X387),bool),X388),s('type.list.list'(X387),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,X390)) )),
  inference(subsumption_resolution,[],[f9110,f350])).
fof(f9110,plain,(
  ( ! [X389,X387,X388,X390] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun('type.list.list'(X387),bool),X388),s('type.list.list'(X387),'const.list.CONS_2'(s(X387,sK18(X388,X387)),s('type.list.list'(X387),sK17(X388,X387))))))) | p(s(bool,happ(s(fun('type.list.list'(X387),bool),X388),s('type.list.list'(X387),X389)))) | s(bool,X390) = s(bool,happ(s(fun('type.list.list'(X387),bool),X388),s('type.list.list'(X387),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,X390)) )),
  inference(superposition,[],[f411,f983])).
fof(f3576,plain,(
  ( ! [X10,X8,X7,X9] : (~p(s(bool,happ(s(fun('type.list.list'(X7),bool),X8),s('type.list.list'(X7),'const.list.NIL_0')))) | ~p(s(bool,happ(s(fun('type.list.list'(X7),bool),X8),s('type.list.list'(X7),'const.list.CONS_2'(s(X7,sK18('const.list.HD_1'(s('type.list.list'(fun('type.list.list'(X7),bool)),'const.list.CONS_2'(s(fun('type.list.list'(X7),bool),X8),s('type.list.list'(fun('type.list.list'(X7),bool)),X9)))),X7)),s('type.list.list'(X7),sK17('const.list.HD_1'(s('type.list.list'(fun('type.list.list'(X7),bool)),'const.list.CONS_2'(s(fun('type.list.list'(X7),bool),X8),s('type.list.list'(fun('type.list.list'(X7),bool)),X9)))),X7))))))) | p(s(bool,happ(s(fun('type.list.list'(X7),bool),X8),s('type.list.list'(X7),X10))))) )),
  inference(forward_demodulation,[],[f3575,f419])).
fof(f3575,plain,(
  ( ! [X10,X8,X7,X9] : (~p(s(bool,happ(s(fun('type.list.list'(X7),bool),X8),s('type.list.list'(X7),'const.list.CONS_2'(s(X7,sK18('const.list.HD_1'(s('type.list.list'(fun('type.list.list'(X7),bool)),'const.list.CONS_2'(s(fun('type.list.list'(X7),bool),X8),s('type.list.list'(fun('type.list.list'(X7),bool)),X9)))),X7)),s('type.list.list'(X7),sK17('const.list.HD_1'(s('type.list.list'(fun('type.list.list'(X7),bool)),'const.list.CONS_2'(s(fun('type.list.list'(X7),bool),X8),s('type.list.list'(fun('type.list.list'(X7),bool)),X9)))),X7))))))) | ~p(s(bool,happ(s(fun('type.list.list'(X7),bool),X8),s('type.list.list'(X7),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X7),bool),'const.list.HD_1'(s('type.list.list'(fun('type.list.list'(X7),bool)),'const.list.CONS_2'(s(fun('type.list.list'(X7),bool),X8),s('type.list.list'(fun('type.list.list'(X7),bool)),X9))))),s('type.list.list'(X7),X10))))) )),
  inference(forward_demodulation,[],[f3569,f419])).
fof(f3569,plain,(
  ( ! [X10,X8,X7,X9] : (~p(s(bool,happ(s(fun('type.list.list'(X7),bool),X8),s('type.list.list'(X7),'const.list.NIL_0')))) | ~p(s(bool,happ(s(fun('type.list.list'(X7),bool),'const.list.HD_1'(s('type.list.list'(fun('type.list.list'(X7),bool)),'const.list.CONS_2'(s(fun('type.list.list'(X7),bool),X8),s('type.list.list'(fun('type.list.list'(X7),bool)),X9))))),s('type.list.list'(X7),'const.list.CONS_2'(s(X7,sK18('const.list.HD_1'(s('type.list.list'(fun('type.list.list'(X7),bool)),'const.list.CONS_2'(s(fun('type.list.list'(X7),bool),X8),s('type.list.list'(fun('type.list.list'(X7),bool)),X9)))),X7)),s('type.list.list'(X7),sK17('const.list.HD_1'(s('type.list.list'(fun('type.list.list'(X7),bool)),'const.list.CONS_2'(s(fun('type.list.list'(X7),bool),X8),s('type.list.list'(fun('type.list.list'(X7),bool)),X9)))),X7))))))) | p(s(bool,happ(s(fun('type.list.list'(X7),bool),'const.list.HD_1'(s('type.list.list'(fun('type.list.list'(X7),bool)),'const.list.CONS_2'(s(fun('type.list.list'(X7),bool),X8),s('type.list.list'(fun('type.list.list'(X7),bool)),X9))))),s('type.list.list'(X7),X10))))) )),
  inference(superposition,[],[f411,f419])).
fof(f411,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK18(X1,X0)),s('type.list.list'(X0),sK17(X1,X0))))))) | p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2))))) )),
  inference(cnf_transformation,[],[f265])).
fof(f265,plain,(
  ! [X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK17(X1,X0))))) & ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK18(X1,X0)),s('type.list.list'(X0),sK17(X1,X0)))))))) | ! [X2] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK17,sK18])],[f264])).
fof(f264,plain,(
  ! [X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | ? [X3] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X3)))) & ? [X4] : ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X3))))))) | ! [X2] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))))),
  inference(rectify,[],[f197])).
fof(f197,plain,(
  ! [X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | ? [X2] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))) & ? [X3] : ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X2))))))) | ! [X4] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X4)))))),
  inference(flattening,[],[f196])).
fof(f196,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | ? [X2] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))) & ? [X3] : ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X2)))))))) | ! [X4] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X4)))))),
  inference(ennf_transformation,[],[f115])).
fof(f115,plain,(
  ! [X0,X1] : ((p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) & ! [X2] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))) => ! [X3] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X2)))))))) => ! [X4] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X4)))))),
  inference(flattening,[],[f114])).
fof(f114,plain,(
  ! [X0] : ! [X1] : ((p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) & ! [X2] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))) => ! [X3] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X2)))))))) => ! [X4] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X4)))))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X6] : ! [X9] : ((p(s(bool,happ(s(fun('type.list.list'(X6),bool),X9),s('type.list.list'(X6),'const.list.NIL_0')))) & ! [X15] : (p(s(bool,happ(s(fun('type.list.list'(X6),bool),X9),s('type.list.list'(X6),X15)))) => ! [X17] : p(s(bool,happ(s(fun('type.list.list'(X6),bool),X9),s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X17),s('type.list.list'(X6),X15)))))))) => ! [X7] : p(s(bool,happ(s(fun('type.list.list'(X6),bool),X9),s('type.list.list'(X6),X7)))))),
  file('thm.list.MONO_EVERY','thm.list.list_INDUCT')).
fof(f10067,plain,(
  ( ! [X391,X393,X392] : (p(s(bool,happ(s(fun('type.list.list'(X391),bool),X392),s('type.list.list'(X391),sK17(X392,X391))))) | s(bool,X393) = s(bool,happ(s(fun('type.list.list'(X391),bool),X392),s('type.list.list'(X391),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,X393)) )),
  inference(subsumption_resolution,[],[f9111,f350])).
fof(f9111,plain,(
  ( ! [X391,X393,X392] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.list.list'(X391),bool),X392),s('type.list.list'(X391),sK17(X392,X391))))) | s(bool,X393) = s(bool,happ(s(fun('type.list.list'(X391),bool),X392),s('type.list.list'(X391),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,X393)) )),
  inference(superposition,[],[f576,f983])).
fof(f2727,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,happ(s(fun('type.list.list'(X5),bool),X6),s('type.list.list'(X5),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X5),bool),X6),s('type.list.list'(X5),sK17('const.list.HD_1'(s('type.list.list'(fun('type.list.list'(X5),bool)),'const.list.CONS_2'(s(fun('type.list.list'(X5),bool),X6),s('type.list.list'(fun('type.list.list'(X5),bool)),X7)))),X5)))))) )),
  inference(forward_demodulation,[],[f2723,f419])).
fof(f2723,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,happ(s(fun('type.list.list'(X5),bool),X6),s('type.list.list'(X5),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X5),bool),'const.list.HD_1'(s('type.list.list'(fun('type.list.list'(X5),bool)),'const.list.CONS_2'(s(fun('type.list.list'(X5),bool),X6),s('type.list.list'(fun('type.list.list'(X5),bool)),X7))))),s('type.list.list'(X5),sK17('const.list.HD_1'(s('type.list.list'(fun('type.list.list'(X5),bool)),'const.list.CONS_2'(s(fun('type.list.list'(X5),bool),X6),s('type.list.list'(fun('type.list.list'(X5),bool)),X7)))),X5)))))) )),
  inference(superposition,[],[f576,f419])).
fof(f576,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK17(X1,X0)))))) )),
  inference(condensation,[],[f410])).
fof(f410,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))) | p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK17(X1,X0))))) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0'))))) )),
  inference(cnf_transformation,[],[f265])).
fof(f3572,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),'const.list.HD_1'(s('type.list.list'(fun('type.list.list'(X0),bool)),X1))),s('type.list.list'(X0),'const.list.NIL_0')))) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),'const.list.HD_1'(s('type.list.list'(fun('type.list.list'(X0),bool)),X1))),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK18('const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(fun('type.list.list'(X0),bool)),X1)),X0)),s('type.list.list'(X0),sK17('const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(fun('type.list.list'(X0),bool)),X1)),X0))))))) | p(s(bool,happ(s(fun('type.list.list'(X0),bool),'const.list.HD_1'(s('type.list.list'(fun('type.list.list'(X0),bool)),X1))),s('type.list.list'(X0),X2))))) )),
  inference(forward_demodulation,[],[f3571,f409])).
fof(f3571,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),'const.list.HD_1'(s('type.list.list'(fun('type.list.list'(X0),bool)),X1))),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK18('const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(fun('type.list.list'(X0),bool)),X1)),X0)),s('type.list.list'(X0),sK17('const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(fun('type.list.list'(X0),bool)),X1)),X0))))))) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),'const.list.HD_1'(s('type.list.list'(fun('type.list.list'(X0),bool)),X1))),s('type.list.list'(X0),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X0),bool),'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(fun('type.list.list'(X0),bool)),X1))),s('type.list.list'(X0),X2))))) )),
  inference(forward_demodulation,[],[f3567,f409])).
fof(f3567,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),'const.list.HD_1'(s('type.list.list'(fun('type.list.list'(X0),bool)),X1))),s('type.list.list'(X0),'const.list.NIL_0')))) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(fun('type.list.list'(X0),bool)),X1))),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK18('const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(fun('type.list.list'(X0),bool)),X1)),X0)),s('type.list.list'(X0),sK17('const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(fun('type.list.list'(X0),bool)),X1)),X0))))))) | p(s(bool,happ(s(fun('type.list.list'(X0),bool),'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(fun('type.list.list'(X0),bool)),X1))),s('type.list.list'(X0),X2))))) )),
  inference(superposition,[],[f411,f409])).
fof(f2725,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),'const.list.HD_1'(s('type.list.list'(fun('type.list.list'(X0),bool)),X1))),s('type.list.list'(X0),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X0),bool),'const.list.HD_1'(s('type.list.list'(fun('type.list.list'(X0),bool)),X1))),s('type.list.list'(X0),sK17('const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(fun('type.list.list'(X0),bool)),X1)),X0)))))) )),
  inference(forward_demodulation,[],[f2721,f409])).
fof(f2721,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),'const.list.HD_1'(s('type.list.list'(fun('type.list.list'(X0),bool)),X1))),s('type.list.list'(X0),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X0),bool),'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(fun('type.list.list'(X0),bool)),X1))),s('type.list.list'(X0),sK17('const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(fun('type.list.list'(X0),bool)),X1)),X0)))))) )),
  inference(superposition,[],[f576,f409])).
fof(f3574,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,happ(s(fun('type.list.list'(X3),bool),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))),s('type.list.list'(fun('type.list.list'(X3),bool)),X5))),s('type.list.list'(X3),'const.list.NIL_0')))) | ~p(s(bool,happ(s(fun('type.list.list'(X3),bool),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))),s('type.list.list'(fun('type.list.list'(X3),bool)),X5))),s('type.list.list'(X3),'const.list.CONS_2'(s(X3,sK18('const.list.EL_2'(s('type.num.num',X4),s('type.list.list'(fun('type.list.list'(X3),bool)),'const.list.TL_1'(s('type.list.list'(fun('type.list.list'(X3),bool)),X5)))),X3)),s('type.list.list'(X3),sK17('const.list.EL_2'(s('type.num.num',X4),s('type.list.list'(fun('type.list.list'(X3),bool)),'const.list.TL_1'(s('type.list.list'(fun('type.list.list'(X3),bool)),X5)))),X3))))))) | p(s(bool,happ(s(fun('type.list.list'(X3),bool),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))),s('type.list.list'(fun('type.list.list'(X3),bool)),X5))),s('type.list.list'(X3),X6))))) )),
  inference(forward_demodulation,[],[f3573,f422])).
fof(f3573,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,happ(s(fun('type.list.list'(X3),bool),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))),s('type.list.list'(fun('type.list.list'(X3),bool)),X5))),s('type.list.list'(X3),'const.list.CONS_2'(s(X3,sK18('const.list.EL_2'(s('type.num.num',X4),s('type.list.list'(fun('type.list.list'(X3),bool)),'const.list.TL_1'(s('type.list.list'(fun('type.list.list'(X3),bool)),X5)))),X3)),s('type.list.list'(X3),sK17('const.list.EL_2'(s('type.num.num',X4),s('type.list.list'(fun('type.list.list'(X3),bool)),'const.list.TL_1'(s('type.list.list'(fun('type.list.list'(X3),bool)),X5)))),X3))))))) | ~p(s(bool,happ(s(fun('type.list.list'(X3),bool),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))),s('type.list.list'(fun('type.list.list'(X3),bool)),X5))),s('type.list.list'(X3),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X3),bool),'const.list.EL_2'(s('type.num.num',X4),s('type.list.list'(fun('type.list.list'(X3),bool)),'const.list.TL_1'(s('type.list.list'(fun('type.list.list'(X3),bool)),X5))))),s('type.list.list'(X3),X6))))) )),
  inference(forward_demodulation,[],[f3568,f422])).
fof(f3568,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,happ(s(fun('type.list.list'(X3),bool),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))),s('type.list.list'(fun('type.list.list'(X3),bool)),X5))),s('type.list.list'(X3),'const.list.NIL_0')))) | ~p(s(bool,happ(s(fun('type.list.list'(X3),bool),'const.list.EL_2'(s('type.num.num',X4),s('type.list.list'(fun('type.list.list'(X3),bool)),'const.list.TL_1'(s('type.list.list'(fun('type.list.list'(X3),bool)),X5))))),s('type.list.list'(X3),'const.list.CONS_2'(s(X3,sK18('const.list.EL_2'(s('type.num.num',X4),s('type.list.list'(fun('type.list.list'(X3),bool)),'const.list.TL_1'(s('type.list.list'(fun('type.list.list'(X3),bool)),X5)))),X3)),s('type.list.list'(X3),sK17('const.list.EL_2'(s('type.num.num',X4),s('type.list.list'(fun('type.list.list'(X3),bool)),'const.list.TL_1'(s('type.list.list'(fun('type.list.list'(X3),bool)),X5)))),X3))))))) | p(s(bool,happ(s(fun('type.list.list'(X3),bool),'const.list.EL_2'(s('type.num.num',X4),s('type.list.list'(fun('type.list.list'(X3),bool)),'const.list.TL_1'(s('type.list.list'(fun('type.list.list'(X3),bool)),X5))))),s('type.list.list'(X3),X6))))) )),
  inference(superposition,[],[f411,f422])).
fof(f2726,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,happ(s(fun('type.list.list'(X2),bool),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))),s('type.list.list'(fun('type.list.list'(X2),bool)),X4))),s('type.list.list'(X2),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X2),bool),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))),s('type.list.list'(fun('type.list.list'(X2),bool)),X4))),s('type.list.list'(X2),sK17('const.list.EL_2'(s('type.num.num',X3),s('type.list.list'(fun('type.list.list'(X2),bool)),'const.list.TL_1'(s('type.list.list'(fun('type.list.list'(X2),bool)),X4)))),X2)))))) )),
  inference(forward_demodulation,[],[f2722,f422])).
fof(f2722,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,happ(s(fun('type.list.list'(X2),bool),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))),s('type.list.list'(fun('type.list.list'(X2),bool)),X4))),s('type.list.list'(X2),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X2),bool),'const.list.EL_2'(s('type.num.num',X3),s('type.list.list'(fun('type.list.list'(X2),bool)),'const.list.TL_1'(s('type.list.list'(fun('type.list.list'(X2),bool)),X4))))),s('type.list.list'(X2),sK17('const.list.EL_2'(s('type.num.num',X3),s('type.list.list'(fun('type.list.list'(X2),bool)),'const.list.TL_1'(s('type.list.list'(fun('type.list.list'(X2),bool)),X4)))),X2)))))) )),
  inference(superposition,[],[f576,f422])).
fof(f10068,plain,(
  ( ! [X395,X396,X394] : (~p(s(bool,happ(s(fun('type.num.num',bool),X394),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16(X394))))))) | p(s(bool,happ(s(fun('type.num.num',bool),X394),s('type.num.num',X395)))) | s(bool,X396) = s(bool,happ(s(fun('type.num.num',bool),X394),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X396)) )),
  inference(subsumption_resolution,[],[f9112,f350])).
fof(f9112,plain,(
  ( ! [X395,X396,X394] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun('type.num.num',bool),X394),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16(X394))))))) | p(s(bool,happ(s(fun('type.num.num',bool),X394),s('type.num.num',X395)))) | s(bool,X396) = s(bool,happ(s(fun('type.num.num',bool),X394),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X396)) )),
  inference(superposition,[],[f365,f983])).
fof(f3268,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,happ(s(fun('type.num.num',bool),X5),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,happ(s(fun('type.num.num',bool),X5),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16('const.list.HD_1'(s('type.list.list'(fun('type.num.num',bool)),'const.list.CONS_2'(s(fun('type.num.num',bool),X5),s('type.list.list'(fun('type.num.num',bool)),X6))))))))))) | p(s(bool,happ(s(fun('type.num.num',bool),X5),s('type.num.num',X7))))) )),
  inference(forward_demodulation,[],[f3267,f419])).
fof(f3267,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,happ(s(fun('type.num.num',bool),X5),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16('const.list.HD_1'(s('type.list.list'(fun('type.num.num',bool)),'const.list.CONS_2'(s(fun('type.num.num',bool),X5),s('type.list.list'(fun('type.num.num',bool)),X6))))))))))) | ~p(s(bool,happ(s(fun('type.num.num',bool),X5),s('type.num.num','const.num.|48|_0')))) | p(s(bool,happ(s(fun('type.num.num',bool),'const.list.HD_1'(s('type.list.list'(fun('type.num.num',bool)),'const.list.CONS_2'(s(fun('type.num.num',bool),X5),s('type.list.list'(fun('type.num.num',bool)),X6))))),s('type.num.num',X7))))) )),
  inference(forward_demodulation,[],[f3261,f419])).
fof(f3261,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,happ(s(fun('type.num.num',bool),X5),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,happ(s(fun('type.num.num',bool),'const.list.HD_1'(s('type.list.list'(fun('type.num.num',bool)),'const.list.CONS_2'(s(fun('type.num.num',bool),X5),s('type.list.list'(fun('type.num.num',bool)),X6))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16('const.list.HD_1'(s('type.list.list'(fun('type.num.num',bool)),'const.list.CONS_2'(s(fun('type.num.num',bool),X5),s('type.list.list'(fun('type.num.num',bool)),X6))))))))))) | p(s(bool,happ(s(fun('type.num.num',bool),'const.list.HD_1'(s('type.list.list'(fun('type.num.num',bool)),'const.list.CONS_2'(s(fun('type.num.num',bool),X5),s('type.list.list'(fun('type.num.num',bool)),X6))))),s('type.num.num',X7))))) )),
  inference(superposition,[],[f365,f419])).
fof(f365,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16(X0))))))) | p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f242])).
fof(f242,plain,(
  ! [X0] : (~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.|48|_0')))) | (p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',sK16(X0))))) & ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16(X0)))))))) | ! [X1] : p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK16])],[f241])).
fof(f241,plain,(
  ! [X0] : (~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.|48|_0')))) | ? [X2] : (p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2)))) & ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))))))) | ! [X1] : p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X1)))))),
  inference(rectify,[],[f186])).
fof(f186,plain,(
  ! [X0] : (~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.|48|_0')))) | ? [X1] : (p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X1)))) & ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))))) | ! [X2] : p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2)))))),
  inference(flattening,[],[f185])).
fof(f185,plain,(
  ! [X0] : ((~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.|48|_0')))) | ? [X1] : (p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X1)))) & ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))))) | ! [X2] : p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2)))))),
  inference(ennf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0] : ((p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.|48|_0')))) & ! [X1] : (p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X1)))) => p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))))) => ! [X2] : p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2)))))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X9] : ((p(s(bool,happ(s(fun('type.num.num',bool),X9),s('type.num.num','const.num.|48|_0')))) & ! [X11] : (p(s(bool,happ(s(fun('type.num.num',bool),X9),s('type.num.num',X11)))) => p(s(bool,happ(s(fun('type.num.num',bool),X9),s('type.num.num','const.num.SUC_1'(s('type.num.num',X11)))))))) => ! [X11] : p(s(bool,happ(s(fun('type.num.num',bool),X9),s('type.num.num',X11)))))),
  file('thm.list.MONO_EVERY','thm.num.INDUCTION')).
fof(f10069,plain,(
  ( ! [X397,X398] : (p(s(bool,happ(s(fun('type.num.num',bool),X397),s('type.num.num',sK16(X397))))) | s(bool,X398) = s(bool,happ(s(fun('type.num.num',bool),X397),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X398)) )),
  inference(subsumption_resolution,[],[f9113,f350])).
fof(f9113,plain,(
  ( ! [X397,X398] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.num.num',bool),X397),s('type.num.num',sK16(X397))))) | s(bool,X398) = s(bool,happ(s(fun('type.num.num',bool),X397),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X398)) )),
  inference(superposition,[],[f584,f983])).
fof(f2127,plain,(
  ( ! [X2,X1] : (~p(s(bool,happ(s(fun('type.num.num',bool),X1),s('type.num.num','const.num.|48|_0')))) | p(s(bool,happ(s(fun('type.num.num',bool),X1),s('type.num.num',sK16('const.list.HD_1'(s('type.list.list'(fun('type.num.num',bool)),'const.list.CONS_2'(s(fun('type.num.num',bool),X1),s('type.list.list'(fun('type.num.num',bool)),X2)))))))))) )),
  inference(forward_demodulation,[],[f2124,f419])).
fof(f2124,plain,(
  ( ! [X2,X1] : (~p(s(bool,happ(s(fun('type.num.num',bool),X1),s('type.num.num','const.num.|48|_0')))) | p(s(bool,happ(s(fun('type.num.num',bool),'const.list.HD_1'(s('type.list.list'(fun('type.num.num',bool)),'const.list.CONS_2'(s(fun('type.num.num',bool),X1),s('type.list.list'(fun('type.num.num',bool)),X2))))),s('type.num.num',sK16('const.list.HD_1'(s('type.list.list'(fun('type.num.num',bool)),'const.list.CONS_2'(s(fun('type.num.num',bool),X1),s('type.list.list'(fun('type.num.num',bool)),X2)))))))))) )),
  inference(superposition,[],[f584,f419])).
fof(f584,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.|48|_0')))) | p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',sK16(X0)))))) )),
  inference(condensation,[],[f364])).
fof(f364,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X1)))) | p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',sK16(X0))))) | ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.num.|48|_0'))))) )),
  inference(cnf_transformation,[],[f242])).
fof(f3264,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun('type.num.num',bool),'const.list.HD_1'(s('type.list.list'(fun('type.num.num',bool)),X0))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,happ(s(fun('type.num.num',bool),'const.list.HD_1'(s('type.list.list'(fun('type.num.num',bool)),X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16('const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(fun('type.num.num',bool)),X0))))))))) | p(s(bool,happ(s(fun('type.num.num',bool),'const.list.HD_1'(s('type.list.list'(fun('type.num.num',bool)),X0))),s('type.num.num',X1))))) )),
  inference(forward_demodulation,[],[f3263,f409])).
fof(f3263,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun('type.num.num',bool),'const.list.HD_1'(s('type.list.list'(fun('type.num.num',bool)),X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16('const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(fun('type.num.num',bool)),X0))))))))) | ~p(s(bool,happ(s(fun('type.num.num',bool),'const.list.HD_1'(s('type.list.list'(fun('type.num.num',bool)),X0))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,happ(s(fun('type.num.num',bool),'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(fun('type.num.num',bool)),X0))),s('type.num.num',X1))))) )),
  inference(forward_demodulation,[],[f3259,f409])).
fof(f3259,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun('type.num.num',bool),'const.list.HD_1'(s('type.list.list'(fun('type.num.num',bool)),X0))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,happ(s(fun('type.num.num',bool),'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(fun('type.num.num',bool)),X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16('const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(fun('type.num.num',bool)),X0))))))))) | p(s(bool,happ(s(fun('type.num.num',bool),'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(fun('type.num.num',bool)),X0))),s('type.num.num',X1))))) )),
  inference(superposition,[],[f365,f409])).
fof(f2126,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun('type.num.num',bool),'const.list.HD_1'(s('type.list.list'(fun('type.num.num',bool)),X0))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,happ(s(fun('type.num.num',bool),'const.list.HD_1'(s('type.list.list'(fun('type.num.num',bool)),X0))),s('type.num.num',sK16('const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(fun('type.num.num',bool)),X0)))))))) )),
  inference(forward_demodulation,[],[f2123,f409])).
fof(f2123,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun('type.num.num',bool),'const.list.HD_1'(s('type.list.list'(fun('type.num.num',bool)),X0))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,happ(s(fun('type.num.num',bool),'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(fun('type.num.num',bool)),X0))),s('type.num.num',sK16('const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(fun('type.num.num',bool)),X0)))))))) )),
  inference(superposition,[],[f584,f409])).
fof(f3266,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,happ(s(fun('type.num.num',bool),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s('type.list.list'(fun('type.num.num',bool)),X3))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,happ(s(fun('type.num.num',bool),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s('type.list.list'(fun('type.num.num',bool)),X3))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16('const.list.EL_2'(s('type.num.num',X2),s('type.list.list'(fun('type.num.num',bool)),'const.list.TL_1'(s('type.list.list'(fun('type.num.num',bool)),X3))))))))))) | p(s(bool,happ(s(fun('type.num.num',bool),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s('type.list.list'(fun('type.num.num',bool)),X3))),s('type.num.num',X4))))) )),
  inference(forward_demodulation,[],[f3265,f422])).
fof(f3265,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,happ(s(fun('type.num.num',bool),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s('type.list.list'(fun('type.num.num',bool)),X3))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16('const.list.EL_2'(s('type.num.num',X2),s('type.list.list'(fun('type.num.num',bool)),'const.list.TL_1'(s('type.list.list'(fun('type.num.num',bool)),X3))))))))))) | ~p(s(bool,happ(s(fun('type.num.num',bool),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s('type.list.list'(fun('type.num.num',bool)),X3))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,happ(s(fun('type.num.num',bool),'const.list.EL_2'(s('type.num.num',X2),s('type.list.list'(fun('type.num.num',bool)),'const.list.TL_1'(s('type.list.list'(fun('type.num.num',bool)),X3))))),s('type.num.num',X4))))) )),
  inference(forward_demodulation,[],[f3260,f422])).
fof(f3260,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,happ(s(fun('type.num.num',bool),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s('type.list.list'(fun('type.num.num',bool)),X3))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,happ(s(fun('type.num.num',bool),'const.list.EL_2'(s('type.num.num',X2),s('type.list.list'(fun('type.num.num',bool)),'const.list.TL_1'(s('type.list.list'(fun('type.num.num',bool)),X3))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16('const.list.EL_2'(s('type.num.num',X2),s('type.list.list'(fun('type.num.num',bool)),'const.list.TL_1'(s('type.list.list'(fun('type.num.num',bool)),X3))))))))))) | p(s(bool,happ(s(fun('type.num.num',bool),'const.list.EL_2'(s('type.num.num',X2),s('type.list.list'(fun('type.num.num',bool)),'const.list.TL_1'(s('type.list.list'(fun('type.num.num',bool)),X3))))),s('type.num.num',X4))))) )),
  inference(superposition,[],[f365,f422])).
fof(f2660,plain,(
  ( ! [X151,X150] : (~p(s(bool,happ(s(fun('type.num.num',bool),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X150))),s('type.list.list'(fun('type.num.num',bool)),X151))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,happ(s(fun('type.num.num',bool),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X150))),s('type.list.list'(fun('type.num.num',bool)),X151))),s('type.num.num',sK16('const.list.EL_2'(s('type.num.num',X150),s('type.list.list'(fun('type.num.num',bool)),'const.list.TL_1'(s('type.list.list'(fun('type.num.num',bool)),X151)))))))))) )),
  inference(forward_demodulation,[],[f2617,f422])).
fof(f2617,plain,(
  ( ! [X151,X150] : (~p(s(bool,happ(s(fun('type.num.num',bool),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X150))),s('type.list.list'(fun('type.num.num',bool)),X151))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,happ(s(fun('type.num.num',bool),'const.list.EL_2'(s('type.num.num',X150),s('type.list.list'(fun('type.num.num',bool)),'const.list.TL_1'(s('type.list.list'(fun('type.num.num',bool)),X151))))),s('type.num.num',sK16('const.list.EL_2'(s('type.num.num',X150),s('type.list.list'(fun('type.num.num',bool)),'const.list.TL_1'(s('type.list.list'(fun('type.num.num',bool)),X151)))))))))) )),
  inference(superposition,[],[f584,f422])).
fof(f10070,plain,(
  ( ! [X399,X400] : (p(s(bool,happ(s(fun(sK12,bool),sK14),s(sK12,X399)))) | s(bool,X400) = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X399))) | s(bool,'T_0') = s(bool,X400)) )),
  inference(subsumption_resolution,[],[f9114,f350])).
fof(f9114,plain,(
  ( ! [X399,X400] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK12,bool),sK14),s(sK12,X399)))) | s(bool,X400) = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X399))) | s(bool,'T_0') = s(bool,X400)) )),
  inference(superposition,[],[f347,f983])).
fof(f347,plain,(
  ( ! [X4] : (~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X4)))) | p(s(bool,happ(s(fun(sK12,bool),sK14),s(sK12,X4))))) )),
  inference(cnf_transformation,[],[f238])).
fof(f238,plain,(
  ! [X4] : (~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X4)))) | p(s(bool,happ(s(fun(sK12,bool),sK14),s(sK12,X4))))) & p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) & ~p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13,sK14,sK15])],[f181])).
fof(f181,plain,(
  ? [X0,X1,X2,X3] : (! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) & p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X3),s('type.list.list'(X0),X1)))) & ~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X2),s('type.list.list'(X0),X1)))))),
  inference(flattening,[],[f180])).
fof(f180,plain,(
  ? [X0,X1,X2,X3] : (! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) & (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X3),s('type.list.list'(X0),X1)))) & ~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X2),s('type.list.list'(X0),X1))))))),
  inference(ennf_transformation,[],[f72])).
fof(f72,plain,(
  ~! [X0,X1,X2,X3] : (! [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) => p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) => (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X3),s('type.list.list'(X0),X1)))) => p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X2),s('type.list.list'(X0),X1))))))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ~! [X0] : ! [X1,X2,X3] : (! [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) => p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) => (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X3),s('type.list.list'(X0),X1)))) => p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X2),s('type.list.list'(X0),X1))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : (! [X10] : (p(s(bool,happ(s(fun(X6,bool),X9),s(X6,X10)))) => p(s(bool,happ(s(fun(X6,bool),X8),s(X6,X10))))) => (p(s(bool,'const.list.EVERY_2'(s(fun(X6,bool),X9),s('type.list.list'(X6),X7)))) => p(s(bool,'const.list.EVERY_2'(s(fun(X6,bool),X8),s('type.list.list'(X6),X7))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : (! [X10] : (p(s(bool,happ(s(fun(X6,bool),X9),s(X6,X10)))) => p(s(bool,happ(s(fun(X6,bool),X8),s(X6,X10))))) => (p(s(bool,'const.list.EVERY_2'(s(fun(X6,bool),X9),s('type.list.list'(X6),X7)))) => p(s(bool,'const.list.EVERY_2'(s(fun(X6,bool),X8),s('type.list.list'(X6),X7))))))),
  file('thm.list.MONO_EVERY','thm.list.MONO_EVERY')).
fof(f3499,plain,(
  ( ! [X0] : (p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X0),s('type.list.list'(sK12),sK13)))))) | ~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X0))))) )),
  inference(subsumption_resolution,[],[f3485,f350])).
fof(f3485,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X0),s('type.list.list'(sK12),sK13)))))) | ~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X0))))) )),
  inference(superposition,[],[f445,f647])).
fof(f3302,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X0))))) )),
  inference(subsumption_resolution,[],[f3291,f350])).
fof(f3291,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X0))))) )),
  inference(superposition,[],[f428,f647])).
fof(f3542,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK12),sK13))))))) )),
  inference(subsumption_resolution,[],[f3526,f350])).
fof(f3526,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK12),sK13))))))) )),
  inference(superposition,[],[f425,f647])).
fof(f426,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK21(X2,X1,X0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1)))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X2),s('type.list.list'(X0),X1))))) )),
  inference(cnf_transformation,[],[f275])).
fof(f429,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK22(X2,X1,X0)),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))))) )),
  inference(cnf_transformation,[],[f278])).
fof(f444,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3)))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X3))))) )),
  inference(cnf_transformation,[],[f292])).
fof(f2701,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK21(X3,'const.list.TL_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))),X0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2)))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X3),s('type.list.list'(X0),X2))))) )),
  inference(forward_demodulation,[],[f2694,f420])).
fof(f2694,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK21(X3,'const.list.TL_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))),X0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2)))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X3),s('type.list.list'(X0),'const.list.TL_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2))))))))) )),
  inference(superposition,[],[f426,f420])).
fof(f2704,plain,(
  ( ! [X14,X12,X13,X11] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK21(X14,'const.list.HD_1'(s('type.list.list'('type.list.list'(X11)),'const.list.CONS_2'(s('type.list.list'(X11),X12),s('type.list.list'('type.list.list'(X11)),X13)))),X11)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X11),X12)))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X11,bool),X14),s('type.list.list'(X11),X12))))) )),
  inference(forward_demodulation,[],[f2697,f419])).
fof(f2697,plain,(
  ( ! [X14,X12,X13,X11] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK21(X14,'const.list.HD_1'(s('type.list.list'('type.list.list'(X11)),'const.list.CONS_2'(s('type.list.list'(X11),X12),s('type.list.list'('type.list.list'(X11)),X13)))),X11)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X11),X12)))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X11,bool),X14),s('type.list.list'(X11),'const.list.HD_1'(s('type.list.list'('type.list.list'(X11)),'const.list.CONS_2'(s('type.list.list'(X11),X12),s('type.list.list'('type.list.list'(X11)),X13))))))))) )),
  inference(superposition,[],[f426,f419])).
fof(f2708,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.list.EVERY_2'(s(fun(X2,bool),X0),s('type.list.list'(X2),X1)))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK21(X0,X1,X2)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X2),X1)))))) )),
  inference(subsumption_resolution,[],[f2700,f585])).
fof(f2700,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | p(s(bool,'const.list.EVERY_2'(s(fun(X2,bool),X0),s('type.list.list'(X2),X1)))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK21(X0,X1,X2)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X2),X1)))))) )),
  inference(superposition,[],[f426,f360])).
fof(f2898,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK22('const.list.TL_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))),X3,X0)),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X3),s('type.list.list'(X0),X2))))) )),
  inference(forward_demodulation,[],[f2893,f420])).
fof(f2893,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK22('const.list.TL_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))),X3,X0)),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X3),s('type.list.list'(X0),'const.list.TL_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2))))))))) )),
  inference(superposition,[],[f429,f420])).
fof(f2901,plain,(
  ( ! [X14,X12,X13,X11] : (p(s(bool,'const.bool.IN_2'(s(X11,sK22('const.list.HD_1'(s('type.list.list'('type.list.list'(X11)),'const.list.CONS_2'(s('type.list.list'(X11),X12),s('type.list.list'('type.list.list'(X11)),X13)))),X14,X11)),s(fun(X11,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X11),X12)))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X11,bool),X14),s('type.list.list'(X11),X12))))) )),
  inference(forward_demodulation,[],[f2896,f419])).
fof(f2896,plain,(
  ( ! [X14,X12,X13,X11] : (p(s(bool,'const.bool.IN_2'(s(X11,sK22('const.list.HD_1'(s('type.list.list'('type.list.list'(X11)),'const.list.CONS_2'(s('type.list.list'(X11),X12),s('type.list.list'('type.list.list'(X11)),X13)))),X14,X11)),s(fun(X11,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X11),X12)))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X11,bool),X14),s('type.list.list'(X11),'const.list.HD_1'(s('type.list.list'('type.list.list'(X11)),'const.list.CONS_2'(s('type.list.list'(X11),X12),s('type.list.list'('type.list.list'(X11)),X13))))))))) )),
  inference(superposition,[],[f429,f419])).
fof(f2902,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X2),s('type.list.list'(X0),X1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,sK22(X1,X2,X0)),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1)))))) )),
  inference(subsumption_resolution,[],[f2897,f585])).
fof(f2897,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X2),s('type.list.list'(X0),X1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,sK22(X1,X2,X0)),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1)))))) )),
  inference(superposition,[],[f429,f360])).
fof(f10071,plain,(
  ( ! [X408,X410,X407,X409,X411] : (p(s(bool,'const.list.EVERY_2'(s(fun(X407,bool),X408),s('type.list.list'(X407),X410)))) | s(bool,X411) = s(bool,'const.list.EVERY_2'(s(fun(X407,bool),X408),s('type.list.list'(X407),'const.list.CONS_2'(s(X407,X409),s('type.list.list'(X407),X410))))) | s(bool,'T_0') = s(bool,X411)) )),
  inference(subsumption_resolution,[],[f9118,f350])).
fof(f9118,plain,(
  ( ! [X408,X410,X407,X409,X411] : (~p(s(bool,'T_0')) | p(s(bool,'const.list.EVERY_2'(s(fun(X407,bool),X408),s('type.list.list'(X407),X410)))) | s(bool,X411) = s(bool,'const.list.EVERY_2'(s(fun(X407,bool),X408),s('type.list.list'(X407),'const.list.CONS_2'(s(X407,X409),s('type.list.list'(X407),X410))))) | s(bool,'T_0') = s(bool,X411)) )),
  inference(superposition,[],[f444,f983])).
fof(f2900,plain,(
  ( ! [X10,X8,X7,X9] : (p(s(bool,'const.bool.IN_2'(s(X7,sK22('const.list.EL_2'(s('type.num.num',X8),s('type.list.list'('type.list.list'(X7)),'const.list.TL_1'(s('type.list.list'('type.list.list'(X7)),X9)))),X10,X7)),s(fun(X7,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X7),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.list.list'('type.list.list'(X7)),X9)))))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X7,bool),X10),s('type.list.list'(X7),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.list.list'('type.list.list'(X7)),X9))))))) )),
  inference(forward_demodulation,[],[f2895,f422])).
fof(f2895,plain,(
  ( ! [X10,X8,X7,X9] : (p(s(bool,'const.bool.IN_2'(s(X7,sK22('const.list.EL_2'(s('type.num.num',X8),s('type.list.list'('type.list.list'(X7)),'const.list.TL_1'(s('type.list.list'('type.list.list'(X7)),X9)))),X10,X7)),s(fun(X7,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X7),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.list.list'('type.list.list'(X7)),X9)))))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X7,bool),X10),s('type.list.list'(X7),'const.list.EL_2'(s('type.num.num',X8),s('type.list.list'('type.list.list'(X7)),'const.list.TL_1'(s('type.list.list'('type.list.list'(X7)),X9))))))))) )),
  inference(superposition,[],[f429,f422])).
fof(f2703,plain,(
  ( ! [X10,X8,X7,X9] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK21(X10,'const.list.EL_2'(s('type.num.num',X8),s('type.list.list'('type.list.list'(X7)),'const.list.TL_1'(s('type.list.list'('type.list.list'(X7)),X9)))),X7)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X7),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.list.list'('type.list.list'(X7)),X9)))))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X7,bool),X10),s('type.list.list'(X7),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.list.list'('type.list.list'(X7)),X9))))))) )),
  inference(forward_demodulation,[],[f2696,f422])).
fof(f2696,plain,(
  ( ! [X10,X8,X7,X9] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK21(X10,'const.list.EL_2'(s('type.num.num',X8),s('type.list.list'('type.list.list'(X7)),'const.list.TL_1'(s('type.list.list'('type.list.list'(X7)),X9)))),X7)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X7),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.list.list'('type.list.list'(X7)),X9)))))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X7,bool),X10),s('type.list.list'(X7),'const.list.EL_2'(s('type.num.num',X8),s('type.list.list'('type.list.list'(X7)),'const.list.TL_1'(s('type.list.list'('type.list.list'(X7)),X9))))))))) )),
  inference(superposition,[],[f426,f422])).
fof(f2699,plain,(
  ( ! [X4,X2,X5,X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK21(X5,'const.list.CONS_2'(s(X2,X3),s('type.list.list'(X2),X4)),X2)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X2),X4)))))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X2,bool),X5),s('type.list.list'(X2),'const.list.CONS_2'(s(X2,X3),s('type.list.list'(X2),X4))))))) )),
  inference(superposition,[],[f426,f421])).
fof(f2899,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,'const.bool.IN_2'(s(X4,sK22('const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'(X4)),X5)),X6,X4)),s(fun(X4,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X4),'const.list.HD_1'(s('type.list.list'('type.list.list'(X4)),X5)))))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X4,bool),X6),s('type.list.list'(X4),'const.list.HD_1'(s('type.list.list'('type.list.list'(X4)),X5))))))) )),
  inference(forward_demodulation,[],[f2894,f409])).
fof(f2894,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,'const.bool.IN_2'(s(X4,sK22('const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'(X4)),X5)),X6,X4)),s(fun(X4,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X4),'const.list.HD_1'(s('type.list.list'('type.list.list'(X4)),X5)))))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X4,bool),X6),s('type.list.list'(X4),'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'(X4)),X5))))))) )),
  inference(superposition,[],[f429,f409])).
fof(f2702,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK21(X6,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'(X4)),X5)),X4)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X4),'const.list.HD_1'(s('type.list.list'('type.list.list'(X4)),X5)))))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X4,bool),X6),s('type.list.list'(X4),'const.list.HD_1'(s('type.list.list'('type.list.list'(X4)),X5))))))) )),
  inference(forward_demodulation,[],[f2695,f409])).
fof(f2695,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK21(X6,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'(X4)),X5)),X4)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X4),'const.list.HD_1'(s('type.list.list'('type.list.list'(X4)),X5)))))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X4,bool),X6),s('type.list.list'(X4),'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'(X4)),X5))))))) )),
  inference(superposition,[],[f426,f409])).
fof(f8637,plain,(
  ( ! [X685,X683,X684] : (p(s(bool,'const.bool.IN_2'(s(bool,X685),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X683)))))) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X684),s('type.list.list'(bool),X683)))) | s(bool,'T_0') = s(bool,X685) | s(bool,'T_0') = s(bool,sK22(X683,X684,bool))) )),
  inference(superposition,[],[f429,f983])).
fof(f9202,plain,(
  ( ! [X675,X676,X674] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X674)))))) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X675),s('type.list.list'(bool),X674)))) | s(bool,X676) = s(bool,sK22(X674,X675,bool)) | s(bool,'T_0') = s(bool,X676)) )),
  inference(superposition,[],[f429,f983])).
fof(f449,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0,X1] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X18,X11] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X18),s('type.num.num',X11))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X18))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X11)))))),
  file('thm.list.MONO_EVERY','thm.prim_rec.LESS_MONO_EQ')).
fof(f2076,plain,(
  ( ! [X21,X19,X20] : (sP0(X21,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X19))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X20))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X19),s('type.num.num',X20))))) )),
  inference(superposition,[],[f464,f449])).
fof(f2077,plain,(
  ( ! [X24,X23,X22] : (sP0('const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X22))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X23)))),X24) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X22),s('type.num.num',X23))))) )),
  inference(superposition,[],[f465,f449])).
fof(f2078,plain,(
  ( ! [X26,X27,X25] : (sP1('const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X25))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X26)))),X27) | p(s(bool,X27)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X25),s('type.num.num',X26))))) )),
  inference(superposition,[],[f480,f449])).
fof(f2079,plain,(
  ( ! [X30,X28,X29] : (sP1(X30,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X28))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X29))))) | p(s(bool,X30)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X28),s('type.num.num',X29))))) )),
  inference(superposition,[],[f483,f449])).
fof(f2080,plain,(
  ( ! [X33,X31,X34,X32] : (~sP5(X33,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X31))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X32)))),X34) | ~p(s(bool,X33)) | p(s(bool,X34)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X31),s('type.num.num',X32))))) )),
  inference(superposition,[],[f501,f449])).
fof(f2081,plain,(
  ( ! [X37,X35,X38,X36] : (~sP4(X37,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X35))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X36)))),X38) | ~p(s(bool,X37)) | p(s(bool,X38)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X35),s('type.num.num',X36))))) )),
  inference(superposition,[],[f505,f449])).
fof(f2082,plain,(
  ( ! [X39,X41,X42,X40] : (~sP3(X41,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X39))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X40)))),X42) | ~p(s(bool,X41)) | p(s(bool,X42)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X39),s('type.num.num',X40))))) )),
  inference(superposition,[],[f509,f449])).
fof(f2083,plain,(
  ( ! [X45,X43,X46,X44] : (~sP6(X45,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X43))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X44)))),X46) | ~p(s(bool,X45)) | p(s(bool,X46)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X43),s('type.num.num',X44))))) )),
  inference(superposition,[],[f528,f449])).
fof(f2084,plain,(
  ( ! [X47,X50,X48,X49] : (~sP7(X50,X49,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X47))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X48))))) | ~p(s(bool,X49)) | p(s(bool,X50)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X47),s('type.num.num',X48))))) )),
  inference(superposition,[],[f537,f449])).
fof(f2085,plain,(
  ( ! [X54,X52,X53,X51] : (~sP8('const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X51))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X52)))),X53,X54) | p(s(bool,X53)) | p(s(bool,X54)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X51),s('type.num.num',X52))))) )),
  inference(superposition,[],[f545,f449])).
fof(f2086,plain,(
  ( ! [X57,X58,X56,X55] : (sP9('const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X55))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X56)))),X58,X57) | ~p(s(bool,X57)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X55),s('type.num.num',X56))))) )),
  inference(superposition,[],[f550,f449])).
fof(f2087,plain,(
  ( ! [X61,X59,X62,X60] : (sP9('const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X59))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X60)))),X61,X62) | ~p(s(bool,X61)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X59),s('type.num.num',X60))))) )),
  inference(superposition,[],[f551,f449])).
fof(f2088,plain,(
  ( ! [X66,X64,X65,X63] : (~sP9('const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X63))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X64)))),X66,X65) | p(s(bool,X65)) | p(s(bool,X66)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X63),s('type.num.num',X64))))) )),
  inference(superposition,[],[f552,f449])).
fof(f2089,plain,(
  ( ! [X70,X68,X69,X67] : (~sP10(X70,X69,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X67))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X68))))) | p(s(bool,X69)) | ~p(s(bool,X70)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X67),s('type.num.num',X68))))) )),
  inference(superposition,[],[f562,f449])).
fof(f2090,plain,(
  ( ! [X74,X72,X71,X73] : (sP11('const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X71))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X72)))),X74,X73) | p(s(bool,X73)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X71),s('type.num.num',X72))))) )),
  inference(superposition,[],[f567,f449])).
fof(f2091,plain,(
  ( ! [X78,X76,X77,X75] : (sP11('const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X75))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X76)))),X77,X78) | ~p(s(bool,X77)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X75),s('type.num.num',X76))))) )),
  inference(superposition,[],[f568,f449])).
fof(f2092,plain,(
  ( ! [X80,X81,X79,X82] : (~sP11('const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X79))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X80)))),X82,X81) | ~p(s(bool,X81)) | p(s(bool,X82)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X79),s('type.num.num',X80))))) )),
  inference(superposition,[],[f569,f449])).
fof(f2093,plain,(
  ( ! [X85,X83,X86,X84] : (sP11('const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X83))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X84)))),X85,X86) | ~sP10('const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X83))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X84)))),X85,X86) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X83),s('type.num.num',X84))))) )),
  inference(superposition,[],[f572,f449])).
fof(f2094,plain,(
  ( ! [X90,X88,X87,X89] : (sP9('const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X87))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X88)))),X89,X90) | ~sP8('const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X87))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X88)))),X89,X90) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X87),s('type.num.num',X88))))) )),
  inference(superposition,[],[f573,f449])).
fof(f5439,plain,(
  ( ! [X10,X11,X9] : (~sP5('F_0','const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num',X10))),s('type.list.list'(bool),X11)))),'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X10)))))) )),
  inference(superposition,[],[f1662,f449])).
fof(f5485,plain,(
  ( ! [X10,X11,X9] : (sP5('T_0','const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num',X10))),s('type.list.list'(bool),X11)))),'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X10)))))) )),
  inference(superposition,[],[f1663,f449])).
fof(f5534,plain,(
  ( ! [X10,X11,X9] : (~sP5('F_0','const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X10)))),'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num',X10))),s('type.list.list'(bool),X11)))))) )),
  inference(superposition,[],[f1664,f449])).
fof(f5588,plain,(
  ( ! [X10,X11,X9] : (sP5('T_0','const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X10)))),'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num',X10))),s('type.list.list'(bool),X11)))))) )),
  inference(superposition,[],[f1665,f449])).
fof(f8342,plain,(
  ( ! [X30,X31,X29] : (s(bool,X31) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X29),s('type.num.num',X30))) | s(bool,'T_0') = s(bool,X31) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X29))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X30)))))) )),
  inference(superposition,[],[f983,f449])).
fof(f8384,plain,(
  ( ! [X23,X21,X22] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X21),s('type.num.num',X22))) | s(bool,X23) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X21))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X22))))) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f983,f449])).
fof(f2595,plain,(
  ( ! [X68,X69,X67] : (sP0(X69,'const.list.EL_2'(s('type.num.num',X67),s('type.list.list'(bool),'const.list.TL_1'(s('type.list.list'(bool),X68))))) | ~p(s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X67))),s('type.list.list'(bool),X68))))) )),
  inference(superposition,[],[f464,f422])).
fof(f2596,plain,(
  ( ! [X70,X72,X71] : (sP0('const.list.EL_2'(s('type.num.num',X70),s('type.list.list'(bool),'const.list.TL_1'(s('type.list.list'(bool),X71)))),X72) | ~p(s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X70))),s('type.list.list'(bool),X71))))) )),
  inference(superposition,[],[f465,f422])).
fof(f2597,plain,(
  ( ! [X74,X75,X73] : (sP1('const.list.EL_2'(s('type.num.num',X73),s('type.list.list'(bool),'const.list.TL_1'(s('type.list.list'(bool),X74)))),X75) | p(s(bool,X75)) | ~p(s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X73))),s('type.list.list'(bool),X74))))) )),
  inference(superposition,[],[f480,f422])).
fof(f2598,plain,(
  ( ! [X78,X76,X77] : (sP1(X78,'const.list.EL_2'(s('type.num.num',X76),s('type.list.list'(bool),'const.list.TL_1'(s('type.list.list'(bool),X77))))) | p(s(bool,X78)) | ~p(s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X76))),s('type.list.list'(bool),X77))))) )),
  inference(superposition,[],[f483,f422])).
fof(f2599,plain,(
  ( ! [X80,X81,X79,X82] : (~sP5(X81,'const.list.EL_2'(s('type.num.num',X79),s('type.list.list'(bool),'const.list.TL_1'(s('type.list.list'(bool),X80)))),X82) | ~p(s(bool,X81)) | p(s(bool,X82)) | ~p(s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X79))),s('type.list.list'(bool),X80))))) )),
  inference(superposition,[],[f501,f422])).
fof(f2600,plain,(
  ( ! [X85,X83,X86,X84] : (~sP4(X85,'const.list.EL_2'(s('type.num.num',X83),s('type.list.list'(bool),'const.list.TL_1'(s('type.list.list'(bool),X84)))),X86) | ~p(s(bool,X85)) | p(s(bool,X86)) | ~p(s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X83))),s('type.list.list'(bool),X84))))) )),
  inference(superposition,[],[f505,f422])).
fof(f2601,plain,(
  ( ! [X90,X88,X87,X89] : (~sP3(X89,'const.list.EL_2'(s('type.num.num',X87),s('type.list.list'(bool),'const.list.TL_1'(s('type.list.list'(bool),X88)))),X90) | ~p(s(bool,X89)) | p(s(bool,X90)) | ~p(s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X87))),s('type.list.list'(bool),X88))))) )),
  inference(superposition,[],[f509,f422])).
fof(f2602,plain,(
  ( ! [X94,X92,X93,X91] : (~sP6(X93,'const.list.EL_2'(s('type.num.num',X91),s('type.list.list'(bool),'const.list.TL_1'(s('type.list.list'(bool),X92)))),X94) | ~p(s(bool,X93)) | p(s(bool,X94)) | ~p(s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X91))),s('type.list.list'(bool),X92))))) )),
  inference(superposition,[],[f528,f422])).
fof(f2603,plain,(
  ( ! [X97,X95,X98,X96] : (~sP7(X98,X97,'const.list.EL_2'(s('type.num.num',X95),s('type.list.list'(bool),'const.list.TL_1'(s('type.list.list'(bool),X96))))) | ~p(s(bool,X97)) | p(s(bool,X98)) | ~p(s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X95))),s('type.list.list'(bool),X96))))) )),
  inference(superposition,[],[f537,f422])).
fof(f2604,plain,(
  ( ! [X101,X99,X102,X100] : (~sP8('const.list.EL_2'(s('type.num.num',X99),s('type.list.list'(bool),'const.list.TL_1'(s('type.list.list'(bool),X100)))),X101,X102) | p(s(bool,X101)) | p(s(bool,X102)) | ~p(s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X99))),s('type.list.list'(bool),X100))))) )),
  inference(superposition,[],[f545,f422])).
fof(f2605,plain,(
  ( ! [X103,X105,X106,X104] : (sP9('const.list.EL_2'(s('type.num.num',X103),s('type.list.list'(bool),'const.list.TL_1'(s('type.list.list'(bool),X104)))),X106,X105) | ~p(s(bool,X105)) | ~p(s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X103))),s('type.list.list'(bool),X104))))) )),
  inference(superposition,[],[f550,f422])).
fof(f2606,plain,(
  ( ! [X109,X107,X110,X108] : (sP9('const.list.EL_2'(s('type.num.num',X107),s('type.list.list'(bool),'const.list.TL_1'(s('type.list.list'(bool),X108)))),X109,X110) | ~p(s(bool,X109)) | ~p(s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X107))),s('type.list.list'(bool),X108))))) )),
  inference(superposition,[],[f551,f422])).
fof(f2607,plain,(
  ( ! [X111,X114,X112,X113] : (~sP9('const.list.EL_2'(s('type.num.num',X111),s('type.list.list'(bool),'const.list.TL_1'(s('type.list.list'(bool),X112)))),X114,X113) | p(s(bool,X113)) | p(s(bool,X114)) | ~p(s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X111))),s('type.list.list'(bool),X112))))) )),
  inference(superposition,[],[f552,f422])).
fof(f2608,plain,(
  ( ! [X118,X116,X117,X115] : (~sP10(X118,X117,'const.list.EL_2'(s('type.num.num',X115),s('type.list.list'(bool),'const.list.TL_1'(s('type.list.list'(bool),X116))))) | p(s(bool,X117)) | ~p(s(bool,X118)) | ~p(s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X115))),s('type.list.list'(bool),X116))))) )),
  inference(superposition,[],[f562,f422])).
fof(f2609,plain,(
  ( ! [X121,X122,X120,X119] : (sP11('const.list.EL_2'(s('type.num.num',X119),s('type.list.list'(bool),'const.list.TL_1'(s('type.list.list'(bool),X120)))),X122,X121) | p(s(bool,X121)) | ~p(s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X119))),s('type.list.list'(bool),X120))))) )),
  inference(superposition,[],[f567,f422])).
fof(f2610,plain,(
  ( ! [X125,X123,X126,X124] : (sP11('const.list.EL_2'(s('type.num.num',X123),s('type.list.list'(bool),'const.list.TL_1'(s('type.list.list'(bool),X124)))),X125,X126) | ~p(s(bool,X125)) | ~p(s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X123))),s('type.list.list'(bool),X124))))) )),
  inference(superposition,[],[f568,f422])).
fof(f2611,plain,(
  ( ! [X127,X130,X128,X129] : (~sP11('const.list.EL_2'(s('type.num.num',X127),s('type.list.list'(bool),'const.list.TL_1'(s('type.list.list'(bool),X128)))),X130,X129) | ~p(s(bool,X129)) | p(s(bool,X130)) | ~p(s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X127))),s('type.list.list'(bool),X128))))) )),
  inference(superposition,[],[f569,f422])).
fof(f2612,plain,(
  ( ! [X134,X132,X133,X131] : (sP11('const.list.EL_2'(s('type.num.num',X131),s('type.list.list'(bool),'const.list.TL_1'(s('type.list.list'(bool),X132)))),X133,X134) | ~sP10('const.list.EL_2'(s('type.num.num',X131),s('type.list.list'(bool),'const.list.TL_1'(s('type.list.list'(bool),X132)))),X133,X134) | ~p(s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X131))),s('type.list.list'(bool),X132))))) )),
  inference(superposition,[],[f572,f422])).
fof(f2613,plain,(
  ( ! [X138,X136,X135,X137] : (sP9('const.list.EL_2'(s('type.num.num',X135),s('type.list.list'(bool),'const.list.TL_1'(s('type.list.list'(bool),X136)))),X137,X138) | ~sP8('const.list.EL_2'(s('type.num.num',X135),s('type.list.list'(bool),'const.list.TL_1'(s('type.list.list'(bool),X136)))),X137,X138) | ~p(s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X135))),s('type.list.list'(bool),X136))))) )),
  inference(superposition,[],[f573,f422])).
fof(f5444,plain,(
  ( ! [X23,X21,X22] : (~sP5('F_0','const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X21))),s('type.list.list'(bool),X22))),s('type.list.list'(bool),X23)))),'const.list.EL_2'(s('type.num.num',X21),s('type.list.list'(bool),'const.list.TL_1'(s('type.list.list'(bool),X22)))))) )),
  inference(superposition,[],[f1662,f422])).
fof(f5490,plain,(
  ( ! [X23,X21,X22] : (sP5('T_0','const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X21))),s('type.list.list'(bool),X22))),s('type.list.list'(bool),X23)))),'const.list.EL_2'(s('type.num.num',X21),s('type.list.list'(bool),'const.list.TL_1'(s('type.list.list'(bool),X22)))))) )),
  inference(superposition,[],[f1663,f422])).
fof(f5539,plain,(
  ( ! [X23,X21,X22] : (~sP5('F_0','const.list.EL_2'(s('type.num.num',X21),s('type.list.list'(bool),'const.list.TL_1'(s('type.list.list'(bool),X22)))),'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X21))),s('type.list.list'(bool),X22))),s('type.list.list'(bool),X23)))))) )),
  inference(superposition,[],[f1664,f422])).
fof(f5593,plain,(
  ( ! [X23,X21,X22] : (sP5('T_0','const.list.EL_2'(s('type.num.num',X21),s('type.list.list'(bool),'const.list.TL_1'(s('type.list.list'(bool),X22)))),'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X21))),s('type.list.list'(bool),X22))),s('type.list.list'(bool),X23)))))) )),
  inference(superposition,[],[f1665,f422])).
fof(f8348,plain,(
  ( ! [X45,X43,X44] : (s(bool,X45) = s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X43))),s('type.list.list'(bool),X44))) | s(bool,'T_0') = s(bool,X45) | s(bool,'T_0') = s(bool,'const.list.EL_2'(s('type.num.num',X43),s('type.list.list'(bool),'const.list.TL_1'(s('type.list.list'(bool),X44)))))) )),
  inference(superposition,[],[f983,f422])).
fof(f8388,plain,(
  ( ! [X35,X33,X34] : (s(bool,'T_0') = s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X33))),s('type.list.list'(bool),X34))) | s(bool,X35) = s(bool,'const.list.EL_2'(s('type.num.num',X33),s('type.list.list'(bool),'const.list.TL_1'(s('type.list.list'(bool),X34))))) | s(bool,'T_0') = s(bool,X35)) )),
  inference(superposition,[],[f983,f422])).
fof(f10077,plain,(
  ( ! [X453,X455,X457,X454,X456] : (s(X453,X454) = s(X453,X455) | p(s(bool,'const.bool.IN_2'(s(X453,X454),s(fun(X453,bool),X456)))) | s(bool,X457) = s(bool,'const.bool.IN_2'(s(X453,X454),s(fun(X453,bool),'const.pred_set.INSERT_2'(s(X453,X455),s(fun(X453,bool),X456))))) | s(bool,'T_0') = s(bool,X457)) )),
  inference(subsumption_resolution,[],[f9129,f350])).
fof(f9129,plain,(
  ( ! [X453,X455,X457,X454,X456] : (~p(s(bool,'T_0')) | s(X453,X454) = s(X453,X455) | p(s(bool,'const.bool.IN_2'(s(X453,X454),s(fun(X453,bool),X456)))) | s(bool,X457) = s(bool,'const.bool.IN_2'(s(X453,X454),s(fun(X453,bool),'const.pred_set.INSERT_2'(s(X453,X455),s(fun(X453,bool),X456))))) | s(bool,'T_0') = s(bool,X457)) )),
  inference(superposition,[],[f446,f983])).
fof(f10076,plain,(
  ( ! [X449,X451,X452,X448,X450] : (p(s(bool,'const.bool.IN_2'(s(X448,X449),s(fun(X448,bool),'const.pred_set.INSERT_2'(s(X448,X451),s(fun(X448,bool),X450)))))) | s(bool,X452) = s(bool,'const.bool.IN_2'(s(X448,X449),s(fun(X448,bool),X450))) | s(bool,'T_0') = s(bool,X452)) )),
  inference(subsumption_resolution,[],[f9128,f350])).
fof(f9128,plain,(
  ( ! [X449,X451,X452,X448,X450] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(X448,X449),s(fun(X448,bool),'const.pred_set.INSERT_2'(s(X448,X451),s(fun(X448,bool),X450)))))) | s(bool,X452) = s(bool,'const.bool.IN_2'(s(X448,X449),s(fun(X448,bool),X450))) | s(bool,'T_0') = s(bool,X452)) )),
  inference(superposition,[],[f448,f983])).
fof(f446,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) | s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f294])).
fof(f294,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) | s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))) & ((s(X0,X1) != s(X0,X2) & ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3))))))))),
  inference(flattening,[],[f293])).
fof(f293,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) | (s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3)))))) & ((s(X0,X1) != s(X0,X2) & ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3))))))))),
  inference(nnf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) <=> (s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))))),
  inference(flattening,[],[f145])).
fof(f145,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) <=> (s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X6] : ! [X10,X21,X22] : (p(s(bool,'const.bool.IN_2'(s(X6,X10),s(fun(X6,bool),'const.pred_set.INSERT_2'(s(X6,X21),s(fun(X6,bool),X22)))))) <=> (s(X6,X10) = s(X6,X21) | p(s(bool,'const.bool.IN_2'(s(X6,X10),s(fun(X6,bool),X22))))))),
  file('thm.list.MONO_EVERY','thm.pred_set.IN_INSERT')).
fof(f448,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3))))))) )),
  inference(cnf_transformation,[],[f294])).
fof(f447,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X1) != s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3))))))) )),
  inference(cnf_transformation,[],[f294])).
fof(f7608,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X3))))))))) )),
  inference(unit_resulting_resolution,[],[f2232,f448])).
fof(f9147,plain,(
  ( ! [X524,X522,X525,X523] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X523)))) | p(s(bool,'const.bool.IN_2'(s(bool,X522),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X524),s(fun(bool,bool),X523)))))) | s(bool,X522) = s(bool,X525) | s(bool,'T_0') = s(bool,X525)) )),
  inference(superposition,[],[f448,f983])).
fof(f8926,plain,(
  ( ! [X198,X200,X199,X197] : (~p(s(bool,'const.bool.IN_2'(s(bool,X197),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X198),s(fun(bool,bool),X199)))))) | s(bool,'T_0') = s(bool,X198) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X199)))) | s(bool,X197) = s(bool,X200) | s(bool,'T_0') = s(bool,X200)) )),
  inference(superposition,[],[f446,f983])).
fof(f8578,plain,(
  ( ! [X516,X514,X517,X515] : (~p(s(bool,'const.bool.IN_2'(s(bool,X515),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X516),s(fun(bool,bool),X517)))))) | s(bool,X514) = s(bool,X516) | p(s(bool,'const.bool.IN_2'(s(bool,X514),s(fun(bool,bool),X517)))) | s(bool,'T_0') = s(bool,X515) | s(bool,'T_0') = s(bool,X514)) )),
  inference(superposition,[],[f446,f983])).
fof(f8579,plain,(
  ( ! [X520,X518,X521,X519] : (~p(s(bool,'const.bool.IN_2'(s(bool,X520),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X519),s(fun(bool,bool),X521)))))) | s(bool,X518) = s(bool,X520) | p(s(bool,'const.bool.IN_2'(s(bool,X520),s(fun(bool,bool),X521)))) | s(bool,'T_0') = s(bool,X519) | s(bool,'T_0') = s(bool,X518)) )),
  inference(superposition,[],[f446,f983])).
fof(f8582,plain,(
  ( ! [X532,X530,X533,X531] : (~p(s(bool,'const.bool.IN_2'(s(bool,X531),s(fun(bool,bool),X532)))) | p(s(bool,'const.bool.IN_2'(s(bool,X530),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X533),s(fun(bool,bool),X532)))))) | s(bool,'T_0') = s(bool,X531) | s(bool,'T_0') = s(bool,X530)) )),
  inference(superposition,[],[f448,f983])).
fof(f8927,plain,(
  ( ! [X204,X202,X203,X201] : (~p(s(bool,'const.bool.IN_2'(s(bool,X202),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X201),s(fun(bool,bool),X203)))))) | s(bool,'T_0') = s(bool,X202) | p(s(bool,'const.bool.IN_2'(s(bool,X202),s(fun(bool,bool),X203)))) | s(bool,X201) = s(bool,X204) | s(bool,'T_0') = s(bool,X204)) )),
  inference(superposition,[],[f446,f983])).
fof(f8930,plain,(
  ( ! [X216,X215,X213,X214] : (~p(s(bool,'const.bool.IN_2'(s(bool,X213),s(fun(bool,bool),X214)))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X215),s(fun(bool,bool),X214)))))) | s(bool,X213) = s(bool,X216) | s(bool,'T_0') = s(bool,X216)) )),
  inference(superposition,[],[f448,f983])).
fof(f9143,plain,(
  ( ! [X509,X507,X508,X506] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X507),s(fun(bool,bool),X508)))))) | s(bool,X506) = s(bool,X507) | p(s(bool,'const.bool.IN_2'(s(bool,X506),s(fun(bool,bool),X508)))) | s(bool,X506) = s(bool,X509) | s(bool,'T_0') = s(bool,X509)) )),
  inference(superposition,[],[f446,f983])).
fof(f9144,plain,(
  ( ! [X511,X512,X510,X513] : (~p(s(bool,'const.bool.IN_2'(s(bool,X511),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X512)))))) | s(bool,X510) = s(bool,X511) | p(s(bool,'const.bool.IN_2'(s(bool,X511),s(fun(bool,bool),X512)))) | s(bool,X510) = s(bool,X513) | s(bool,'T_0') = s(bool,X513)) )),
  inference(superposition,[],[f446,f983])).
fof(f8580,plain,(
  ( ! [X524,X522,X525,X523] : (s(bool,X523) != s(bool,X524) | p(s(bool,'const.bool.IN_2'(s(bool,X522),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X524),s(fun(bool,bool),X525)))))) | s(bool,'T_0') = s(bool,X523) | s(bool,'T_0') = s(bool,X522)) )),
  inference(superposition,[],[f447,f983])).
fof(f8581,plain,(
  ( ! [X528,X526,X529,X527] : (s(bool,X527) != s(bool,X528) | p(s(bool,'const.bool.IN_2'(s(bool,X528),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X526),s(fun(bool,bool),X529)))))) | s(bool,'T_0') = s(bool,X527) | s(bool,'T_0') = s(bool,X526)) )),
  inference(superposition,[],[f447,f983])).
fof(f9145,plain,(
  ( ! [X516,X514,X517,X515] : (s(bool,'T_0') != s(bool,X515) | p(s(bool,'const.bool.IN_2'(s(bool,X514),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X515),s(fun(bool,bool),X516)))))) | s(bool,X514) = s(bool,X517) | s(bool,'T_0') = s(bool,X517)) )),
  inference(superposition,[],[f447,f983])).
fof(f9146,plain,(
  ( ! [X520,X518,X521,X519] : (s(bool,'T_0') != s(bool,X519) | p(s(bool,'const.bool.IN_2'(s(bool,X519),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X518),s(fun(bool,bool),X520)))))) | s(bool,X518) = s(bool,X521) | s(bool,'T_0') = s(bool,X521)) )),
  inference(superposition,[],[f447,f983])).
fof(f8929,plain,(
  ( ! [X212,X210,X211,X209] : (s(bool,X209) != s(bool,X210) | p(s(bool,'const.bool.IN_2'(s(bool,X210),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X211)))))) | s(bool,X209) = s(bool,X212) | s(bool,'T_0') = s(bool,X212)) )),
  inference(superposition,[],[f447,f983])).
fof(f7083,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X1))))))))) )),
  inference(unit_resulting_resolution,[],[f2229,f448])).
fof(f8928,plain,(
  ( ! [X208,X206,X207,X205] : (s(bool,X205) != s(bool,X206) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X206),s(fun(bool,bool),X207)))))) | s(bool,X205) = s(bool,X208) | s(bool,'T_0') = s(bool,X208)) )),
  inference(superposition,[],[f447,f983])).
fof(f6551,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num',X0),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X1))))))))) )),
  inference(unit_resulting_resolution,[],[f2226,f448])).
fof(f1601,plain,(
  ( ! [X28,X27] : (sP0(X28,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(bool),X27))) | ~p(s(bool,'const.list.HD_1'(s('type.list.list'(bool),X27))))) )),
  inference(superposition,[],[f464,f409])).
fof(f1602,plain,(
  ( ! [X30,X29] : (sP0('const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(bool),X29)),X30) | ~p(s(bool,'const.list.HD_1'(s('type.list.list'(bool),X29))))) )),
  inference(superposition,[],[f465,f409])).
fof(f1603,plain,(
  ( ! [X31,X32] : (sP1('const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(bool),X31)),X32) | p(s(bool,X32)) | ~p(s(bool,'const.list.HD_1'(s('type.list.list'(bool),X31))))) )),
  inference(superposition,[],[f480,f409])).
fof(f1604,plain,(
  ( ! [X33,X34] : (sP1(X34,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(bool),X33))) | p(s(bool,X34)) | ~p(s(bool,'const.list.HD_1'(s('type.list.list'(bool),X33))))) )),
  inference(superposition,[],[f483,f409])).
fof(f1605,plain,(
  ( ! [X37,X35,X36] : (~sP5(X36,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(bool),X35)),X37) | ~p(s(bool,X36)) | p(s(bool,X37)) | ~p(s(bool,'const.list.HD_1'(s('type.list.list'(bool),X35))))) )),
  inference(superposition,[],[f501,f409])).
fof(f1606,plain,(
  ( ! [X39,X38,X40] : (~sP4(X39,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(bool),X38)),X40) | ~p(s(bool,X39)) | p(s(bool,X40)) | ~p(s(bool,'const.list.HD_1'(s('type.list.list'(bool),X38))))) )),
  inference(superposition,[],[f505,f409])).
fof(f1607,plain,(
  ( ! [X43,X41,X42] : (~sP3(X42,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(bool),X41)),X43) | ~p(s(bool,X42)) | p(s(bool,X43)) | ~p(s(bool,'const.list.HD_1'(s('type.list.list'(bool),X41))))) )),
  inference(superposition,[],[f509,f409])).
fof(f1608,plain,(
  ( ! [X45,X46,X44] : (~sP6(X45,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(bool),X44)),X46) | ~p(s(bool,X45)) | p(s(bool,X46)) | ~p(s(bool,'const.list.HD_1'(s('type.list.list'(bool),X44))))) )),
  inference(superposition,[],[f528,f409])).
fof(f1609,plain,(
  ( ! [X47,X48,X49] : (~sP7(X49,X48,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(bool),X47))) | ~p(s(bool,X48)) | p(s(bool,X49)) | ~p(s(bool,'const.list.HD_1'(s('type.list.list'(bool),X47))))) )),
  inference(superposition,[],[f537,f409])).
fof(f1610,plain,(
  ( ! [X52,X50,X51] : (~sP8('const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(bool),X50)),X51,X52) | p(s(bool,X51)) | p(s(bool,X52)) | ~p(s(bool,'const.list.HD_1'(s('type.list.list'(bool),X50))))) )),
  inference(superposition,[],[f545,f409])).
fof(f1611,plain,(
  ( ! [X54,X55,X53] : (sP9('const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(bool),X53)),X55,X54) | ~p(s(bool,X54)) | ~p(s(bool,'const.list.HD_1'(s('type.list.list'(bool),X53))))) )),
  inference(superposition,[],[f550,f409])).
fof(f1612,plain,(
  ( ! [X57,X58,X56] : (sP9('const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(bool),X56)),X57,X58) | ~p(s(bool,X57)) | ~p(s(bool,'const.list.HD_1'(s('type.list.list'(bool),X56))))) )),
  inference(superposition,[],[f551,f409])).
fof(f1613,plain,(
  ( ! [X61,X59,X60] : (~sP9('const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(bool),X59)),X61,X60) | p(s(bool,X60)) | p(s(bool,X61)) | ~p(s(bool,'const.list.HD_1'(s('type.list.list'(bool),X59))))) )),
  inference(superposition,[],[f552,f409])).
fof(f1614,plain,(
  ( ! [X64,X62,X63] : (~sP10(X64,X63,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(bool),X62))) | p(s(bool,X63)) | ~p(s(bool,X64)) | ~p(s(bool,'const.list.HD_1'(s('type.list.list'(bool),X62))))) )),
  inference(superposition,[],[f562,f409])).
fof(f1615,plain,(
  ( ! [X66,X67,X65] : (sP11('const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(bool),X65)),X67,X66) | p(s(bool,X66)) | ~p(s(bool,'const.list.HD_1'(s('type.list.list'(bool),X65))))) )),
  inference(superposition,[],[f567,f409])).
fof(f1616,plain,(
  ( ! [X70,X68,X69] : (sP11('const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(bool),X68)),X69,X70) | ~p(s(bool,X69)) | ~p(s(bool,'const.list.HD_1'(s('type.list.list'(bool),X68))))) )),
  inference(superposition,[],[f568,f409])).
fof(f1617,plain,(
  ( ! [X72,X71,X73] : (~sP11('const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(bool),X71)),X73,X72) | ~p(s(bool,X72)) | p(s(bool,X73)) | ~p(s(bool,'const.list.HD_1'(s('type.list.list'(bool),X71))))) )),
  inference(superposition,[],[f569,f409])).
fof(f1618,plain,(
  ( ! [X76,X74,X75] : (sP11('const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(bool),X74)),X75,X76) | ~sP10('const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(bool),X74)),X75,X76) | ~p(s(bool,'const.list.HD_1'(s('type.list.list'(bool),X74))))) )),
  inference(superposition,[],[f572,f409])).
fof(f1619,plain,(
  ( ! [X78,X79,X77] : (sP9('const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(bool),X77)),X78,X79) | ~sP8('const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(bool),X77)),X78,X79) | ~p(s(bool,'const.list.HD_1'(s('type.list.list'(bool),X77))))) )),
  inference(superposition,[],[f573,f409])).
fof(f5443,plain,(
  ( ! [X19,X20] : (~sP5('F_0','const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'const.list.HD_1'(s('type.list.list'(bool),X19))),s('type.list.list'(bool),X20)))),'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(bool),X19)))) )),
  inference(superposition,[],[f1662,f409])).
fof(f5489,plain,(
  ( ! [X19,X20] : (sP5('T_0','const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'const.list.HD_1'(s('type.list.list'(bool),X19))),s('type.list.list'(bool),X20)))),'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(bool),X19)))) )),
  inference(superposition,[],[f1663,f409])).
fof(f5538,plain,(
  ( ! [X19,X20] : (~sP5('F_0','const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(bool),X19)),'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'const.list.HD_1'(s('type.list.list'(bool),X19))),s('type.list.list'(bool),X20)))))) )),
  inference(superposition,[],[f1664,f409])).
fof(f5592,plain,(
  ( ! [X19,X20] : (sP5('T_0','const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(bool),X19)),'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'const.list.HD_1'(s('type.list.list'(bool),X19))),s('type.list.list'(bool),X20)))))) )),
  inference(superposition,[],[f1665,f409])).
fof(f8347,plain,(
  ( ! [X41,X42] : (s(bool,X42) = s(bool,'const.list.HD_1'(s('type.list.list'(bool),X41))) | s(bool,'T_0') = s(bool,X42) | s(bool,'T_0') = s(bool,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(bool),X41)))) )),
  inference(superposition,[],[f983,f409])).
fof(f8387,plain,(
  ( ! [X31,X32] : (s(bool,'T_0') = s(bool,'const.list.HD_1'(s('type.list.list'(bool),X31))) | s(bool,X32) = s(bool,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(bool),X31))) | s(bool,'T_0') = s(bool,X32)) )),
  inference(superposition,[],[f983,f409])).
fof(f10119,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('T_0',X1,X2) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f10118])).
fof(f10118,plain,(
  ( ! [X130,X128,X131,X129] : (~sP8('T_0',X129,X130) | sP9('T_0',X129,X130) | s(bool,X128) = s(bool,X131) | s(bool,'T_0') = s(bool,X131)) )),
  inference(subsumption_resolution,[],[f9241,f1059])).
fof(f9241,plain,(
  ( ! [X130,X128,X131,X129] : (~p(s(bool,X128)) | ~sP8('T_0',X129,X130) | sP9('T_0',X129,X130) | s(bool,X128) = s(bool,X131) | s(bool,'T_0') = s(bool,X131)) )),
  inference(superposition,[],[f573,f983])).
fof(f10117,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP10('T_0',X1,X2) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f10116])).
fof(f10116,plain,(
  ( ! [X127,X125,X126,X124] : (~sP10('T_0',X125,X126) | sP11('T_0',X125,X126) | s(bool,X124) = s(bool,X127) | s(bool,'T_0') = s(bool,X127)) )),
  inference(subsumption_resolution,[],[f9240,f1059])).
fof(f9240,plain,(
  ( ! [X127,X125,X126,X124] : (~p(s(bool,X124)) | ~sP10('T_0',X125,X126) | sP11('T_0',X125,X126) | s(bool,X124) = s(bool,X127) | s(bool,'T_0') = s(bool,X127)) )),
  inference(superposition,[],[f572,f983])).
fof(f10115,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f10114])).
fof(f10114,plain,(
  ( ! [X123,X121,X122,X120] : (~p(s(bool,X121)) | p(s(bool,X122)) | ~sP11('T_0',X122,X121) | s(bool,X120) = s(bool,X123) | s(bool,'T_0') = s(bool,X123)) )),
  inference(subsumption_resolution,[],[f9239,f1059])).
fof(f9239,plain,(
  ( ! [X123,X121,X122,X120] : (~p(s(bool,X120)) | ~p(s(bool,X121)) | p(s(bool,X122)) | ~sP11('T_0',X122,X121) | s(bool,X120) = s(bool,X123) | s(bool,'T_0') = s(bool,X123)) )),
  inference(superposition,[],[f569,f983])).
fof(f10113,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f10112])).
fof(f10112,plain,(
  ( ! [X118,X116,X119,X117] : (~p(s(bool,X117)) | sP11('T_0',X117,X118) | s(bool,X116) = s(bool,X119) | s(bool,'T_0') = s(bool,X119)) )),
  inference(subsumption_resolution,[],[f9238,f1059])).
fof(f9238,plain,(
  ( ! [X118,X116,X119,X117] : (~p(s(bool,X116)) | ~p(s(bool,X117)) | sP11('T_0',X117,X118) | s(bool,X116) = s(bool,X119) | s(bool,'T_0') = s(bool,X119)) )),
  inference(superposition,[],[f568,f983])).
fof(f10111,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f10110])).
fof(f10110,plain,(
  ( ! [X114,X112,X115,X113] : (p(s(bool,X113)) | sP11('T_0',X114,X113) | s(bool,X112) = s(bool,X115) | s(bool,'T_0') = s(bool,X115)) )),
  inference(subsumption_resolution,[],[f9237,f1059])).
fof(f9237,plain,(
  ( ! [X114,X112,X115,X113] : (~p(s(bool,X112)) | p(s(bool,X113)) | sP11('T_0',X114,X113) | s(bool,X112) = s(bool,X115) | s(bool,'T_0') = s(bool,X115)) )),
  inference(superposition,[],[f567,f983])).
fof(f10109,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP10(X2,X1,'T_0')) )),
  inference(condensation,[],[f10108])).
fof(f10108,plain,(
  ( ! [X111,X109,X110,X108] : (p(s(bool,X109)) | ~p(s(bool,X110)) | ~sP10(X110,X109,'T_0') | s(bool,X108) = s(bool,X111) | s(bool,'T_0') = s(bool,X111)) )),
  inference(subsumption_resolution,[],[f9236,f1059])).
fof(f9236,plain,(
  ( ! [X111,X109,X110,X108] : (~p(s(bool,X108)) | p(s(bool,X109)) | ~p(s(bool,X110)) | ~sP10(X110,X109,'T_0') | s(bool,X108) = s(bool,X111) | s(bool,'T_0') = s(bool,X111)) )),
  inference(superposition,[],[f562,f983])).
fof(f10107,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f10106])).
fof(f10106,plain,(
  ( ! [X107,X105,X106,X104] : (p(s(bool,X105)) | p(s(bool,X106)) | ~sP9('T_0',X106,X105) | s(bool,X104) = s(bool,X107) | s(bool,'T_0') = s(bool,X107)) )),
  inference(subsumption_resolution,[],[f9235,f1059])).
fof(f9235,plain,(
  ( ! [X107,X105,X106,X104] : (~p(s(bool,X104)) | p(s(bool,X105)) | p(s(bool,X106)) | ~sP9('T_0',X106,X105) | s(bool,X104) = s(bool,X107) | s(bool,'T_0') = s(bool,X107)) )),
  inference(superposition,[],[f552,f983])).
fof(f10105,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f10104])).
fof(f10104,plain,(
  ( ! [X103,X101,X102,X100] : (~p(s(bool,X101)) | sP9('T_0',X101,X102) | s(bool,X100) = s(bool,X103) | s(bool,'T_0') = s(bool,X103)) )),
  inference(subsumption_resolution,[],[f9234,f1059])).
fof(f9234,plain,(
  ( ! [X103,X101,X102,X100] : (~p(s(bool,X100)) | ~p(s(bool,X101)) | sP9('T_0',X101,X102) | s(bool,X100) = s(bool,X103) | s(bool,'T_0') = s(bool,X103)) )),
  inference(superposition,[],[f551,f983])).
fof(f10103,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f10102])).
fof(f10102,plain,(
  ( ! [X99,X97,X98,X96] : (~p(s(bool,X97)) | sP9('T_0',X98,X97) | s(bool,X96) = s(bool,X99) | s(bool,'T_0') = s(bool,X99)) )),
  inference(subsumption_resolution,[],[f9233,f1059])).
fof(f9233,plain,(
  ( ! [X99,X97,X98,X96] : (~p(s(bool,X96)) | ~p(s(bool,X97)) | sP9('T_0',X98,X97) | s(bool,X96) = s(bool,X99) | s(bool,'T_0') = s(bool,X99)) )),
  inference(superposition,[],[f550,f983])).
fof(f10101,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP8('T_0',X1,X2)) )),
  inference(condensation,[],[f10100])).
fof(f10100,plain,(
  ( ! [X94,X92,X95,X93] : (p(s(bool,X93)) | p(s(bool,X94)) | ~sP8('T_0',X93,X94) | s(bool,X92) = s(bool,X95) | s(bool,'T_0') = s(bool,X95)) )),
  inference(subsumption_resolution,[],[f9232,f1059])).
fof(f9232,plain,(
  ( ! [X94,X92,X95,X93] : (~p(s(bool,X92)) | p(s(bool,X93)) | p(s(bool,X94)) | ~sP8('T_0',X93,X94) | s(bool,X92) = s(bool,X95) | s(bool,'T_0') = s(bool,X95)) )),
  inference(superposition,[],[f545,f983])).
fof(f10099,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,X1,'T_0')) )),
  inference(condensation,[],[f10098])).
fof(f10098,plain,(
  ( ! [X90,X88,X91,X89] : (~p(s(bool,X89)) | p(s(bool,X90)) | ~sP7(X90,X89,'T_0') | s(bool,X88) = s(bool,X91) | s(bool,'T_0') = s(bool,X91)) )),
  inference(subsumption_resolution,[],[f9231,f1059])).
fof(f9231,plain,(
  ( ! [X90,X88,X91,X89] : (~p(s(bool,X88)) | ~p(s(bool,X89)) | p(s(bool,X90)) | ~sP7(X90,X89,'T_0') | s(bool,X88) = s(bool,X91) | s(bool,'T_0') = s(bool,X91)) )),
  inference(superposition,[],[f537,f983])).
fof(f10097,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X1,'T_0',X2)) )),
  inference(condensation,[],[f10096])).
fof(f10096,plain,(
  ( ! [X87,X85,X86,X84] : (~p(s(bool,X85)) | p(s(bool,X86)) | ~sP6(X85,'T_0',X86) | s(bool,X84) = s(bool,X87) | s(bool,'T_0') = s(bool,X87)) )),
  inference(subsumption_resolution,[],[f9230,f1059])).
fof(f9230,plain,(
  ( ! [X87,X85,X86,X84] : (~p(s(bool,X84)) | ~p(s(bool,X85)) | p(s(bool,X86)) | ~sP6(X85,'T_0',X86) | s(bool,X84) = s(bool,X87) | s(bool,'T_0') = s(bool,X87)) )),
  inference(superposition,[],[f528,f983])).
fof(f10095,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP3(X1,'T_0',X2)) )),
  inference(condensation,[],[f10094])).
fof(f10094,plain,(
  ( ! [X80,X83,X81,X82] : (~p(s(bool,X81)) | p(s(bool,X82)) | ~sP3(X81,'T_0',X82) | s(bool,X80) = s(bool,X83) | s(bool,'T_0') = s(bool,X83)) )),
  inference(subsumption_resolution,[],[f9229,f1059])).
fof(f9229,plain,(
  ( ! [X80,X83,X81,X82] : (~p(s(bool,X80)) | ~p(s(bool,X81)) | p(s(bool,X82)) | ~sP3(X81,'T_0',X82) | s(bool,X80) = s(bool,X83) | s(bool,'T_0') = s(bool,X83)) )),
  inference(superposition,[],[f509,f983])).
fof(f10093,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X1,'T_0',X2)) )),
  inference(condensation,[],[f10092])).
fof(f10092,plain,(
  ( ! [X78,X76,X79,X77] : (~p(s(bool,X77)) | p(s(bool,X78)) | ~sP4(X77,'T_0',X78) | s(bool,X76) = s(bool,X79) | s(bool,'T_0') = s(bool,X79)) )),
  inference(subsumption_resolution,[],[f9228,f1059])).
fof(f9228,plain,(
  ( ! [X78,X76,X79,X77] : (~p(s(bool,X76)) | ~p(s(bool,X77)) | p(s(bool,X78)) | ~sP4(X77,'T_0',X78) | s(bool,X76) = s(bool,X79) | s(bool,'T_0') = s(bool,X79)) )),
  inference(superposition,[],[f505,f983])).
fof(f10091,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5(X1,'T_0',X2)) )),
  inference(condensation,[],[f10090])).
fof(f10090,plain,(
  ( ! [X74,X72,X75,X73] : (~p(s(bool,X73)) | p(s(bool,X74)) | ~sP5(X73,'T_0',X74) | s(bool,X72) = s(bool,X75) | s(bool,'T_0') = s(bool,X75)) )),
  inference(subsumption_resolution,[],[f9227,f1059])).
fof(f9227,plain,(
  ( ! [X74,X72,X75,X73] : (~p(s(bool,X72)) | ~p(s(bool,X73)) | p(s(bool,X74)) | ~sP5(X73,'T_0',X74) | s(bool,X72) = s(bool,X75) | s(bool,'T_0') = s(bool,X75)) )),
  inference(superposition,[],[f501,f983])).
fof(f10089,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f10088])).
fof(f10088,plain,(
  ( ! [X70,X71,X69] : (p(s(bool,X70)) | sP1(X70,'T_0') | s(bool,X69) = s(bool,X71) | s(bool,'T_0') = s(bool,X71)) )),
  inference(subsumption_resolution,[],[f9226,f1059])).
fof(f9226,plain,(
  ( ! [X70,X71,X69] : (~p(s(bool,X69)) | p(s(bool,X70)) | sP1(X70,'T_0') | s(bool,X69) = s(bool,X71) | s(bool,'T_0') = s(bool,X71)) )),
  inference(superposition,[],[f483,f983])).
fof(f10087,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f10086])).
fof(f10086,plain,(
  ( ! [X68,X66,X67] : (p(s(bool,X67)) | sP1('T_0',X67) | s(bool,X66) = s(bool,X68) | s(bool,'T_0') = s(bool,X68)) )),
  inference(subsumption_resolution,[],[f9225,f1059])).
fof(f9225,plain,(
  ( ! [X68,X66,X67] : (~p(s(bool,X66)) | p(s(bool,X67)) | sP1('T_0',X67) | s(bool,X66) = s(bool,X68) | s(bool,'T_0') = s(bool,X68)) )),
  inference(superposition,[],[f480,f983])).
fof(f10081,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,happ(s(fun(X1,bool),X2),s(X1,X3)))) | p(s(bool,happ(s(fun(X1,bool),X2),s(X1,sK26(X2,'T_0',X1)))))) )),
  inference(condensation,[],[f10080])).
fof(f10080,plain,(
  ( ! [X47,X50,X48,X46,X49] : (~p(s(bool,happ(s(fun(X47,bool),X48),s(X47,X49)))) | p(s(bool,happ(s(fun(X47,bool),X48),s(X47,sK26(X48,'T_0',X47))))) | s(bool,X46) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f9219,f1059])).
fof(f9219,plain,(
  ( ! [X47,X50,X48,X46,X49] : (~p(s(bool,X46)) | ~p(s(bool,happ(s(fun(X47,bool),X48),s(X47,X49)))) | p(s(bool,happ(s(fun(X47,bool),X48),s(X47,sK26(X48,'T_0',X47))))) | s(bool,X46) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f438,f983])).
fof(f10079,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,happ(s(fun(X1,bool),X2),s(X1,X3)))) | p(s(bool,happ(s(fun(X1,bool),X2),s(X1,sK25(X2,'T_0',X1)))))) )),
  inference(condensation,[],[f10078])).
fof(f10078,plain,(
  ( ! [X45,X43,X41,X44,X42] : (~p(s(bool,happ(s(fun(X42,bool),X43),s(X42,X44)))) | p(s(bool,happ(s(fun(X42,bool),X43),s(X42,sK25(X43,'T_0',X42))))) | s(bool,X41) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) )),
  inference(subsumption_resolution,[],[f9218,f1059])).
fof(f9218,plain,(
  ( ! [X45,X43,X41,X44,X42] : (~p(s(bool,X41)) | ~p(s(bool,happ(s(fun(X42,bool),X43),s(X42,X44)))) | p(s(bool,happ(s(fun(X42,bool),X43),s(X42,sK25(X43,'T_0',X42))))) | s(bool,X41) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) )),
  inference(superposition,[],[f436,f983])).
fof(f9204,plain,(
  ( ! [X681,X683,X680,X682] : (s(X682,f(s(bool,'T_0'))) != s(X682,g(s(bool,'T_0'))) | s(fun(bool,X682),X680) = s(fun(bool,X682),X681) | s(bool,X683) = s(bool,sK30(X680,X681,X682,bool)) | s(bool,'T_0') = s(bool,X683)) )),
  inference(superposition,[],[f497,f983])).
fof(f9011,plain,(
  ( ! [X114,X115,X113] : (sP5('T_0',X113,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'T_0'),s('type.list.list'(bool),X114))))) | s(bool,X113) = s(bool,X115) | s(bool,'T_0') = s(bool,X115)) )),
  inference(superposition,[],[f1665,f983])).
fof(f9010,plain,(
  ( ! [X111,X112,X110] : (~sP5('F_0',X110,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'T_0'),s('type.list.list'(bool),X111))))) | s(bool,X110) = s(bool,X112) | s(bool,'T_0') = s(bool,X112)) )),
  inference(superposition,[],[f1664,f983])).
fof(f9009,plain,(
  ( ! [X109,X107,X108] : (sP5('T_0','const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'T_0'),s('type.list.list'(bool),X108)))),X107) | s(bool,X107) = s(bool,X109) | s(bool,'T_0') = s(bool,X109)) )),
  inference(superposition,[],[f1663,f983])).
fof(f9008,plain,(
  ( ! [X105,X106,X104] : (~sP5('F_0','const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'T_0'),s('type.list.list'(bool),X105)))),X104) | s(bool,X104) = s(bool,X106) | s(bool,'T_0') = s(bool,X106)) )),
  inference(superposition,[],[f1662,f983])).
fof(f10055,plain,(
  ( ! [X99,X97,X100,X98] : (~sP8(X97,X98,X99) | sP9(X97,X98,X99) | s(bool,X97) = s(bool,X100) | s(bool,'T_0') = s(bool,X100)) )),
  inference(subsumption_resolution,[],[f9006,f350])).
fof(f9006,plain,(
  ( ! [X99,X97,X100,X98] : (~p(s(bool,'T_0')) | ~sP8(X97,X98,X99) | sP9(X97,X98,X99) | s(bool,X97) = s(bool,X100) | s(bool,'T_0') = s(bool,X100)) )),
  inference(superposition,[],[f573,f983])).
fof(f10054,plain,(
  ( ! [X94,X95,X93,X96] : (~sP10(X93,X94,X95) | sP11(X93,X94,X95) | s(bool,X93) = s(bool,X96) | s(bool,'T_0') = s(bool,X96)) )),
  inference(subsumption_resolution,[],[f9005,f350])).
fof(f9005,plain,(
  ( ! [X94,X95,X93,X96] : (~p(s(bool,'T_0')) | ~sP10(X93,X94,X95) | sP11(X93,X94,X95) | s(bool,X93) = s(bool,X96) | s(bool,'T_0') = s(bool,X96)) )),
  inference(superposition,[],[f572,f983])).
fof(f10053,plain,(
  ( ! [X92,X90,X91,X89] : (~p(s(bool,X90)) | p(s(bool,X91)) | ~sP11(X89,X91,X90) | s(bool,X89) = s(bool,X92) | s(bool,'T_0') = s(bool,X92)) )),
  inference(subsumption_resolution,[],[f9004,f350])).
fof(f9004,plain,(
  ( ! [X92,X90,X91,X89] : (~p(s(bool,'T_0')) | ~p(s(bool,X90)) | p(s(bool,X91)) | ~sP11(X89,X91,X90) | s(bool,X89) = s(bool,X92) | s(bool,'T_0') = s(bool,X92)) )),
  inference(superposition,[],[f569,f983])).
fof(f10052,plain,(
  ( ! [X88,X87,X85,X86] : (~p(s(bool,X86)) | sP11(X85,X86,X87) | s(bool,X85) = s(bool,X88) | s(bool,'T_0') = s(bool,X88)) )),
  inference(subsumption_resolution,[],[f9003,f350])).
fof(f9003,plain,(
  ( ! [X88,X87,X85,X86] : (~p(s(bool,'T_0')) | ~p(s(bool,X86)) | sP11(X85,X86,X87) | s(bool,X85) = s(bool,X88) | s(bool,'T_0') = s(bool,X88)) )),
  inference(superposition,[],[f568,f983])).
fof(f10051,plain,(
  ( ! [X83,X81,X84,X82] : (p(s(bool,X82)) | sP11(X81,X83,X82) | s(bool,X81) = s(bool,X84) | s(bool,'T_0') = s(bool,X84)) )),
  inference(subsumption_resolution,[],[f9002,f350])).
fof(f9002,plain,(
  ( ! [X83,X81,X84,X82] : (~p(s(bool,'T_0')) | p(s(bool,X82)) | sP11(X81,X83,X82) | s(bool,X81) = s(bool,X84) | s(bool,'T_0') = s(bool,X84)) )),
  inference(superposition,[],[f567,f983])).
fof(f10050,plain,(
  ( ! [X80,X78,X79,X77] : (p(s(bool,X78)) | ~p(s(bool,X79)) | ~sP10(X79,X78,X77) | s(bool,X77) = s(bool,X80) | s(bool,'T_0') = s(bool,X80)) )),
  inference(subsumption_resolution,[],[f9001,f350])).
fof(f9001,plain,(
  ( ! [X80,X78,X79,X77] : (~p(s(bool,'T_0')) | p(s(bool,X78)) | ~p(s(bool,X79)) | ~sP10(X79,X78,X77) | s(bool,X77) = s(bool,X80) | s(bool,'T_0') = s(bool,X80)) )),
  inference(superposition,[],[f562,f983])).
fof(f10049,plain,(
  ( ! [X76,X74,X75,X73] : (p(s(bool,X74)) | p(s(bool,X75)) | ~sP9(X73,X75,X74) | s(bool,X73) = s(bool,X76) | s(bool,'T_0') = s(bool,X76)) )),
  inference(subsumption_resolution,[],[f9000,f350])).
fof(f9000,plain,(
  ( ! [X76,X74,X75,X73] : (~p(s(bool,'T_0')) | p(s(bool,X74)) | p(s(bool,X75)) | ~sP9(X73,X75,X74) | s(bool,X73) = s(bool,X76) | s(bool,'T_0') = s(bool,X76)) )),
  inference(superposition,[],[f552,f983])).
fof(f10048,plain,(
  ( ! [X70,X72,X71,X69] : (~p(s(bool,X70)) | sP9(X69,X70,X71) | s(bool,X69) = s(bool,X72) | s(bool,'T_0') = s(bool,X72)) )),
  inference(subsumption_resolution,[],[f8999,f350])).
fof(f8999,plain,(
  ( ! [X70,X72,X71,X69] : (~p(s(bool,'T_0')) | ~p(s(bool,X70)) | sP9(X69,X70,X71) | s(bool,X69) = s(bool,X72) | s(bool,'T_0') = s(bool,X72)) )),
  inference(superposition,[],[f551,f983])).
fof(f10047,plain,(
  ( ! [X68,X66,X67,X65] : (~p(s(bool,X66)) | sP9(X65,X67,X66) | s(bool,X65) = s(bool,X68) | s(bool,'T_0') = s(bool,X68)) )),
  inference(subsumption_resolution,[],[f8998,f350])).
fof(f8998,plain,(
  ( ! [X68,X66,X67,X65] : (~p(s(bool,'T_0')) | ~p(s(bool,X66)) | sP9(X65,X67,X66) | s(bool,X65) = s(bool,X68) | s(bool,'T_0') = s(bool,X68)) )),
  inference(superposition,[],[f550,f983])).
fof(f10046,plain,(
  ( ! [X61,X64,X62,X63] : (p(s(bool,X62)) | p(s(bool,X63)) | ~sP8(X61,X62,X63) | s(bool,X61) = s(bool,X64) | s(bool,'T_0') = s(bool,X64)) )),
  inference(subsumption_resolution,[],[f8997,f350])).
fof(f8997,plain,(
  ( ! [X61,X64,X62,X63] : (~p(s(bool,'T_0')) | p(s(bool,X62)) | p(s(bool,X63)) | ~sP8(X61,X62,X63) | s(bool,X61) = s(bool,X64) | s(bool,'T_0') = s(bool,X64)) )),
  inference(superposition,[],[f545,f983])).
fof(f10045,plain,(
  ( ! [X59,X57,X60,X58] : (~p(s(bool,X58)) | p(s(bool,X59)) | ~sP7(X59,X58,X57) | s(bool,X57) = s(bool,X60) | s(bool,'T_0') = s(bool,X60)) )),
  inference(subsumption_resolution,[],[f8996,f350])).
fof(f8996,plain,(
  ( ! [X59,X57,X60,X58] : (~p(s(bool,'T_0')) | ~p(s(bool,X58)) | p(s(bool,X59)) | ~sP7(X59,X58,X57) | s(bool,X57) = s(bool,X60) | s(bool,'T_0') = s(bool,X60)) )),
  inference(superposition,[],[f537,f983])).
fof(f10044,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X54)) | p(s(bool,X55)) | ~sP6(X54,X53,X55) | s(bool,X53) = s(bool,X56) | s(bool,'T_0') = s(bool,X56)) )),
  inference(subsumption_resolution,[],[f8995,f350])).
fof(f8995,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,'T_0')) | ~p(s(bool,X54)) | p(s(bool,X55)) | ~sP6(X54,X53,X55) | s(bool,X53) = s(bool,X56) | s(bool,'T_0') = s(bool,X56)) )),
  inference(superposition,[],[f528,f983])).
fof(f10043,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X50)) | p(s(bool,X51)) | ~sP3(X50,X49,X51) | s(bool,X49) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(subsumption_resolution,[],[f8994,f350])).
fof(f8994,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,X50)) | p(s(bool,X51)) | ~sP3(X50,X49,X51) | s(bool,X49) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(superposition,[],[f509,f983])).
fof(f10042,plain,(
  ( ! [X47,X45,X48,X46] : (~p(s(bool,X46)) | p(s(bool,X47)) | ~sP4(X46,X45,X47) | s(bool,X45) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) )),
  inference(subsumption_resolution,[],[f8993,f350])).
fof(f8993,plain,(
  ( ! [X47,X45,X48,X46] : (~p(s(bool,'T_0')) | ~p(s(bool,X46)) | p(s(bool,X47)) | ~sP4(X46,X45,X47) | s(bool,X45) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) )),
  inference(superposition,[],[f505,f983])).
fof(f10041,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,X42)) | p(s(bool,X43)) | ~sP5(X42,X41,X43) | s(bool,X41) = s(bool,X44) | s(bool,'T_0') = s(bool,X44)) )),
  inference(subsumption_resolution,[],[f8992,f350])).
fof(f8992,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,'T_0')) | ~p(s(bool,X42)) | p(s(bool,X43)) | ~sP5(X42,X41,X43) | s(bool,X41) = s(bool,X44) | s(bool,'T_0') = s(bool,X44)) )),
  inference(superposition,[],[f501,f983])).
fof(f10040,plain,(
  ( ! [X39,X38,X40] : (p(s(bool,X39)) | sP1(X39,X38) | s(bool,X38) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f8991,f350])).
fof(f8991,plain,(
  ( ! [X39,X38,X40] : (~p(s(bool,'T_0')) | p(s(bool,X39)) | sP1(X39,X38) | s(bool,X38) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f483,f983])).
fof(f10039,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,X36)) | sP1(X35,X36) | s(bool,X35) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) )),
  inference(subsumption_resolution,[],[f8990,f350])).
fof(f8990,plain,(
  ( ! [X37,X35,X36] : (~p(s(bool,'T_0')) | p(s(bool,X36)) | sP1(X35,X36) | s(bool,X35) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) )),
  inference(superposition,[],[f480,f983])).
fof(f10038,plain,(
  ( ! [X33,X34,X32] : (sP0(X32,X33) | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f8989,f350])).
fof(f8989,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,'T_0')) | sP0(X32,X33) | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f465,f983])).
fof(f10037,plain,(
  ( ! [X30,X31,X29] : (sP0(X30,X29) | s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(subsumption_resolution,[],[f8988,f350])).
fof(f8988,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,'T_0')) | sP0(X30,X29) | s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(superposition,[],[f464,f983])).
fof(f10036,plain,(
  ( ! [X28,X26,X27] : (~p(s(bool,X27)) | s(bool,X26) = s(bool,X27) | s(bool,X26) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f8987,f350])).
fof(f8987,plain,(
  ( ! [X28,X26,X27] : (~p(s(bool,'T_0')) | ~p(s(bool,X27)) | s(bool,X26) = s(bool,X27) | s(bool,X26) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f454,f983])).
fof(f10033,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f8981,f350])).
fof(f8981,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f377,f983])).
fof(f8931,plain,(
  ( ! [X218,X217] : (s(bool,X217) != s(bool,'const.bool.IN_2'(s(bool,X217),s(fun(bool,bool),'const.pred_set.EMPTY_0'))) | s(bool,X217) = s(bool,X218) | s(bool,'T_0') = s(bool,X218)) )),
  inference(superposition,[],[f715,f983])).
fof(f10032,plain,(
  ( ! [X154,X152,X151,X153] : (sP0('const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X151),s('type.list.list'(bool),X152)))),X153) | s(bool,X151) = s(bool,X154) | s(bool,'T_0') = s(bool,X154)) )),
  inference(subsumption_resolution,[],[f8914,f350])).
fof(f8914,plain,(
  ( ! [X154,X152,X151,X153] : (sP0('const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X151),s('type.list.list'(bool),X152)))),X153) | ~p(s(bool,'T_0')) | s(bool,X151) = s(bool,X154) | s(bool,'T_0') = s(bool,X154)) )),
  inference(superposition,[],[f1694,f983])).
fof(f10031,plain,(
  ( ! [X149,X147,X150,X148] : (sP0(X148,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X147),s('type.list.list'(bool),X149))))) | s(bool,X147) = s(bool,X150) | s(bool,'T_0') = s(bool,X150)) )),
  inference(subsumption_resolution,[],[f8913,f350])).
fof(f8913,plain,(
  ( ! [X149,X147,X150,X148] : (sP0(X148,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X147),s('type.list.list'(bool),X149))))) | ~p(s(bool,'T_0')) | s(bool,X147) = s(bool,X150) | s(bool,'T_0') = s(bool,X150)) )),
  inference(superposition,[],[f1693,f983])).
fof(f8912,plain,(
  ( ! [X146,X144,X145] : (sP5('T_0','T_0','const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X144),s('type.list.list'(bool),X145))))) | s(bool,X144) = s(bool,X146) | s(bool,'T_0') = s(bool,X146)) )),
  inference(superposition,[],[f1665,f983])).
fof(f8911,plain,(
  ( ! [X142,X143,X141] : (~sP5('F_0','T_0','const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X141),s('type.list.list'(bool),X142))))) | s(bool,X141) = s(bool,X143) | s(bool,'T_0') = s(bool,X143)) )),
  inference(superposition,[],[f1664,f983])).
fof(f8910,plain,(
  ( ! [X140,X138,X139] : (sP5('T_0','const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X138),s('type.list.list'(bool),X139)))),'T_0') | s(bool,X138) = s(bool,X140) | s(bool,'T_0') = s(bool,X140)) )),
  inference(superposition,[],[f1663,f983])).
fof(f8909,plain,(
  ( ! [X136,X135,X137] : (~sP5('F_0','const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X135),s('type.list.list'(bool),X136)))),'T_0') | s(bool,X135) = s(bool,X137) | s(bool,'T_0') = s(bool,X137)) )),
  inference(superposition,[],[f1662,f983])).
fof(f10029,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('T_0',X1,X2) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f10028])).
fof(f10028,plain,(
  ( ! [X130,X128,X131,X129] : (~sP8('T_0',X129,X130) | sP9('T_0',X129,X130) | s(bool,X128) = s(bool,X131) | s(bool,'T_0') = s(bool,X131)) )),
  inference(subsumption_resolution,[],[f8907,f1059])).
fof(f8907,plain,(
  ( ! [X130,X128,X131,X129] : (~p(s(bool,X128)) | ~sP8('T_0',X129,X130) | sP9('T_0',X129,X130) | s(bool,X128) = s(bool,X131) | s(bool,'T_0') = s(bool,X131)) )),
  inference(superposition,[],[f573,f983])).
fof(f10026,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP10('T_0',X1,X2) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f10025])).
fof(f10025,plain,(
  ( ! [X127,X125,X126,X124] : (~sP10('T_0',X125,X126) | sP11('T_0',X125,X126) | s(bool,X124) = s(bool,X127) | s(bool,'T_0') = s(bool,X127)) )),
  inference(subsumption_resolution,[],[f8906,f1059])).
fof(f8906,plain,(
  ( ! [X127,X125,X126,X124] : (~p(s(bool,X124)) | ~sP10('T_0',X125,X126) | sP11('T_0',X125,X126) | s(bool,X124) = s(bool,X127) | s(bool,'T_0') = s(bool,X127)) )),
  inference(superposition,[],[f572,f983])).
fof(f10023,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f10022])).
fof(f10022,plain,(
  ( ! [X123,X121,X122,X120] : (~p(s(bool,X121)) | p(s(bool,X122)) | ~sP11('T_0',X122,X121) | s(bool,X120) = s(bool,X123) | s(bool,'T_0') = s(bool,X123)) )),
  inference(subsumption_resolution,[],[f8905,f1059])).
fof(f8905,plain,(
  ( ! [X123,X121,X122,X120] : (~p(s(bool,X120)) | ~p(s(bool,X121)) | p(s(bool,X122)) | ~sP11('T_0',X122,X121) | s(bool,X120) = s(bool,X123) | s(bool,'T_0') = s(bool,X123)) )),
  inference(superposition,[],[f569,f983])).
fof(f10020,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f10019])).
fof(f10019,plain,(
  ( ! [X118,X116,X119,X117] : (~p(s(bool,X117)) | sP11('T_0',X117,X118) | s(bool,X116) = s(bool,X119) | s(bool,'T_0') = s(bool,X119)) )),
  inference(subsumption_resolution,[],[f8904,f1059])).
fof(f8904,plain,(
  ( ! [X118,X116,X119,X117] : (~p(s(bool,X116)) | ~p(s(bool,X117)) | sP11('T_0',X117,X118) | s(bool,X116) = s(bool,X119) | s(bool,'T_0') = s(bool,X119)) )),
  inference(superposition,[],[f568,f983])).
fof(f10017,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f10016])).
fof(f10016,plain,(
  ( ! [X114,X112,X115,X113] : (p(s(bool,X113)) | sP11('T_0',X114,X113) | s(bool,X112) = s(bool,X115) | s(bool,'T_0') = s(bool,X115)) )),
  inference(subsumption_resolution,[],[f8903,f1059])).
fof(f8903,plain,(
  ( ! [X114,X112,X115,X113] : (~p(s(bool,X112)) | p(s(bool,X113)) | sP11('T_0',X114,X113) | s(bool,X112) = s(bool,X115) | s(bool,'T_0') = s(bool,X115)) )),
  inference(superposition,[],[f567,f983])).
fof(f10014,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP10(X2,X1,'T_0')) )),
  inference(condensation,[],[f10013])).
fof(f10013,plain,(
  ( ! [X111,X109,X110,X108] : (p(s(bool,X109)) | ~p(s(bool,X110)) | ~sP10(X110,X109,'T_0') | s(bool,X108) = s(bool,X111) | s(bool,'T_0') = s(bool,X111)) )),
  inference(subsumption_resolution,[],[f8902,f1059])).
fof(f8902,plain,(
  ( ! [X111,X109,X110,X108] : (~p(s(bool,X108)) | p(s(bool,X109)) | ~p(s(bool,X110)) | ~sP10(X110,X109,'T_0') | s(bool,X108) = s(bool,X111) | s(bool,'T_0') = s(bool,X111)) )),
  inference(superposition,[],[f562,f983])).
fof(f10011,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f10010])).
fof(f10010,plain,(
  ( ! [X107,X105,X106,X104] : (p(s(bool,X105)) | p(s(bool,X106)) | ~sP9('T_0',X106,X105) | s(bool,X104) = s(bool,X107) | s(bool,'T_0') = s(bool,X107)) )),
  inference(subsumption_resolution,[],[f8901,f1059])).
fof(f8901,plain,(
  ( ! [X107,X105,X106,X104] : (~p(s(bool,X104)) | p(s(bool,X105)) | p(s(bool,X106)) | ~sP9('T_0',X106,X105) | s(bool,X104) = s(bool,X107) | s(bool,'T_0') = s(bool,X107)) )),
  inference(superposition,[],[f552,f983])).
fof(f10008,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f10007])).
fof(f10007,plain,(
  ( ! [X103,X101,X102,X100] : (~p(s(bool,X101)) | sP9('T_0',X101,X102) | s(bool,X100) = s(bool,X103) | s(bool,'T_0') = s(bool,X103)) )),
  inference(subsumption_resolution,[],[f8900,f1059])).
fof(f8900,plain,(
  ( ! [X103,X101,X102,X100] : (~p(s(bool,X100)) | ~p(s(bool,X101)) | sP9('T_0',X101,X102) | s(bool,X100) = s(bool,X103) | s(bool,'T_0') = s(bool,X103)) )),
  inference(superposition,[],[f551,f983])).
fof(f10005,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f10004])).
fof(f10004,plain,(
  ( ! [X99,X97,X98,X96] : (~p(s(bool,X97)) | sP9('T_0',X98,X97) | s(bool,X96) = s(bool,X99) | s(bool,'T_0') = s(bool,X99)) )),
  inference(subsumption_resolution,[],[f8899,f1059])).
fof(f8899,plain,(
  ( ! [X99,X97,X98,X96] : (~p(s(bool,X96)) | ~p(s(bool,X97)) | sP9('T_0',X98,X97) | s(bool,X96) = s(bool,X99) | s(bool,'T_0') = s(bool,X99)) )),
  inference(superposition,[],[f550,f983])).
fof(f10002,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP8('T_0',X1,X2)) )),
  inference(condensation,[],[f10001])).
fof(f10001,plain,(
  ( ! [X94,X92,X95,X93] : (p(s(bool,X93)) | p(s(bool,X94)) | ~sP8('T_0',X93,X94) | s(bool,X92) = s(bool,X95) | s(bool,'T_0') = s(bool,X95)) )),
  inference(subsumption_resolution,[],[f8898,f1059])).
fof(f8898,plain,(
  ( ! [X94,X92,X95,X93] : (~p(s(bool,X92)) | p(s(bool,X93)) | p(s(bool,X94)) | ~sP8('T_0',X93,X94) | s(bool,X92) = s(bool,X95) | s(bool,'T_0') = s(bool,X95)) )),
  inference(superposition,[],[f545,f983])).
fof(f9999,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,X1,'T_0')) )),
  inference(condensation,[],[f9998])).
fof(f9998,plain,(
  ( ! [X90,X88,X91,X89] : (~p(s(bool,X89)) | p(s(bool,X90)) | ~sP7(X90,X89,'T_0') | s(bool,X88) = s(bool,X91) | s(bool,'T_0') = s(bool,X91)) )),
  inference(subsumption_resolution,[],[f8897,f1059])).
fof(f8897,plain,(
  ( ! [X90,X88,X91,X89] : (~p(s(bool,X88)) | ~p(s(bool,X89)) | p(s(bool,X90)) | ~sP7(X90,X89,'T_0') | s(bool,X88) = s(bool,X91) | s(bool,'T_0') = s(bool,X91)) )),
  inference(superposition,[],[f537,f983])).
fof(f9996,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X1,'T_0',X2)) )),
  inference(condensation,[],[f9995])).
fof(f9995,plain,(
  ( ! [X87,X85,X86,X84] : (~p(s(bool,X85)) | p(s(bool,X86)) | ~sP6(X85,'T_0',X86) | s(bool,X84) = s(bool,X87) | s(bool,'T_0') = s(bool,X87)) )),
  inference(subsumption_resolution,[],[f8896,f1059])).
fof(f8896,plain,(
  ( ! [X87,X85,X86,X84] : (~p(s(bool,X84)) | ~p(s(bool,X85)) | p(s(bool,X86)) | ~sP6(X85,'T_0',X86) | s(bool,X84) = s(bool,X87) | s(bool,'T_0') = s(bool,X87)) )),
  inference(superposition,[],[f528,f983])).
fof(f9993,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP3(X1,'T_0',X2)) )),
  inference(condensation,[],[f9992])).
fof(f9992,plain,(
  ( ! [X80,X83,X81,X82] : (~p(s(bool,X81)) | p(s(bool,X82)) | ~sP3(X81,'T_0',X82) | s(bool,X80) = s(bool,X83) | s(bool,'T_0') = s(bool,X83)) )),
  inference(subsumption_resolution,[],[f8895,f1059])).
fof(f8895,plain,(
  ( ! [X80,X83,X81,X82] : (~p(s(bool,X80)) | ~p(s(bool,X81)) | p(s(bool,X82)) | ~sP3(X81,'T_0',X82) | s(bool,X80) = s(bool,X83) | s(bool,'T_0') = s(bool,X83)) )),
  inference(superposition,[],[f509,f983])).
fof(f9990,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X1,'T_0',X2)) )),
  inference(condensation,[],[f9989])).
fof(f9989,plain,(
  ( ! [X78,X76,X79,X77] : (~p(s(bool,X77)) | p(s(bool,X78)) | ~sP4(X77,'T_0',X78) | s(bool,X76) = s(bool,X79) | s(bool,'T_0') = s(bool,X79)) )),
  inference(subsumption_resolution,[],[f8894,f1059])).
fof(f8894,plain,(
  ( ! [X78,X76,X79,X77] : (~p(s(bool,X76)) | ~p(s(bool,X77)) | p(s(bool,X78)) | ~sP4(X77,'T_0',X78) | s(bool,X76) = s(bool,X79) | s(bool,'T_0') = s(bool,X79)) )),
  inference(superposition,[],[f505,f983])).
fof(f9987,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5(X1,'T_0',X2)) )),
  inference(condensation,[],[f9986])).
fof(f9986,plain,(
  ( ! [X74,X72,X75,X73] : (~p(s(bool,X73)) | p(s(bool,X74)) | ~sP5(X73,'T_0',X74) | s(bool,X72) = s(bool,X75) | s(bool,'T_0') = s(bool,X75)) )),
  inference(subsumption_resolution,[],[f8893,f1059])).
fof(f8893,plain,(
  ( ! [X74,X72,X75,X73] : (~p(s(bool,X72)) | ~p(s(bool,X73)) | p(s(bool,X74)) | ~sP5(X73,'T_0',X74) | s(bool,X72) = s(bool,X75) | s(bool,'T_0') = s(bool,X75)) )),
  inference(superposition,[],[f501,f983])).
fof(f9984,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f9983])).
fof(f9983,plain,(
  ( ! [X70,X71,X69] : (p(s(bool,X70)) | sP1(X70,'T_0') | s(bool,X69) = s(bool,X71) | s(bool,'T_0') = s(bool,X71)) )),
  inference(subsumption_resolution,[],[f8892,f1059])).
fof(f8892,plain,(
  ( ! [X70,X71,X69] : (~p(s(bool,X69)) | p(s(bool,X70)) | sP1(X70,'T_0') | s(bool,X69) = s(bool,X71) | s(bool,'T_0') = s(bool,X71)) )),
  inference(superposition,[],[f483,f983])).
fof(f9981,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f9980])).
fof(f9980,plain,(
  ( ! [X68,X66,X67] : (p(s(bool,X67)) | sP1('T_0',X67) | s(bool,X66) = s(bool,X68) | s(bool,'T_0') = s(bool,X68)) )),
  inference(subsumption_resolution,[],[f8891,f1059])).
fof(f8891,plain,(
  ( ! [X68,X66,X67] : (~p(s(bool,X66)) | p(s(bool,X67)) | sP1('T_0',X67) | s(bool,X66) = s(bool,X68) | s(bool,'T_0') = s(bool,X68)) )),
  inference(superposition,[],[f480,f983])).
fof(f9974,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,happ(s(fun(X1,bool),X2),s(X1,X3)))) | p(s(bool,happ(s(fun(X1,bool),X2),s(X1,sK26(X2,'T_0',X1)))))) )),
  inference(condensation,[],[f9973])).
fof(f9973,plain,(
  ( ! [X47,X50,X48,X46,X49] : (~p(s(bool,happ(s(fun(X47,bool),X48),s(X47,X49)))) | p(s(bool,happ(s(fun(X47,bool),X48),s(X47,sK26(X48,'T_0',X47))))) | s(bool,X46) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f8885,f1059])).
fof(f8885,plain,(
  ( ! [X47,X50,X48,X46,X49] : (~p(s(bool,X46)) | ~p(s(bool,happ(s(fun(X47,bool),X48),s(X47,X49)))) | p(s(bool,happ(s(fun(X47,bool),X48),s(X47,sK26(X48,'T_0',X47))))) | s(bool,X46) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f438,f983])).
fof(f9970,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,happ(s(fun(X1,bool),X2),s(X1,X3)))) | p(s(bool,happ(s(fun(X1,bool),X2),s(X1,sK25(X2,'T_0',X1)))))) )),
  inference(condensation,[],[f9969])).
fof(f9969,plain,(
  ( ! [X45,X43,X41,X44,X42] : (~p(s(bool,happ(s(fun(X42,bool),X43),s(X42,X44)))) | p(s(bool,happ(s(fun(X42,bool),X43),s(X42,sK25(X43,'T_0',X42))))) | s(bool,X41) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) )),
  inference(subsumption_resolution,[],[f8884,f1059])).
fof(f8884,plain,(
  ( ! [X45,X43,X41,X44,X42] : (~p(s(bool,X41)) | ~p(s(bool,happ(s(fun(X42,bool),X43),s(X42,X44)))) | p(s(bool,happ(s(fun(X42,bool),X43),s(X42,sK25(X43,'T_0',X42))))) | s(bool,X41) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) )),
  inference(superposition,[],[f436,f983])).
fof(f8879,plain,(
  ( ! [X30,X31,X32] : (sP4('T_0','const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X30),s('type.list.list'(bool),X31))))) | s(bool,X30) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(superposition,[],[f5603,f983])).
fof(f8878,plain,(
  ( ! [X28,X29,X27] : (sP5('T_0','const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X27),s('type.list.list'(bool),X28))))) | s(bool,X27) = s(bool,X29) | s(bool,'T_0') = s(bool,X29)) )),
  inference(superposition,[],[f5585,f983])).
fof(f8877,plain,(
  ( ! [X26,X24,X25] : (~sP3('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X24),s('type.list.list'(bool),X25)))),'F_0') | s(bool,X24) = s(bool,X26) | s(bool,'T_0') = s(bool,X26)) )),
  inference(superposition,[],[f5547,f983])).
fof(f8876,plain,(
  ( ! [X23,X21,X22] : (~sP5('F_0','const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X21),s('type.list.list'(bool),X22))))) | s(bool,X21) = s(bool,X23) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f5531,f983])).
fof(f8875,plain,(
  ( ! [X19,X20,X18] : (sP5('T_0','const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X18),s('type.list.list'(bool),X19)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | s(bool,X18) = s(bool,X20) | s(bool,'T_0') = s(bool,X20)) )),
  inference(superposition,[],[f5482,f983])).
fof(f8874,plain,(
  ( ! [X17,X15,X16] : (~sP3('const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X15),s('type.list.list'(bool),X16)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'F_0') | s(bool,X15) = s(bool,X17) | s(bool,'T_0') = s(bool,X17)) )),
  inference(superposition,[],[f5458,f983])).
fof(f8873,plain,(
  ( ! [X14,X12,X13] : (~sP5('F_0','const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X12),s('type.list.list'(bool),X13)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | s(bool,X12) = s(bool,X14) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f5436,f983])).
fof(f8872,plain,(
  ( ! [X4,X2,X3] : (s(bool,X2) != s(bool,X3) | p(s(bool,X3)) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f373,f983])).
fof(f8639,plain,(
  ( ! [X689,X691,X692,X690] : (s(X691,f(s(bool,X692))) != s(X691,g(s(bool,X692))) | s(fun(bool,X691),X689) = s(fun(bool,X691),X690) | s(bool,'T_0') = s(bool,X692) | s(bool,'T_0') = s(bool,sK30(X689,X690,X691,bool))) )),
  inference(superposition,[],[f497,f983])).
fof(f8442,plain,(
  ( ! [X114,X115,X113] : (sP5('T_0',X113,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X114),s('type.list.list'(bool),X115))))) | s(bool,'T_0') = s(bool,X114) | s(bool,'T_0') = s(bool,X113)) )),
  inference(superposition,[],[f1665,f983])).
fof(f8441,plain,(
  ( ! [X111,X112,X110] : (~sP5('F_0',X110,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X111),s('type.list.list'(bool),X112))))) | s(bool,'T_0') = s(bool,X111) | s(bool,'T_0') = s(bool,X110)) )),
  inference(superposition,[],[f1664,f983])).
fof(f8440,plain,(
  ( ! [X109,X107,X108] : (sP5('T_0','const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X108),s('type.list.list'(bool),X109)))),X107) | s(bool,'T_0') = s(bool,X108) | s(bool,'T_0') = s(bool,X107)) )),
  inference(superposition,[],[f1663,f983])).
fof(f8439,plain,(
  ( ! [X105,X106,X104] : (~sP5('F_0','const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X105),s('type.list.list'(bool),X106)))),X104) | s(bool,'T_0') = s(bool,X105) | s(bool,'T_0') = s(bool,X104)) )),
  inference(superposition,[],[f1662,f983])).
fof(f8389,plain,(
  ( ! [X37,X38,X36] : (s(bool,'T_0') = s(bool,X36) | s(bool,X38) = s(bool,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X36),s('type.list.list'(bool),X37))))) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f983,f419])).
fof(f8371,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f983,f983])).
fof(f8349,plain,(
  ( ! [X47,X48,X46] : (s(bool,X46) = s(bool,X48) | s(bool,'T_0') = s(bool,X48) | s(bool,'T_0') = s(bool,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X46),s('type.list.list'(bool),X47)))))) )),
  inference(superposition,[],[f983,f419])).
fof(f7613,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f350,f2232,f454])).
fof(f7614,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2))))) = s(bool,'const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),X5)))))) )),
  inference(unit_resulting_resolution,[],[f2232,f2232,f454])).
fof(f7615,plain,(
  ( ! [X2,X0,X3,X1] : (sP0(X0,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.INSERT_2'(s(X1,X2),s(fun(X1,bool),X3)))))) )),
  inference(unit_resulting_resolution,[],[f2232,f464])).
fof(f7616,plain,(
  ( ! [X2,X0,X3,X1] : (sP0('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),X3)) )),
  inference(unit_resulting_resolution,[],[f2232,f465])).
fof(f7617,plain,(
  ( ! [X2,X0,X1] : (~sP1('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f2232,f479])).
fof(f7620,plain,(
  ( ! [X2,X0,X1] : (~sP1('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f350,f2232,f479])).
fof(f7621,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~sP1('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),X5)))))) )),
  inference(unit_resulting_resolution,[],[f2232,f2232,f479])).
fof(f7623,plain,(
  ( ! [X2,X0,X1] : (sP1('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f2232,f480])).
fof(f7625,plain,(
  ( ! [X2,X0,X1] : (sP1('F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f585,f2232,f483])).
fof(f7631,plain,(
  ( ! [X2,X0,X1] : (~sP5('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f350,f585,f2232,f501])).
fof(f7632,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),X5)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f2232,f585,f2232,f501])).
fof(f7692,plain,(
  ( ! [X2,X0,X1] : (sP5('T_0','T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f598,f592,f606,f2232,f502])).
fof(f7693,plain,(
  ( ! [X2,X0,X1] : (sP5('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f599,f592,f606,f2232,f502])).
fof(f7694,plain,(
  ( ! [X2,X0,X3,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'T_0','const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.INSERT_2'(s(X1,X2),s(fun(X1,bool),X3)))))) )),
  inference(unit_resulting_resolution,[],[f918,f592,f606,f2232,f502])).
fof(f7697,plain,(
  ( ! [X2,X0,X1] : (sP5('T_0','const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f598,f593,f607,f2232,f502])).
fof(f7698,plain,(
  ( ! [X2,X0,X1] : (sP5('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f599,f593,f607,f2232,f502])).
fof(f7699,plain,(
  ( ! [X2,X0,X3,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.INSERT_2'(s(X1,X2),s(fun(X1,bool),X3)))))) )),
  inference(unit_resulting_resolution,[],[f918,f593,f607,f2232,f502])).
fof(f7701,plain,(
  ( ! [X2,X0,X3,X1] : (sP5('T_0','const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.INSERT_2'(s(X1,X2),s(fun(X1,bool),X3)))))) )),
  inference(unit_resulting_resolution,[],[f598,f917,f920,f2232,f502])).
fof(f7702,plain,(
  ( ! [X2,X0,X3,X1] : (sP5('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.INSERT_2'(s(X1,X2),s(fun(X1,bool),X3)))))) )),
  inference(unit_resulting_resolution,[],[f599,f917,f920,f2232,f502])).
fof(f7703,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))),'const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.pred_set.INSERT_2'(s(X2,X3),s(fun(X2,bool),X4)))))) )),
  inference(unit_resulting_resolution,[],[f918,f917,f920,f2232,f502])).
fof(f7709,plain,(
  ( ! [X2,X0,X1] : (~sP4('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f350,f585,f2232,f505])).
fof(f7721,plain,(
  ( ! [X2,X0,X1] : (~sP4('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f350,f585,f2232,f505])).
fof(f7722,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~sP4('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),X5)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f2232,f585,f2232,f505])).
fof(f7733,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP4(X0,X1,'const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.pred_set.INSERT_2'(s(X2,X3),s(fun(X2,bool),X4)))))) )),
  inference(unit_resulting_resolution,[],[f2232,f506])).
fof(f7734,plain,(
  ( ! [X2,X0,X1] : (~sP3('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f350,f585,f2232,f509])).
fof(f7746,plain,(
  ( ! [X2,X0,X1] : (~sP3('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f350,f585,f2232,f509])).
fof(f7747,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~sP3('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),X5)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f2232,f585,f2232,f509])).
fof(f7758,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP3(X0,X1,'const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.pred_set.INSERT_2'(s(X2,X3),s(fun(X2,bool),X4)))))) )),
  inference(unit_resulting_resolution,[],[f2232,f510])).
fof(f7759,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP2(X0,X1,'const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.pred_set.INSERT_2'(s(X2,X3),s(fun(X2,bool),X4)))))) )),
  inference(unit_resulting_resolution,[],[f2232,f514])).
fof(f7760,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP2(X0,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.INSERT_2'(s(X1,X2),s(fun(X1,bool),X3)))),X4)) )),
  inference(unit_resulting_resolution,[],[f2232,f515])).
fof(f7761,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP2('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),X3,X4)) )),
  inference(unit_resulting_resolution,[],[f2232,f516])).
fof(f7762,plain,(
  ( ! [X2,X0,X1] : (sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f647,f2232,f518])).
fof(f7763,plain,(
  ( ! [X2,X0,X1] : (sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f667,f2232,f518])).
fof(f7764,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.list.EVERY_2'(s(fun(X3,bool),X4),s('type.list.list'(X3),'const.list.NIL_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f408,f2232,f518])).
fof(f7765,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X4)))),'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num',X4)))) )),
  inference(unit_resulting_resolution,[],[f449,f2232,f518])).
fof(f7766,plain,(
  ( ! [X2,X0,X3,X1] : (sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.num.|48|_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f668,f2232,f518])).
fof(f7767,plain,(
  ( ! [X2,X0,X3,X1] : (sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f907,f2232,f518])).
fof(f7769,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.EMPTY_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f716,f2232,f518])).
fof(f7770,plain,(
  ( ! [X2,X0,X3,X1] : (sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(bool),X3)),'const.list.HD_1'(s('type.list.list'(bool),X3)))) )),
  inference(unit_resulting_resolution,[],[f409,f2232,f518])).
fof(f7771,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X3),s('type.list.list'(bool),X4)))),X3)) )),
  inference(unit_resulting_resolution,[],[f419,f2232,f518])).
fof(f7772,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.list.EL_2'(s('type.num.num',X3),s('type.list.list'(bool),'const.list.TL_1'(s('type.list.list'(bool),X4)))),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))),s('type.list.list'(bool),X4)))) )),
  inference(unit_resulting_resolution,[],[f422,f2232,f518])).
fof(f7773,plain,(
  ( ! [X2,X0,X1] : (sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'T_0','const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) )),
  inference(unit_resulting_resolution,[],[f647,f2232,f518])).
fof(f7774,plain,(
  ( ! [X2,X0,X1] : (sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'F_0','const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)))) )),
  inference(unit_resulting_resolution,[],[f667,f2232,f518])).
fof(f7775,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'T_0','const.list.EVERY_2'(s(fun(X3,bool),X4),s('type.list.list'(X3),'const.list.NIL_0')))) )),
  inference(unit_resulting_resolution,[],[f408,f2232,f518])).
fof(f7776,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num',X4)),'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X4)))))) )),
  inference(unit_resulting_resolution,[],[f449,f2232,f518])).
fof(f7777,plain,(
  ( ! [X2,X0,X3,X1] : (sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'F_0','const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f668,f2232,f518])).
fof(f7778,plain,(
  ( ! [X2,X0,X3,X1] : (sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'T_0','const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))) )),
  inference(unit_resulting_resolution,[],[f907,f2232,f518])).
fof(f7780,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'F_0','const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f716,f2232,f518])).
fof(f7781,plain,(
  ( ! [X2,X0,X3,X1] : (sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.list.HD_1'(s('type.list.list'(bool),X3)),'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(bool),X3)))) )),
  inference(unit_resulting_resolution,[],[f409,f2232,f518])).
fof(f7782,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),X3,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X3),s('type.list.list'(bool),X4)))))) )),
  inference(unit_resulting_resolution,[],[f419,f2232,f518])).
fof(f7783,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))),s('type.list.list'(bool),X4)),'const.list.EL_2'(s('type.num.num',X3),s('type.list.list'(bool),'const.list.TL_1'(s('type.list.list'(bool),X4)))))) )),
  inference(unit_resulting_resolution,[],[f422,f2232,f518])).
fof(f7786,plain,(
  ( ! [X2,X0,X1] : (~sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f642,f2232,f519])).
fof(f7787,plain,(
  ( ! [X2,X0,X1] : (~sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'T_0','const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)))) )),
  inference(unit_resulting_resolution,[],[f643,f2232,f519])).
fof(f7788,plain,(
  ( ! [X2,X0,X3,X1] : (~sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'T_0','const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f644,f2232,f519])).
fof(f7789,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'T_0','const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f715,f2232,f519])).
fof(f7790,plain,(
  ( ! [X2,X0,X1] : (~sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f642,f2232,f519])).
fof(f7791,plain,(
  ( ! [X2,X0,X1] : (~sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f643,f2232,f519])).
fof(f7792,plain,(
  ( ! [X2,X0,X3,X1] : (~sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.num.|48|_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f644,f2232,f519])).
fof(f7793,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.EMPTY_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f715,f2232,f519])).
fof(f7796,plain,(
  ( ! [X2,X0,X3,X1] : (~sP7('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),X3,'F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f2232,f522])).
fof(f7798,plain,(
  ( ! [X2,X0,X3,X1] : (~sP7('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'F_0',X3)) )),
  inference(unit_resulting_resolution,[],[f585,f2232,f523])).
fof(f7802,plain,(
  ( ! [X2,X0,X1] : (sP7('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f611,f2232,f524])).
fof(f7804,plain,(
  ( ! [X2,X0,X3,X1] : (sP7('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.INSERT_2'(s(X1,X2),s(fun(X1,bool),X3)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f924,f2232,f524])).
fof(f7814,plain,(
  ( ! [X2,X0,X1] : (sP7('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f350,f611,f2232,f524])).
fof(f7815,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP7('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),X5)))))) )),
  inference(unit_resulting_resolution,[],[f2232,f611,f2232,f524])).
fof(f7816,plain,(
  ( ! [X2,X0,X3,X1] : (sP7('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'T_0','const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.INSERT_2'(s(X1,X2),s(fun(X1,bool),X3)))))) )),
  inference(unit_resulting_resolution,[],[f350,f924,f2232,f524])).
fof(f7817,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (sP7('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.INSERT_2'(s(X1,X2),s(fun(X1,bool),X3)))),'const.bool.IN_2'(s(X4,X5),s(fun(X4,bool),'const.pred_set.INSERT_2'(s(X4,X5),s(fun(X4,bool),X6)))))) )),
  inference(unit_resulting_resolution,[],[f2232,f924,f2232,f524])).
fof(f7824,plain,(
  ( ! [X2,X0,X1] : (~sP6('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f350,f585,f2232,f528])).
fof(f7836,plain,(
  ( ! [X2,X0,X1] : (~sP6('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f350,f585,f2232,f528])).
fof(f7837,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~sP6('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),X5)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f2232,f585,f2232,f528])).
fof(f7848,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP6(X0,X1,'const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.pred_set.INSERT_2'(s(X2,X3),s(fun(X2,bool),X4)))))) )),
  inference(unit_resulting_resolution,[],[f2232,f529])).
fof(f7849,plain,(
  ( ! [X2,X0,X1] : (sP7('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f350,f2232,f534])).
fof(f7856,plain,(
  ( ! [X2,X0,X1] : (sP7('T_0','T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f350,f350,f2232,f534])).
fof(f7858,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP7('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),X5)))))) )),
  inference(unit_resulting_resolution,[],[f350,f2232,f2232,f534])).
fof(f7863,plain,(
  ( ! [X2,X0,X1] : (sP7('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'T_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f350,f2232,f534])).
fof(f7864,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP7('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),X5)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f2232,f350,f2232,f534])).
fof(f7865,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP7('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'T_0','const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),X5)))))) )),
  inference(unit_resulting_resolution,[],[f350,f2232,f2232,f534])).
fof(f7866,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (sP7('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),X5)))),'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),'const.pred_set.INSERT_2'(s(X6,X7),s(fun(X6,bool),X8)))))) )),
  inference(unit_resulting_resolution,[],[f2232,f2232,f2232,f534])).
fof(f7869,plain,(
  ( ! [X2,X0,X1] : (~sP7('F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f585,f2232,f537])).
fof(f7876,plain,(
  ( ! [X2,X0,X1] : (~sP7('F_0','T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f350,f585,f2232,f537])).
fof(f7877,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~sP7('F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),X5)))))) )),
  inference(unit_resulting_resolution,[],[f2232,f585,f2232,f537])).
fof(f7883,plain,(
  ( ! [X2,X0,X3,X1] : (~sP9('F_0',X0,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.INSERT_2'(s(X1,X2),s(fun(X1,bool),X3)))))) )),
  inference(unit_resulting_resolution,[],[f585,f2232,f538])).
fof(f7896,plain,(
  ( ! [X2,X0,X3,X1] : (~sP9('F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),X3)) )),
  inference(unit_resulting_resolution,[],[f585,f2232,f539])).
fof(f7919,plain,(
  ( ! [X2,X0,X1] : (~sP8('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f585,f2232,f545])).
fof(f7928,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP8(X0,X1,'const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.pred_set.INSERT_2'(s(X2,X3),s(fun(X2,bool),X4)))))) )),
  inference(unit_resulting_resolution,[],[f2232,f546])).
fof(f7929,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP8(X0,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.INSERT_2'(s(X1,X2),s(fun(X1,bool),X3)))),X4)) )),
  inference(unit_resulting_resolution,[],[f2232,f547])).
fof(f7930,plain,(
  ( ! [X2,X0,X3,X1] : (sP9('T_0',X0,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.INSERT_2'(s(X1,X2),s(fun(X1,bool),X3)))))) )),
  inference(unit_resulting_resolution,[],[f350,f2232,f550])).
fof(f7937,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (sP9('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),X3,'const.bool.IN_2'(s(X4,X5),s(fun(X4,bool),'const.pred_set.INSERT_2'(s(X4,X5),s(fun(X4,bool),X6)))))) )),
  inference(unit_resulting_resolution,[],[f2232,f2232,f550])).
fof(f7939,plain,(
  ( ! [X2,X0,X3,X1] : (sP9('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),X3)) )),
  inference(unit_resulting_resolution,[],[f350,f2232,f551])).
fof(f7946,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (sP9('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),X5)))),X6)) )),
  inference(unit_resulting_resolution,[],[f2232,f2232,f551])).
fof(f7948,plain,(
  ( ! [X2,X0,X1] : (~sP9('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f585,f2232,f552])).
fof(f7951,plain,(
  ( ! [X2,X0,X3,X1] : (~sP11('F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),X3)) )),
  inference(unit_resulting_resolution,[],[f585,f2232,f556])).
fof(f7999,plain,(
  ( ! [X2,X0,X1] : (sP11('const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)),'F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f585,f627,f2232,f561])).
fof(f8000,plain,(
  ( ! [X2,X0,X3,X1] : (sP11('const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'F_0','const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.INSERT_2'(s(X1,X2),s(fun(X1,bool),X3)))))) )),
  inference(unit_resulting_resolution,[],[f585,f636,f2232,f561])).
fof(f8001,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP11('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0','const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.pred_set.INSERT_2'(s(X2,X3),s(fun(X2,bool),X4)))))) )),
  inference(unit_resulting_resolution,[],[f585,f731,f2232,f561])).
fof(f8002,plain,(
  ( ! [X2,X0,X1] : (~sP10('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f585,f350,f2232,f562])).
fof(f8014,plain,(
  ( ! [X2,X0,X1] : (~sP10('T_0','F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f585,f350,f2232,f562])).
fof(f8015,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~sP10('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'F_0','const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),X5)))))) )),
  inference(unit_resulting_resolution,[],[f585,f2232,f2232,f562])).
fof(f8026,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(X0,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.INSERT_2'(s(X1,X2),s(fun(X1,bool),X3)))),X4)) )),
  inference(unit_resulting_resolution,[],[f2232,f564])).
fof(f8027,plain,(
  ( ! [X2,X0,X1] : (sP11('F_0','F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f585,f585,f2232,f566])).
fof(f8034,plain,(
  ( ! [X2,X0,X3,X1] : (sP11('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),X3)) )),
  inference(unit_resulting_resolution,[],[f350,f2232,f568])).
fof(f8039,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (sP11('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),X5)))),X6)) )),
  inference(unit_resulting_resolution,[],[f2232,f2232,f568])).
fof(f8041,plain,(
  ( ! [X2,X0,X1] : (~sP11('T_0','F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f350,f585,f2232,f569])).
fof(f8049,plain,(
  ( ! [X2,X0,X1] : (~sP11('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f585,f2232,f569])).
fof(f8050,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~sP11('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'F_0','const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),X5)))))) )),
  inference(unit_resulting_resolution,[],[f2232,f585,f2232,f569])).
fof(f8054,plain,(
  ( ! [X2,X0,X3,X1] : (sP11('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),X3,'F_0')) )),
  inference(unit_resulting_resolution,[],[f622,f2232,f572])).
fof(f8055,plain,(
  ( ! [X2,X0,X3,X1] : (sP11('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),X3,'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)))) )),
  inference(unit_resulting_resolution,[],[f623,f2232,f572])).
fof(f8056,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP11('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),X3,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f635,f2232,f572])).
fof(f8057,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP11('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),X3,'const.bool.IN_2'(s(X4,X5),s(fun(X4,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f730,f2232,f572])).
fof(f8058,plain,(
  ( ! [X2,X0,X3,X1] : (sP11('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'T_0',X3)) )),
  inference(unit_resulting_resolution,[],[f624,f2232,f572])).
fof(f8059,plain,(
  ( ! [X2,X0,X3,X1] : (sP11('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X3)) )),
  inference(unit_resulting_resolution,[],[f625,f2232,f572])).
fof(f8060,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP11('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))),X4)) )),
  inference(unit_resulting_resolution,[],[f949,f2232,f572])).
fof(f8061,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP11('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X3)))),X4)) )),
  inference(unit_resulting_resolution,[],[f2464,f2232,f572])).
fof(f8063,plain,(
  ( ! [X2,X0,X3,X1] : (sP9('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),X3,'T_0')) )),
  inference(unit_resulting_resolution,[],[f616,f2232,f573])).
fof(f8064,plain,(
  ( ! [X2,X0,X3,X1] : (sP9('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),X3,'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) )),
  inference(unit_resulting_resolution,[],[f617,f2232,f573])).
fof(f8065,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP9('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),X3,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X4)))))) )),
  inference(unit_resulting_resolution,[],[f933,f2232,f573])).
fof(f8066,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP9('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),X3,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X4)))))) )),
  inference(unit_resulting_resolution,[],[f2413,f2232,f573])).
fof(f8067,plain,(
  ( ! [X2,X0,X3,X1] : (sP9('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'T_0',X3)) )),
  inference(unit_resulting_resolution,[],[f618,f2232,f573])).
fof(f8068,plain,(
  ( ! [X2,X0,X3,X1] : (sP9('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X3)) )),
  inference(unit_resulting_resolution,[],[f619,f2232,f573])).
fof(f8069,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP9('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))),X4)) )),
  inference(unit_resulting_resolution,[],[f934,f2232,f573])).
fof(f8070,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP9('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X3)))),X4)) )),
  inference(unit_resulting_resolution,[],[f2414,f2232,f573])).
fof(f8071,plain,(
  ( ! [X2,X0,X3,X1] : (sP9('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X0,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.INSERT_2'(s(X1,X2),s(fun(X1,bool),X3)))))) )),
  inference(unit_resulting_resolution,[],[f2232,f806])).
fof(f8072,plain,(
  ( ! [X2,X0,X3,X1] : (sP9('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),X3)) )),
  inference(unit_resulting_resolution,[],[f2232,f816])).
fof(f8073,plain,(
  ( ! [X2,X0,X3,X1] : (sP11('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),X3)) )),
  inference(unit_resulting_resolution,[],[f2232,f848])).
fof(f8074,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP9('const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')),X2,'const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),X5)))))) )),
  inference(unit_resulting_resolution,[],[f2232,f1189])).
fof(f8075,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP9('const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')),'const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.pred_set.INSERT_2'(s(X2,X3),s(fun(X2,bool),X4)))),X5)) )),
  inference(unit_resulting_resolution,[],[f2232,f1190])).
fof(f8076,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP11('const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')),'const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.pred_set.INSERT_2'(s(X2,X3),s(fun(X2,bool),X4)))),X5)) )),
  inference(unit_resulting_resolution,[],[f2232,f1192])).
fof(f8077,plain,(
  ( ! [X2,X0,X1] : (~sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f2232,f1216])).
fof(f8078,plain,(
  ( ! [X2,X0,X1] : (~sP4('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f2232,f1241])).
fof(f8083,plain,(
  ( ! [X2,X0,X1] : (~sP3('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f2232,f1268])).
fof(f8088,plain,(
  ( ! [X2,X0,X1] : (~sP6('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f2232,f1318])).
fof(f8093,plain,(
  ( ! [X2,X0,X1] : (~sP7('F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) )),
  inference(unit_resulting_resolution,[],[f585,f2232,f1353])).
fof(f8094,plain,(
  ( ! [X2,X0,X1] : (~sP10('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'F_0','const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) )),
  inference(unit_resulting_resolution,[],[f585,f2232,f1527])).
fof(f8099,plain,(
  ( ! [X2,X0,X1] : (~sP11('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f585,f2232,f1570])).
fof(f8103,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP0(X0,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.INSERT_2'(s(X1,X2),s(fun(X1,bool),X3))))),s('type.list.list'(bool),X4)))))) )),
  inference(unit_resulting_resolution,[],[f2232,f1693])).
fof(f8104,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP0('const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2))))),s('type.list.list'(bool),X3)))),X4)) )),
  inference(unit_resulting_resolution,[],[f2232,f1694])).
fof(f8105,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP9('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),X1,'const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.pred_set.INSERT_2'(s(X2,X3),s(fun(X2,bool),X4)))))) )),
  inference(unit_resulting_resolution,[],[f2232,f4197])).
fof(f8106,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP9('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.INSERT_2'(s(X1,X2),s(fun(X1,bool),X3)))),X4)) )),
  inference(unit_resulting_resolution,[],[f2232,f4198])).
fof(f8107,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP11('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.INSERT_2'(s(X1,X2),s(fun(X1,bool),X3)))),X4)) )),
  inference(unit_resulting_resolution,[],[f2232,f4202])).
fof(f7088,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))))) )),
  inference(unit_resulting_resolution,[],[f350,f2229,f454])).
fof(f7089,plain,(
  ( ! [X0,X1] : (s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0))))) = s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f2229,f2229,f454])).
fof(f7090,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f2229,f464])).
fof(f7091,plain,(
  ( ! [X0,X1] : (sP0('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),X1)) )),
  inference(unit_resulting_resolution,[],[f2229,f465])).
fof(f7092,plain,(
  ( ! [X0] : (~sP1('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f2229,f479])).
fof(f7094,plain,(
  ( ! [X0] : (~sP1('T_0','const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))))) )),
  inference(unit_resulting_resolution,[],[f350,f2229,f479])).
fof(f7095,plain,(
  ( ! [X0,X1] : (~sP1('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f2229,f2229,f479])).
fof(f7096,plain,(
  ( ! [X0] : (sP1('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f2229,f480])).
fof(f7097,plain,(
  ( ! [X0] : (sP1('F_0','const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))))) )),
  inference(unit_resulting_resolution,[],[f585,f2229,f483])).
fof(f7100,plain,(
  ( ! [X0] : (~sP5('T_0','const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f350,f585,f2229,f501])).
fof(f7101,plain,(
  ( ! [X0,X1] : (~sP5('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X1)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f2229,f585,f2229,f501])).
fof(f7119,plain,(
  ( ! [X0] : (sP5('T_0','T_0','const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))))) )),
  inference(unit_resulting_resolution,[],[f598,f592,f606,f2229,f502])).
fof(f7120,plain,(
  ( ! [X0] : (sP5('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'T_0','const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))))) )),
  inference(unit_resulting_resolution,[],[f599,f592,f606,f2229,f502])).
fof(f7121,plain,(
  ( ! [X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'T_0','const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f918,f592,f606,f2229,f502])).
fof(f7122,plain,(
  ( ! [X0,X1] : (sP5('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'T_0','const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f2325,f592,f606,f2229,f502])).
fof(f7123,plain,(
  ( ! [X0] : (sP5('T_0','const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))))) )),
  inference(unit_resulting_resolution,[],[f598,f593,f607,f2229,f502])).
fof(f7124,plain,(
  ( ! [X0] : (sP5('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))))) )),
  inference(unit_resulting_resolution,[],[f599,f593,f607,f2229,f502])).
fof(f7125,plain,(
  ( ! [X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f918,f593,f607,f2229,f502])).
fof(f7126,plain,(
  ( ! [X0,X1] : (sP5('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f2325,f593,f607,f2229,f502])).
fof(f7127,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f598,f917,f920,f2229,f502])).
fof(f7128,plain,(
  ( ! [X0,X1] : (sP5('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f599,f917,f920,f2229,f502])).
fof(f7129,plain,(
  ( ! [X2,X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))),'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f918,f917,f920,f2229,f502])).
fof(f7130,plain,(
  ( ! [X2,X0,X1] : (sP5('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))),'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f2325,f917,f920,f2229,f502])).
fof(f7131,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f598,f2310,f2327,f2229,f502])).
fof(f7132,plain,(
  ( ! [X0,X1] : (sP5('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f599,f2310,f2327,f2229,f502])).
fof(f7133,plain,(
  ( ! [X2,X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))),'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f918,f2310,f2327,f2229,f502])).
fof(f7134,plain,(
  ( ! [X2,X0,X1] : (sP5('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))),'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f2325,f2310,f2327,f2229,f502])).
fof(f7135,plain,(
  ( ! [X0] : (~sP4('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f350,f585,f2229,f505])).
fof(f7147,plain,(
  ( ! [X0] : (~sP4('T_0','const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f350,f585,f2229,f505])).
fof(f7148,plain,(
  ( ! [X0,X1] : (~sP4('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X1)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f2229,f585,f2229,f505])).
fof(f7159,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f2229,f506])).
fof(f7160,plain,(
  ( ! [X0] : (~sP3('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f350,f585,f2229,f509])).
fof(f7172,plain,(
  ( ! [X0] : (~sP3('T_0','const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f350,f585,f2229,f509])).
fof(f7173,plain,(
  ( ! [X0,X1] : (~sP3('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X1)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f2229,f585,f2229,f509])).
fof(f7184,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f2229,f510])).
fof(f7185,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f2229,f514])).
fof(f7186,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X1)))),X2)) )),
  inference(unit_resulting_resolution,[],[f2229,f515])).
fof(f7187,plain,(
  ( ! [X2,X0,X1] : (sP2('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),X1,X2)) )),
  inference(unit_resulting_resolution,[],[f2229,f516])).
fof(f7188,plain,(
  ( ! [X0] : (sP5('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f647,f2229,f518])).
fof(f7189,plain,(
  ( ! [X0] : (sP5('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f667,f2229,f518])).
fof(f7190,plain,(
  ( ! [X2,X0,X1] : (sP5('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'const.list.EVERY_2'(s(fun(X1,bool),X2),s('type.list.list'(X1),'const.list.NIL_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f408,f2229,f518])).
fof(f7191,plain,(
  ( ! [X2,X0,X1] : (sP5('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2)))),'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X2)))) )),
  inference(unit_resulting_resolution,[],[f449,f2229,f518])).
fof(f7192,plain,(
  ( ! [X0,X1] : (sP5('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f668,f2229,f518])).
fof(f7193,plain,(
  ( ! [X0,X1] : (sP5('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f907,f2229,f518])).
fof(f7194,plain,(
  ( ! [X0,X1] : (sP5('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f2274,f2229,f518])).
fof(f7195,plain,(
  ( ! [X2,X0,X1] : (sP5('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f716,f2229,f518])).
fof(f7196,plain,(
  ( ! [X0,X1] : (sP5('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(bool),X1)),'const.list.HD_1'(s('type.list.list'(bool),X1)))) )),
  inference(unit_resulting_resolution,[],[f409,f2229,f518])).
fof(f7197,plain,(
  ( ! [X2,X0,X1] : (sP5('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X1),s('type.list.list'(bool),X2)))),X1)) )),
  inference(unit_resulting_resolution,[],[f419,f2229,f518])).
fof(f7198,plain,(
  ( ! [X2,X0,X1] : (sP5('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'const.list.EL_2'(s('type.num.num',X1),s('type.list.list'(bool),'const.list.TL_1'(s('type.list.list'(bool),X2)))),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.list.list'(bool),X2)))) )),
  inference(unit_resulting_resolution,[],[f422,f2229,f518])).
fof(f7199,plain,(
  ( ! [X0] : (sP5('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'T_0','const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) )),
  inference(unit_resulting_resolution,[],[f647,f2229,f518])).
fof(f7200,plain,(
  ( ! [X0] : (sP5('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'F_0','const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)))) )),
  inference(unit_resulting_resolution,[],[f667,f2229,f518])).
fof(f7201,plain,(
  ( ! [X2,X0,X1] : (sP5('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'T_0','const.list.EVERY_2'(s(fun(X1,bool),X2),s('type.list.list'(X1),'const.list.NIL_0')))) )),
  inference(unit_resulting_resolution,[],[f408,f2229,f518])).
fof(f7202,plain,(
  ( ! [X2,X0,X1] : (sP5('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X2)),'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2)))))) )),
  inference(unit_resulting_resolution,[],[f449,f2229,f518])).
fof(f7203,plain,(
  ( ! [X0,X1] : (sP5('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'F_0','const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f668,f2229,f518])).
fof(f7204,plain,(
  ( ! [X0,X1] : (sP5('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'T_0','const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f907,f2229,f518])).
fof(f7205,plain,(
  ( ! [X0,X1] : (sP5('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'T_0','const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f2274,f2229,f518])).
fof(f7206,plain,(
  ( ! [X2,X0,X1] : (sP5('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'F_0','const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f716,f2229,f518])).
fof(f7207,plain,(
  ( ! [X0,X1] : (sP5('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'const.list.HD_1'(s('type.list.list'(bool),X1)),'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(bool),X1)))) )),
  inference(unit_resulting_resolution,[],[f409,f2229,f518])).
fof(f7208,plain,(
  ( ! [X2,X0,X1] : (sP5('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),X1,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X1),s('type.list.list'(bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f419,f2229,f518])).
fof(f7209,plain,(
  ( ! [X2,X0,X1] : (sP5('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.list.list'(bool),X2)),'const.list.EL_2'(s('type.num.num',X1),s('type.list.list'(bool),'const.list.TL_1'(s('type.list.list'(bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f422,f2229,f518])).
fof(f7210,plain,(
  ( ! [X0] : (~sP5('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f642,f2229,f519])).
fof(f7211,plain,(
  ( ! [X0] : (~sP5('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'T_0','const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)))) )),
  inference(unit_resulting_resolution,[],[f643,f2229,f519])).
fof(f7212,plain,(
  ( ! [X0,X1] : (~sP5('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'T_0','const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f644,f2229,f519])).
fof(f7213,plain,(
  ( ! [X2,X0,X1] : (~sP5('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'T_0','const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f715,f2229,f519])).
fof(f7214,plain,(
  ( ! [X0] : (~sP5('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f642,f2229,f519])).
fof(f7215,plain,(
  ( ! [X0] : (~sP5('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f643,f2229,f519])).
fof(f7216,plain,(
  ( ! [X0,X1] : (~sP5('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f644,f2229,f519])).
fof(f7217,plain,(
  ( ! [X2,X0,X1] : (~sP5('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f715,f2229,f519])).
fof(f7218,plain,(
  ( ! [X0,X1] : (~sP7('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f2229,f522])).
fof(f7219,plain,(
  ( ! [X0,X1] : (~sP7('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f585,f2229,f523])).
fof(f7222,plain,(
  ( ! [X0] : (sP7('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f611,f2229,f524])).
fof(f7224,plain,(
  ( ! [X0,X1] : (sP7('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X1)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f924,f2229,f524])).
fof(f7226,plain,(
  ( ! [X0,X1] : (sP7('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X1)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f2371,f2229,f524])).
fof(f7233,plain,(
  ( ! [X0] : (sP7('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'T_0','const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))))) )),
  inference(unit_resulting_resolution,[],[f350,f611,f2229,f524])).
fof(f7234,plain,(
  ( ! [X0,X1] : (sP7('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f2229,f611,f2229,f524])).
fof(f7235,plain,(
  ( ! [X0,X1] : (sP7('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'T_0','const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f350,f924,f2229,f524])).
fof(f7236,plain,(
  ( ! [X2,X0,X1] : (sP7('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X1)))),'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f2229,f924,f2229,f524])).
fof(f7237,plain,(
  ( ! [X0,X1] : (sP7('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'T_0','const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f350,f2371,f2229,f524])).
fof(f7238,plain,(
  ( ! [X2,X0,X1] : (sP7('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X1)))),'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f2229,f2371,f2229,f524])).
fof(f7242,plain,(
  ( ! [X0] : (~sP6('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f350,f585,f2229,f528])).
fof(f7254,plain,(
  ( ! [X0] : (~sP6('T_0','const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f350,f585,f2229,f528])).
fof(f7255,plain,(
  ( ! [X0,X1] : (~sP6('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X1)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f2229,f585,f2229,f528])).
fof(f7266,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f2229,f529])).
fof(f7267,plain,(
  ( ! [X0] : (sP7('T_0','const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f350,f2229,f534])).
fof(f7272,plain,(
  ( ! [X0] : (sP7('T_0','T_0','const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))))) )),
  inference(unit_resulting_resolution,[],[f350,f350,f2229,f534])).
fof(f7274,plain,(
  ( ! [X0,X1] : (sP7('T_0','const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f350,f2229,f2229,f534])).
fof(f7277,plain,(
  ( ! [X0] : (sP7('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'T_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f350,f2229,f534])).
fof(f7278,plain,(
  ( ! [X0,X1] : (sP7('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X1)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f2229,f350,f2229,f534])).
fof(f7279,plain,(
  ( ! [X0,X1] : (sP7('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'T_0','const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f350,f2229,f2229,f534])).
fof(f7280,plain,(
  ( ! [X2,X0,X1] : (sP7('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X1)))),'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f2229,f2229,f2229,f534])).
fof(f7281,plain,(
  ( ! [X0] : (~sP7('F_0','const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f585,f2229,f537])).
fof(f7287,plain,(
  ( ! [X0] : (~sP7('F_0','T_0','const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))))) )),
  inference(unit_resulting_resolution,[],[f350,f585,f2229,f537])).
fof(f7288,plain,(
  ( ! [X0,X1] : (~sP7('F_0','const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f2229,f585,f2229,f537])).
fof(f7293,plain,(
  ( ! [X0,X1] : (~sP9('F_0',X0,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f585,f2229,f538])).
fof(f7304,plain,(
  ( ! [X0,X1] : (~sP9('F_0','const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),X1)) )),
  inference(unit_resulting_resolution,[],[f585,f2229,f539])).
fof(f7323,plain,(
  ( ! [X0] : (~sP8('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f585,f2229,f545])).
fof(f7332,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f2229,f546])).
fof(f7333,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X1)))),X2)) )),
  inference(unit_resulting_resolution,[],[f2229,f547])).
fof(f7334,plain,(
  ( ! [X0,X1] : (sP9('T_0',X0,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f350,f2229,f550])).
fof(f7340,plain,(
  ( ! [X2,X0,X1] : (sP9('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),X1,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f2229,f2229,f550])).
fof(f7341,plain,(
  ( ! [X0,X1] : (sP9('T_0','const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),X1)) )),
  inference(unit_resulting_resolution,[],[f350,f2229,f551])).
fof(f7347,plain,(
  ( ! [X2,X0,X1] : (sP9('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X1)))),X2)) )),
  inference(unit_resulting_resolution,[],[f2229,f2229,f551])).
fof(f7348,plain,(
  ( ! [X0] : (~sP9('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f585,f2229,f552])).
fof(f7349,plain,(
  ( ! [X0,X1] : (~sP11('F_0','const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),X1)) )),
  inference(unit_resulting_resolution,[],[f585,f2229,f556])).
fof(f7392,plain,(
  ( ! [X0] : (sP11('const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)),'F_0','const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))))) )),
  inference(unit_resulting_resolution,[],[f585,f627,f2229,f561])).
fof(f7393,plain,(
  ( ! [X0,X1] : (sP11('const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'F_0','const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f585,f636,f2229,f561])).
fof(f7394,plain,(
  ( ! [X2,X0,X1] : (sP11('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0','const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f585,f731,f2229,f561])).
fof(f7395,plain,(
  ( ! [X0] : (~sP10('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f585,f350,f2229,f562])).
fof(f7407,plain,(
  ( ! [X0] : (~sP10('T_0','F_0','const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))))) )),
  inference(unit_resulting_resolution,[],[f585,f350,f2229,f562])).
fof(f7408,plain,(
  ( ! [X0,X1] : (~sP10('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'F_0','const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f585,f2229,f2229,f562])).
fof(f7419,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X1)))),X2)) )),
  inference(unit_resulting_resolution,[],[f2229,f564])).
fof(f7420,plain,(
  ( ! [X0] : (sP11('F_0','F_0','const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))))) )),
  inference(unit_resulting_resolution,[],[f585,f585,f2229,f566])).
fof(f7425,plain,(
  ( ! [X0,X1] : (sP11('T_0','const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),X1)) )),
  inference(unit_resulting_resolution,[],[f350,f2229,f568])).
fof(f7429,plain,(
  ( ! [X2,X0,X1] : (sP11('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X1)))),X2)) )),
  inference(unit_resulting_resolution,[],[f2229,f2229,f568])).
fof(f7430,plain,(
  ( ! [X0] : (~sP11('T_0','F_0','const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))))) )),
  inference(unit_resulting_resolution,[],[f350,f585,f2229,f569])).
fof(f7436,plain,(
  ( ! [X0] : (~sP11('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f585,f2229,f569])).
fof(f7437,plain,(
  ( ! [X0,X1] : (~sP11('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'F_0','const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f2229,f585,f2229,f569])).
fof(f7438,plain,(
  ( ! [X0,X1] : (sP11('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f622,f2229,f572])).
fof(f7439,plain,(
  ( ! [X0,X1] : (sP11('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),X1,'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)))) )),
  inference(unit_resulting_resolution,[],[f623,f2229,f572])).
fof(f7440,plain,(
  ( ! [X2,X0,X1] : (sP11('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),X1,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f635,f2229,f572])).
fof(f7441,plain,(
  ( ! [X2,X0,X3,X1] : (sP11('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),X1,'const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f730,f2229,f572])).
fof(f7442,plain,(
  ( ! [X0,X1] : (sP11('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f624,f2229,f572])).
fof(f7443,plain,(
  ( ! [X0,X1] : (sP11('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X1)) )),
  inference(unit_resulting_resolution,[],[f625,f2229,f572])).
fof(f7444,plain,(
  ( ! [X2,X0,X1] : (sP11('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))),X2)) )),
  inference(unit_resulting_resolution,[],[f949,f2229,f572])).
fof(f7445,plain,(
  ( ! [X2,X0,X1] : (sP11('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))),X2)) )),
  inference(unit_resulting_resolution,[],[f2464,f2229,f572])).
fof(f7446,plain,(
  ( ! [X0,X1] : (sP9('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f616,f2229,f573])).
fof(f7447,plain,(
  ( ! [X0,X1] : (sP9('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),X1,'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) )),
  inference(unit_resulting_resolution,[],[f617,f2229,f573])).
fof(f7448,plain,(
  ( ! [X2,X0,X1] : (sP9('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),X1,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2)))))) )),
  inference(unit_resulting_resolution,[],[f933,f2229,f573])).
fof(f7449,plain,(
  ( ! [X2,X0,X1] : (sP9('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),X1,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f2413,f2229,f573])).
fof(f7450,plain,(
  ( ! [X0,X1] : (sP9('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f618,f2229,f573])).
fof(f7451,plain,(
  ( ! [X0,X1] : (sP9('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X1)) )),
  inference(unit_resulting_resolution,[],[f619,f2229,f573])).
fof(f7452,plain,(
  ( ! [X2,X0,X1] : (sP9('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))),X2)) )),
  inference(unit_resulting_resolution,[],[f934,f2229,f573])).
fof(f7453,plain,(
  ( ! [X2,X0,X1] : (sP9('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))),X2)) )),
  inference(unit_resulting_resolution,[],[f2414,f2229,f573])).
fof(f7454,plain,(
  ( ! [X0,X1] : (sP9('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X0,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f2229,f806])).
fof(f7455,plain,(
  ( ! [X0,X1] : (sP9('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),X1)) )),
  inference(unit_resulting_resolution,[],[f2229,f816])).
fof(f7456,plain,(
  ( ! [X0,X1] : (sP11('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),X1)) )),
  inference(unit_resulting_resolution,[],[f2229,f848])).
fof(f7457,plain,(
  ( ! [X2,X0,X3,X1] : (sP9('const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')),X2,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X3)))))) )),
  inference(unit_resulting_resolution,[],[f2229,f1189])).
fof(f7458,plain,(
  ( ! [X2,X0,X3,X1] : (sP9('const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')),'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X2)))),X3)) )),
  inference(unit_resulting_resolution,[],[f2229,f1190])).
fof(f7459,plain,(
  ( ! [X2,X0,X3,X1] : (sP11('const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')),'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X2)))),X3)) )),
  inference(unit_resulting_resolution,[],[f2229,f1192])).
fof(f7460,plain,(
  ( ! [X0] : (~sP5('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f2229,f1216])).
fof(f7461,plain,(
  ( ! [X0] : (~sP4('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f2229,f1241])).
fof(f7466,plain,(
  ( ! [X0] : (~sP3('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f2229,f1268])).
fof(f7471,plain,(
  ( ! [X0] : (~sP6('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f2229,f1318])).
fof(f7476,plain,(
  ( ! [X0] : (~sP7('F_0','const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) )),
  inference(unit_resulting_resolution,[],[f585,f2229,f1353])).
fof(f7477,plain,(
  ( ! [X0] : (~sP10('const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))),'F_0','const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) )),
  inference(unit_resulting_resolution,[],[f585,f2229,f1527])).
fof(f7482,plain,(
  ( ! [X0] : (~sP11('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'F_0','const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))))) )),
  inference(unit_resulting_resolution,[],[f585,f2229,f1570])).
fof(f7486,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X1))))),s('type.list.list'(bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f2229,f1693])).
fof(f7487,plain,(
  ( ! [X2,X0,X1] : (sP0('const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0))))),s('type.list.list'(bool),X1)))),X2)) )),
  inference(unit_resulting_resolution,[],[f2229,f1694])).
fof(f7488,plain,(
  ( ! [X2,X0,X1] : (sP9('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),X1,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f2229,f4197])).
fof(f7489,plain,(
  ( ! [X2,X0,X1] : (sP9('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X1)))),X2)) )),
  inference(unit_resulting_resolution,[],[f2229,f4198])).
fof(f7490,plain,(
  ( ! [X2,X0,X1] : (sP11('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X1)))),X2)) )),
  inference(unit_resulting_resolution,[],[f2229,f4202])).
fof(f7063,plain,(
  ( ! [X0] : (~sP3('T_0','const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f606,f592,f5627,f504])).
fof(f6556,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))))) )),
  inference(unit_resulting_resolution,[],[f350,f2226,f454])).
fof(f6557,plain,(
  ( ! [X0,X1] : (s(bool,'const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0))))) = s(bool,'const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f2226,f2226,f454])).
fof(f6558,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f2226,f464])).
fof(f6559,plain,(
  ( ! [X0,X1] : (sP0('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),X1)) )),
  inference(unit_resulting_resolution,[],[f2226,f465])).
fof(f6560,plain,(
  ( ! [X0] : (~sP1('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f2226,f479])).
fof(f6562,plain,(
  ( ! [X0] : (~sP1('T_0','const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))))) )),
  inference(unit_resulting_resolution,[],[f350,f2226,f479])).
fof(f6563,plain,(
  ( ! [X0,X1] : (~sP1('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f2226,f2226,f479])).
fof(f6564,plain,(
  ( ! [X0] : (sP1('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f2226,f480])).
fof(f6565,plain,(
  ( ! [X0] : (sP1('F_0','const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))))) )),
  inference(unit_resulting_resolution,[],[f585,f2226,f483])).
fof(f6568,plain,(
  ( ! [X0] : (~sP5('T_0','const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f350,f585,f2226,f501])).
fof(f6569,plain,(
  ( ! [X0,X1] : (~sP5('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X1)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f2226,f585,f2226,f501])).
fof(f6587,plain,(
  ( ! [X0] : (sP5('T_0','T_0','const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))))) )),
  inference(unit_resulting_resolution,[],[f598,f592,f606,f2226,f502])).
fof(f6588,plain,(
  ( ! [X0] : (sP5('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'T_0','const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))))) )),
  inference(unit_resulting_resolution,[],[f599,f592,f606,f2226,f502])).
fof(f6589,plain,(
  ( ! [X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'T_0','const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f918,f592,f606,f2226,f502])).
fof(f6590,plain,(
  ( ! [X0,X1] : (sP5('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'T_0','const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f2325,f592,f606,f2226,f502])).
fof(f6591,plain,(
  ( ! [X0] : (sP5('T_0','const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))))) )),
  inference(unit_resulting_resolution,[],[f598,f593,f607,f2226,f502])).
fof(f6592,plain,(
  ( ! [X0] : (sP5('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))))) )),
  inference(unit_resulting_resolution,[],[f599,f593,f607,f2226,f502])).
fof(f6593,plain,(
  ( ! [X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f918,f593,f607,f2226,f502])).
fof(f6594,plain,(
  ( ! [X0,X1] : (sP5('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f2325,f593,f607,f2226,f502])).
fof(f6595,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f598,f917,f920,f2226,f502])).
fof(f6596,plain,(
  ( ! [X0,X1] : (sP5('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f599,f917,f920,f2226,f502])).
fof(f6597,plain,(
  ( ! [X2,X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))),'const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f918,f917,f920,f2226,f502])).
fof(f6598,plain,(
  ( ! [X2,X0,X1] : (sP5('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))),'const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f2325,f917,f920,f2226,f502])).
fof(f6599,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f598,f2310,f2327,f2226,f502])).
fof(f6600,plain,(
  ( ! [X0,X1] : (sP5('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f599,f2310,f2327,f2226,f502])).
fof(f6601,plain,(
  ( ! [X2,X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))),'const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f918,f2310,f2327,f2226,f502])).
fof(f6602,plain,(
  ( ! [X2,X0,X1] : (sP5('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))),'const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f2325,f2310,f2327,f2226,f502])).
fof(f6603,plain,(
  ( ! [X0] : (~sP4('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f350,f585,f2226,f505])).
fof(f6615,plain,(
  ( ! [X0] : (~sP4('T_0','const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f350,f585,f2226,f505])).
fof(f6616,plain,(
  ( ! [X0,X1] : (~sP4('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X1)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f2226,f585,f2226,f505])).
fof(f6627,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,'const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f2226,f506])).
fof(f6628,plain,(
  ( ! [X0] : (~sP3('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f350,f585,f2226,f509])).
fof(f6640,plain,(
  ( ! [X0] : (~sP3('T_0','const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f350,f585,f2226,f509])).
fof(f6641,plain,(
  ( ! [X0,X1] : (~sP3('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X1)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f2226,f585,f2226,f509])).
fof(f6652,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,'const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f2226,f510])).
fof(f6653,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,'const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f2226,f514])).
fof(f6654,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,'const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X1)))),X2)) )),
  inference(unit_resulting_resolution,[],[f2226,f515])).
fof(f6655,plain,(
  ( ! [X2,X0,X1] : (sP2('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),X1,X2)) )),
  inference(unit_resulting_resolution,[],[f2226,f516])).
fof(f6656,plain,(
  ( ! [X0] : (sP5('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f647,f2226,f518])).
fof(f6657,plain,(
  ( ! [X0] : (sP5('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f667,f2226,f518])).
fof(f6658,plain,(
  ( ! [X2,X0,X1] : (sP5('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'const.list.EVERY_2'(s(fun(X1,bool),X2),s('type.list.list'(X1),'const.list.NIL_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f408,f2226,f518])).
fof(f6659,plain,(
  ( ! [X2,X0,X1] : (sP5('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2)))),'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X2)))) )),
  inference(unit_resulting_resolution,[],[f449,f2226,f518])).
fof(f6660,plain,(
  ( ! [X0,X1] : (sP5('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f668,f2226,f518])).
fof(f6661,plain,(
  ( ! [X0,X1] : (sP5('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f907,f2226,f518])).
fof(f6662,plain,(
  ( ! [X0,X1] : (sP5('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f2274,f2226,f518])).
fof(f6663,plain,(
  ( ! [X2,X0,X1] : (sP5('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f716,f2226,f518])).
fof(f6664,plain,(
  ( ! [X0,X1] : (sP5('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(bool),X1)),'const.list.HD_1'(s('type.list.list'(bool),X1)))) )),
  inference(unit_resulting_resolution,[],[f409,f2226,f518])).
fof(f6665,plain,(
  ( ! [X2,X0,X1] : (sP5('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X1),s('type.list.list'(bool),X2)))),X1)) )),
  inference(unit_resulting_resolution,[],[f419,f2226,f518])).
fof(f6666,plain,(
  ( ! [X2,X0,X1] : (sP5('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'const.list.EL_2'(s('type.num.num',X1),s('type.list.list'(bool),'const.list.TL_1'(s('type.list.list'(bool),X2)))),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.list.list'(bool),X2)))) )),
  inference(unit_resulting_resolution,[],[f422,f2226,f518])).
fof(f6667,plain,(
  ( ! [X0] : (sP5('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'T_0','const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) )),
  inference(unit_resulting_resolution,[],[f647,f2226,f518])).
fof(f6668,plain,(
  ( ! [X0] : (sP5('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'F_0','const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)))) )),
  inference(unit_resulting_resolution,[],[f667,f2226,f518])).
fof(f6669,plain,(
  ( ! [X2,X0,X1] : (sP5('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'T_0','const.list.EVERY_2'(s(fun(X1,bool),X2),s('type.list.list'(X1),'const.list.NIL_0')))) )),
  inference(unit_resulting_resolution,[],[f408,f2226,f518])).
fof(f6670,plain,(
  ( ! [X2,X0,X1] : (sP5('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X2)),'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2)))))) )),
  inference(unit_resulting_resolution,[],[f449,f2226,f518])).
fof(f6671,plain,(
  ( ! [X0,X1] : (sP5('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'F_0','const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f668,f2226,f518])).
fof(f6672,plain,(
  ( ! [X0,X1] : (sP5('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'T_0','const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f907,f2226,f518])).
fof(f6673,plain,(
  ( ! [X0,X1] : (sP5('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'T_0','const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f2274,f2226,f518])).
fof(f6674,plain,(
  ( ! [X2,X0,X1] : (sP5('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'F_0','const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f716,f2226,f518])).
fof(f6675,plain,(
  ( ! [X0,X1] : (sP5('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'const.list.HD_1'(s('type.list.list'(bool),X1)),'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(bool),X1)))) )),
  inference(unit_resulting_resolution,[],[f409,f2226,f518])).
fof(f6676,plain,(
  ( ! [X2,X0,X1] : (sP5('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),X1,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X1),s('type.list.list'(bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f419,f2226,f518])).
fof(f6677,plain,(
  ( ! [X2,X0,X1] : (sP5('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.list.list'(bool),X2)),'const.list.EL_2'(s('type.num.num',X1),s('type.list.list'(bool),'const.list.TL_1'(s('type.list.list'(bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f422,f2226,f518])).
fof(f6678,plain,(
  ( ! [X0] : (~sP5('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f642,f2226,f519])).
fof(f6679,plain,(
  ( ! [X0] : (~sP5('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'T_0','const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)))) )),
  inference(unit_resulting_resolution,[],[f643,f2226,f519])).
fof(f6680,plain,(
  ( ! [X0,X1] : (~sP5('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'T_0','const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f644,f2226,f519])).
fof(f6681,plain,(
  ( ! [X2,X0,X1] : (~sP5('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'T_0','const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f715,f2226,f519])).
fof(f6682,plain,(
  ( ! [X0] : (~sP5('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f642,f2226,f519])).
fof(f6683,plain,(
  ( ! [X0] : (~sP5('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f643,f2226,f519])).
fof(f6684,plain,(
  ( ! [X0,X1] : (~sP5('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f644,f2226,f519])).
fof(f6685,plain,(
  ( ! [X2,X0,X1] : (~sP5('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f715,f2226,f519])).
fof(f6686,plain,(
  ( ! [X0,X1] : (~sP7('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f2226,f522])).
fof(f6687,plain,(
  ( ! [X0,X1] : (~sP7('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f585,f2226,f523])).
fof(f6690,plain,(
  ( ! [X0] : (sP7('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f611,f2226,f524])).
fof(f6692,plain,(
  ( ! [X0,X1] : (sP7('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X1)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f924,f2226,f524])).
fof(f6694,plain,(
  ( ! [X0,X1] : (sP7('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X1)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f2371,f2226,f524])).
fof(f6701,plain,(
  ( ! [X0] : (sP7('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'T_0','const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))))) )),
  inference(unit_resulting_resolution,[],[f350,f611,f2226,f524])).
fof(f6702,plain,(
  ( ! [X0,X1] : (sP7('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f2226,f611,f2226,f524])).
fof(f6703,plain,(
  ( ! [X0,X1] : (sP7('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'T_0','const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f350,f924,f2226,f524])).
fof(f6704,plain,(
  ( ! [X2,X0,X1] : (sP7('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X1)))),'const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f2226,f924,f2226,f524])).
fof(f6705,plain,(
  ( ! [X0,X1] : (sP7('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'T_0','const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f350,f2371,f2226,f524])).
fof(f6706,plain,(
  ( ! [X2,X0,X1] : (sP7('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X1)))),'const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f2226,f2371,f2226,f524])).
fof(f6710,plain,(
  ( ! [X0] : (~sP6('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f350,f585,f2226,f528])).
fof(f6722,plain,(
  ( ! [X0] : (~sP6('T_0','const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f350,f585,f2226,f528])).
fof(f6723,plain,(
  ( ! [X0,X1] : (~sP6('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X1)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f2226,f585,f2226,f528])).
fof(f6734,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,'const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f2226,f529])).
fof(f6735,plain,(
  ( ! [X0] : (sP7('T_0','const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f350,f2226,f534])).
fof(f6740,plain,(
  ( ! [X0] : (sP7('T_0','T_0','const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))))) )),
  inference(unit_resulting_resolution,[],[f350,f350,f2226,f534])).
fof(f6742,plain,(
  ( ! [X0,X1] : (sP7('T_0','const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f350,f2226,f2226,f534])).
fof(f6745,plain,(
  ( ! [X0] : (sP7('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'T_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f350,f2226,f534])).
fof(f6746,plain,(
  ( ! [X0,X1] : (sP7('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X1)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f2226,f350,f2226,f534])).
fof(f6747,plain,(
  ( ! [X0,X1] : (sP7('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'T_0','const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f350,f2226,f2226,f534])).
fof(f6748,plain,(
  ( ! [X2,X0,X1] : (sP7('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X1)))),'const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f2226,f2226,f2226,f534])).
fof(f6749,plain,(
  ( ! [X0] : (~sP7('F_0','const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f585,f2226,f537])).
fof(f6755,plain,(
  ( ! [X0] : (~sP7('F_0','T_0','const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))))) )),
  inference(unit_resulting_resolution,[],[f350,f585,f2226,f537])).
fof(f6756,plain,(
  ( ! [X0,X1] : (~sP7('F_0','const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f2226,f585,f2226,f537])).
fof(f6761,plain,(
  ( ! [X0,X1] : (~sP9('F_0',X0,'const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f585,f2226,f538])).
fof(f6772,plain,(
  ( ! [X0,X1] : (~sP9('F_0','const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),X1)) )),
  inference(unit_resulting_resolution,[],[f585,f2226,f539])).
fof(f6791,plain,(
  ( ! [X0] : (~sP8('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f585,f2226,f545])).
fof(f6800,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,'const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f2226,f546])).
fof(f6801,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,'const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X1)))),X2)) )),
  inference(unit_resulting_resolution,[],[f2226,f547])).
fof(f6802,plain,(
  ( ! [X0,X1] : (sP9('T_0',X0,'const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f350,f2226,f550])).
fof(f6808,plain,(
  ( ! [X2,X0,X1] : (sP9('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),X1,'const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f2226,f2226,f550])).
fof(f6809,plain,(
  ( ! [X0,X1] : (sP9('T_0','const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),X1)) )),
  inference(unit_resulting_resolution,[],[f350,f2226,f551])).
fof(f6815,plain,(
  ( ! [X2,X0,X1] : (sP9('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X1)))),X2)) )),
  inference(unit_resulting_resolution,[],[f2226,f2226,f551])).
fof(f6816,plain,(
  ( ! [X0] : (~sP9('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f585,f2226,f552])).
fof(f6817,plain,(
  ( ! [X0,X1] : (~sP11('F_0','const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),X1)) )),
  inference(unit_resulting_resolution,[],[f585,f2226,f556])).
fof(f6860,plain,(
  ( ! [X0] : (sP11('const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)),'F_0','const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))))) )),
  inference(unit_resulting_resolution,[],[f585,f627,f2226,f561])).
fof(f6861,plain,(
  ( ! [X0,X1] : (sP11('const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'F_0','const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f585,f636,f2226,f561])).
fof(f6862,plain,(
  ( ! [X2,X0,X1] : (sP11('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0','const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f585,f731,f2226,f561])).
fof(f6863,plain,(
  ( ! [X0] : (~sP10('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f585,f350,f2226,f562])).
fof(f6875,plain,(
  ( ! [X0] : (~sP10('T_0','F_0','const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))))) )),
  inference(unit_resulting_resolution,[],[f585,f350,f2226,f562])).
fof(f6876,plain,(
  ( ! [X0,X1] : (~sP10('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'F_0','const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f585,f2226,f2226,f562])).
fof(f6887,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,'const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X1)))),X2)) )),
  inference(unit_resulting_resolution,[],[f2226,f564])).
fof(f6888,plain,(
  ( ! [X0] : (sP11('F_0','F_0','const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))))) )),
  inference(unit_resulting_resolution,[],[f585,f585,f2226,f566])).
fof(f6893,plain,(
  ( ! [X0,X1] : (sP11('T_0','const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),X1)) )),
  inference(unit_resulting_resolution,[],[f350,f2226,f568])).
fof(f6897,plain,(
  ( ! [X2,X0,X1] : (sP11('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X1)))),X2)) )),
  inference(unit_resulting_resolution,[],[f2226,f2226,f568])).
fof(f6898,plain,(
  ( ! [X0] : (~sP11('T_0','F_0','const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))))) )),
  inference(unit_resulting_resolution,[],[f350,f585,f2226,f569])).
fof(f6904,plain,(
  ( ! [X0] : (~sP11('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f585,f2226,f569])).
fof(f6905,plain,(
  ( ! [X0,X1] : (~sP11('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'F_0','const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f2226,f585,f2226,f569])).
fof(f6906,plain,(
  ( ! [X0,X1] : (sP11('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f622,f2226,f572])).
fof(f6907,plain,(
  ( ! [X0,X1] : (sP11('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),X1,'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)))) )),
  inference(unit_resulting_resolution,[],[f623,f2226,f572])).
fof(f6908,plain,(
  ( ! [X2,X0,X1] : (sP11('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),X1,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f635,f2226,f572])).
fof(f6909,plain,(
  ( ! [X2,X0,X3,X1] : (sP11('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),X1,'const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f730,f2226,f572])).
fof(f6910,plain,(
  ( ! [X0,X1] : (sP11('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f624,f2226,f572])).
fof(f6911,plain,(
  ( ! [X0,X1] : (sP11('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X1)) )),
  inference(unit_resulting_resolution,[],[f625,f2226,f572])).
fof(f6912,plain,(
  ( ! [X2,X0,X1] : (sP11('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))),X2)) )),
  inference(unit_resulting_resolution,[],[f949,f2226,f572])).
fof(f6913,plain,(
  ( ! [X2,X0,X1] : (sP11('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))),X2)) )),
  inference(unit_resulting_resolution,[],[f2464,f2226,f572])).
fof(f6914,plain,(
  ( ! [X0,X1] : (sP9('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f616,f2226,f573])).
fof(f6915,plain,(
  ( ! [X0,X1] : (sP9('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),X1,'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) )),
  inference(unit_resulting_resolution,[],[f617,f2226,f573])).
fof(f6916,plain,(
  ( ! [X2,X0,X1] : (sP9('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),X1,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2)))))) )),
  inference(unit_resulting_resolution,[],[f933,f2226,f573])).
fof(f6917,plain,(
  ( ! [X2,X0,X1] : (sP9('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),X1,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f2413,f2226,f573])).
fof(f6918,plain,(
  ( ! [X0,X1] : (sP9('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f618,f2226,f573])).
fof(f6919,plain,(
  ( ! [X0,X1] : (sP9('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X1)) )),
  inference(unit_resulting_resolution,[],[f619,f2226,f573])).
fof(f6920,plain,(
  ( ! [X2,X0,X1] : (sP9('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))),X2)) )),
  inference(unit_resulting_resolution,[],[f934,f2226,f573])).
fof(f6921,plain,(
  ( ! [X2,X0,X1] : (sP9('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))),X2)) )),
  inference(unit_resulting_resolution,[],[f2414,f2226,f573])).
fof(f6922,plain,(
  ( ! [X0,X1] : (sP9('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X0,'const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f2226,f806])).
fof(f6923,plain,(
  ( ! [X0,X1] : (sP9('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),X1)) )),
  inference(unit_resulting_resolution,[],[f2226,f816])).
fof(f6924,plain,(
  ( ! [X0,X1] : (sP11('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),X1)) )),
  inference(unit_resulting_resolution,[],[f2226,f848])).
fof(f6925,plain,(
  ( ! [X2,X0,X3,X1] : (sP9('const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')),X2,'const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X3)))))) )),
  inference(unit_resulting_resolution,[],[f2226,f1189])).
fof(f6926,plain,(
  ( ! [X2,X0,X3,X1] : (sP9('const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')),'const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X2)))),X3)) )),
  inference(unit_resulting_resolution,[],[f2226,f1190])).
fof(f6927,plain,(
  ( ! [X2,X0,X3,X1] : (sP11('const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')),'const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X2)))),X3)) )),
  inference(unit_resulting_resolution,[],[f2226,f1192])).
fof(f6928,plain,(
  ( ! [X0] : (~sP5('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f2226,f1216])).
fof(f6929,plain,(
  ( ! [X0] : (~sP4('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f2226,f1241])).
fof(f6934,plain,(
  ( ! [X0] : (~sP3('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f2226,f1268])).
fof(f6939,plain,(
  ( ! [X0] : (~sP6('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f2226,f1318])).
fof(f6944,plain,(
  ( ! [X0] : (~sP7('F_0','const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) )),
  inference(unit_resulting_resolution,[],[f585,f2226,f1353])).
fof(f6945,plain,(
  ( ! [X0] : (~sP10('const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))),'F_0','const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) )),
  inference(unit_resulting_resolution,[],[f585,f2226,f1527])).
fof(f6950,plain,(
  ( ! [X0] : (~sP11('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'F_0','const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0)))))) )),
  inference(unit_resulting_resolution,[],[f585,f2226,f1570])).
fof(f6954,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X1))))),s('type.list.list'(bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f2226,f1693])).
fof(f6955,plain,(
  ( ! [X2,X0,X1] : (sP0('const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0))))),s('type.list.list'(bool),X1)))),X2)) )),
  inference(unit_resulting_resolution,[],[f2226,f1694])).
fof(f6956,plain,(
  ( ! [X2,X0,X1] : (sP9('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),X1,'const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f2226,f4197])).
fof(f6957,plain,(
  ( ! [X2,X0,X1] : (sP9('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X1)))),X2)) )),
  inference(unit_resulting_resolution,[],[f2226,f4198])).
fof(f6958,plain,(
  ( ! [X2,X0,X1] : (sP11('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X1)))),X2)) )),
  inference(unit_resulting_resolution,[],[f2226,f4202])).
fof(f6543,plain,(
  ( ! [X0] : (~sP3('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f350,f920,f917,f5623,f502])).
fof(f6391,plain,(
  ( ! [X0,X1] : (sP0('const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'T_0'),s('type.list.list'(bool),X0)))),X1)) )),
  inference(unit_resulting_resolution,[],[f350,f1694])).
fof(f6345,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'T_0'),s('type.list.list'(bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f350,f1693])).
fof(f5926,plain,(
  ( ! [X19,X20,X18] : (s(bool,X20) = s(bool,'const.bool.IN_2'(s(X18,X19),s(fun(X18,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,X20)) )),
  inference(subsumption_resolution,[],[f5918,f585])).
fof(f5918,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,'F_0')) | s(bool,X20) = s(bool,'const.bool.IN_2'(s(X18,X19),s(fun(X18,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,X20)) )),
  inference(superposition,[],[f1059,f716])).
fof(f5925,plain,(
  ( ! [X12,X13] : (s(bool,X13) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X12),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X13)) )),
  inference(subsumption_resolution,[],[f5915,f585])).
fof(f5915,plain,(
  ( ! [X12,X13] : (p(s(bool,'F_0')) | s(bool,X13) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X12),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f1059,f668])).
fof(f5923,plain,(
  ( ! [X5] : (s(bool,X5) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f5912,f585])).
fof(f5912,plain,(
  ( ! [X5] : (p(s(bool,'F_0')) | s(bool,X5) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1059,f667])).
fof(f5881,plain,(
  ( ! [X2,X0,X1] : (sP11('const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')),'T_0',X2)) )),
  inference(unit_resulting_resolution,[],[f350,f1192])).
fof(f5868,plain,(
  ( ! [X0] : (sP9('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X0,'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) )),
  inference(unit_resulting_resolution,[],[f617,f890])).
fof(f5872,plain,(
  ( ! [X0] : (sP9('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X0)) )),
  inference(unit_resulting_resolution,[],[f619,f890])).
fof(f5863,plain,(
  ( ! [X2,X0,X1] : (sP11('const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')),X2,'F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f1191])).
fof(f5857,plain,(
  ( ! [X0,X1] : (sP4('T_0','const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X0),s('type.list.list'(bool),X1)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f5823,f360])).
fof(f5847,plain,(
  ( ! [X2,X0,X1] : (sP9('const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')),'T_0',X2)) )),
  inference(unit_resulting_resolution,[],[f350,f1190])).
fof(f5833,plain,(
  ( ! [X2,X0,X1] : (sP9('const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')),X2,'T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f1189])).
fof(f5826,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)),'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X0),s('type.list.list'(bool),X1))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f5586,f360])).
fof(f5807,plain,(
  ( ! [X0,X1] : (sP1('F_0','const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(unit_resulting_resolution,[],[f585,f1188])).
fof(f5799,plain,(
  ( ! [X0,X1] : (~sP2('const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X0),s('type.list.list'(bool),X1)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)),'F_0') | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f5550,f360])).
fof(f5781,plain,(
  ( ! [X0,X1] : (sP1('const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f1187])).
fof(f5774,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)),'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X0),s('type.list.list'(bool),X1))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f5532,f360])).
fof(f5668,plain,(
  ( ! [X78,X79] : (sP5('T_0','const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X78)))),'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'T_0'),s('type.list.list'(bool),X79)))))) )),
  inference(superposition,[],[f1665,f907])).
fof(f5667,plain,(
  ( ! [X76,X77] : (~sP5('F_0','const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X76)))),'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'T_0'),s('type.list.list'(bool),X77)))))) )),
  inference(superposition,[],[f1664,f907])).
fof(f5666,plain,(
  ( ! [X74,X75] : (sP5('T_0','const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'T_0'),s('type.list.list'(bool),X75)))),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X74)))))) )),
  inference(superposition,[],[f1663,f907])).
fof(f5665,plain,(
  ( ! [X72,X73] : (~sP5('F_0','const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'T_0'),s('type.list.list'(bool),X73)))),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X72)))))) )),
  inference(superposition,[],[f1662,f907])).
fof(f5756,plain,(
  ( ! [X70,X71,X69] : (~sP8('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X69)))),X70,X71) | sP9('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X69)))),X70,X71)) )),
  inference(subsumption_resolution,[],[f5664,f350])).
fof(f5664,plain,(
  ( ! [X70,X71,X69] : (~p(s(bool,'T_0')) | ~sP8('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X69)))),X70,X71) | sP9('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X69)))),X70,X71)) )),
  inference(superposition,[],[f573,f907])).
fof(f5755,plain,(
  ( ! [X68,X66,X67] : (~sP10('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X66)))),X67,X68) | sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X66)))),X67,X68)) )),
  inference(subsumption_resolution,[],[f5663,f350])).
fof(f5663,plain,(
  ( ! [X68,X66,X67] : (~p(s(bool,'T_0')) | ~sP10('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X66)))),X67,X68) | sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X66)))),X67,X68)) )),
  inference(superposition,[],[f572,f907])).
fof(f5754,plain,(
  ( ! [X64,X65,X63] : (~p(s(bool,X64)) | p(s(bool,X65)) | ~sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X63)))),X65,X64)) )),
  inference(subsumption_resolution,[],[f5662,f350])).
fof(f5662,plain,(
  ( ! [X64,X65,X63] : (~p(s(bool,'T_0')) | ~p(s(bool,X64)) | p(s(bool,X65)) | ~sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X63)))),X65,X64)) )),
  inference(superposition,[],[f569,f907])).
fof(f5753,plain,(
  ( ! [X61,X62,X60] : (~p(s(bool,X61)) | sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X60)))),X61,X62)) )),
  inference(subsumption_resolution,[],[f5661,f350])).
fof(f5661,plain,(
  ( ! [X61,X62,X60] : (~p(s(bool,'T_0')) | ~p(s(bool,X61)) | sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X60)))),X61,X62)) )),
  inference(superposition,[],[f568,f907])).
fof(f5752,plain,(
  ( ! [X59,X57,X58] : (p(s(bool,X58)) | sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X57)))),X59,X58)) )),
  inference(subsumption_resolution,[],[f5660,f350])).
fof(f5660,plain,(
  ( ! [X59,X57,X58] : (~p(s(bool,'T_0')) | p(s(bool,X58)) | sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X57)))),X59,X58)) )),
  inference(superposition,[],[f567,f907])).
fof(f5751,plain,(
  ( ! [X54,X56,X55] : (p(s(bool,X55)) | ~p(s(bool,X56)) | ~sP10(X56,X55,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X54)))))) )),
  inference(subsumption_resolution,[],[f5659,f350])).
fof(f5659,plain,(
  ( ! [X54,X56,X55] : (~p(s(bool,'T_0')) | p(s(bool,X55)) | ~p(s(bool,X56)) | ~sP10(X56,X55,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X54)))))) )),
  inference(superposition,[],[f562,f907])).
fof(f5750,plain,(
  ( ! [X52,X53,X51] : (p(s(bool,X52)) | p(s(bool,X53)) | ~sP9('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X51)))),X53,X52)) )),
  inference(subsumption_resolution,[],[f5658,f350])).
fof(f5658,plain,(
  ( ! [X52,X53,X51] : (~p(s(bool,'T_0')) | p(s(bool,X52)) | p(s(bool,X53)) | ~sP9('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X51)))),X53,X52)) )),
  inference(superposition,[],[f552,f907])).
fof(f5749,plain,(
  ( ! [X50,X48,X49] : (~p(s(bool,X49)) | sP9('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X48)))),X49,X50)) )),
  inference(subsumption_resolution,[],[f5657,f350])).
fof(f5657,plain,(
  ( ! [X50,X48,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,X49)) | sP9('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X48)))),X49,X50)) )),
  inference(superposition,[],[f551,f907])).
fof(f5748,plain,(
  ( ! [X47,X45,X46] : (~p(s(bool,X46)) | sP9('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45)))),X47,X46)) )),
  inference(subsumption_resolution,[],[f5656,f350])).
fof(f5656,plain,(
  ( ! [X47,X45,X46] : (~p(s(bool,'T_0')) | ~p(s(bool,X46)) | sP9('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45)))),X47,X46)) )),
  inference(superposition,[],[f550,f907])).
fof(f5747,plain,(
  ( ! [X43,X44,X42] : (p(s(bool,X43)) | p(s(bool,X44)) | ~sP8('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X42)))),X43,X44)) )),
  inference(subsumption_resolution,[],[f5655,f350])).
fof(f5655,plain,(
  ( ! [X43,X44,X42] : (~p(s(bool,'T_0')) | p(s(bool,X43)) | p(s(bool,X44)) | ~sP8('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X42)))),X43,X44)) )),
  inference(superposition,[],[f545,f907])).
fof(f5746,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,X40)) | p(s(bool,X41)) | ~sP7(X41,X40,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X39)))))) )),
  inference(subsumption_resolution,[],[f5654,f350])).
fof(f5654,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X40)) | p(s(bool,X41)) | ~sP7(X41,X40,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X39)))))) )),
  inference(superposition,[],[f537,f907])).
fof(f5745,plain,(
  ( ! [X37,X38,X36] : (~p(s(bool,X37)) | p(s(bool,X38)) | ~sP6(X37,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X36)))),X38)) )),
  inference(subsumption_resolution,[],[f5653,f350])).
fof(f5653,plain,(
  ( ! [X37,X38,X36] : (~p(s(bool,'T_0')) | ~p(s(bool,X37)) | p(s(bool,X38)) | ~sP6(X37,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X36)))),X38)) )),
  inference(superposition,[],[f528,f907])).
fof(f5744,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X34)) | p(s(bool,X35)) | ~sP3(X34,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X33)))),X35)) )),
  inference(subsumption_resolution,[],[f5652,f350])).
fof(f5652,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,'T_0')) | ~p(s(bool,X34)) | p(s(bool,X35)) | ~sP3(X34,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X33)))),X35)) )),
  inference(superposition,[],[f509,f907])).
fof(f5743,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,X31)) | p(s(bool,X32)) | ~sP4(X31,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X30)))),X32)) )),
  inference(subsumption_resolution,[],[f5651,f350])).
fof(f5651,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X31)) | p(s(bool,X32)) | ~sP4(X31,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X30)))),X32)) )),
  inference(superposition,[],[f505,f907])).
fof(f5742,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,X28)) | p(s(bool,X29)) | ~sP5(X28,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X27)))),X29)) )),
  inference(subsumption_resolution,[],[f5650,f350])).
fof(f5650,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,'T_0')) | ~p(s(bool,X28)) | p(s(bool,X29)) | ~sP5(X28,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X27)))),X29)) )),
  inference(superposition,[],[f501,f907])).
fof(f5741,plain,(
  ( ! [X26,X25] : (p(s(bool,X26)) | sP1(X26,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X25)))))) )),
  inference(subsumption_resolution,[],[f5649,f350])).
fof(f5649,plain,(
  ( ! [X26,X25] : (~p(s(bool,'T_0')) | p(s(bool,X26)) | sP1(X26,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X25)))))) )),
  inference(superposition,[],[f483,f907])).
fof(f5740,plain,(
  ( ! [X24,X23] : (p(s(bool,X24)) | sP1('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X23)))),X24)) )),
  inference(subsumption_resolution,[],[f5648,f350])).
fof(f5648,plain,(
  ( ! [X24,X23] : (~p(s(bool,'T_0')) | p(s(bool,X24)) | sP1('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X23)))),X24)) )),
  inference(superposition,[],[f480,f907])).
fof(f5623,plain,(
  ( ! [X0] : (~sP5('F_0','const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f585,f907,f520])).
fof(f5627,plain,(
  ( ! [X0] : (~sP5('F_0','T_0','const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f585,f907,f520])).
fof(f5628,plain,(
  ( ! [X0] : (sP5('T_0','T_0','const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f350,f907,f518])).
fof(f5615,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X0),s('type.list.list'(bool),X1)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f5483,f360])).
fof(f5603,plain,(
  ( ! [X0] : (sP4('T_0','const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'T_0'),s('type.list.list'(bool),X0)))))) )),
  inference(unit_resulting_resolution,[],[f5482,f500])).
fof(f5587,plain,(
  ( ! [X6,X8,X7] : (sP5('T_0','const.list.EVERY_2'(s(fun(X6,bool),X7),s('type.list.list'(X6),'const.list.NIL_0')),'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'T_0'),s('type.list.list'(bool),X8)))))) )),
  inference(superposition,[],[f1665,f408])).
fof(f5585,plain,(
  ( ! [X4] : (sP5('T_0','const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'T_0'),s('type.list.list'(bool),X4)))))) )),
  inference(superposition,[],[f1665,f647])).
fof(f5584,plain,(
  ( ! [X2,X3] : (sP5('T_0','F_0','const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X2),s('type.list.list'(bool),X3))))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1665,f360])).
fof(f5579,plain,(
  ( ! [X0,X1] : (sP4('T_0','const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X0),s('type.list.list'(bool),X1)))),X0)) )),
  inference(unit_resulting_resolution,[],[f1665,f500])).
fof(f5558,plain,(
  ( ! [X0,X1] : (~sP2('const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)),'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X0),s('type.list.list'(bool),X1)))),'F_0') | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f5452,f360])).
fof(f5533,plain,(
  ( ! [X6,X8,X7] : (~sP5('F_0','const.list.EVERY_2'(s(fun(X6,bool),X7),s('type.list.list'(X6),'const.list.NIL_0')),'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'T_0'),s('type.list.list'(bool),X8)))))) )),
  inference(superposition,[],[f1664,f408])).
fof(f5531,plain,(
  ( ! [X4] : (~sP5('F_0','const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'T_0'),s('type.list.list'(bool),X4)))))) )),
  inference(superposition,[],[f1664,f647])).
fof(f5530,plain,(
  ( ! [X2,X3] : (~sP5('F_0','F_0','const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X2),s('type.list.list'(bool),X3))))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1664,f360])).
fof(f5522,plain,(
  ( ! [X0] : (~sP3('T_0','const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'T_0'),s('type.list.list'(bool),X0)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f606,f592,f1664,f504])).
fof(f5547,plain,(
  ( ! [X0] : (~sP3('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'T_0'),s('type.list.list'(bool),X0)))),'F_0')) )),
  inference(forward_demodulation,[],[f5523,f647])).
fof(f5523,plain,(
  ( ! [X0] : (~sP3('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))),s('type.list.list'(bool),X0)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f607,f593,f1664,f504])).
fof(f5546,plain,(
  ( ! [X0,X1] : (~sP3('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'T_0'),s('type.list.list'(bool),X1)))),'F_0')) )),
  inference(forward_demodulation,[],[f5524,f907])).
fof(f5524,plain,(
  ( ! [X0,X1] : (~sP3('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))),s('type.list.list'(bool),X1)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f920,f917,f1664,f504])).
fof(f5507,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X0),s('type.list.list'(bool),X1)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f5437,f360])).
fof(f5491,plain,(
  ( ! [X26,X24,X25] : (sP5('T_0','const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X24),s('type.list.list'(bool),X26)))),'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X24),s('type.list.list'(bool),X25)))))) )),
  inference(superposition,[],[f1663,f419])).
fof(f5484,plain,(
  ( ! [X6,X8,X7] : (sP5('T_0','const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'T_0'),s('type.list.list'(bool),X8)))),'const.list.EVERY_2'(s(fun(X6,bool),X7),s('type.list.list'(X6),'const.list.NIL_0')))) )),
  inference(superposition,[],[f1663,f408])).
fof(f5482,plain,(
  ( ! [X4] : (sP5('T_0','const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'T_0'),s('type.list.list'(bool),X4)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) )),
  inference(superposition,[],[f1663,f647])).
fof(f5481,plain,(
  ( ! [X2,X3] : (sP5('T_0','const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X2),s('type.list.list'(bool),X3)))),'F_0') | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1663,f360])).
fof(f5477,plain,(
  ( ! [X0,X1] : (sP4('T_0',X0,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X0),s('type.list.list'(bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f1663,f500])).
fof(f5445,plain,(
  ( ! [X26,X24,X25] : (~sP5('F_0','const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X24),s('type.list.list'(bool),X26)))),'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X24),s('type.list.list'(bool),X25)))))) )),
  inference(superposition,[],[f1662,f419])).
fof(f5438,plain,(
  ( ! [X6,X8,X7] : (~sP5('F_0','const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'T_0'),s('type.list.list'(bool),X8)))),'const.list.EVERY_2'(s(fun(X6,bool),X7),s('type.list.list'(X6),'const.list.NIL_0')))) )),
  inference(superposition,[],[f1662,f408])).
fof(f5436,plain,(
  ( ! [X4] : (~sP5('F_0','const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'T_0'),s('type.list.list'(bool),X4)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) )),
  inference(superposition,[],[f1662,f647])).
fof(f5435,plain,(
  ( ! [X2,X3] : (~sP5('F_0','const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X2),s('type.list.list'(bool),X3)))),'F_0') | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1662,f360])).
fof(f5458,plain,(
  ( ! [X0] : (~sP3('const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'T_0'),s('type.list.list'(bool),X0)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'F_0')) )),
  inference(forward_demodulation,[],[f5420,f647])).
fof(f5420,plain,(
  ( ! [X0] : (~sP3('const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))),s('type.list.list'(bool),X0)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f609,f596,f1662,f504])).
fof(f5457,plain,(
  ( ! [X0,X1] : (~sP3('const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'T_0'),s('type.list.list'(bool),X1)))),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'F_0')) )),
  inference(forward_demodulation,[],[f5421,f907])).
fof(f5421,plain,(
  ( ! [X0,X1] : (~sP3('const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))),s('type.list.list'(bool),X1)))),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f921,f596,f1662,f504])).
fof(f5431,plain,(
  ( ! [X0] : (~sP3('const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'T_0'),s('type.list.list'(bool),X0)))),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f350,f608,f596,f1662,f502])).
fof(f5387,plain,(
  ( ! [X0] : (sP4('T_0','const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(bool),X0)),'const.list.HD_1'(s('type.list.list'(bool),X0)))) )),
  inference(unit_resulting_resolution,[],[f1582,f500])).
fof(f5367,plain,(
  ( ! [X0] : (sP11('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X0,'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)))) )),
  inference(unit_resulting_resolution,[],[f623,f872])).
fof(f5368,plain,(
  ( ! [X0,X1] : (sP11('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X0,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f635,f872])).
fof(f5369,plain,(
  ( ! [X2,X0,X1] : (sP11('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X0,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f730,f872])).
fof(f5371,plain,(
  ( ! [X0] : (sP11('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X0)) )),
  inference(unit_resulting_resolution,[],[f625,f872])).
fof(f5372,plain,(
  ( ! [X0,X1] : (sP11('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),X1)) )),
  inference(unit_resulting_resolution,[],[f949,f872])).
fof(f5313,plain,(
  ( ! [X0] : (sP4('T_0','const.list.HD_1'(s('type.list.list'(bool),X0)),'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(bool),X0)))) )),
  inference(unit_resulting_resolution,[],[f1580,f500])).
fof(f5273,plain,(
  ( ! [X0,X1] : (~sP2('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f725,f723,f5101,f503])).
fof(f5220,plain,(
  ( ! [X0,X1] : (~sP2('F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f600,f594,f5098,f504])).
fof(f5110,plain,(
  ( ! [X9] : (s(bool,X9) = s(bool,'const.bool.IN_2'(s(bool,X9),s(fun(bool,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f716,f360])).
fof(f5098,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f716,f520])).
fof(f5099,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f350,f716,f518])).
fof(f5101,plain,(
  ( ! [X0,X1] : (~sP5('F_0','F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f585,f716,f520])).
fof(f5040,plain,(
  ( ! [X0,X1] : (~sP4('T_0','T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f350,f604,f598,f5009,f502])).
fof(f5009,plain,(
  ( ! [X0,X1] : (~sP5('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f715,f519])).
fof(f5014,plain,(
  ( ! [X0,X1] : (~sP5('T_0','T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f350,f715,f519])).
fof(f4991,plain,(
  ( ! [X0,X1] : (~sP3('T_0','const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f606,f592,f1149,f504])).
fof(f4984,plain,(
  ( ! [X0,X1] : (sP4('T_0','T_0','const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(unit_resulting_resolution,[],[f1148,f500])).
fof(f4976,plain,(
  ( ! [X0,X1] : (~sP3('const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f350,f608,f596,f1147,f502])).
fof(f4897,plain,(
  ( ! [X2,X0,X3,X1] : (sP11('const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X1)))),X2,X3) | ~p(s(bool,X2)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f4202,f360])).
fof(f4872,plain,(
  ( ! [X2,X0,X3,X1] : (sP11('const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X1)))),X2,X3) | p(s(bool,X3)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f4201,f360])).
fof(f4865,plain,(
  ( ! [X2,X0,X3,X1] : (sP9('const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X1)))),X2,X3) | ~p(s(bool,X2)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f4198,f360])).
fof(f4846,plain,(
  ( ! [X2,X0,X3,X1] : (sP9('const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X1)))),X2,X3) | ~p(s(bool,X3)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f4197,f360])).
fof(f4765,plain,(
  ( ! [X0] : (sP7('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'T_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f350,f924,f524])).
fof(f4618,plain,(
  ( ! [X0] : (sP5('F_0','const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f585,f603,f596,f921,f504])).
fof(f4619,plain,(
  ( ! [X0,X1] : (sP5('F_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f585,f631,f596,f921,f504])).
fof(f4620,plain,(
  ( ! [X2,X0,X1] : (sP5('F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2)))))) )),
  inference(unit_resulting_resolution,[],[f585,f726,f596,f921,f504])).
fof(f4621,plain,(
  ( ! [X0] : (sP5('F_0','F_0','const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f585,f602,f596,f921,f503])).
fof(f4622,plain,(
  ( ! [X0] : (sP5('const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)),'F_0','const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f585,f602,f597,f921,f503])).
fof(f4623,plain,(
  ( ! [X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'F_0','const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f585,f602,f629,f921,f503])).
fof(f4624,plain,(
  ( ! [X2,X0,X1] : (sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0','const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2)))))) )),
  inference(unit_resulting_resolution,[],[f585,f602,f724,f921,f503])).
fof(f4582,plain,(
  ( ! [X0] : (sP5('F_0','const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f600,f596,f920,f504])).
fof(f4583,plain,(
  ( ! [X0] : (sP5('F_0','const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)))) )),
  inference(unit_resulting_resolution,[],[f585,f601,f596,f920,f504])).
fof(f4584,plain,(
  ( ! [X0,X1] : (sP5('F_0','const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f585,f630,f596,f920,f504])).
fof(f4585,plain,(
  ( ! [X2,X0,X1] : (sP5('F_0','const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f585,f725,f596,f920,f504])).
fof(f4535,plain,(
  ( ! [X0] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f602,f594,f919,f503])).
fof(f4538,plain,(
  ( ! [X0] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'F_0','const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)))) )),
  inference(unit_resulting_resolution,[],[f585,f602,f595,f919,f503])).
fof(f4541,plain,(
  ( ! [X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'F_0','const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f585,f602,f628,f919,f503])).
fof(f4544,plain,(
  ( ! [X2,X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'F_0','const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f585,f602,f723,f919,f503])).
fof(f4498,plain,(
  ( ! [X0] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'T_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f608,f592,f918,f502])).
fof(f4500,plain,(
  ( ! [X0] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f608,f593,f918,f502])).
fof(f4501,plain,(
  ( ! [X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f608,f917,f918,f502])).
fof(f4479,plain,(
  ( ! [X0] : (~sP2('const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f630,f628,f4339,f503])).
fof(f4460,plain,(
  ( ! [X0] : (sP5('T_0','const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f598,f608,f917,f502])).
fof(f4461,plain,(
  ( ! [X0] : (sP5('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f599,f608,f917,f502])).
fof(f4462,plain,(
  ( ! [X0,X1] : (sP5('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f2325,f608,f917,f502])).
fof(f4437,plain,(
  ( ! [X0] : (~sP2('F_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f600,f594,f4336,f504])).
fof(f4336,plain,(
  ( ! [X0] : (~sP5('F_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f668,f520])).
fof(f4337,plain,(
  ( ! [X0] : (sP5('T_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f350,f668,f518])).
fof(f4339,plain,(
  ( ! [X0] : (~sP5('F_0','F_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f585,f668,f520])).
fof(f4296,plain,(
  ( ! [X0] : (~sP4('T_0','T_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f350,f604,f598,f4274,f502])).
fof(f4274,plain,(
  ( ! [X0] : (~sP5('T_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f644,f519])).
fof(f4279,plain,(
  ( ! [X0] : (~sP5('T_0','T_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f350,f644,f519])).
fof(f4227,plain,(
  ( ! [X2,X0,X1] : (sP1(X1,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X2))))) | p(s(bool,X1)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f4190,f360])).
fof(f4220,plain,(
  ( ! [X2,X0,X1] : (sP1('const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X1)))),X2) | p(s(bool,X2)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f4189,f360])).
fof(f4182,plain,(
  ( ! [X2,X1] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X2),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X1))))))))) )),
  inference(subsumption_resolution,[],[f4110,f350])).
fof(f4110,plain,(
  ( ! [X2,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X2),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X1))))))))) )),
  inference(superposition,[],[f448,f2274])).
fof(f4103,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X1))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2274,f360])).
fof(f3921,plain,(
  ( ! [X0,X1] : (sP11('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f585,f731,f561])).
fof(f3872,plain,(
  ( ! [X2,X0,X1] : (sP11('T_0',X0,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f350,f730,f572])).
fof(f3849,plain,(
  ( ! [X0,X1] : (sP5('F_0','F_0','const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X1))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f3037,f360])).
fof(f3831,plain,(
  ( ! [X0,X1] : (sP9('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f585,f729,f544])).
fof(f3824,plain,(
  ( ! [X0,X1] : (sP5('F_0','const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X1)))),'F_0') | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2985,f360])).
fof(f3815,plain,(
  ( ! [X0,X1] : (sP5('const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X1)))),'F_0','F_0') | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2959,f360])).
fof(f3794,plain,(
  ( ! [X2,X0,X1] : (sP7('F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),X2)) )),
  inference(unit_resulting_resolution,[],[f585,f728,f574])).
fof(f3788,plain,(
  ( ! [X0,X1] : (sP5('const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X1)))),'T_0','T_0') | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2913,f360])).
fof(f3780,plain,(
  ( ! [X0,X1] : (~sP5('F_0','T_0','const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X1))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2906,f360])).
fof(f3759,plain,(
  ( ! [X2,X0,X1] : (sP7('F_0',X0,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f585,f727,f574])).
fof(f3753,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X1)))),'T_0') | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2886,f360])).
fof(f3739,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X1)))),'T_0') | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2874,f360])).
fof(f3713,plain,(
  ( ! [X0,X1] : (sP5('F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) )),
  inference(unit_resulting_resolution,[],[f585,f609,f596,f726,f504])).
fof(f3715,plain,(
  ( ! [X0,X1] : (sP5('F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f608,f596,f726,f502])).
fof(f3716,plain,(
  ( ! [X0,X1] : (sP5('const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)),'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f608,f597,f726,f502])).
fof(f3717,plain,(
  ( ! [X2,X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f608,f629,f726,f502])).
fof(f3718,plain,(
  ( ! [X2,X0,X3,X1] : (sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.pred_set.EMPTY_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f608,f724,f726,f502])).
fof(f3706,plain,(
  ( ! [X0,X1] : (~sP5('T_0','F_0','const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X1))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2849,f360])).
fof(f3696,plain,(
  ( ! [X2,X0,X1] : (sP9('const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X1)))),'T_0',X2) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2486,f360])).
fof(f3673,plain,(
  ( ! [X0,X1] : (sP5('F_0','T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f585,f606,f596,f725,f504])).
fof(f3674,plain,(
  ( ! [X0,X1] : (sP5('F_0','const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f585,f607,f596,f725,f504])).
fof(f3662,plain,(
  ( ! [X2,X0,X1] : (sP9('const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X1)))),X2,'T_0') | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2484,f360])).
fof(f3654,plain,(
  ( ! [X2,X0,X1] : (sP11('const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X1)))),'T_0',X2) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2482,f360])).
fof(f3630,plain,(
  ( ! [X0,X1] : (sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0','const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) )),
  inference(unit_resulting_resolution,[],[f585,f602,f609,f724,f503])).
fof(f3632,plain,(
  ( ! [X0,X1] : (sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f602,f608,f724,f502])).
fof(f3633,plain,(
  ( ! [X0,X1] : (sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f603,f608,f724,f502])).
fof(f3634,plain,(
  ( ! [X2,X0,X1] : (sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num','const.num.|48|_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f631,f608,f724,f502])).
fof(f3624,plain,(
  ( ! [X2,X0,X1] : (sP11('const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X1)))),X2,'F_0') | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2479,f360])).
fof(f3616,plain,(
  ( ! [X0,X1] : (~sP11('const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X1)))),'F_0','T_0') | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2477,f360])).
fof(f3595,plain,(
  ( ! [X0,X1] : (sP5('T_0','F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f585,f602,f604,f723,f503])).
fof(f3597,plain,(
  ( ! [X0,X1] : (sP5('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f585,f602,f605,f723,f503])).
fof(f3599,plain,(
  ( ! [X2,X0,X1] : (sP5('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'F_0','const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f585,f602,f2326,f723,f503])).
fof(f3589,plain,(
  ( ! [X0,X1] : (~sP11('T_0','F_0','const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X1))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2473,f360])).
fof(f3580,plain,(
  ( ! [X2,X0,X1] : (sP11('T_0','const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X1)))),X2) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2468,f360])).
fof(f3562,plain,(
  ( ! [X0,X1] : (sP11('F_0','F_0','const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X1))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2465,f360])).
fof(f3554,plain,(
  ( ! [X2,X0,X3,X1] : (sP10(X1,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X2)))),X3) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2464,f360])).
fof(f3514,plain,(
  ( ! [X0,X1] : (~sP10('T_0','F_0','const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X1))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2457,f360])).
fof(f3505,plain,(
  ( ! [X0,X1] : (~sP10('const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X1)))),'F_0','T_0') | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2450,f360])).
fof(f3473,plain,(
  ( ! [X2,X0,X1] : (~sP11('F_0','const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X1)))),X2) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2426,f360])).
fof(f3463,plain,(
  ( ! [X0,X1] : (~sP9('const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X1)))),'F_0','F_0') | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2425,f360])).
fof(f3437,plain,(
  ( ! [X2,X0,X1] : (sP9('T_0','const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X1)))),X2) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2420,f360])).
fof(f3429,plain,(
  ( ! [X2,X0,X1] : (sP9('T_0',X1,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X2))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2415,f360])).
fof(f3402,plain,(
  ( ! [X2,X0,X3,X1] : (sP8(X1,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X2)))),X3) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2414,f360])).
fof(f3388,plain,(
  ( ! [X2,X0,X3,X1] : (sP8(X1,X2,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X3))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2413,f360])).
fof(f3355,plain,(
  ( ! [X0,X1] : (~sP8('const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X1)))),'F_0','F_0') | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2408,f360])).
fof(f3346,plain,(
  ( ! [X2,X0,X1] : (~sP9('F_0','const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X1)))),X2) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2400,f360])).
fof(f3317,plain,(
  ( ! [X2,X0,X1] : (~sP9('F_0',X1,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X2))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2396,f360])).
fof(f3307,plain,(
  ( ! [X0,X1] : (~sP7('F_0','T_0','const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X1))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2391,f360])).
fof(f3279,plain,(
  ( ! [X0,X1] : (~sP7('F_0','const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X1)))),'T_0') | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2386,f360])).
fof(f3271,plain,(
  ( ! [X0,X1] : (sP7('const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X1)))),'T_0','T_0') | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2382,f360])).
fof(f3254,plain,(
  ( ! [X0,X1] : (sP7('T_0','T_0','const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X1))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2377,f360])).
fof(f3246,plain,(
  ( ! [X0,X1] : (sP7('T_0','const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X1)))),'T_0') | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2372,f360])).
fof(f3177,plain,(
  ( ! [X2,X0,X3,X1] : (sP6(X1,X2,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X3))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2371,f360])).
fof(f3165,plain,(
  ( ! [X0,X1] : (~sP6('T_0','const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X1)))),'F_0') | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2364,f360])).
fof(f3139,plain,(
  ( ! [X0,X1] : (~sP6('const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X1)))),'T_0','F_0') | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2357,f360])).
fof(f3130,plain,(
  ( ! [X2,X0,X1] : (~sP7('const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X1)))),'F_0',X2) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2346,f360])).
fof(f3091,plain,(
  ( ! [X2,X0,X1] : (~sP7('const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X1)))),X2,'F_0') | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2345,f360])).
fof(f3079,plain,(
  ( ! [X0,X1] : (~sP5('const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X1)))),'F_0','T_0') | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2343,f360])).
fof(f3063,plain,(
  ( ! [X0,X1] : (~sP5('const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X1)))),'T_0','F_0') | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2341,f360])).
fof(f3041,plain,(
  ( ! [X2,X0,X3,X1] : (sP2('const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X1)))),X2,X3) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2328,f360])).
fof(f3000,plain,(
  ( ! [X2,X0,X3,X1] : (sP2(X1,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X2)))),X3) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2327,f360])).
fof(f2970,plain,(
  ( ! [X2,X0,X3,X1] : (sP2(X1,X2,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X3))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2326,f360])).
fof(f2965,plain,(
  ( ! [X0,X1] : (sP5('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'F_0','const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f585,f602,f628,f2326,f503])).
fof(f2919,plain,(
  ( ! [X2,X0,X3,X1] : (sP3(X1,X2,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X3))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2325,f360])).
fof(f2916,plain,(
  ( ! [X0,X1] : (sP5('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f608,f2310,f2325,f502])).
fof(f2907,plain,(
  ( ! [X0,X1] : (~sP3('T_0','const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X1)))),'F_0') | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2318,f360])).
fof(f2887,plain,(
  ( ! [X0,X1] : (~sP3('const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X1)))),'T_0','F_0') | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2311,f360])).
fof(f2878,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X1,X2,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X3))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2310,f360])).
fof(f2850,plain,(
  ( ! [X0,X1] : (~sP4('T_0','const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X1)))),'F_0') | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2303,f360])).
fof(f2841,plain,(
  ( ! [X0,X1] : (~sP4('const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X1)))),'T_0','F_0') | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2296,f360])).
fof(f2821,plain,(
  ( ! [X0,X1] : (sP5('T_0','T_0','const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X1))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2292,f360])).
fof(f2810,plain,(
  ( ! [X0,X1] : (~sP5('T_0','const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X1)))),'F_0') | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2286,f360])).
fof(f2716,plain,(
  ( ! [X0,X1] : (sP1('F_0','const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X1))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2283,f360])).
fof(f2710,plain,(
  ( ! [X0,X1] : (sP1('const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X1)))),'F_0') | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2282,f360])).
fof(f2689,plain,(
  ( ! [X0,X1] : (~sP1('T_0','const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X1))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2280,f360])).
fof(f2683,plain,(
  ( ! [X0,X1] : (~sP1('const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X1)))),'T_0') | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2278,f360])).
fof(f2669,plain,(
  ( ! [X2,X0,X1] : (sP0('const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X1)))),X2) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2277,f360])).
fof(f2663,plain,(
  ( ! [X2,X0,X1] : (sP0(X1,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X2))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2276,f360])).
fof(f2571,plain,(
  ( ! [X0,X1] : (s(X0,'const.list.HD_1'(s('type.list.list'(X0),'const.list.TL_1'(s('type.list.list'(X0),X1))))) = s(X0,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.list.list'(X0),X1)))) )),
  inference(superposition,[],[f422,f409])).
fof(f2570,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.list.EL_2'(s('type.num.num',X3),s('type.list.list'(X0),X2))) = s(X0,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))))) )),
  inference(superposition,[],[f422,f420])).
fof(f2555,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(bool),'const.list.TL_1'(s('type.list.list'(bool),X1)))),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.list.list'(bool),X1)))) )),
  inference(unit_resulting_resolution,[],[f585,f422,f520])).
fof(f2556,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(bool),'const.list.TL_1'(s('type.list.list'(bool),X1)))),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.list.list'(bool),X1)))) )),
  inference(unit_resulting_resolution,[],[f350,f422,f518])).
fof(f2558,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.list.list'(bool),X1)),'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(bool),'const.list.TL_1'(s('type.list.list'(bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f585,f422,f520])).
fof(f2559,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.list.list'(bool),X1)),'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(bool),'const.list.TL_1'(s('type.list.list'(bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f350,f422,f518])).
fof(f2274,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))))) )),
  inference(unit_resulting_resolution,[],[f350,f2230,f454])).
fof(f2275,plain,(
  ( ! [X0,X1] : (s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0))))) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f2230,f2230,f454])).
fof(f2276,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f2230,f464])).
fof(f2277,plain,(
  ( ! [X0,X1] : (sP0('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),X1)) )),
  inference(unit_resulting_resolution,[],[f2230,f465])).
fof(f2278,plain,(
  ( ! [X0] : (~sP1('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f2230,f479])).
fof(f2280,plain,(
  ( ! [X0] : (~sP1('T_0','const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))))) )),
  inference(unit_resulting_resolution,[],[f350,f2230,f479])).
fof(f2281,plain,(
  ( ! [X0,X1] : (~sP1('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f2230,f2230,f479])).
fof(f2282,plain,(
  ( ! [X0] : (sP1('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f2230,f480])).
fof(f2283,plain,(
  ( ! [X0] : (sP1('F_0','const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))))) )),
  inference(unit_resulting_resolution,[],[f585,f2230,f483])).
fof(f2286,plain,(
  ( ! [X0] : (~sP5('T_0','const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f350,f585,f2230,f501])).
fof(f2287,plain,(
  ( ! [X0,X1] : (~sP5('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f2230,f585,f2230,f501])).
fof(f2292,plain,(
  ( ! [X0] : (sP5('T_0','T_0','const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))))) )),
  inference(unit_resulting_resolution,[],[f598,f592,f606,f2230,f502])).
fof(f2293,plain,(
  ( ! [X0] : (sP5('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'T_0','const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))))) )),
  inference(unit_resulting_resolution,[],[f599,f592,f606,f2230,f502])).
fof(f2294,plain,(
  ( ! [X0] : (sP5('T_0','const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))))) )),
  inference(unit_resulting_resolution,[],[f598,f593,f607,f2230,f502])).
fof(f2295,plain,(
  ( ! [X0] : (sP5('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))))) )),
  inference(unit_resulting_resolution,[],[f599,f593,f607,f2230,f502])).
fof(f2296,plain,(
  ( ! [X0] : (~sP4('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f350,f585,f2230,f505])).
fof(f2303,plain,(
  ( ! [X0] : (~sP4('T_0','const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f350,f585,f2230,f505])).
fof(f2304,plain,(
  ( ! [X0,X1] : (~sP4('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f2230,f585,f2230,f505])).
fof(f2310,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f2230,f506])).
fof(f2311,plain,(
  ( ! [X0] : (~sP3('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f350,f585,f2230,f509])).
fof(f2318,plain,(
  ( ! [X0] : (~sP3('T_0','const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f350,f585,f2230,f509])).
fof(f2319,plain,(
  ( ! [X0,X1] : (~sP3('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f2230,f585,f2230,f509])).
fof(f2325,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f2230,f510])).
fof(f2326,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f2230,f514])).
fof(f2327,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))),X2)) )),
  inference(unit_resulting_resolution,[],[f2230,f515])).
fof(f2328,plain,(
  ( ! [X2,X0,X1] : (sP2('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),X1,X2)) )),
  inference(unit_resulting_resolution,[],[f2230,f516])).
fof(f2329,plain,(
  ( ! [X0] : (sP5('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f647,f2230,f518])).
fof(f2330,plain,(
  ( ! [X0] : (sP5('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f667,f2230,f518])).
fof(f2331,plain,(
  ( ! [X2,X0,X1] : (sP5('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'const.list.EVERY_2'(s(fun(X1,bool),X2),s('type.list.list'(X1),'const.list.NIL_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f408,f2230,f518])).
fof(f2332,plain,(
  ( ! [X2,X0,X1] : (sP5('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2)))),'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X2)))) )),
  inference(unit_resulting_resolution,[],[f449,f2230,f518])).
fof(f2333,plain,(
  ( ! [X0,X1] : (sP5('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(bool),X1)),'const.list.HD_1'(s('type.list.list'(bool),X1)))) )),
  inference(unit_resulting_resolution,[],[f409,f2230,f518])).
fof(f2334,plain,(
  ( ! [X2,X0,X1] : (sP5('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X1),s('type.list.list'(bool),X2)))),X1)) )),
  inference(unit_resulting_resolution,[],[f419,f2230,f518])).
fof(f2335,plain,(
  ( ! [X0] : (sP5('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'T_0','const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) )),
  inference(unit_resulting_resolution,[],[f647,f2230,f518])).
fof(f2336,plain,(
  ( ! [X0] : (sP5('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'F_0','const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)))) )),
  inference(unit_resulting_resolution,[],[f667,f2230,f518])).
fof(f2337,plain,(
  ( ! [X2,X0,X1] : (sP5('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'T_0','const.list.EVERY_2'(s(fun(X1,bool),X2),s('type.list.list'(X1),'const.list.NIL_0')))) )),
  inference(unit_resulting_resolution,[],[f408,f2230,f518])).
fof(f2338,plain,(
  ( ! [X2,X0,X1] : (sP5('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X2)),'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2)))))) )),
  inference(unit_resulting_resolution,[],[f449,f2230,f518])).
fof(f2339,plain,(
  ( ! [X0,X1] : (sP5('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'const.list.HD_1'(s('type.list.list'(bool),X1)),'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(bool),X1)))) )),
  inference(unit_resulting_resolution,[],[f409,f2230,f518])).
fof(f2340,plain,(
  ( ! [X2,X0,X1] : (sP5('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),X1,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X1),s('type.list.list'(bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f419,f2230,f518])).
fof(f2341,plain,(
  ( ! [X0] : (~sP5('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f642,f2230,f519])).
fof(f2342,plain,(
  ( ! [X0] : (~sP5('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'T_0','const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)))) )),
  inference(unit_resulting_resolution,[],[f643,f2230,f519])).
fof(f2343,plain,(
  ( ! [X0] : (~sP5('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f642,f2230,f519])).
fof(f2344,plain,(
  ( ! [X0] : (~sP5('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f643,f2230,f519])).
fof(f2345,plain,(
  ( ! [X0,X1] : (~sP7('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f2230,f522])).
fof(f2346,plain,(
  ( ! [X0,X1] : (~sP7('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f585,f2230,f523])).
fof(f2349,plain,(
  ( ! [X0] : (sP7('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f611,f2230,f524])).
fof(f2354,plain,(
  ( ! [X0] : (sP7('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'T_0','const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))))) )),
  inference(unit_resulting_resolution,[],[f350,f611,f2230,f524])).
fof(f2355,plain,(
  ( ! [X0,X1] : (sP7('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f2230,f611,f2230,f524])).
fof(f2357,plain,(
  ( ! [X0] : (~sP6('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f350,f585,f2230,f528])).
fof(f2364,plain,(
  ( ! [X0] : (~sP6('T_0','const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f350,f585,f2230,f528])).
fof(f2365,plain,(
  ( ! [X0,X1] : (~sP6('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f2230,f585,f2230,f528])).
fof(f2371,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f2230,f529])).
fof(f2372,plain,(
  ( ! [X0] : (sP7('T_0','const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f350,f2230,f534])).
fof(f2377,plain,(
  ( ! [X0] : (sP7('T_0','T_0','const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))))) )),
  inference(unit_resulting_resolution,[],[f350,f350,f2230,f534])).
fof(f2379,plain,(
  ( ! [X0,X1] : (sP7('T_0','const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f350,f2230,f2230,f534])).
fof(f2382,plain,(
  ( ! [X0] : (sP7('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'T_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f350,f2230,f534])).
fof(f2383,plain,(
  ( ! [X0,X1] : (sP7('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f2230,f350,f2230,f534])).
fof(f2384,plain,(
  ( ! [X0,X1] : (sP7('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'T_0','const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f350,f2230,f2230,f534])).
fof(f2385,plain,(
  ( ! [X2,X0,X1] : (sP7('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))),'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f2230,f2230,f2230,f534])).
fof(f2386,plain,(
  ( ! [X0] : (~sP7('F_0','const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f585,f2230,f537])).
fof(f2391,plain,(
  ( ! [X0] : (~sP7('F_0','T_0','const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))))) )),
  inference(unit_resulting_resolution,[],[f350,f585,f2230,f537])).
fof(f2392,plain,(
  ( ! [X0,X1] : (~sP7('F_0','const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f2230,f585,f2230,f537])).
fof(f2396,plain,(
  ( ! [X0,X1] : (~sP9('F_0',X0,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f585,f2230,f538])).
fof(f2400,plain,(
  ( ! [X0,X1] : (~sP9('F_0','const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),X1)) )),
  inference(unit_resulting_resolution,[],[f585,f2230,f539])).
fof(f2408,plain,(
  ( ! [X0] : (~sP8('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f585,f2230,f545])).
fof(f2413,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f2230,f546])).
fof(f2414,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))),X2)) )),
  inference(unit_resulting_resolution,[],[f2230,f547])).
fof(f2415,plain,(
  ( ! [X0,X1] : (sP9('T_0',X0,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f350,f2230,f550])).
fof(f2419,plain,(
  ( ! [X2,X0,X1] : (sP9('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),X1,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f2230,f2230,f550])).
fof(f2420,plain,(
  ( ! [X0,X1] : (sP9('T_0','const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),X1)) )),
  inference(unit_resulting_resolution,[],[f350,f2230,f551])).
fof(f2424,plain,(
  ( ! [X2,X0,X1] : (sP9('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))),X2)) )),
  inference(unit_resulting_resolution,[],[f2230,f2230,f551])).
fof(f2425,plain,(
  ( ! [X0] : (~sP9('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f585,f2230,f552])).
fof(f2426,plain,(
  ( ! [X0,X1] : (~sP11('F_0','const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),X1)) )),
  inference(unit_resulting_resolution,[],[f585,f2230,f556])).
fof(f2448,plain,(
  ( ! [X0] : (sP11('const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)),'F_0','const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))))) )),
  inference(unit_resulting_resolution,[],[f585,f627,f2230,f561])).
fof(f2449,plain,(
  ( ! [X0,X1] : (sP11('const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'F_0','const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f585,f636,f2230,f561])).
fof(f2450,plain,(
  ( ! [X0] : (~sP10('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f585,f350,f2230,f562])).
fof(f2457,plain,(
  ( ! [X0] : (~sP10('T_0','F_0','const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))))) )),
  inference(unit_resulting_resolution,[],[f585,f350,f2230,f562])).
fof(f2458,plain,(
  ( ! [X0,X1] : (~sP10('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'F_0','const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f585,f2230,f2230,f562])).
fof(f2464,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))),X2)) )),
  inference(unit_resulting_resolution,[],[f2230,f564])).
fof(f2465,plain,(
  ( ! [X0] : (sP11('F_0','F_0','const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))))) )),
  inference(unit_resulting_resolution,[],[f585,f585,f2230,f566])).
fof(f2468,plain,(
  ( ! [X0,X1] : (sP11('T_0','const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),X1)) )),
  inference(unit_resulting_resolution,[],[f350,f2230,f568])).
fof(f2472,plain,(
  ( ! [X2,X0,X1] : (sP11('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))),X2)) )),
  inference(unit_resulting_resolution,[],[f2230,f2230,f568])).
fof(f2473,plain,(
  ( ! [X0] : (~sP11('T_0','F_0','const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))))) )),
  inference(unit_resulting_resolution,[],[f350,f585,f2230,f569])).
fof(f2477,plain,(
  ( ! [X0] : (~sP11('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f585,f2230,f569])).
fof(f2478,plain,(
  ( ! [X0,X1] : (~sP11('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'F_0','const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f2230,f585,f2230,f569])).
fof(f2479,plain,(
  ( ! [X0,X1] : (sP11('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f622,f2230,f572])).
fof(f2480,plain,(
  ( ! [X0,X1] : (sP11('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),X1,'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)))) )),
  inference(unit_resulting_resolution,[],[f623,f2230,f572])).
fof(f2481,plain,(
  ( ! [X2,X0,X1] : (sP11('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),X1,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f635,f2230,f572])).
fof(f2482,plain,(
  ( ! [X0,X1] : (sP11('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f624,f2230,f572])).
fof(f2483,plain,(
  ( ! [X0,X1] : (sP11('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X1)) )),
  inference(unit_resulting_resolution,[],[f625,f2230,f572])).
fof(f2484,plain,(
  ( ! [X0,X1] : (sP9('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f616,f2230,f573])).
fof(f2485,plain,(
  ( ! [X0,X1] : (sP9('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),X1,'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) )),
  inference(unit_resulting_resolution,[],[f617,f2230,f573])).
fof(f2486,plain,(
  ( ! [X0,X1] : (sP9('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f618,f2230,f573])).
fof(f2487,plain,(
  ( ! [X0,X1] : (sP9('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X1)) )),
  inference(unit_resulting_resolution,[],[f619,f2230,f573])).
fof(f2488,plain,(
  ( ! [X0,X1] : (sP9('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X0,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f2230,f806])).
fof(f2489,plain,(
  ( ! [X0,X1] : (sP9('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),X1)) )),
  inference(unit_resulting_resolution,[],[f2230,f816])).
fof(f2490,plain,(
  ( ! [X0,X1] : (sP11('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),X1)) )),
  inference(unit_resulting_resolution,[],[f2230,f848])).
fof(f2491,plain,(
  ( ! [X0] : (~sP5('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f2230,f1216])).
fof(f2492,plain,(
  ( ! [X0] : (~sP4('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f2230,f1241])).
fof(f2495,plain,(
  ( ! [X0] : (~sP3('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f2230,f1268])).
fof(f2498,plain,(
  ( ! [X0] : (~sP6('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f2230,f1318])).
fof(f2501,plain,(
  ( ! [X0] : (~sP7('F_0','const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) )),
  inference(unit_resulting_resolution,[],[f585,f2230,f1353])).
fof(f2502,plain,(
  ( ! [X0] : (~sP10('const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))),'F_0','const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) )),
  inference(unit_resulting_resolution,[],[f585,f2230,f1527])).
fof(f2505,plain,(
  ( ! [X0] : (~sP11('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'F_0','const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))))) )),
  inference(unit_resulting_resolution,[],[f585,f2230,f1570])).
fof(f2232,plain,(
  ( ! [X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X3))))))) )),
  inference(forward_demodulation,[],[f2194,f419])).
fof(f2194,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,'const.list.HD_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2))))),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X3))))))) )),
  inference(unit_resulting_resolution,[],[f419,f447])).
fof(f2230,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0))))))) )),
  inference(forward_demodulation,[],[f2196,f667])).
fof(f2196,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0))))))) )),
  inference(unit_resulting_resolution,[],[f667,f447])).
fof(f2229,plain,(
  ( ! [X2] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X2))))))) )),
  inference(forward_demodulation,[],[f2197,f408])).
fof(f2197,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0'))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X2))))))) )),
  inference(unit_resulting_resolution,[],[f408,f447])).
fof(f2228,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))),s(fun(bool,bool),X2))))))) )),
  inference(forward_demodulation,[],[f2198,f449])).
fof(f2198,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))),s(fun(bool,bool),X2))))))) )),
  inference(unit_resulting_resolution,[],[f449,f447])).
fof(f2227,plain,(
  ( ! [X2,X0,X3] : (p(s(bool,'const.bool.IN_2'(s('type.list.list'(X0),X2),s(fun('type.list.list'(X0),bool),'const.pred_set.INSERT_2'(s('type.list.list'(X0),X2),s(fun('type.list.list'(X0),bool),X3))))))) )),
  inference(forward_demodulation,[],[f2199,f420])).
fof(f2199,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s('type.list.list'(X0),'const.list.TL_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2))))),s(fun('type.list.list'(X0),bool),'const.pred_set.INSERT_2'(s('type.list.list'(X0),X2),s(fun('type.list.list'(X0),bool),X3))))))) )),
  inference(unit_resulting_resolution,[],[f420,f447])).
fof(f2226,plain,(
  ( ! [X1] : (p(s(bool,'const.bool.IN_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X1))))))) )),
  inference(forward_demodulation,[],[f2200,f357])).
fof(f2200,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.NIL_0'))),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X1))))))) )),
  inference(unit_resulting_resolution,[],[f357,f447])).
fof(f2225,plain,(
  ( ! [X2,X0,X3] : (p(s(bool,'const.bool.IN_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2))))),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2))))),s(fun('type.num.num',bool),X3))))))) )),
  inference(forward_demodulation,[],[f2201,f421])).
fof(f2201,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2))))),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2))))),s(fun('type.num.num',bool),X3))))))) )),
  inference(unit_resulting_resolution,[],[f421,f447])).
fof(f2055,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))),'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) )),
  inference(unit_resulting_resolution,[],[f585,f449,f520])).
fof(f2056,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))),'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) )),
  inference(unit_resulting_resolution,[],[f350,f449,f518])).
fof(f2057,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)),'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f585,f449,f520])).
fof(f2058,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)),'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f350,f449,f518])).
fof(f1958,plain,(
  ~sP2('F_0','const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)),'F_0')),
  inference(unit_resulting_resolution,[],[f1133,f603,f594,f585,f504])).
fof(f1973,plain,(
  ( ! [X0] : (sP5('F_0','T_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f630,f628,f606,f585,f504])).
fof(f1982,plain,(
  sP5('F_0','const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'F_0')),
  inference(unit_resulting_resolution,[],[f600,f594,f607,f585,f504])).
fof(f1983,plain,(
  sP5('F_0','const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)))),
  inference(unit_resulting_resolution,[],[f601,f595,f607,f585,f504])).
fof(f1984,plain,(
  ( ! [X0] : (sP5('F_0','const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f630,f628,f607,f585,f504])).
fof(f1994,plain,(
  sP5('F_0','const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)),'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))),
  inference(unit_resulting_resolution,[],[f603,f596,f609,f585,f504])).
fof(f1995,plain,(
  ( ! [X0] : (sP5('F_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) )),
  inference(unit_resulting_resolution,[],[f631,f596,f609,f585,f504])).
fof(f1888,plain,(
  ~sP2('const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1135,f601,f596,f585,f503])).
fof(f1893,plain,(
  sP5('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f602,f594,f604,f585,f503])).
fof(f1897,plain,(
  ( ! [X0] : (sP5('T_0','F_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f598,f628,f604,f585,f503])).
fof(f1898,plain,(
  sP5('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f602,f594,f605,f585,f503])).
fof(f1899,plain,(
  sP5('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'F_0','const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)))),
  inference(unit_resulting_resolution,[],[f602,f595,f605,f585,f503])).
fof(f1902,plain,(
  ( ! [X0] : (sP5('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'F_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f599,f628,f605,f585,f503])).
fof(f1910,plain,(
  sP5('F_0','F_0','const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))),
  inference(unit_resulting_resolution,[],[f602,f596,f609,f585,f503])).
fof(f1911,plain,(
  sP5('const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)),'F_0','const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))),
  inference(unit_resulting_resolution,[],[f602,f597,f609,f585,f503])).
fof(f1912,plain,(
  ( ! [X0] : (sP5('const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'F_0','const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) )),
  inference(unit_resulting_resolution,[],[f602,f629,f609,f585,f503])).
fof(f1818,plain,(
  ( ! [X0] : (sP5('F_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f631,f596,f608,f350,f502])).
fof(f1819,plain,(
  sP5('const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f602,f597,f608,f350,f502])).
fof(f1820,plain,(
  sP5('const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)),'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)),'T_0')),
  inference(unit_resulting_resolution,[],[f603,f597,f608,f350,f502])).
fof(f1821,plain,(
  ( ! [X0] : (sP5('const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)),'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f631,f597,f608,f350,f502])).
fof(f1822,plain,(
  ( ! [X0] : (sP5('const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f602,f629,f608,f350,f502])).
fof(f1823,plain,(
  ( ! [X0] : (sP5('const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f603,f629,f608,f350,f502])).
fof(f1824,plain,(
  ( ! [X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f631,f629,f608,f350,f502])).
fof(f1828,plain,(
  ~sP4('T_0','T_0','const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)))),
  inference(unit_resulting_resolution,[],[f1118,f603,f604,f350,f502])).
fof(f1834,plain,(
  sP5('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f598,f592,f606,f350,f502])).
fof(f1835,plain,(
  sP5('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'T_0','T_0')),
  inference(unit_resulting_resolution,[],[f599,f592,f606,f350,f502])).
fof(f1839,plain,(
  sP5('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'T_0')),
  inference(unit_resulting_resolution,[],[f599,f593,f607,f350,f502])).
fof(f1840,plain,(
  ~sP3('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'T_0','F_0')),
  inference(unit_resulting_resolution,[],[f1132,f596,f607,f350,f502])).
fof(f1774,plain,(
  ~sP10('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f350,f585,f1746,f561])).
fof(f1746,plain,(
  ~sP11('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f350,f585,f1570])).
fof(f1711,plain,(
  ( ! [X111,X114,X112,X113] : (~p(s(bool,X111)) | ~sP8('const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X111),s('type.list.list'(bool),X112)))),X113,X114) | sP9('const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X111),s('type.list.list'(bool),X112)))),X113,X114)) )),
  inference(superposition,[],[f573,f419])).
fof(f1710,plain,(
  ( ! [X109,X107,X110,X108] : (~p(s(bool,X107)) | ~sP10('const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X107),s('type.list.list'(bool),X108)))),X109,X110) | sP11('const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X107),s('type.list.list'(bool),X108)))),X109,X110)) )),
  inference(superposition,[],[f572,f419])).
fof(f1709,plain,(
  ( ! [X103,X105,X106,X104] : (~p(s(bool,X105)) | ~p(s(bool,X103)) | p(s(bool,X106)) | ~sP11('const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X103),s('type.list.list'(bool),X104)))),X106,X105)) )),
  inference(superposition,[],[f569,f419])).
fof(f1708,plain,(
  ( ! [X101,X99,X102,X100] : (sP11('const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X99),s('type.list.list'(bool),X100)))),X101,X102) | ~p(s(bool,X101)) | ~p(s(bool,X99))) )),
  inference(superposition,[],[f568,f419])).
fof(f1707,plain,(
  ( ! [X97,X95,X98,X96] : (sP11('const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X95),s('type.list.list'(bool),X96)))),X98,X97) | p(s(bool,X97)) | ~p(s(bool,X95))) )),
  inference(superposition,[],[f567,f419])).
fof(f1706,plain,(
  ( ! [X94,X92,X93,X91] : (~p(s(bool,X94)) | p(s(bool,X93)) | ~p(s(bool,X91)) | ~sP10(X94,X93,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X91),s('type.list.list'(bool),X92)))))) )),
  inference(superposition,[],[f562,f419])).
fof(f1705,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,X87)) | p(s(bool,X89)) | p(s(bool,X90)) | ~sP9('const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X87),s('type.list.list'(bool),X88)))),X90,X89)) )),
  inference(superposition,[],[f552,f419])).
fof(f1704,plain,(
  ( ! [X85,X83,X86,X84] : (sP9('const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X83),s('type.list.list'(bool),X84)))),X85,X86) | ~p(s(bool,X85)) | ~p(s(bool,X83))) )),
  inference(superposition,[],[f551,f419])).
fof(f1703,plain,(
  ( ! [X80,X81,X79,X82] : (sP9('const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X79),s('type.list.list'(bool),X80)))),X82,X81) | ~p(s(bool,X81)) | ~p(s(bool,X79))) )),
  inference(superposition,[],[f550,f419])).
fof(f1702,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,X75)) | p(s(bool,X77)) | p(s(bool,X78)) | ~sP8('const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X75),s('type.list.list'(bool),X76)))),X77,X78)) )),
  inference(superposition,[],[f545,f419])).
fof(f1701,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,X73)) | ~p(s(bool,X71)) | p(s(bool,X74)) | ~sP7(X74,X73,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X71),s('type.list.list'(bool),X72)))))) )),
  inference(superposition,[],[f537,f419])).
fof(f1700,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,X69)) | ~p(s(bool,X67)) | p(s(bool,X70)) | ~sP6(X69,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X67),s('type.list.list'(bool),X68)))),X70)) )),
  inference(superposition,[],[f528,f419])).
fof(f1699,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X65)) | ~p(s(bool,X63)) | p(s(bool,X66)) | ~sP3(X65,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X63),s('type.list.list'(bool),X64)))),X66)) )),
  inference(superposition,[],[f509,f419])).
fof(f1698,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X61)) | ~p(s(bool,X59)) | p(s(bool,X62)) | ~sP4(X61,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X59),s('type.list.list'(bool),X60)))),X62)) )),
  inference(superposition,[],[f505,f419])).
fof(f1697,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X57)) | ~p(s(bool,X55)) | p(s(bool,X58)) | ~sP5(X57,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X55),s('type.list.list'(bool),X56)))),X58)) )),
  inference(superposition,[],[f501,f419])).
fof(f1696,plain,(
  ( ! [X54,X52,X53] : (sP1(X54,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X52),s('type.list.list'(bool),X53))))) | p(s(bool,X54)) | ~p(s(bool,X52))) )),
  inference(superposition,[],[f483,f419])).
fof(f1695,plain,(
  ( ! [X50,X51,X49] : (sP1('const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X49),s('type.list.list'(bool),X50)))),X51) | p(s(bool,X51)) | ~p(s(bool,X49))) )),
  inference(superposition,[],[f480,f419])).
fof(f1694,plain,(
  ( ! [X47,X48,X46] : (sP0('const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X46),s('type.list.list'(bool),X47)))),X48) | ~p(s(bool,X46))) )),
  inference(superposition,[],[f465,f419])).
fof(f1693,plain,(
  ( ! [X45,X43,X44] : (sP0(X45,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X43),s('type.list.list'(bool),X44))))) | ~p(s(bool,X43))) )),
  inference(superposition,[],[f464,f419])).
fof(f1662,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X0),s('type.list.list'(bool),X1)))),X0)) )),
  inference(unit_resulting_resolution,[],[f585,f419,f520])).
fof(f1663,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X0),s('type.list.list'(bool),X1)))),X0)) )),
  inference(unit_resulting_resolution,[],[f350,f419,f518])).
fof(f1664,plain,(
  ( ! [X0,X1] : (~sP5('F_0',X0,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X0),s('type.list.list'(bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f585,f419,f520])).
fof(f1665,plain,(
  ( ! [X0,X1] : (sP5('T_0',X0,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X0),s('type.list.list'(bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f350,f419,f518])).
fof(f1658,plain,(
  ~sP11('T_0','F_0','const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))),
  inference(unit_resulting_resolution,[],[f1646,f557])).
fof(f1646,plain,(
  ~sP10('T_0','F_0','const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))),
  inference(unit_resulting_resolution,[],[f585,f350,f1527])).
fof(f1579,plain,(
  ( ! [X0] : (~sP5('F_0','const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(bool),X0)),'const.list.HD_1'(s('type.list.list'(bool),X0)))) )),
  inference(unit_resulting_resolution,[],[f585,f409,f520])).
fof(f1580,plain,(
  ( ! [X0] : (sP5('T_0','const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(bool),X0)),'const.list.HD_1'(s('type.list.list'(bool),X0)))) )),
  inference(unit_resulting_resolution,[],[f350,f409,f518])).
fof(f1581,plain,(
  ( ! [X0] : (~sP5('F_0','const.list.HD_1'(s('type.list.list'(bool),X0)),'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(bool),X0)))) )),
  inference(unit_resulting_resolution,[],[f585,f409,f520])).
fof(f1582,plain,(
  ( ! [X0] : (sP5('T_0','const.list.HD_1'(s('type.list.list'(bool),X0)),'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(bool),X0)))) )),
  inference(unit_resulting_resolution,[],[f350,f409,f518])).
fof(f1571,plain,(
  ( ! [X6,X4,X7,X5] : (~sP11('const.list.EVERY_2'(s(fun(X4,bool),X5),s('type.list.list'(X4),'const.list.NIL_0')),X7,X6) | p(s(bool,X7)) | ~p(s(bool,X6))) )),
  inference(subsumption_resolution,[],[f1563,f350])).
fof(f1563,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'T_0')) | ~p(s(bool,X6)) | p(s(bool,X7)) | ~sP11('const.list.EVERY_2'(s(fun(X4,bool),X5),s('type.list.list'(X4),'const.list.NIL_0')),X7,X6)) )),
  inference(superposition,[],[f569,f408])).
fof(f1570,plain,(
  ( ! [X0,X1] : (~sP11('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X1,X0) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f1561,f350])).
fof(f1561,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X1)) | ~sP11('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X1,X0)) )),
  inference(superposition,[],[f569,f647])).
fof(f1546,plain,(
  ~sP9('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1492,f540])).
fof(f1532,plain,(
  ~sP11('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f1516,f557])).
fof(f1528,plain,(
  ( ! [X6,X4,X7,X5] : (~sP10(X7,X6,'const.list.EVERY_2'(s(fun(X4,bool),X5),s('type.list.list'(X4),'const.list.NIL_0'))) | ~p(s(bool,X7)) | p(s(bool,X6))) )),
  inference(subsumption_resolution,[],[f1519,f350])).
fof(f1519,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'T_0')) | p(s(bool,X6)) | ~p(s(bool,X7)) | ~sP10(X7,X6,'const.list.EVERY_2'(s(fun(X4,bool),X5),s('type.list.list'(X4),'const.list.NIL_0')))) )),
  inference(superposition,[],[f562,f408])).
fof(f1527,plain,(
  ( ! [X0,X1] : (~sP10(X1,X0,'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | ~p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f1517,f350])).
fof(f1517,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | ~p(s(bool,X1)) | ~sP10(X1,X0,'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) )),
  inference(superposition,[],[f562,f647])).
fof(f1516,plain,(
  ~sP10('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f350,f585,f350,f562])).
fof(f1501,plain,(
  sP11('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f626,f350,f585,f561])).
fof(f1502,plain,(
  sP11('const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f627,f350,f585,f561])).
fof(f1503,plain,(
  ( ! [X0] : (sP11('const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f636,f350,f585,f561])).
fof(f1492,plain,(
  ~sP8('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f585,f585,f1410])).
fof(f1476,plain,(
  ~sP7('F_0','T_0','const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))),
  inference(unit_resulting_resolution,[],[f1461,f521])).
fof(f1461,plain,(
  ~sP6('T_0','const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'F_0')),
  inference(unit_resulting_resolution,[],[f350,f585,f1318])).
fof(f1455,plain,(
  ( ! [X6,X4,X7,X5] : (~sP9('const.list.EVERY_2'(s(fun(X4,bool),X5),s('type.list.list'(X4),'const.list.NIL_0')),X7,X6) | p(s(bool,X7)) | p(s(bool,X6))) )),
  inference(subsumption_resolution,[],[f1447,f350])).
fof(f1447,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'T_0')) | p(s(bool,X6)) | p(s(bool,X7)) | ~sP9('const.list.EVERY_2'(s(fun(X4,bool),X5),s('type.list.list'(X4),'const.list.NIL_0')),X7,X6)) )),
  inference(superposition,[],[f552,f408])).
fof(f1454,plain,(
  ( ! [X0,X1] : (~sP9('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X1,X0) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f1445,f350])).
fof(f1445,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP9('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X1,X0)) )),
  inference(superposition,[],[f552,f647])).
fof(f1428,plain,(
  ~sP3('T_0','const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'F_0')),
  inference(unit_resulting_resolution,[],[f350,f585,f1268])).
fof(f1415,plain,(
  ~sP9('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1400,f540])).
fof(f1411,plain,(
  ( ! [X6,X4,X7,X5] : (~sP8('const.list.EVERY_2'(s(fun(X4,bool),X5),s('type.list.list'(X4),'const.list.NIL_0')),X6,X7) | p(s(bool,X7)) | p(s(bool,X6))) )),
  inference(subsumption_resolution,[],[f1403,f350])).
fof(f1403,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'T_0')) | p(s(bool,X6)) | p(s(bool,X7)) | ~sP8('const.list.EVERY_2'(s(fun(X4,bool),X5),s('type.list.list'(X4),'const.list.NIL_0')),X6,X7)) )),
  inference(superposition,[],[f545,f408])).
fof(f1410,plain,(
  ( ! [X0,X1] : (~sP8('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f1401,f350])).
fof(f1401,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP8('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X0,X1)) )),
  inference(superposition,[],[f545,f647])).
fof(f1400,plain,(
  ~sP8('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f585,f585,f350,f545])).
fof(f1386,plain,(
  sP9('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f620,f585,f585,f544])).
fof(f1387,plain,(
  sP9('const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f621,f585,f585,f544])).
fof(f1388,plain,(
  ( ! [X0] : (sP9('const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f634,f585,f585,f544])).
fof(f1380,plain,(
  ~sP5('T_0','F_0','const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))),
  inference(unit_resulting_resolution,[],[f1365,f500])).
fof(f1365,plain,(
  ~sP4('T_0','const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'F_0')),
  inference(unit_resulting_resolution,[],[f350,f585,f1241])).
fof(f1358,plain,(
  ~sP5('T_0','const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'F_0')),
  inference(unit_resulting_resolution,[],[f350,f585,f1216])).
fof(f1354,plain,(
  ( ! [X6,X4,X7,X5] : (~sP7(X7,X6,'const.list.EVERY_2'(s(fun(X4,bool),X5),s('type.list.list'(X4),'const.list.NIL_0'))) | p(s(bool,X7)) | ~p(s(bool,X6))) )),
  inference(subsumption_resolution,[],[f1346,f350])).
fof(f1346,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'T_0')) | ~p(s(bool,X6)) | p(s(bool,X7)) | ~sP7(X7,X6,'const.list.EVERY_2'(s(fun(X4,bool),X5),s('type.list.list'(X4),'const.list.NIL_0')))) )),
  inference(superposition,[],[f537,f408])).
fof(f1353,plain,(
  ( ! [X0,X1] : (~sP7(X1,X0,'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f1344,f350])).
fof(f1344,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X1)) | ~sP7(X1,X0,'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) )),
  inference(superposition,[],[f537,f647])).
fof(f1323,plain,(
  ~sP7('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f1307,f521])).
fof(f1319,plain,(
  ( ! [X6,X4,X7,X5] : (~sP6(X6,'const.list.EVERY_2'(s(fun(X4,bool),X5),s('type.list.list'(X4),'const.list.NIL_0')),X7) | p(s(bool,X7)) | ~p(s(bool,X6))) )),
  inference(subsumption_resolution,[],[f1310,f350])).
fof(f1310,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'T_0')) | ~p(s(bool,X6)) | p(s(bool,X7)) | ~sP6(X6,'const.list.EVERY_2'(s(fun(X4,bool),X5),s('type.list.list'(X4),'const.list.NIL_0')),X7)) )),
  inference(superposition,[],[f528,f408])).
fof(f1318,plain,(
  ( ! [X0,X1] : (~sP6(X0,'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X1) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f1308,f350])).
fof(f1308,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X1)) | ~sP6(X0,'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X1)) )),
  inference(superposition,[],[f528,f647])).
fof(f1307,plain,(
  ~sP6('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f585,f350,f350,f528])).
fof(f1292,plain,(
  sP7('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f610,f350,f350,f524])).
fof(f1293,plain,(
  sP7('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'T_0','T_0')),
  inference(unit_resulting_resolution,[],[f611,f350,f350,f524])).
fof(f1288,plain,(
  ~sP5('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1279,f498])).
fof(f1279,plain,(
  ~sP2('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f585,f585,f585,f513])).
fof(f1273,plain,(
  ~sP5('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f1257,f499])).
fof(f1269,plain,(
  ( ! [X6,X4,X7,X5] : (~sP3(X6,'const.list.EVERY_2'(s(fun(X4,bool),X5),s('type.list.list'(X4),'const.list.NIL_0')),X7) | p(s(bool,X7)) | ~p(s(bool,X6))) )),
  inference(subsumption_resolution,[],[f1260,f350])).
fof(f1260,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'T_0')) | ~p(s(bool,X6)) | p(s(bool,X7)) | ~sP3(X6,'const.list.EVERY_2'(s(fun(X4,bool),X5),s('type.list.list'(X4),'const.list.NIL_0')),X7)) )),
  inference(superposition,[],[f509,f408])).
fof(f1268,plain,(
  ( ! [X0,X1] : (~sP3(X0,'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X1) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f1258,f350])).
fof(f1258,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X1)) | ~sP3(X0,'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X1)) )),
  inference(superposition,[],[f509,f647])).
fof(f1257,plain,(
  ~sP3('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f585,f350,f350,f509])).
fof(f1242,plain,(
  ( ! [X6,X4,X7,X5] : (~sP4(X6,'const.list.EVERY_2'(s(fun(X4,bool),X5),s('type.list.list'(X4),'const.list.NIL_0')),X7) | p(s(bool,X7)) | ~p(s(bool,X6))) )),
  inference(subsumption_resolution,[],[f1233,f350])).
fof(f1233,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'T_0')) | ~p(s(bool,X6)) | p(s(bool,X7)) | ~sP4(X6,'const.list.EVERY_2'(s(fun(X4,bool),X5),s('type.list.list'(X4),'const.list.NIL_0')),X7)) )),
  inference(superposition,[],[f505,f408])).
fof(f1241,plain,(
  ( ! [X0,X1] : (~sP4(X0,'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X1) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f1231,f350])).
fof(f1231,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X1)) | ~sP4(X0,'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X1)) )),
  inference(superposition,[],[f505,f647])).
fof(f1230,plain,(
  ~sP4('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f585,f350,f350,f505])).
fof(f1217,plain,(
  ( ! [X6,X4,X7,X5] : (~sP5(X6,'const.list.EVERY_2'(s(fun(X4,bool),X5),s('type.list.list'(X4),'const.list.NIL_0')),X7) | p(s(bool,X7)) | ~p(s(bool,X6))) )),
  inference(subsumption_resolution,[],[f1211,f350])).
fof(f1211,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'T_0')) | ~p(s(bool,X6)) | p(s(bool,X7)) | ~sP5(X6,'const.list.EVERY_2'(s(fun(X4,bool),X5),s('type.list.list'(X4),'const.list.NIL_0')),X7)) )),
  inference(superposition,[],[f501,f408])).
fof(f1216,plain,(
  ( ! [X0,X1] : (~sP5(X0,'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X1) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f1209,f350])).
fof(f1209,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X1)) | ~sP5(X0,'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X1)) )),
  inference(superposition,[],[f501,f647])).
fof(f1194,plain,(
  ( ! [X54,X52,X53,X51] : (sP9('const.list.EVERY_2'(s(fun(X51,bool),X52),s('type.list.list'(X51),'const.list.NIL_0')),X53,X54) | ~sP8('const.list.EVERY_2'(s(fun(X51,bool),X52),s('type.list.list'(X51),'const.list.NIL_0')),X53,X54)) )),
  inference(subsumption_resolution,[],[f1169,f350])).
fof(f1169,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,'T_0')) | ~sP8('const.list.EVERY_2'(s(fun(X51,bool),X52),s('type.list.list'(X51),'const.list.NIL_0')),X53,X54) | sP9('const.list.EVERY_2'(s(fun(X51,bool),X52),s('type.list.list'(X51),'const.list.NIL_0')),X53,X54)) )),
  inference(superposition,[],[f573,f408])).
fof(f1193,plain,(
  ( ! [X47,X50,X48,X49] : (sP11('const.list.EVERY_2'(s(fun(X47,bool),X48),s('type.list.list'(X47),'const.list.NIL_0')),X49,X50) | ~sP10('const.list.EVERY_2'(s(fun(X47,bool),X48),s('type.list.list'(X47),'const.list.NIL_0')),X49,X50)) )),
  inference(subsumption_resolution,[],[f1168,f350])).
fof(f1168,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,'T_0')) | ~sP10('const.list.EVERY_2'(s(fun(X47,bool),X48),s('type.list.list'(X47),'const.list.NIL_0')),X49,X50) | sP11('const.list.EVERY_2'(s(fun(X47,bool),X48),s('type.list.list'(X47),'const.list.NIL_0')),X49,X50)) )),
  inference(superposition,[],[f572,f408])).
fof(f1192,plain,(
  ( ! [X45,X43,X46,X44] : (sP11('const.list.EVERY_2'(s(fun(X43,bool),X44),s('type.list.list'(X43),'const.list.NIL_0')),X45,X46) | ~p(s(bool,X45))) )),
  inference(subsumption_resolution,[],[f1167,f350])).
fof(f1167,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,'T_0')) | ~p(s(bool,X45)) | sP11('const.list.EVERY_2'(s(fun(X43,bool),X44),s('type.list.list'(X43),'const.list.NIL_0')),X45,X46)) )),
  inference(superposition,[],[f568,f408])).
fof(f1191,plain,(
  ( ! [X39,X41,X42,X40] : (sP11('const.list.EVERY_2'(s(fun(X39,bool),X40),s('type.list.list'(X39),'const.list.NIL_0')),X42,X41) | p(s(bool,X41))) )),
  inference(subsumption_resolution,[],[f1166,f350])).
fof(f1166,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,'T_0')) | p(s(bool,X41)) | sP11('const.list.EVERY_2'(s(fun(X39,bool),X40),s('type.list.list'(X39),'const.list.NIL_0')),X42,X41)) )),
  inference(superposition,[],[f567,f408])).
fof(f1190,plain,(
  ( ! [X37,X35,X38,X36] : (sP9('const.list.EVERY_2'(s(fun(X35,bool),X36),s('type.list.list'(X35),'const.list.NIL_0')),X37,X38) | ~p(s(bool,X37))) )),
  inference(subsumption_resolution,[],[f1165,f350])).
fof(f1165,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'T_0')) | ~p(s(bool,X37)) | sP9('const.list.EVERY_2'(s(fun(X35,bool),X36),s('type.list.list'(X35),'const.list.NIL_0')),X37,X38)) )),
  inference(superposition,[],[f551,f408])).
fof(f1189,plain,(
  ( ! [X33,X31,X34,X32] : (sP9('const.list.EVERY_2'(s(fun(X31,bool),X32),s('type.list.list'(X31),'const.list.NIL_0')),X34,X33) | ~p(s(bool,X33))) )),
  inference(subsumption_resolution,[],[f1164,f350])).
fof(f1164,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X33)) | sP9('const.list.EVERY_2'(s(fun(X31,bool),X32),s('type.list.list'(X31),'const.list.NIL_0')),X34,X33)) )),
  inference(superposition,[],[f550,f408])).
fof(f1188,plain,(
  ( ! [X30,X28,X29] : (sP1(X30,'const.list.EVERY_2'(s(fun(X28,bool),X29),s('type.list.list'(X28),'const.list.NIL_0'))) | p(s(bool,X30))) )),
  inference(subsumption_resolution,[],[f1163,f350])).
fof(f1163,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'T_0')) | p(s(bool,X30)) | sP1(X30,'const.list.EVERY_2'(s(fun(X28,bool),X29),s('type.list.list'(X28),'const.list.NIL_0')))) )),
  inference(superposition,[],[f483,f408])).
fof(f1187,plain,(
  ( ! [X26,X27,X25] : (sP1('const.list.EVERY_2'(s(fun(X25,bool),X26),s('type.list.list'(X25),'const.list.NIL_0')),X27) | p(s(bool,X27))) )),
  inference(subsumption_resolution,[],[f1162,f350])).
fof(f1162,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'T_0')) | p(s(bool,X27)) | sP1('const.list.EVERY_2'(s(fun(X25,bool),X26),s('type.list.list'(X25),'const.list.NIL_0')),X27)) )),
  inference(superposition,[],[f480,f408])).
fof(f1186,plain,(
  ( ! [X24,X23,X22] : (sP0('const.list.EVERY_2'(s(fun(X22,bool),X23),s('type.list.list'(X22),'const.list.NIL_0')),X24)) )),
  inference(subsumption_resolution,[],[f1161,f350])).
fof(f1161,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | sP0('const.list.EVERY_2'(s(fun(X22,bool),X23),s('type.list.list'(X22),'const.list.NIL_0')),X24)) )),
  inference(superposition,[],[f465,f408])).
fof(f1185,plain,(
  ( ! [X21,X19,X20] : (sP0(X21,'const.list.EVERY_2'(s(fun(X19,bool),X20),s('type.list.list'(X19),'const.list.NIL_0')))) )),
  inference(subsumption_resolution,[],[f1160,f350])).
fof(f1160,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | sP0(X21,'const.list.EVERY_2'(s(fun(X19,bool),X20),s('type.list.list'(X19),'const.list.NIL_0')))) )),
  inference(superposition,[],[f464,f408])).
fof(f1147,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f585,f408,f520])).
fof(f1148,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f408,f518])).
fof(f1149,plain,(
  ( ! [X0,X1] : (~sP5('F_0','T_0','const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(unit_resulting_resolution,[],[f585,f408,f520])).
fof(f1150,plain,(
  ( ! [X0,X1] : (sP5('T_0','T_0','const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(unit_resulting_resolution,[],[f350,f408,f518])).
fof(f1132,plain,(
  ~sP5('F_0','const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'T_0')),
  inference(unit_resulting_resolution,[],[f585,f647,f520])).
fof(f1133,plain,(
  ~sP5('F_0','const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)),'F_0')),
  inference(unit_resulting_resolution,[],[f585,f667,f520])).
fof(f1134,plain,(
  ~sP5('F_0','T_0','const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))),
  inference(unit_resulting_resolution,[],[f585,f647,f520])).
fof(f1135,plain,(
  ~sP5('F_0','F_0','const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)))),
  inference(unit_resulting_resolution,[],[f585,f667,f520])).
fof(f1117,plain,(
  ~sP5('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f350,f642,f519])).
fof(f1118,plain,(
  ~sP5('T_0','const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)),'T_0')),
  inference(unit_resulting_resolution,[],[f350,f643,f519])).
fof(f1121,plain,(
  ~sP5('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f350,f642,f519])).
fof(f1122,plain,(
  ~sP5('T_0','T_0','const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)))),
  inference(unit_resulting_resolution,[],[f350,f643,f519])).
fof(f1106,plain,(
  sP5('T_0','const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'T_0')),
  inference(unit_resulting_resolution,[],[f350,f647,f518])).
fof(f1107,plain,(
  sP5('T_0','const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)),'F_0')),
  inference(unit_resulting_resolution,[],[f350,f667,f518])).
fof(f1108,plain,(
  sP5('T_0','T_0','const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))),
  inference(unit_resulting_resolution,[],[f350,f647,f518])).
fof(f1109,plain,(
  sP5('T_0','F_0','const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)))),
  inference(unit_resulting_resolution,[],[f350,f667,f518])).
fof(f1096,plain,(
  sP5('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f585,f642,f517])).
fof(f1097,plain,(
  sP5('F_0','const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)),'T_0')),
  inference(unit_resulting_resolution,[],[f585,f643,f517])).
fof(f1098,plain,(
  sP5('F_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f585,f642,f517])).
fof(f1099,plain,(
  sP5('F_0','T_0','const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)))),
  inference(unit_resulting_resolution,[],[f585,f643,f517])).
fof(f1059,plain,(
  ( ! [X4,X5] : (p(s(bool,X5)) | s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f1052,f585])).
fof(f1052,plain,(
  ( ! [X4,X5] : (p(s(bool,'F_0')) | p(s(bool,X5)) | s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f451,f360])).
fof(f983,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f360,f360])).
fof(f974,plain,(
  ( ! [X0] : (sP11('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f350,f848])).
fof(f907,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f358,f374])).
fof(f908,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f358,f377])).
fof(f909,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f358,f464])).
fof(f910,plain,(
  ( ! [X0,X1] : (sP0('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),X1)) )),
  inference(unit_resulting_resolution,[],[f358,f465])).
fof(f911,plain,(
  ( ! [X0] : (~sP1('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f358,f479])).
fof(f913,plain,(
  ( ! [X0] : (~sP1('T_0','const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f350,f358,f479])).
fof(f914,plain,(
  ( ! [X0,X1] : (~sP1('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f358,f358,f479])).
fof(f915,plain,(
  ( ! [X0] : (sP1('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f358,f480])).
fof(f916,plain,(
  ( ! [X0] : (sP1('F_0','const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f585,f358,f483])).
fof(f917,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2)))))) )),
  inference(unit_resulting_resolution,[],[f358,f506])).
fof(f918,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2)))))) )),
  inference(unit_resulting_resolution,[],[f358,f510])).
fof(f919,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2)))))) )),
  inference(unit_resulting_resolution,[],[f358,f514])).
fof(f920,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))),X2)) )),
  inference(unit_resulting_resolution,[],[f358,f515])).
fof(f921,plain,(
  ( ! [X2,X0,X1] : (sP2('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),X1,X2)) )),
  inference(unit_resulting_resolution,[],[f358,f516])).
fof(f922,plain,(
  ( ! [X0,X1] : (~sP7('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f358,f522])).
fof(f923,plain,(
  ( ! [X0,X1] : (~sP7('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f585,f358,f523])).
fof(f924,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2)))))) )),
  inference(unit_resulting_resolution,[],[f358,f529])).
fof(f925,plain,(
  ( ! [X0,X1] : (~sP9('F_0',X0,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f585,f358,f538])).
fof(f929,plain,(
  ( ! [X0,X1] : (~sP9('F_0','const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),X1)) )),
  inference(unit_resulting_resolution,[],[f585,f358,f539])).
fof(f933,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2)))))) )),
  inference(unit_resulting_resolution,[],[f358,f546])).
fof(f934,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))),X2)) )),
  inference(unit_resulting_resolution,[],[f358,f547])).
fof(f935,plain,(
  ( ! [X0,X1] : (sP9('T_0',X0,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f350,f358,f550])).
fof(f939,plain,(
  ( ! [X2,X0,X1] : (sP9('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),X1,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2)))))) )),
  inference(unit_resulting_resolution,[],[f358,f358,f550])).
fof(f940,plain,(
  ( ! [X0,X1] : (sP9('T_0','const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),X1)) )),
  inference(unit_resulting_resolution,[],[f350,f358,f551])).
fof(f944,plain,(
  ( ! [X2,X0,X1] : (sP9('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))),X2)) )),
  inference(unit_resulting_resolution,[],[f358,f358,f551])).
fof(f945,plain,(
  ( ! [X0,X1] : (~sP11('F_0','const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),X1)) )),
  inference(unit_resulting_resolution,[],[f585,f358,f556])).
fof(f949,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))),X2)) )),
  inference(unit_resulting_resolution,[],[f358,f564])).
fof(f951,plain,(
  ( ! [X0,X1] : (sP11('T_0','const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),X1)) )),
  inference(unit_resulting_resolution,[],[f350,f358,f568])).
fof(f955,plain,(
  ( ! [X2,X0,X1] : (sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))),X2)) )),
  inference(unit_resulting_resolution,[],[f358,f358,f568])).
fof(f956,plain,(
  ( ! [X0,X1] : (sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f622,f358,f572])).
fof(f957,plain,(
  ( ! [X0,X1] : (sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),X1,'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)))) )),
  inference(unit_resulting_resolution,[],[f623,f358,f572])).
fof(f958,plain,(
  ( ! [X2,X0,X1] : (sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),X1,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f635,f358,f572])).
fof(f959,plain,(
  ( ! [X0,X1] : (sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f624,f358,f572])).
fof(f960,plain,(
  ( ! [X0,X1] : (sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X1)) )),
  inference(unit_resulting_resolution,[],[f625,f358,f572])).
fof(f961,plain,(
  ( ! [X0,X1] : (sP9('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f616,f358,f573])).
fof(f962,plain,(
  ( ! [X0,X1] : (sP9('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),X1,'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) )),
  inference(unit_resulting_resolution,[],[f617,f358,f573])).
fof(f963,plain,(
  ( ! [X0,X1] : (sP9('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f618,f358,f573])).
fof(f964,plain,(
  ( ! [X0,X1] : (sP9('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X1)) )),
  inference(unit_resulting_resolution,[],[f619,f358,f573])).
fof(f965,plain,(
  ( ! [X0,X1] : (sP9('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X0,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f358,f806])).
fof(f966,plain,(
  ( ! [X0,X1] : (sP9('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),X1)) )),
  inference(unit_resulting_resolution,[],[f358,f816])).
fof(f898,plain,(
  ( ! [X0] : (sP7('F_0',X0,'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)))) )),
  inference(unit_resulting_resolution,[],[f613,f585,f574])).
fof(f899,plain,(
  ( ! [X0,X1] : (sP7('F_0',X0,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f632,f585,f574])).
fof(f901,plain,(
  ( ! [X0] : (sP7('F_0','const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)),X0)) )),
  inference(unit_resulting_resolution,[],[f615,f585,f574])).
fof(f902,plain,(
  ( ! [X0,X1] : (sP7('F_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f633,f585,f574])).
fof(f890,plain,(
  ( ! [X0,X1] : (sP9('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X0,X1) | ~sP8('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X0,X1)) )),
  inference(subsumption_resolution,[],[f888,f350])).
fof(f888,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~sP8('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X0,X1) | sP9('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X0,X1)) )),
  inference(superposition,[],[f573,f647])).
fof(f885,plain,(
  ( ! [X0] : (sP9('T_0',X0,'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) )),
  inference(unit_resulting_resolution,[],[f617,f350,f573])).
fof(f887,plain,(
  ( ! [X0] : (sP9('T_0','const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X0)) )),
  inference(unit_resulting_resolution,[],[f619,f350,f573])).
fof(f872,plain,(
  ( ! [X0,X1] : (sP11('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X0,X1) | ~sP10('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X0,X1)) )),
  inference(subsumption_resolution,[],[f870,f350])).
fof(f870,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~sP10('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X0,X1) | sP11('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X0,X1)) )),
  inference(superposition,[],[f572,f647])).
fof(f866,plain,(
  ( ! [X0] : (sP11('T_0',X0,'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)))) )),
  inference(unit_resulting_resolution,[],[f623,f350,f572])).
fof(f867,plain,(
  ( ! [X0,X1] : (sP11('T_0',X0,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f635,f350,f572])).
fof(f869,plain,(
  ( ! [X0] : (sP11('T_0','const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X0)) )),
  inference(unit_resulting_resolution,[],[f625,f350,f572])).
fof(f848,plain,(
  ( ! [X0,X1] : (sP11('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X0,X1) | ~p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f846,f350])).
fof(f846,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | sP11('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X0,X1)) )),
  inference(superposition,[],[f568,f647])).
fof(f845,plain,(
  ( ! [X0] : (sP11('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f350,f350,f568])).
fof(f842,plain,(
  ( ! [X0] : (sP11('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f838])).
fof(f838,plain,(
  ( ! [X0,X1] : (sP11('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X1,X0) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f836,f350])).
fof(f836,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | sP11('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X1,X0)) )),
  inference(superposition,[],[f567,f647])).
fof(f835,plain,(
  ( ! [X0] : (sP11('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f350,f567])).
fof(f827,plain,(
  ( ! [X0] : (~sP11('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f350,f585,f556])).
fof(f826,plain,(
  ( ! [X0] : (sP9('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f350,f816])).
fof(f825,plain,(
  ( ! [X0] : (~sP11('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f585,f555])).
fof(f816,plain,(
  ( ! [X0,X1] : (sP9('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X0,X1) | ~p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f814,f350])).
fof(f814,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | sP9('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X0,X1)) )),
  inference(superposition,[],[f551,f647])).
fof(f813,plain,(
  ( ! [X0] : (sP9('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f350,f350,f551])).
fof(f810,plain,(
  ( ! [X0] : (sP9('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f806])).
fof(f806,plain,(
  ( ! [X0,X1] : (sP9('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X1,X0) | ~p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f804,f350])).
fof(f804,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | sP9('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X1,X0)) )),
  inference(superposition,[],[f550,f647])).
fof(f803,plain,(
  ( ! [X0] : (sP9('T_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f350,f550])).
fof(f794,plain,(
  ( ! [X0] : (~sP9('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f350,f585,f539])).
fof(f788,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f585,f538])).
fof(f779,plain,(
  ( ! [X0] : (sP7('F_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f585,f585,f533])).
fof(f768,plain,(
  ( ! [X0] : (sP7('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f585,f532])).
fof(f760,plain,(
  ( ! [X0] : (~sP7('T_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f350,f585,f523])).
fof(f756,plain,(
  ( ! [X0] : (~sP7('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f350,f585,f522])).
fof(f715,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f418,f373])).
fof(f717,plain,(
  ( ! [X0,X1] : (~sP1('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f418,f478])).
fof(f719,plain,(
  ( ! [X0,X1] : (~sP1('F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f585,f418,f478])).
fof(f720,plain,(
  ( ! [X2,X0,X3,X1] : (~sP1('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f418,f418,f478])).
fof(f721,plain,(
  ( ! [X0,X1] : (sP1('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f350,f418,f480])).
fof(f722,plain,(
  ( ! [X0,X1] : (sP1('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f418,f483])).
fof(f723,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X0,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')),X3)) )),
  inference(unit_resulting_resolution,[],[f418,f507])).
fof(f724,plain,(
  ( ! [X2,X0,X3,X1] : (sP4('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),X2,X3)) )),
  inference(unit_resulting_resolution,[],[f418,f508])).
fof(f725,plain,(
  ( ! [X2,X0,X3,X1] : (sP3(X0,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')),X3)) )),
  inference(unit_resulting_resolution,[],[f418,f511])).
fof(f726,plain,(
  ( ! [X2,X0,X3,X1] : (sP3('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),X2,X3)) )),
  inference(unit_resulting_resolution,[],[f418,f512])).
fof(f727,plain,(
  ( ! [X2,X0,X3,X1] : (sP6(X0,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')),X3)) )),
  inference(unit_resulting_resolution,[],[f418,f530])).
fof(f728,plain,(
  ( ! [X2,X0,X3,X1] : (sP6('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),X2,X3)) )),
  inference(unit_resulting_resolution,[],[f418,f531])).
fof(f729,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),X2,X3)) )),
  inference(unit_resulting_resolution,[],[f418,f548])).
fof(f730,plain,(
  ( ! [X2,X0,X3,X1] : (sP10(X0,X1,'const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f418,f563])).
fof(f731,plain,(
  ( ! [X2,X0,X3,X1] : (sP10('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),X2,X3)) )),
  inference(unit_resulting_resolution,[],[f418,f565])).
fof(f732,plain,(
  ( ! [X0,X1] : (~sP0('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f418,f587])).
fof(f736,plain,(
  ( ! [X0,X1] : (~sP0('F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f585,f418,f587])).
fof(f737,plain,(
  ( ! [X2,X0,X3,X1] : (~sP0('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f418,f418,f587])).
fof(f740,plain,(
  ( ! [X0,X1] : (sP1('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f418,f691])).
fof(f741,plain,(
  ( ! [X0,X1] : (sP1('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) )),
  inference(unit_resulting_resolution,[],[f418,f697])).
fof(f714,plain,(
  sP1('F_0','const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))),
  inference(unit_resulting_resolution,[],[f585,f697])).
fof(f711,plain,(
  sP1('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'F_0')),
  inference(unit_resulting_resolution,[],[f585,f691])).
fof(f706,plain,(
  ~sP0('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f585,f585,f587])).
fof(f697,plain,(
  ( ! [X0] : (sP1(X0,'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f695,f350])).
fof(f695,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | sP1(X0,'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) )),
  inference(superposition,[],[f483,f647])).
fof(f694,plain,(
  sP1('F_0','T_0')),
  inference(unit_resulting_resolution,[],[f585,f350,f483])).
fof(f691,plain,(
  ( ! [X0] : (sP1('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X0) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f689,f350])).
fof(f689,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | sP1('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X0)) )),
  inference(superposition,[],[f480,f647])).
fof(f688,plain,(
  sP1('T_0','F_0')),
  inference(unit_resulting_resolution,[],[f585,f350,f480])).
fof(f685,plain,(
  ~sP1('T_0','T_0')),
  inference(unit_resulting_resolution,[],[f350,f350,f479])).
fof(f683,plain,(
  ~sP1('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f585,f585,f478])).
fof(f647,plain,(
  s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))),
  inference(unit_resulting_resolution,[],[f348,f374])).
fof(f642,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f585,f373])).
fof(f643,plain,(
  s(bool,'T_0') != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)))),
  inference(unit_resulting_resolution,[],[f349,f373])).
fof(f644,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f407,f373])).
fof(f628,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0')),X2)) )),
  inference(unit_resulting_resolution,[],[f407,f507])).
fof(f629,plain,(
  ( ! [X2,X0,X1] : (sP4('const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),X1,X2)) )),
  inference(unit_resulting_resolution,[],[f407,f508])).
fof(f630,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0')),X2)) )),
  inference(unit_resulting_resolution,[],[f407,f511])).
fof(f631,plain,(
  ( ! [X2,X0,X1] : (sP3('const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),X1,X2)) )),
  inference(unit_resulting_resolution,[],[f407,f512])).
fof(f632,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0')),X2)) )),
  inference(unit_resulting_resolution,[],[f407,f530])).
fof(f633,plain,(
  ( ! [X2,X0,X1] : (sP6('const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),X1,X2)) )),
  inference(unit_resulting_resolution,[],[f407,f531])).
fof(f634,plain,(
  ( ! [X2,X0,X1] : (sP8('const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),X1,X2)) )),
  inference(unit_resulting_resolution,[],[f407,f548])).
fof(f635,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f407,f563])).
fof(f636,plain,(
  ( ! [X2,X0,X1] : (sP10('const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),X1,X2)) )),
  inference(unit_resulting_resolution,[],[f407,f565])).
fof(f626,plain,(
  ( ! [X0,X1] : (sP10('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f585,f565])).
fof(f627,plain,(
  ( ! [X0,X1] : (sP10('const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f349,f565])).
fof(f624,plain,(
  ( ! [X0,X1] : (sP10(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f350,f564])).
fof(f625,plain,(
  ( ! [X0,X1] : (sP10(X0,'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X1)) )),
  inference(unit_resulting_resolution,[],[f348,f564])).
fof(f622,plain,(
  ( ! [X0,X1] : (sP10(X0,X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f585,f563])).
fof(f623,plain,(
  ( ! [X0,X1] : (sP10(X0,X1,'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)))) )),
  inference(unit_resulting_resolution,[],[f349,f563])).
fof(f620,plain,(
  ( ! [X0,X1] : (sP8('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f585,f548])).
fof(f621,plain,(
  ( ! [X0,X1] : (sP8('const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f349,f548])).
fof(f618,plain,(
  ( ! [X0,X1] : (sP8(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f350,f547])).
fof(f619,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X1)) )),
  inference(unit_resulting_resolution,[],[f348,f547])).
fof(f616,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f546])).
fof(f617,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) )),
  inference(unit_resulting_resolution,[],[f348,f546])).
fof(f614,plain,(
  ( ! [X0,X1] : (sP6('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f585,f531])).
fof(f615,plain,(
  ( ! [X0,X1] : (sP6('const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f349,f531])).
fof(f612,plain,(
  ( ! [X0,X1] : (sP6(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f585,f530])).
fof(f613,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)),X1)) )),
  inference(unit_resulting_resolution,[],[f349,f530])).
fof(f610,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f529])).
fof(f611,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) )),
  inference(unit_resulting_resolution,[],[f348,f529])).
fof(f608,plain,(
  ( ! [X0,X1] : (sP2('T_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f350,f516])).
fof(f609,plain,(
  ( ! [X0,X1] : (sP2('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f348,f516])).
fof(f606,plain,(
  ( ! [X0,X1] : (sP2(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f350,f515])).
fof(f607,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X1)) )),
  inference(unit_resulting_resolution,[],[f348,f515])).
fof(f604,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f514])).
fof(f605,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) )),
  inference(unit_resulting_resolution,[],[f348,f514])).
fof(f602,plain,(
  ( ! [X0,X1] : (sP3('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f585,f512])).
fof(f603,plain,(
  ( ! [X0,X1] : (sP3('const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f349,f512])).
fof(f600,plain,(
  ( ! [X0,X1] : (sP3(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f585,f511])).
fof(f601,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)),X1)) )),
  inference(unit_resulting_resolution,[],[f349,f511])).
fof(f598,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f510])).
fof(f599,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) )),
  inference(unit_resulting_resolution,[],[f348,f510])).
fof(f596,plain,(
  ( ! [X0,X1] : (sP4('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f585,f508])).
fof(f597,plain,(
  ( ! [X0,X1] : (sP4('const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f349,f508])).
fof(f594,plain,(
  ( ! [X0,X1] : (sP4(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f585,f507])).
fof(f595,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13)),X1)) )),
  inference(unit_resulting_resolution,[],[f349,f507])).
fof(f592,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f506])).
fof(f593,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) )),
  inference(unit_resulting_resolution,[],[f348,f506])).
fof(f590,plain,(
  ( ! [X0] : (sP0('T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f350,f465])).
fof(f591,plain,(
  ( ! [X0] : (sP0('const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X0)) )),
  inference(unit_resulting_resolution,[],[f348,f465])).
fof(f588,plain,(
  ( ! [X0] : (sP0(X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f350,f464])).
fof(f589,plain,(
  ( ! [X0] : (sP0(X0,'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) )),
  inference(unit_resulting_resolution,[],[f348,f464])).
fof(f566,plain,(
  ( ! [X2,X0,X1] : (sP11(X2,X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f346])).
fof(f346,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f345])).
fof(f345,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f237])).
fof(f237,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP11(X2,X0,X1))),
  inference(definition_folding,[],[f219,f236,f235])).
fof(f235,plain,(
  ! [X2,X0,X1] : (sP10(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f236,plain,(
  ! [X2,X0,X1] : (sP11(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f219,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f179])).
fof(f179,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f178])).
fof(f178,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X23,X20,X19] : ((p(s(bool,X19)) <=> (p(s(bool,X20)) => p(s(bool,X23)))) <=> ((p(s(bool,X19)) | p(s(bool,X20))) & (p(s(bool,X19)) | ~p(s(bool,X23))) & (~p(s(bool,X20)) | p(s(bool,X23)) | ~p(s(bool,X19)))))),
  file('thm.list.MONO_EVERY','thm.sat.sat_27_prove_ltb')).
fof(f567,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f346])).
fof(f568,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f346])).
fof(f569,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f346])).
fof(f562,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f344])).
fof(f344,plain,(
  ! [X0,X1,X2] : ((~sP10(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP10(X0,X1,X2)))),
  inference(rectify,[],[f343])).
fof(f343,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(flattening,[],[f342])).
fof(f342,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(nnf_transformation,[],[f235])).
fof(f563,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f344])).
fof(f564,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f344])).
fof(f565,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f344])).
fof(f555,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f341])).
fof(f341,plain,(
  ! [X0,X1,X2] : ((~sP11(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP10(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP10(X0,X1,X2) | sP11(X0,X1,X2)))),
  inference(rectify,[],[f340])).
fof(f340,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1) | sP11(X2,X0,X1)))),
  inference(flattening,[],[f339])).
fof(f339,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1)) | sP11(X2,X0,X1)))),
  inference(nnf_transformation,[],[f236])).
fof(f556,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f341])).
fof(f557,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f341])).
fof(f572,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP10(X0,X1,X2) | sP11(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f558])).
fof(f558,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f341])).
fof(f559,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f341])).
fof(f560,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f341])).
fof(f561,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f341])).
fof(f549,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f338])).
fof(f338,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f337])).
fof(f337,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f234])).
fof(f234,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP9(X2,X0,X1))),
  inference(definition_folding,[],[f177,f233,f232])).
fof(f232,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f233,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f177,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f176])).
fof(f176,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X23,X20,X19] : ((p(s(bool,X19)) <=> (p(s(bool,X20)) | p(s(bool,X23)))) <=> ((p(s(bool,X19)) | ~p(s(bool,X20))) & (p(s(bool,X19)) | ~p(s(bool,X23))) & (p(s(bool,X20)) | p(s(bool,X23)) | ~p(s(bool,X19)))))),
  file('thm.list.MONO_EVERY','thm.sat.sat_25_prove_ltb')).
fof(f550,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f338])).
fof(f551,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f338])).
fof(f552,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)) | ~sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f338])).
fof(f545,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f336])).
fof(f336,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f335])).
fof(f335,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f334])).
fof(f334,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f232])).
fof(f546,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f336])).
fof(f547,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f336])).
fof(f548,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f336])).
fof(f538,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f333])).
fof(f333,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP8(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f332])).
fof(f332,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f331])).
fof(f331,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1)) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f233])).
fof(f539,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f333])).
fof(f540,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f333])).
fof(f573,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f541])).
fof(f541,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f333])).
fof(f542,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f333])).
fof(f543,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f333])).
fof(f544,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f333])).
fof(f532,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f330])).
fof(f330,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f329])).
fof(f329,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f231])).
fof(f231,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP7(X2,X0,X1))),
  inference(definition_folding,[],[f175,f230,f229])).
fof(f229,plain,(
  ! [X0,X1,X2] : (sP6(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f230,plain,(
  ! [X2,X0,X1] : (sP7(X2,X0,X1) <=> (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f175,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f174])).
fof(f174,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f50])).
fof(f50,axiom,(
  ! [X23,X20,X19] : ((p(s(bool,X19)) <=> (p(s(bool,X20)) & p(s(bool,X23)))) <=> ((p(s(bool,X19)) | ~p(s(bool,X20)) | ~p(s(bool,X23))) & (p(s(bool,X20)) | ~p(s(bool,X19))) & (p(s(bool,X23)) | ~p(s(bool,X19)))))),
  file('thm.list.MONO_EVERY','thm.sat.sat_23_prove_ltb')).
fof(f533,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f330])).
fof(f534,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f330])).
fof(f537,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP7(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f330])).
fof(f528,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f328])).
fof(f328,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(flattening,[],[f327])).
fof(f327,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(nnf_transformation,[],[f229])).
fof(f529,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f328])).
fof(f530,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f328])).
fof(f531,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f328])).
fof(f521,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f326])).
fof(f326,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | (sP6(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP6(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f325])).
fof(f325,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP7(X2,X0,X1)))),
  inference(flattening,[],[f324])).
fof(f324,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP7(X2,X0,X1)))),
  inference(nnf_transformation,[],[f230])).
fof(f522,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f326])).
fof(f523,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f326])).
fof(f524,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f326])).
fof(f574,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(duplicate_literal_removal,[],[f527])).
fof(f527,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f326])).
fof(f517,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f323])).
fof(f323,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP5(X2,X1,X0)) & (~sP5(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f228])).
fof(f228,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP5(X2,X1,X0))),
  inference(definition_folding,[],[f173,f227,f226,f225,f224])).
fof(f224,plain,(
  ! [X0,X1,X2] : (sP2(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f225,plain,(
  ! [X1,X0,X2] : (sP3(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f226,plain,(
  ! [X2,X0,X1] : (sP4(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f227,plain,(
  ! [X2,X1,X0] : (sP5(X2,X1,X0) <=> (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f173,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f172])).
fof(f172,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f48])).
fof(f48,axiom,(
  ! [X23,X20,X19] : ((p(s(bool,X19)) <=> s(bool,X20) = s(bool,X23)) <=> ((p(s(bool,X19)) | p(s(bool,X20)) | p(s(bool,X23))) & (p(s(bool,X19)) | ~p(s(bool,X23)) | ~p(s(bool,X20))) & (p(s(bool,X20)) | ~p(s(bool,X23)) | ~p(s(bool,X19))) & (p(s(bool,X23)) | ~p(s(bool,X20)) | ~p(s(bool,X19)))))),
  file('thm.list.MONO_EVERY','thm.sat.sat_2|49|_prove_ltb')).
fof(f518,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f323])).
fof(f519,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f323])).
fof(f520,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f323])).
fof(f513,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f322])).
fof(f322,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(flattening,[],[f321])).
fof(f321,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(nnf_transformation,[],[f224])).
fof(f514,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f322])).
fof(f515,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f322])).
fof(f516,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f322])).
fof(f509,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f320])).
fof(f320,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f319])).
fof(f319,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(flattening,[],[f318])).
fof(f318,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(nnf_transformation,[],[f225])).
fof(f510,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f320])).
fof(f511,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f320])).
fof(f512,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f320])).
fof(f505,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f317])).
fof(f317,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f316])).
fof(f316,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(flattening,[],[f315])).
fof(f315,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(nnf_transformation,[],[f226])).
fof(f506,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f317])).
fof(f507,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f317])).
fof(f508,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f317])).
fof(f498,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f314])).
fof(f314,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | (sP2(X2,X1,X0) & sP3(X1,X2,X0) & sP4(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP2(X2,X1,X0) | ~sP3(X1,X2,X0) | ~sP4(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f313])).
fof(f313,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP5(X2,X1,X0)))),
  inference(flattening,[],[f312])).
fof(f312,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP5(X2,X1,X0)))),
  inference(nnf_transformation,[],[f227])).
fof(f499,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP3(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f314])).
fof(f500,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP4(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f314])).
fof(f501,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f314])).
fof(f502,plain,(
  ( ! [X2,X0,X1] : (~sP2(X2,X1,X0) | ~p(s(bool,X2)) | ~sP4(X0,X2,X1) | ~sP3(X1,X2,X0) | sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f314])).
fof(f503,plain,(
  ( ! [X2,X0,X1] : (~sP2(X2,X1,X0) | p(s(bool,X1)) | ~sP4(X0,X2,X1) | ~sP3(X1,X2,X0) | sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f314])).
fof(f504,plain,(
  ( ! [X2,X0,X1] : (~sP2(X2,X1,X0) | p(s(bool,X0)) | ~sP4(X0,X2,X1) | ~sP3(X1,X2,X0) | sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f314])).
fof(f497,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK30(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK30(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f311])).
fof(f311,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK30(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK30(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK30])],[f218])).
fof(f218,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f171])).
fof(f171,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f170])).
fof(f170,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.list.MONO_EVERY',aHL_EXT)).
fof(f478,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f305])).
fof(f305,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1)))),
  inference(rectify,[],[f304])).
fof(f304,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0)))),
  inference(flattening,[],[f303])).
fof(f303,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP1(X1,X0)))),
  inference(nnf_transformation,[],[f222])).
fof(f222,plain,(
  ! [X1,X0] : (sP1(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f479,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f305])).
fof(f480,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f305])).
fof(f483,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f305])).
fof(f587,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f463,f585])).
fof(f463,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f298])).
fof(f298,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f297])).
fof(f297,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f220])).
fof(f220,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f464,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f298])).
fof(f465,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f298])).
fof(f451,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f203])).
fof(f203,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f202])).
fof(f202,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f149])).
fof(f149,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X13,X14] : ((p(s(bool,X13)) => p(s(bool,X14))) => ((p(s(bool,X14)) => p(s(bool,X13))) => s(bool,X13) = s(bool,X14)))),
  file('thm.list.MONO_EVERY','thm.bool.IMP_ANTISYM_AX')).
fof(f454,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f203])).
fof(f422,plain,(
  ( ! [X2,X0,X1] : (s(X0,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s('type.list.list'(X0),X1))) = s(X0,'const.list.EL_2'(s('type.num.num',X2),s('type.list.list'(X0),'const.list.TL_1'(s('type.list.list'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X0,X1,X2] : s(X0,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s('type.list.list'(X0),X1))) = s(X0,'const.list.EL_2'(s('type.num.num',X2),s('type.list.list'(X0),'const.list.TL_1'(s('type.list.list'(X0),X1)))))),
  inference(flattening,[],[f129])).
fof(f129,plain,(
  ! [X0] : ! [X1,X2] : s(X0,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s('type.list.list'(X0),X1))) = s(X0,'const.list.EL_2'(s('type.num.num',X2),s('type.list.list'(X0),'const.list.TL_1'(s('type.list.list'(X0),X1)))))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X6] : ! [X7,X11] : s(X6,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X11))),s('type.list.list'(X6),X7))) = s(X6,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(X6),'const.list.TL_1'(s('type.list.list'(X6),X7)))))),
  file('thm.list.MONO_EVERY','thm.list.EL|split|1')).
fof(f421,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2))))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))))) )),
  inference(cnf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0,X1,X2] : s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2))))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))))),
  inference(flattening,[],[f127])).
fof(f127,plain,(
  ! [X0] : ! [X1,X2] : s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2))))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X6] : ! [X17,X15] : s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X15))))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X17),s('type.list.list'(X6),X15)))))),
  file('thm.list.MONO_EVERY','thm.list.LENGTH|split|1')).
fof(f420,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'(X0),X2) = s('type.list.list'(X0),'const.list.TL_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))))) )),
  inference(cnf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0,X1,X2] : s('type.list.list'(X0),X2) = s('type.list.list'(X0),'const.list.TL_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))))),
  inference(flattening,[],[f125])).
fof(f125,plain,(
  ! [X0] : ! [X1,X2] : s('type.list.list'(X0),X2) = s('type.list.list'(X0),'const.list.TL_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X6] : ! [X17,X15] : s('type.list.list'(X6),X15) = s('type.list.list'(X6),'const.list.TL_1'(s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X17),s('type.list.list'(X6),X15)))))),
  file('thm.list.MONO_EVERY','thm.list.TL')).
fof(f419,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.list.HD_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))))) )),
  inference(cnf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.list.HD_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))))),
  inference(flattening,[],[f123])).
fof(f123,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.list.HD_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X6] : ! [X17,X15] : s(X6,X17) = s(X6,'const.list.HD_1'(s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X17),s('type.list.list'(X6),X15)))))),
  file('thm.list.MONO_EVERY','thm.list.HD')).
fof(f418,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(cnf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0,X1] : ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),
  inference(flattening,[],[f121])).
fof(f121,plain,(
  ! [X0] : ! [X1] : ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X6] : ! [X10] : ~p(s(bool,'const.bool.IN_2'(s(X6,X10),s(fun(X6,bool),'const.pred_set.EMPTY_0'))))),
  file('thm.list.MONO_EVERY','thm.pred_set.NOT_IN_EMPTY')).
fof(f409,plain,(
  ( ! [X0,X1] : (s(X0,'const.list.HD_1'(s('type.list.list'(X0),X1))) = s(X0,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(X0),X1)))) )),
  inference(cnf_transformation,[],[f113])).
fof(f113,plain,(
  ! [X0,X1] : s(X0,'const.list.HD_1'(s('type.list.list'(X0),X1))) = s(X0,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(X0),X1)))),
  inference(flattening,[],[f112])).
fof(f112,plain,(
  ! [X0] : ! [X1] : s(X0,'const.list.HD_1'(s('type.list.list'(X0),X1))) = s(X0,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(X0),X1)))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X6] : ! [X7] : s(X6,'const.list.HD_1'(s('type.list.list'(X6),X7))) = s(X6,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(X6),X7)))),
  file('thm.list.MONO_EVERY','thm.list.EL|split|0')).
fof(f408,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(cnf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X0,X1] : s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))),
  inference(flattening,[],[f110])).
fof(f110,plain,(
  ! [X0] : ! [X1] : s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X6] : ! [X9] : s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X6,bool),X9),s('type.list.list'(X6),'const.list.NIL_0')))),
  file('thm.list.MONO_EVERY','thm.list.EVERY_DEF|split|0')).
fof(f407,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))) )),
  inference(cnf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))),
  inference(flattening,[],[f108])).
fof(f108,plain,(
  ! [X0] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X11] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.num.|48|_0'))))),
  file('thm.list.MONO_EVERY','thm.prim_rec.NOT_LESS_|48|')).
fof(f373,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f243])).
fof(f243,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f91])).
fof(f91,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X15] : (s(bool,'T_0') = s(bool,X15) <=> p(s(bool,X15)))),
  file('thm.list.MONO_EVERY','thm.bool.EQ_CLAUSES|split|0')).
fof(f374,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f243])).
fof(f360,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.list.MONO_EVERY',aHL_BOOL_CASES)).
fof(f358,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))) )),
  inference(cnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X11] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X11))))))),
  file('thm.list.MONO_EVERY','thm.prim_rec.LESS_|48|')).
fof(f357,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.NIL_0')))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X6] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),'const.list.NIL_0')))),
  file('thm.list.MONO_EVERY','thm.list.LENGTH|split|0')).
fof(f350,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.list.MONO_EVERY',aHL_TRUTH)).
fof(f348,plain,(
  p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))))),
  inference(cnf_transformation,[],[f238])).
fof(f349,plain,(
  ~p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13))))),
  inference(cnf_transformation,[],[f238])).
fof(f58051,plain,(
  p(s(bool,'F_0'))),
  inference(forward_demodulation,[],[f57072,f667])).
fof(f57072,plain,(
  p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK13))))),
  inference(unit_resulting_resolution,[],[f55933,f430])).
fof(f55933,plain,(
  p(s(bool,happ(s(fun(sK12,bool),sK14),s(sK12,sK22(sK13,sK14,sK12)))))),
  inference(unit_resulting_resolution,[],[f55931,f347])).
fof(f55931,plain,(
  p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(sK13,sK14,sK12)))))),
  inference(subsumption_resolution,[],[f55919,f350])).
fof(f55919,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(sK13,sK14,sK12)))))),
  inference(superposition,[],[f3302,f26815])).
fof(f26815,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK22(sK13,sK14,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))),
  inference(subsumption_resolution,[],[f26801,f12261])).
fof(f26801,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK22(sK13,sK14,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))),
  inference(superposition,[],[f2902,f667])).
% SZS output end Proof for thm.list
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 1585
% Passive clauses: 6580
% Generated clauses: 58850
% Final active clauses: 1123
% Final passive clauses: 2365
% Input formulas: 69
% Initial clauses: 225
% 
% Duplicate literals: 1153
% Trivial inequalities: 66
% Fw subsumption resolutions: 1373
% Fw demodulations: 10594
% Bw demodulations: 1706
% Condensations: 1469
% Global subsumptions: 215
% 
% Simple tautologies: 126
% Equational tautologies: 81
% Forward subsumptions: 36816
% Bw demodulations to eq. taut.: 115
% 
% Unit resulting resolution: 14244
% Factoring: 1
% Forward superposition: 19162
% Backward superposition: 6559
% Self superposition: 41
% Equality factoring: 28
% Equality resolution: 8
% 
% Split clauses: 150
% Split components: 322
% SAT solver clauses: 7317
% SAT solver unit clauses: 1957
% SAT solver binary clauses: 2197
% 
% Sat splits: 150
% Sat splitting refutations: 220
% 
% Memory used [KB]: 38506
% Time elapsed: 2.283 s
% ------------------------------
% ------------------------------
% Success in time 3.62 s
