Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.pred_set
% (56178)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 644
% Passive clauses: 9863
% Generated clauses: 18088
% Final active clauses: 568
% Final passive clauses: 8701
% Input formulas: 45
% Initial clauses: 202
% 
% Duplicate literals: 278
% Trivial inequalities: 87
% Fw subsumption resolutions: 1508
% Fw demodulations: 107
% Bw demodulations: 75
% 
% Simple tautologies: 226
% Equational tautologies: 65
% Forward subsumptions: 5778
% 
% Binary resolution: 11337
% Factoring: 984
% Forward superposition: 1439
% Backward superposition: 1880
% Self superposition: 24
% Equality factoring: 26
% Equality resolution: 31
% 
% Split clauses: 294
% Split components: 590
% SAT solver clauses: 619
% SAT solver unit clauses: 24
% SAT solver binary clauses: 143
% 
% Sat splits: 294
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 617
% TWLsolver calls for satisfiability: 354
% 
% Memory used [KB]: 8827
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 294 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pred_set
% (56197)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 434
% Passive clauses: 2922
% Generated clauses: 12594
% Final active clauses: 411
% Final passive clauses: 2096
% Input formulas: 45
% Initial clauses: 202
% 
% Duplicate literals: 221
% Trivial inequalities: 55
% Fw subsumption resolutions: 1008
% Fw demodulations: 597
% Bw demodulations: 412
% Condensations: 355
% Global subsumptions: 13
% 
% Simple tautologies: 293
% Equational tautologies: 18
% Forward subsumptions: 4152
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 4064
% Factoring: 225
% Forward superposition: 1378
% Backward superposition: 1070
% Self superposition: 26
% Equality factoring: 30
% Equality resolution: 13
% 
% Split clauses: 46
% Split components: 93
% SAT solver clauses: 3087
% SAT solver unit clauses: 182
% SAT solver binary clauses: 83
% SAT solver learnt clauses: 26
% SAT solver learnt literals: 10
% 
% Sat splits: 46
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 3035
% TWLsolver calls for satisfiability: 1808
% 
% Memory used [KB]: 13048
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f31693,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f268,f482,f480,f478,f267,f481,f266,f479,f477,f2368,f2372,f2375,f2379,f2382,f2385,f2388,f2391,f2394,f2397,f2400,f2403,f2406,f2409,f2412,f2415,f2418,f2421,f2469,f2369,f2471,f2373,f2473,f2376,f2476,f2380,f2478,f2383,f2480,f2386,f2482,f2389,f2484,f2392,f2486,f2395,f2488,f2398,f2490,f2401,f2492,f2404,f2494,f2407,f2496,f2410,f2498,f2413,f2500,f2416,f2502,f2419,f2504,f2370,f2422,f2623,f2624,f3744,f3747,f3745,f3811,f3812,f3819,f3816,f3813,f3814,f3822,f3815,f3824,f3817,f3821,f4214,f4215,f4228,f4226,f4897,f4901,f4899,f12413,f12414,f12415,f12416,f12418,f3746,f12419,f12422,f12424,f12426,f12434,f12436,f12438,f12440,f12442,f12444,f12446,f12448,f12452,f12454,f12456,f12458,f12460,f12462,f12464,f12466,f12468,f12470,f12506,f12508,f12544,f12546,f12567,f12588,f12628,f12703,f12744,f12819,f13248,f13252,f13256,f13260,f13269,f13274,f13278,f13282,f13287,f13293,f13323,f13328,f13360,f13377,f13399,f13416,f13418,f13438,f13453,f13473,f13488,f13490,f13492,f13494,f13496,f16259,f16260,f16261,f16262,f16263,f16264,f16265,f16266,f16267,f16268,f13556,f13558,f13560,f13562,f13564,f13566,f13568,f13570,f13572,f13574,f13576,f13578,f13580,f13582,f13657,f13659,f13661,f13663,f13665,f13667,f13669,f13671,f13706,f13741,f13799,f13801,f13803,f13805,f13807,f13812,f13829,f13831,f13833,f13835,f13837,f13839,f13841,f13843,f13845,f13847,f13857,f13859,f13920,f13922,f13924,f13926,f13928,f13936,f13954,f13956,f13958,f13960,f13962,f13964,f13966,f13968,f13970,f13972,f13982,f13984,f14007,f14022,f14042,f14057,f14059,f14105,f14107,f14132,f14134,f14184,f14186,f14211,f14213,f14264,f14266,f14291,f14293,f14322,f14357,f14386,f14421,f14467,f14513,f14563,f14565,f14567,f14573,f14576,f14578,f14634,f14637,f14639,f14641,f14643,f14645,f14647,f14665,f14667,f14669,f14698,f14700,f14729,f14731,f14760,f14762,f14791,f14793,f14839,f14873,f14940,f14942,f14944,f14950,f14953,f14955,f14978,f14980,f14982,f14984,f14986,f14988,f14990,f14992,f14994,f14996,f15003,f15005,f15008,f15010,f15012,f15014,f15016,f15150,f15152,f15154,f15156,f15158,f15160,f15162,f15164,f15166,f15168,f15170,f15172,f15174,f15256,f15258,f15260,f15262,f15264,f15266,f15268,f15270,f15272,f15274,f15356,f15358,f15360,f15362,f15364,f15366,f15368,f15370,f15372,f15374,f15407,f15409,f15442,f15444,f15446,f15474,f15507,f15540,f15542,f15575,f15577,f15601,f15635,f15673,f15707,f15740,f15742,f15744,f15746,f15748,f15750,f15752,f15754,f15756,f15758,f15760,f15762,f15764,f15793,f15795,f15797,f15799,f15801,f15803,f15805,f15807,f15809,f15811,f15813,f15815,f15817,f15819,f15825,f15830,f15833,f15840,f15842,f15844,f15846,f15848,f15854,f15856,f15858,f15860,f15862,f15868,f15870,f15872,f15874,f15876,f15879,f15881,f15883,f15885,f15887,f15899,f15901,f15903,f15905,f15907,f15910,f15912,f15914,f15937,f15949,f15959,f15969,f15988,f15990,f15992,f16002,f16004,f16013,f16020,f16027,f16029,f16031,f16033,f16035,f16041,f16043,f16045,f16047,f16049,f16055,f16057,f16059,f16061,f16063,f16068,f16070,f16072,f16074,f16076,f16087,f16089,f16095,f16097,f16099,f16101,f16103,f16105,f16108,f16119,f16121,f16123,f16125,f16127,f16133,f16135,f16137,f16139,f16141,f16147,f16149,f16151,f16153,f16155,f16158,f16160,f16162,f16164,f16166,f16174,f16176,f16182,f16184,f16186,f16188,f16190,f16192,f16195,f16269,f16270,f16271,f16272,f16273,f16274,f16275,f16276,f16277,f16278,f16279,f16280,f16281,f16282,f16283,f16284,f16285,f16286,f16287,f16288,f16289,f16290,f16291,f16292,f16293,f16294,f16295,f16296,f16297,f16298,f16299,f16300,f16301,f16302,f16303,f16304,f16305,f16306,f16307,f16308,f16309,f16310,f16311,f16312,f16313,f16314,f16315,f16316,f16317,f16318,f16319,f16320,f16321,f16322,f16323,f16324,f16325,f16326,f16327,f16328,f16329,f16330,f16331,f16332,f16333,f16334,f16335,f16336,f16337,f16338,f16339,f16340,f16341,f16342,f16343,f16344,f16345,f16346,f16347,f16348,f16349,f16350,f16351,f16352,f16353,f16354,f16355,f16356,f16357,f16358,f16359,f16360,f16361,f16362,f16363,f16364,f16365,f16366,f16367,f16368,f16369,f16370,f16371,f16372,f16373,f16374,f16375,f16376,f16377,f16378,f16379,f16380,f16381,f16382,f16383,f16384,f16385,f16386,f16387,f16388,f16389,f16390,f16391,f16392,f16393,f16394,f16395,f16396,f16397,f16398,f16399,f16400,f16401,f16402,f16403,f16404,f16405,f16406,f16407,f16408,f16409,f16410,f16411,f16412,f16413,f16414,f16415,f16416,f16417,f16418,f16419,f16420,f16421,f16422,f16423,f16424,f16425,f16426,f16427,f16428,f16429,f16430,f16431,f16432,f16433,f16434,f16435,f16436,f16437,f16438,f16439,f16440,f16441,f16442,f16443,f16444,f16445,f16446,f16447,f16448,f16449,f16450,f16451,f16452,f16453,f16454,f16455,f16456,f16457,f16458,f16459,f16460,f16461,f16462,f16463,f16464,f16465,f16466,f16467,f16468,f16469,f16470,f16471,f16472,f16473,f16474,f16475,f16476,f16477,f16478,f16479,f16480,f16481,f16482,f16483,f16484,f16485,f16486,f16487,f16488,f16489,f16490,f16491,f16492,f16493,f16494,f16495,f16496,f16497,f16498,f16499,f16500,f16501,f16502,f16503,f16504,f16505,f16506,f16507,f16508,f16509,f16510,f16511,f16512,f16513,f16514,f16515,f16516,f16517,f16518,f16519,f16520,f16521,f16522,f16523,f16524,f16525,f16526,f16527,f16528,f16529,f16530,f16531,f16532,f16533,f16534,f16535,f16536,f16537,f16538,f16539,f16540,f16541,f16542,f16543,f16544,f16545,f16546,f16547,f16548,f16549,f16550,f16551,f16552,f16553,f16554,f16555,f16556,f16557,f16558,f16559,f16560,f16561,f16562,f16563,f16564,f16565,f16566,f16567,f16568,f16569,f16570,f16571,f16572,f16573,f16574,f16575,f16576,f16577,f16578,f16579,f16580,f16581,f16582,f16583,f16584,f16585,f16586,f16587,f16588,f16589,f16590,f16591,f16592,f16593,f16594,f16595,f16596,f16597,f16598,f16599,f16600,f16601,f16602,f16603,f16604,f16605,f16606,f16607,f16608,f16609,f16610,f16611,f16612,f16613,f16614,f16615,f16616,f16617,f16618,f16619,f16620,f16621,f16622,f16623,f16624,f16625,f16626,f16627,f16628,f16629,f16630,f16631,f16632,f16633,f16634,f16635,f16636,f16637,f16638,f16639,f16640,f16641,f16642,f16643,f16644,f16645,f16646,f16647,f16648,f16649,f16650,f16651,f16652,f16653,f16654,f16655,f16656,f16657,f16658,f16659,f16660,f16661,f16662,f16663,f16664,f16665,f16666,f16667,f16668,f16669,f16670,f16671,f16672,f16673,f16674,f16675,f16676,f16677,f16678,f16679,f16680,f16681,f16682,f16683,f16684,f16685,f16686,f16687,f16688,f16689,f16690,f16691,f16692,f16693,f16694,f16695,f16696,f16697,f16698,f16699,f16700,f16701,f16702,f16703,f16704,f16705,f16706,f16707,f16708,f16709,f16710,f16711,f16712,f16713,f16714,f16715,f16716,f16717,f16718,f16719,f16720,f16721,f16722,f16723,f16724,f16725,f16726,f16727,f16728,f16729,f16730,f16731,f16732,f16733,f16734,f16735,f16736,f16737,f16738,f16739,f16740,f16741,f16742,f16743,f16744,f16745,f16746,f16747,f16748,f16749,f16750,f16751,f16752,f16753,f16754,f16755,f16756,f16757,f16758,f16759,f16760,f16761,f16762,f16763,f16764,f16765,f16766,f16767,f16768,f16769,f16770,f16771,f16772,f16773,f16774,f16775,f16776,f16777,f16778,f16779,f16780,f16781,f16782,f16783,f16784,f16785,f16786,f16787,f16788,f16789,f16790,f16791,f16792,f16793,f16794,f16795,f16796,f16797,f16798,f16799,f16800,f16801,f16802,f16803,f16804,f16805,f16806,f16807,f16808,f16809,f16810,f16811,f16812,f16813,f16814,f16815,f16816,f16817,f16818,f16819,f16820,f16821,f16822,f16823,f16824,f16825,f16826,f16827,f16828,f16829,f16830,f16831,f16832,f16833,f16834,f16835,f16836,f16837,f16838,f16839,f16840,f16841,f16842,f16843,f16844,f16845,f16846,f16847,f16848,f16849,f16850,f16851,f16852,f16853,f16854,f16855,f16856,f16857,f16858,f16859,f16860,f16861,f16862,f16863,f16864,f16865,f16866,f16867,f16868,f16869,f16870,f16871,f16872,f16873,f16874,f16875,f16876,f16877,f16878,f16879,f16880,f16881,f16882,f16883,f16884,f16885,f16886,f16887,f16888,f16889,f16890,f16891,f16892,f16893,f16894,f16895,f16896,f16897,f16898,f16899,f16900,f16901,f16902,f16903,f16904,f16905,f16906,f16907,f16908,f16909,f16910,f16911,f16912,f16913,f16914,f16915,f16916,f16917,f16918,f16919,f16920,f16921,f16922,f16923,f16924,f16925,f16926,f16927,f16928,f16929,f16930,f16931,f16932,f16933,f16934,f16935,f16936,f16937,f16938,f16939,f16940,f16941,f16942,f16943,f16944,f16945,f16946,f16947,f16948,f16949,f16950,f16951,f16952,f16953,f16954,f16955,f16956,f16957,f16958,f16959,f16960,f16961,f16962,f16963,f16964,f16965,f16966,f16967,f16968,f16969,f16970,f16971,f16972,f16973,f16974,f16975,f16976,f16977,f16978,f16979,f16980,f16981,f16983,f16984,f16986,f16987,f16988,f16990,f16991,f16993,f16994,f16996,f16998,f17000,f17001,f17002,f17005,f17006,f17007,f17010,f17012,f17014,f17016,f17018,f17020,f17022,f17023,f17025,f17027,f17029,f17031,f17032,f17033,f17034,f17035,f17036,f17039,f17040,f17041,f17042,f17043,f17044,f17046,f17047,f17048,f17049,f17050,f17051,f17052,f17053,f17054,f17055,f17056,f17057,f17058,f17059,f17060,f17061,f17062,f17063,f17064,f17065,f17078,f17079,f17080,f17081,f17082,f17083,f17084,f17085,f17086,f17087,f17088,f17089,f17090,f17091,f17092,f17093,f17094,f17095,f17096,f17097,f17098,f17099,f17100,f17101,f17102,f17103,f17104,f17105,f17106,f17107,f17108,f17109,f17110,f17111,f17112,f17113,f17114,f17115,f17116,f17117,f17118,f17119,f17120,f17121,f17122,f17123,f17124,f17125,f17126,f17127,f17128,f17129,f17130,f17131,f17132,f17133,f17134,f17135,f17136,f17137,f17138,f17139,f17140,f17141,f17142,f17143,f17144,f17145,f17146,f17147,f17148,f17149,f17150,f17151,f17152,f17153,f17154,f17155,f17156,f17157,f17158,f17159,f17160,f17161,f17162,f17163,f17164,f17165,f17166,f17167,f17168,f17169,f17170,f17171,f17172,f17173,f17174,f17175,f17176,f17177,f17178,f17179,f17180,f17181,f17182,f17183,f17184,f17185,f17186,f17187,f17188,f17189,f17190,f17191,f17192,f17193,f17194,f17195,f17196,f17197,f17198,f17199,f17200,f17201,f17202,f17203,f17204,f17205,f17206,f17207,f17208,f17209,f17210,f17211,f17212,f17213,f17214,f17215,f17216,f17217,f17218,f17219,f17220,f17221,f17222,f17223,f17224,f17225,f17226,f17227,f17228,f17229,f17230,f17231,f17232,f17233,f17234,f17235,f17236,f17237,f17238,f17239,f17240,f17241,f17242,f17243,f17244,f17245,f17246,f17247,f17248,f17249,f17250,f17251,f17252,f17253,f17254,f17255,f17256,f17257,f17258,f17259,f17260,f17261,f17262,f17263,f17264,f17265,f17266,f17267,f17268,f17269,f17270,f17271,f17272,f17273,f17274,f17275,f17276,f17277,f17278,f17279,f17280,f17281,f17282,f17283,f17284,f17285,f17286,f17287,f17288,f17289,f17290,f17291,f17292,f17293,f17294,f17295,f17296,f17297,f17298,f17299,f17300,f17301,f22593,f26614,f30231,f31688])).
fof(f31688,plain,(
  $false | ($spl1 | $spl2 | $spl66)),
  inference(subsumption_resolution,[],[f31687,f483])).
fof(f483,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f272,f269])).
fof(f269,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.UNION_SUBSET','thm.bool.TRUTH')).
fof(f272,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f164])).
fof(f164,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f49])).
fof(f49,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f19])).
fof(f19,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.pred_set.UNION_SUBSET','thm.bool.NOT_CLAUSES|split|2')).
fof(f31687,plain,(
  p(s(bool,'F_0')) | ($spl1 | $spl2 | $spl66)),
  inference(forward_demodulation,[],[f31181,f28946])).
fof(f28946,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),sK12)),s(fun(sK12,bool),sK15))) | $spl1),
  inference(unit_resulting_resolution,[],[f527,f26977,f832])).
fof(f832,plain,(
  ( ! [X4,X5] : (p(s(bool,X5)) | s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f827,f483])).
fof(f827,plain,(
  ( ! [X4,X5] : (p(s(bool,'F_0')) | p(s(bool,X5)) | s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f338,f276])).
fof(f276,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.pred_set.UNION_SUBSET',aHL_BOOL_CASES)).
fof(f338,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f133])).
fof(f133,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f92])).
fof(f92,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X10,X11] : ((p(s(bool,X10)) => p(s(bool,X11))) => ((p(s(bool,X11)) => p(s(bool,X10))) => s(bool,X10) = s(bool,X11)))),
  file('thm.pred_set.UNION_SUBSET','thm.bool.IMP_ANTISYM_AX')).
fof(f26977,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),sK12)),s(fun(sK12,bool),sK15)))) | $spl1),
  inference(unit_resulting_resolution,[],[f478,f318])).
fof(f318,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,sK18(X2,X1,X0)),s(fun(X0,bool),X2)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f188])).
fof(f188,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))))) & ((p(s(bool,'const.bool.IN_2'(s(X0,sK18(X2,X1,X0)),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,sK18(X2,X1,X0)),s(fun(X0,bool),X2))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK18])],[f187])).
fof(f187,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))))) & (? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f186])).
fof(f186,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & (? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(nnf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(ennf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) => p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(flattening,[],[f80])).
fof(f80,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) => p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X6] : ! [X7,X8] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))) <=> ! [X13] : (p(s(bool,'const.bool.IN_2'(s(X6,X13),s(fun(X6,bool),X7)))) => p(s(bool,'const.bool.IN_2'(s(X6,X13),s(fun(X6,bool),X8))))))),
  file('thm.pred_set.UNION_SUBSET','thm.pred_set.SUBSET_DEF')).
fof(f527,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f483,f285])).
fof(f285,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f166])).
fof(f166,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X8] : (s(bool,'T_0') = s(bool,X8) <=> p(s(bool,X8)))),
  file('thm.pred_set.UNION_SUBSET','thm.bool.EQ_CLAUSES|split|1')).
fof(f31181,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),sK12)),s(fun(sK12,bool),sK15)))) | ($spl2 | $spl66)),
  inference(unit_resulting_resolution,[],[f4901,f27663])).
fof(f27663,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13)))) | p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK15))))) ) | $spl2),
  inference(subsumption_resolution,[],[f27605,f269])).
fof(f27605,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13)))) | p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK15))))) ) | $spl2),
  inference(superposition,[],[f316,f26624])).
fof(f26624,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | $spl2),
  inference(unit_resulting_resolution,[],[f269,f479,f341])).
fof(f341,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f134])).
fof(f316,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f188])).
fof(f30231,plain,(
  $false | ($spl1 | $spl4 | $spl64)),
  inference(subsumption_resolution,[],[f30230,f483])).
fof(f30230,plain,(
  p(s(bool,'F_0')) | ($spl1 | $spl4 | $spl64)),
  inference(forward_demodulation,[],[f29742,f28946])).
fof(f29742,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),sK12)),s(fun(sK12,bool),sK15)))) | ($spl4 | $spl64)),
  inference(unit_resulting_resolution,[],[f4899,f27769])).
fof(f27769,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))) | p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK15))))) ) | $spl4),
  inference(subsumption_resolution,[],[f27711,f269])).
fof(f27711,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))) | p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK15))))) ) | $spl4),
  inference(superposition,[],[f316,f26800])).
fof(f26800,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))) | $spl4),
  inference(unit_resulting_resolution,[],[f269,f481,f341])).
fof(f26614,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f26613,f269])).
fof(f26613,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl3)),
  inference(forward_demodulation,[],[f26052,f22599])).
fof(f22599,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK13))) | $spl3),
  inference(unit_resulting_resolution,[],[f480,f1366])).
fof(f1366,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,sK18(X1,X2,X0)),s(fun(X0,bool),X2)))) )),
  inference(subsumption_resolution,[],[f1365,f483])).
fof(f1365,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,sK18(X1,X2,X0)),s(fun(X0,bool),X2)))) )),
  inference(superposition,[],[f317,f276])).
fof(f317,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK18(X2,X1,X0)),s(fun(X0,bool),X1)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f188])).
fof(f26052,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) | ($spl0 | $spl3)),
  inference(unit_resulting_resolution,[],[f23117,f336])).
fof(f336,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))) )),
  inference(cnf_transformation,[],[f206])).
fof(f206,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) & ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))))),
  inference(flattening,[],[f205])).
fof(f205,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))))),
  inference(nnf_transformation,[],[f91])).
fof(f91,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(flattening,[],[f90])).
fof(f90,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X6] : ! [X7,X8,X13] : (p(s(bool,'const.bool.IN_2'(s(X6,X13),s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X6,X13),s(fun(X6,bool),X7)))) | p(s(bool,'const.bool.IN_2'(s(X6,X13),s(fun(X6,bool),X8))))))),
  file('thm.pred_set.UNION_SUBSET','thm.pred_set.IN_UNION')).
fof(f23117,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))) | ($spl0 | $spl3)),
  inference(unit_resulting_resolution,[],[f22600,f7512])).
fof(f7512,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))) | p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK15))))) ) | $spl0),
  inference(subsumption_resolution,[],[f7456,f269])).
fof(f7456,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))) | p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK15))))) ) | $spl0),
  inference(superposition,[],[f316,f5212])).
fof(f5212,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15))) | $spl0),
  inference(unit_resulting_resolution,[],[f269,f477,f341])).
fof(f22600,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | $spl3),
  inference(unit_resulting_resolution,[],[f480,f318])).
fof(f22593,plain,(
  $false | ($spl0 | $spl5)),
  inference(subsumption_resolution,[],[f22592,f269])).
fof(f22592,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl5)),
  inference(forward_demodulation,[],[f21901,f17578])).
fof(f17578,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK14,sK12)),s(fun(sK12,bool),sK14))) | $spl5),
  inference(unit_resulting_resolution,[],[f482,f1366])).
fof(f21901,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK14,sK12)),s(fun(sK12,bool),sK14)))) | ($spl0 | $spl5)),
  inference(unit_resulting_resolution,[],[f18333,f337])).
fof(f337,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))) )),
  inference(cnf_transformation,[],[f206])).
fof(f18333,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK14,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))) | ($spl0 | $spl5)),
  inference(unit_resulting_resolution,[],[f17579,f7512])).
fof(f17579,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK14,sK12)),s(fun(sK12,bool),sK15)))) | $spl5),
  inference(unit_resulting_resolution,[],[f482,f318])).
fof(f17301,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f12425,f12450])).
fof(f12450,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1)) ) | $spl46),
  inference(unit_resulting_resolution,[],[f3746,f3746,f341])).
fof(f12425,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | $spl46),
  inference(unit_resulting_resolution,[],[f3746,f289])).
fof(f289,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f168])).
fof(f168,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f63])).
fof(f63,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f62])).
fof(f62,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X8] : (s(bool,'F_0') = s(bool,X8) <=> ~p(s(bool,X8)))),
  file('thm.pred_set.UNION_SUBSET','thm.bool.EQ_CLAUSES|split|2')).
fof(f17300,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f12433,f269])).
fof(f12433,plain,(
  ~p(s(bool,'T_0')) | $spl46),
  inference(unit_resulting_resolution,[],[f527,f3746,f341])).
fof(f17299,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f12435,f269])).
fof(f12435,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f4936,f3746,f341])).
fof(f4936,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | $spl3),
  inference(unit_resulting_resolution,[],[f480,f285])).
fof(f17298,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f12437,f269])).
fof(f12437,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5686,f3746,f341])).
fof(f5686,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | $spl3),
  inference(unit_resulting_resolution,[],[f4934,f285])).
fof(f4934,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | $spl3),
  inference(unit_resulting_resolution,[],[f480,f318])).
fof(f17297,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f12439,f269])).
fof(f12439,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6640,f3746,f341])).
fof(f6640,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14))) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f5683,f285])).
fof(f5683,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f4934,f1549])).
fof(f1549,plain,(
  ( ! [X1] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))) | p(s(bool,'const.bool.IN_2'(s(sK12,X1),s(fun(sK12,bool),sK15))))) ) | $spl4),
  inference(subsumption_resolution,[],[f1546,f269])).
fof(f1546,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))) | p(s(bool,'const.bool.IN_2'(s(sK12,X1),s(fun(sK12,bool),sK15))))) ) | $spl4),
  inference(superposition,[],[f316,f531])).
fof(f531,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))) | $spl4),
  inference(unit_resulting_resolution,[],[f481,f286])).
fof(f286,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f166])).
fof(f17296,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f12443,f3746])).
fof(f12443,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f4936,f3746,f341])).
fof(f17295,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f12445,f3746])).
fof(f12445,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5686,f3746,f341])).
fof(f17294,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f12447,f3746])).
fof(f12447,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6640,f3746,f341])).
fof(f17293,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f12451,f269])).
fof(f12451,plain,(
  ~p(s(bool,'T_0')) | $spl46),
  inference(unit_resulting_resolution,[],[f527,f3746,f341])).
fof(f17292,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f12453,f269])).
fof(f12453,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f4936,f3746,f341])).
fof(f17291,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f12455,f269])).
fof(f12455,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5686,f3746,f341])).
fof(f17290,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f12457,f269])).
fof(f12457,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6640,f3746,f341])).
fof(f17289,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f12461,f3746])).
fof(f12461,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f4936,f3746,f341])).
fof(f17288,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f12463,f3746])).
fof(f12463,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5686,f3746,f341])).
fof(f17287,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f12465,f3746])).
fof(f12465,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6640,f3746,f341])).
fof(f17286,plain,(
  $false | ($spl3 | $spl6 | $spl46)),
  inference(subsumption_resolution,[],[f12472,f3746])).
fof(f12472,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f5141,f3746,f362])).
fof(f362,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f217])).
fof(f217,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1)))),
  inference(rectify,[],[f216])).
fof(f216,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0)))),
  inference(flattening,[],[f215])).
fof(f215,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP1(X1,X0)))),
  inference(nnf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X1,X0] : (sP1(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f5141,plain,(
  sP1('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | ($spl3 | $spl6)),
  inference(unit_resulting_resolution,[],[f480,f2369])).
fof(f17285,plain,(
  $false | ($spl3 | $spl6 | $spl46)),
  inference(subsumption_resolution,[],[f12473,f3746])).
fof(f12473,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f5939,f3746,f362])).
fof(f5939,plain,(
  sP1('T_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | ($spl3 | $spl6)),
  inference(unit_resulting_resolution,[],[f4934,f2369])).
fof(f17284,plain,(
  $false | ($spl3 | $spl4 | $spl6 | $spl46)),
  inference(subsumption_resolution,[],[f12474,f3746])).
fof(f12474,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f6975,f3746,f362])).
fof(f6975,plain,(
  sP1('T_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14))) | ($spl3 | $spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f5683,f2369])).
fof(f17283,plain,(
  $false | ($spl3 | $spl6 | $spl46)),
  inference(subsumption_resolution,[],[f12475,f3746])).
fof(f12475,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | ($spl3 | $spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f9482,f3746,f362])).
fof(f9482,plain,(
  sP1('T_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15))))) | ($spl3 | $spl6)),
  inference(unit_resulting_resolution,[],[f5685,f2369])).
fof(f5685,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | $spl3),
  inference(unit_resulting_resolution,[],[f4934,f4934,f335])).
fof(f335,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f206])).
fof(f17282,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f12476,f269])).
fof(f12476,plain,(
  ~p(s(bool,'T_0')) | $spl46),
  inference(unit_resulting_resolution,[],[f575,f3746,f362])).
fof(f575,plain,(
  sP1('F_0','T_0')),
  inference(unit_resulting_resolution,[],[f483,f269,f366])).
fof(f366,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f217])).
fof(f17281,plain,(
  $false | ($spl46 | $spl63)),
  inference(subsumption_resolution,[],[f12477,f3746])).
fof(f12477,plain,(
  ~p(s(bool,f(s(bool,'T_0')))) | ($spl46 | $spl63)),
  inference(unit_resulting_resolution,[],[f4809,f3746,f362])).
fof(f4809,plain,(
  sP1('F_0',f(s(bool,'T_0'))) | $spl63),
  inference(unit_resulting_resolution,[],[f483,f4532])).
fof(f4532,plain,(
  ( ! [X6] : (sP1(X6,f(s(bool,'T_0'))) | p(s(bool,X6))) ) | $spl63),
  inference(subsumption_resolution,[],[f4479,f269])).
fof(f4479,plain,(
  ( ! [X6] : (~p(s(bool,'T_0')) | p(s(bool,X6)) | sP1(X6,f(s(bool,'T_0')))) ) | $spl63),
  inference(superposition,[],[f366,f4246])).
fof(f4246,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | $spl63),
  inference(unit_resulting_resolution,[],[f483,f4228,f832])).
fof(f17280,plain,(
  $false | ($spl4 | $spl46)),
  inference(subsumption_resolution,[],[f12478,f3746])).
fof(f12478,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))) | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f645,f3746,f362])).
fof(f645,plain,(
  sP1('F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))) | $spl4),
  inference(unit_resulting_resolution,[],[f483,f579])).
fof(f579,plain,(
  ( ! [X1] : (sP1(X1,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))) | p(s(bool,X1))) ) | $spl4),
  inference(subsumption_resolution,[],[f577,f269])).
fof(f577,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | p(s(bool,X1)) | sP1(X1,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))) ) | $spl4),
  inference(superposition,[],[f366,f531])).
fof(f17279,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f12479,f3746])).
fof(f12479,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)))) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5221,f3746,f362])).
fof(f5221,plain,(
  sP1('F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15))) | $spl0),
  inference(unit_resulting_resolution,[],[f483,f477,f366])).
fof(f17278,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f12480,f3746])).
fof(f12480,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f8666,f3746,f362])).
fof(f8666,plain,(
  sP1('F_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK13))) | $spl3),
  inference(unit_resulting_resolution,[],[f483,f8008])).
fof(f8008,plain,(
  ( ! [X8] : (sP1(X8,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK13))) | p(s(bool,X8))) ) | $spl3),
  inference(subsumption_resolution,[],[f7954,f269])).
fof(f7954,plain,(
  ( ! [X8] : (~p(s(bool,'T_0')) | p(s(bool,X8)) | sP1(X8,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | $spl3),
  inference(superposition,[],[f366,f4933])).
fof(f4933,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK13))) | $spl3),
  inference(unit_resulting_resolution,[],[f480,f1366])).
fof(f17277,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f12481,f3746])).
fof(f12481,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),X0),s(fun(sK12,bool),sK13))))))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f10665,f3746,f362])).
fof(f10665,plain,(
  ( ! [X0] : (sP1('F_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),X0),s(fun(sK12,bool),sK13)))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f483,f7998,f366])).
fof(f7998,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),X0),s(fun(sK12,bool),sK13))))))) ) | $spl3),
  inference(subsumption_resolution,[],[f7941,f269])).
fof(f7941,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),X0),s(fun(sK12,bool),sK13))))))) ) | $spl3),
  inference(superposition,[],[f337,f4933])).
fof(f17276,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f12482,f3746])).
fof(f12482,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),X0))))))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f11286,f3746,f362])).
fof(f11286,plain,(
  ( ! [X0] : (sP1('F_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),X0)))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f483,f7999,f366])).
fof(f7999,plain,(
  ( ! [X1] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),X1))))))) ) | $spl3),
  inference(subsumption_resolution,[],[f7942,f269])).
fof(f7942,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),X1))))))) ) | $spl3),
  inference(superposition,[],[f336,f4933])).
fof(f17275,plain,(
  $false | ($spl3 | $spl46 | $spl63)),
  inference(subsumption_resolution,[],[f12484,f3746])).
fof(f12484,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46 | $spl63)),
  inference(unit_resulting_resolution,[],[f5172,f3746,f362])).
fof(f5172,plain,(
  sP1(f(s(bool,'T_0')),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | ($spl3 | $spl63)),
  inference(unit_resulting_resolution,[],[f480,f4531])).
fof(f4531,plain,(
  ( ! [X5] : (sP1(f(s(bool,'T_0')),X5) | p(s(bool,X5))) ) | $spl63),
  inference(subsumption_resolution,[],[f4478,f269])).
fof(f4478,plain,(
  ( ! [X5] : (~p(s(bool,'T_0')) | p(s(bool,X5)) | sP1(f(s(bool,'T_0')),X5)) ) | $spl63),
  inference(superposition,[],[f363,f4246])).
fof(f363,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f217])).
fof(f17274,plain,(
  $false | ($spl3 | $spl46 | $spl63)),
  inference(subsumption_resolution,[],[f12485,f3746])).
fof(f12485,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46 | $spl63)),
  inference(unit_resulting_resolution,[],[f5974,f3746,f362])).
fof(f5974,plain,(
  sP1(f(s(bool,'T_0')),'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | ($spl3 | $spl63)),
  inference(unit_resulting_resolution,[],[f4934,f4531])).
fof(f17273,plain,(
  $false | ($spl3 | $spl4 | $spl46 | $spl63)),
  inference(subsumption_resolution,[],[f12486,f3746])).
fof(f12486,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl46 | $spl63)),
  inference(unit_resulting_resolution,[],[f7015,f3746,f362])).
fof(f7015,plain,(
  sP1(f(s(bool,'T_0')),'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14))) | ($spl3 | $spl4 | $spl63)),
  inference(unit_resulting_resolution,[],[f5683,f4531])).
fof(f17272,plain,(
  $false | ($spl3 | $spl10 | $spl46)),
  inference(subsumption_resolution,[],[f12487,f269])).
fof(f12487,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl10 | $spl46)),
  inference(unit_resulting_resolution,[],[f5142,f3746,f362])).
fof(f5142,plain,(
  sP1('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'T_0') | ($spl3 | $spl10)),
  inference(unit_resulting_resolution,[],[f480,f2373])).
fof(f17271,plain,(
  $false | ($spl3 | $spl46 | $spl63)),
  inference(subsumption_resolution,[],[f12490,f3746])).
fof(f12490,plain,(
  ~p(s(bool,f(s(bool,'T_0')))) | ($spl3 | $spl46 | $spl63)),
  inference(unit_resulting_resolution,[],[f5173,f3746,f362])).
fof(f5173,plain,(
  sP1('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),f(s(bool,'T_0'))) | ($spl3 | $spl63)),
  inference(unit_resulting_resolution,[],[f480,f4532])).
fof(f17270,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f12491,f3746])).
fof(f12491,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5109,f3746,f362])).
fof(f5109,plain,(
  sP1('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f480,f569])).
fof(f569,plain,(
  ( ! [X1] : (sP1('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),X1) | p(s(bool,X1))) ) | $spl4),
  inference(subsumption_resolution,[],[f567,f269])).
fof(f567,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | p(s(bool,X1)) | sP1('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),X1)) ) | $spl4),
  inference(superposition,[],[f363,f531])).
fof(f17269,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f12492,f3746])).
fof(f12492,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5110,f3746,f362])).
fof(f5110,plain,(
  sP1('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f480,f579])).
fof(f17268,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f12493,f3746])).
fof(f12493,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5903,f3746,f362])).
fof(f5903,plain,(
  sP1('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f4934,f569])).
fof(f17267,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f12494,f3746])).
fof(f12494,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6934,f3746,f362])).
fof(f6934,plain,(
  sP1('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14))) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f5683,f569])).
fof(f17266,plain,(
  $false | ($spl3 | $spl10 | $spl46)),
  inference(subsumption_resolution,[],[f12495,f269])).
fof(f12495,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl10 | $spl46)),
  inference(unit_resulting_resolution,[],[f5940,f3746,f362])).
fof(f5940,plain,(
  sP1('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),'T_0') | ($spl3 | $spl10)),
  inference(unit_resulting_resolution,[],[f4934,f2373])).
fof(f17265,plain,(
  $false | ($spl3 | $spl4 | $spl10 | $spl46)),
  inference(subsumption_resolution,[],[f12496,f269])).
fof(f12496,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl4 | $spl10 | $spl46)),
  inference(unit_resulting_resolution,[],[f6976,f3746,f362])).
fof(f6976,plain,(
  sP1('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),'T_0') | ($spl3 | $spl4 | $spl10)),
  inference(unit_resulting_resolution,[],[f5683,f2373])).
fof(f17264,plain,(
  $false | ($spl3 | $spl10 | $spl46)),
  inference(subsumption_resolution,[],[f12497,f269])).
fof(f12497,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl10 | $spl46)),
  inference(unit_resulting_resolution,[],[f9483,f3746,f362])).
fof(f9483,plain,(
  sP1('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))),'T_0') | ($spl3 | $spl10)),
  inference(unit_resulting_resolution,[],[f5685,f2373])).
fof(f17263,plain,(
  $false | ($spl3 | $spl46 | $spl63)),
  inference(subsumption_resolution,[],[f12501,f3746])).
fof(f12501,plain,(
  ~p(s(bool,f(s(bool,'T_0')))) | ($spl3 | $spl46 | $spl63)),
  inference(unit_resulting_resolution,[],[f5975,f3746,f362])).
fof(f5975,plain,(
  sP1('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),f(s(bool,'T_0'))) | ($spl3 | $spl63)),
  inference(unit_resulting_resolution,[],[f4934,f4532])).
fof(f17262,plain,(
  $false | ($spl3 | $spl4 | $spl46 | $spl63)),
  inference(subsumption_resolution,[],[f12502,f3746])).
fof(f12502,plain,(
  ~p(s(bool,f(s(bool,'T_0')))) | ($spl3 | $spl4 | $spl46 | $spl63)),
  inference(unit_resulting_resolution,[],[f7016,f3746,f362])).
fof(f7016,plain,(
  sP1('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),f(s(bool,'T_0'))) | ($spl3 | $spl4 | $spl63)),
  inference(unit_resulting_resolution,[],[f5683,f4532])).
fof(f17261,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f12503,f3746])).
fof(f12503,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5904,f3746,f362])).
fof(f5904,plain,(
  sP1('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f4934,f579])).
fof(f17260,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f12504,f3746])).
fof(f12504,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6935,f3746,f362])).
fof(f6935,plain,(
  sP1('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f5683,f579])).
fof(f17259,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f12509,f269])).
fof(f12509,plain,(
  ~p(s(bool,'T_0')) | $spl46),
  inference(unit_resulting_resolution,[],[f565,f3746,f362])).
fof(f565,plain,(
  sP1('T_0','F_0')),
  inference(unit_resulting_resolution,[],[f483,f269,f363])).
fof(f17258,plain,(
  $false | ($spl3 | $spl6 | $spl46)),
  inference(subsumption_resolution,[],[f12510,f269])).
fof(f12510,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f5141,f3746,f362])).
fof(f17257,plain,(
  $false | ($spl3 | $spl6 | $spl46)),
  inference(subsumption_resolution,[],[f12511,f269])).
fof(f12511,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f5939,f3746,f362])).
fof(f17256,plain,(
  $false | ($spl3 | $spl4 | $spl6 | $spl46)),
  inference(subsumption_resolution,[],[f12512,f269])).
fof(f12512,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl4 | $spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f6975,f3746,f362])).
fof(f17255,plain,(
  $false | ($spl3 | $spl6 | $spl46)),
  inference(subsumption_resolution,[],[f12513,f269])).
fof(f12513,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f9482,f3746,f362])).
fof(f17254,plain,(
  $false | ($spl46 | $spl63)),
  inference(subsumption_resolution,[],[f12521,f3746])).
fof(f12521,plain,(
  ~p(s(bool,f(s(bool,'T_0')))) | ($spl46 | $spl63)),
  inference(unit_resulting_resolution,[],[f4778,f3746,f362])).
fof(f4778,plain,(
  sP1(f(s(bool,'T_0')),'F_0') | $spl63),
  inference(unit_resulting_resolution,[],[f483,f4531])).
fof(f17253,plain,(
  $false | ($spl3 | $spl46 | $spl63)),
  inference(subsumption_resolution,[],[f12522,f3746])).
fof(f12522,plain,(
  ~p(s(bool,f(s(bool,'T_0')))) | ($spl3 | $spl46 | $spl63)),
  inference(unit_resulting_resolution,[],[f5172,f3746,f362])).
fof(f17252,plain,(
  $false | ($spl3 | $spl46 | $spl63)),
  inference(subsumption_resolution,[],[f12523,f3746])).
fof(f12523,plain,(
  ~p(s(bool,f(s(bool,'T_0')))) | ($spl3 | $spl46 | $spl63)),
  inference(unit_resulting_resolution,[],[f5974,f3746,f362])).
fof(f17251,plain,(
  $false | ($spl3 | $spl4 | $spl46 | $spl63)),
  inference(subsumption_resolution,[],[f12524,f3746])).
fof(f12524,plain,(
  ~p(s(bool,f(s(bool,'T_0')))) | ($spl3 | $spl4 | $spl46 | $spl63)),
  inference(unit_resulting_resolution,[],[f7015,f3746,f362])).
fof(f17250,plain,(
  $false | ($spl3 | $spl10 | $spl46)),
  inference(subsumption_resolution,[],[f12525,f3746])).
fof(f12525,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl10 | $spl46)),
  inference(unit_resulting_resolution,[],[f5142,f3746,f362])).
fof(f17249,plain,(
  $false | ($spl4 | $spl46)),
  inference(subsumption_resolution,[],[f12526,f3746])).
fof(f12526,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))) | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f631,f3746,f362])).
fof(f631,plain,(
  sP1('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f483,f569])).
fof(f17248,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f12527,f3746])).
fof(f12527,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)))) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5220,f3746,f362])).
fof(f5220,plain,(
  sP1('const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f483,f477,f363])).
fof(f17247,plain,(
  $false | ($spl3 | $spl46 | $spl63)),
  inference(subsumption_resolution,[],[f12528,f3746])).
fof(f12528,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46 | $spl63)),
  inference(unit_resulting_resolution,[],[f5173,f3746,f362])).
fof(f17246,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f12529,f3746])).
fof(f12529,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5109,f3746,f362])).
fof(f17245,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f12530,f3746])).
fof(f12530,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5110,f3746,f362])).
fof(f17244,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f12531,f3746])).
fof(f12531,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5903,f3746,f362])).
fof(f17243,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f12532,f3746])).
fof(f12532,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6934,f3746,f362])).
fof(f17242,plain,(
  $false | ($spl3 | $spl10 | $spl46)),
  inference(subsumption_resolution,[],[f12533,f3746])).
fof(f12533,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl10 | $spl46)),
  inference(unit_resulting_resolution,[],[f5940,f3746,f362])).
fof(f17241,plain,(
  $false | ($spl3 | $spl4 | $spl10 | $spl46)),
  inference(subsumption_resolution,[],[f12534,f3746])).
fof(f12534,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl10 | $spl46)),
  inference(unit_resulting_resolution,[],[f6976,f3746,f362])).
fof(f17240,plain,(
  $false | ($spl3 | $spl10 | $spl46)),
  inference(subsumption_resolution,[],[f12535,f3746])).
fof(f12535,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | ($spl3 | $spl10 | $spl46)),
  inference(unit_resulting_resolution,[],[f9483,f3746,f362])).
fof(f17239,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f12536,f3746])).
fof(f12536,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f8658,f3746,f362])).
fof(f8658,plain,(
  sP1('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),'F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f483,f8007])).
fof(f8007,plain,(
  ( ! [X7] : (sP1('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X7) | p(s(bool,X7))) ) | $spl3),
  inference(subsumption_resolution,[],[f7953,f269])).
fof(f7953,plain,(
  ( ! [X7] : (~p(s(bool,'T_0')) | p(s(bool,X7)) | sP1('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X7)) ) | $spl3),
  inference(superposition,[],[f363,f4933])).
fof(f17238,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f12537,f3746])).
fof(f12537,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),X0),s(fun(sK12,bool),sK13))))))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f10664,f3746,f362])).
fof(f10664,plain,(
  ( ! [X0] : (sP1('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),X0),s(fun(sK12,bool),sK13)))),'F_0')) ) | $spl3),
  inference(unit_resulting_resolution,[],[f483,f7998,f363])).
fof(f17237,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f12538,f3746])).
fof(f12538,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),X0))))))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f11284,f3746,f362])).
fof(f11284,plain,(
  ( ! [X0] : (sP1('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),X0)))),'F_0')) ) | $spl3),
  inference(unit_resulting_resolution,[],[f483,f7999,f363])).
fof(f17236,plain,(
  $false | ($spl3 | $spl46 | $spl63)),
  inference(subsumption_resolution,[],[f12539,f3746])).
fof(f12539,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46 | $spl63)),
  inference(unit_resulting_resolution,[],[f5975,f3746,f362])).
fof(f17235,plain,(
  $false | ($spl3 | $spl4 | $spl46 | $spl63)),
  inference(subsumption_resolution,[],[f12540,f3746])).
fof(f12540,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl46 | $spl63)),
  inference(unit_resulting_resolution,[],[f7016,f3746,f362])).
fof(f17234,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f12541,f3746])).
fof(f12541,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5904,f3746,f362])).
fof(f17233,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f12542,f3746])).
fof(f12542,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6935,f3746,f362])).
fof(f17232,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f12590,f3746])).
fof(f12590,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f4951,f483,f3746,f397])).
fof(f397,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f233])).
fof(f233,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | (sP2(X2,X1,X0) & sP3(X1,X2,X0) & sP4(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP2(X2,X1,X0) | ~sP3(X1,X2,X0) | ~sP4(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f232])).
fof(f232,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP5(X2,X1,X0)))),
  inference(flattening,[],[f231])).
fof(f231,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP5(X2,X1,X0)))),
  inference(nnf_transformation,[],[f150])).
fof(f150,plain,(
  ! [X2,X1,X0] : (sP5(X2,X1,X0) <=> (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f4951,plain,(
  sP5('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f501,f496,f494,f480,f399])).
fof(f399,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X2,X1) | p(s(bool,X1)) | sP5(X0,X1,X2) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f233])).
fof(f494,plain,(
  ( ! [X0,X1] : (sP4(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f483,f403])).
fof(f403,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f236])).
fof(f236,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f235])).
fof(f235,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(flattening,[],[f234])).
fof(f234,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(nnf_transformation,[],[f149])).
fof(f149,plain,(
  ! [X2,X0,X1] : (sP4(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f496,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f269,f406])).
fof(f406,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f239])).
fof(f239,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f238])).
fof(f238,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(flattening,[],[f237])).
fof(f237,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(nnf_transformation,[],[f148])).
fof(f148,plain,(
  ! [X1,X0,X2] : (sP3(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f501,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f269,f410])).
fof(f410,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f241])).
fof(f241,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(flattening,[],[f240])).
fof(f240,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(nnf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0,X1,X2] : (sP2(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f17231,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f12591,f3746])).
fof(f12591,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5701,f483,f3746,f397])).
fof(f5701,plain,(
  sP5('T_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),'F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f501,f496,f494,f4934,f399])).
fof(f17230,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f12592,f3746])).
fof(f12592,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6656,f483,f3746,f397])).
fof(f6656,plain,(
  sP5('T_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),'F_0') | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f501,f496,f494,f5683,f399])).
fof(f17229,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f12593,f3746])).
fof(f12593,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9070,f483,f3746,f397])).
fof(f9070,plain,(
  sP5('T_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))),'F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f501,f496,f494,f5685,f399])).
fof(f17228,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f12594,f269])).
fof(f12594,plain,(
  ~p(s(bool,'T_0')) | $spl46),
  inference(unit_resulting_resolution,[],[f859,f483,f3746,f397])).
fof(f859,plain,(
  sP5('F_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f483,f527,f413])).
fof(f413,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f242])).
fof(f242,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP5(X2,X1,X0)) & (~sP5(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f151])).
fof(f151,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP5(X2,X1,X0))),
  inference(definition_folding,[],[f111,f150,f149,f148,f147])).
fof(f111,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f110])).
fof(f110,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X20,X21,X22] : ((p(s(bool,X22)) <=> s(bool,X20) = s(bool,X21)) <=> ((p(s(bool,X22)) | p(s(bool,X21)) | p(s(bool,X20))) & (p(s(bool,X22)) | ~p(s(bool,X20)) | ~p(s(bool,X21))) & (p(s(bool,X21)) | ~p(s(bool,X20)) | ~p(s(bool,X22))) & (p(s(bool,X20)) | ~p(s(bool,X21)) | ~p(s(bool,X22)))))),
  file('thm.pred_set.UNION_SUBSET','thm.sat.sat_2|49|_prove_ltb')).
fof(f17227,plain,(
  $false | ($spl46 | $spl63)),
  inference(subsumption_resolution,[],[f12595,f3746])).
fof(f12595,plain,(
  ~p(s(bool,f(s(bool,'T_0')))) | ($spl46 | $spl63)),
  inference(unit_resulting_resolution,[],[f4237,f483,f3746,f397])).
fof(f4237,plain,(
  sP5('F_0',f(s(bool,'T_0')),'F_0') | $spl63),
  inference(unit_resulting_resolution,[],[f483,f4228,f413])).
fof(f17226,plain,(
  $false | ($spl4 | $spl46)),
  inference(subsumption_resolution,[],[f12596,f3746])).
fof(f12596,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))) | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f1330,f483,f3746,f397])).
fof(f1330,plain,(
  sP5('F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f506,f499,f494,f483,f400])).
fof(f400,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X2,X1) | p(s(bool,X0)) | sP5(X0,X1,X2) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f233])).
fof(f499,plain,(
  ( ! [X0,X1] : (sP3(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f483,f407])).
fof(f407,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f239])).
fof(f506,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f481,f411])).
fof(f411,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f241])).
fof(f17225,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f12597,f3746])).
fof(f12597,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)))) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f6523,f483,f3746,f397])).
fof(f6523,plain,(
  sP5('F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f483,f495,f499,f5253,f400])).
fof(f5253,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),X1)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f477,f411])).
fof(f495,plain,(
  ( ! [X0,X1] : (sP4('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f483,f404])).
fof(f404,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f236])).
fof(f17224,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f12598,f269])).
fof(f12598,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f4992,f483,f3746,f397])).
fof(f4992,plain,(
  sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'T_0','F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f527,f480,f413])).
fof(f17223,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f12601,f3746])).
fof(f12601,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f4952,f483,f3746,f397])).
fof(f4952,plain,(
  sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'F_0') | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f503,f498,f494,f480,f399])).
fof(f498,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f481,f406])).
fof(f503,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f481,f410])).
fof(f17222,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f12602,f3746])).
fof(f12602,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f4956,f483,f3746,f397])).
fof(f4956,plain,(
  sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),'F_0') | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f506,f499,f494,f480,f400])).
fof(f17221,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f12603,f3746])).
fof(f12603,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5702,f483,f3746,f397])).
fof(f5702,plain,(
  sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),'F_0') | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f503,f498,f494,f4934,f399])).
fof(f17220,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f12604,f269])).
fof(f12604,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5754,f483,f3746,f397])).
fof(f5754,plain,(
  sP5('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),'T_0','F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f527,f4934,f413])).
fof(f17219,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f12605,f3746])).
fof(f12605,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5712,f483,f3746,f397])).
fof(f5712,plain,(
  sP5('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),'F_0') | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f506,f499,f494,f4934,f400])).
fof(f17218,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f12606,f269])).
fof(f12606,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6739,f483,f3746,f397])).
fof(f6739,plain,(
  sP5('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),'T_0','F_0') | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f527,f5683,f413])).
fof(f17217,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f12607,f269])).
fof(f12607,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9171,f483,f3746,f397])).
fof(f9171,plain,(
  sP5('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))),'T_0','F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f527,f5685,f413])).
fof(f17216,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f12704,f269])).
fof(f12704,plain,(
  ~p(s(bool,'T_0')) | $spl46),
  inference(unit_resulting_resolution,[],[f1282,f483,f3746,f397])).
fof(f1282,plain,(
  sP5('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f501,f496,f494,f483,f399])).
fof(f17215,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f12705,f269])).
fof(f12705,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f4951,f483,f3746,f397])).
fof(f17214,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f12706,f269])).
fof(f12706,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5701,f483,f3746,f397])).
fof(f17213,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f12707,f269])).
fof(f12707,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6656,f483,f3746,f397])).
fof(f17212,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f12708,f269])).
fof(f12708,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9070,f483,f3746,f397])).
fof(f17211,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f12713,f3746])).
fof(f12713,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f4992,f483,f3746,f397])).
fof(f17210,plain,(
  $false | ($spl4 | $spl46)),
  inference(subsumption_resolution,[],[f12714,f3746])).
fof(f12714,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))) | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f1284,f483,f3746,f397])).
fof(f1284,plain,(
  sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),'F_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f503,f498,f494,f483,f399])).
fof(f17209,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f12715,f3746])).
fof(f12715,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)))) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f6506,f483,f3746,f397])).
fof(f6506,plain,(
  sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),'F_0','F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f483,f494,f500,f5252,f399])).
fof(f5252,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f477,f410])).
fof(f500,plain,(
  ( ! [X0,X1] : (sP3('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f483,f408])).
fof(f408,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f239])).
fof(f17208,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f12716,f3746])).
fof(f12716,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f4952,f483,f3746,f397])).
fof(f17207,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f12717,f3746])).
fof(f12717,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f4956,f483,f3746,f397])).
fof(f17206,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f12718,f3746])).
fof(f12718,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5702,f483,f3746,f397])).
fof(f17205,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f12719,f3746])).
fof(f12719,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5754,f483,f3746,f397])).
fof(f17204,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f12720,f3746])).
fof(f12720,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5712,f483,f3746,f397])).
fof(f17203,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f12721,f3746])).
fof(f12721,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6739,f483,f3746,f397])).
fof(f17202,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f12722,f3746])).
fof(f12722,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9171,f483,f3746,f397])).
fof(f17201,plain,(
  $false | ($spl12 | $spl46)),
  inference(subsumption_resolution,[],[f12835,f15936])).
fof(f15936,plain,(
  ( ! [X0] : (~sP5(X0,'T_0','F_0')) ) | ($spl12 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f3746,f2376])).
fof(f12835,plain,(
  ( ! [X0] : (sP5(X0,'T_0','F_0')) ) | $spl46),
  inference(unit_resulting_resolution,[],[f504,f499,f491,f3746,f398])).
fof(f398,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP5(X0,X1,X2) | ~sP4(X0,X2,X1) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f233])).
fof(f491,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f269,f402])).
fof(f402,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f236])).
fof(f504,plain,(
  ( ! [X0,X1] : (sP2(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f269,f411])).
fof(f17200,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f12839,f13658])).
fof(f13658,plain,(
  ( ! [X0] : (~sP5(X0,'T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f4936,f3746,f415])).
fof(f415,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f242])).
fof(f12839,plain,(
  ( ! [X0] : (sP5(X0,'T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f504,f4965,f491,f3746,f398])).
fof(f4965,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),X1)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f480,f407])).
fof(f17199,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f12843,f13660])).
fof(f13660,plain,(
  ( ! [X0] : (~sP5(X0,'T_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5686,f3746,f415])).
fof(f12843,plain,(
  ( ! [X0] : (sP5(X0,'T_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f504,f5727,f491,f3746,f398])).
fof(f5727,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),X1)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f4934,f407])).
fof(f17198,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f12847,f13662])).
fof(f13662,plain,(
  ( ! [X0] : (~sP5(X0,'T_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6640,f3746,f415])).
fof(f12847,plain,(
  ( ! [X0] : (sP5(X0,'T_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f504,f6705,f491,f3746,f398])).
fof(f6705,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),X1)) ) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f5683,f407])).
fof(f17197,plain,(
  $false | ($spl4 | $spl46)),
  inference(subsumption_resolution,[],[f12870,f15820])).
fof(f15820,plain,(
  ( ! [X0] : (~sP5(X0,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),'F_0')) ) | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f3746,f915])).
fof(f915,plain,(
  ( ! [X2,X3] : (~sP5(X2,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),X3) | p(s(bool,X3)) | ~p(s(bool,X2))) ) | $spl4),
  inference(subsumption_resolution,[],[f909,f269])).
fof(f909,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | ~p(s(bool,X2)) | p(s(bool,X3)) | ~sP5(X2,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),X3)) ) | $spl4),
  inference(superposition,[],[f397,f531])).
fof(f12870,plain,(
  ( ! [X0] : (sP5(X0,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),'F_0')) ) | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f506,f499,f493,f3746,f398])).
fof(f493,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f481,f402])).
fof(f17196,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f12905,f16017])).
fof(f16017,plain,(
  ( ! [X0] : (~sP5(X0,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),'F_0')) ) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f3746,f7520])).
fof(f7520,plain,(
  ( ! [X8,X9] : (~sP5(X8,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),X9) | p(s(bool,X9)) | ~p(s(bool,X8))) ) | $spl0),
  inference(subsumption_resolution,[],[f7469,f269])).
fof(f7469,plain,(
  ( ! [X8,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,X8)) | p(s(bool,X9)) | ~sP5(X8,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),X9)) ) | $spl0),
  inference(superposition,[],[f397,f5212])).
fof(f12905,plain,(
  ( ! [X0] : (sP5(X0,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),'F_0')) ) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5253,f499,f5240,f3746,f398])).
fof(f5240,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f477,f402])).
fof(f17195,plain,(
  $false | ($spl12 | $spl46)),
  inference(subsumption_resolution,[],[f12929,f15936])).
fof(f12929,plain,(
  ( ! [X0] : (sP5(X0,'T_0','F_0')) ) | $spl46),
  inference(unit_resulting_resolution,[],[f504,f499,f494,f3746,f398])).
fof(f17194,plain,(
  $false | ($spl4 | $spl46)),
  inference(subsumption_resolution,[],[f12930,f15820])).
fof(f12930,plain,(
  ( ! [X0] : (sP5(X0,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),'F_0')) ) | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f506,f499,f494,f3746,f398])).
fof(f17193,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f12931,f16017])).
fof(f12931,plain,(
  ( ! [X0] : (sP5(X0,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),'F_0')) ) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5253,f499,f494,f3746,f398])).
fof(f17192,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f12933,f881])).
fof(f881,plain,(
  ~sP5('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f269,f527,f415])).
fof(f12933,plain,(
  sP5('T_0','T_0','F_0') | $spl46),
  inference(unit_resulting_resolution,[],[f504,f496,f494,f3746,f398])).
fof(f17191,plain,(
  $false | ($spl4 | $spl46)),
  inference(subsumption_resolution,[],[f12934,f1062])).
fof(f1062,plain,(
  ~sP5('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f269,f483,f915])).
fof(f12934,plain,(
  sP5('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),'F_0') | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f506,f496,f494,f3746,f398])).
fof(f17190,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f12935,f5224])).
fof(f5224,plain,(
  ~sP5('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f269,f483,f477,f397])).
fof(f12935,plain,(
  sP5('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),'F_0') | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5253,f496,f494,f3746,f398])).
fof(f17189,plain,(
  $false | ($spl4 | $spl12 | $spl46)),
  inference(subsumption_resolution,[],[f12937,f15936])).
fof(f12937,plain,(
  sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),'T_0','F_0') | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f504,f498,f494,f3746,f398])).
fof(f17188,plain,(
  $false | ($spl4 | $spl46)),
  inference(subsumption_resolution,[],[f12938,f15820])).
fof(f12938,plain,(
  sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),'F_0') | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f506,f498,f494,f3746,f398])).
fof(f17187,plain,(
  $false | ($spl0 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f12939,f16017])).
fof(f12939,plain,(
  sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),'F_0') | ($spl0 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5253,f498,f494,f3746,f398])).
fof(f17186,plain,(
  $false | ($spl0 | $spl12 | $spl46)),
  inference(subsumption_resolution,[],[f12941,f5405])).
fof(f5405,plain,(
  ~sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),'T_0','F_0') | ($spl0 | $spl12)),
  inference(unit_resulting_resolution,[],[f483,f477,f2376])).
fof(f12941,plain,(
  sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),'T_0','F_0') | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f504,f5251,f494,f3746,f398])).
fof(f5251,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f477,f406])).
fof(f17185,plain,(
  $false | ($spl0 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f12942,f5388])).
fof(f5388,plain,(
  ~sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),'F_0') | ($spl0 | $spl4)),
  inference(unit_resulting_resolution,[],[f483,f477,f915])).
fof(f12942,plain,(
  sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),'F_0') | ($spl0 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f506,f5251,f494,f3746,f398])).
fof(f17184,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f12943,f5225])).
fof(f5225,plain,(
  ~sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f477,f483,f477,f397])).
fof(f12943,plain,(
  sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),'F_0') | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5253,f5251,f494,f3746,f398])).
fof(f17183,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f12962,f13658])).
fof(f12962,plain,(
  ( ! [X0] : (sP5(X0,'T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f504,f4965,f4960,f3746,f398])).
fof(f4960,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),X1)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f480,f403])).
fof(f17182,plain,(
  $false | ($spl3 | $spl12 | $spl46)),
  inference(subsumption_resolution,[],[f12966,f5143])).
fof(f5143,plain,(
  ~sP5('T_0','T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | ($spl3 | $spl12)),
  inference(unit_resulting_resolution,[],[f269,f480,f2376])).
fof(f12966,plain,(
  sP5('T_0','T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f504,f496,f4960,f3746,f398])).
fof(f17181,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f12967,f5113])).
fof(f5113,plain,(
  ~sP5('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f269,f480,f915])).
fof(f12967,plain,(
  sP5('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f506,f496,f4960,f3746,f398])).
fof(f17180,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f12970,f13658])).
fof(f12970,plain,(
  sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),'T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f504,f498,f4960,f3746,f398])).
fof(f17179,plain,(
  $false | ($spl0 | $spl3 | $spl46)),
  inference(subsumption_resolution,[],[f12974,f13658])).
fof(f12974,plain,(
  sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),'T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | ($spl0 | $spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f504,f5251,f4960,f3746,f398])).
fof(f17178,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f12995,f13660])).
fof(f12995,plain,(
  ( ! [X0] : (sP5(X0,'T_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f504,f5727,f5720,f3746,f398])).
fof(f5720,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),X1)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f4934,f403])).
fof(f17177,plain,(
  $false | ($spl3 | $spl12 | $spl46)),
  inference(subsumption_resolution,[],[f12999,f5941])).
fof(f5941,plain,(
  ~sP5('T_0','T_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | ($spl3 | $spl12)),
  inference(unit_resulting_resolution,[],[f269,f4934,f2376])).
fof(f12999,plain,(
  sP5('T_0','T_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f504,f496,f5720,f3746,f398])).
fof(f17176,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13000,f5907])).
fof(f5907,plain,(
  ~sP5('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f269,f4934,f915])).
fof(f13000,plain,(
  sP5('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f506,f496,f5720,f3746,f398])).
fof(f17175,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13003,f13660])).
fof(f13003,plain,(
  sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),'T_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f504,f498,f5720,f3746,f398])).
fof(f17174,plain,(
  $false | ($spl0 | $spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13007,f13660])).
fof(f13007,plain,(
  sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),'T_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | ($spl0 | $spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f504,f5251,f5720,f3746,f398])).
fof(f17173,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13028,f13662])).
fof(f13028,plain,(
  ( ! [X0] : (sP5(X0,'T_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f504,f6705,f6696,f3746,f398])).
fof(f6696,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),X1)) ) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f5683,f403])).
fof(f17172,plain,(
  $false | ($spl3 | $spl4 | $spl12 | $spl46)),
  inference(subsumption_resolution,[],[f13032,f6977])).
fof(f6977,plain,(
  ~sP5('T_0','T_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14))) | ($spl3 | $spl4 | $spl12)),
  inference(unit_resulting_resolution,[],[f269,f5683,f2376])).
fof(f13032,plain,(
  sP5('T_0','T_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f504,f496,f6696,f3746,f398])).
fof(f17171,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13033,f6939])).
fof(f6939,plain,(
  ~sP5('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14))) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f269,f5683,f915])).
fof(f13033,plain,(
  sP5('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f506,f496,f6696,f3746,f398])).
fof(f17170,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13036,f13662])).
fof(f13036,plain,(
  sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),'T_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f504,f498,f6696,f3746,f398])).
fof(f17169,plain,(
  $false | ($spl0 | $spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13040,f13662])).
fof(f13040,plain,(
  sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),'T_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14))) | ($spl0 | $spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f504,f5251,f6696,f3746,f398])).
fof(f17168,plain,(
  $false | ($spl3 | $spl12 | $spl46)),
  inference(subsumption_resolution,[],[f13065,f9484])).
fof(f9484,plain,(
  ~sP5('T_0','T_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15))))) | ($spl3 | $spl12)),
  inference(unit_resulting_resolution,[],[f269,f5685,f2376])).
fof(f13065,plain,(
  sP5('T_0','T_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f504,f496,f9124,f3746,f398])).
fof(f9124,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))),X1)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f5685,f403])).
fof(f17167,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13066,f9440])).
fof(f9440,plain,(
  ~sP5('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15))))) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f269,f5685,f915])).
fof(f13066,plain,(
  sP5('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15))))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f506,f496,f9124,f3746,f398])).
fof(f17166,plain,(
  $false | ($spl0 | $spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13127,f16017])).
fof(f13127,plain,(
  sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),'F_0') | ($spl0 | $spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5253,f499,f4961,f3746,f398])).
fof(f4961,plain,(
  ( ! [X0,X1] : (sP4('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),X0,X1)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f480,f404])).
fof(f17165,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13128,f13658])).
fof(f13128,plain,(
  sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f504,f4965,f4961,f3746,f398])).
fof(f17164,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13131,f13660])).
fof(f13131,plain,(
  sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'T_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f504,f5727,f4961,f3746,f398])).
fof(f17163,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13134,f13662])).
fof(f13134,plain,(
  sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'T_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f504,f6705,f4961,f3746,f398])).
fof(f17162,plain,(
  $false | ($spl0 | $spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13157,f16017])).
fof(f13157,plain,(
  sP5('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),'F_0') | ($spl0 | $spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5253,f499,f5721,f3746,f398])).
fof(f5721,plain,(
  ( ! [X0,X1] : (sP4('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),X0,X1)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f4934,f404])).
fof(f17161,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13161,f13660])).
fof(f13161,plain,(
  sP5('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),'T_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f504,f5727,f5721,f3746,f398])).
fof(f17160,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13164,f13662])).
fof(f13164,plain,(
  sP5('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),'T_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f504,f6705,f5721,f3746,f398])).
fof(f17159,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13194,f13662])).
fof(f13194,plain,(
  sP5('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),'T_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f504,f6705,f6697,f3746,f398])).
fof(f6697,plain,(
  ( ! [X0,X1] : (sP4('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),X0,X1)) ) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f5683,f404])).
fof(f17158,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f13245,f13417])).
fof(f13417,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2)) ) | $spl46),
  inference(unit_resulting_resolution,[],[f3746,f402])).
fof(f13245,plain,(
  ~sP4('T_0','F_0','T_0') | $spl46),
  inference(unit_resulting_resolution,[],[f504,f499,f881,f3746,f398])).
fof(f17157,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f13246,f13489])).
fof(f13489,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2)) ) | $spl46),
  inference(unit_resulting_resolution,[],[f3746,f406])).
fof(f13246,plain,(
  ~sP3('T_0','F_0','T_0') | $spl46),
  inference(unit_resulting_resolution,[],[f504,f494,f881,f3746,f398])).
fof(f17156,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f13247,f13495])).
fof(f13495,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2)) ) | $spl46),
  inference(unit_resulting_resolution,[],[f3746,f412])).
fof(f412,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f241])).
fof(f13247,plain,(
  ~sP2('F_0','T_0','T_0') | $spl46),
  inference(unit_resulting_resolution,[],[f499,f494,f881,f3746,f398])).
fof(f17155,plain,(
  $false | ($spl3 | $spl12 | $spl46)),
  inference(subsumption_resolution,[],[f13249,f13417])).
fof(f13249,plain,(
  ~sP4('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'T_0') | ($spl3 | $spl12 | $spl46)),
  inference(unit_resulting_resolution,[],[f504,f4965,f5143,f3746,f398])).
fof(f17154,plain,(
  $false | ($spl3 | $spl12 | $spl46)),
  inference(subsumption_resolution,[],[f13250,f13489])).
fof(f13250,plain,(
  ~sP3('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'T_0') | ($spl3 | $spl12 | $spl46)),
  inference(unit_resulting_resolution,[],[f504,f4960,f5143,f3746,f398])).
fof(f17153,plain,(
  $false | ($spl3 | $spl12 | $spl46)),
  inference(subsumption_resolution,[],[f13251,f13495])).
fof(f13251,plain,(
  ~sP2('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'T_0','T_0') | ($spl3 | $spl12 | $spl46)),
  inference(unit_resulting_resolution,[],[f4965,f4960,f5143,f3746,f398])).
fof(f17152,plain,(
  $false | ($spl3 | $spl12 | $spl46)),
  inference(subsumption_resolution,[],[f13253,f13417])).
fof(f13253,plain,(
  ~sP4('T_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),'T_0') | ($spl3 | $spl12 | $spl46)),
  inference(unit_resulting_resolution,[],[f504,f5727,f5941,f3746,f398])).
fof(f17151,plain,(
  $false | ($spl3 | $spl12 | $spl46)),
  inference(subsumption_resolution,[],[f13254,f13489])).
fof(f13254,plain,(
  ~sP3('T_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),'T_0') | ($spl3 | $spl12 | $spl46)),
  inference(unit_resulting_resolution,[],[f504,f5720,f5941,f3746,f398])).
fof(f17150,plain,(
  $false | ($spl3 | $spl12 | $spl46)),
  inference(subsumption_resolution,[],[f13255,f13495])).
fof(f13255,plain,(
  ~sP2('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),'T_0','T_0') | ($spl3 | $spl12 | $spl46)),
  inference(unit_resulting_resolution,[],[f5727,f5720,f5941,f3746,f398])).
fof(f17149,plain,(
  $false | ($spl3 | $spl4 | $spl12 | $spl46)),
  inference(subsumption_resolution,[],[f13257,f13417])).
fof(f13257,plain,(
  ~sP4('T_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),'T_0') | ($spl3 | $spl4 | $spl12 | $spl46)),
  inference(unit_resulting_resolution,[],[f504,f6705,f6977,f3746,f398])).
fof(f17148,plain,(
  $false | ($spl3 | $spl4 | $spl12 | $spl46)),
  inference(subsumption_resolution,[],[f13258,f13489])).
fof(f13258,plain,(
  ~sP3('T_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),'T_0') | ($spl3 | $spl4 | $spl12 | $spl46)),
  inference(unit_resulting_resolution,[],[f504,f6696,f6977,f3746,f398])).
fof(f17147,plain,(
  $false | ($spl3 | $spl4 | $spl12 | $spl46)),
  inference(subsumption_resolution,[],[f13259,f13495])).
fof(f13259,plain,(
  ~sP2('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),'T_0','T_0') | ($spl3 | $spl4 | $spl12 | $spl46)),
  inference(unit_resulting_resolution,[],[f6705,f6696,f6977,f3746,f398])).
fof(f17146,plain,(
  $false | ($spl46 | $spl63)),
  inference(subsumption_resolution,[],[f13266,f13417])).
fof(f13266,plain,(
  ~sP4('T_0','F_0',f(s(bool,'T_0'))) | ($spl46 | $spl63)),
  inference(unit_resulting_resolution,[],[f501,f499,f4238,f3746,f398])).
fof(f4238,plain,(
  ~sP5('T_0',f(s(bool,'T_0')),'F_0') | $spl63),
  inference(unit_resulting_resolution,[],[f269,f4228,f415])).
fof(f17145,plain,(
  $false | ($spl46 | $spl63)),
  inference(subsumption_resolution,[],[f13267,f13489])).
fof(f13267,plain,(
  ~sP3(f(s(bool,'T_0')),'F_0','T_0') | ($spl46 | $spl63)),
  inference(unit_resulting_resolution,[],[f501,f494,f4238,f3746,f398])).
fof(f17144,plain,(
  $false | ($spl46 | $spl63)),
  inference(subsumption_resolution,[],[f13268,f13495])).
fof(f13268,plain,(
  ~sP2('F_0',f(s(bool,'T_0')),'T_0') | ($spl46 | $spl63)),
  inference(unit_resulting_resolution,[],[f499,f494,f4238,f3746,f398])).
fof(f17143,plain,(
  $false | ($spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13271,f13417])).
fof(f13271,plain,(
  ~sP4('T_0','F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))) | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f506,f499,f1062,f3746,f398])).
fof(f17142,plain,(
  $false | ($spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13272,f13489])).
fof(f13272,plain,(
  ~sP3('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),'F_0','T_0') | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f506,f494,f1062,f3746,f398])).
fof(f17141,plain,(
  $false | ($spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13273,f13495])).
fof(f13273,plain,(
  ~sP2('F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),'T_0') | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f499,f494,f1062,f3746,f398])).
fof(f17140,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f13275,f13417])).
fof(f13275,plain,(
  ~sP4('T_0','F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15))) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5253,f499,f5224,f3746,f398])).
fof(f17139,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f13276,f13489])).
fof(f13276,plain,(
  ~sP3('const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),'F_0','T_0') | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5253,f494,f5224,f3746,f398])).
fof(f17138,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f13277,f13495])).
fof(f13277,plain,(
  ~sP2('F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),'T_0') | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f499,f494,f5224,f3746,f398])).
fof(f17137,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13279,f13417])).
fof(f13279,plain,(
  ~sP4('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f506,f4965,f5113,f3746,f398])).
fof(f17136,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13280,f13489])).
fof(f13280,plain,(
  ~sP3('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'T_0') | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f506,f4960,f5113,f3746,f398])).
fof(f17135,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13281,f13495])).
fof(f13281,plain,(
  ~sP2('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),'T_0') | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f4965,f4960,f5113,f3746,f398])).
fof(f17134,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13284,f13417])).
fof(f13284,plain,(
  ~sP4('T_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f506,f5727,f5907,f3746,f398])).
fof(f17133,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13285,f13489])).
fof(f13285,plain,(
  ~sP3('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),'T_0') | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f506,f5720,f5907,f3746,f398])).
fof(f17132,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13286,f13495])).
fof(f13286,plain,(
  ~sP2('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),'T_0') | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5727,f5720,f5907,f3746,f398])).
fof(f17131,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13290,f13417])).
fof(f13290,plain,(
  ~sP4('T_0','F_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK13))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f501,f499,f9703,f3746,f398])).
fof(f9703,plain,(
  ~sP5('T_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),'F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f269,f483,f8009])).
fof(f8009,plain,(
  ( ! [X10,X9] : (~sP5(X9,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X10) | p(s(bool,X10)) | ~p(s(bool,X9))) ) | $spl3),
  inference(subsumption_resolution,[],[f7955,f269])).
fof(f7955,plain,(
  ( ! [X10,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,X9)) | p(s(bool,X10)) | ~sP5(X9,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X10)) ) | $spl3),
  inference(superposition,[],[f397,f4933])).
fof(f17130,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13291,f13489])).
fof(f13291,plain,(
  ~sP3('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),'F_0','T_0') | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f501,f494,f9703,f3746,f398])).
fof(f17129,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13292,f13495])).
fof(f13292,plain,(
  ~sP2('F_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),'T_0') | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f499,f494,f9703,f3746,f398])).
fof(f17128,plain,(
  $false | ($spl3 | $spl46 | $spl63)),
  inference(subsumption_resolution,[],[f13313,f13489])).
fof(f13313,plain,(
  ~sP3('T_0',f(s(bool,'T_0')),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | ($spl3 | $spl46 | $spl63)),
  inference(unit_resulting_resolution,[],[f504,f4961,f4996,f3746,f398])).
fof(f4996,plain,(
  ~sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'T_0',f(s(bool,'T_0'))) | ($spl3 | $spl63)),
  inference(unit_resulting_resolution,[],[f4246,f480,f416])).
fof(f416,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f242])).
fof(f17127,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f13320,f13417])).
fof(f13320,plain,(
  ~sP4('const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),'F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15))) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5253,f499,f5225,f3746,f398])).
fof(f17126,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f13321,f13489])).
fof(f13321,plain,(
  ~sP3('const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),'F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15))) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5253,f494,f5225,f3746,f398])).
fof(f17125,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f13322,f13495])).
fof(f13322,plain,(
  ~sP2('F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15))) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f499,f494,f5225,f3746,f398])).
fof(f17124,plain,(
  $false | ($spl0 | $spl12 | $spl46)),
  inference(subsumption_resolution,[],[f13325,f13417])).
fof(f13325,plain,(
  ~sP4('const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),'F_0','T_0') | ($spl0 | $spl12 | $spl46)),
  inference(unit_resulting_resolution,[],[f504,f499,f5405,f3746,f398])).
fof(f17123,plain,(
  $false | ($spl0 | $spl12 | $spl46)),
  inference(subsumption_resolution,[],[f13326,f13489])).
fof(f13326,plain,(
  ~sP3('T_0','F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15))) | ($spl0 | $spl12 | $spl46)),
  inference(unit_resulting_resolution,[],[f504,f494,f5405,f3746,f398])).
fof(f17122,plain,(
  $false | ($spl0 | $spl12 | $spl46)),
  inference(subsumption_resolution,[],[f13327,f13495])).
fof(f13327,plain,(
  ~sP2('F_0','T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15))) | ($spl0 | $spl12 | $spl46)),
  inference(unit_resulting_resolution,[],[f499,f494,f5405,f3746,f398])).
fof(f17121,plain,(
  $false | ($spl3 | $spl46 | $spl63)),
  inference(subsumption_resolution,[],[f13330,f13489])).
fof(f13330,plain,(
  ~sP3('T_0',f(s(bool,'T_0')),'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | ($spl3 | $spl46 | $spl63)),
  inference(unit_resulting_resolution,[],[f504,f5721,f5758,f3746,f398])).
fof(f5758,plain,(
  ~sP5('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),'T_0',f(s(bool,'T_0'))) | ($spl3 | $spl63)),
  inference(unit_resulting_resolution,[],[f4246,f4934,f416])).
fof(f17120,plain,(
  $false | ($spl3 | $spl4 | $spl46 | $spl63)),
  inference(subsumption_resolution,[],[f13337,f13489])).
fof(f13337,plain,(
  ~sP3('T_0',f(s(bool,'T_0')),'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14))) | ($spl3 | $spl4 | $spl46 | $spl63)),
  inference(unit_resulting_resolution,[],[f504,f6697,f6744,f3746,f398])).
fof(f6744,plain,(
  ~sP5('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),'T_0',f(s(bool,'T_0'))) | ($spl3 | $spl4 | $spl63)),
  inference(unit_resulting_resolution,[],[f4246,f5683,f416])).
fof(f17119,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13340,f3746])).
fof(f13340,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f4960,f483,f3746,f401])).
fof(f401,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f236])).
fof(f17118,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13341,f3746])).
fof(f13341,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5720,f483,f3746,f401])).
fof(f17117,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13342,f3746])).
fof(f13342,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6696,f483,f3746,f401])).
fof(f17116,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13343,f3746])).
fof(f13343,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9124,f483,f3746,f401])).
fof(f17115,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f13344,f3746])).
fof(f13344,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl46),
  inference(unit_resulting_resolution,[],[f495,f483,f3746,f401])).
fof(f17114,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13345,f3746])).
fof(f13345,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f4961,f483,f3746,f401])).
fof(f17113,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13346,f3746])).
fof(f13346,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5721,f483,f3746,f401])).
fof(f17112,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13347,f3746])).
fof(f13347,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6697,f483,f3746,f401])).
fof(f17111,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13348,f3746])).
fof(f13348,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9125,f483,f3746,f401])).
fof(f9125,plain,(
  ( ! [X0,X1] : (sP4('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))),X0,X1)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f5685,f404])).
fof(f17110,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f13376,f13417])).
fof(f13376,plain,(
  ( ! [X0,X1] : (~sP4(X0,X1,'F_0')) ) | $spl46),
  inference(unit_resulting_resolution,[],[f483,f3746,f3746,f401])).
fof(f17109,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f13378,f3746])).
fof(f13378,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl46),
  inference(unit_resulting_resolution,[],[f494,f483,f3746,f401])).
fof(f17108,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13379,f3746])).
fof(f13379,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f4960,f483,f3746,f401])).
fof(f17107,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13380,f3746])).
fof(f13380,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5720,f483,f3746,f401])).
fof(f17106,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13381,f3746])).
fof(f13381,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6696,f483,f3746,f401])).
fof(f17105,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13382,f3746])).
fof(f13382,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9124,f483,f3746,f401])).
fof(f17104,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13384,f3746])).
fof(f13384,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f4961,f483,f3746,f401])).
fof(f17103,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13385,f3746])).
fof(f13385,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5721,f483,f3746,f401])).
fof(f17102,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13386,f3746])).
fof(f13386,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6697,f483,f3746,f401])).
fof(f17101,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13387,f3746])).
fof(f13387,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9125,f483,f3746,f401])).
fof(f17100,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f13398,f13417])).
fof(f13398,plain,(
  ( ! [X0] : (~sP4('T_0',X0,'F_0')) ) | $spl46),
  inference(unit_resulting_resolution,[],[f483,f269,f3746,f401])).
fof(f17099,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f13415,f13417])).
fof(f13415,plain,(
  ( ! [X0,X1] : (~sP4(X0,X1,'F_0')) ) | $spl46),
  inference(unit_resulting_resolution,[],[f483,f3746,f3746,f401])).
fof(f17098,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13420,f3746])).
fof(f13420,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f4965,f483,f3746,f405])).
fof(f405,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f239])).
fof(f17097,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13421,f3746])).
fof(f13421,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5727,f483,f3746,f405])).
fof(f17096,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13422,f3746])).
fof(f13422,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6705,f483,f3746,f405])).
fof(f17095,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13423,f3746])).
fof(f13423,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9135,f483,f3746,f405])).
fof(f9135,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))),X1)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f5685,f407])).
fof(f17094,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f13424,f3746])).
fof(f13424,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl46),
  inference(unit_resulting_resolution,[],[f500,f483,f3746,f405])).
fof(f17093,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13425,f3746])).
fof(f13425,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f4966,f483,f3746,f405])).
fof(f4966,plain,(
  ( ! [X0,X1] : (sP3('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),X0,X1)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f480,f408])).
fof(f17092,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13426,f3746])).
fof(f13426,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5728,f483,f3746,f405])).
fof(f5728,plain,(
  ( ! [X0,X1] : (sP3('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),X0,X1)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f4934,f408])).
fof(f17091,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13427,f3746])).
fof(f13427,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6706,f483,f3746,f405])).
fof(f6706,plain,(
  ( ! [X0,X1] : (sP3('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),X0,X1)) ) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f5683,f408])).
fof(f17090,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13428,f3746])).
fof(f13428,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9136,f483,f3746,f405])).
fof(f9136,plain,(
  ( ! [X0,X1] : (sP3('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))),X0,X1)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f5685,f408])).
fof(f17089,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f13452,f13489])).
fof(f13452,plain,(
  ( ! [X0,X1] : (~sP3(X0,X1,'F_0')) ) | $spl46),
  inference(unit_resulting_resolution,[],[f483,f3746,f3746,f405])).
fof(f17088,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f13454,f3746])).
fof(f13454,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl46),
  inference(unit_resulting_resolution,[],[f499,f483,f3746,f405])).
fof(f17087,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13455,f3746])).
fof(f13455,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f4965,f483,f3746,f405])).
fof(f17086,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13456,f3746])).
fof(f13456,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5727,f483,f3746,f405])).
fof(f17085,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13457,f3746])).
fof(f13457,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6705,f483,f3746,f405])).
fof(f17084,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13458,f3746])).
fof(f13458,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9135,f483,f3746,f405])).
fof(f17083,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13460,f3746])).
fof(f13460,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f4966,f483,f3746,f405])).
fof(f17082,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13461,f3746])).
fof(f13461,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5728,f483,f3746,f405])).
fof(f17081,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13462,f3746])).
fof(f13462,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6706,f483,f3746,f405])).
fof(f17080,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13463,f3746])).
fof(f13463,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9136,f483,f3746,f405])).
fof(f17079,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f13472,f13489])).
fof(f13472,plain,(
  ( ! [X0] : (~sP3('T_0',X0,'F_0')) ) | $spl46),
  inference(unit_resulting_resolution,[],[f483,f269,f3746,f405])).
fof(f17078,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f13487,f13489])).
fof(f13487,plain,(
  ( ! [X0,X1] : (~sP3(X0,X1,'F_0')) ) | $spl46),
  inference(unit_resulting_resolution,[],[f483,f3746,f3746,f405])).
fof(f17065,plain,(
  $false | ($spl46 | $spl63)),
  inference(subsumption_resolution,[],[f13517,f4246])).
fof(f13517,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl46 | $spl63)),
  inference(unit_resulting_resolution,[],[f4457,f3746,f414])).
fof(f414,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f242])).
fof(f4457,plain,(
  ~sP5('F_0','T_0',f(s(bool,'T_0'))) | $spl63),
  inference(unit_resulting_resolution,[],[f483,f4246,f416])).
fof(f17064,plain,(
  $false | ($spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13518,f531])).
fof(f13518,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))) | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f893,f3746,f414])).
fof(f893,plain,(
  ~sP5('F_0','T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))) | $spl4),
  inference(unit_resulting_resolution,[],[f483,f531,f416])).
fof(f17063,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f13519,f5212])).
fof(f13519,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15))) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f6451,f3746,f414])).
fof(f6451,plain,(
  ~sP5('F_0','T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15))) | $spl0),
  inference(unit_resulting_resolution,[],[f5246,f395])).
fof(f395,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP3(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f233])).
fof(f5246,plain,(
  ~sP3('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f269,f483,f477,f405])).
fof(f17062,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13520,f4933])).
fof(f13520,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK13))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f7935,f3746,f414])).
fof(f7935,plain,(
  ~sP5('F_0','T_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK13))) | $spl3),
  inference(unit_resulting_resolution,[],[f483,f4933,f416])).
fof(f17061,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13522,f5112])).
fof(f5112,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | $spl3),
  inference(unit_resulting_resolution,[],[f527,f480,f832])).
fof(f13522,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5521,f3746,f414])).
fof(f5521,plain,(
  ~sP5('F_0','F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | $spl3),
  inference(unit_resulting_resolution,[],[f4967,f394])).
fof(f394,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f233])).
fof(f4967,plain,(
  ~sP2('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'F_0','F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f483,f483,f480,f409])).
fof(f409,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f241])).
fof(f17060,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13523,f5906])).
fof(f5906,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | $spl3),
  inference(unit_resulting_resolution,[],[f527,f4934,f832])).
fof(f13523,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f6336,f3746,f414])).
fof(f6336,plain,(
  ~sP5('F_0','F_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | $spl3),
  inference(unit_resulting_resolution,[],[f5729,f394])).
fof(f5729,plain,(
  ~sP2('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),'F_0','F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f483,f483,f4934,f409])).
fof(f17059,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13524,f7047])).
fof(f7047,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14))) | ($spl3 | $spl4)),
  inference(forward_demodulation,[],[f6938,f5112])).
fof(f6938,plain,(
  s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) = s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14))) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f4936,f5683,f832])).
fof(f13524,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f7432,f3746,f414])).
fof(f7432,plain,(
  ~sP5('F_0','F_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14))) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f6707,f394])).
fof(f6707,plain,(
  ~sP2('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),'F_0','F_0') | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f483,f483,f5683,f409])).
fof(f17058,plain,(
  $false | ($spl46 | $spl63)),
  inference(subsumption_resolution,[],[f13525,f4246])).
fof(f13525,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl46 | $spl63)),
  inference(unit_resulting_resolution,[],[f4455,f3746,f414])).
fof(f4455,plain,(
  ~sP5('F_0',f(s(bool,'T_0')),'T_0') | $spl63),
  inference(unit_resulting_resolution,[],[f483,f4246,f416])).
fof(f17057,plain,(
  $false | ($spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13526,f531])).
fof(f13526,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))) | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f890,f3746,f414])).
fof(f890,plain,(
  ~sP5('F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f483,f531,f416])).
fof(f17056,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f13527,f5212])).
fof(f13527,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15))) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f7448,f3746,f414])).
fof(f7448,plain,(
  ~sP5('F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),'T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f483,f5212,f416])).
fof(f17055,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13528,f5112])).
fof(f13528,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5526,f3746,f414])).
fof(f5526,plain,(
  ~sP5('F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f4975,f394])).
fof(f4975,plain,(
  ~sP2('F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f483,f483,f480,f409])).
fof(f17054,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13530,f5906])).
fof(f13530,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f6353,f3746,f414])).
fof(f6353,plain,(
  ~sP5('F_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),'F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f5737,f394])).
fof(f5737,plain,(
  ~sP2('F_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),'F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f483,f483,f4934,f409])).
fof(f17053,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13531,f7047])).
fof(f13531,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f7440,f3746,f414])).
fof(f7440,plain,(
  ~sP5('F_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),'F_0') | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f6717,f394])).
fof(f6717,plain,(
  ~sP2('F_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),'F_0') | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f483,f483,f5683,f409])).
fof(f17052,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13532,f4933])).
fof(f13532,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK13))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f7933,f3746,f414])).
fof(f7933,plain,(
  ~sP5('F_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),'T_0') | $spl3),
  inference(unit_resulting_resolution,[],[f483,f4933,f416])).
fof(f17051,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13534,f531])).
fof(f13534,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f4995,f3746,f414])).
fof(f4995,plain,(
  ~sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f531,f480,f416])).
fof(f17050,plain,(
  $false | ($spl3 | $spl46 | $spl63)),
  inference(subsumption_resolution,[],[f13535,f4246])).
fof(f13535,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl3 | $spl46 | $spl63)),
  inference(unit_resulting_resolution,[],[f4996,f3746,f414])).
fof(f17049,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13537,f5112])).
fof(f13537,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f6223,f3746,f414])).
fof(f6223,plain,(
  ~sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | $spl3),
  inference(unit_resulting_resolution,[],[f4984,f394])).
fof(f4984,plain,(
  ~sP2('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | $spl3),
  inference(unit_resulting_resolution,[],[f480,f483,f480,f409])).
fof(f17048,plain,(
  $false | ($spl3 | $spl46 | $spl63)),
  inference(subsumption_resolution,[],[f13538,f4246])).
fof(f13538,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl3 | $spl46 | $spl63)),
  inference(unit_resulting_resolution,[],[f4994,f3746,f414])).
fof(f4994,plain,(
  ~sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),f(s(bool,'T_0')),'T_0') | ($spl3 | $spl63)),
  inference(unit_resulting_resolution,[],[f4246,f480,f416])).
fof(f17047,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13539,f531])).
fof(f13539,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f4993,f3746,f414])).
fof(f4993,plain,(
  ~sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),'T_0') | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f531,f480,f416])).
fof(f17046,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13540,f5112])).
fof(f13540,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f6243,f3746,f414])).
fof(f6243,plain,(
  ~sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f4985,f394])).
fof(f4985,plain,(
  ~sP2('F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | $spl3),
  inference(unit_resulting_resolution,[],[f483,f480,f480,f409])).
fof(f17044,plain,(
  $false | ($spl3 | $spl46 | $spl63)),
  inference(subsumption_resolution,[],[f13546,f4246])).
fof(f13546,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl3 | $spl46 | $spl63)),
  inference(unit_resulting_resolution,[],[f5758,f3746,f414])).
fof(f17043,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13547,f531])).
fof(f13547,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5757,f3746,f414])).
fof(f5757,plain,(
  ~sP5('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),'T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f531,f4934,f416])).
fof(f17042,plain,(
  $false | ($spl3 | $spl46 | $spl63)),
  inference(subsumption_resolution,[],[f13549,f4246])).
fof(f13549,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl3 | $spl46 | $spl63)),
  inference(unit_resulting_resolution,[],[f5756,f3746,f414])).
fof(f5756,plain,(
  ~sP5('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),f(s(bool,'T_0')),'T_0') | ($spl3 | $spl63)),
  inference(unit_resulting_resolution,[],[f4246,f4934,f416])).
fof(f17041,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13550,f531])).
fof(f13550,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5755,f3746,f414])).
fof(f5755,plain,(
  ~sP5('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),'T_0') | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f531,f4934,f416])).
fof(f17040,plain,(
  $false | ($spl3 | $spl4 | $spl46 | $spl63)),
  inference(subsumption_resolution,[],[f13553,f4246])).
fof(f13553,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl3 | $spl4 | $spl46 | $spl63)),
  inference(unit_resulting_resolution,[],[f6744,f3746,f414])).
fof(f17039,plain,(
  $false | ($spl3 | $spl4 | $spl46 | $spl63)),
  inference(subsumption_resolution,[],[f13554,f4246])).
fof(f13554,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl3 | $spl4 | $spl46 | $spl63)),
  inference(unit_resulting_resolution,[],[f6742,f3746,f414])).
fof(f6742,plain,(
  ~sP5('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),f(s(bool,'T_0')),'T_0') | ($spl3 | $spl4 | $spl63)),
  inference(unit_resulting_resolution,[],[f4246,f5683,f416])).
fof(f17036,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f13603,f527])).
fof(f13603,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl46),
  inference(unit_resulting_resolution,[],[f859,f3746,f415])).
fof(f17035,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13604,f4936])).
fof(f13604,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5500,f3746,f415])).
fof(f5500,plain,(
  sP5('F_0','T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | $spl3),
  inference(unit_resulting_resolution,[],[f483,f504,f495,f4965,f400])).
fof(f17034,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13605,f5686])).
fof(f13605,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f6312,f3746,f415])).
fof(f6312,plain,(
  sP5('F_0','T_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | $spl3),
  inference(unit_resulting_resolution,[],[f483,f504,f495,f5727,f400])).
fof(f17033,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13606,f6640])).
fof(f13606,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f7400,f3746,f415])).
fof(f7400,plain,(
  sP5('F_0','T_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14))) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f483,f504,f495,f6705,f400])).
fof(f17032,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f13607,f527])).
fof(f13607,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl46),
  inference(unit_resulting_resolution,[],[f858,f3746,f415])).
fof(f858,plain,(
  sP5('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f483,f527,f413])).
fof(f17031,plain,(
  $false | ($spl4 | $spl46)),
  inference(subsumption_resolution,[],[f17030,f527])).
fof(f17030,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl4 | $spl46)),
  inference(forward_demodulation,[],[f13608,f531])).
fof(f13608,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))) | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f1288,f3746,f415])).
fof(f1288,plain,(
  sP5('F_0','F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))) | $spl4),
  inference(unit_resulting_resolution,[],[f509,f500,f495,f483,f399])).
fof(f509,plain,(
  ( ! [X0,X1] : (sP2('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),X0,X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f481,f412])).
fof(f17029,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f17028,f527])).
fof(f17028,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl0 | $spl46)),
  inference(forward_demodulation,[],[f13609,f5212])).
fof(f13609,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15))) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f6558,f3746,f415])).
fof(f6558,plain,(
  sP5('F_0','F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15))) | $spl0),
  inference(unit_resulting_resolution,[],[f483,f495,f500,f5254,f399])).
fof(f5254,plain,(
  ( ! [X0,X1] : (sP2('const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),X0,X1)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f477,f412])).
fof(f17027,plain,(
  $false | ($spl46 | $spl63)),
  inference(subsumption_resolution,[],[f17026,f527])).
fof(f17026,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl46 | $spl63)),
  inference(forward_demodulation,[],[f13610,f4246])).
fof(f13610,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | ($spl46 | $spl63)),
  inference(unit_resulting_resolution,[],[f4243,f3746,f415])).
fof(f4243,plain,(
  sP5('F_0','F_0',f(s(bool,'T_0'))) | $spl63),
  inference(unit_resulting_resolution,[],[f483,f4228,f413])).
fof(f17025,plain,(
  $false | ($spl46 | $spl63)),
  inference(subsumption_resolution,[],[f17024,f527])).
fof(f17024,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl46 | $spl63)),
  inference(forward_demodulation,[],[f13611,f4246])).
fof(f13611,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | ($spl46 | $spl63)),
  inference(unit_resulting_resolution,[],[f4237,f3746,f415])).
fof(f17023,plain,(
  $false | ($spl3 | $spl14 | $spl46)),
  inference(subsumption_resolution,[],[f13612,f4936])).
fof(f13612,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | ($spl3 | $spl14 | $spl46)),
  inference(unit_resulting_resolution,[],[f5513,f3746,f415])).
fof(f5513,plain,(
  sP5('F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'T_0') | ($spl3 | $spl14)),
  inference(unit_resulting_resolution,[],[f495,f4966,f2380])).
fof(f17022,plain,(
  $false | ($spl4 | $spl46)),
  inference(subsumption_resolution,[],[f17021,f527])).
fof(f17021,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl4 | $spl46)),
  inference(forward_demodulation,[],[f13613,f531])).
fof(f13613,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))) | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f1330,f3746,f415])).
fof(f17020,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f17019,f527])).
fof(f17019,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl0 | $spl46)),
  inference(forward_demodulation,[],[f13614,f5212])).
fof(f13614,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15))) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f6523,f3746,f415])).
fof(f17018,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f17017,f4936])).
fof(f17017,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | ($spl3 | $spl4 | $spl46)),
  inference(forward_demodulation,[],[f13615,f531])).
fof(f13615,plain,(
  s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5501,f3746,f415])).
fof(f5501,plain,(
  sP5('F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f483,f506,f495,f4965,f400])).
fof(f17016,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f17015,f4936])).
fof(f17015,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | ($spl3 | $spl4 | $spl46)),
  inference(forward_demodulation,[],[f13616,f531])).
fof(f13616,plain,(
  s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5507,f3746,f415])).
fof(f5507,plain,(
  sP5('F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f483,f509,f495,f4966,f400])).
fof(f17014,plain,(
  $false | ($spl0 | $spl3 | $spl46)),
  inference(subsumption_resolution,[],[f17013,f4936])).
fof(f17013,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | ($spl0 | $spl3 | $spl46)),
  inference(forward_demodulation,[],[f13617,f5212])).
fof(f13617,plain,(
  s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15))) | ($spl0 | $spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f6556,f3746,f415])).
fof(f6556,plain,(
  sP5('F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15))) | ($spl0 | $spl3)),
  inference(unit_resulting_resolution,[],[f483,f495,f4966,f5254,f400])).
fof(f17012,plain,(
  $false | ($spl0 | $spl3 | $spl46)),
  inference(subsumption_resolution,[],[f17011,f4936])).
fof(f17011,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | ($spl0 | $spl3 | $spl46)),
  inference(forward_demodulation,[],[f13618,f5212])).
fof(f13618,plain,(
  s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15))) | ($spl0 | $spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f6526,f3746,f415])).
fof(f6526,plain,(
  sP5('F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | ($spl0 | $spl3)),
  inference(unit_resulting_resolution,[],[f483,f495,f4965,f5253,f400])).
fof(f17010,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f17009,f527])).
fof(f17009,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl3 | $spl4 | $spl46)),
  inference(forward_demodulation,[],[f17008,f531])).
fof(f17008,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))) | ($spl3 | $spl4 | $spl46)),
  inference(forward_demodulation,[],[f13619,f5906])).
fof(f13619,plain,(
  s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))) = s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6313,f3746,f415])).
fof(f6313,plain,(
  sP5('F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f483,f506,f495,f5727,f400])).
fof(f17007,plain,(
  $false | ($spl3 | $spl14 | $spl46)),
  inference(subsumption_resolution,[],[f13620,f5686])).
fof(f13620,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | ($spl3 | $spl14 | $spl46)),
  inference(unit_resulting_resolution,[],[f6327,f3746,f415])).
fof(f6327,plain,(
  sP5('F_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),'T_0') | ($spl3 | $spl14)),
  inference(unit_resulting_resolution,[],[f495,f5728,f2380])).
fof(f17006,plain,(
  $false | ($spl3 | $spl4 | $spl14 | $spl46)),
  inference(subsumption_resolution,[],[f13621,f6640])).
fof(f13621,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14))) | ($spl3 | $spl4 | $spl14 | $spl46)),
  inference(unit_resulting_resolution,[],[f7422,f3746,f415])).
fof(f7422,plain,(
  sP5('F_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),'T_0') | ($spl3 | $spl4 | $spl14)),
  inference(unit_resulting_resolution,[],[f495,f6706,f2380])).
fof(f17005,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f17004,f527])).
fof(f17004,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl3 | $spl4 | $spl46)),
  inference(forward_demodulation,[],[f17003,f531])).
fof(f17003,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))) | ($spl3 | $spl4 | $spl46)),
  inference(forward_demodulation,[],[f13622,f5906])).
fof(f13622,plain,(
  s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))) = s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6320,f3746,f415])).
fof(f6320,plain,(
  sP5('F_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f483,f509,f495,f5728,f400])).
fof(f17002,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13625,f527])).
fof(f13625,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f4992,f3746,f415])).
fof(f17001,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13629,f527])).
fof(f13629,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f4991,f3746,f415])).
fof(f4991,plain,(
  sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'F_0','T_0') | $spl3),
  inference(unit_resulting_resolution,[],[f527,f480,f413])).
fof(f17000,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f16999,f527])).
fof(f16999,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl3 | $spl4 | $spl46)),
  inference(forward_demodulation,[],[f13633,f531])).
fof(f13633,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5487,f3746,f415])).
fof(f5487,plain,(
  sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f483,f509,f500,f4961,f399])).
fof(f16998,plain,(
  $false | ($spl0 | $spl3 | $spl46)),
  inference(subsumption_resolution,[],[f16997,f527])).
fof(f16997,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl0 | $spl3 | $spl46)),
  inference(forward_demodulation,[],[f13634,f5212])).
fof(f13634,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15))) | ($spl0 | $spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f6559,f3746,f415])).
fof(f6559,plain,(
  sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15))) | ($spl0 | $spl3)),
  inference(unit_resulting_resolution,[],[f483,f4961,f500,f5254,f399])).
fof(f16996,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f16995,f527])).
fof(f16995,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl3 | $spl4 | $spl46)),
  inference(forward_demodulation,[],[f13639,f531])).
fof(f13639,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f4956,f3746,f415])).
fof(f16994,plain,(
  $false | ($spl3 | $spl14 | $spl46)),
  inference(subsumption_resolution,[],[f13640,f4936])).
fof(f13640,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | ($spl3 | $spl14 | $spl46)),
  inference(unit_resulting_resolution,[],[f5514,f3746,f415])).
fof(f5514,plain,(
  sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'T_0') | ($spl3 | $spl14)),
  inference(unit_resulting_resolution,[],[f4961,f4966,f2380])).
fof(f16993,plain,(
  $false | ($spl3 | $spl14 | $spl46)),
  inference(subsumption_resolution,[],[f13643,f5686])).
fof(f13643,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | ($spl3 | $spl14 | $spl46)),
  inference(unit_resulting_resolution,[],[f6328,f3746,f415])).
fof(f6328,plain,(
  sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),'T_0') | ($spl3 | $spl14)),
  inference(unit_resulting_resolution,[],[f4961,f5728,f2380])).
fof(f16991,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13645,f527])).
fof(f13645,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5753,f3746,f415])).
fof(f5753,plain,(
  sP5('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),'F_0','T_0') | $spl3),
  inference(unit_resulting_resolution,[],[f527,f4934,f413])).
fof(f16990,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f16989,f527])).
fof(f16989,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl3 | $spl4 | $spl46)),
  inference(forward_demodulation,[],[f13646,f531])).
fof(f13646,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6297,f3746,f415])).
fof(f6297,plain,(
  sP5('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),'F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f483,f509,f500,f5721,f399])).
fof(f16988,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13647,f527])).
fof(f13647,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5754,f3746,f415])).
fof(f16987,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13648,f4936])).
fof(f13648,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5713,f3746,f415])).
fof(f5713,plain,(
  sP5('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),'T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | $spl3),
  inference(unit_resulting_resolution,[],[f504,f4965,f4960,f4934,f400])).
fof(f16986,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f16985,f527])).
fof(f16985,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl3 | $spl4 | $spl46)),
  inference(forward_demodulation,[],[f13649,f531])).
fof(f13649,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5712,f3746,f415])).
fof(f16984,plain,(
  $false | ($spl3 | $spl14 | $spl46)),
  inference(subsumption_resolution,[],[f13650,f4936])).
fof(f13650,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | ($spl3 | $spl14 | $spl46)),
  inference(unit_resulting_resolution,[],[f6304,f3746,f415])).
fof(f6304,plain,(
  sP5('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'T_0') | ($spl3 | $spl14)),
  inference(unit_resulting_resolution,[],[f4966,f5721,f2380])).
fof(f16983,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f16982,f4936])).
fof(f16982,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | ($spl3 | $spl4 | $spl46)),
  inference(forward_demodulation,[],[f13651,f531])).
fof(f13651,plain,(
  s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5714,f3746,f415])).
fof(f5714,plain,(
  sP5('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f506,f4965,f4960,f4934,f400])).
fof(f16981,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13652,f527])).
fof(f13652,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6737,f3746,f415])).
fof(f6737,plain,(
  sP5('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),'F_0','T_0') | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f527,f5683,f413])).
fof(f16980,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13653,f527])).
fof(f13653,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6739,f3746,f415])).
fof(f16979,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13654,f527])).
fof(f13654,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9167,f3746,f415])).
fof(f9167,plain,(
  sP5('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))),'F_0','T_0') | $spl3),
  inference(unit_resulting_resolution,[],[f527,f5685,f413])).
fof(f16978,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13655,f527])).
fof(f13655,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9171,f3746,f415])).
fof(f16977,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f13705,f14292])).
fof(f14292,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2)) ) | $spl46),
  inference(unit_resulting_resolution,[],[f3746,f3746,f3746,f430])).
fof(f430,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f249])).
fof(f249,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f248])).
fof(f248,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f154])).
fof(f154,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP7(X2,X0,X1))),
  inference(definition_folding,[],[f113,f153,f152])).
fof(f152,plain,(
  ! [X0,X1,X2] : (sP6(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f153,plain,(
  ! [X2,X0,X1] : (sP7(X2,X0,X1) <=> (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f113,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f112])).
fof(f112,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X20,X21,X22] : ((p(s(bool,X22)) <=> (p(s(bool,X21)) & p(s(bool,X20)))) <=> ((p(s(bool,X22)) | ~p(s(bool,X21)) | ~p(s(bool,X20))) & (p(s(bool,X21)) | ~p(s(bool,X22))) & (p(s(bool,X20)) | ~p(s(bool,X22)))))),
  file('thm.pred_set.UNION_SUBSET','thm.sat.sat_23_prove_ltb')).
fof(f13705,plain,(
  ( ! [X0,X1] : (~sP7(X0,X1,'F_0')) ) | $spl46),
  inference(unit_resulting_resolution,[],[f483,f3746,f418])).
fof(f418,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f245])).
fof(f245,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | (sP6(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP6(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f244])).
fof(f244,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP7(X2,X0,X1)))),
  inference(flattening,[],[f243])).
fof(f243,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP7(X2,X0,X1)))),
  inference(nnf_transformation,[],[f153])).
fof(f16976,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f13740,f14292])).
fof(f13740,plain,(
  ( ! [X0,X1] : (~sP7(X0,'F_0',X1)) ) | $spl46),
  inference(unit_resulting_resolution,[],[f483,f3746,f419])).
fof(f419,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f245])).
fof(f16975,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13743,f3746])).
fof(f13743,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5004,f4997,f3746,f420])).
fof(f420,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f245])).
fof(f4997,plain,(
  ( ! [X0] : (~sP7('T_0',X0,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f269,f480,f418])).
fof(f5004,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),X1)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f480,f426])).
fof(f426,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f247])).
fof(f247,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(flattening,[],[f246])).
fof(f246,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(nnf_transformation,[],[f152])).
fof(f16974,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13744,f3746])).
fof(f13744,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5770,f5759,f3746,f420])).
fof(f5759,plain,(
  ( ! [X0] : (~sP7('T_0',X0,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f269,f4934,f418])).
fof(f5770,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),X1)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f4934,f426])).
fof(f16973,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13745,f3746])).
fof(f13745,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6766,f6745,f3746,f420])).
fof(f6745,plain,(
  ( ! [X0] : (~sP7('T_0',X0,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f269,f5683,f418])).
fof(f6766,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),X1)) ) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f5683,f426])).
fof(f16972,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13746,f3746])).
fof(f13746,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9216,f9189,f3746,f420])).
fof(f9189,plain,(
  ( ! [X0] : (~sP7('T_0',X0,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f269,f5685,f418])).
fof(f9216,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))),X1)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f5685,f426])).
fof(f16971,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f13747,f3746])).
fof(f13747,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl46),
  inference(unit_resulting_resolution,[],[f514,f636,f3746,f420])).
fof(f636,plain,(
  ( ! [X0] : (~sP7('T_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f269,f483,f419])).
fof(f514,plain,(
  ( ! [X0,X1] : (sP6('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f483,f427])).
fof(f427,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f247])).
fof(f16970,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f13748,f3746])).
fof(f13748,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl46),
  inference(unit_resulting_resolution,[],[f510,f636,f3746,f420])).
fof(f510,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f269,f425])).
fof(f425,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f247])).
fof(f16969,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13750,f3746])).
fof(f13750,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5004,f636,f3746,f420])).
fof(f16968,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13751,f3746])).
fof(f13751,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5770,f636,f3746,f420])).
fof(f16967,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13752,f3746])).
fof(f13752,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6766,f636,f3746,f420])).
fof(f16966,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13753,f3746])).
fof(f13753,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9216,f636,f3746,f420])).
fof(f16965,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13754,f3746])).
fof(f13754,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5005,f4999,f3746,f420])).
fof(f4999,plain,(
  ( ! [X0] : (~sP7('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),X0)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f269,f480,f419])).
fof(f5005,plain,(
  ( ! [X0,X1] : (sP6('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),X0,X1)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f480,f427])).
fof(f16964,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13755,f3746])).
fof(f13755,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f510,f4999,f3746,f420])).
fof(f16963,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13757,f3746])).
fof(f13757,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5004,f4999,f3746,f420])).
fof(f16962,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13758,f3746])).
fof(f13758,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5770,f4999,f3746,f420])).
fof(f16961,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13759,f3746])).
fof(f13759,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6766,f4999,f3746,f420])).
fof(f16960,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13760,f3746])).
fof(f13760,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9216,f4999,f3746,f420])).
fof(f16959,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13761,f3746])).
fof(f13761,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5771,f5762,f3746,f420])).
fof(f5762,plain,(
  ( ! [X0] : (~sP7('T_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),X0)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f269,f4934,f419])).
fof(f5771,plain,(
  ( ! [X0,X1] : (sP6('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),X0,X1)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f4934,f427])).
fof(f16958,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13762,f3746])).
fof(f13762,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f510,f5762,f3746,f420])).
fof(f16957,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13764,f3746])).
fof(f13764,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5004,f5762,f3746,f420])).
fof(f16956,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13765,f3746])).
fof(f13765,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5770,f5762,f3746,f420])).
fof(f16955,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13766,f3746])).
fof(f13766,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6766,f5762,f3746,f420])).
fof(f16954,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13767,f3746])).
fof(f13767,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9216,f5762,f3746,f420])).
fof(f16953,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13768,f3746])).
fof(f13768,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6767,f6752,f3746,f420])).
fof(f6752,plain,(
  ( ! [X0] : (~sP7('T_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),X0)) ) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f269,f5683,f419])).
fof(f6767,plain,(
  ( ! [X0,X1] : (sP6('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),X0,X1)) ) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f5683,f427])).
fof(f16952,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13769,f3746])).
fof(f13769,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f510,f6752,f3746,f420])).
fof(f16951,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13771,f3746])).
fof(f13771,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5004,f6752,f3746,f420])).
fof(f16950,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13772,f3746])).
fof(f13772,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5770,f6752,f3746,f420])).
fof(f16949,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13773,f3746])).
fof(f13773,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6766,f6752,f3746,f420])).
fof(f16948,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13774,f3746])).
fof(f13774,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f9216,f6752,f3746,f420])).
fof(f16947,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13775,f3746])).
fof(f13775,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9217,f9198,f3746,f420])).
fof(f9198,plain,(
  ( ! [X0] : (~sP7('T_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))),X0)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f269,f5685,f419])).
fof(f9217,plain,(
  ( ! [X0,X1] : (sP6('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))),X0,X1)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f5685,f427])).
fof(f16946,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13776,f3746])).
fof(f13776,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f510,f9198,f3746,f420])).
fof(f16945,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13778,f3746])).
fof(f13778,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5004,f9198,f3746,f420])).
fof(f16944,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13779,f3746])).
fof(f13779,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5770,f9198,f3746,f420])).
fof(f16943,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13780,f3746])).
fof(f13780,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6766,f9198,f3746,f420])).
fof(f16942,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13781,f3746])).
fof(f13781,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9216,f9198,f3746,f420])).
fof(f16941,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f13783,f3746])).
fof(f13783,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f514,f5262,f3746,f420])).
fof(f5262,plain,(
  ( ! [X0] : (~sP7('const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),'F_0',X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f483,f477,f419])).
fof(f16940,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f13784,f3746])).
fof(f13784,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5283,f5262,f3746,f420])).
fof(f5283,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f477,f425])).
fof(f16939,plain,(
  $false | ($spl0 | $spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13786,f3746])).
fof(f13786,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl0 | $spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5004,f5262,f3746,f420])).
fof(f16938,plain,(
  $false | ($spl0 | $spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13787,f3746])).
fof(f13787,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl0 | $spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5770,f5262,f3746,f420])).
fof(f16937,plain,(
  $false | ($spl0 | $spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13788,f3746])).
fof(f13788,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl0 | $spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6766,f5262,f3746,f420])).
fof(f16936,plain,(
  $false | ($spl0 | $spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13789,f3746])).
fof(f13789,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | ($spl0 | $spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9216,f5262,f3746,f420])).
fof(f16935,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f13798,f14058])).
fof(f14058,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2)) ) | $spl46),
  inference(unit_resulting_resolution,[],[f3746,f425])).
fof(f13798,plain,(
  ~sP6('F_0','T_0','T_0') | $spl46),
  inference(unit_resulting_resolution,[],[f636,f269,f3746,f420])).
fof(f16934,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13800,f14058])).
fof(f13800,plain,(
  ~sP6('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'T_0','T_0') | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f4999,f269,f3746,f420])).
fof(f16933,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13802,f14058])).
fof(f13802,plain,(
  ~sP6('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),'T_0','T_0') | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5762,f269,f3746,f420])).
fof(f16932,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13804,f14058])).
fof(f13804,plain,(
  ~sP6('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),'T_0','T_0') | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6752,f269,f3746,f420])).
fof(f16931,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13806,f14058])).
fof(f13806,plain,(
  ~sP6('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))),'T_0','T_0') | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9198,f269,f3746,f420])).
fof(f16930,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f13811,f14058])).
fof(f13811,plain,(
  ~sP6('F_0','T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15))) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5262,f269,f3746,f420])).
fof(f16929,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f13828,f14058])).
fof(f13828,plain,(
  ( ! [X0] : (~sP6(X0,'F_0','T_0')) ) | $spl46),
  inference(unit_resulting_resolution,[],[f634,f3746,f3746,f420])).
fof(f634,plain,(
  ( ! [X0] : (~sP7('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f269,f483,f418])).
fof(f16928,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13830,f14058])).
fof(f13830,plain,(
  ( ! [X0] : (~sP6(X0,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'T_0')) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f4997,f3746,f3746,f420])).
fof(f16927,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13832,f14058])).
fof(f13832,plain,(
  ( ! [X0] : (~sP6(X0,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),'T_0')) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5759,f3746,f3746,f420])).
fof(f16926,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13834,f14058])).
fof(f13834,plain,(
  ( ! [X0] : (~sP6(X0,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),'T_0')) ) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6745,f3746,f3746,f420])).
fof(f16925,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13836,f14058])).
fof(f13836,plain,(
  ( ! [X0] : (~sP6(X0,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))),'T_0')) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9189,f3746,f3746,f420])).
fof(f16924,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f13838,f14058])).
fof(f13838,plain,(
  ( ! [X0] : (~sP6('F_0',X0,'T_0')) ) | $spl46),
  inference(unit_resulting_resolution,[],[f636,f3746,f3746,f420])).
fof(f16923,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13840,f14058])).
fof(f13840,plain,(
  ( ! [X0] : (~sP6('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),X0,'T_0')) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f4999,f3746,f3746,f420])).
fof(f16922,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13842,f14058])).
fof(f13842,plain,(
  ( ! [X0] : (~sP6('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),X0,'T_0')) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5762,f3746,f3746,f420])).
fof(f16921,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13844,f14058])).
fof(f13844,plain,(
  ( ! [X0] : (~sP6('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),X0,'T_0')) ) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6752,f3746,f3746,f420])).
fof(f16920,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13846,f14058])).
fof(f13846,plain,(
  ( ! [X0] : (~sP6('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))),X0,'T_0')) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9198,f3746,f3746,f420])).
fof(f16919,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f13856,f14058])).
fof(f13856,plain,(
  ( ! [X0] : (~sP6(X0,'F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)))) ) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5261,f3746,f3746,f420])).
fof(f5261,plain,(
  ( ! [X0] : (~sP7('const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),X0,'F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f483,f477,f418])).
fof(f16918,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f13858,f14058])).
fof(f13858,plain,(
  ( ! [X0] : (~sP6('F_0',X0,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)))) ) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5262,f3746,f3746,f420])).
fof(f16917,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f13863,f3746])).
fof(f13863,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl46),
  inference(unit_resulting_resolution,[],[f513,f634,f3746,f420])).
fof(f513,plain,(
  ( ! [X0,X1] : (sP6(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f483,f426])).
fof(f16916,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f13864,f3746])).
fof(f13864,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl46),
  inference(unit_resulting_resolution,[],[f510,f634,f3746,f420])).
fof(f16915,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13866,f3746])).
fof(f13866,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5005,f634,f3746,f420])).
fof(f16914,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13867,f3746])).
fof(f13867,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5771,f634,f3746,f420])).
fof(f16913,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13868,f3746])).
fof(f13868,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6767,f634,f3746,f420])).
fof(f16912,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13869,f3746])).
fof(f13869,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9217,f634,f3746,f420])).
fof(f16911,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13870,f3746])).
fof(f13870,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5004,f4997,f3746,f420])).
fof(f16910,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13871,f3746])).
fof(f13871,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f510,f4997,f3746,f420])).
fof(f16909,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13873,f3746])).
fof(f13873,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5005,f4997,f3746,f420])).
fof(f16908,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13874,f3746])).
fof(f13874,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5771,f4997,f3746,f420])).
fof(f16907,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13875,f3746])).
fof(f13875,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6767,f4997,f3746,f420])).
fof(f16906,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13876,f3746])).
fof(f13876,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9217,f4997,f3746,f420])).
fof(f16905,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13877,f3746])).
fof(f13877,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5770,f5759,f3746,f420])).
fof(f16904,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13878,f3746])).
fof(f13878,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f510,f5759,f3746,f420])).
fof(f16903,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13880,f3746])).
fof(f13880,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5005,f5759,f3746,f420])).
fof(f16902,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13881,f3746])).
fof(f13881,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5771,f5759,f3746,f420])).
fof(f16901,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13882,f3746])).
fof(f13882,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6767,f5759,f3746,f420])).
fof(f16900,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13883,f3746])).
fof(f13883,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9217,f5759,f3746,f420])).
fof(f16899,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13884,f3746])).
fof(f13884,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6766,f6745,f3746,f420])).
fof(f16898,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13885,f3746])).
fof(f13885,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f510,f6745,f3746,f420])).
fof(f16897,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13887,f3746])).
fof(f13887,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5005,f6745,f3746,f420])).
fof(f16896,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13888,f3746])).
fof(f13888,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5771,f6745,f3746,f420])).
fof(f16895,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13889,f3746])).
fof(f13889,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6767,f6745,f3746,f420])).
fof(f16894,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13890,f3746])).
fof(f13890,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f9217,f6745,f3746,f420])).
fof(f16893,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13891,f3746])).
fof(f13891,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9216,f9189,f3746,f420])).
fof(f16892,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13892,f3746])).
fof(f13892,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f510,f9189,f3746,f420])).
fof(f16891,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13894,f3746])).
fof(f13894,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5005,f9189,f3746,f420])).
fof(f16890,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13895,f3746])).
fof(f13895,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5771,f9189,f3746,f420])).
fof(f16889,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13896,f3746])).
fof(f13896,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6767,f9189,f3746,f420])).
fof(f16888,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13897,f3746])).
fof(f13897,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9217,f9189,f3746,f420])).
fof(f16887,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13899,f3746])).
fof(f13899,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5005,f4999,f3746,f420])).
fof(f16886,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13900,f3746])).
fof(f13900,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5771,f5762,f3746,f420])).
fof(f16885,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13901,f3746])).
fof(f13901,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6767,f6752,f3746,f420])).
fof(f16884,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13902,f3746])).
fof(f13902,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9217,f9198,f3746,f420])).
fof(f16883,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f13903,f3746])).
fof(f13903,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f513,f5261,f3746,f420])).
fof(f16882,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f13904,f3746])).
fof(f13904,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5283,f5261,f3746,f420])).
fof(f16881,plain,(
  $false | ($spl0 | $spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13906,f3746])).
fof(f13906,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl0 | $spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5005,f5261,f3746,f420])).
fof(f16880,plain,(
  $false | ($spl0 | $spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13907,f3746])).
fof(f13907,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl0 | $spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5771,f5261,f3746,f420])).
fof(f16879,plain,(
  $false | ($spl0 | $spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13908,f3746])).
fof(f13908,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl0 | $spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6767,f5261,f3746,f420])).
fof(f16878,plain,(
  $false | ($spl0 | $spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13909,f3746])).
fof(f13909,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | ($spl0 | $spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9217,f5261,f3746,f420])).
fof(f16877,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f13919,f14058])).
fof(f13919,plain,(
  ~sP6('T_0','F_0','T_0') | $spl46),
  inference(unit_resulting_resolution,[],[f634,f269,f3746,f420])).
fof(f16876,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13921,f14058])).
fof(f13921,plain,(
  ~sP6('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'T_0') | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f4997,f269,f3746,f420])).
fof(f16875,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13923,f14058])).
fof(f13923,plain,(
  ~sP6('T_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),'T_0') | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5759,f269,f3746,f420])).
fof(f16874,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13925,f14058])).
fof(f13925,plain,(
  ~sP6('T_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),'T_0') | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6745,f269,f3746,f420])).
fof(f16873,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13927,f14058])).
fof(f13927,plain,(
  ~sP6('T_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))),'T_0') | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9189,f269,f3746,f420])).
fof(f16872,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f13935,f14058])).
fof(f13935,plain,(
  ~sP6('T_0','F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15))) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5261,f269,f3746,f420])).
fof(f16871,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f13953,f14058])).
fof(f13953,plain,(
  ( ! [X0] : (~sP6(X0,'F_0','T_0')) ) | $spl46),
  inference(unit_resulting_resolution,[],[f634,f3746,f3746,f420])).
fof(f16870,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13955,f14058])).
fof(f13955,plain,(
  ( ! [X0] : (~sP6(X0,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'T_0')) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f4997,f3746,f3746,f420])).
fof(f16869,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13957,f14058])).
fof(f13957,plain,(
  ( ! [X0] : (~sP6(X0,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),'T_0')) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5759,f3746,f3746,f420])).
fof(f16868,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13959,f14058])).
fof(f13959,plain,(
  ( ! [X0] : (~sP6(X0,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),'T_0')) ) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6745,f3746,f3746,f420])).
fof(f16867,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13961,f14058])).
fof(f13961,plain,(
  ( ! [X0] : (~sP6(X0,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))),'T_0')) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9189,f3746,f3746,f420])).
fof(f16866,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f13963,f14058])).
fof(f13963,plain,(
  ( ! [X0] : (~sP6('F_0',X0,'T_0')) ) | $spl46),
  inference(unit_resulting_resolution,[],[f636,f3746,f3746,f420])).
fof(f16865,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13965,f14058])).
fof(f13965,plain,(
  ( ! [X0] : (~sP6('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),X0,'T_0')) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f4999,f3746,f3746,f420])).
fof(f16864,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13967,f14058])).
fof(f13967,plain,(
  ( ! [X0] : (~sP6('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),X0,'T_0')) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5762,f3746,f3746,f420])).
fof(f16863,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13969,f14058])).
fof(f13969,plain,(
  ( ! [X0] : (~sP6('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),X0,'T_0')) ) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6752,f3746,f3746,f420])).
fof(f16862,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13971,f14058])).
fof(f13971,plain,(
  ( ! [X0] : (~sP6('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))),X0,'T_0')) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9198,f3746,f3746,f420])).
fof(f16861,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f13981,f14058])).
fof(f13981,plain,(
  ( ! [X0] : (~sP6(X0,'F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)))) ) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5261,f3746,f3746,f420])).
fof(f16860,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f13983,f14058])).
fof(f13983,plain,(
  ( ! [X0] : (~sP6('F_0',X0,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)))) ) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5262,f3746,f3746,f420])).
fof(f16859,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13989,f3746])).
fof(f13989,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5004,f483,f3746,f424])).
fof(f424,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f247])).
fof(f16858,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13990,f3746])).
fof(f13990,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5770,f483,f3746,f424])).
fof(f16857,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13991,f3746])).
fof(f13991,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6766,f483,f3746,f424])).
fof(f16856,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13992,f3746])).
fof(f13992,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9216,f483,f3746,f424])).
fof(f16855,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f13993,f3746])).
fof(f13993,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl46),
  inference(unit_resulting_resolution,[],[f514,f483,f3746,f424])).
fof(f16854,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13994,f3746])).
fof(f13994,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5005,f483,f3746,f424])).
fof(f16853,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13995,f3746])).
fof(f13995,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5771,f483,f3746,f424])).
fof(f16852,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f13996,f3746])).
fof(f13996,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6767,f483,f3746,f424])).
fof(f16851,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f13997,f3746])).
fof(f13997,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9217,f483,f3746,f424])).
fof(f16850,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f14021,f14058])).
fof(f14021,plain,(
  ( ! [X0,X1] : (~sP6(X0,X1,'F_0')) ) | $spl46),
  inference(unit_resulting_resolution,[],[f483,f3746,f3746,f424])).
fof(f16849,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f14023,f3746])).
fof(f14023,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl46),
  inference(unit_resulting_resolution,[],[f513,f483,f3746,f424])).
fof(f16848,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14024,f3746])).
fof(f14024,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5004,f483,f3746,f424])).
fof(f16847,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14025,f3746])).
fof(f14025,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5770,f483,f3746,f424])).
fof(f16846,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14026,f3746])).
fof(f14026,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6766,f483,f3746,f424])).
fof(f16845,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14027,f3746])).
fof(f14027,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9216,f483,f3746,f424])).
fof(f16844,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14029,f3746])).
fof(f14029,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5005,f483,f3746,f424])).
fof(f16843,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14030,f3746])).
fof(f14030,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5771,f483,f3746,f424])).
fof(f16842,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14031,f3746])).
fof(f14031,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6767,f483,f3746,f424])).
fof(f16841,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14032,f3746])).
fof(f14032,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9217,f483,f3746,f424])).
fof(f16840,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f14041,f14058])).
fof(f14041,plain,(
  ( ! [X0] : (~sP6('T_0',X0,'F_0')) ) | $spl46),
  inference(unit_resulting_resolution,[],[f483,f269,f3746,f424])).
fof(f16839,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f14056,f14058])).
fof(f14056,plain,(
  ( ! [X0,X1] : (~sP6(X0,X1,'F_0')) ) | $spl46),
  inference(unit_resulting_resolution,[],[f483,f3746,f3746,f424])).
fof(f16838,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14061,f3746])).
fof(f14061,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f4997,f269,f3746,f430])).
fof(f16837,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14062,f3746])).
fof(f14062,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5759,f269,f3746,f430])).
fof(f16836,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14063,f3746])).
fof(f14063,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6745,f269,f3746,f430])).
fof(f16835,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14064,f3746])).
fof(f14064,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9189,f269,f3746,f430])).
fof(f16834,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f14065,f3746])).
fof(f14065,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl46),
  inference(unit_resulting_resolution,[],[f636,f269,f3746,f430])).
fof(f16833,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14066,f3746])).
fof(f14066,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f4999,f269,f3746,f430])).
fof(f16832,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14067,f3746])).
fof(f14067,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5762,f269,f3746,f430])).
fof(f16831,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14068,f3746])).
fof(f14068,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6752,f269,f3746,f430])).
fof(f16830,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14069,f3746])).
fof(f14069,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9198,f269,f3746,f430])).
fof(f16829,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14071,f3746])).
fof(f14071,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f4997,f3746,f3746,f430])).
fof(f16828,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14072,f3746])).
fof(f14072,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5759,f3746,f3746,f430])).
fof(f16827,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14073,f3746])).
fof(f14073,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6745,f3746,f3746,f430])).
fof(f16826,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14074,f3746])).
fof(f14074,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9189,f3746,f3746,f430])).
fof(f16825,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f14075,f3746])).
fof(f14075,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl46),
  inference(unit_resulting_resolution,[],[f636,f3746,f3746,f430])).
fof(f16824,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14076,f3746])).
fof(f14076,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f4999,f3746,f3746,f430])).
fof(f16823,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14077,f3746])).
fof(f14077,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5762,f3746,f3746,f430])).
fof(f16822,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14078,f3746])).
fof(f14078,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6752,f3746,f3746,f430])).
fof(f16821,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14079,f3746])).
fof(f14079,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9198,f3746,f3746,f430])).
fof(f16820,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f14080,f269])).
fof(f14080,plain,(
  ~p(s(bool,'T_0')) | $spl46),
  inference(unit_resulting_resolution,[],[f1007,f3746,f3746,f430])).
fof(f1007,plain,(
  ~sP7('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f995,f417])).
fof(f417,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f245])).
fof(f995,plain,(
  ~sP6('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f483,f269,f269,f424])).
fof(f16819,plain,(
  $false | ($spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14081,f3746])).
fof(f14081,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))) | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f1395,f3746,f3746,f430])).
fof(f1395,plain,(
  ~sP7('F_0','T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))) | $spl4),
  inference(unit_resulting_resolution,[],[f1383,f417])).
fof(f1383,plain,(
  ~sP6('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f269,f483,f1003])).
fof(f1003,plain,(
  ( ! [X2,X3] : (~sP6(X2,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),X3) | p(s(bool,X3)) | ~p(s(bool,X2))) ) | $spl4),
  inference(subsumption_resolution,[],[f997,f269])).
fof(f997,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | ~p(s(bool,X2)) | p(s(bool,X3)) | ~sP6(X2,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),X3)) ) | $spl4),
  inference(superposition,[],[f424,f531])).
fof(f16818,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f14082,f3746])).
fof(f14082,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)))) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5301,f3746,f3746,f430])).
fof(f5301,plain,(
  ~sP7('F_0','T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15))) | $spl0),
  inference(unit_resulting_resolution,[],[f269,f483,f477,f433])).
fof(f433,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP7(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f249])).
fof(f16817,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14083,f3746])).
fof(f14083,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9776,f3746,f3746,f430])).
fof(f9776,plain,(
  ~sP7('F_0','T_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK13))) | $spl3),
  inference(unit_resulting_resolution,[],[f9755,f417])).
fof(f9755,plain,(
  ~sP6('T_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),'F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f269,f483,f8013])).
fof(f8013,plain,(
  ( ! [X17,X18] : (~sP6(X17,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X18) | p(s(bool,X18)) | ~p(s(bool,X17))) ) | $spl3),
  inference(subsumption_resolution,[],[f7959,f269])).
fof(f7959,plain,(
  ( ! [X17,X18] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | p(s(bool,X18)) | ~sP6(X17,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X18)) ) | $spl3),
  inference(superposition,[],[f424,f4933])).
fof(f16816,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f14084,f3746])).
fof(f14084,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)))) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5302,f3746,f3746,f430])).
fof(f5302,plain,(
  ~sP7('F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15))) | $spl0),
  inference(unit_resulting_resolution,[],[f477,f483,f477,f433])).
fof(f16815,plain,(
  $false | ($spl0 | $spl22 | $spl46)),
  inference(subsumption_resolution,[],[f14085,f269])).
fof(f14085,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl22 | $spl46)),
  inference(unit_resulting_resolution,[],[f5415,f3746,f3746,f430])).
fof(f5415,plain,(
  ~sP7('F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),'T_0') | ($spl0 | $spl22)),
  inference(unit_resulting_resolution,[],[f483,f477,f2392])).
fof(f16814,plain,(
  $false | ($spl3 | $spl22 | $spl46)),
  inference(subsumption_resolution,[],[f14086,f269])).
fof(f14086,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl22 | $spl46)),
  inference(unit_resulting_resolution,[],[f5150,f3746,f3746,f430])).
fof(f5150,plain,(
  ~sP7('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'T_0','T_0') | ($spl3 | $spl22)),
  inference(unit_resulting_resolution,[],[f269,f480,f2392])).
fof(f16813,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14087,f3746])).
fof(f14087,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5120,f3746,f3746,f430])).
fof(f5120,plain,(
  ~sP7('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f269,f480,f1031])).
fof(f1031,plain,(
  ( ! [X2,X3] : (~sP7(X3,X2,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))) | p(s(bool,X3)) | ~p(s(bool,X2))) ) | $spl4),
  inference(subsumption_resolution,[],[f1025,f269])).
fof(f1025,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | ~p(s(bool,X2)) | p(s(bool,X3)) | ~sP7(X3,X2,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))) ) | $spl4),
  inference(superposition,[],[f433,f531])).
fof(f16812,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f14089,f3746])).
fof(f14089,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5262,f3746,f3746,f430])).
fof(f16811,plain,(
  $false | ($spl3 | $spl22 | $spl46)),
  inference(subsumption_resolution,[],[f14090,f269])).
fof(f14090,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl22 | $spl46)),
  inference(unit_resulting_resolution,[],[f5951,f3746,f3746,f430])).
fof(f5951,plain,(
  ~sP7('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),'T_0','T_0') | ($spl3 | $spl22)),
  inference(unit_resulting_resolution,[],[f269,f4934,f2392])).
fof(f16810,plain,(
  $false | ($spl3 | $spl4 | $spl22 | $spl46)),
  inference(subsumption_resolution,[],[f14091,f269])).
fof(f14091,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl4 | $spl22 | $spl46)),
  inference(unit_resulting_resolution,[],[f6990,f3746,f3746,f430])).
fof(f6990,plain,(
  ~sP7('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),'T_0','T_0') | ($spl3 | $spl4 | $spl22)),
  inference(unit_resulting_resolution,[],[f269,f5683,f2392])).
fof(f16809,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14092,f3746])).
fof(f14092,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5917,f3746,f3746,f430])).
fof(f5917,plain,(
  ~sP7('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),'T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f269,f4934,f1031])).
fof(f16808,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f14093,f269])).
fof(f14093,plain,(
  ~p(s(bool,'T_0')) | $spl46),
  inference(unit_resulting_resolution,[],[f636,f269,f3746,f430])).
fof(f16807,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14094,f269])).
fof(f14094,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f4999,f269,f3746,f430])).
fof(f16806,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14095,f269])).
fof(f14095,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5762,f269,f3746,f430])).
fof(f16805,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14096,f269])).
fof(f14096,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6752,f269,f3746,f430])).
fof(f16804,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14097,f269])).
fof(f14097,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9198,f269,f3746,f430])).
fof(f16803,plain,(
  $false | ($spl3 | $spl22 | $spl46)),
  inference(subsumption_resolution,[],[f14100,f3746])).
fof(f14100,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl22 | $spl46)),
  inference(unit_resulting_resolution,[],[f5150,f269,f3746,f430])).
fof(f16802,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f14101,f3746])).
fof(f14101,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)))) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5262,f269,f3746,f430])).
fof(f16801,plain,(
  $false | ($spl3 | $spl22 | $spl46)),
  inference(subsumption_resolution,[],[f14102,f3746])).
fof(f14102,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl22 | $spl46)),
  inference(unit_resulting_resolution,[],[f5951,f269,f3746,f430])).
fof(f16800,plain,(
  $false | ($spl3 | $spl4 | $spl22 | $spl46)),
  inference(subsumption_resolution,[],[f14103,f3746])).
fof(f14103,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl22 | $spl46)),
  inference(unit_resulting_resolution,[],[f6990,f269,f3746,f430])).
fof(f16799,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f14108,f269])).
fof(f14108,plain,(
  ~p(s(bool,'T_0')) | $spl46),
  inference(unit_resulting_resolution,[],[f634,f3746,f3746,f430])).
fof(f16798,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14109,f269])).
fof(f14109,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f4997,f3746,f3746,f430])).
fof(f16797,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14110,f269])).
fof(f14110,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5759,f3746,f3746,f430])).
fof(f16796,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14111,f269])).
fof(f14111,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6745,f3746,f3746,f430])).
fof(f16795,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14112,f269])).
fof(f14112,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9189,f3746,f3746,f430])).
fof(f16794,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f14113,f269])).
fof(f14113,plain,(
  ~p(s(bool,'T_0')) | $spl46),
  inference(unit_resulting_resolution,[],[f636,f3746,f3746,f430])).
fof(f16793,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14114,f269])).
fof(f14114,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f4999,f3746,f3746,f430])).
fof(f16792,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14115,f269])).
fof(f14115,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5762,f3746,f3746,f430])).
fof(f16791,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14116,f269])).
fof(f14116,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6752,f3746,f3746,f430])).
fof(f16790,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14117,f269])).
fof(f14117,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9198,f3746,f3746,f430])).
fof(f16789,plain,(
  $false | ($spl3 | $spl22 | $spl46)),
  inference(subsumption_resolution,[],[f14124,f3746])).
fof(f14124,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl22 | $spl46)),
  inference(unit_resulting_resolution,[],[f5150,f3746,f3746,f430])).
fof(f16788,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14125,f3746])).
fof(f14125,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5120,f3746,f3746,f430])).
fof(f16787,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f14126,f3746])).
fof(f14126,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)))) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5261,f3746,f3746,f430])).
fof(f16786,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f14127,f3746])).
fof(f14127,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)))) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5262,f3746,f3746,f430])).
fof(f16785,plain,(
  $false | ($spl3 | $spl22 | $spl46)),
  inference(subsumption_resolution,[],[f14128,f3746])).
fof(f14128,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl22 | $spl46)),
  inference(unit_resulting_resolution,[],[f5951,f3746,f3746,f430])).
fof(f16784,plain,(
  $false | ($spl3 | $spl4 | $spl22 | $spl46)),
  inference(subsumption_resolution,[],[f14129,f3746])).
fof(f14129,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl22 | $spl46)),
  inference(unit_resulting_resolution,[],[f6990,f3746,f3746,f430])).
fof(f16783,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14130,f3746])).
fof(f14130,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5917,f3746,f3746,f430])).
fof(f16782,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f14135,f3746])).
fof(f14135,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl46),
  inference(unit_resulting_resolution,[],[f634,f269,f3746,f430])).
fof(f16781,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14136,f3746])).
fof(f14136,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f4997,f269,f3746,f430])).
fof(f16780,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14137,f3746])).
fof(f14137,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5759,f269,f3746,f430])).
fof(f16779,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14138,f3746])).
fof(f14138,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6745,f269,f3746,f430])).
fof(f16778,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14139,f3746])).
fof(f14139,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9189,f269,f3746,f430])).
fof(f16777,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14141,f3746])).
fof(f14141,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f4999,f269,f3746,f430])).
fof(f16776,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14142,f3746])).
fof(f14142,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5762,f269,f3746,f430])).
fof(f16775,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14143,f3746])).
fof(f14143,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6752,f269,f3746,f430])).
fof(f16774,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14144,f3746])).
fof(f14144,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9198,f269,f3746,f430])).
fof(f16773,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f14145,f3746])).
fof(f14145,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl46),
  inference(unit_resulting_resolution,[],[f634,f3746,f3746,f430])).
fof(f16772,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14146,f3746])).
fof(f14146,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f4997,f3746,f3746,f430])).
fof(f16771,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14147,f3746])).
fof(f14147,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5759,f3746,f3746,f430])).
fof(f16770,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14148,f3746])).
fof(f14148,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6745,f3746,f3746,f430])).
fof(f16769,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14149,f3746])).
fof(f14149,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9189,f3746,f3746,f430])).
fof(f16768,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14151,f3746])).
fof(f14151,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f4999,f3746,f3746,f430])).
fof(f16767,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14152,f3746])).
fof(f14152,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5762,f3746,f3746,f430])).
fof(f16766,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14153,f3746])).
fof(f14153,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6752,f3746,f3746,f430])).
fof(f16765,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14154,f3746])).
fof(f14154,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9198,f3746,f3746,f430])).
fof(f16764,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f14155,f269])).
fof(f14155,plain,(
  ~p(s(bool,'T_0')) | $spl46),
  inference(unit_resulting_resolution,[],[f1007,f3746,f3746,f430])).
fof(f16763,plain,(
  $false | ($spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14156,f269])).
fof(f14156,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f1395,f3746,f3746,f430])).
fof(f16762,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f14157,f269])).
fof(f14157,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5301,f3746,f3746,f430])).
fof(f16761,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14158,f269])).
fof(f14158,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9776,f3746,f3746,f430])).
fof(f16760,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f14159,f3746])).
fof(f14159,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)))) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5302,f3746,f3746,f430])).
fof(f16759,plain,(
  $false | ($spl0 | $spl22 | $spl46)),
  inference(subsumption_resolution,[],[f14160,f3746])).
fof(f14160,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)))) | ($spl0 | $spl22 | $spl46)),
  inference(unit_resulting_resolution,[],[f5415,f3746,f3746,f430])).
fof(f16758,plain,(
  $false | ($spl3 | $spl22 | $spl46)),
  inference(subsumption_resolution,[],[f14161,f269])).
fof(f14161,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl22 | $spl46)),
  inference(unit_resulting_resolution,[],[f5150,f3746,f3746,f430])).
fof(f16757,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14162,f269])).
fof(f14162,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5120,f3746,f3746,f430])).
fof(f16756,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f14163,f3746])).
fof(f14163,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5261,f3746,f3746,f430])).
fof(f16755,plain,(
  $false | ($spl3 | $spl22 | $spl46)),
  inference(subsumption_resolution,[],[f14165,f269])).
fof(f14165,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl22 | $spl46)),
  inference(unit_resulting_resolution,[],[f5951,f3746,f3746,f430])).
fof(f16754,plain,(
  $false | ($spl3 | $spl4 | $spl22 | $spl46)),
  inference(subsumption_resolution,[],[f14166,f269])).
fof(f14166,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl4 | $spl22 | $spl46)),
  inference(unit_resulting_resolution,[],[f6990,f3746,f3746,f430])).
fof(f16753,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14167,f269])).
fof(f14167,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5917,f3746,f3746,f430])).
fof(f16752,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f14168,f269])).
fof(f14168,plain,(
  ~p(s(bool,'T_0')) | $spl46),
  inference(unit_resulting_resolution,[],[f634,f269,f3746,f430])).
fof(f16751,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14169,f269])).
fof(f14169,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f4997,f269,f3746,f430])).
fof(f16750,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14170,f269])).
fof(f14170,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5759,f269,f3746,f430])).
fof(f16749,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14171,f269])).
fof(f14171,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6745,f269,f3746,f430])).
fof(f16748,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14172,f269])).
fof(f14172,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9189,f269,f3746,f430])).
fof(f16747,plain,(
  $false | ($spl3 | $spl22 | $spl46)),
  inference(subsumption_resolution,[],[f14177,f3746])).
fof(f14177,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl22 | $spl46)),
  inference(unit_resulting_resolution,[],[f5150,f269,f3746,f430])).
fof(f16746,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14178,f3746])).
fof(f14178,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5120,f269,f3746,f430])).
fof(f16745,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f14179,f3746])).
fof(f14179,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)))) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5261,f269,f3746,f430])).
fof(f16744,plain,(
  $false | ($spl3 | $spl22 | $spl46)),
  inference(subsumption_resolution,[],[f14180,f3746])).
fof(f14180,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl22 | $spl46)),
  inference(unit_resulting_resolution,[],[f5951,f269,f3746,f430])).
fof(f16743,plain,(
  $false | ($spl3 | $spl4 | $spl22 | $spl46)),
  inference(subsumption_resolution,[],[f14181,f3746])).
fof(f14181,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl22 | $spl46)),
  inference(unit_resulting_resolution,[],[f6990,f269,f3746,f430])).
fof(f16742,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14182,f3746])).
fof(f14182,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5917,f269,f3746,f430])).
fof(f16741,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f14187,f269])).
fof(f14187,plain,(
  ~p(s(bool,'T_0')) | $spl46),
  inference(unit_resulting_resolution,[],[f634,f3746,f3746,f430])).
fof(f16740,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14188,f269])).
fof(f14188,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f4997,f3746,f3746,f430])).
fof(f16739,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14189,f269])).
fof(f14189,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5759,f3746,f3746,f430])).
fof(f16738,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14190,f269])).
fof(f14190,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6745,f3746,f3746,f430])).
fof(f16737,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14191,f269])).
fof(f14191,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9189,f3746,f3746,f430])).
fof(f16736,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f14192,f269])).
fof(f14192,plain,(
  ~p(s(bool,'T_0')) | $spl46),
  inference(unit_resulting_resolution,[],[f636,f3746,f3746,f430])).
fof(f16735,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14193,f269])).
fof(f14193,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f4999,f3746,f3746,f430])).
fof(f16734,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14194,f269])).
fof(f14194,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5762,f3746,f3746,f430])).
fof(f16733,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14195,f269])).
fof(f14195,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6752,f3746,f3746,f430])).
fof(f16732,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14196,f269])).
fof(f14196,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9198,f3746,f3746,f430])).
fof(f16731,plain,(
  $false | ($spl3 | $spl22 | $spl46)),
  inference(subsumption_resolution,[],[f14203,f3746])).
fof(f14203,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl22 | $spl46)),
  inference(unit_resulting_resolution,[],[f5150,f3746,f3746,f430])).
fof(f16730,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14204,f3746])).
fof(f14204,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5120,f3746,f3746,f430])).
fof(f16729,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f14205,f3746])).
fof(f14205,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)))) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5261,f3746,f3746,f430])).
fof(f16728,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f14206,f3746])).
fof(f14206,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)))) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5262,f3746,f3746,f430])).
fof(f16727,plain,(
  $false | ($spl3 | $spl22 | $spl46)),
  inference(subsumption_resolution,[],[f14207,f3746])).
fof(f14207,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl22 | $spl46)),
  inference(unit_resulting_resolution,[],[f5951,f3746,f3746,f430])).
fof(f16726,plain,(
  $false | ($spl3 | $spl4 | $spl22 | $spl46)),
  inference(subsumption_resolution,[],[f14208,f3746])).
fof(f14208,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl22 | $spl46)),
  inference(unit_resulting_resolution,[],[f6990,f3746,f3746,f430])).
fof(f16725,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14209,f3746])).
fof(f14209,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5917,f3746,f3746,f430])).
fof(f16724,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14215,f3746])).
fof(f14215,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f4999,f269,f3746,f430])).
fof(f16723,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14216,f3746])).
fof(f14216,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5762,f269,f3746,f430])).
fof(f16722,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14217,f3746])).
fof(f14217,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6752,f269,f3746,f430])).
fof(f16721,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14218,f3746])).
fof(f14218,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9198,f269,f3746,f430])).
fof(f16720,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f14219,f269])).
fof(f14219,plain,(
  ~p(s(bool,'T_0')) | $spl46),
  inference(unit_resulting_resolution,[],[f1007,f269,f3746,f430])).
fof(f16719,plain,(
  $false | ($spl0 | $spl22 | $spl46)),
  inference(subsumption_resolution,[],[f14220,f3746])).
fof(f14220,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)))) | ($spl0 | $spl22 | $spl46)),
  inference(unit_resulting_resolution,[],[f5415,f269,f3746,f430])).
fof(f16718,plain,(
  $false | ($spl3 | $spl22 | $spl46)),
  inference(subsumption_resolution,[],[f14221,f269])).
fof(f14221,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl22 | $spl46)),
  inference(unit_resulting_resolution,[],[f5150,f269,f3746,f430])).
fof(f16717,plain,(
  $false | ($spl3 | $spl22 | $spl46)),
  inference(subsumption_resolution,[],[f14223,f269])).
fof(f14223,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl22 | $spl46)),
  inference(unit_resulting_resolution,[],[f5951,f269,f3746,f430])).
fof(f16716,plain,(
  $false | ($spl3 | $spl4 | $spl22 | $spl46)),
  inference(subsumption_resolution,[],[f14224,f269])).
fof(f14224,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl4 | $spl22 | $spl46)),
  inference(unit_resulting_resolution,[],[f6990,f269,f3746,f430])).
fof(f16715,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f14225,f3746])).
fof(f14225,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl46),
  inference(unit_resulting_resolution,[],[f634,f3746,f3746,f430])).
fof(f16714,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14226,f3746])).
fof(f14226,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f4997,f3746,f3746,f430])).
fof(f16713,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14227,f3746])).
fof(f14227,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5759,f3746,f3746,f430])).
fof(f16712,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14228,f3746])).
fof(f14228,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6745,f3746,f3746,f430])).
fof(f16711,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14229,f3746])).
fof(f14229,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9189,f3746,f3746,f430])).
fof(f16710,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14231,f3746])).
fof(f14231,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f4999,f3746,f3746,f430])).
fof(f16709,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14232,f3746])).
fof(f14232,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5762,f3746,f3746,f430])).
fof(f16708,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14233,f3746])).
fof(f14233,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6752,f3746,f3746,f430])).
fof(f16707,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14234,f3746])).
fof(f14234,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9198,f3746,f3746,f430])).
fof(f16706,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f14235,f269])).
fof(f14235,plain,(
  ~p(s(bool,'T_0')) | $spl46),
  inference(unit_resulting_resolution,[],[f1007,f3746,f3746,f430])).
fof(f16705,plain,(
  $false | ($spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14236,f269])).
fof(f14236,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f1395,f3746,f3746,f430])).
fof(f16704,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f14237,f269])).
fof(f14237,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5301,f3746,f3746,f430])).
fof(f16703,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14238,f269])).
fof(f14238,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9776,f3746,f3746,f430])).
fof(f16702,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f14239,f3746])).
fof(f14239,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)))) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5302,f3746,f3746,f430])).
fof(f16701,plain,(
  $false | ($spl0 | $spl22 | $spl46)),
  inference(subsumption_resolution,[],[f14240,f3746])).
fof(f14240,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)))) | ($spl0 | $spl22 | $spl46)),
  inference(unit_resulting_resolution,[],[f5415,f3746,f3746,f430])).
fof(f16700,plain,(
  $false | ($spl3 | $spl22 | $spl46)),
  inference(subsumption_resolution,[],[f14241,f269])).
fof(f14241,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl22 | $spl46)),
  inference(unit_resulting_resolution,[],[f5150,f3746,f3746,f430])).
fof(f16699,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14242,f269])).
fof(f14242,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5120,f3746,f3746,f430])).
fof(f16698,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f14243,f3746])).
fof(f14243,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5261,f3746,f3746,f430])).
fof(f16697,plain,(
  $false | ($spl3 | $spl22 | $spl46)),
  inference(subsumption_resolution,[],[f14245,f269])).
fof(f14245,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl22 | $spl46)),
  inference(unit_resulting_resolution,[],[f5951,f3746,f3746,f430])).
fof(f16696,plain,(
  $false | ($spl3 | $spl4 | $spl22 | $spl46)),
  inference(subsumption_resolution,[],[f14246,f269])).
fof(f14246,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl4 | $spl22 | $spl46)),
  inference(unit_resulting_resolution,[],[f6990,f3746,f3746,f430])).
fof(f16695,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14247,f269])).
fof(f14247,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5917,f3746,f3746,f430])).
fof(f16694,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14249,f3746])).
fof(f14249,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f4997,f269,f3746,f430])).
fof(f16693,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14250,f3746])).
fof(f14250,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5759,f269,f3746,f430])).
fof(f16692,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14251,f3746])).
fof(f14251,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6745,f269,f3746,f430])).
fof(f16691,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14252,f3746])).
fof(f14252,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9189,f269,f3746,f430])).
fof(f16690,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f14253,f269])).
fof(f14253,plain,(
  ~p(s(bool,'T_0')) | $spl46),
  inference(unit_resulting_resolution,[],[f1007,f269,f3746,f430])).
fof(f16689,plain,(
  $false | ($spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14254,f3746])).
fof(f14254,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))) | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f1395,f269,f3746,f430])).
fof(f16688,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f14255,f3746])).
fof(f14255,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)))) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5301,f269,f3746,f430])).
fof(f16687,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14256,f3746])).
fof(f14256,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9776,f269,f3746,f430])).
fof(f16686,plain,(
  $false | ($spl3 | $spl22 | $spl46)),
  inference(subsumption_resolution,[],[f14257,f269])).
fof(f14257,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl22 | $spl46)),
  inference(unit_resulting_resolution,[],[f5150,f269,f3746,f430])).
fof(f16685,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14258,f3746])).
fof(f14258,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5120,f269,f3746,f430])).
fof(f16684,plain,(
  $false | ($spl3 | $spl22 | $spl46)),
  inference(subsumption_resolution,[],[f14260,f269])).
fof(f14260,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl22 | $spl46)),
  inference(unit_resulting_resolution,[],[f5951,f269,f3746,f430])).
fof(f16683,plain,(
  $false | ($spl3 | $spl4 | $spl22 | $spl46)),
  inference(subsumption_resolution,[],[f14261,f269])).
fof(f14261,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl4 | $spl22 | $spl46)),
  inference(unit_resulting_resolution,[],[f6990,f269,f3746,f430])).
fof(f16682,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14262,f3746])).
fof(f14262,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5917,f269,f3746,f430])).
fof(f16681,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14268,f3746])).
fof(f14268,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f4997,f3746,f3746,f430])).
fof(f16680,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14269,f3746])).
fof(f14269,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5759,f3746,f3746,f430])).
fof(f16679,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14270,f3746])).
fof(f14270,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6745,f3746,f3746,f430])).
fof(f16678,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14271,f3746])).
fof(f14271,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9189,f3746,f3746,f430])).
fof(f16677,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f14272,f3746])).
fof(f14272,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl46),
  inference(unit_resulting_resolution,[],[f636,f3746,f3746,f430])).
fof(f16676,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14273,f3746])).
fof(f14273,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f4999,f3746,f3746,f430])).
fof(f16675,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14274,f3746])).
fof(f14274,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5762,f3746,f3746,f430])).
fof(f16674,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14275,f3746])).
fof(f14275,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6752,f3746,f3746,f430])).
fof(f16673,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14276,f3746])).
fof(f14276,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9198,f3746,f3746,f430])).
fof(f16672,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f14277,f269])).
fof(f14277,plain,(
  ~p(s(bool,'T_0')) | $spl46),
  inference(unit_resulting_resolution,[],[f1007,f3746,f3746,f430])).
fof(f16671,plain,(
  $false | ($spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14278,f3746])).
fof(f14278,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))) | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f1395,f3746,f3746,f430])).
fof(f16670,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f14279,f3746])).
fof(f14279,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)))) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5301,f3746,f3746,f430])).
fof(f16669,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14280,f3746])).
fof(f14280,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9776,f3746,f3746,f430])).
fof(f16668,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f14281,f3746])).
fof(f14281,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)))) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5302,f3746,f3746,f430])).
fof(f16667,plain,(
  $false | ($spl0 | $spl22 | $spl46)),
  inference(subsumption_resolution,[],[f14282,f269])).
fof(f14282,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl22 | $spl46)),
  inference(unit_resulting_resolution,[],[f5415,f3746,f3746,f430])).
fof(f16666,plain,(
  $false | ($spl3 | $spl22 | $spl46)),
  inference(subsumption_resolution,[],[f14283,f269])).
fof(f14283,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl22 | $spl46)),
  inference(unit_resulting_resolution,[],[f5150,f3746,f3746,f430])).
fof(f16665,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14284,f3746])).
fof(f14284,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5120,f3746,f3746,f430])).
fof(f16664,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f14286,f3746])).
fof(f14286,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5262,f3746,f3746,f430])).
fof(f16663,plain,(
  $false | ($spl3 | $spl22 | $spl46)),
  inference(subsumption_resolution,[],[f14287,f269])).
fof(f14287,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl22 | $spl46)),
  inference(unit_resulting_resolution,[],[f5951,f3746,f3746,f430])).
fof(f16662,plain,(
  $false | ($spl3 | $spl4 | $spl22 | $spl46)),
  inference(subsumption_resolution,[],[f14288,f269])).
fof(f14288,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl4 | $spl22 | $spl46)),
  inference(unit_resulting_resolution,[],[f6990,f3746,f3746,f430])).
fof(f16661,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14289,f3746])).
fof(f14289,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5917,f3746,f3746,f430])).
fof(f16660,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14295,f3746])).
fof(f14295,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5006,f483,f3746,f433])).
fof(f5006,plain,(
  ( ! [X0] : (sP7('F_0',X0,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f483,f480,f428])).
fof(f428,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f249])).
fof(f16659,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14296,f3746])).
fof(f14296,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5772,f483,f3746,f433])).
fof(f5772,plain,(
  ( ! [X0] : (sP7('F_0',X0,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f483,f4934,f428])).
fof(f16658,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14297,f3746])).
fof(f14297,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6768,f483,f3746,f433])).
fof(f6768,plain,(
  ( ! [X0] : (sP7('F_0',X0,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f483,f5683,f428])).
fof(f16657,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14298,f3746])).
fof(f14298,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9218,f483,f3746,f433])).
fof(f9218,plain,(
  ( ! [X0] : (sP7('F_0',X0,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f483,f5685,f428])).
fof(f16656,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f14299,f3746])).
fof(f14299,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl46),
  inference(unit_resulting_resolution,[],[f648,f483,f3746,f433])).
fof(f648,plain,(
  ( ! [X0] : (sP7('F_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f483,f483,f429])).
fof(f429,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f249])).
fof(f16655,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14300,f3746])).
fof(f14300,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5011,f483,f3746,f433])).
fof(f5011,plain,(
  ( ! [X0] : (sP7('F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),X0)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f483,f480,f429])).
fof(f16654,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14301,f3746])).
fof(f14301,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5778,f483,f3746,f433])).
fof(f5778,plain,(
  ( ! [X0] : (sP7('F_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),X0)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f483,f4934,f429])).
fof(f16653,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14302,f3746])).
fof(f14302,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6776,f483,f3746,f433])).
fof(f6776,plain,(
  ( ! [X0] : (sP7('F_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),X0)) ) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f483,f5683,f429])).
fof(f16652,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14303,f3746])).
fof(f14303,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9227,f483,f3746,f433])).
fof(f9227,plain,(
  ( ! [X0] : (sP7('F_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))),X0)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f483,f5685,f429])).
fof(f16651,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f14358,f3746])).
fof(f14358,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl46),
  inference(unit_resulting_resolution,[],[f641,f483,f3746,f433])).
fof(f641,plain,(
  ( ! [X0] : (sP7('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f483,f483,f428])).
fof(f16650,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14359,f3746])).
fof(f14359,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5006,f483,f3746,f433])).
fof(f16649,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14360,f3746])).
fof(f14360,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5772,f483,f3746,f433])).
fof(f16648,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14361,f3746])).
fof(f14361,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6768,f483,f3746,f433])).
fof(f16647,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14362,f3746])).
fof(f14362,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9218,f483,f3746,f433])).
fof(f16646,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14364,f3746])).
fof(f14364,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5011,f483,f3746,f433])).
fof(f16645,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14365,f3746])).
fof(f14365,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5778,f483,f3746,f433])).
fof(f16644,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14366,f3746])).
fof(f14366,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6776,f483,f3746,f433])).
fof(f16643,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14367,f3746])).
fof(f14367,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9227,f483,f3746,f433])).
fof(f16642,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f14466,f15804])).
fof(f15804,plain,(
  ( ! [X0,X1] : (sP9('F_0',X0,X1)) ) | $spl46),
  inference(unit_resulting_resolution,[],[f521,f3746,f469])).
fof(f469,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f437])).
fof(f437,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f252])).
fof(f252,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP8(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f251])).
fof(f251,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f250])).
fof(f250,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1)) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f156])).
fof(f156,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f521,plain,(
  ( ! [X0,X1] : (sP8('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f483,f444])).
fof(f444,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f255])).
fof(f255,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f254])).
fof(f254,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f253])).
fof(f253,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f155])).
fof(f155,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f14466,plain,(
  ( ! [X0,X1] : (~sP9('F_0',X0,X1)) ) | $spl46),
  inference(unit_resulting_resolution,[],[f483,f3746,f434])).
fof(f434,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f252])).
fof(f16641,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f14512,f15804])).
fof(f14512,plain,(
  ( ! [X0,X1] : (~sP9('F_0',X0,X1)) ) | $spl46),
  inference(unit_resulting_resolution,[],[f483,f3746,f435])).
fof(f435,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f252])).
fof(f16640,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f14562,f14668])).
fof(f14668,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2)) ) | $spl46),
  inference(unit_resulting_resolution,[],[f3746,f443])).
fof(f443,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f255])).
fof(f14562,plain,(
  ~sP8('F_0','T_0','F_0') | $spl46),
  inference(unit_resulting_resolution,[],[f657,f483,f3746,f438])).
fof(f438,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f252])).
fof(f657,plain,(
  ( ! [X0] : (~sP9('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f269,f483,f435])).
fof(f16639,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f14564,f14668])).
fof(f14564,plain,(
  ~sP8('F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),'F_0') | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5309,f483,f3746,f438])).
fof(f5309,plain,(
  ( ! [X0] : (~sP9('F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f483,f477,f435])).
fof(f16638,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14566,f14668])).
fof(f14566,plain,(
  ~sP8('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'T_0','F_0') | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5018,f483,f3746,f438])).
fof(f5018,plain,(
  ( ! [X0] : (~sP9('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'T_0',X0)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f269,f480,f435])).
fof(f16637,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14572,f14668])).
fof(f14572,plain,(
  ~sP8('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))),'T_0','F_0') | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9238,f483,f3746,f438])).
fof(f9238,plain,(
  ( ! [X0] : (~sP9('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))),'T_0',X0)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f269,f5685,f435])).
fof(f16636,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14575,f14668])).
fof(f14575,plain,(
  ~sP8('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),'T_0','F_0') | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6786,f483,f3746,f438])).
fof(f6786,plain,(
  ( ! [X0] : (~sP9('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),'T_0',X0)) ) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f269,f5683,f435])).
fof(f16635,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14577,f14668])).
fof(f14577,plain,(
  ~sP8('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),'T_0','F_0') | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5786,f483,f3746,f438])).
fof(f5786,plain,(
  ( ! [X0] : (~sP9('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),'T_0',X0)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f269,f4934,f435])).
fof(f16634,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f14622,f14668])).
fof(f14622,plain,(
  ~sP8('F_0','F_0','T_0') | $spl46),
  inference(unit_resulting_resolution,[],[f654,f483,f3746,f439])).
fof(f439,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f252])).
fof(f654,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f269,f483,f434])).
fof(f16633,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f14623,f14668])).
fof(f14623,plain,(
  ~sP8('F_0','F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15))) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5304,f483,f3746,f439])).
fof(f5304,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f483,f477,f434])).
fof(f16632,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14628,f14668])).
fof(f14628,plain,(
  ~sP8('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'F_0','T_0') | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5017,f483,f3746,f439])).
fof(f5017,plain,(
  ( ! [X0] : (~sP9('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),X0,'T_0')) ) | $spl3),
  inference(unit_resulting_resolution,[],[f269,f480,f434])).
fof(f16631,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14629,f14668])).
fof(f14629,plain,(
  ~sP8('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))),'F_0','T_0') | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9237,f483,f3746,f439])).
fof(f9237,plain,(
  ( ! [X0] : (~sP9('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))),X0,'T_0')) ) | $spl3),
  inference(unit_resulting_resolution,[],[f269,f5685,f434])).
fof(f16630,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14631,f14668])).
fof(f14631,plain,(
  ~sP8('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),'F_0','T_0') | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6785,f483,f3746,f439])).
fof(f6785,plain,(
  ( ! [X0] : (~sP9('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),X0,'T_0')) ) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f269,f5683,f434])).
fof(f16629,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14632,f14668])).
fof(f14632,plain,(
  ~sP8('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),'F_0','T_0') | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5785,f483,f3746,f439])).
fof(f5785,plain,(
  ( ! [X0] : (~sP9('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),X0,'T_0')) ) | $spl3),
  inference(unit_resulting_resolution,[],[f269,f4934,f434])).
fof(f16628,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f14664,f14668])).
fof(f14664,plain,(
  ( ! [X0] : (~sP8(X0,'F_0','F_0')) ) | $spl46),
  inference(unit_resulting_resolution,[],[f483,f483,f3746,f441])).
fof(f441,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f255])).
fof(f16627,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f14670,f269])).
fof(f14670,plain,(
  ~p(s(bool,'T_0')) | $spl46),
  inference(unit_resulting_resolution,[],[f1085,f3746,f446])).
fof(f446,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f257])).
fof(f257,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f256])).
fof(f256,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f157])).
fof(f157,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP9(X2,X0,X1))),
  inference(definition_folding,[],[f115,f156,f155])).
fof(f115,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f114])).
fof(f114,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X20,X21,X22] : ((p(s(bool,X22)) <=> (p(s(bool,X21)) | p(s(bool,X20)))) <=> ((p(s(bool,X22)) | ~p(s(bool,X21))) & (p(s(bool,X22)) | ~p(s(bool,X20))) & (p(s(bool,X21)) | p(s(bool,X20)) | ~p(s(bool,X22)))))),
  file('thm.pred_set.UNION_SUBSET','thm.sat.sat_25_prove_ltb')).
fof(f1085,plain,(
  ~sP9('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1071,f436])).
fof(f436,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f252])).
fof(f1071,plain,(
  ~sP8('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f483,f483,f269,f441])).
fof(f16626,plain,(
  $false | ($spl3 | $spl30 | $spl46)),
  inference(subsumption_resolution,[],[f14671,f269])).
fof(f14671,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl30 | $spl46)),
  inference(unit_resulting_resolution,[],[f5159,f3746,f446])).
fof(f5159,plain,(
  ~sP9('T_0','F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | ($spl3 | $spl30)),
  inference(unit_resulting_resolution,[],[f483,f480,f2404])).
fof(f16625,plain,(
  $false | ($spl3 | $spl30 | $spl46)),
  inference(subsumption_resolution,[],[f14672,f269])).
fof(f14672,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl30 | $spl46)),
  inference(unit_resulting_resolution,[],[f5960,f3746,f446])).
fof(f5960,plain,(
  ~sP9('T_0','F_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | ($spl3 | $spl30)),
  inference(unit_resulting_resolution,[],[f483,f4934,f2404])).
fof(f16624,plain,(
  $false | ($spl3 | $spl4 | $spl30 | $spl46)),
  inference(subsumption_resolution,[],[f14673,f269])).
fof(f14673,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl4 | $spl30 | $spl46)),
  inference(unit_resulting_resolution,[],[f6999,f3746,f446])).
fof(f6999,plain,(
  ~sP9('T_0','F_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14))) | ($spl3 | $spl4 | $spl30)),
  inference(unit_resulting_resolution,[],[f483,f5683,f2404])).
fof(f16623,plain,(
  $false | ($spl3 | $spl30 | $spl46)),
  inference(subsumption_resolution,[],[f14674,f269])).
fof(f14674,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl30 | $spl46)),
  inference(unit_resulting_resolution,[],[f5163,f3746,f446])).
fof(f5163,plain,(
  ~sP9('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'F_0') | ($spl3 | $spl30)),
  inference(unit_resulting_resolution,[],[f483,f480,f2404])).
fof(f16622,plain,(
  $false | ($spl3 | $spl30 | $spl46)),
  inference(subsumption_resolution,[],[f14675,f269])).
fof(f14675,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl30 | $spl46)),
  inference(unit_resulting_resolution,[],[f5164,f3746,f446])).
fof(f5164,plain,(
  ~sP9('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | ($spl3 | $spl30)),
  inference(unit_resulting_resolution,[],[f480,f480,f2404])).
fof(f16621,plain,(
  $false | ($spl3 | $spl30 | $spl46)),
  inference(subsumption_resolution,[],[f14676,f269])).
fof(f14676,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl30 | $spl46)),
  inference(unit_resulting_resolution,[],[f5964,f3746,f446])).
fof(f5964,plain,(
  ~sP9('T_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),'F_0') | ($spl3 | $spl30)),
  inference(unit_resulting_resolution,[],[f483,f4934,f2404])).
fof(f16620,plain,(
  $false | ($spl3 | $spl4 | $spl30 | $spl46)),
  inference(subsumption_resolution,[],[f14677,f269])).
fof(f14677,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl4 | $spl30 | $spl46)),
  inference(unit_resulting_resolution,[],[f7003,f3746,f446])).
fof(f7003,plain,(
  ~sP9('T_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),'F_0') | ($spl3 | $spl4 | $spl30)),
  inference(unit_resulting_resolution,[],[f483,f5683,f2404])).
fof(f16619,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14682,f3746])).
fof(f14682,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5017,f3746,f446])).
fof(f16618,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14683,f3746])).
fof(f14683,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5018,f3746,f446])).
fof(f16617,plain,(
  $false | ($spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14684,f3746])).
fof(f14684,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))) | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f1454,f3746,f446])).
fof(f1454,plain,(
  ~sP9('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),'F_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f1438,f436])).
fof(f1438,plain,(
  ~sP8('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),'F_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f483,f483,f1081])).
fof(f1081,plain,(
  ( ! [X2,X3] : (~sP8('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),X2,X3) | p(s(bool,X3)) | p(s(bool,X2))) ) | $spl4),
  inference(subsumption_resolution,[],[f1073,f269])).
fof(f1073,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,X2)) | p(s(bool,X3)) | ~sP8('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),X2,X3)) ) | $spl4),
  inference(superposition,[],[f441,f531])).
fof(f16616,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f14685,f3746])).
fof(f14685,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)))) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5335,f3746,f446])).
fof(f5335,plain,(
  ~sP9('const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),'F_0','F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f483,f483,f477,f448])).
fof(f448,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)) | ~sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f257])).
fof(f16615,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14686,f3746])).
fof(f14686,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5129,f3746,f446])).
fof(f5129,plain,(
  ~sP9('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),'F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f483,f480,f1122])).
fof(f1122,plain,(
  ( ! [X2,X3] : (~sP9('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),X3,X2) | p(s(bool,X3)) | p(s(bool,X2))) ) | $spl4),
  inference(subsumption_resolution,[],[f1114,f269])).
fof(f1114,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,X2)) | p(s(bool,X3)) | ~sP9('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),X3,X2)) ) | $spl4),
  inference(superposition,[],[f448,f531])).
fof(f16614,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14687,f3746])).
fof(f14687,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5926,f3746,f446])).
fof(f5926,plain,(
  ~sP9('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),'F_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f483,f4934,f1122])).
fof(f16613,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14688,f3746])).
fof(f14688,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5133,f3746,f446])).
fof(f5133,plain,(
  ~sP9('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'F_0') | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f483,f480,f1122])).
fof(f16612,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14689,f3746])).
fof(f14689,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5930,f3746,f446])).
fof(f5930,plain,(
  ~sP9('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),'F_0') | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f483,f4934,f1122])).
fof(f16611,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14690,f3746])).
fof(f14690,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9237,f3746,f446])).
fof(f16610,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14691,f3746])).
fof(f14691,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9238,f3746,f446])).
fof(f16609,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14692,f3746])).
fof(f14692,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9813,f3746,f446])).
fof(f9813,plain,(
  ~sP9('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),'F_0','F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f9795,f436])).
fof(f9795,plain,(
  ~sP8('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),'F_0','F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f483,f483,f8015])).
fof(f8015,plain,(
  ( ! [X21,X22] : (~sP8('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X21,X22) | p(s(bool,X22)) | p(s(bool,X21))) ) | $spl3),
  inference(subsumption_resolution,[],[f7961,f269])).
fof(f7961,plain,(
  ( ! [X21,X22] : (~p(s(bool,'T_0')) | p(s(bool,X21)) | p(s(bool,X22)) | ~sP8('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X21,X22)) ) | $spl3),
  inference(superposition,[],[f441,f4933])).
fof(f16608,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14693,f3746])).
fof(f14693,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6785,f3746,f446])).
fof(f16607,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14694,f3746])).
fof(f14694,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6786,f3746,f446])).
fof(f16606,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14695,f3746])).
fof(f14695,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5785,f3746,f446])).
fof(f16605,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14696,f3746])).
fof(f14696,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5786,f3746,f446])).
fof(f16604,plain,(
  $false | ($spl3 | $spl30 | $spl46)),
  inference(subsumption_resolution,[],[f14702,f3746])).
fof(f14702,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl30 | $spl46)),
  inference(unit_resulting_resolution,[],[f5159,f3746,f446])).
fof(f16603,plain,(
  $false | ($spl3 | $spl30 | $spl46)),
  inference(subsumption_resolution,[],[f14703,f3746])).
fof(f14703,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl30 | $spl46)),
  inference(unit_resulting_resolution,[],[f5960,f3746,f446])).
fof(f16602,plain,(
  $false | ($spl3 | $spl4 | $spl30 | $spl46)),
  inference(subsumption_resolution,[],[f14704,f3746])).
fof(f14704,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl30 | $spl46)),
  inference(unit_resulting_resolution,[],[f6999,f3746,f446])).
fof(f16601,plain,(
  $false | ($spl3 | $spl30 | $spl46)),
  inference(subsumption_resolution,[],[f14706,f3746])).
fof(f14706,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl30 | $spl46)),
  inference(unit_resulting_resolution,[],[f5164,f3746,f446])).
fof(f16600,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f14709,f269])).
fof(f14709,plain,(
  ~p(s(bool,'T_0')) | $spl46),
  inference(unit_resulting_resolution,[],[f654,f3746,f446])).
fof(f16599,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f14710,f3746])).
fof(f14710,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)))) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5304,f3746,f446])).
fof(f16598,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f14711,f3746])).
fof(f14711,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl46),
  inference(unit_resulting_resolution,[],[f657,f3746,f446])).
fof(f16597,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f14712,f3746])).
fof(f14712,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5309,f3746,f446])).
fof(f16596,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14713,f269])).
fof(f14713,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5017,f3746,f446])).
fof(f16595,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14714,f3746])).
fof(f14714,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5018,f3746,f446])).
fof(f16594,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14717,f3746])).
fof(f14717,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5129,f3746,f446])).
fof(f16593,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14718,f3746])).
fof(f14718,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5926,f3746,f446])).
fof(f16592,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14721,f269])).
fof(f14721,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9237,f3746,f446])).
fof(f16591,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14722,f3746])).
fof(f14722,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9238,f3746,f446])).
fof(f16590,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14724,f269])).
fof(f14724,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6785,f3746,f446])).
fof(f16589,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14725,f3746])).
fof(f14725,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6786,f3746,f446])).
fof(f16588,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14726,f269])).
fof(f14726,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5785,f3746,f446])).
fof(f16587,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14727,f3746])).
fof(f14727,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5786,f3746,f446])).
fof(f16586,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f14732,f269])).
fof(f14732,plain,(
  ~p(s(bool,'T_0')) | $spl46),
  inference(unit_resulting_resolution,[],[f1085,f3746,f447])).
fof(f447,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f257])).
fof(f16585,plain,(
  $false | ($spl3 | $spl30 | $spl46)),
  inference(subsumption_resolution,[],[f14733,f269])).
fof(f14733,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl30 | $spl46)),
  inference(unit_resulting_resolution,[],[f5159,f3746,f447])).
fof(f16584,plain,(
  $false | ($spl3 | $spl30 | $spl46)),
  inference(subsumption_resolution,[],[f14734,f269])).
fof(f14734,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl30 | $spl46)),
  inference(unit_resulting_resolution,[],[f5960,f3746,f447])).
fof(f16583,plain,(
  $false | ($spl3 | $spl4 | $spl30 | $spl46)),
  inference(subsumption_resolution,[],[f14735,f269])).
fof(f14735,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl4 | $spl30 | $spl46)),
  inference(unit_resulting_resolution,[],[f6999,f3746,f447])).
fof(f16582,plain,(
  $false | ($spl3 | $spl30 | $spl46)),
  inference(subsumption_resolution,[],[f14736,f269])).
fof(f14736,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl30 | $spl46)),
  inference(unit_resulting_resolution,[],[f5163,f3746,f447])).
fof(f16581,plain,(
  $false | ($spl3 | $spl30 | $spl46)),
  inference(subsumption_resolution,[],[f14737,f269])).
fof(f14737,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl30 | $spl46)),
  inference(unit_resulting_resolution,[],[f5164,f3746,f447])).
fof(f16580,plain,(
  $false | ($spl3 | $spl30 | $spl46)),
  inference(subsumption_resolution,[],[f14738,f269])).
fof(f14738,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl30 | $spl46)),
  inference(unit_resulting_resolution,[],[f5964,f3746,f447])).
fof(f16579,plain,(
  $false | ($spl3 | $spl4 | $spl30 | $spl46)),
  inference(subsumption_resolution,[],[f14739,f269])).
fof(f14739,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl4 | $spl30 | $spl46)),
  inference(unit_resulting_resolution,[],[f7003,f3746,f447])).
fof(f16578,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14744,f3746])).
fof(f14744,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5017,f3746,f447])).
fof(f16577,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14745,f3746])).
fof(f14745,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5018,f3746,f447])).
fof(f16576,plain,(
  $false | ($spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14746,f3746])).
fof(f14746,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))) | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f1454,f3746,f447])).
fof(f16575,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f14747,f3746])).
fof(f14747,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)))) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5335,f3746,f447])).
fof(f16574,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14748,f3746])).
fof(f14748,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5129,f3746,f447])).
fof(f16573,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14749,f3746])).
fof(f14749,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5926,f3746,f447])).
fof(f16572,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14750,f3746])).
fof(f14750,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5133,f3746,f447])).
fof(f16571,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14751,f3746])).
fof(f14751,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5930,f3746,f447])).
fof(f16570,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14752,f3746])).
fof(f14752,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9237,f3746,f447])).
fof(f16569,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14753,f3746])).
fof(f14753,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9238,f3746,f447])).
fof(f16568,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14754,f3746])).
fof(f14754,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9813,f3746,f447])).
fof(f16567,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14755,f3746])).
fof(f14755,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6785,f3746,f447])).
fof(f16566,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14756,f3746])).
fof(f14756,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6786,f3746,f447])).
fof(f16565,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14757,f3746])).
fof(f14757,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5785,f3746,f447])).
fof(f16564,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14758,f3746])).
fof(f14758,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5786,f3746,f447])).
fof(f16563,plain,(
  $false | ($spl3 | $spl30 | $spl46)),
  inference(subsumption_resolution,[],[f14767,f3746])).
fof(f14767,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl30 | $spl46)),
  inference(unit_resulting_resolution,[],[f5163,f3746,f447])).
fof(f16562,plain,(
  $false | ($spl3 | $spl30 | $spl46)),
  inference(subsumption_resolution,[],[f14768,f3746])).
fof(f14768,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl30 | $spl46)),
  inference(unit_resulting_resolution,[],[f5164,f3746,f447])).
fof(f16561,plain,(
  $false | ($spl3 | $spl30 | $spl46)),
  inference(subsumption_resolution,[],[f14769,f3746])).
fof(f14769,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl30 | $spl46)),
  inference(unit_resulting_resolution,[],[f5964,f3746,f447])).
fof(f16560,plain,(
  $false | ($spl3 | $spl4 | $spl30 | $spl46)),
  inference(subsumption_resolution,[],[f14770,f3746])).
fof(f14770,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl30 | $spl46)),
  inference(unit_resulting_resolution,[],[f7003,f3746,f447])).
fof(f16559,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f14771,f3746])).
fof(f14771,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl46),
  inference(unit_resulting_resolution,[],[f654,f3746,f447])).
fof(f16558,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f14772,f3746])).
fof(f14772,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5304,f3746,f447])).
fof(f16557,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f14773,f269])).
fof(f14773,plain,(
  ~p(s(bool,'T_0')) | $spl46),
  inference(unit_resulting_resolution,[],[f657,f3746,f447])).
fof(f16556,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f14774,f3746])).
fof(f14774,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)))) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5309,f3746,f447])).
fof(f16555,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14775,f3746])).
fof(f14775,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5017,f3746,f447])).
fof(f16554,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14776,f269])).
fof(f14776,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5018,f3746,f447])).
fof(f16553,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14781,f3746])).
fof(f14781,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5133,f3746,f447])).
fof(f16552,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14782,f3746])).
fof(f14782,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5930,f3746,f447])).
fof(f16551,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14783,f3746])).
fof(f14783,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9237,f3746,f447])).
fof(f16550,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14784,f269])).
fof(f14784,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9238,f3746,f447])).
fof(f16549,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14786,f3746])).
fof(f14786,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6785,f3746,f447])).
fof(f16548,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14787,f269])).
fof(f14787,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6786,f3746,f447])).
fof(f16547,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14788,f3746])).
fof(f14788,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5785,f3746,f447])).
fof(f16546,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14789,f269])).
fof(f14789,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5786,f3746,f447])).
fof(f16545,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f14872,f15755])).
fof(f15755,plain,(
  ( ! [X0,X1] : (sP11('F_0',X0,X1)) ) | $spl46),
  inference(unit_resulting_resolution,[],[f526,f3746,f468])).
fof(f468,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP10(X0,X1,X2) | sP11(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f454])).
fof(f454,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f260])).
fof(f260,plain,(
  ! [X0,X1,X2] : ((~sP11(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP10(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP10(X0,X1,X2) | sP11(X0,X1,X2)))),
  inference(rectify,[],[f259])).
fof(f259,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1) | sP11(X2,X0,X1)))),
  inference(flattening,[],[f258])).
fof(f258,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1)) | sP11(X2,X0,X1)))),
  inference(nnf_transformation,[],[f159])).
fof(f159,plain,(
  ! [X2,X0,X1] : (sP11(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f526,plain,(
  ( ! [X0,X1] : (sP10('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f483,f461])).
fof(f461,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f263])).
fof(f263,plain,(
  ! [X0,X1,X2] : ((~sP10(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP10(X0,X1,X2)))),
  inference(rectify,[],[f262])).
fof(f262,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(flattening,[],[f261])).
fof(f261,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(nnf_transformation,[],[f158])).
fof(f158,plain,(
  ! [X2,X0,X1] : (sP10(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f14872,plain,(
  ( ! [X0,X1] : (~sP11('F_0',X0,X1)) ) | $spl46),
  inference(unit_resulting_resolution,[],[f483,f3746,f452])).
fof(f452,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f260])).
fof(f16544,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14875,f3746])).
fof(f14875,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f526,f5070,f3746,f455])).
fof(f455,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f260])).
fof(f5070,plain,(
  ( ! [X0] : (~sP11('F_0',X0,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f483,f480,f451])).
fof(f451,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f260])).
fof(f16543,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14876,f3746])).
fof(f14876,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f526,f5850,f3746,f455])).
fof(f5850,plain,(
  ( ! [X0] : (~sP11('F_0',X0,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f483,f4934,f451])).
fof(f16542,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14877,f3746])).
fof(f14877,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f526,f6862,f3746,f455])).
fof(f6862,plain,(
  ( ! [X0] : (~sP11('F_0',X0,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f483,f5683,f451])).
fof(f16541,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14878,f3746])).
fof(f14878,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f526,f9338,f3746,f455])).
fof(f9338,plain,(
  ( ! [X0] : (~sP11('F_0',X0,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f483,f5685,f451])).
fof(f16540,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f14879,f3746])).
fof(f14879,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl46),
  inference(unit_resulting_resolution,[],[f526,f688,f3746,f455])).
fof(f688,plain,(
  ( ! [X0] : (~sP11('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f269,f483,f452])).
fof(f16539,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f14880,f3746])).
fof(f14880,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl46),
  inference(unit_resulting_resolution,[],[f523,f688,f3746,f455])).
fof(f523,plain,(
  ( ! [X0,X1] : (sP10(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f269,f460])).
fof(f460,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f263])).
fof(f16538,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14882,f3746])).
fof(f14882,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5086,f688,f3746,f455])).
fof(f5086,plain,(
  ( ! [X0,X1] : (sP10(X0,X1,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f480,f459])).
fof(f459,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f263])).
fof(f16537,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14883,f3746])).
fof(f14883,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5876,f688,f3746,f455])).
fof(f5876,plain,(
  ( ! [X0,X1] : (sP10(X0,X1,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f4934,f459])).
fof(f16536,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14884,f3746])).
fof(f14884,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6900,f688,f3746,f455])).
fof(f6900,plain,(
  ( ! [X0,X1] : (sP10(X0,X1,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f5683,f459])).
fof(f16535,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14885,f3746])).
fof(f14885,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9390,f688,f3746,f455])).
fof(f9390,plain,(
  ( ! [X0,X1] : (sP10(X0,X1,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f5685,f459])).
fof(f16534,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f14886,f3746])).
fof(f14886,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f526,f5336,f3746,f455])).
fof(f5336,plain,(
  ( ! [X0] : (~sP11('F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f483,f477,f452])).
fof(f16533,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f14887,f3746])).
fof(f14887,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5363,f5336,f3746,f455])).
fof(f5363,plain,(
  ( ! [X0,X1] : (sP10(X0,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),X1)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f477,f460])).
fof(f16532,plain,(
  $false | ($spl0 | $spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14889,f3746])).
fof(f14889,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl0 | $spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5086,f5336,f3746,f455])).
fof(f16531,plain,(
  $false | ($spl0 | $spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14890,f3746])).
fof(f14890,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl0 | $spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5876,f5336,f3746,f455])).
fof(f16530,plain,(
  $false | ($spl0 | $spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14891,f3746])).
fof(f14891,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl0 | $spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6900,f5336,f3746,f455])).
fof(f16529,plain,(
  $false | ($spl0 | $spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14892,f3746])).
fof(f14892,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | ($spl0 | $spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9390,f5336,f3746,f455])).
fof(f16528,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14894,f3746])).
fof(f14894,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5087,f5071,f3746,f455])).
fof(f5071,plain,(
  ( ! [X0] : (~sP11('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),X0,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f480,f480,f451])).
fof(f5087,plain,(
  ( ! [X0,X1] : (sP10('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),X0,X1)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f480,f461])).
fof(f16527,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14895,f3746])).
fof(f14895,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5087,f5076,f3746,f455])).
fof(f5076,plain,(
  ( ! [X0] : (~sP11('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'T_0',X0)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f269,f480,f452])).
fof(f16526,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14896,f3746])).
fof(f14896,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f523,f5076,f3746,f455])).
fof(f16525,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14898,f3746])).
fof(f14898,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5086,f5076,f3746,f455])).
fof(f16524,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14899,f3746])).
fof(f14899,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5876,f5076,f3746,f455])).
fof(f16523,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14900,f3746])).
fof(f14900,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6900,f5076,f3746,f455])).
fof(f16522,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14901,f3746])).
fof(f14901,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9390,f5076,f3746,f455])).
fof(f16521,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14903,f3746])).
fof(f14903,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9391,f9348,f3746,f455])).
fof(f9348,plain,(
  ( ! [X0] : (~sP11('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))),'T_0',X0)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f269,f5685,f452])).
fof(f9391,plain,(
  ( ! [X0,X1] : (sP10('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))),X0,X1)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f5685,f461])).
fof(f16520,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14904,f3746])).
fof(f14904,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f523,f9348,f3746,f455])).
fof(f16519,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14906,f3746])).
fof(f14906,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5086,f9348,f3746,f455])).
fof(f16518,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14907,f3746])).
fof(f14907,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5876,f9348,f3746,f455])).
fof(f16517,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14908,f3746])).
fof(f14908,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6900,f9348,f3746,f455])).
fof(f16516,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14909,f3746])).
fof(f14909,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9390,f9348,f3746,f455])).
fof(f16515,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14911,f3746])).
fof(f14911,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6901,f6868,f3746,f455])).
fof(f6868,plain,(
  ( ! [X0] : (~sP11('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),'T_0',X0)) ) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f269,f5683,f452])).
fof(f6901,plain,(
  ( ! [X0,X1] : (sP10('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),X0,X1)) ) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f5683,f461])).
fof(f16514,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14912,f3746])).
fof(f14912,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f523,f6868,f3746,f455])).
fof(f16513,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14914,f3746])).
fof(f14914,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5086,f6868,f3746,f455])).
fof(f16512,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14915,f3746])).
fof(f14915,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5876,f6868,f3746,f455])).
fof(f16511,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14916,f3746])).
fof(f14916,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6900,f6868,f3746,f455])).
fof(f16510,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14917,f3746])).
fof(f14917,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f9390,f6868,f3746,f455])).
fof(f16509,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14919,f3746])).
fof(f14919,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5877,f5851,f3746,f455])).
fof(f5851,plain,(
  ( ! [X0] : (~sP11('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),X0,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f4934,f4934,f451])).
fof(f5877,plain,(
  ( ! [X0,X1] : (sP10('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),X0,X1)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f4934,f461])).
fof(f16508,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14920,f3746])).
fof(f14920,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5877,f5856,f3746,f455])).
fof(f5856,plain,(
  ( ! [X0] : (~sP11('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),'T_0',X0)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f269,f4934,f452])).
fof(f16507,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14921,f3746])).
fof(f14921,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f523,f5856,f3746,f455])).
fof(f16506,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14923,f3746])).
fof(f14923,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5086,f5856,f3746,f455])).
fof(f16505,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14924,f3746])).
fof(f14924,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5876,f5856,f3746,f455])).
fof(f16504,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14925,f3746])).
fof(f14925,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6900,f5856,f3746,f455])).
fof(f16503,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14926,f3746])).
fof(f14926,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9390,f5856,f3746,f455])).
fof(f16502,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f14939,f15445])).
fof(f15445,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2)) ) | $spl46),
  inference(unit_resulting_resolution,[],[f3746,f460])).
fof(f14939,plain,(
  ~sP10('F_0','T_0','T_0') | $spl46),
  inference(unit_resulting_resolution,[],[f688,f269,f3746,f455])).
fof(f16501,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f14941,f15445])).
fof(f14941,plain,(
  ~sP10('F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),'T_0') | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5336,f269,f3746,f455])).
fof(f16500,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14943,f15445])).
fof(f14943,plain,(
  ~sP10('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'T_0','T_0') | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5076,f269,f3746,f455])).
fof(f16499,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14949,f15445])).
fof(f14949,plain,(
  ~sP10('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))),'T_0','T_0') | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9348,f269,f3746,f455])).
fof(f16498,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14952,f15445])).
fof(f14952,plain,(
  ~sP10('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),'T_0','T_0') | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6868,f269,f3746,f455])).
fof(f16497,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14954,f15445])).
fof(f14954,plain,(
  ~sP10('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),'T_0','T_0') | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5856,f269,f3746,f455])).
fof(f16496,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f14977,f15445])).
fof(f14977,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'F_0')) ) | $spl46),
  inference(unit_resulting_resolution,[],[f686,f3746,f3746,f455])).
fof(f686,plain,(
  ( ! [X0] : (~sP11('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f483,f483,f451])).
fof(f16495,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14979,f15445])).
fof(f14979,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5070,f3746,f3746,f455])).
fof(f16494,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14981,f15445])).
fof(f14981,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5850,f3746,f3746,f455])).
fof(f16493,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f14983,f15445])).
fof(f14983,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6862,f3746,f3746,f455])).
fof(f16492,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14985,f15445])).
fof(f14985,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9338,f3746,f3746,f455])).
fof(f16491,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f14987,f15445])).
fof(f14987,plain,(
  ( ! [X0] : (~sP10('F_0','T_0',X0)) ) | $spl46),
  inference(unit_resulting_resolution,[],[f688,f3746,f3746,f455])).
fof(f16490,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f14989,f15445])).
fof(f14989,plain,(
  ( ! [X0] : (~sP10('F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),X0)) ) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5336,f3746,f3746,f455])).
fof(f16489,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14991,f15445])).
fof(f14991,plain,(
  ( ! [X0] : (~sP10('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),X0,'F_0')) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5068,f3746,f3746,f455])).
fof(f5068,plain,(
  ( ! [X0] : (~sP11('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),X0,'F_0')) ) | $spl3),
  inference(unit_resulting_resolution,[],[f483,f480,f451])).
fof(f16488,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14993,f15445])).
fof(f14993,plain,(
  ( ! [X0] : (~sP10('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),X0,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5071,f3746,f3746,f455])).
fof(f16487,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f14995,f15445])).
fof(f14995,plain,(
  ( ! [X0] : (~sP10('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'T_0',X0)) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5076,f3746,f3746,f455])).
fof(f16486,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15002,f15445])).
fof(f15002,plain,(
  ( ! [X0] : (~sP10('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))),X0,'F_0')) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9336,f3746,f3746,f455])).
fof(f9336,plain,(
  ( ! [X0] : (~sP11('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))),X0,'F_0')) ) | $spl3),
  inference(unit_resulting_resolution,[],[f483,f5685,f451])).
fof(f16485,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15004,f15445])).
fof(f15004,plain,(
  ( ! [X0] : (~sP10('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))),'T_0',X0)) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9348,f3746,f3746,f455])).
fof(f16484,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f15007,f15445])).
fof(f15007,plain,(
  ( ! [X0] : (~sP10('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),X0,'F_0')) ) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6860,f3746,f3746,f455])).
fof(f6860,plain,(
  ( ! [X0] : (~sP11('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),X0,'F_0')) ) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f483,f5683,f451])).
fof(f16483,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f15009,f15445])).
fof(f15009,plain,(
  ( ! [X0] : (~sP10('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),'T_0',X0)) ) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6868,f3746,f3746,f455])).
fof(f16482,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15011,f15445])).
fof(f15011,plain,(
  ( ! [X0] : (~sP10('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),X0,'F_0')) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5848,f3746,f3746,f455])).
fof(f5848,plain,(
  ( ! [X0] : (~sP11('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),X0,'F_0')) ) | $spl3),
  inference(unit_resulting_resolution,[],[f483,f4934,f451])).
fof(f16481,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15013,f15445])).
fof(f15013,plain,(
  ( ! [X0] : (~sP10('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),X0,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5851,f3746,f3746,f455])).
fof(f16480,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15015,f15445])).
fof(f15015,plain,(
  ( ! [X0] : (~sP10('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),'T_0',X0)) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5856,f3746,f3746,f455])).
fof(f16479,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15020,f3746])).
fof(f15020,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5068,f522,f3746,f455])).
fof(f522,plain,(
  ( ! [X0,X1] : (sP10(X0,X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f483,f459])).
fof(f16478,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15021,f3746])).
fof(f15021,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5076,f522,f3746,f455])).
fof(f16477,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15022,f3746])).
fof(f15022,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9336,f522,f3746,f455])).
fof(f16476,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15023,f3746])).
fof(f15023,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9348,f522,f3746,f455])).
fof(f16475,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f15024,f3746])).
fof(f15024,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6860,f522,f3746,f455])).
fof(f16474,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f15025,f3746])).
fof(f15025,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6868,f522,f3746,f455])).
fof(f16473,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15026,f3746])).
fof(f15026,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5848,f522,f3746,f455])).
fof(f16472,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15027,f3746])).
fof(f15027,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5856,f522,f3746,f455])).
fof(f16471,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15031,f3746])).
fof(f15031,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5071,f5086,f3746,f455])).
fof(f16470,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15032,f3746])).
fof(f15032,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5076,f5086,f3746,f455])).
fof(f16469,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15033,f3746])).
fof(f15033,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9348,f5086,f3746,f455])).
fof(f16468,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f15034,f3746])).
fof(f15034,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6868,f5086,f3746,f455])).
fof(f16467,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15035,f3746])).
fof(f15035,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5856,f5086,f3746,f455])).
fof(f16466,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15039,f3746])).
fof(f15039,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5076,f5876,f3746,f455])).
fof(f16465,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15040,f3746])).
fof(f15040,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9348,f5876,f3746,f455])).
fof(f16464,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f15041,f3746])).
fof(f15041,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6868,f5876,f3746,f455])).
fof(f16463,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15042,f3746])).
fof(f15042,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5851,f5876,f3746,f455])).
fof(f16462,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15043,f3746])).
fof(f15043,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5856,f5876,f3746,f455])).
fof(f16461,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f15047,f3746])).
fof(f15047,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5076,f6900,f3746,f455])).
fof(f16460,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f15048,f3746])).
fof(f15048,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f9348,f6900,f3746,f455])).
fof(f16459,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f15049,f3746])).
fof(f15049,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6868,f6900,f3746,f455])).
fof(f16458,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f15050,f3746])).
fof(f15050,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5856,f6900,f3746,f455])).
fof(f16457,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15054,f3746])).
fof(f15054,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5076,f9390,f3746,f455])).
fof(f16456,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15055,f3746])).
fof(f15055,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9348,f9390,f3746,f455])).
fof(f16455,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f15056,f3746])).
fof(f15056,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6868,f9390,f3746,f455])).
fof(f16454,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15057,f3746])).
fof(f15057,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5856,f9390,f3746,f455])).
fof(f16453,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15064,f3746])).
fof(f15064,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5076,f523,f3746,f455])).
fof(f16452,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15065,f3746])).
fof(f15065,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5068,f523,f3746,f455])).
fof(f16451,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15066,f3746])).
fof(f15066,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5071,f523,f3746,f455])).
fof(f16450,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15067,f3746])).
fof(f15067,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9348,f523,f3746,f455])).
fof(f16449,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15068,f3746])).
fof(f15068,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9336,f523,f3746,f455])).
fof(f16448,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f15069,f3746])).
fof(f15069,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6868,f523,f3746,f455])).
fof(f16447,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f15070,f3746])).
fof(f15070,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6860,f523,f3746,f455])).
fof(f16446,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15071,f3746])).
fof(f15071,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5856,f523,f3746,f455])).
fof(f16445,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15072,f3746])).
fof(f15072,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5848,f523,f3746,f455])).
fof(f16444,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15073,f3746])).
fof(f15073,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5851,f523,f3746,f455])).
fof(f16443,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f15079,f3746])).
fof(f15079,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5068,f525,f3746,f455])).
fof(f525,plain,(
  ( ! [X0,X1] : (sP10(X0,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f481,f460])).
fof(f16442,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f15080,f3746])).
fof(f15080,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5071,f525,f3746,f455])).
fof(f16441,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f15081,f3746])).
fof(f15081,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f9336,f525,f3746,f455])).
fof(f16440,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f15082,f3746])).
fof(f15082,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6860,f525,f3746,f455])).
fof(f16439,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f15083,f3746])).
fof(f15083,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5848,f525,f3746,f455])).
fof(f16438,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f15084,f3746])).
fof(f15084,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5851,f525,f3746,f455])).
fof(f16437,plain,(
  $false | ($spl0 | $spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15091,f3746])).
fof(f15091,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl0 | $spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5068,f5363,f3746,f455])).
fof(f16436,plain,(
  $false | ($spl0 | $spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15092,f3746])).
fof(f15092,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl0 | $spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5071,f5363,f3746,f455])).
fof(f16435,plain,(
  $false | ($spl0 | $spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15093,f3746])).
fof(f15093,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | ($spl0 | $spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9336,f5363,f3746,f455])).
fof(f16434,plain,(
  $false | ($spl0 | $spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f15094,f3746])).
fof(f15094,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl0 | $spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6860,f5363,f3746,f455])).
fof(f16433,plain,(
  $false | ($spl0 | $spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15095,f3746])).
fof(f15095,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl0 | $spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5848,f5363,f3746,f455])).
fof(f16432,plain,(
  $false | ($spl0 | $spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15096,f3746])).
fof(f15096,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl0 | $spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5851,f5363,f3746,f455])).
fof(f16431,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15098,f3746])).
fof(f15098,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5068,f5087,f3746,f455])).
fof(f16430,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15099,f3746])).
fof(f15099,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5848,f5877,f3746,f455])).
fof(f16429,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f15100,f3746])).
fof(f15100,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6860,f6901,f3746,f455])).
fof(f16428,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15101,f3746])).
fof(f15101,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9336,f9391,f3746,f455])).
fof(f16427,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f15126,f15445])).
fof(f15126,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'F_0')) ) | $spl46),
  inference(unit_resulting_resolution,[],[f686,f3746,f3746,f455])).
fof(f16426,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15127,f15445])).
fof(f15127,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5070,f3746,f3746,f455])).
fof(f16425,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15128,f15445])).
fof(f15128,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5850,f3746,f3746,f455])).
fof(f16424,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f15129,f15445])).
fof(f15129,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6862,f3746,f3746,f455])).
fof(f16423,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15130,f15445])).
fof(f15130,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9338,f3746,f3746,f455])).
fof(f16422,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f15131,f15445])).
fof(f15131,plain,(
  ( ! [X0] : (~sP10('F_0','T_0',X0)) ) | $spl46),
  inference(unit_resulting_resolution,[],[f688,f3746,f3746,f455])).
fof(f16421,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f15132,f15445])).
fof(f15132,plain,(
  ( ! [X0] : (~sP10('F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),X0)) ) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5336,f3746,f3746,f455])).
fof(f16420,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15133,f15445])).
fof(f15133,plain,(
  ( ! [X0] : (~sP10('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),X0,'F_0')) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5068,f3746,f3746,f455])).
fof(f16419,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15134,f15445])).
fof(f15134,plain,(
  ( ! [X0] : (~sP10('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),X0,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5071,f3746,f3746,f455])).
fof(f16418,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15135,f15445])).
fof(f15135,plain,(
  ( ! [X0] : (~sP10('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'T_0',X0)) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5076,f3746,f3746,f455])).
fof(f16417,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15141,f15445])).
fof(f15141,plain,(
  ( ! [X0] : (~sP10('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))),X0,'F_0')) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9336,f3746,f3746,f455])).
fof(f16416,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15142,f15445])).
fof(f15142,plain,(
  ( ! [X0] : (~sP10('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))),'T_0',X0)) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9348,f3746,f3746,f455])).
fof(f16415,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f15144,f15445])).
fof(f15144,plain,(
  ( ! [X0] : (~sP10('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),X0,'F_0')) ) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6860,f3746,f3746,f455])).
fof(f16414,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f15145,f15445])).
fof(f15145,plain,(
  ( ! [X0] : (~sP10('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),'T_0',X0)) ) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6868,f3746,f3746,f455])).
fof(f16413,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15146,f15445])).
fof(f15146,plain,(
  ( ! [X0] : (~sP10('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),X0,'F_0')) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5848,f3746,f3746,f455])).
fof(f16412,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15147,f15445])).
fof(f15147,plain,(
  ( ! [X0] : (~sP10('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),X0,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5851,f3746,f3746,f455])).
fof(f16411,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15148,f15445])).
fof(f15148,plain,(
  ( ! [X0] : (~sP10('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),'T_0',X0)) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5856,f3746,f3746,f455])).
fof(f16410,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f15240,f15445])).
fof(f15240,plain,(
  ~sP10('F_0','F_0','F_0') | $spl46),
  inference(unit_resulting_resolution,[],[f686,f483,f3746,f456])).
fof(f456,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f260])).
fof(f16409,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15241,f15445])).
fof(f15241,plain,(
  ~sP10('F_0','F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5070,f483,f3746,f456])).
fof(f16408,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15242,f15445])).
fof(f15242,plain,(
  ~sP10('F_0','F_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5850,f483,f3746,f456])).
fof(f16407,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f15243,f15445])).
fof(f15243,plain,(
  ~sP10('F_0','F_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6862,f483,f3746,f456])).
fof(f16406,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15244,f15445])).
fof(f15244,plain,(
  ~sP10('F_0','F_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9338,f483,f3746,f456])).
fof(f16405,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15248,f15445])).
fof(f15248,plain,(
  ~sP10('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'F_0','F_0') | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5068,f483,f3746,f456])).
fof(f16404,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15249,f15445])).
fof(f15249,plain,(
  ~sP10('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5071,f483,f3746,f456])).
fof(f16403,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15250,f15445])).
fof(f15250,plain,(
  ~sP10('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))),'F_0','F_0') | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9336,f483,f3746,f456])).
fof(f16402,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f15252,f15445])).
fof(f15252,plain,(
  ~sP10('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),'F_0','F_0') | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6860,f483,f3746,f456])).
fof(f16401,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15253,f15445])).
fof(f15253,plain,(
  ~sP10('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),'F_0','F_0') | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5848,f483,f3746,f456])).
fof(f16400,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15254,f15445])).
fof(f15254,plain,(
  ~sP10('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),'F_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5851,f483,f3746,f456])).
fof(f16399,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f15340,f15445])).
fof(f15340,plain,(
  ~sP10('F_0','F_0','F_0') | $spl46),
  inference(unit_resulting_resolution,[],[f686,f483,f3746,f457])).
fof(f457,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f260])).
fof(f16398,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15341,f15445])).
fof(f15341,plain,(
  ~sP10('F_0','F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5070,f483,f3746,f457])).
fof(f16397,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15342,f15445])).
fof(f15342,plain,(
  ~sP10('F_0','F_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5850,f483,f3746,f457])).
fof(f16396,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f15343,f15445])).
fof(f15343,plain,(
  ~sP10('F_0','F_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6862,f483,f3746,f457])).
fof(f16395,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15344,f15445])).
fof(f15344,plain,(
  ~sP10('F_0','F_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9338,f483,f3746,f457])).
fof(f16394,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15348,f15445])).
fof(f15348,plain,(
  ~sP10('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'F_0','F_0') | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5068,f483,f3746,f457])).
fof(f16393,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15349,f15445])).
fof(f15349,plain,(
  ~sP10('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5071,f483,f3746,f457])).
fof(f16392,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15350,f15445])).
fof(f15350,plain,(
  ~sP10('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))),'F_0','F_0') | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9336,f483,f3746,f457])).
fof(f16391,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f15352,f15445])).
fof(f15352,plain,(
  ~sP10('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),'F_0','F_0') | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6860,f483,f3746,f457])).
fof(f16390,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15353,f15445])).
fof(f15353,plain,(
  ~sP10('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),'F_0','F_0') | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5848,f483,f3746,f457])).
fof(f16389,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15354,f15445])).
fof(f15354,plain,(
  ~sP10('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),'F_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5851,f483,f3746,f457])).
fof(f16388,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15397,f3746])).
fof(f15397,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5086,f483,f3746,f458])).
fof(f458,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f263])).
fof(f16387,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15398,f3746])).
fof(f15398,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5876,f483,f3746,f458])).
fof(f16386,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f15399,f3746])).
fof(f15399,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6900,f483,f3746,f458])).
fof(f16385,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15400,f3746])).
fof(f15400,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9390,f483,f3746,f458])).
fof(f16384,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f15401,f3746])).
fof(f15401,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl46),
  inference(unit_resulting_resolution,[],[f526,f483,f3746,f458])).
fof(f16383,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15402,f3746])).
fof(f15402,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5087,f483,f3746,f458])).
fof(f16382,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15403,f3746])).
fof(f15403,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5877,f483,f3746,f458])).
fof(f16381,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f15404,f3746])).
fof(f15404,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6901,f483,f3746,f458])).
fof(f16380,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15405,f3746])).
fof(f15405,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9391,f483,f3746,f458])).
fof(f16379,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f15408,f15445])).
fof(f15408,plain,(
  ( ! [X0,X1] : (~sP10(X0,'F_0',X1)) ) | $spl46),
  inference(unit_resulting_resolution,[],[f3746,f483,f3746,f458])).
fof(f16378,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f15431,f3746])).
fof(f15431,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl46),
  inference(unit_resulting_resolution,[],[f522,f483,f3746,f458])).
fof(f16377,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15432,f3746])).
fof(f15432,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5086,f483,f3746,f458])).
fof(f16376,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15433,f3746])).
fof(f15433,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5876,f483,f3746,f458])).
fof(f16375,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f15434,f3746])).
fof(f15434,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6900,f483,f3746,f458])).
fof(f16374,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15435,f3746])).
fof(f15435,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9390,f483,f3746,f458])).
fof(f16373,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15437,f3746])).
fof(f15437,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5087,f483,f3746,f458])).
fof(f16372,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15438,f3746])).
fof(f15438,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5877,f483,f3746,f458])).
fof(f16371,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f15439,f3746])).
fof(f15439,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6901,f483,f3746,f458])).
fof(f16370,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15440,f3746])).
fof(f15440,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9391,f483,f3746,f458])).
fof(f16369,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f15441,f15445])).
fof(f15441,plain,(
  ( ! [X0] : (~sP10('T_0','F_0',X0)) ) | $spl46),
  inference(unit_resulting_resolution,[],[f269,f483,f3746,f458])).
fof(f16368,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f15443,f15445])).
fof(f15443,plain,(
  ( ! [X0,X1] : (~sP10(X0,'F_0',X1)) ) | $spl46),
  inference(unit_resulting_resolution,[],[f3746,f483,f3746,f458])).
fof(f16367,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f15508,f269])).
fof(f15508,plain,(
  ~p(s(bool,'T_0')) | $spl46),
  inference(unit_resulting_resolution,[],[f1185,f3746,f464])).
fof(f464,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f265])).
fof(f265,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f264])).
fof(f264,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f160])).
fof(f160,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP11(X2,X0,X1))),
  inference(definition_folding,[],[f142,f159,f158])).
fof(f142,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f117])).
fof(f117,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f116])).
fof(f116,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X20,X21,X22] : ((p(s(bool,X22)) <=> (p(s(bool,X21)) => p(s(bool,X20)))) <=> ((p(s(bool,X22)) | p(s(bool,X21))) & (p(s(bool,X22)) | ~p(s(bool,X20))) & (~p(s(bool,X21)) | p(s(bool,X20)) | ~p(s(bool,X22)))))),
  file('thm.pred_set.UNION_SUBSET','thm.sat.sat_27_prove_ltb')).
fof(f1185,plain,(
  ~sP11('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f1173,f453])).
fof(f453,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f260])).
fof(f1173,plain,(
  ~sP10('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f269,f483,f269,f458])).
fof(f16366,plain,(
  $false | ($spl4 | $spl46)),
  inference(subsumption_resolution,[],[f15509,f269])).
fof(f15509,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f1508,f3746,f464])).
fof(f1508,plain,(
  ~sP11('T_0','F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))) | $spl4),
  inference(unit_resulting_resolution,[],[f1496,f453])).
fof(f1496,plain,(
  ~sP10('T_0','F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))) | $spl4),
  inference(unit_resulting_resolution,[],[f483,f269,f1181])).
fof(f1181,plain,(
  ( ! [X2,X3] : (~sP10(X3,X2,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))) | ~p(s(bool,X3)) | p(s(bool,X2))) ) | $spl4),
  inference(subsumption_resolution,[],[f1175,f269])).
fof(f1175,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,X2)) | ~p(s(bool,X3)) | ~sP10(X3,X2,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))) ) | $spl4),
  inference(superposition,[],[f458,f531])).
fof(f16365,plain,(
  $false | ($spl0 | $spl38 | $spl46)),
  inference(subsumption_resolution,[],[f15510,f269])).
fof(f15510,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl38 | $spl46)),
  inference(unit_resulting_resolution,[],[f5422,f3746,f464])).
fof(f5422,plain,(
  ~sP11('T_0','F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15))) | ($spl0 | $spl38)),
  inference(unit_resulting_resolution,[],[f483,f477,f2416])).
fof(f16364,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15511,f269])).
fof(f15511,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9842,f3746,f464])).
fof(f9842,plain,(
  ~sP11('T_0','F_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK13))) | $spl3),
  inference(unit_resulting_resolution,[],[f9825,f453])).
fof(f9825,plain,(
  ~sP10('T_0','F_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK13))) | $spl3),
  inference(unit_resulting_resolution,[],[f483,f269,f8019])).
fof(f8019,plain,(
  ( ! [X30,X29] : (~sP10(X30,X29,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK13))) | ~p(s(bool,X30)) | p(s(bool,X29))) ) | $spl3),
  inference(subsumption_resolution,[],[f7965,f269])).
fof(f7965,plain,(
  ( ! [X30,X29] : (~p(s(bool,'T_0')) | p(s(bool,X29)) | ~p(s(bool,X30)) | ~sP10(X30,X29,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | $spl3),
  inference(superposition,[],[f458,f4933])).
fof(f16363,plain,(
  $false | ($spl3 | $spl38 | $spl46)),
  inference(subsumption_resolution,[],[f15512,f269])).
fof(f15512,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl38 | $spl46)),
  inference(unit_resulting_resolution,[],[f5170,f3746,f464])).
fof(f5170,plain,(
  ~sP11('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'T_0') | ($spl3 | $spl38)),
  inference(unit_resulting_resolution,[],[f269,f480,f2416])).
fof(f16362,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f15513,f269])).
fof(f15513,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f10253,f3746,f464])).
fof(f10253,plain,(
  ~sP11('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f5137,f453])).
fof(f5137,plain,(
  ~sP10('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f269,f480,f1181])).
fof(f16361,plain,(
  $false | ($spl3 | $spl38 | $spl46)),
  inference(subsumption_resolution,[],[f15514,f269])).
fof(f15514,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl38 | $spl46)),
  inference(unit_resulting_resolution,[],[f5972,f3746,f464])).
fof(f5972,plain,(
  ~sP11('T_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),'T_0') | ($spl3 | $spl38)),
  inference(unit_resulting_resolution,[],[f269,f4934,f2416])).
fof(f16360,plain,(
  $false | ($spl3 | $spl4 | $spl38 | $spl46)),
  inference(subsumption_resolution,[],[f15515,f269])).
fof(f15515,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl4 | $spl38 | $spl46)),
  inference(unit_resulting_resolution,[],[f7012,f3746,f464])).
fof(f7012,plain,(
  ~sP11('T_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),'T_0') | ($spl3 | $spl4 | $spl38)),
  inference(unit_resulting_resolution,[],[f269,f5683,f2416])).
fof(f16359,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15523,f3746])).
fof(f15523,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5068,f3746,f464])).
fof(f16358,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15524,f3746])).
fof(f15524,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5071,f3746,f464])).
fof(f16357,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15525,f3746])).
fof(f15525,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5076,f3746,f464])).
fof(f16356,plain,(
  $false | ($spl4 | $spl46)),
  inference(subsumption_resolution,[],[f15526,f3746])).
fof(f15526,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))) | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f1542,f3746,f464])).
fof(f1542,plain,(
  ~sP11('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),'F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f269,f483,f1221])).
fof(f1221,plain,(
  ( ! [X2,X3] : (~sP11('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),X3,X2) | p(s(bool,X3)) | ~p(s(bool,X2))) ) | $spl4),
  inference(subsumption_resolution,[],[f1215,f269])).
fof(f1215,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | ~p(s(bool,X2)) | p(s(bool,X3)) | ~sP11('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),X3,X2)) ) | $spl4),
  inference(superposition,[],[f465,f531])).
fof(f465,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f265])).
fof(f16355,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f15527,f3746])).
fof(f15527,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)))) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5376,f3746,f464])).
fof(f5376,plain,(
  ~sP11('const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),'F_0','T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f269,f483,f477,f465])).
fof(f16354,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f15528,f3746])).
fof(f15528,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)))) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5377,f3746,f464])).
fof(f5377,plain,(
  ~sP11('const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),'F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15))) | $spl0),
  inference(unit_resulting_resolution,[],[f477,f483,f477,f465])).
fof(f16353,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f15529,f3746])).
fof(f15529,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5139,f3746,f464])).
fof(f5139,plain,(
  ~sP11('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'T_0') | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f269,f480,f1221])).
fof(f16352,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f15530,f3746])).
fof(f15530,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5937,f3746,f464])).
fof(f5937,plain,(
  ~sP11('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),'T_0') | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f269,f4934,f1221])).
fof(f16351,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15531,f3746])).
fof(f15531,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9336,f3746,f464])).
fof(f16350,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15532,f3746])).
fof(f15532,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9348,f3746,f464])).
fof(f16349,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15533,f3746])).
fof(f15533,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9854,f3746,f464])).
fof(f9854,plain,(
  ~sP11('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),'F_0','T_0') | $spl3),
  inference(unit_resulting_resolution,[],[f269,f483,f8022])).
fof(f8022,plain,(
  ( ! [X35,X36] : (~sP11('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X36,X35) | p(s(bool,X36)) | ~p(s(bool,X35))) ) | $spl3),
  inference(subsumption_resolution,[],[f7968,f269])).
fof(f7968,plain,(
  ( ! [X35,X36] : (~p(s(bool,'T_0')) | ~p(s(bool,X35)) | p(s(bool,X36)) | ~sP11('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X36,X35)) ) | $spl3),
  inference(superposition,[],[f465,f4933])).
fof(f16348,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f15534,f3746])).
fof(f15534,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6860,f3746,f464])).
fof(f16347,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f15535,f3746])).
fof(f15535,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6868,f3746,f464])).
fof(f16346,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15536,f3746])).
fof(f15536,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5848,f3746,f464])).
fof(f16345,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15537,f3746])).
fof(f15537,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5851,f3746,f464])).
fof(f16344,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15538,f3746])).
fof(f15538,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5856,f3746,f464])).
fof(f16343,plain,(
  $false | ($spl3 | $spl38 | $spl46)),
  inference(subsumption_resolution,[],[f15547,f3746])).
fof(f15547,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl38 | $spl46)),
  inference(unit_resulting_resolution,[],[f5170,f3746,f464])).
fof(f16342,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f15548,f3746])).
fof(f15548,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f10253,f3746,f464])).
fof(f16341,plain,(
  $false | ($spl3 | $spl38 | $spl46)),
  inference(subsumption_resolution,[],[f15549,f3746])).
fof(f15549,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl38 | $spl46)),
  inference(unit_resulting_resolution,[],[f5972,f3746,f464])).
fof(f16340,plain,(
  $false | ($spl3 | $spl4 | $spl38 | $spl46)),
  inference(subsumption_resolution,[],[f15550,f3746])).
fof(f15550,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl38 | $spl46)),
  inference(unit_resulting_resolution,[],[f7012,f3746,f464])).
fof(f16339,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f15551,f3746])).
fof(f15551,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl46),
  inference(unit_resulting_resolution,[],[f686,f3746,f464])).
fof(f16338,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15552,f3746])).
fof(f15552,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5070,f3746,f464])).
fof(f16337,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15553,f3746])).
fof(f15553,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5850,f3746,f464])).
fof(f16336,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f15554,f3746])).
fof(f15554,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6862,f3746,f464])).
fof(f16335,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15555,f3746])).
fof(f15555,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9338,f3746,f464])).
fof(f16334,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f15556,f269])).
fof(f15556,plain,(
  ~p(s(bool,'T_0')) | $spl46),
  inference(unit_resulting_resolution,[],[f688,f3746,f464])).
fof(f16333,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f15557,f3746])).
fof(f15557,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)))) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5336,f3746,f464])).
fof(f16332,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15558,f3746])).
fof(f15558,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5068,f3746,f464])).
fof(f16331,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15559,f3746])).
fof(f15559,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5071,f3746,f464])).
fof(f16330,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15560,f269])).
fof(f15560,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5076,f3746,f464])).
fof(f16329,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f15564,f3746])).
fof(f15564,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5139,f3746,f464])).
fof(f16328,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f15565,f3746])).
fof(f15565,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5937,f3746,f464])).
fof(f16327,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15566,f3746])).
fof(f15566,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9336,f3746,f464])).
fof(f16326,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15567,f269])).
fof(f15567,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9348,f3746,f464])).
fof(f16325,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f15569,f3746])).
fof(f15569,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6860,f3746,f464])).
fof(f16324,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f15570,f269])).
fof(f15570,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6868,f3746,f464])).
fof(f16323,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15571,f3746])).
fof(f15571,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5848,f3746,f464])).
fof(f16322,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15572,f3746])).
fof(f15572,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5851,f3746,f464])).
fof(f16321,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15573,f269])).
fof(f15573,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5856,f3746,f464])).
fof(f16320,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f15578,f269])).
fof(f15578,plain,(
  ~p(s(bool,'T_0')) | $spl46),
  inference(unit_resulting_resolution,[],[f696,f483,f3746,f465])).
fof(f696,plain,(
  ( ! [X0] : (sP11('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f483,f269,f463])).
fof(f463,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f265])).
fof(f16319,plain,(
  $false | ($spl3 | $spl34 | $spl46)),
  inference(subsumption_resolution,[],[f15579,f269])).
fof(f15579,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl34 | $spl46)),
  inference(unit_resulting_resolution,[],[f5169,f483,f3746,f465])).
fof(f5169,plain,(
  ( ! [X0] : (sP11('T_0',X0,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) ) | ($spl3 | $spl34)),
  inference(unit_resulting_resolution,[],[f480,f2410])).
fof(f16318,plain,(
  $false | ($spl3 | $spl34 | $spl46)),
  inference(subsumption_resolution,[],[f15580,f269])).
fof(f15580,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl34 | $spl46)),
  inference(unit_resulting_resolution,[],[f5971,f483,f3746,f465])).
fof(f5971,plain,(
  ( ! [X0] : (sP11('T_0',X0,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) | ($spl3 | $spl34)),
  inference(unit_resulting_resolution,[],[f4934,f2410])).
fof(f16317,plain,(
  $false | ($spl3 | $spl4 | $spl34 | $spl46)),
  inference(subsumption_resolution,[],[f15581,f269])).
fof(f15581,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl4 | $spl34 | $spl46)),
  inference(unit_resulting_resolution,[],[f7011,f483,f3746,f465])).
fof(f7011,plain,(
  ( ! [X0] : (sP11('T_0',X0,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) | ($spl3 | $spl4 | $spl34)),
  inference(unit_resulting_resolution,[],[f5683,f2410])).
fof(f16316,plain,(
  $false | ($spl46 | $spl63)),
  inference(subsumption_resolution,[],[f15584,f3746])).
fof(f15584,plain,(
  ~p(s(bool,f(s(bool,'T_0')))) | ($spl46 | $spl63)),
  inference(unit_resulting_resolution,[],[f8996,f483,f3746,f465])).
fof(f8996,plain,(
  ( ! [X0] : (sP11(f(s(bool,'T_0')),X0,'F_0')) ) | $spl63),
  inference(unit_resulting_resolution,[],[f483,f4544])).
fof(f4544,plain,(
  ( ! [X30,X29] : (sP11(f(s(bool,'T_0')),X30,X29) | p(s(bool,X29))) ) | $spl63),
  inference(subsumption_resolution,[],[f4491,f269])).
fof(f4491,plain,(
  ( ! [X30,X29] : (~p(s(bool,'T_0')) | p(s(bool,X29)) | sP11(f(s(bool,'T_0')),X30,X29)) ) | $spl63),
  inference(superposition,[],[f463,f4246])).
fof(f16315,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15585,f3746])).
fof(f15585,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5091,f483,f3746,f465])).
fof(f5091,plain,(
  sP11('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'F_0','T_0') | $spl3),
  inference(unit_resulting_resolution,[],[f483,f269,f480,f462])).
fof(f462,plain,(
  ( ! [X2,X0,X1] : (sP11(X2,X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f265])).
fof(f16314,plain,(
  $false | ($spl4 | $spl46)),
  inference(subsumption_resolution,[],[f15586,f3746])).
fof(f15586,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))) | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f852,f483,f3746,f465])).
fof(f852,plain,(
  ( ! [X0] : (sP11('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),X0,'F_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f483,f700])).
fof(f700,plain,(
  ( ! [X2,X3] : (sP11('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),X3,X2) | p(s(bool,X2))) ) | $spl4),
  inference(subsumption_resolution,[],[f698,f269])).
fof(f698,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,X2)) | sP11('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),X3,X2)) ) | $spl4),
  inference(superposition,[],[f463,f531])).
fof(f16313,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f15587,f3746])).
fof(f15587,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)))) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5378,f483,f3746,f465])).
fof(f5378,plain,(
  ( ! [X0] : (sP11('const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),X0,'F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f522,f477,f468])).
fof(f16312,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f15588,f3746])).
fof(f15588,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5111,f483,f3746,f465])).
fof(f5111,plain,(
  ( ! [X0] : (sP11('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),X0,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) ) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f480,f700])).
fof(f16311,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f15589,f3746])).
fof(f15589,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5905,f483,f3746,f465])).
fof(f5905,plain,(
  ( ! [X0] : (sP11('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),X0,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f4934,f700])).
fof(f16310,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15590,f3746])).
fof(f15590,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5882,f483,f3746,f465])).
fof(f5882,plain,(
  sP11('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),'F_0','T_0') | $spl3),
  inference(unit_resulting_resolution,[],[f483,f269,f4934,f462])).
fof(f16309,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f15591,f3746])).
fof(f15591,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6907,f483,f3746,f465])).
fof(f6907,plain,(
  sP11('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),'F_0','T_0') | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f483,f269,f5683,f462])).
fof(f16308,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15592,f3746])).
fof(f15592,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f8926,f483,f3746,f465])).
fof(f8926,plain,(
  ( ! [X0] : (sP11('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X0,'F_0')) ) | $spl3),
  inference(unit_resulting_resolution,[],[f483,f8020])).
fof(f8020,plain,(
  ( ! [X31,X32] : (sP11('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X32,X31) | p(s(bool,X31))) ) | $spl3),
  inference(subsumption_resolution,[],[f7966,f269])).
fof(f7966,plain,(
  ( ! [X31,X32] : (~p(s(bool,'T_0')) | p(s(bool,X31)) | sP11('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X32,X31)) ) | $spl3),
  inference(superposition,[],[f463,f4933])).
fof(f16307,plain,(
  $false | ($spl3 | $spl34 | $spl46)),
  inference(subsumption_resolution,[],[f15637,f3746])).
fof(f15637,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl34 | $spl46)),
  inference(unit_resulting_resolution,[],[f5169,f483,f3746,f465])).
fof(f16306,plain,(
  $false | ($spl3 | $spl34 | $spl46)),
  inference(subsumption_resolution,[],[f15638,f3746])).
fof(f15638,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl34 | $spl46)),
  inference(unit_resulting_resolution,[],[f5971,f483,f3746,f465])).
fof(f16305,plain,(
  $false | ($spl3 | $spl4 | $spl34 | $spl46)),
  inference(subsumption_resolution,[],[f15639,f3746])).
fof(f15639,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4 | $spl34 | $spl46)),
  inference(unit_resulting_resolution,[],[f7011,f483,f3746,f465])).
fof(f16304,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f15640,f269])).
fof(f15640,plain,(
  ~p(s(bool,'T_0')) | $spl46),
  inference(unit_resulting_resolution,[],[f1159,f483,f3746,f465])).
fof(f1159,plain,(
  sP11('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f526,f269,f483,f457])).
fof(f16303,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f15641,f3746])).
fof(f15641,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)))) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5364,f483,f3746,f465])).
fof(f5364,plain,(
  sP11('F_0','F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15))) | $spl0),
  inference(unit_resulting_resolution,[],[f483,f483,f477,f462])).
fof(f16302,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15643,f269])).
fof(f15643,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5091,f483,f3746,f465])).
fof(f16301,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f15646,f3746])).
fof(f15646,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5111,f483,f3746,f465])).
fof(f16300,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f15647,f3746])).
fof(f15647,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f5905,f483,f3746,f465])).
fof(f16299,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15648,f269])).
fof(f15648,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5882,f483,f3746,f465])).
fof(f16298,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f15649,f269])).
fof(f15649,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6907,f483,f3746,f465])).
fof(f16297,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f15716,f522])).
fof(f15716,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'F_0')) ) | $spl46),
  inference(unit_resulting_resolution,[],[f686,f3746,f468])).
fof(f16296,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15717,f5086])).
fof(f15717,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5070,f3746,f468])).
fof(f16295,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15718,f5876])).
fof(f15718,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5850,f3746,f468])).
fof(f16294,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f15719,f6900])).
fof(f15719,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6862,f3746,f468])).
fof(f16293,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15720,f9390])).
fof(f15720,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9338,f3746,f468])).
fof(f16292,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f15721,f523])).
fof(f15721,plain,(
  ( ! [X0] : (~sP10('F_0','T_0',X0)) ) | $spl46),
  inference(unit_resulting_resolution,[],[f688,f3746,f468])).
fof(f16291,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f15722,f5363])).
fof(f15722,plain,(
  ( ! [X0] : (~sP10('F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),X0)) ) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5336,f3746,f468])).
fof(f16290,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15723,f522])).
fof(f15723,plain,(
  ( ! [X0] : (~sP10('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),X0,'F_0')) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5068,f3746,f468])).
fof(f16289,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15724,f5086])).
fof(f15724,plain,(
  ( ! [X0] : (~sP10('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),X0,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5071,f3746,f468])).
fof(f16288,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15725,f523])).
fof(f15725,plain,(
  ( ! [X0] : (~sP10('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'T_0',X0)) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5076,f3746,f468])).
fof(f16287,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15731,f522])).
fof(f15731,plain,(
  ( ! [X0] : (~sP10('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))),X0,'F_0')) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9336,f3746,f468])).
fof(f16286,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15732,f523])).
fof(f15732,plain,(
  ( ! [X0] : (~sP10('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))),'T_0',X0)) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9348,f3746,f468])).
fof(f16285,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f15734,f522])).
fof(f15734,plain,(
  ( ! [X0] : (~sP10('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),X0,'F_0')) ) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6860,f3746,f468])).
fof(f16284,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f15735,f523])).
fof(f15735,plain,(
  ( ! [X0] : (~sP10('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),'T_0',X0)) ) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6868,f3746,f468])).
fof(f16283,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15736,f522])).
fof(f15736,plain,(
  ( ! [X0] : (~sP10('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),X0,'F_0')) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5848,f3746,f468])).
fof(f16282,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15737,f5876])).
fof(f15737,plain,(
  ( ! [X0] : (~sP10('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),X0,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5851,f3746,f468])).
fof(f16281,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15738,f523])).
fof(f15738,plain,(
  ( ! [X0] : (~sP10('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),'T_0',X0)) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5856,f3746,f468])).
fof(f16280,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f15773,f515])).
fof(f515,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f269,f442])).
fof(f442,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f255])).
fof(f15773,plain,(
  ( ! [X0] : (~sP8('F_0',X0,'T_0')) ) | $spl46),
  inference(unit_resulting_resolution,[],[f654,f3746,f469])).
fof(f16279,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f15774,f5323])).
fof(f5323,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f477,f442])).
fof(f15774,plain,(
  ( ! [X0] : (~sP8('F_0',X0,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)))) ) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5304,f3746,f469])).
fof(f16278,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f15775,f518])).
fof(f518,plain,(
  ( ! [X0,X1] : (sP8(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f269,f443])).
fof(f15775,plain,(
  ( ! [X0] : (~sP8('F_0','T_0',X0)) ) | $spl46),
  inference(unit_resulting_resolution,[],[f657,f3746,f469])).
fof(f16277,plain,(
  $false | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f15776,f5324])).
fof(f5324,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),X1)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f477,f443])).
fof(f15776,plain,(
  ( ! [X0] : (~sP8('F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),X0)) ) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5309,f3746,f469])).
fof(f16276,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15777,f515])).
fof(f15777,plain,(
  ( ! [X0] : (~sP8('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),X0,'T_0')) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5017,f3746,f469])).
fof(f16275,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15778,f518])).
fof(f15778,plain,(
  ( ! [X0] : (~sP8('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'T_0',X0)) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5018,f3746,f469])).
fof(f16274,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15785,f515])).
fof(f15785,plain,(
  ( ! [X0] : (~sP8('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))),X0,'T_0')) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9237,f3746,f469])).
fof(f16273,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15786,f518])).
fof(f15786,plain,(
  ( ! [X0] : (~sP8('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))),'T_0',X0)) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9238,f3746,f469])).
fof(f16272,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f15788,f515])).
fof(f15788,plain,(
  ( ! [X0] : (~sP8('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),X0,'T_0')) ) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6785,f3746,f469])).
fof(f16271,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f15789,f518])).
fof(f15789,plain,(
  ( ! [X0] : (~sP8('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),'T_0',X0)) ) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6786,f3746,f469])).
fof(f16270,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15790,f515])).
fof(f15790,plain,(
  ( ! [X0] : (~sP8('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),X0,'T_0')) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5785,f3746,f469])).
fof(f16269,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f15791,f518])).
fof(f15791,plain,(
  ( ! [X0] : (~sP8('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),'T_0',X0)) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5786,f3746,f469])).
fof(f16195,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f8926,f3746,f8022])).
fof(f16192,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9854,f3746,f8021])).
fof(f8021,plain,(
  ( ! [X33,X34] : (sP11('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X33,X34) | ~p(s(bool,X33))) ) | $spl3),
  inference(subsumption_resolution,[],[f7967,f269])).
fof(f7967,plain,(
  ( ! [X33,X34] : (~p(s(bool,'T_0')) | ~p(s(bool,X33)) | sP11('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X33,X34)) ) | $spl3),
  inference(superposition,[],[f464,f4933])).
fof(f16190,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f9391,f3746,f8019])).
fof(f16188,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f6901,f3746,f8019])).
fof(f16186,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f5877,f3746,f8019])).
fof(f16184,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f5087,f3746,f8019])).
fof(f16182,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f526,f3746,f8019])).
fof(f16176,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9813,f3746,f8017])).
fof(f8017,plain,(
  ( ! [X26,X25] : (sP9('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X25,X26) | ~p(s(bool,X25))) ) | $spl3),
  inference(subsumption_resolution,[],[f7963,f269])).
fof(f7963,plain,(
  ( ! [X26,X25] : (~p(s(bool,'T_0')) | ~p(s(bool,X25)) | sP9('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X25,X26)) ) | $spl3),
  inference(superposition,[],[f447,f4933])).
fof(f16174,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9813,f3746,f8016])).
fof(f8016,plain,(
  ( ! [X24,X23] : (sP9('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X24,X23) | ~p(s(bool,X23))) ) | $spl3),
  inference(subsumption_resolution,[],[f7962,f269])).
fof(f7962,plain,(
  ( ! [X24,X23] : (~p(s(bool,'T_0')) | ~p(s(bool,X23)) | sP9('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X24,X23)) ) | $spl3),
  inference(superposition,[],[f446,f4933])).
fof(f16166,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f9227,f3746,f8014])).
fof(f8014,plain,(
  ( ! [X19,X20] : (~sP7(X20,X19,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK13))) | p(s(bool,X20)) | ~p(s(bool,X19))) ) | $spl3),
  inference(subsumption_resolution,[],[f7960,f269])).
fof(f7960,plain,(
  ( ! [X19,X20] : (~p(s(bool,'T_0')) | ~p(s(bool,X19)) | p(s(bool,X20)) | ~sP7(X20,X19,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | $spl3),
  inference(superposition,[],[f433,f4933])).
fof(f16164,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f6776,f3746,f8014])).
fof(f16162,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f5778,f3746,f8014])).
fof(f16160,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f5011,f3746,f8014])).
fof(f16158,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f648,f3746,f8014])).
fof(f16155,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f9217,f3746,f8013])).
fof(f16153,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f6767,f3746,f8013])).
fof(f16151,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f5771,f3746,f8013])).
fof(f16149,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f5005,f3746,f8013])).
fof(f16147,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f514,f3746,f8013])).
fof(f16141,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f9136,f3746,f8012])).
fof(f8012,plain,(
  ( ! [X15,X16] : (~sP3(X15,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X16) | p(s(bool,X16)) | ~p(s(bool,X15))) ) | $spl3),
  inference(subsumption_resolution,[],[f7958,f269])).
fof(f7958,plain,(
  ( ! [X15,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X15)) | p(s(bool,X16)) | ~sP3(X15,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X16)) ) | $spl3),
  inference(superposition,[],[f405,f4933])).
fof(f16139,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f6706,f3746,f8012])).
fof(f16137,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f5728,f3746,f8012])).
fof(f16135,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f4966,f3746,f8012])).
fof(f16133,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f500,f3746,f8012])).
fof(f16127,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f9125,f3746,f8011])).
fof(f8011,plain,(
  ( ! [X14,X13] : (~sP4(X13,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X14) | p(s(bool,X14)) | ~p(s(bool,X13))) ) | $spl3),
  inference(subsumption_resolution,[],[f7957,f269])).
fof(f7957,plain,(
  ( ! [X14,X13] : (~p(s(bool,'T_0')) | ~p(s(bool,X13)) | p(s(bool,X14)) | ~sP4(X13,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X14)) ) | $spl3),
  inference(superposition,[],[f401,f4933])).
fof(f16125,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f6697,f3746,f8011])).
fof(f16123,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f5721,f3746,f8011])).
fof(f16121,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f4961,f3746,f8011])).
fof(f16119,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f495,f3746,f8011])).
fof(f16108,plain,(
  $false | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f5378,f3746,f7534])).
fof(f7534,plain,(
  ( ! [X35,X34] : (~sP11('const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),X35,X34) | p(s(bool,X35)) | ~p(s(bool,X34))) ) | $spl0),
  inference(subsumption_resolution,[],[f7482,f269])).
fof(f7482,plain,(
  ( ! [X35,X34] : (~p(s(bool,'T_0')) | ~p(s(bool,X34)) | p(s(bool,X35)) | ~sP11('const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),X35,X34)) ) | $spl0),
  inference(superposition,[],[f465,f5212])).
fof(f16105,plain,(
  $false | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5376,f3746,f7533])).
fof(f7533,plain,(
  ( ! [X33,X32] : (sP11('const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),X32,X33) | ~p(s(bool,X32))) ) | $spl0),
  inference(subsumption_resolution,[],[f7481,f269])).
fof(f7481,plain,(
  ( ! [X33,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X32)) | sP11('const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),X32,X33)) ) | $spl0),
  inference(superposition,[],[f464,f5212])).
fof(f16103,plain,(
  $false | ($spl0 | $spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f9391,f3746,f7531])).
fof(f7531,plain,(
  ( ! [X28,X29] : (~sP10(X29,X28,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15))) | ~p(s(bool,X29)) | p(s(bool,X28))) ) | $spl0),
  inference(subsumption_resolution,[],[f7479,f269])).
fof(f7479,plain,(
  ( ! [X28,X29] : (~p(s(bool,'T_0')) | p(s(bool,X28)) | ~p(s(bool,X29)) | ~sP10(X29,X28,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)))) ) | $spl0),
  inference(superposition,[],[f458,f5212])).
fof(f16101,plain,(
  $false | ($spl0 | $spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f6901,f3746,f7531])).
fof(f16099,plain,(
  $false | ($spl0 | $spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f5877,f3746,f7531])).
fof(f16097,plain,(
  $false | ($spl0 | $spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f5087,f3746,f7531])).
fof(f16095,plain,(
  $false | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f526,f3746,f7531])).
fof(f16089,plain,(
  $false | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5335,f3746,f7529])).
fof(f7529,plain,(
  ( ! [X24,X25] : (sP9('const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),X24,X25) | ~p(s(bool,X24))) ) | $spl0),
  inference(subsumption_resolution,[],[f7477,f269])).
fof(f7477,plain,(
  ( ! [X24,X25] : (~p(s(bool,'T_0')) | ~p(s(bool,X24)) | sP9('const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),X24,X25)) ) | $spl0),
  inference(superposition,[],[f447,f5212])).
fof(f16087,plain,(
  $false | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5335,f3746,f7528])).
fof(f7528,plain,(
  ( ! [X23,X22] : (sP9('const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),X23,X22) | ~p(s(bool,X22))) ) | $spl0),
  inference(subsumption_resolution,[],[f7476,f269])).
fof(f7476,plain,(
  ( ! [X23,X22] : (~p(s(bool,'T_0')) | ~p(s(bool,X22)) | sP9('const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),X23,X22)) ) | $spl0),
  inference(superposition,[],[f446,f5212])).
fof(f16076,plain,(
  $false | ($spl0 | $spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f9227,f3746,f7526])).
fof(f7526,plain,(
  ( ! [X19,X18] : (~sP7(X19,X18,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15))) | p(s(bool,X19)) | ~p(s(bool,X18))) ) | $spl0),
  inference(subsumption_resolution,[],[f7474,f269])).
fof(f7474,plain,(
  ( ! [X19,X18] : (~p(s(bool,'T_0')) | ~p(s(bool,X18)) | p(s(bool,X19)) | ~sP7(X19,X18,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)))) ) | $spl0),
  inference(superposition,[],[f433,f5212])).
fof(f16074,plain,(
  $false | ($spl0 | $spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f6776,f3746,f7526])).
fof(f16072,plain,(
  $false | ($spl0 | $spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f5778,f3746,f7526])).
fof(f16070,plain,(
  $false | ($spl0 | $spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f5011,f3746,f7526])).
fof(f16068,plain,(
  $false | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f648,f3746,f7526])).
fof(f16063,plain,(
  $false | ($spl0 | $spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f9217,f3746,f7525])).
fof(f7525,plain,(
  ( ! [X17,X16] : (~sP6(X16,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),X17) | p(s(bool,X17)) | ~p(s(bool,X16))) ) | $spl0),
  inference(subsumption_resolution,[],[f7473,f269])).
fof(f7473,plain,(
  ( ! [X17,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X16)) | p(s(bool,X17)) | ~sP6(X16,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),X17)) ) | $spl0),
  inference(superposition,[],[f424,f5212])).
fof(f16061,plain,(
  $false | ($spl0 | $spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f6767,f3746,f7525])).
fof(f16059,plain,(
  $false | ($spl0 | $spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f5771,f3746,f7525])).
fof(f16057,plain,(
  $false | ($spl0 | $spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f5005,f3746,f7525])).
fof(f16055,plain,(
  $false | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f514,f3746,f7525])).
fof(f16049,plain,(
  $false | ($spl0 | $spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f9136,f3746,f7524])).
fof(f7524,plain,(
  ( ! [X14,X15] : (~sP3(X14,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),X15) | p(s(bool,X15)) | ~p(s(bool,X14))) ) | $spl0),
  inference(subsumption_resolution,[],[f7472,f269])).
fof(f7472,plain,(
  ( ! [X14,X15] : (~p(s(bool,'T_0')) | ~p(s(bool,X14)) | p(s(bool,X15)) | ~sP3(X14,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),X15)) ) | $spl0),
  inference(superposition,[],[f405,f5212])).
fof(f16047,plain,(
  $false | ($spl0 | $spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f6706,f3746,f7524])).
fof(f16045,plain,(
  $false | ($spl0 | $spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f5728,f3746,f7524])).
fof(f16043,plain,(
  $false | ($spl0 | $spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f4966,f3746,f7524])).
fof(f16041,plain,(
  $false | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f500,f3746,f7524])).
fof(f16035,plain,(
  $false | ($spl0 | $spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f9125,f3746,f7523])).
fof(f7523,plain,(
  ( ! [X12,X13] : (~sP4(X12,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),X13) | p(s(bool,X13)) | ~p(s(bool,X12))) ) | $spl0),
  inference(subsumption_resolution,[],[f7471,f269])).
fof(f7471,plain,(
  ( ! [X12,X13] : (~p(s(bool,'T_0')) | ~p(s(bool,X12)) | p(s(bool,X13)) | ~sP4(X12,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),X13)) ) | $spl0),
  inference(superposition,[],[f401,f5212])).
fof(f16033,plain,(
  $false | ($spl0 | $spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f6697,f3746,f7523])).
fof(f16031,plain,(
  $false | ($spl0 | $spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f5721,f3746,f7523])).
fof(f16029,plain,(
  $false | ($spl0 | $spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f4961,f3746,f7523])).
fof(f16027,plain,(
  $false | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f495,f3746,f7523])).
fof(f16020,plain,(
  $false | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f6523,f3746,f7520])).
fof(f16013,plain,(
  $false | ($spl38 | $spl46)),
  inference(unit_resulting_resolution,[],[f696,f483,f3746,f2416])).
fof(f16004,plain,(
  $false | ($spl36 | $spl46)),
  inference(unit_resulting_resolution,[],[f1185,f3746,f2413])).
fof(f16002,plain,(
  $false | ($spl32 | $spl46)),
  inference(unit_resulting_resolution,[],[f526,f483,f3746,f2407])).
fof(f15992,plain,(
  $false | ($spl28 | $spl46)),
  inference(unit_resulting_resolution,[],[f1085,f3746,f2401])).
fof(f15990,plain,(
  $false | ($spl26 | $spl46)),
  inference(unit_resulting_resolution,[],[f1085,f3746,f2398])).
fof(f15988,plain,(
  $false | ($spl22 | $spl46)),
  inference(unit_resulting_resolution,[],[f648,f483,f3746,f2392])).
fof(f15969,plain,(
  $false | ($spl20 | $spl46)),
  inference(unit_resulting_resolution,[],[f514,f483,f3746,f2389])).
fof(f15959,plain,(
  $false | ($spl18 | $spl46)),
  inference(unit_resulting_resolution,[],[f500,f483,f3746,f2386])).
fof(f15949,plain,(
  $false | ($spl16 | $spl46)),
  inference(unit_resulting_resolution,[],[f495,f483,f3746,f2383])).
fof(f15937,plain,(
  $false | ($spl12 | $spl46)),
  inference(unit_resulting_resolution,[],[f859,f483,f3746,f2376])).
fof(f15914,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f5905,f3746,f1221])).
fof(f15912,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f5111,f3746,f1221])).
fof(f15910,plain,(
  $false | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f852,f3746,f1221])).
fof(f15907,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f9391,f3746,f1181])).
fof(f15905,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f6901,f3746,f1181])).
fof(f15903,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f5877,f3746,f1181])).
fof(f15901,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f5087,f3746,f1181])).
fof(f15899,plain,(
  $false | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f526,f3746,f1181])).
fof(f15887,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f9227,f3746,f1031])).
fof(f15885,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f6776,f3746,f1031])).
fof(f15883,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f5778,f3746,f1031])).
fof(f15881,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f5011,f3746,f1031])).
fof(f15879,plain,(
  $false | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f648,f3746,f1031])).
fof(f15876,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f9217,f3746,f1003])).
fof(f15874,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f6767,f3746,f1003])).
fof(f15872,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f5771,f3746,f1003])).
fof(f15870,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f5005,f3746,f1003])).
fof(f15868,plain,(
  $false | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f514,f3746,f1003])).
fof(f15862,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f9136,f3746,f953])).
fof(f953,plain,(
  ( ! [X2,X3] : (~sP3(X2,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),X3) | p(s(bool,X3)) | ~p(s(bool,X2))) ) | $spl4),
  inference(subsumption_resolution,[],[f947,f269])).
fof(f947,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | ~p(s(bool,X2)) | p(s(bool,X3)) | ~sP3(X2,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),X3)) ) | $spl4),
  inference(superposition,[],[f405,f531])).
fof(f15860,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f6706,f3746,f953])).
fof(f15858,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f5728,f3746,f953])).
fof(f15856,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f4966,f3746,f953])).
fof(f15854,plain,(
  $false | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f500,f3746,f953])).
fof(f15848,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f9125,f3746,f933])).
fof(f933,plain,(
  ( ! [X2,X3] : (~sP4(X2,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),X3) | p(s(bool,X3)) | ~p(s(bool,X2))) ) | $spl4),
  inference(subsumption_resolution,[],[f927,f269])).
fof(f927,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | ~p(s(bool,X2)) | p(s(bool,X3)) | ~sP4(X2,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),X3)) ) | $spl4),
  inference(superposition,[],[f401,f531])).
fof(f15846,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f6697,f3746,f933])).
fof(f15844,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f5721,f3746,f933])).
fof(f15842,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f4961,f3746,f933])).
fof(f15840,plain,(
  $false | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f495,f3746,f933])).
fof(f15833,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f5712,f3746,f915])).
fof(f15830,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f4956,f3746,f915])).
fof(f15825,plain,(
  $false | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f1330,f3746,f915])).
fof(f15819,plain,(
  $false | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f1542,f3746,f711])).
fof(f711,plain,(
  ( ! [X2,X3] : (sP11('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),X2,X3) | ~p(s(bool,X2))) ) | $spl4),
  inference(subsumption_resolution,[],[f709,f269])).
fof(f709,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | ~p(s(bool,X2)) | sP11('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),X2,X3)) ) | $spl4),
  inference(superposition,[],[f464,f531])).
fof(f15817,plain,(
  $false | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f1454,f3746,f677])).
fof(f677,plain,(
  ( ! [X2,X3] : (sP9('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),X2,X3) | ~p(s(bool,X2))) ) | $spl4),
  inference(subsumption_resolution,[],[f675,f269])).
fof(f675,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | ~p(s(bool,X2)) | sP9('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),X2,X3)) ) | $spl4),
  inference(superposition,[],[f447,f531])).
fof(f15815,plain,(
  $false | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f1454,f3746,f666])).
fof(f666,plain,(
  ( ! [X2,X3] : (sP9('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),X3,X2) | ~p(s(bool,X2))) ) | $spl4),
  inference(subsumption_resolution,[],[f664,f269])).
fof(f664,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | ~p(s(bool,X2)) | sP9('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),X3,X2)) ) | $spl4),
  inference(superposition,[],[f446,f531])).
fof(f15813,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9237,f9311,f3746,f469])).
fof(f9311,plain,(
  ( ! [X0,X1] : (sP8('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))),X0,X1)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f5685,f444])).
fof(f15811,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6785,f6837,f3746,f469])).
fof(f6837,plain,(
  ( ! [X0,X1] : (sP8('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),X0,X1)) ) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f5683,f444])).
fof(f15809,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5785,f5825,f3746,f469])).
fof(f5825,plain,(
  ( ! [X0,X1] : (sP8('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),X0,X1)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f4934,f444])).
fof(f15807,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5017,f5045,f3746,f469])).
fof(f5045,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),X0,X1)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f480,f444])).
fof(f15805,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f654,f521,f3746,f469])).
fof(f15803,plain,(
  $false | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5309,f5324,f3746,f469])).
fof(f15801,plain,(
  $false | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f654,f520,f3746,f469])).
fof(f520,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f481,f443])).
fof(f15799,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f657,f518,f3746,f469])).
fof(f15797,plain,(
  $false | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5304,f5323,f3746,f469])).
fof(f15795,plain,(
  $false | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f657,f517,f3746,f469])).
fof(f517,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f481,f442])).
fof(f15793,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f654,f515,f3746,f469])).
fof(f15764,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9336,f9391,f3746,f468])).
fof(f15762,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6860,f6901,f3746,f468])).
fof(f15760,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5848,f5877,f3746,f468])).
fof(f15758,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5068,f5087,f3746,f468])).
fof(f15756,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f686,f526,f3746,f468])).
fof(f15754,plain,(
  $false | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5336,f5363,f3746,f468])).
fof(f15752,plain,(
  $false | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f686,f525,f3746,f468])).
fof(f15750,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f688,f523,f3746,f468])).
fof(f15748,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9338,f9390,f3746,f468])).
fof(f15746,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6862,f6900,f3746,f468])).
fof(f15744,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5850,f5876,f3746,f468])).
fof(f15742,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5070,f5086,f3746,f468])).
fof(f15740,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f686,f522,f3746,f468])).
fof(f15707,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f696,f483,f3746,f3746,f465])).
fof(f15673,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f1159,f483,f269,f3746,f465])).
fof(f15635,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f696,f483,f3746,f3746,f465])).
fof(f15601,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f696,f483,f269,f3746,f465])).
fof(f15577,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f1185,f3746,f3746,f464])).
fof(f15575,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f688,f269,f3746,f464])).
fof(f15542,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f1185,f3746,f3746,f464])).
fof(f15540,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f1185,f269,f3746,f464])).
fof(f15507,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f686,f483,f3746,f463])).
fof(f15474,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f686,f483,f483,f3746,f462])).
fof(f15446,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f1173,f3746,f460])).
fof(f15444,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f522,f3746,f483,f3746,f458])).
fof(f15442,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f522,f269,f483,f3746,f458])).
fof(f15409,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f522,f3746,f483,f3746,f458])).
fof(f15407,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f526,f269,f483,f3746,f458])).
fof(f15374,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9336,f9391,f483,f3746,f457])).
fof(f15372,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6860,f6901,f483,f3746,f457])).
fof(f15370,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5848,f5877,f483,f3746,f457])).
fof(f15368,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5068,f5087,f483,f3746,f457])).
fof(f15366,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f686,f526,f483,f3746,f457])).
fof(f15364,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9338,f9390,f483,f3746,f457])).
fof(f15362,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6862,f6900,f483,f3746,f457])).
fof(f15360,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5850,f5876,f483,f3746,f457])).
fof(f15358,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5070,f5086,f483,f3746,f457])).
fof(f15356,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f686,f522,f483,f3746,f457])).
fof(f15274,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9336,f9391,f483,f3746,f456])).
fof(f15272,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6860,f6901,f483,f3746,f456])).
fof(f15270,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5848,f5877,f483,f3746,f456])).
fof(f15268,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5068,f5087,f483,f3746,f456])).
fof(f15266,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f686,f526,f483,f3746,f456])).
fof(f15264,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9338,f9390,f483,f3746,f456])).
fof(f15262,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6862,f6900,f483,f3746,f456])).
fof(f15260,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5850,f5876,f483,f3746,f456])).
fof(f15258,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5070,f5086,f483,f3746,f456])).
fof(f15256,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f686,f522,f483,f3746,f456])).
fof(f15174,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9336,f9391,f3746,f3746,f455])).
fof(f15172,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6860,f6901,f3746,f3746,f455])).
fof(f15170,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5848,f5877,f3746,f3746,f455])).
fof(f15168,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5068,f5087,f3746,f3746,f455])).
fof(f15166,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f686,f526,f3746,f3746,f455])).
fof(f15164,plain,(
  $false | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5336,f5363,f3746,f3746,f455])).
fof(f15162,plain,(
  $false | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f686,f525,f3746,f3746,f455])).
fof(f15160,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f688,f523,f3746,f3746,f455])).
fof(f15158,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9338,f9390,f3746,f3746,f455])).
fof(f15156,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6862,f6900,f3746,f3746,f455])).
fof(f15154,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5850,f5876,f3746,f3746,f455])).
fof(f15152,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5070,f5086,f3746,f3746,f455])).
fof(f15150,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f686,f522,f3746,f3746,f455])).
fof(f15016,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5877,f5856,f3746,f3746,f455])).
fof(f15014,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5877,f5851,f3746,f3746,f455])).
fof(f15012,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5877,f5848,f3746,f3746,f455])).
fof(f15010,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6901,f6868,f3746,f3746,f455])).
fof(f15008,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6901,f6860,f3746,f3746,f455])).
fof(f15005,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9391,f9348,f3746,f3746,f455])).
fof(f15003,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9391,f9336,f3746,f3746,f455])).
fof(f14996,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5087,f5076,f3746,f3746,f455])).
fof(f14994,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5087,f5071,f3746,f3746,f455])).
fof(f14992,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5087,f5068,f3746,f3746,f455])).
fof(f14990,plain,(
  $false | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f526,f5336,f3746,f3746,f455])).
fof(f14988,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f526,f688,f3746,f3746,f455])).
fof(f14986,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f526,f9338,f3746,f3746,f455])).
fof(f14984,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f526,f6862,f3746,f3746,f455])).
fof(f14982,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f526,f5850,f3746,f3746,f455])).
fof(f14980,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f526,f5070,f3746,f3746,f455])).
fof(f14978,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f526,f686,f3746,f3746,f455])).
fof(f14955,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5877,f5856,f269,f3746,f455])).
fof(f14953,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6901,f6868,f269,f3746,f455])).
fof(f14950,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9391,f9348,f269,f3746,f455])).
fof(f14944,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5087,f5076,f269,f3746,f455])).
fof(f14942,plain,(
  $false | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f526,f5336,f269,f3746,f455])).
fof(f14940,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f526,f688,f269,f3746,f455])).
fof(f14873,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f1159,f483,f3746,f452])).
fof(f14839,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f1058,f483,f483,f3746,f448])).
fof(f1058,plain,(
  sP9('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f521,f483,f483,f440])).
fof(f440,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f252])).
fof(f14793,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f1085,f3746,f3746,f447])).
fof(f14791,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f657,f269,f3746,f447])).
fof(f14762,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f1085,f3746,f3746,f447])).
fof(f14760,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f1085,f269,f3746,f447])).
fof(f14731,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f1085,f3746,f3746,f446])).
fof(f14729,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f654,f269,f3746,f446])).
fof(f14700,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f1085,f3746,f3746,f446])).
fof(f14698,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f1085,f269,f3746,f446])).
fof(f14669,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f1071,f3746,f443])).
fof(f14667,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f1071,f3746,f442])).
fof(f14665,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f521,f483,f483,f3746,f441])).
fof(f14647,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9237,f9311,f483,f3746,f439])).
fof(f14645,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6785,f6837,f483,f3746,f439])).
fof(f14643,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5785,f5825,f483,f3746,f439])).
fof(f14641,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5017,f5045,f483,f3746,f439])).
fof(f14639,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f654,f521,f483,f3746,f439])).
fof(f14637,plain,(
  $false | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5304,f5323,f483,f3746,f439])).
fof(f14634,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f654,f515,f483,f3746,f439])).
fof(f14578,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5825,f5786,f483,f3746,f438])).
fof(f14576,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6837,f6786,f483,f3746,f438])).
fof(f14573,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9311,f9238,f483,f3746,f438])).
fof(f14567,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5045,f5018,f483,f3746,f438])).
fof(f14565,plain,(
  $false | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f521,f5309,f483,f3746,f438])).
fof(f14563,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f521,f657,f483,f3746,f438])).
fof(f14513,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f1058,f483,f3746,f435])).
fof(f14467,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f1058,f483,f3746,f434])).
fof(f14421,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f641,f483,f3746,f3746,f433])).
fof(f14386,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f641,f483,f269,f3746,f433])).
fof(f14357,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f641,f483,f3746,f3746,f433])).
fof(f14322,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f648,f483,f269,f3746,f433])).
fof(f14293,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f634,f3746,f3746,f3746,f430])).
fof(f14291,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f636,f269,f3746,f3746,f430])).
fof(f14266,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f634,f3746,f269,f3746,f430])).
fof(f14264,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f1007,f269,f269,f3746,f430])).
fof(f14213,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f634,f3746,f3746,f3746,f430])).
fof(f14211,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f634,f269,f3746,f3746,f430])).
fof(f14186,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f634,f3746,f269,f3746,f430])).
fof(f14184,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f634,f269,f269,f3746,f430])).
fof(f14134,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f634,f3746,f3746,f3746,f430])).
fof(f14132,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f634,f269,f3746,f3746,f430])).
fof(f14107,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f636,f3746,f269,f3746,f430])).
fof(f14105,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f636,f269,f269,f3746,f430])).
fof(f14059,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f995,f3746,f425])).
fof(f14057,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f513,f483,f3746,f3746,f424])).
fof(f14042,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f513,f483,f269,f3746,f424])).
fof(f14022,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f513,f483,f3746,f3746,f424])).
fof(f14007,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f514,f483,f269,f3746,f424])).
fof(f13984,plain,(
  $false | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f514,f5262,f3746,f3746,f420])).
fof(f13982,plain,(
  $false | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f513,f5261,f3746,f3746,f420])).
fof(f13972,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9217,f9198,f3746,f3746,f420])).
fof(f13970,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6767,f6752,f3746,f3746,f420])).
fof(f13968,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5771,f5762,f3746,f3746,f420])).
fof(f13966,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5005,f4999,f3746,f3746,f420])).
fof(f13964,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f514,f636,f3746,f3746,f420])).
fof(f13962,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9216,f9189,f3746,f3746,f420])).
fof(f13960,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6766,f6745,f3746,f3746,f420])).
fof(f13958,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5770,f5759,f3746,f3746,f420])).
fof(f13956,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5004,f4997,f3746,f3746,f420])).
fof(f13954,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f513,f634,f3746,f3746,f420])).
fof(f13936,plain,(
  $false | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f513,f5261,f269,f3746,f420])).
fof(f13928,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9216,f9189,f269,f3746,f420])).
fof(f13926,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6766,f6745,f269,f3746,f420])).
fof(f13924,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5770,f5759,f269,f3746,f420])).
fof(f13922,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5004,f4997,f269,f3746,f420])).
fof(f13920,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f513,f634,f269,f3746,f420])).
fof(f13859,plain,(
  $false | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f514,f5262,f3746,f3746,f420])).
fof(f13857,plain,(
  $false | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f513,f5261,f3746,f3746,f420])).
fof(f13847,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9217,f9198,f3746,f3746,f420])).
fof(f13845,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6767,f6752,f3746,f3746,f420])).
fof(f13843,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5771,f5762,f3746,f3746,f420])).
fof(f13841,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5005,f4999,f3746,f3746,f420])).
fof(f13839,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f514,f636,f3746,f3746,f420])).
fof(f13837,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9216,f9189,f3746,f3746,f420])).
fof(f13835,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6766,f6745,f3746,f3746,f420])).
fof(f13833,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5770,f5759,f3746,f3746,f420])).
fof(f13831,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5004,f4997,f3746,f3746,f420])).
fof(f13829,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f513,f634,f3746,f3746,f420])).
fof(f13812,plain,(
  $false | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f514,f5262,f269,f3746,f420])).
fof(f13807,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f9217,f9198,f269,f3746,f420])).
fof(f13805,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f6767,f6752,f269,f3746,f420])).
fof(f13803,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5771,f5762,f269,f3746,f420])).
fof(f13801,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5005,f4999,f269,f3746,f420])).
fof(f13799,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f514,f636,f269,f3746,f420])).
fof(f13741,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f648,f483,f3746,f419])).
fof(f13706,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f641,f483,f3746,f418])).
fof(f13671,plain,(
  $false | ($spl3 | $spl4 | $spl14 | $spl46)),
  inference(unit_resulting_resolution,[],[f7422,f6640,f3746,f415])).
fof(f13669,plain,(
  $false | ($spl3 | $spl14 | $spl46)),
  inference(unit_resulting_resolution,[],[f6327,f5686,f3746,f415])).
fof(f13667,plain,(
  $false | ($spl3 | $spl14 | $spl46)),
  inference(unit_resulting_resolution,[],[f5513,f4936,f3746,f415])).
fof(f13665,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f858,f527,f3746,f415])).
fof(f13663,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f7400,f6640,f3746,f415])).
fof(f13661,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f6312,f5686,f3746,f415])).
fof(f13659,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5500,f4936,f3746,f415])).
fof(f13657,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f859,f527,f3746,f415])).
fof(f13582,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f7432,f7047,f3746,f414])).
fof(f13580,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f6336,f5906,f3746,f414])).
fof(f13578,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f7935,f4933,f3746,f414])).
fof(f13576,plain,(
  $false | ($spl46 | $spl63)),
  inference(unit_resulting_resolution,[],[f4457,f4246,f3746,f414])).
fof(f13574,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5521,f5112,f3746,f414])).
fof(f13572,plain,(
  $false | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f6451,f5212,f3746,f414])).
fof(f13570,plain,(
  $false | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f893,f531,f3746,f414])).
fof(f13568,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f7440,f7047,f3746,f414])).
fof(f13566,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f6353,f5906,f3746,f414])).
fof(f13564,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f7933,f4933,f3746,f414])).
fof(f13562,plain,(
  $false | ($spl46 | $spl63)),
  inference(unit_resulting_resolution,[],[f4455,f4246,f3746,f414])).
fof(f13560,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5526,f5112,f3746,f414])).
fof(f13558,plain,(
  $false | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f7448,f5212,f3746,f414])).
fof(f13556,plain,(
  $false | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f890,f531,f3746,f414])).
fof(f16268,plain,(
  $false | ($spl3 | $spl4 | $spl18 | $spl46)),
  inference(trivial_inequality_removal,[],[f13552])).
fof(f13552,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl3 | $spl4 | $spl18 | $spl46)),
  inference(unit_resulting_resolution,[],[f7781,f3746,f414])).
fof(f7781,plain,(
  ~sP5('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),'T_0','T_0') | ($spl3 | $spl4 | $spl18)),
  inference(unit_resulting_resolution,[],[f6982,f395])).
fof(f6982,plain,(
  ~sP3('T_0','T_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14))) | ($spl3 | $spl4 | $spl18)),
  inference(unit_resulting_resolution,[],[f269,f5683,f2386])).
fof(f16267,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(trivial_inequality_removal,[],[f13551])).
fof(f13551,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f7445,f3746,f414])).
fof(f7445,plain,(
  ~sP5('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),'F_0','F_0') | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f6727,f394])).
fof(f6727,plain,(
  ~sP2('F_0','F_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK14))) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f483,f483,f5683,f409])).
fof(f16266,plain,(
  $false | ($spl3 | $spl46)),
  inference(trivial_inequality_removal,[],[f13548])).
fof(f13548,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f6358,f3746,f414])).
fof(f6358,plain,(
  ~sP5('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),'F_0','F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f5745,f394])).
fof(f5745,plain,(
  ~sP2('F_0','F_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | $spl3),
  inference(unit_resulting_resolution,[],[f483,f483,f4934,f409])).
fof(f16265,plain,(
  $false | ($spl3 | $spl18 | $spl46)),
  inference(trivial_inequality_removal,[],[f13545])).
fof(f13545,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl3 | $spl18 | $spl46)),
  inference(unit_resulting_resolution,[],[f7192,f3746,f414])).
fof(f7192,plain,(
  ~sP5('const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),'T_0','T_0') | ($spl3 | $spl18)),
  inference(unit_resulting_resolution,[],[f5945,f395])).
fof(f5945,plain,(
  ~sP3('T_0','T_0','const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | ($spl3 | $spl18)),
  inference(unit_resulting_resolution,[],[f269,f4934,f2386])).
fof(f16264,plain,(
  $false | ($spl3 | $spl46)),
  inference(trivial_inequality_removal,[],[f13541])).
fof(f13541,plain,(
  s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f7558,f3746,f414])).
fof(f7558,plain,(
  ~sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | $spl3),
  inference(unit_resulting_resolution,[],[f4986,f394])).
fof(f4986,plain,(
  ~sP2('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | $spl3),
  inference(unit_resulting_resolution,[],[f480,f480,f480,f409])).
fof(f16263,plain,(
  $false | ($spl3 | $spl46)),
  inference(trivial_inequality_removal,[],[f13536])).
fof(f13536,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5531,f3746,f414])).
fof(f5531,plain,(
  ~sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'F_0','F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f4983,f394])).
fof(f4983,plain,(
  ~sP2('F_0','F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | $spl3),
  inference(unit_resulting_resolution,[],[f483,f483,f480,f409])).
fof(f16262,plain,(
  $false | ($spl3 | $spl18 | $spl46)),
  inference(trivial_inequality_removal,[],[f13533])).
fof(f13533,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl3 | $spl18 | $spl46)),
  inference(unit_resulting_resolution,[],[f6183,f3746,f414])).
fof(f6183,plain,(
  ~sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'T_0','T_0') | ($spl3 | $spl18)),
  inference(unit_resulting_resolution,[],[f5146,f395])).
fof(f5146,plain,(
  ~sP3('T_0','T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | ($spl3 | $spl18)),
  inference(unit_resulting_resolution,[],[f269,f480,f2386])).
fof(f16261,plain,(
  $false | ($spl3 | $spl46)),
  inference(trivial_inequality_removal,[],[f13529])).
fof(f13529,plain,(
  s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f6203,f3746,f414])).
fof(f6203,plain,(
  ~sP5('F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | $spl3),
  inference(unit_resulting_resolution,[],[f4976,f394])).
fof(f4976,plain,(
  ~sP2('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),'F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f480,f483,f480,f409])).
fof(f16260,plain,(
  $false | $spl46),
  inference(trivial_inequality_removal,[],[f13521])).
fof(f13521,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl46),
  inference(unit_resulting_resolution,[],[f978,f3746,f414])).
fof(f978,plain,(
  ~sP5('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f965,f394])).
fof(f965,plain,(
  ~sP2('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f483,f483,f483,f409])).
fof(f16259,plain,(
  $false | $spl46),
  inference(trivial_inequality_removal,[],[f13516])).
fof(f13516,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl46),
  inference(unit_resulting_resolution,[],[f957,f3746,f414])).
fof(f957,plain,(
  ~sP5('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f945,f395])).
fof(f945,plain,(
  ~sP3('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f483,f269,f269,f405])).
fof(f13496,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f965,f3746,f412])).
fof(f13494,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f965,f3746,f411])).
fof(f13492,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f965,f3746,f410])).
fof(f13490,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f945,f3746,f406])).
fof(f13488,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f499,f483,f3746,f3746,f405])).
fof(f13473,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f499,f483,f269,f3746,f405])).
fof(f13453,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f499,f483,f3746,f3746,f405])).
fof(f13438,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f500,f483,f269,f3746,f405])).
fof(f13418,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f925,f3746,f402])).
fof(f925,plain,(
  ~sP4('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f483,f269,f269,f401])).
fof(f13416,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f494,f483,f3746,f3746,f401])).
fof(f13399,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f494,f483,f269,f3746,f401])).
fof(f13377,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f494,f483,f3746,f3746,f401])).
fof(f13360,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f495,f483,f269,f3746,f401])).
fof(f13328,plain,(
  $false | ($spl0 | $spl12 | $spl46)),
  inference(unit_resulting_resolution,[],[f504,f499,f494,f5405,f3746,f398])).
fof(f13323,plain,(
  $false | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5253,f499,f494,f5225,f3746,f398])).
fof(f13293,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f501,f499,f494,f9703,f3746,f398])).
fof(f13287,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f506,f5727,f5720,f5907,f3746,f398])).
fof(f13282,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f506,f4965,f4960,f5113,f3746,f398])).
fof(f13278,plain,(
  $false | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f5253,f499,f494,f5224,f3746,f398])).
fof(f13274,plain,(
  $false | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f506,f499,f494,f1062,f3746,f398])).
fof(f13269,plain,(
  $false | ($spl46 | $spl63)),
  inference(unit_resulting_resolution,[],[f501,f499,f494,f4238,f3746,f398])).
fof(f13260,plain,(
  $false | ($spl3 | $spl4 | $spl12 | $spl46)),
  inference(unit_resulting_resolution,[],[f504,f6705,f6696,f6977,f3746,f398])).
fof(f13256,plain,(
  $false | ($spl3 | $spl12 | $spl46)),
  inference(unit_resulting_resolution,[],[f504,f5727,f5720,f5941,f3746,f398])).
fof(f13252,plain,(
  $false | ($spl3 | $spl12 | $spl46)),
  inference(unit_resulting_resolution,[],[f504,f4965,f4960,f5143,f3746,f398])).
fof(f13248,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f504,f499,f494,f881,f3746,f398])).
fof(f12819,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f1282,f483,f3746,f3746,f397])).
fof(f12744,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f1282,f483,f269,f3746,f397])).
fof(f12703,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f1282,f483,f3746,f3746,f397])).
fof(f12628,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f859,f483,f269,f3746,f397])).
fof(f12588,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f562,f483,f3746,f366])).
fof(f562,plain,(
  ~sP1('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f483,f483,f361])).
fof(f361,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f217])).
fof(f12567,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f562,f483,f3746,f363])).
fof(f12546,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f565,f3746,f3746,f362])).
fof(f12544,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f565,f269,f3746,f362])).
fof(f12508,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f565,f3746,f3746,f362])).
fof(f12506,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f575,f269,f3746,f362])).
fof(f12470,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f589,f3746,f348])).
fof(f348,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f210])).
fof(f210,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f209])).
fof(f209,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f589,plain,(
  ~sP0('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f483,f483,f484])).
fof(f484,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f346,f483])).
fof(f346,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f210])).
fof(f12468,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f589,f3746,f347])).
fof(f347,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f210])).
fof(f12466,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f3746,f6640,f3746,f341])).
fof(f12464,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f3746,f5686,f3746,f341])).
fof(f12462,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f3746,f4936,f3746,f341])).
fof(f12460,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f3746,f527,f3746,f341])).
fof(f12458,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f269,f6640,f3746,f341])).
fof(f12456,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f269,f5686,f3746,f341])).
fof(f12454,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f269,f4936,f3746,f341])).
fof(f12452,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f269,f527,f3746,f341])).
fof(f12448,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f3746,f6640,f3746,f341])).
fof(f12446,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f3746,f5686,f3746,f341])).
fof(f12444,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f3746,f4936,f3746,f341])).
fof(f12442,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f3746,f527,f3746,f341])).
fof(f12440,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f269,f6640,f3746,f341])).
fof(f12438,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f269,f5686,f3746,f341])).
fof(f12436,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f269,f4936,f3746,f341])).
fof(f12434,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f269,f527,f3746,f341])).
fof(f12426,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f5112,f3746,f289])).
fof(f12424,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f527,f3746,f286])).
fof(f12422,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f483,f3746])).
fof(f12419,plain,(
  ( ! [X50,X51] : (sP1(X50,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | s(bool,'T_0') = s(bool,X51) | s(bool,X50) = s(bool,X51)) ) | $spl72),
  inference(cnf_transformation,[],[f12419_D])).
fof(f12419_D,plain,(
  ( ! [X50,X51] : (sP1(X50,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | s(bool,'T_0') = s(bool,X51) | s(bool,X50) = s(bool,X51)) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f3746,plain,(
  ( ! [X11] : (p(s(bool,X11))) ) | $spl46),
  inference(cnf_transformation,[],[f3746_D])).
fof(f3746_D,plain,(
  ( ! [X11] : (p(s(bool,X11))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f12418,plain,(
  ( ! [X50,X51,X49] : (p(s(bool,X49)) | sP1(X50,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | s(bool,X50) = s(bool,X51) | s(bool,'T_0') = s(bool,X51)) ) | $spl3),
  inference(subsumption_resolution,[],[f12386,f285])).
fof(f12386,plain,(
  ( ! [X50,X51,X49] : (p(s(bool,X49)) | sP1(X50,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | s(bool,X50) = s(bool,X51) | s(bool,'T_0') = s(bool,X51) | s(bool,'T_0') = s(bool,X49)) ) | $spl3),
  inference(superposition,[],[f2427,f6034])).
fof(f6034,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(subsumption_resolution,[],[f5976,f269])).
fof(f5976,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(superposition,[],[f4934,f772])).
fof(f772,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f276,f276])).
fof(f2427,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X26)) | sP1(X25,X26) | s(bool,X25) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(subsumption_resolution,[],[f2023,f269])).
fof(f2023,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'T_0')) | p(s(bool,X26)) | sP1(X25,X26) | s(bool,X25) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f363,f772])).
fof(f12416,plain,(
  ( ! [X43,X44] : (sP1(X43,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | s(bool,'T_0') = s(bool,X44) | s(bool,X43) = s(bool,X44)) ) | $spl70),
  inference(cnf_transformation,[],[f12416_D])).
fof(f12416_D,plain,(
  ( ! [X43,X44] : (sP1(X43,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | s(bool,'T_0') = s(bool,X44) | s(bool,X43) = s(bool,X44)) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f12415,plain,(
  ( ! [X43,X44,X42] : (p(s(bool,X42)) | sP1(X43,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | s(bool,X43) = s(bool,X44) | s(bool,'T_0') = s(bool,X44)) ) | $spl3),
  inference(subsumption_resolution,[],[f12383,f285])).
fof(f12383,plain,(
  ( ! [X43,X44,X42] : (p(s(bool,X42)) | sP1(X43,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | s(bool,X43) = s(bool,X44) | s(bool,'T_0') = s(bool,X44) | s(bool,'T_0') = s(bool,X42)) ) | $spl3),
  inference(superposition,[],[f2427,f5205])).
fof(f5205,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(subsumption_resolution,[],[f5174,f269])).
fof(f5174,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(superposition,[],[f480,f772])).
fof(f12414,plain,(
  ( ! [X35,X34] : (sP1(X34,'F_0') | s(bool,'T_0') = s(bool,X35) | s(bool,X34) = s(bool,X35)) ) | $spl68),
  inference(cnf_transformation,[],[f12414_D])).
fof(f12414_D,plain,(
  ( ! [X35,X34] : (sP1(X34,'F_0') | s(bool,'T_0') = s(bool,X35) | s(bool,X34) = s(bool,X35)) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f12413,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X33)) | sP1(X34,'F_0') | s(bool,X34) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) )),
  inference(subsumption_resolution,[],[f12379,f285])).
fof(f12379,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X33)) | sP1(X34,'F_0') | s(bool,X34) = s(bool,X35) | s(bool,'T_0') = s(bool,X35) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f2427,f276])).
fof(f4899,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),sK12)),s(fun(sK12,bool),sK14)))) | $spl64),
  inference(cnf_transformation,[],[f4899_D])).
fof(f4899_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),sK12)),s(fun(sK12,bool),sK14)))) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f4901,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),sK12)),s(fun(sK12,bool),sK13)))) | $spl66),
  inference(cnf_transformation,[],[f4901_D])).
fof(f4901_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),sK12)),s(fun(sK12,bool),sK13)))) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f4897,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),sK12)),s(fun(sK12,bool),sK13)))) | p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),sK12)),s(fun(sK12,bool),sK14)))) | $spl1),
  inference(subsumption_resolution,[],[f4838,f269])).
fof(f4838,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),sK12)),s(fun(sK12,bool),sK13)))) | p(s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),sK12)),s(fun(sK12,bool),sK14)))) | $spl1),
  inference(superposition,[],[f335,f3743])).
fof(f3743,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))))) | $spl1),
  inference(subsumption_resolution,[],[f3737,f483])).
fof(f3737,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK18(sK15,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))))) | $spl1),
  inference(superposition,[],[f1366,f599])).
fof(f599,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15))) | $spl1),
  inference(unit_resulting_resolution,[],[f478,f290])).
fof(f290,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f168])).
fof(f4226,plain,(
  ( ! [X4,X2,X3] : (s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,X4) | s(bool,X4) = s(bool,sK31(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK31(X2,X3,bool,bool))))) ) | $spl60),
  inference(cnf_transformation,[],[f4226_D])).
fof(f4226_D,plain,(
  ( ! [X4,X2,X3] : (s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,X4) | s(bool,X4) = s(bool,sK31(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK31(X2,X3,bool,bool))))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f4228,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | $spl63),
  inference(cnf_transformation,[],[f4228_D])).
fof(f4228_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl63),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl63])])).
fof(f4215,plain,(
  ( ! [X6,X7,X5] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X5) = s(fun(bool,bool),X6) | s(bool,'T_0') = s(bool,g(s(bool,sK31(X5,X6,bool,bool)))) | s(bool,X7) = s(bool,sK31(X5,X6,bool,bool)) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f1504,f772])).
fof(f1504,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK31(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK31(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f381,f276])).
fof(f381,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK31(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK31(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f226])).
fof(f226,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK31(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK31(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK31])],[f141])).
fof(f141,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f107])).
fof(f107,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f106])).
fof(f106,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.pred_set.UNION_SUBSET',aHL_EXT)).
fof(f4214,plain,(
  ( ! [X4,X2,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,g(s(bool,sK31(X2,X3,bool,bool)))) | s(bool,X4) = s(bool,sK31(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f1504,f772])).
fof(f3821,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl54),
  inference(cnf_transformation,[],[f3821_D])).
fof(f3821_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f3817,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK31(X0,X1,bool,bool))) ) | $spl50),
  inference(cnf_transformation,[],[f3817_D])).
fof(f3817_D,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK31(X0,X1,bool,bool))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f3824,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl59),
  inference(cnf_transformation,[],[f3824_D])).
fof(f3824_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl59),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl59])])).
fof(f3815,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X12) = s(fun(bool,bool),X13) | s(bool,'T_0') = s(bool,sK31(X12,X13,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f1503,f276])).
fof(f1503,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK31(X0,X1,X2,bool))) )),
  inference(superposition,[],[f381,f276])).
fof(f3822,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl56),
  inference(cnf_transformation,[],[f3822_D])).
fof(f3822_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f3814,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X10) = s(fun(bool,bool),X11) | s(bool,'T_0') = s(bool,sK31(X10,X11,bool,bool)) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f1503,f772])).
fof(f3813,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X7) = s(fun(bool,bool),X8) | s(bool,'T_0') = s(bool,sK31(X7,X8,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f1503,f772])).
fof(f3816,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl48),
  inference(cnf_transformation,[],[f3816_D])).
fof(f3816_D,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f3819,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl53),
  inference(cnf_transformation,[],[f3819_D])).
fof(f3819_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl53),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl53])])).
fof(f3812,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,sK31(X3,X4,bool,bool)) | s(bool,X5) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1503,f772])).
fof(f3811,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK31(X0,X1,bool,bool)) | s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1503,f772])).
fof(f3745,plain,(
  ( ! [X10,X8,X9] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X8,sK18(X10,X9,X8)),s(fun(X8,bool),X9))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(X8,bool),X9),s(fun(X8,bool),X10)))) ) | $spl44),
  inference(cnf_transformation,[],[f3745_D])).
fof(f3745_D,plain,(
  ( ! [X10,X8,X9] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X8,sK18(X10,X9,X8)),s(fun(X8,bool),X9))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(X8,bool),X9),s(fun(X8,bool),X10)))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f3747,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X15)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X12,sK18(X14,X13,X12)),s(fun(X12,bool),X13))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(X12,bool),X13),s(fun(X12,bool),X14)))) )),
  inference(subsumption_resolution,[],[f3741,f285])).
fof(f3741,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X15)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X12,sK18(X14,X13,X12)),s(fun(X12,bool),X13))) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(X12,bool),X13),s(fun(X12,bool),X14)))) )),
  inference(superposition,[],[f1366,f772])).
fof(f3744,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X11)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X8,sK18(X10,X9,X8)),s(fun(X8,bool),X9))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(X8,bool),X9),s(fun(X8,bool),X10)))) )),
  inference(subsumption_resolution,[],[f3740,f285])).
fof(f3740,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X11)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X8,sK18(X10,X9,X8)),s(fun(X8,bool),X9))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(X8,bool),X9),s(fun(X8,bool),X10))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f1366,f772])).
fof(f2624,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f2507,f269])).
fof(f2507,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(backward_demodulation,[],[f2370,f483])).
fof(f2623,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f268,f267,f266,f269,f483,f276,f286,f285,f290,f289,f313,f312,f318,f317,f316,f322,f320,f326,f324,f330,f328,f333,f331,f337,f336,f335,f341,f338,f348,f347,f484,f366,f363,f362,f361,f380,f379,f381,f400,f399,f398,f397,f396,f395,f394,f404,f403,f402,f401,f408,f407,f406,f405,f412,f411,f410,f409,f416,f415,f414,f413,f470,f420,f419,f418,f417,f427,f426,f425,f424,f433,f430,f429,f428,f440,f439,f438,f469,f436,f435,f434,f444,f443,f442,f441,f448,f447,f446,f445,f457,f456,f455,f468,f453,f452,f451,f461,f460,f459,f458,f465,f464,f463,f462,f485,f488,f491,f494,f495,f496,f499,f500,f501,f504,f507,f510,f513,f514,f515,f518,f521,f522,f523,f526,f527,f562,f563,f565,f575,f589,f634,f636,f641,f648,f654,f657,f662,f673,f686,f688,f696,f707,f772,f832,f859,f858,f881,f879,f925,f945,f957,f965,f978,f982,f995,f1007,f1058,f1071,f1085,f1159,f1173,f1185,f1237,f1282,f1338,f1350,f1364,f1366,f1377,f1390,f1401,f1449,f1461,f1503,f1504,f1517,f1518,f1532,f1533,f1562,f1563,f1573,f1585,f1586,f1795,f1879,f1880,f1881,f1883,f1884,f1885,f1886,f1887,f1888,f1889,f1890,f1891,f1892,f1893,f1894,f1895,f1975,f2368,f2372,f2375,f2379,f2382,f2385,f2388,f2391,f2394,f2397,f2400,f2403,f2406,f2409,f2412,f2415,f2418,f2421,f2009,f2010,f2011,f2423,f2424,f2425,f2426,f2427,f2428,f2429,f2430,f2431,f2432,f2433,f2434,f2435,f2436,f2437,f2438,f2439,f2440,f2441,f2442,f2443,f2444,f2445,f2446,f2447,f2448,f2449,f2450,f2451,f2452,f2453,f2454,f2455,f2456,f2457,f2458,f2459,f2460,f2461,f2073,f2074,f2075,f2077,f2078,f2079,f2080,f2081,f2082,f2083,f2084,f2085,f2086,f2087,f2088,f2089,f2469,f2471,f2473,f2476,f2478,f2480,f2482,f2484,f2486,f2488,f2490,f2492,f2494,f2496,f2498,f2500,f2502,f2504,f2370,f2622])).
fof(f2622,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(subsumption_resolution,[],[f2505,f2370])).
fof(f2505,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl8),
  inference(backward_demodulation,[],[f2370,f289])).
fof(f2089,plain,(
  ( ! [X282,X285,X283,X284] : (s(X284,f(s(bool,'T_0'))) != s(X284,g(s(bool,'T_0'))) | s(fun(bool,X284),X282) = s(fun(bool,X284),X283) | s(bool,X285) = s(bool,sK31(X282,X283,X284,bool)) | s(bool,'T_0') = s(bool,X285)) )),
  inference(superposition,[],[f381,f772])).
fof(f2088,plain,(
  ( ! [X280,X277,X279,X281,X278] : (~p(s(bool,happ(s(fun(X279,bool),happ(s(fun(bool,fun(X279,bool)),X278),s(bool,'T_0'))),s(X279,happ(s(fun(bool,X279),X277),s(bool,'T_0')))))) | p(s(bool,happ(s(fun(X279,bool),happ(s(fun(bool,fun(X279,bool)),X278),s(bool,X280))),s(X279,sK30(X280,X278,bool,X279))))) | s(bool,X281) = s(bool,sK29(X277,X278,bool,X279)) | s(bool,'T_0') = s(bool,X281)) )),
  inference(superposition,[],[f380,f772])).
fof(f2087,plain,(
  ( ! [X273,X275,X276,X272,X274] : (~p(s(bool,happ(s(fun(X273,bool),happ(s(fun(bool,fun(X273,bool)),X272),s(bool,'T_0'))),s(X273,X274)))) | p(s(bool,happ(s(fun(X273,bool),happ(s(fun(bool,fun(X273,bool)),X272),s(bool,X275))),s(X273,happ(s(fun(bool,X273),sK28(X272,bool,X273)),s(bool,X275)))))) | s(bool,X276) = s(bool,sK27(X272,bool,X273)) | s(bool,'T_0') = s(bool,X276)) )),
  inference(superposition,[],[f379,f772])).
fof(f2086,plain,(
  ( ! [X269,X271,X268,X270] : (~p(s(bool,happ(s(fun(bool,bool),X268),s(bool,'T_0')))) | p(s(bool,X269)) | p(s(bool,happ(s(fun(bool,bool),X268),s(bool,X270)))) | s(bool,X271) = s(bool,sK24(X268,X269,bool)) | s(bool,'T_0') = s(bool,X271)) )),
  inference(superposition,[],[f330,f772])).
fof(f2085,plain,(
  ( ! [X265,X267,X264,X266] : (~p(s(bool,happ(s(fun(bool,bool),X264),s(bool,'T_0')))) | p(s(bool,X265)) | p(s(bool,happ(s(fun(bool,bool),X264),s(bool,X266)))) | s(bool,X267) = s(bool,sK23(X264,X265,bool)) | s(bool,'T_0') = s(bool,X267)) )),
  inference(superposition,[],[f328,f772])).
fof(f2084,plain,(
  ( ! [X261,X263,X260,X262] : (~p(s(bool,happ(s(fun(bool,bool),X260),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X260),s(bool,X262)))) | ~p(s(bool,X261)) | s(bool,X263) = s(bool,sK20(X260,X261,bool)) | s(bool,'T_0') = s(bool,X263)) )),
  inference(superposition,[],[f322,f772])).
fof(f2083,plain,(
  ( ! [X257,X259,X256,X258] : (~p(s(bool,happ(s(fun(bool,bool),X256),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X256),s(bool,X258)))) | ~p(s(bool,X257)) | s(bool,X259) = s(bool,sK19(X256,X257,bool)) | s(bool,'T_0') = s(bool,X259)) )),
  inference(superposition,[],[f320,f772])).
fof(f2082,plain,(
  ( ! [X255,X253,X254] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X254)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(bool,bool),X254),s(fun(bool,bool),X253)))) | s(bool,X255) = s(bool,sK18(X253,X254,bool)) | s(bool,'T_0') = s(bool,X255)) )),
  inference(superposition,[],[f317,f772])).
fof(f2081,plain,(
  ( ! [X251,X252,X250] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X250)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(bool,bool),X251),s(fun(bool,bool),X250)))) | s(bool,X252) = s(bool,sK18(X250,X251,bool)) | s(bool,'T_0') = s(bool,X252)) )),
  inference(superposition,[],[f318,f772])).
fof(f2080,plain,(
  ( ! [X249,X248,X247] : (~p(s(bool,happ(s(fun(bool,bool),X247),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X247),s(bool,X248)))) | s(bool,X249) = s(bool,sK17(X247,bool)) | s(bool,'T_0') = s(bool,X249)) )),
  inference(superposition,[],[f313,f772])).
fof(f2079,plain,(
  ( ! [X246,X244,X245] : (~p(s(bool,happ(s(fun(bool,bool),X244),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X244),s(bool,X245)))) | s(bool,X246) = s(bool,sK16(X244,bool)) | s(bool,'T_0') = s(bool,X246)) )),
  inference(superposition,[],[f312,f772])).
fof(f2078,plain,(
  ( ! [X243,X241,X239,X242,X240] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(X239,fun(bool,bool)),X241),s(X239,sK29(X240,X241,X239,bool)))),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(X239,fun(bool,bool)),X241),s(X239,X242))),s(bool,sK30(X242,X241,X239,bool))))) | s(bool,X243) = s(bool,happ(s(fun(X239,bool),X240),s(X239,sK29(X240,X241,X239,bool)))) | s(bool,'T_0') = s(bool,X243)) )),
  inference(superposition,[],[f380,f772])).
fof(f2077,plain,(
  ( ! [X237,X235,X238,X236] : (s(bool,'T_0') != s(bool,f(s(X235,sK31(X236,X237,bool,X235)))) | s(fun(X235,bool),X236) = s(fun(X235,bool),X237) | s(bool,X238) = s(bool,g(s(X235,sK31(X236,X237,bool,X235)))) | s(bool,'T_0') = s(bool,X238)) )),
  inference(superposition,[],[f381,f772])).
fof(f2075,plain,(
  ( ! [X225,X223,X226,X224] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X224)))) | p(s(bool,'const.bool.IN_2'(s(bool,X223),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X225),s(fun(bool,bool),X224)))))) | s(bool,X223) = s(bool,X226) | s(bool,'T_0') = s(bool,X226)) )),
  inference(superposition,[],[f337,f772])).
fof(f2074,plain,(
  ( ! [X222,X220,X219,X221] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X220)))) | p(s(bool,'const.bool.IN_2'(s(bool,X219),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X220),s(fun(bool,bool),X221)))))) | s(bool,X219) = s(bool,X222) | s(bool,'T_0') = s(bool,X222)) )),
  inference(superposition,[],[f336,f772])).
fof(f2073,plain,(
  ( ! [X218,X216,X215,X217] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X216),s(fun(bool,bool),X217)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X215),s(fun(bool,bool),X216)))) | p(s(bool,'const.bool.IN_2'(s(bool,X215),s(fun(bool,bool),X217)))) | s(bool,X215) = s(bool,X218) | s(bool,'T_0') = s(bool,X218)) )),
  inference(superposition,[],[f335,f772])).
fof(f2461,plain,(
  ( ! [X191,X194,X192,X190,X193] : (p(s(bool,'const.bool.IN_2'(s(X190,X191),s(fun(X190,bool),X192)))) | p(s(bool,'const.bool.IN_2'(s(X190,X191),s(fun(X190,bool),X193)))) | s(bool,X194) = s(bool,'const.bool.IN_2'(s(X190,X191),s(fun(X190,bool),'const.pred_set.UNION_2'(s(fun(X190,bool),X192),s(fun(X190,bool),X193))))) | s(bool,'T_0') = s(bool,X194)) )),
  inference(subsumption_resolution,[],[f2066,f269])).
fof(f2066,plain,(
  ( ! [X191,X194,X192,X190,X193] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(X190,X191),s(fun(X190,bool),X192)))) | p(s(bool,'const.bool.IN_2'(s(X190,X191),s(fun(X190,bool),X193)))) | s(bool,X194) = s(bool,'const.bool.IN_2'(s(X190,X191),s(fun(X190,bool),'const.pred_set.UNION_2'(s(fun(X190,bool),X192),s(fun(X190,bool),X193))))) | s(bool,'T_0') = s(bool,X194)) )),
  inference(superposition,[],[f335,f772])).
fof(f2460,plain,(
  ( ! [X189,X187,X185,X188,X186] : (p(s(bool,'const.bool.IN_2'(s(X185,X186),s(fun(X185,bool),'const.pred_set.UNION_2'(s(fun(X185,bool),X187),s(fun(X185,bool),X188)))))) | s(bool,X189) = s(bool,'const.bool.IN_2'(s(X185,X186),s(fun(X185,bool),X187))) | s(bool,'T_0') = s(bool,X189)) )),
  inference(subsumption_resolution,[],[f2065,f269])).
fof(f2065,plain,(
  ( ! [X189,X187,X185,X188,X186] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(X185,X186),s(fun(X185,bool),'const.pred_set.UNION_2'(s(fun(X185,bool),X187),s(fun(X185,bool),X188)))))) | s(bool,X189) = s(bool,'const.bool.IN_2'(s(X185,X186),s(fun(X185,bool),X187))) | s(bool,'T_0') = s(bool,X189)) )),
  inference(superposition,[],[f336,f772])).
fof(f2459,plain,(
  ( ! [X182,X180,X184,X183,X181] : (p(s(bool,'const.bool.IN_2'(s(X180,X181),s(fun(X180,bool),'const.pred_set.UNION_2'(s(fun(X180,bool),X183),s(fun(X180,bool),X182)))))) | s(bool,X184) = s(bool,'const.bool.IN_2'(s(X180,X181),s(fun(X180,bool),X182))) | s(bool,'T_0') = s(bool,X184)) )),
  inference(subsumption_resolution,[],[f2064,f269])).
fof(f2064,plain,(
  ( ! [X182,X180,X184,X183,X181] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(X180,X181),s(fun(X180,bool),'const.pred_set.UNION_2'(s(fun(X180,bool),X183),s(fun(X180,bool),X182)))))) | s(bool,X184) = s(bool,'const.bool.IN_2'(s(X180,X181),s(fun(X180,bool),X182))) | s(bool,'T_0') = s(bool,X184)) )),
  inference(superposition,[],[f337,f772])).
fof(f2458,plain,(
  ( ! [X177,X178,X176,X179] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X176,bool),X178),s(fun(X176,bool),X177)))) | s(bool,X179) = s(bool,'const.bool.IN_2'(s(X176,sK18(X177,X178,X176)),s(fun(X176,bool),X177))) | s(bool,'T_0') = s(bool,X179)) )),
  inference(subsumption_resolution,[],[f2063,f269])).
fof(f2063,plain,(
  ( ! [X177,X178,X176,X179] : (~p(s(bool,'T_0')) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X176,bool),X178),s(fun(X176,bool),X177)))) | s(bool,X179) = s(bool,'const.bool.IN_2'(s(X176,sK18(X177,X178,X176)),s(fun(X176,bool),X177))) | s(bool,'T_0') = s(bool,X179)) )),
  inference(superposition,[],[f318,f772])).
fof(f2457,plain,(
  ( ! [X167,X171,X169,X166,X170,X168] : (p(s(bool,happ(s(fun(X166,bool),happ(s(fun(X167,fun(X166,bool)),X168),s(X167,X170))),s(X166,sK30(X170,X168,X167,X166))))) | s(bool,X171) = s(bool,happ(s(fun(X166,bool),happ(s(fun(X167,fun(X166,bool)),X168),s(X167,sK29(X169,X168,X167,X166)))),s(X166,happ(s(fun(X167,X166),X169),s(X167,sK29(X169,X168,X167,X166)))))) | s(bool,'T_0') = s(bool,X171)) )),
  inference(subsumption_resolution,[],[f2061,f269])).
fof(f2061,plain,(
  ( ! [X167,X171,X169,X166,X170,X168] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X166,bool),happ(s(fun(X167,fun(X166,bool)),X168),s(X167,X170))),s(X166,sK30(X170,X168,X167,X166))))) | s(bool,X171) = s(bool,happ(s(fun(X166,bool),happ(s(fun(X167,fun(X166,bool)),X168),s(X167,sK29(X169,X168,X167,X166)))),s(X166,happ(s(fun(X167,X166),X169),s(X167,sK29(X169,X168,X167,X166)))))) | s(bool,'T_0') = s(bool,X171)) )),
  inference(superposition,[],[f380,f772])).
fof(f2456,plain,(
  ( ! [X165,X163,X161,X164,X162,X160] : (p(s(bool,happ(s(fun(X160,bool),happ(s(fun(X161,fun(X160,bool)),X162),s(X161,X164))),s(X160,happ(s(fun(X161,X160),sK28(X162,X161,X160)),s(X161,X164)))))) | s(bool,X165) = s(bool,happ(s(fun(X160,bool),happ(s(fun(X161,fun(X160,bool)),X162),s(X161,sK27(X162,X161,X160)))),s(X160,X163))) | s(bool,'T_0') = s(bool,X165)) )),
  inference(subsumption_resolution,[],[f2060,f269])).
fof(f2060,plain,(
  ( ! [X165,X163,X161,X164,X162,X160] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X160,bool),happ(s(fun(X161,fun(X160,bool)),X162),s(X161,X164))),s(X160,happ(s(fun(X161,X160),sK28(X162,X161,X160)),s(X161,X164)))))) | s(bool,X165) = s(bool,happ(s(fun(X160,bool),happ(s(fun(X161,fun(X160,bool)),X162),s(X161,sK27(X162,X161,X160)))),s(X160,X163))) | s(bool,'T_0') = s(bool,X165)) )),
  inference(superposition,[],[f379,f772])).
fof(f2455,plain,(
  ( ! [X158,X156,X159,X157,X155] : (p(s(bool,X157)) | p(s(bool,happ(s(fun(X155,bool),X156),s(X155,X158)))) | s(bool,X159) = s(bool,happ(s(fun(X155,bool),X156),s(X155,sK24(X156,X157,X155)))) | s(bool,'T_0') = s(bool,X159)) )),
  inference(subsumption_resolution,[],[f2059,f269])).
fof(f2059,plain,(
  ( ! [X158,X156,X159,X157,X155] : (~p(s(bool,'T_0')) | p(s(bool,X157)) | p(s(bool,happ(s(fun(X155,bool),X156),s(X155,X158)))) | s(bool,X159) = s(bool,happ(s(fun(X155,bool),X156),s(X155,sK24(X156,X157,X155)))) | s(bool,'T_0') = s(bool,X159)) )),
  inference(superposition,[],[f330,f772])).
fof(f2454,plain,(
  ( ! [X154,X152,X151,X153,X150] : (p(s(bool,X152)) | p(s(bool,happ(s(fun(X150,bool),X151),s(X150,X153)))) | s(bool,X154) = s(bool,happ(s(fun(X150,bool),X151),s(X150,sK23(X151,X152,X150)))) | s(bool,'T_0') = s(bool,X154)) )),
  inference(subsumption_resolution,[],[f2058,f269])).
fof(f2058,plain,(
  ( ! [X154,X152,X151,X153,X150] : (~p(s(bool,'T_0')) | p(s(bool,X152)) | p(s(bool,happ(s(fun(X150,bool),X151),s(X150,X153)))) | s(bool,X154) = s(bool,happ(s(fun(X150,bool),X151),s(X150,sK23(X151,X152,X150)))) | s(bool,'T_0') = s(bool,X154)) )),
  inference(superposition,[],[f328,f772])).
fof(f2453,plain,(
  ( ! [X146,X149,X147,X145,X148] : (p(s(bool,happ(s(fun(X145,bool),X146),s(X145,X148)))) | ~p(s(bool,X147)) | s(bool,X149) = s(bool,happ(s(fun(X145,bool),X146),s(X145,sK20(X146,X147,X145)))) | s(bool,'T_0') = s(bool,X149)) )),
  inference(subsumption_resolution,[],[f2057,f269])).
fof(f2057,plain,(
  ( ! [X146,X149,X147,X145,X148] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X145,bool),X146),s(X145,X148)))) | ~p(s(bool,X147)) | s(bool,X149) = s(bool,happ(s(fun(X145,bool),X146),s(X145,sK20(X146,X147,X145)))) | s(bool,'T_0') = s(bool,X149)) )),
  inference(superposition,[],[f322,f772])).
fof(f2452,plain,(
  ( ! [X144,X142,X140,X143,X141] : (p(s(bool,happ(s(fun(X140,bool),X141),s(X140,X143)))) | ~p(s(bool,X142)) | s(bool,X144) = s(bool,happ(s(fun(X140,bool),X141),s(X140,sK19(X141,X142,X140)))) | s(bool,'T_0') = s(bool,X144)) )),
  inference(subsumption_resolution,[],[f2056,f269])).
fof(f2056,plain,(
  ( ! [X144,X142,X140,X143,X141] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X140,bool),X141),s(X140,X143)))) | ~p(s(bool,X142)) | s(bool,X144) = s(bool,happ(s(fun(X140,bool),X141),s(X140,sK19(X141,X142,X140)))) | s(bool,'T_0') = s(bool,X144)) )),
  inference(superposition,[],[f320,f772])).
fof(f2451,plain,(
  ( ! [X138,X136,X139,X137] : (p(s(bool,happ(s(fun(X136,bool),X137),s(X136,X138)))) | s(bool,X139) = s(bool,happ(s(fun(X136,bool),X137),s(X136,sK17(X137,X136)))) | s(bool,'T_0') = s(bool,X139)) )),
  inference(subsumption_resolution,[],[f2055,f269])).
fof(f2055,plain,(
  ( ! [X138,X136,X139,X137] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X136,bool),X137),s(X136,X138)))) | s(bool,X139) = s(bool,happ(s(fun(X136,bool),X137),s(X136,sK17(X137,X136)))) | s(bool,'T_0') = s(bool,X139)) )),
  inference(superposition,[],[f313,f772])).
fof(f2450,plain,(
  ( ! [X134,X132,X135,X133] : (p(s(bool,happ(s(fun(X132,bool),X133),s(X132,X134)))) | s(bool,X135) = s(bool,happ(s(fun(X132,bool),X133),s(X132,sK16(X133,X132)))) | s(bool,'T_0') = s(bool,X135)) )),
  inference(subsumption_resolution,[],[f2054,f269])).
fof(f2054,plain,(
  ( ! [X134,X132,X135,X133] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X132,bool),X133),s(X132,X134)))) | s(bool,X135) = s(bool,happ(s(fun(X132,bool),X133),s(X132,sK16(X133,X132)))) | s(bool,'T_0') = s(bool,X135)) )),
  inference(superposition,[],[f312,f772])).
fof(f2449,plain,(
  ( ! [X127,X130,X128,X131,X129] : (p(s(bool,X130)) | p(s(bool,happ(s(fun(X127,bool),X128),s(X127,sK21(X128,X130,X127))))) | s(bool,X131) = s(bool,happ(s(fun(X127,bool),X128),s(X127,X129))) | s(bool,'T_0') = s(bool,X131)) )),
  inference(subsumption_resolution,[],[f2053,f269])).
fof(f2053,plain,(
  ( ! [X127,X130,X128,X131,X129] : (~p(s(bool,'T_0')) | p(s(bool,X130)) | p(s(bool,happ(s(fun(X127,bool),X128),s(X127,sK21(X128,X130,X127))))) | s(bool,X131) = s(bool,happ(s(fun(X127,bool),X128),s(X127,X129))) | s(bool,'T_0') = s(bool,X131)) )),
  inference(superposition,[],[f324,f772])).
fof(f2448,plain,(
  ( ! [X125,X123,X126,X124,X122] : (p(s(bool,X125)) | p(s(bool,happ(s(fun(X122,bool),X123),s(X122,sK22(X123,X125,X122))))) | s(bool,X126) = s(bool,happ(s(fun(X122,bool),X123),s(X122,X124))) | s(bool,'T_0') = s(bool,X126)) )),
  inference(subsumption_resolution,[],[f2052,f269])).
fof(f2052,plain,(
  ( ! [X125,X123,X126,X124,X122] : (~p(s(bool,'T_0')) | p(s(bool,X125)) | p(s(bool,happ(s(fun(X122,bool),X123),s(X122,sK22(X123,X125,X122))))) | s(bool,X126) = s(bool,happ(s(fun(X122,bool),X123),s(X122,X124))) | s(bool,'T_0') = s(bool,X126)) )),
  inference(superposition,[],[f326,f772])).
fof(f2447,plain,(
  ( ! [X121,X118,X120,X119,X117] : (~p(s(bool,X120)) | p(s(bool,happ(s(fun(X117,bool),X118),s(X117,sK25(X120,X118,X117))))) | s(bool,X121) = s(bool,happ(s(fun(X117,bool),X118),s(X117,X119))) | s(bool,'T_0') = s(bool,X121)) )),
  inference(subsumption_resolution,[],[f2051,f269])).
fof(f2051,plain,(
  ( ! [X121,X118,X120,X119,X117] : (~p(s(bool,'T_0')) | ~p(s(bool,X120)) | p(s(bool,happ(s(fun(X117,bool),X118),s(X117,sK25(X120,X118,X117))))) | s(bool,X121) = s(bool,happ(s(fun(X117,bool),X118),s(X117,X119))) | s(bool,'T_0') = s(bool,X121)) )),
  inference(superposition,[],[f331,f772])).
fof(f2446,plain,(
  ( ! [X116,X114,X112,X115,X113] : (~p(s(bool,X115)) | p(s(bool,happ(s(fun(X112,bool),X113),s(X112,sK26(X115,X113,X112))))) | s(bool,X116) = s(bool,happ(s(fun(X112,bool),X113),s(X112,X114))) | s(bool,'T_0') = s(bool,X116)) )),
  inference(subsumption_resolution,[],[f2050,f269])).
fof(f2050,plain,(
  ( ! [X116,X114,X112,X115,X113] : (~p(s(bool,'T_0')) | ~p(s(bool,X115)) | p(s(bool,happ(s(fun(X112,bool),X113),s(X112,sK26(X115,X113,X112))))) | s(bool,X116) = s(bool,happ(s(fun(X112,bool),X113),s(X112,X114))) | s(bool,'T_0') = s(bool,X116)) )),
  inference(superposition,[],[f333,f772])).
fof(f2445,plain,(
  ( ! [X111,X109,X107,X110,X108] : (~p(s(bool,'const.bool.IN_2'(s(X107,X110),s(fun(X107,bool),X108)))) | p(s(bool,'const.bool.IN_2'(s(X107,X110),s(fun(X107,bool),X109)))) | s(bool,X111) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X107,bool),X108),s(fun(X107,bool),X109))) | s(bool,'T_0') = s(bool,X111)) )),
  inference(subsumption_resolution,[],[f2049,f269])).
fof(f2049,plain,(
  ( ! [X111,X109,X107,X110,X108] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(X107,X110),s(fun(X107,bool),X108)))) | p(s(bool,'const.bool.IN_2'(s(X107,X110),s(fun(X107,bool),X109)))) | s(bool,X111) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X107,bool),X108),s(fun(X107,bool),X109))) | s(bool,'T_0') = s(bool,X111)) )),
  inference(superposition,[],[f316,f772])).
fof(f2444,plain,(
  ( ! [X94,X92,X93,X91] : (~sP8(X91,X92,X93) | sP9(X91,X92,X93) | s(bool,X91) = s(bool,X94) | s(bool,'T_0') = s(bool,X94)) )),
  inference(subsumption_resolution,[],[f2040,f269])).
fof(f2040,plain,(
  ( ! [X94,X92,X93,X91] : (~p(s(bool,'T_0')) | ~sP8(X91,X92,X93) | sP9(X91,X92,X93) | s(bool,X91) = s(bool,X94) | s(bool,'T_0') = s(bool,X94)) )),
  inference(superposition,[],[f469,f772])).
fof(f2443,plain,(
  ( ! [X90,X88,X87,X89] : (~sP10(X87,X88,X89) | sP11(X87,X88,X89) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(subsumption_resolution,[],[f2039,f269])).
fof(f2039,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,'T_0')) | ~sP10(X87,X88,X89) | sP11(X87,X88,X89) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f468,f772])).
fof(f2442,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,X84)) | p(s(bool,X85)) | ~sP11(X83,X85,X84) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(subsumption_resolution,[],[f2038,f269])).
fof(f2038,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,'T_0')) | ~p(s(bool,X84)) | p(s(bool,X85)) | ~sP11(X83,X85,X84) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f465,f772])).
fof(f2441,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,X80)) | sP11(X79,X80,X81) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(subsumption_resolution,[],[f2037,f269])).
fof(f2037,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,'T_0')) | ~p(s(bool,X80)) | sP11(X79,X80,X81) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(superposition,[],[f464,f772])).
fof(f2440,plain,(
  ( ! [X78,X76,X77,X75] : (p(s(bool,X76)) | sP11(X75,X77,X76) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(subsumption_resolution,[],[f2036,f269])).
fof(f2036,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,'T_0')) | p(s(bool,X76)) | sP11(X75,X77,X76) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f463,f772])).
fof(f2439,plain,(
  ( ! [X74,X72,X71,X73] : (p(s(bool,X72)) | ~p(s(bool,X73)) | ~sP10(X73,X72,X71) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(subsumption_resolution,[],[f2035,f269])).
fof(f2035,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,'T_0')) | p(s(bool,X72)) | ~p(s(bool,X73)) | ~sP10(X73,X72,X71) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(superposition,[],[f458,f772])).
fof(f2438,plain,(
  ( ! [X70,X68,X69,X67] : (p(s(bool,X68)) | p(s(bool,X69)) | ~sP9(X67,X69,X68) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f2034,f269])).
fof(f2034,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,'T_0')) | p(s(bool,X68)) | p(s(bool,X69)) | ~sP9(X67,X69,X68) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f448,f772])).
fof(f2437,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X64)) | sP9(X63,X64,X65) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(subsumption_resolution,[],[f2033,f269])).
fof(f2033,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,'T_0')) | ~p(s(bool,X64)) | sP9(X63,X64,X65) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f447,f772])).
fof(f2436,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X60)) | sP9(X59,X61,X60) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(subsumption_resolution,[],[f2032,f269])).
fof(f2032,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,'T_0')) | ~p(s(bool,X60)) | sP9(X59,X61,X60) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f446,f772])).
fof(f2435,plain,(
  ( ! [X57,X58,X56,X55] : (p(s(bool,X56)) | p(s(bool,X57)) | ~sP8(X55,X56,X57) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(subsumption_resolution,[],[f2031,f269])).
fof(f2031,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,'T_0')) | p(s(bool,X56)) | p(s(bool,X57)) | ~sP8(X55,X56,X57) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f441,f772])).
fof(f2434,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X52)) | p(s(bool,X53)) | ~sP7(X53,X52,X51) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(subsumption_resolution,[],[f2030,f269])).
fof(f2030,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,'T_0')) | ~p(s(bool,X52)) | p(s(bool,X53)) | ~sP7(X53,X52,X51) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(superposition,[],[f433,f772])).
fof(f2433,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X48)) | p(s(bool,X49)) | ~sP6(X48,X47,X49) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f2029,f269])).
fof(f2029,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,X48)) | p(s(bool,X49)) | ~sP6(X48,X47,X49) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f424,f772])).
fof(f2432,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X44)) | p(s(bool,X45)) | ~sP3(X44,X43,X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f2028,f269])).
fof(f2028,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,'T_0')) | ~p(s(bool,X44)) | p(s(bool,X45)) | ~sP3(X44,X43,X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f405,f772])).
fof(f2431,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X40)) | p(s(bool,X41)) | ~sP4(X40,X39,X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(subsumption_resolution,[],[f2027,f269])).
fof(f2027,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X40)) | p(s(bool,X41)) | ~sP4(X40,X39,X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f401,f772])).
fof(f2430,plain,(
  ( ! [X37,X35,X38,X36] : (sP5(X36,X37,X35) | ~sP4(X36,X35,X37) | ~sP3(X37,X35,X36) | ~sP2(X35,X37,X36) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(subsumption_resolution,[],[f2026,f269])).
fof(f2026,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'T_0')) | sP5(X36,X37,X35) | ~sP4(X36,X35,X37) | ~sP3(X37,X35,X36) | ~sP2(X35,X37,X36) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f398,f772])).
fof(f2429,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,X32)) | p(s(bool,X33)) | ~sP5(X32,X31,X33) | s(bool,X31) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f2025,f269])).
fof(f2025,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X32)) | p(s(bool,X33)) | ~sP5(X32,X31,X33) | s(bool,X31) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f397,f772])).
fof(f2428,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X29)) | sP1(X29,X28) | s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(subsumption_resolution,[],[f2024,f269])).
fof(f2024,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'T_0')) | p(s(bool,X29)) | sP1(X29,X28) | s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f366,f772])).
fof(f2426,plain,(
  ( ! [X24,X23,X22] : (sP0(X22,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f2022,f269])).
fof(f2022,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | sP0(X22,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f348,f772])).
fof(f2425,plain,(
  ( ! [X21,X19,X20] : (sP0(X20,X19) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(subsumption_resolution,[],[f2021,f269])).
fof(f2021,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | sP0(X20,X19) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f347,f772])).
fof(f2424,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f2020,f269])).
fof(f2020,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f341,f772])).
fof(f2423,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f2016,f269])).
fof(f2016,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f289,f772])).
fof(f2011,plain,(
  ( ! [X127,X130,X128,X129] : (~p(s(bool,'const.bool.IN_2'(s(bool,X127),s(fun(bool,bool),X128)))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X129),s(fun(bool,bool),X128)))))) | s(bool,X127) = s(bool,X130) | s(bool,'T_0') = s(bool,X130)) )),
  inference(superposition,[],[f337,f772])).
fof(f2010,plain,(
  ( ! [X125,X123,X126,X124] : (~p(s(bool,'const.bool.IN_2'(s(bool,X123),s(fun(bool,bool),X124)))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X124),s(fun(bool,bool),X125)))))) | s(bool,X123) = s(bool,X126) | s(bool,'T_0') = s(bool,X126)) )),
  inference(superposition,[],[f336,f772])).
fof(f2009,plain,(
  ( ! [X121,X122,X120,X119] : (~p(s(bool,'const.bool.IN_2'(s(bool,X119),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X120),s(fun(bool,bool),X121)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X120)))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X121)))) | s(bool,X119) = s(bool,X122) | s(bool,'T_0') = s(bool,X122)) )),
  inference(superposition,[],[f335,f772])).
fof(f1975,plain,(
  ( ! [X4,X2,X3] : (s(bool,X2) != s(bool,X3) | p(s(bool,X3)) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f285,f772])).
fof(f1895,plain,(
  ( ! [X288,X290,X287,X289] : (s(X289,f(s(bool,X290))) != s(X289,g(s(bool,X290))) | s(fun(bool,X289),X287) = s(fun(bool,X289),X288) | s(bool,'T_0') = s(bool,X290) | s(bool,'T_0') = s(bool,sK31(X287,X288,X289,bool))) )),
  inference(superposition,[],[f381,f772])).
fof(f1894,plain,(
  ( ! [X282,X285,X283,X284,X286] : (~p(s(bool,happ(s(fun(X284,bool),happ(s(fun(bool,fun(X284,bool)),X283),s(bool,X285))),s(X284,happ(s(fun(bool,X284),X282),s(bool,X285)))))) | p(s(bool,happ(s(fun(X284,bool),happ(s(fun(bool,fun(X284,bool)),X283),s(bool,X286))),s(X284,sK30(X286,X283,bool,X284))))) | s(bool,'T_0') = s(bool,X285) | s(bool,'T_0') = s(bool,sK29(X282,X283,bool,X284))) )),
  inference(superposition,[],[f380,f772])).
fof(f1893,plain,(
  ( ! [X280,X277,X279,X281,X278] : (~p(s(bool,happ(s(fun(X278,bool),happ(s(fun(bool,fun(X278,bool)),X277),s(bool,X279))),s(X278,X280)))) | p(s(bool,happ(s(fun(X278,bool),happ(s(fun(bool,fun(X278,bool)),X277),s(bool,X281))),s(X278,happ(s(fun(bool,X278),sK28(X277,bool,X278)),s(bool,X281)))))) | s(bool,'T_0') = s(bool,X279) | s(bool,'T_0') = s(bool,sK27(X277,bool,X278))) )),
  inference(superposition,[],[f379,f772])).
fof(f1892,plain,(
  ( ! [X273,X275,X276,X274] : (~p(s(bool,happ(s(fun(bool,bool),X273),s(bool,X275)))) | p(s(bool,X274)) | p(s(bool,happ(s(fun(bool,bool),X273),s(bool,X276)))) | s(bool,'T_0') = s(bool,X275) | s(bool,'T_0') = s(bool,sK24(X273,X274,bool))) )),
  inference(superposition,[],[f330,f772])).
fof(f1891,plain,(
  ( ! [X269,X271,X270,X272] : (~p(s(bool,happ(s(fun(bool,bool),X269),s(bool,X271)))) | p(s(bool,X270)) | p(s(bool,happ(s(fun(bool,bool),X269),s(bool,X272)))) | s(bool,'T_0') = s(bool,X271) | s(bool,'T_0') = s(bool,sK23(X269,X270,bool))) )),
  inference(superposition,[],[f328,f772])).
fof(f1890,plain,(
  ( ! [X265,X267,X268,X266] : (~p(s(bool,happ(s(fun(bool,bool),X265),s(bool,X267)))) | p(s(bool,happ(s(fun(bool,bool),X265),s(bool,X268)))) | ~p(s(bool,X266)) | s(bool,'T_0') = s(bool,X267) | s(bool,'T_0') = s(bool,sK20(X265,X266,bool))) )),
  inference(superposition,[],[f322,f772])).
fof(f1889,plain,(
  ( ! [X261,X263,X262,X264] : (~p(s(bool,happ(s(fun(bool,bool),X261),s(bool,X263)))) | p(s(bool,happ(s(fun(bool,bool),X261),s(bool,X264)))) | ~p(s(bool,X262)) | s(bool,'T_0') = s(bool,X263) | s(bool,'T_0') = s(bool,sK19(X261,X262,bool))) )),
  inference(superposition,[],[f320,f772])).
fof(f1888,plain,(
  ( ! [X259,X260,X258] : (p(s(bool,'const.bool.IN_2'(s(bool,X260),s(fun(bool,bool),X259)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(bool,bool),X259),s(fun(bool,bool),X258)))) | s(bool,'T_0') = s(bool,X260) | s(bool,'T_0') = s(bool,sK18(X258,X259,bool))) )),
  inference(superposition,[],[f317,f772])).
fof(f1887,plain,(
  ( ! [X255,X257,X256] : (~p(s(bool,'const.bool.IN_2'(s(bool,X257),s(fun(bool,bool),X255)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(bool,bool),X256),s(fun(bool,bool),X255)))) | s(bool,'T_0') = s(bool,X257) | s(bool,'T_0') = s(bool,sK18(X255,X256,bool))) )),
  inference(superposition,[],[f318,f772])).
fof(f1886,plain,(
  ( ! [X253,X254,X252] : (~p(s(bool,happ(s(fun(bool,bool),X252),s(bool,X253)))) | p(s(bool,happ(s(fun(bool,bool),X252),s(bool,X254)))) | s(bool,'T_0') = s(bool,X253) | s(bool,'T_0') = s(bool,sK17(X252,bool))) )),
  inference(superposition,[],[f313,f772])).
fof(f1885,plain,(
  ( ! [X251,X249,X250] : (~p(s(bool,happ(s(fun(bool,bool),X249),s(bool,X250)))) | p(s(bool,happ(s(fun(bool,bool),X249),s(bool,X251)))) | s(bool,'T_0') = s(bool,X250) | s(bool,'T_0') = s(bool,sK16(X249,bool))) )),
  inference(superposition,[],[f312,f772])).
fof(f1884,plain,(
  ( ! [X246,X244,X248,X247,X245] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(X244,fun(bool,bool)),X246),s(X244,sK29(X245,X246,X244,bool)))),s(bool,X247)))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(X244,fun(bool,bool)),X246),s(X244,X248))),s(bool,sK30(X248,X246,X244,bool))))) | s(bool,'T_0') = s(bool,X247) | s(bool,'T_0') = s(bool,happ(s(fun(X244,bool),X245),s(X244,sK29(X245,X246,X244,bool))))) )),
  inference(superposition,[],[f380,f772])).
fof(f1883,plain,(
  ( ! [X243,X241,X242,X240] : (s(bool,X243) != s(bool,f(s(X240,sK31(X241,X242,bool,X240)))) | s(fun(X240,bool),X241) = s(fun(X240,bool),X242) | s(bool,'T_0') = s(bool,X243) | s(bool,'T_0') = s(bool,g(s(X240,sK31(X241,X242,bool,X240))))) )),
  inference(superposition,[],[f381,f772])).
fof(f1881,plain,(
  ( ! [X229,X227,X230,X228] : (~p(s(bool,'const.bool.IN_2'(s(bool,X228),s(fun(bool,bool),X229)))) | p(s(bool,'const.bool.IN_2'(s(bool,X227),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X230),s(fun(bool,bool),X229)))))) | s(bool,'T_0') = s(bool,X228) | s(bool,'T_0') = s(bool,X227)) )),
  inference(superposition,[],[f337,f772])).
fof(f1880,plain,(
  ( ! [X225,X223,X226,X224] : (~p(s(bool,'const.bool.IN_2'(s(bool,X224),s(fun(bool,bool),X225)))) | p(s(bool,'const.bool.IN_2'(s(bool,X223),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X225),s(fun(bool,bool),X226)))))) | s(bool,'T_0') = s(bool,X224) | s(bool,'T_0') = s(bool,X223)) )),
  inference(superposition,[],[f336,f772])).
fof(f1879,plain,(
  ( ! [X222,X220,X219,X221] : (~p(s(bool,'const.bool.IN_2'(s(bool,X220),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X221),s(fun(bool,bool),X222)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X219),s(fun(bool,bool),X221)))) | p(s(bool,'const.bool.IN_2'(s(bool,X219),s(fun(bool,bool),X222)))) | s(bool,'T_0') = s(bool,X220) | s(bool,'T_0') = s(bool,X219)) )),
  inference(superposition,[],[f335,f772])).
fof(f1795,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f772,f772])).
fof(f1586,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(X0,fun(bool,bool)),X2),s(X0,sK29(X1,X2,X0,bool)))),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(X0,fun(bool,bool)),X2),s(X0,X3))),s(bool,sK30(X3,X2,X0,bool))))) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK29(X1,X2,X0,bool))))) )),
  inference(superposition,[],[f380,f276])).
fof(f1585,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X2,bool),happ(s(fun(bool,fun(X2,bool)),X1),s(bool,'F_0'))),s(X2,happ(s(fun(bool,X2),X0),s(bool,'F_0')))))) | p(s(bool,happ(s(fun(X2,bool),happ(s(fun(bool,fun(X2,bool)),X1),s(bool,X3))),s(X2,sK30(X3,X1,bool,X2))))) | s(bool,'T_0') = s(bool,sK29(X0,X1,bool,X2))) )),
  inference(superposition,[],[f380,f276])).
fof(f1573,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(bool,fun(X1,bool)),X0),s(bool,'F_0'))),s(X1,X2)))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(bool,fun(X1,bool)),X0),s(bool,X3))),s(X1,happ(s(fun(bool,X1),sK28(X0,bool,X1)),s(bool,X3)))))) | s(bool,'T_0') = s(bool,sK27(X0,bool,X1))) )),
  inference(superposition,[],[f379,f276])).
fof(f1563,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X8),s(fun(bool,bool),X9)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X7),s(fun(bool,bool),X8)))) | p(s(bool,'const.bool.IN_2'(s(bool,X7),s(fun(bool,bool),X9)))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f335,f276])).
fof(f1562,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,'const.bool.IN_2'(s(bool,X4),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X5),s(fun(bool,bool),X6)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X5)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X6)))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f335,f276])).
fof(f1533,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X8)))) | p(s(bool,'const.bool.IN_2'(s(bool,X7),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X9),s(fun(bool,bool),X8)))))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f337,f276])).
fof(f1532,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,'const.bool.IN_2'(s(bool,X4),s(fun(bool,bool),X5)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X6),s(fun(bool,bool),X5)))))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f337,f276])).
fof(f1518,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X8)))) | p(s(bool,'const.bool.IN_2'(s(bool,X7),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X8),s(fun(bool,bool),X9)))))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f336,f276])).
fof(f1517,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,'const.bool.IN_2'(s(bool,X4),s(fun(bool,bool),X5)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X5),s(fun(bool,bool),X6)))))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f336,f276])).
fof(f1461,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X2)))) | s(bool,'T_0') = s(bool,sK24(X0,X1,bool))) )),
  inference(superposition,[],[f330,f276])).
fof(f1449,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X2)))) | s(bool,'T_0') = s(bool,sK23(X0,X1,bool))) )),
  inference(superposition,[],[f328,f276])).
fof(f1401,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X2)))) | ~p(s(bool,X1)) | s(bool,'T_0') = s(bool,sK20(X0,X1,bool))) )),
  inference(superposition,[],[f322,f276])).
fof(f1390,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X2)))) | ~p(s(bool,X1)) | s(bool,'T_0') = s(bool,sK19(X0,X1,bool))) )),
  inference(superposition,[],[f320,f276])).
fof(f1377,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(bool,bool),X1),s(fun(bool,bool),X0)))) | s(bool,'T_0') = s(bool,sK18(X0,X1,bool))) )),
  inference(superposition,[],[f318,f276])).
fof(f1364,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(bool,bool),X1),s(fun(bool,bool),X0)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))) | s(bool,'T_0') = s(bool,sK18(X0,X1,bool))) )),
  inference(superposition,[],[f317,f276])).
fof(f1350,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X1)))) | s(bool,'T_0') = s(bool,sK17(X0,bool))) )),
  inference(superposition,[],[f313,f276])).
fof(f1338,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X1)))) | s(bool,'T_0') = s(bool,sK16(X0,bool))) )),
  inference(superposition,[],[f312,f276])).
fof(f1237,plain,(
  sP5('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f507,f496,f491,f269,f398])).
fof(f982,plain,(
  sP7('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f510,f269,f269,f420])).
fof(f879,plain,(
  ~sP5('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f269,f527,f415])).
fof(f707,plain,(
  ( ! [X0] : (sP11('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f269,f269,f464])).
fof(f673,plain,(
  ( ! [X0] : (sP9('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f269,f269,f447])).
fof(f662,plain,(
  ( ! [X0] : (sP9('T_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f269,f269,f446])).
fof(f563,plain,(
  ~sP1('T_0','T_0')),
  inference(unit_resulting_resolution,[],[f269,f269,f362])).
fof(f507,plain,(
  ( ! [X0,X1] : (sP2('T_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f269,f412])).
fof(f488,plain,(
  ( ! [X0] : (sP0('T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f269,f348])).
fof(f485,plain,(
  ( ! [X0] : (sP0(X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f269,f347])).
fof(f445,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f257])).
fof(f470,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(duplicate_literal_removal,[],[f423])).
fof(f423,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f245])).
fof(f396,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP4(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f233])).
fof(f379,plain,(
  ( ! [X2,X0,X10,X8,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,sK27(X2,X1,X0)))),s(X0,X10)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X8))),s(X0,happ(s(fun(X1,X0),sK28(X2,X1,X0)),s(X1,X8))))))) )),
  inference(cnf_transformation,[],[f225])).
fof(f225,plain,(
  ! [X0,X1,X2] : ((! [X10] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,sK27(X2,X1,X0)))),s(X0,X10)))) | ! [X8] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X8))),s(X0,happ(s(fun(X1,X0),sK28(X2,X1,X0)),s(X1,X8))))))) & (! [X5] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,sK29(X5,X2,X1,X0)))),s(X0,happ(s(fun(X1,X0),X5),s(X1,sK29(X5,X2,X1,X0))))))) | ! [X3] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X3))),s(X0,sK30(X3,X2,X1,X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK27,sK28,sK29,sK30])],[f224])).
fof(f224,plain,(
  ! [X0,X1,X2] : ((? [X9] : ! [X10] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X9))),s(X0,X10)))) | ? [X7] : ! [X8] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X8))),s(X0,happ(s(fun(X1,X0),X7),s(X1,X8))))))) & (! [X5] : ? [X6] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X6))),s(X0,happ(s(fun(X1,X0),X5),s(X1,X6)))))) | ! [X3] : ? [X4] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X3))),s(X0,X4))))))),
  inference(rectify,[],[f223])).
fof(f223,plain,(
  ! [X0,X1,X2] : ((? [X3] : ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X3))),s(X0,X4)))) | ? [X5] : ! [X6] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X6))),s(X0,happ(s(fun(X1,X0),X5),s(X1,X6))))))) & (! [X5] : ? [X6] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X6))),s(X0,happ(s(fun(X1,X0),X5),s(X1,X6)))))) | ! [X3] : ? [X4] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X3))),s(X0,X4))))))),
  inference(nnf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0,X1,X2] : (! [X3] : ? [X4] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X3))),s(X0,X4)))) <=> ? [X5] : ! [X6] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X6))),s(X0,happ(s(fun(X1,X0),X5),s(X1,X6)))))))),
  inference(flattening,[],[f104])).
fof(f104,plain,(
  ! [X0,X1] : ! [X2] : (! [X3] : ? [X4] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X3))),s(X0,X4)))) <=> ? [X5] : ! [X6] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X6))),s(X0,happ(s(fun(X1,X0),X5),s(X1,X6)))))))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X17,X6] : ! [X15] : (! [X13] : ? [X14] : p(s(bool,happ(s(fun(X17,bool),happ(s(fun(X6,fun(X17,bool)),X15),s(X6,X13))),s(X17,X14)))) <=> ? [X19] : ! [X13] : p(s(bool,happ(s(fun(X17,bool),happ(s(fun(X6,fun(X17,bool)),X15),s(X6,X13))),s(X17,happ(s(fun(X6,X17),X19),s(X6,X13)))))))),
  file('thm.pred_set.UNION_SUBSET','thm.bool.SKOLEM_THM')).
fof(f380,plain,(
  ( ! [X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,sK29(X5,X2,X1,X0)))),s(X0,happ(s(fun(X1,X0),X5),s(X1,sK29(X5,X2,X1,X0))))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X3))),s(X0,sK30(X3,X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f225])).
fof(f331,plain,(
  ( ! [X6,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) | ~p(s(bool,X2)) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK25(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f204])).
fof(f204,plain,(
  ! [X0,X1,X2] : ((! [X6] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) | ~p(s(bool,X2))) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK25(X2,X1,X0))))) & p(s(bool,X2)))) & (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,X2)) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK26(X2,X1,X0))))) & p(s(bool,X2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK25,sK26])],[f203])).
fof(f203,plain,(
  ! [X0,X1,X2] : ((! [X6] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) | ~p(s(bool,X2))) | (? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & p(s(bool,X2)))) & (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,X2)) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2)))))),
  inference(rectify,[],[f202])).
fof(f202,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,X2))) | (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2)))) & (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,X2)) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2)))))),
  inference(flattening,[],[f201])).
fof(f201,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,X2))) | (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2)))) & ((! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,X2))) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2)))))),
  inference(nnf_transformation,[],[f89])).
fof(f89,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2))) <=> (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2))))),
  inference(flattening,[],[f88])).
fof(f88,plain,(
  ! [X0] : ! [X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2))) <=> (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X6] : ! [X15,X16] : (? [X13] : (p(s(bool,happ(s(fun(X6,bool),X15),s(X6,X13)))) & p(s(bool,X16))) <=> (? [X13] : p(s(bool,happ(s(fun(X6,bool),X15),s(X6,X13)))) & p(s(bool,X16))))),
  file('thm.pred_set.UNION_SUBSET','thm.bool.LEFT_EXISTS_AND_THM')).
fof(f333,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,X2)) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK26(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f204])).
fof(f328,plain,(
  ( ! [X2,X0,X5,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK23(X2,X1,X0))))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5))))) )),
  inference(cnf_transformation,[],[f200])).
fof(f200,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK23(X2,X1,X0)))))) | p(s(bool,X1)) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5))))) & ((~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK24(X2,X1,X0)))))) | ! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK23,sK24])],[f199])).
fof(f199,plain,(
  ! [X0,X1,X2] : ((? [X6] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | p(s(bool,X1)) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5))))) & ((~p(s(bool,X1)) & ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(rectify,[],[f198])).
fof(f198,plain,(
  ! [X0,X1,X2] : ((? [X3] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | p(s(bool,X1)) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) & ((~p(s(bool,X1)) & ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(flattening,[],[f197])).
fof(f197,plain,(
  ! [X0,X1,X2] : ((? [X3] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (p(s(bool,X1)) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & ((~p(s(bool,X1)) & ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X0,X1,X2] : (! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (p(s(bool,X1)) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f86])).
fof(f86,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (p(s(bool,X1)) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X6] : ! [X15,X16] : (! [X13] : (p(s(bool,X15)) | p(s(bool,happ(s(fun(X6,bool),X16),s(X6,X13))))) <=> (p(s(bool,X15)) | ! [X13] : p(s(bool,happ(s(fun(X6,bool),X16),s(X6,X13))))))),
  file('thm.pred_set.UNION_SUBSET','thm.bool.RIGHT_FORALL_OR_THM')).
fof(f330,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK24(X2,X1,X0))))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f200])).
fof(f324,plain,(
  ( ! [X6,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK21(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f196])).
fof(f196,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK21(X2,X1,X0))))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK22(X2,X1,X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK21,sK22])],[f195])).
fof(f195,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | ? [X5] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(rectify,[],[f194])).
fof(f194,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(flattening,[],[f193])).
fof(f193,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f84])).
fof(f84,plain,(
  ! [X0] : ! [X1,X2] : ((p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X6] : ! [X15,X16] : ((p(s(bool,X15)) | ? [X13] : p(s(bool,happ(s(fun(X6,bool),X16),s(X6,X13))))) <=> ? [X13] : (p(s(bool,X15)) | p(s(bool,happ(s(fun(X6,bool),X16),s(X6,X13))))))),
  file('thm.pred_set.UNION_SUBSET','thm.bool.RIGHT_OR_EXISTS_THM')).
fof(f326,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK22(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f196])).
fof(f320,plain,(
  ( ! [X2,X0,X5,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK19(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK19(X2,X1,X0))))) | ! [X5] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))))) & ((~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK20(X2,X1,X0)))))) | (p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK19,sK20])],[f191])).
fof(f191,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,X1)) | ? [X6] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) | ! [X5] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))))) & (? [X4] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(rectify,[],[f190])).
fof(f190,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,X1)) | ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | ! [X4] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (? [X4] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(flattening,[],[f189])).
fof(f189,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) | ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ! [X4] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (? [X4] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> ! [X4] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f82])).
fof(f82,plain,(
  ! [X0] : ! [X1,X2] : ((p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> ! [X4] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X6] : ! [X15,X16] : ((p(s(bool,X15)) & ! [X13] : p(s(bool,happ(s(fun(X6,bool),X16),s(X6,X13))))) <=> ! [X13] : (p(s(bool,X15)) & p(s(bool,happ(s(fun(X6,bool),X16),s(X6,X13))))))),
  file('thm.pred_set.UNION_SUBSET','thm.bool.RIGHT_AND_FORALL_THM')).
fof(f322,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK20(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f192])).
fof(f312,plain,(
  ( ! [X0,X5,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK16(X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))) )),
  inference(cnf_transformation,[],[f184])).
fof(f184,plain,(
  ! [X0,X1] : ((! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK16(X1,X0)))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK17(X1,X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK16,sK17])],[f183])).
fof(f183,plain,(
  ! [X0,X1] : ((! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(rectify,[],[f182])).
fof(f182,plain,(
  ! [X0,X1] : ((! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(nnf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0,X1] : (? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(ennf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0,X1] : (~! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(flattening,[],[f76])).
fof(f76,plain,(
  ! [X0] : ! [X1] : (~! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X6] : ! [X15] : (~! [X13] : p(s(bool,happ(s(fun(X6,bool),X15),s(X6,X13)))) <=> ? [X13] : ~p(s(bool,happ(s(fun(X6,bool),X15),s(X6,X13)))))),
  file('thm.pred_set.UNION_SUBSET','thm.bool.NOT_FORALL_THM')).
fof(f313,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK17(X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f184])).
fof(f2422,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~sP8('T_0',X1,X2)) ) | $spl42),
  inference(cnf_transformation,[],[f2422_D])).
fof(f2422_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~sP8('T_0',X1,X2)) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f2370,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl8),
  inference(cnf_transformation,[],[f2370_D])).
fof(f2370_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f2504,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('T_0',X1,X2) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f2503])).
fof(f2503,plain,(
  ( ! [X99,X97,X98,X96] : (~sP8('T_0',X97,X98) | sP9('T_0',X97,X98) | s(bool,X96) = s(bool,X99) | s(bool,'T_0') = s(bool,X99)) )),
  inference(subsumption_resolution,[],[f2118,f832])).
fof(f2118,plain,(
  ( ! [X99,X97,X98,X96] : (~p(s(bool,X96)) | ~sP8('T_0',X97,X98) | sP9('T_0',X97,X98) | s(bool,X96) = s(bool,X99) | s(bool,'T_0') = s(bool,X99)) )),
  inference(superposition,[],[f469,f772])).
fof(f2419,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~sP10('T_0',X1,X2)) ) | $spl40),
  inference(cnf_transformation,[],[f2419_D])).
fof(f2419_D,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~sP10('T_0',X1,X2)) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f2502,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP10('T_0',X1,X2) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f2501])).
fof(f2501,plain,(
  ( ! [X94,X92,X95,X93] : (~sP10('T_0',X93,X94) | sP11('T_0',X93,X94) | s(bool,X92) = s(bool,X95) | s(bool,'T_0') = s(bool,X95)) )),
  inference(subsumption_resolution,[],[f2117,f832])).
fof(f2117,plain,(
  ( ! [X94,X92,X95,X93] : (~p(s(bool,X92)) | ~sP10('T_0',X93,X94) | sP11('T_0',X93,X94) | s(bool,X92) = s(bool,X95) | s(bool,'T_0') = s(bool,X95)) )),
  inference(superposition,[],[f468,f772])).
fof(f2416,plain,(
  ( ! [X2,X1] : (~sP11('T_0',X2,X1) | p(s(bool,X2)) | ~p(s(bool,X1))) ) | $spl38),
  inference(cnf_transformation,[],[f2416_D])).
fof(f2416_D,plain,(
  ( ! [X2,X1] : (~sP11('T_0',X2,X1) | p(s(bool,X2)) | ~p(s(bool,X1))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f2500,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f2499])).
fof(f2499,plain,(
  ( ! [X90,X88,X91,X89] : (~p(s(bool,X89)) | p(s(bool,X90)) | ~sP11('T_0',X90,X89) | s(bool,X88) = s(bool,X91) | s(bool,'T_0') = s(bool,X91)) )),
  inference(subsumption_resolution,[],[f2116,f832])).
fof(f2116,plain,(
  ( ! [X90,X88,X91,X89] : (~p(s(bool,X88)) | ~p(s(bool,X89)) | p(s(bool,X90)) | ~sP11('T_0',X90,X89) | s(bool,X88) = s(bool,X91) | s(bool,'T_0') = s(bool,X91)) )),
  inference(superposition,[],[f465,f772])).
fof(f2413,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~p(s(bool,X1))) ) | $spl36),
  inference(cnf_transformation,[],[f2413_D])).
fof(f2413_D,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~p(s(bool,X1))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f2498,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f2497])).
fof(f2497,plain,(
  ( ! [X87,X85,X86,X84] : (~p(s(bool,X85)) | sP11('T_0',X85,X86) | s(bool,X84) = s(bool,X87) | s(bool,'T_0') = s(bool,X87)) )),
  inference(subsumption_resolution,[],[f2115,f832])).
fof(f2115,plain,(
  ( ! [X87,X85,X86,X84] : (~p(s(bool,X84)) | ~p(s(bool,X85)) | sP11('T_0',X85,X86) | s(bool,X84) = s(bool,X87) | s(bool,'T_0') = s(bool,X87)) )),
  inference(superposition,[],[f464,f772])).
fof(f2410,plain,(
  ( ! [X2,X1] : (sP11('T_0',X2,X1) | p(s(bool,X1))) ) | $spl34),
  inference(cnf_transformation,[],[f2410_D])).
fof(f2410_D,plain,(
  ( ! [X2,X1] : (sP11('T_0',X2,X1) | p(s(bool,X1))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f2496,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f2495])).
fof(f2495,plain,(
  ( ! [X80,X83,X81,X82] : (p(s(bool,X81)) | sP11('T_0',X82,X81) | s(bool,X80) = s(bool,X83) | s(bool,'T_0') = s(bool,X83)) )),
  inference(subsumption_resolution,[],[f2114,f832])).
fof(f2114,plain,(
  ( ! [X80,X83,X81,X82] : (~p(s(bool,X80)) | p(s(bool,X81)) | sP11('T_0',X82,X81) | s(bool,X80) = s(bool,X83) | s(bool,'T_0') = s(bool,X83)) )),
  inference(superposition,[],[f463,f772])).
fof(f2407,plain,(
  ( ! [X2,X1] : (~sP10(X2,X1,'T_0') | ~p(s(bool,X2)) | p(s(bool,X1))) ) | $spl32),
  inference(cnf_transformation,[],[f2407_D])).
fof(f2407_D,plain,(
  ( ! [X2,X1] : (~sP10(X2,X1,'T_0') | ~p(s(bool,X2)) | p(s(bool,X1))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f2494,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP10(X2,X1,'T_0')) )),
  inference(condensation,[],[f2493])).
fof(f2493,plain,(
  ( ! [X78,X76,X79,X77] : (p(s(bool,X77)) | ~p(s(bool,X78)) | ~sP10(X78,X77,'T_0') | s(bool,X76) = s(bool,X79) | s(bool,'T_0') = s(bool,X79)) )),
  inference(subsumption_resolution,[],[f2113,f832])).
fof(f2113,plain,(
  ( ! [X78,X76,X79,X77] : (~p(s(bool,X76)) | p(s(bool,X77)) | ~p(s(bool,X78)) | ~sP10(X78,X77,'T_0') | s(bool,X76) = s(bool,X79) | s(bool,'T_0') = s(bool,X79)) )),
  inference(superposition,[],[f458,f772])).
fof(f2404,plain,(
  ( ! [X2,X1] : (~sP9('T_0',X2,X1) | p(s(bool,X2)) | p(s(bool,X1))) ) | $spl30),
  inference(cnf_transformation,[],[f2404_D])).
fof(f2404_D,plain,(
  ( ! [X2,X1] : (~sP9('T_0',X2,X1) | p(s(bool,X2)) | p(s(bool,X1))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f2492,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f2491])).
fof(f2491,plain,(
  ( ! [X74,X72,X75,X73] : (p(s(bool,X73)) | p(s(bool,X74)) | ~sP9('T_0',X74,X73) | s(bool,X72) = s(bool,X75) | s(bool,'T_0') = s(bool,X75)) )),
  inference(subsumption_resolution,[],[f2112,f832])).
fof(f2112,plain,(
  ( ! [X74,X72,X75,X73] : (~p(s(bool,X72)) | p(s(bool,X73)) | p(s(bool,X74)) | ~sP9('T_0',X74,X73) | s(bool,X72) = s(bool,X75) | s(bool,'T_0') = s(bool,X75)) )),
  inference(superposition,[],[f448,f772])).
fof(f2401,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~p(s(bool,X1))) ) | $spl28),
  inference(cnf_transformation,[],[f2401_D])).
fof(f2401_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~p(s(bool,X1))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f2490,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f2489])).
fof(f2489,plain,(
  ( ! [X70,X68,X71,X69] : (~p(s(bool,X69)) | sP9('T_0',X69,X70) | s(bool,X68) = s(bool,X71) | s(bool,'T_0') = s(bool,X71)) )),
  inference(subsumption_resolution,[],[f2111,f832])).
fof(f2111,plain,(
  ( ! [X70,X68,X71,X69] : (~p(s(bool,X68)) | ~p(s(bool,X69)) | sP9('T_0',X69,X70) | s(bool,X68) = s(bool,X71) | s(bool,'T_0') = s(bool,X71)) )),
  inference(superposition,[],[f447,f772])).
fof(f2398,plain,(
  ( ! [X2,X1] : (sP9('T_0',X2,X1) | ~p(s(bool,X1))) ) | $spl26),
  inference(cnf_transformation,[],[f2398_D])).
fof(f2398_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X2,X1) | ~p(s(bool,X1))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f2488,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f2487])).
fof(f2487,plain,(
  ( ! [X66,X64,X67,X65] : (~p(s(bool,X65)) | sP9('T_0',X66,X65) | s(bool,X64) = s(bool,X67) | s(bool,'T_0') = s(bool,X67)) )),
  inference(subsumption_resolution,[],[f2110,f832])).
fof(f2110,plain,(
  ( ! [X66,X64,X67,X65] : (~p(s(bool,X64)) | ~p(s(bool,X65)) | sP9('T_0',X66,X65) | s(bool,X64) = s(bool,X67) | s(bool,'T_0') = s(bool,X67)) )),
  inference(superposition,[],[f446,f772])).
fof(f2395,plain,(
  ( ! [X2,X1] : (~sP8('T_0',X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) ) | $spl24),
  inference(cnf_transformation,[],[f2395_D])).
fof(f2395_D,plain,(
  ( ! [X2,X1] : (~sP8('T_0',X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f2486,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP8('T_0',X1,X2)) )),
  inference(condensation,[],[f2485])).
fof(f2485,plain,(
  ( ! [X61,X62,X60,X63] : (p(s(bool,X61)) | p(s(bool,X62)) | ~sP8('T_0',X61,X62) | s(bool,X60) = s(bool,X63) | s(bool,'T_0') = s(bool,X63)) )),
  inference(subsumption_resolution,[],[f2109,f832])).
fof(f2109,plain,(
  ( ! [X61,X62,X60,X63] : (~p(s(bool,X60)) | p(s(bool,X61)) | p(s(bool,X62)) | ~sP8('T_0',X61,X62) | s(bool,X60) = s(bool,X63) | s(bool,'T_0') = s(bool,X63)) )),
  inference(superposition,[],[f441,f772])).
fof(f2392,plain,(
  ( ! [X2,X1] : (~sP7(X2,X1,'T_0') | p(s(bool,X2)) | ~p(s(bool,X1))) ) | $spl22),
  inference(cnf_transformation,[],[f2392_D])).
fof(f2392_D,plain,(
  ( ! [X2,X1] : (~sP7(X2,X1,'T_0') | p(s(bool,X2)) | ~p(s(bool,X1))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f2484,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,X1,'T_0')) )),
  inference(condensation,[],[f2483])).
fof(f2483,plain,(
  ( ! [X59,X57,X58,X56] : (~p(s(bool,X57)) | p(s(bool,X58)) | ~sP7(X58,X57,'T_0') | s(bool,X56) = s(bool,X59) | s(bool,'T_0') = s(bool,X59)) )),
  inference(subsumption_resolution,[],[f2108,f832])).
fof(f2108,plain,(
  ( ! [X59,X57,X58,X56] : (~p(s(bool,X56)) | ~p(s(bool,X57)) | p(s(bool,X58)) | ~sP7(X58,X57,'T_0') | s(bool,X56) = s(bool,X59) | s(bool,'T_0') = s(bool,X59)) )),
  inference(superposition,[],[f433,f772])).
fof(f2389,plain,(
  ( ! [X2,X1] : (~sP6(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl20),
  inference(cnf_transformation,[],[f2389_D])).
fof(f2389_D,plain,(
  ( ! [X2,X1] : (~sP6(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f2482,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X1,'T_0',X2)) )),
  inference(condensation,[],[f2481])).
fof(f2481,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,X53)) | p(s(bool,X54)) | ~sP6(X53,'T_0',X54) | s(bool,X52) = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) )),
  inference(subsumption_resolution,[],[f2107,f832])).
fof(f2107,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,X52)) | ~p(s(bool,X53)) | p(s(bool,X54)) | ~sP6(X53,'T_0',X54) | s(bool,X52) = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) )),
  inference(superposition,[],[f424,f772])).
fof(f2386,plain,(
  ( ! [X2,X1] : (~sP3(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl18),
  inference(cnf_transformation,[],[f2386_D])).
fof(f2386_D,plain,(
  ( ! [X2,X1] : (~sP3(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f2480,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP3(X1,'T_0',X2)) )),
  inference(condensation,[],[f2479])).
fof(f2479,plain,(
  ( ! [X50,X48,X51,X49] : (~p(s(bool,X49)) | p(s(bool,X50)) | ~sP3(X49,'T_0',X50) | s(bool,X48) = s(bool,X51) | s(bool,'T_0') = s(bool,X51)) )),
  inference(subsumption_resolution,[],[f2106,f832])).
fof(f2106,plain,(
  ( ! [X50,X48,X51,X49] : (~p(s(bool,X48)) | ~p(s(bool,X49)) | p(s(bool,X50)) | ~sP3(X49,'T_0',X50) | s(bool,X48) = s(bool,X51) | s(bool,'T_0') = s(bool,X51)) )),
  inference(superposition,[],[f405,f772])).
fof(f2383,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl16),
  inference(cnf_transformation,[],[f2383_D])).
fof(f2383_D,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f2478,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X1,'T_0',X2)) )),
  inference(condensation,[],[f2477])).
fof(f2477,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,X45)) | p(s(bool,X46)) | ~sP4(X45,'T_0',X46) | s(bool,X44) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f2105,f832])).
fof(f2105,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,X44)) | ~p(s(bool,X45)) | p(s(bool,X46)) | ~sP4(X45,'T_0',X46) | s(bool,X44) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f401,f772])).
fof(f2380,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~sP3(X2,'T_0',X1) | sP5(X1,X2,'T_0')) ) | $spl14),
  inference(cnf_transformation,[],[f2380_D])).
fof(f2380_D,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~sP3(X2,'T_0',X1) | sP5(X1,X2,'T_0')) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f2476,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP5(X1,X2,'T_0') | ~sP4(X1,'T_0',X2) | ~sP3(X2,'T_0',X1)) )),
  inference(subsumption_resolution,[],[f2475,f507])).
fof(f2475,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP5(X1,X2,'T_0') | ~sP4(X1,'T_0',X2) | ~sP3(X2,'T_0',X1) | ~sP2('T_0',X2,X1)) )),
  inference(condensation,[],[f2474])).
fof(f2474,plain,(
  ( ! [X43,X41,X42,X40] : (sP5(X41,X42,'T_0') | ~sP4(X41,'T_0',X42) | ~sP3(X42,'T_0',X41) | ~sP2('T_0',X42,X41) | s(bool,X40) = s(bool,X43) | s(bool,'T_0') = s(bool,X43)) )),
  inference(subsumption_resolution,[],[f2104,f832])).
fof(f2104,plain,(
  ( ! [X43,X41,X42,X40] : (~p(s(bool,X40)) | sP5(X41,X42,'T_0') | ~sP4(X41,'T_0',X42) | ~sP3(X42,'T_0',X41) | ~sP2('T_0',X42,X41) | s(bool,X40) = s(bool,X43) | s(bool,'T_0') = s(bool,X43)) )),
  inference(superposition,[],[f398,f772])).
fof(f2376,plain,(
  ( ! [X2,X1] : (~sP5(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl12),
  inference(cnf_transformation,[],[f2376_D])).
fof(f2376_D,plain,(
  ( ! [X2,X1] : (~sP5(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f2473,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5(X1,'T_0',X2)) )),
  inference(condensation,[],[f2472])).
fof(f2472,plain,(
  ( ! [X39,X37,X38,X36] : (~p(s(bool,X37)) | p(s(bool,X38)) | ~sP5(X37,'T_0',X38) | s(bool,X36) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) )),
  inference(subsumption_resolution,[],[f2103,f832])).
fof(f2103,plain,(
  ( ! [X39,X37,X38,X36] : (~p(s(bool,X36)) | ~p(s(bool,X37)) | p(s(bool,X38)) | ~sP5(X37,'T_0',X38) | s(bool,X36) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) )),
  inference(superposition,[],[f397,f772])).
fof(f2373,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) | $spl10),
  inference(cnf_transformation,[],[f2373_D])).
fof(f2373_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f2471,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f2470])).
fof(f2470,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X34)) | sP1(X34,'T_0') | s(bool,X33) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) )),
  inference(subsumption_resolution,[],[f2102,f832])).
fof(f2102,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | p(s(bool,X34)) | sP1(X34,'T_0') | s(bool,X33) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) )),
  inference(superposition,[],[f366,f772])).
fof(f2369,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) | $spl6),
  inference(cnf_transformation,[],[f2369_D])).
fof(f2369_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f2469,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f2468])).
fof(f2468,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X31)) | sP1('T_0',X31) | s(bool,X30) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(subsumption_resolution,[],[f2101,f832])).
fof(f2101,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,X30)) | p(s(bool,X31)) | sP1('T_0',X31) | s(bool,X30) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(superposition,[],[f363,f772])).
fof(f2421,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('T_0',X1,X2) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f2420])).
fof(f2420,plain,(
  ( ! [X99,X97,X98,X96] : (~sP8('T_0',X97,X98) | sP9('T_0',X97,X98) | s(bool,X96) = s(bool,X99) | s(bool,'T_0') = s(bool,X99)) )),
  inference(subsumption_resolution,[],[f2003,f832])).
fof(f2003,plain,(
  ( ! [X99,X97,X98,X96] : (~p(s(bool,X96)) | ~sP8('T_0',X97,X98) | sP9('T_0',X97,X98) | s(bool,X96) = s(bool,X99) | s(bool,'T_0') = s(bool,X99)) )),
  inference(superposition,[],[f469,f772])).
fof(f2418,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP10('T_0',X1,X2) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f2417])).
fof(f2417,plain,(
  ( ! [X94,X92,X95,X93] : (~sP10('T_0',X93,X94) | sP11('T_0',X93,X94) | s(bool,X92) = s(bool,X95) | s(bool,'T_0') = s(bool,X95)) )),
  inference(subsumption_resolution,[],[f2002,f832])).
fof(f2002,plain,(
  ( ! [X94,X92,X95,X93] : (~p(s(bool,X92)) | ~sP10('T_0',X93,X94) | sP11('T_0',X93,X94) | s(bool,X92) = s(bool,X95) | s(bool,'T_0') = s(bool,X95)) )),
  inference(superposition,[],[f468,f772])).
fof(f2415,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f2414])).
fof(f2414,plain,(
  ( ! [X90,X88,X91,X89] : (~p(s(bool,X89)) | p(s(bool,X90)) | ~sP11('T_0',X90,X89) | s(bool,X88) = s(bool,X91) | s(bool,'T_0') = s(bool,X91)) )),
  inference(subsumption_resolution,[],[f2001,f832])).
fof(f2001,plain,(
  ( ! [X90,X88,X91,X89] : (~p(s(bool,X88)) | ~p(s(bool,X89)) | p(s(bool,X90)) | ~sP11('T_0',X90,X89) | s(bool,X88) = s(bool,X91) | s(bool,'T_0') = s(bool,X91)) )),
  inference(superposition,[],[f465,f772])).
fof(f2412,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f2411])).
fof(f2411,plain,(
  ( ! [X87,X85,X86,X84] : (~p(s(bool,X85)) | sP11('T_0',X85,X86) | s(bool,X84) = s(bool,X87) | s(bool,'T_0') = s(bool,X87)) )),
  inference(subsumption_resolution,[],[f2000,f832])).
fof(f2000,plain,(
  ( ! [X87,X85,X86,X84] : (~p(s(bool,X84)) | ~p(s(bool,X85)) | sP11('T_0',X85,X86) | s(bool,X84) = s(bool,X87) | s(bool,'T_0') = s(bool,X87)) )),
  inference(superposition,[],[f464,f772])).
fof(f2409,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f2408])).
fof(f2408,plain,(
  ( ! [X80,X83,X81,X82] : (p(s(bool,X81)) | sP11('T_0',X82,X81) | s(bool,X80) = s(bool,X83) | s(bool,'T_0') = s(bool,X83)) )),
  inference(subsumption_resolution,[],[f1999,f832])).
fof(f1999,plain,(
  ( ! [X80,X83,X81,X82] : (~p(s(bool,X80)) | p(s(bool,X81)) | sP11('T_0',X82,X81) | s(bool,X80) = s(bool,X83) | s(bool,'T_0') = s(bool,X83)) )),
  inference(superposition,[],[f463,f772])).
fof(f2406,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP10(X2,X1,'T_0')) )),
  inference(condensation,[],[f2405])).
fof(f2405,plain,(
  ( ! [X78,X76,X79,X77] : (p(s(bool,X77)) | ~p(s(bool,X78)) | ~sP10(X78,X77,'T_0') | s(bool,X76) = s(bool,X79) | s(bool,'T_0') = s(bool,X79)) )),
  inference(subsumption_resolution,[],[f1998,f832])).
fof(f1998,plain,(
  ( ! [X78,X76,X79,X77] : (~p(s(bool,X76)) | p(s(bool,X77)) | ~p(s(bool,X78)) | ~sP10(X78,X77,'T_0') | s(bool,X76) = s(bool,X79) | s(bool,'T_0') = s(bool,X79)) )),
  inference(superposition,[],[f458,f772])).
fof(f2403,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f2402])).
fof(f2402,plain,(
  ( ! [X74,X72,X75,X73] : (p(s(bool,X73)) | p(s(bool,X74)) | ~sP9('T_0',X74,X73) | s(bool,X72) = s(bool,X75) | s(bool,'T_0') = s(bool,X75)) )),
  inference(subsumption_resolution,[],[f1997,f832])).
fof(f1997,plain,(
  ( ! [X74,X72,X75,X73] : (~p(s(bool,X72)) | p(s(bool,X73)) | p(s(bool,X74)) | ~sP9('T_0',X74,X73) | s(bool,X72) = s(bool,X75) | s(bool,'T_0') = s(bool,X75)) )),
  inference(superposition,[],[f448,f772])).
fof(f2400,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f2399])).
fof(f2399,plain,(
  ( ! [X70,X68,X71,X69] : (~p(s(bool,X69)) | sP9('T_0',X69,X70) | s(bool,X68) = s(bool,X71) | s(bool,'T_0') = s(bool,X71)) )),
  inference(subsumption_resolution,[],[f1996,f832])).
fof(f1996,plain,(
  ( ! [X70,X68,X71,X69] : (~p(s(bool,X68)) | ~p(s(bool,X69)) | sP9('T_0',X69,X70) | s(bool,X68) = s(bool,X71) | s(bool,'T_0') = s(bool,X71)) )),
  inference(superposition,[],[f447,f772])).
fof(f2397,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f2396])).
fof(f2396,plain,(
  ( ! [X66,X64,X67,X65] : (~p(s(bool,X65)) | sP9('T_0',X66,X65) | s(bool,X64) = s(bool,X67) | s(bool,'T_0') = s(bool,X67)) )),
  inference(subsumption_resolution,[],[f1995,f832])).
fof(f1995,plain,(
  ( ! [X66,X64,X67,X65] : (~p(s(bool,X64)) | ~p(s(bool,X65)) | sP9('T_0',X66,X65) | s(bool,X64) = s(bool,X67) | s(bool,'T_0') = s(bool,X67)) )),
  inference(superposition,[],[f446,f772])).
fof(f2394,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP8('T_0',X1,X2)) )),
  inference(condensation,[],[f2393])).
fof(f2393,plain,(
  ( ! [X61,X62,X60,X63] : (p(s(bool,X61)) | p(s(bool,X62)) | ~sP8('T_0',X61,X62) | s(bool,X60) = s(bool,X63) | s(bool,'T_0') = s(bool,X63)) )),
  inference(subsumption_resolution,[],[f1994,f832])).
fof(f1994,plain,(
  ( ! [X61,X62,X60,X63] : (~p(s(bool,X60)) | p(s(bool,X61)) | p(s(bool,X62)) | ~sP8('T_0',X61,X62) | s(bool,X60) = s(bool,X63) | s(bool,'T_0') = s(bool,X63)) )),
  inference(superposition,[],[f441,f772])).
fof(f2391,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,X1,'T_0')) )),
  inference(condensation,[],[f2390])).
fof(f2390,plain,(
  ( ! [X59,X57,X58,X56] : (~p(s(bool,X57)) | p(s(bool,X58)) | ~sP7(X58,X57,'T_0') | s(bool,X56) = s(bool,X59) | s(bool,'T_0') = s(bool,X59)) )),
  inference(subsumption_resolution,[],[f1993,f832])).
fof(f1993,plain,(
  ( ! [X59,X57,X58,X56] : (~p(s(bool,X56)) | ~p(s(bool,X57)) | p(s(bool,X58)) | ~sP7(X58,X57,'T_0') | s(bool,X56) = s(bool,X59) | s(bool,'T_0') = s(bool,X59)) )),
  inference(superposition,[],[f433,f772])).
fof(f2388,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X1,'T_0',X2)) )),
  inference(condensation,[],[f2387])).
fof(f2387,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,X53)) | p(s(bool,X54)) | ~sP6(X53,'T_0',X54) | s(bool,X52) = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) )),
  inference(subsumption_resolution,[],[f1992,f832])).
fof(f1992,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,X52)) | ~p(s(bool,X53)) | p(s(bool,X54)) | ~sP6(X53,'T_0',X54) | s(bool,X52) = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) )),
  inference(superposition,[],[f424,f772])).
fof(f2385,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP3(X1,'T_0',X2)) )),
  inference(condensation,[],[f2384])).
fof(f2384,plain,(
  ( ! [X50,X48,X51,X49] : (~p(s(bool,X49)) | p(s(bool,X50)) | ~sP3(X49,'T_0',X50) | s(bool,X48) = s(bool,X51) | s(bool,'T_0') = s(bool,X51)) )),
  inference(subsumption_resolution,[],[f1991,f832])).
fof(f1991,plain,(
  ( ! [X50,X48,X51,X49] : (~p(s(bool,X48)) | ~p(s(bool,X49)) | p(s(bool,X50)) | ~sP3(X49,'T_0',X50) | s(bool,X48) = s(bool,X51) | s(bool,'T_0') = s(bool,X51)) )),
  inference(superposition,[],[f405,f772])).
fof(f2382,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X1,'T_0',X2)) )),
  inference(condensation,[],[f2381])).
fof(f2381,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,X45)) | p(s(bool,X46)) | ~sP4(X45,'T_0',X46) | s(bool,X44) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f1990,f832])).
fof(f1990,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,X44)) | ~p(s(bool,X45)) | p(s(bool,X46)) | ~sP4(X45,'T_0',X46) | s(bool,X44) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f401,f772])).
fof(f2379,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP5(X1,X2,'T_0') | ~sP4(X1,'T_0',X2) | ~sP3(X2,'T_0',X1)) )),
  inference(subsumption_resolution,[],[f2378,f507])).
fof(f2378,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP5(X1,X2,'T_0') | ~sP4(X1,'T_0',X2) | ~sP3(X2,'T_0',X1) | ~sP2('T_0',X2,X1)) )),
  inference(condensation,[],[f2377])).
fof(f2377,plain,(
  ( ! [X43,X41,X42,X40] : (sP5(X41,X42,'T_0') | ~sP4(X41,'T_0',X42) | ~sP3(X42,'T_0',X41) | ~sP2('T_0',X42,X41) | s(bool,X40) = s(bool,X43) | s(bool,'T_0') = s(bool,X43)) )),
  inference(subsumption_resolution,[],[f1989,f832])).
fof(f1989,plain,(
  ( ! [X43,X41,X42,X40] : (~p(s(bool,X40)) | sP5(X41,X42,'T_0') | ~sP4(X41,'T_0',X42) | ~sP3(X42,'T_0',X41) | ~sP2('T_0',X42,X41) | s(bool,X40) = s(bool,X43) | s(bool,'T_0') = s(bool,X43)) )),
  inference(superposition,[],[f398,f772])).
fof(f2375,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5(X1,'T_0',X2)) )),
  inference(condensation,[],[f2374])).
fof(f2374,plain,(
  ( ! [X39,X37,X38,X36] : (~p(s(bool,X37)) | p(s(bool,X38)) | ~sP5(X37,'T_0',X38) | s(bool,X36) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) )),
  inference(subsumption_resolution,[],[f1988,f832])).
fof(f1988,plain,(
  ( ! [X39,X37,X38,X36] : (~p(s(bool,X36)) | ~p(s(bool,X37)) | p(s(bool,X38)) | ~sP5(X37,'T_0',X38) | s(bool,X36) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) )),
  inference(superposition,[],[f397,f772])).
fof(f2372,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f2371])).
fof(f2371,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X34)) | sP1(X34,'T_0') | s(bool,X33) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) )),
  inference(subsumption_resolution,[],[f1987,f832])).
fof(f1987,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | p(s(bool,X34)) | sP1(X34,'T_0') | s(bool,X33) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) )),
  inference(superposition,[],[f366,f772])).
fof(f2368,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f2367])).
fof(f2367,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X31)) | sP1('T_0',X31) | s(bool,X30) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(subsumption_resolution,[],[f1986,f832])).
fof(f1986,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,X30)) | p(s(bool,X31)) | sP1('T_0',X31) | s(bool,X30) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(superposition,[],[f363,f772])).
fof(f477,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)))) | $spl0),
  inference(cnf_transformation,[],[f477_D])).
fof(f477_D,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f479,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | $spl2),
  inference(cnf_transformation,[],[f479_D])).
fof(f479_D,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f266,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15))))),
  inference(cnf_transformation,[],[f163])).
fof(f163,plain,(
  (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)))) | (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) & (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13,sK14,sK15])],[f162])).
fof(f162,plain,(
  ? [X0,X1,X2,X3] : ((p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X3)))) | (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X3)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))))) & (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X3)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X3)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))))),
  inference(flattening,[],[f161])).
fof(f161,plain,(
  ? [X0,X1,X2,X3] : ((p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X3)))) | (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X3)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))))) & (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X3)))) | (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X3)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))))))),
  inference(nnf_transformation,[],[f118])).
fof(f118,plain,(
  ? [X0,X1,X2,X3] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X3)))) <~> (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X3)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))))),
  inference(ennf_transformation,[],[f48])).
fof(f48,plain,(
  ~! [X0,X1,X2,X3] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X3)))) <=> (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X3)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))))),
  inference(flattening,[],[f47])).
fof(f47,plain,(
  ~! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X3)))) <=> (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X3)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8))),s(fun(X6,bool),X9)))) <=> (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X9)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X9))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8))),s(fun(X6,bool),X9)))) <=> (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X9)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X9))))))),
  file('thm.pred_set.UNION_SUBSET','thm.pred_set.UNION_SUBSET')).
fof(f481,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))) | $spl4),
  inference(cnf_transformation,[],[f481_D])).
fof(f481_D,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f267,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15))))),
  inference(cnf_transformation,[],[f163])).
fof(f478,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)))) | $spl1),
  inference(cnf_transformation,[],[f478_D])).
fof(f478_D,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f480,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | $spl3),
  inference(cnf_transformation,[],[f480_D])).
fof(f480_D,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f482,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))) | $spl5),
  inference(cnf_transformation,[],[f482_D])).
fof(f482_D,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f268,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15))))),
  inference(cnf_transformation,[],[f163])).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 1827
% Passive clauses: 4175
% Generated clauses: 31326
% Final active clauses: 560
% Final passive clauses: 553
% Input formulas: 45
% Initial clauses: 202
% 
% Duplicate literals: 389
% Trivial inequalities: 63
% Fw subsumption resolutions: 1523
% Fw demodulations: 5240
% Bw demodulations: 128
% Condensations: 318
% Global subsumptions: 3
% 
% Simple tautologies: 124
% Equational tautologies: 338
% Forward subsumptions: 17725
% Bw demodulations to eq. taut.: 1
% 
% Unit resulting resolution: 19260
% Factoring: 1
% Forward superposition: 1481
% Backward superposition: 2499
% Self superposition: 36
% Equality factoring: 40
% Equality resolution: 2
% 
% Split clauses: 52
% Split components: 110
% SAT solver clauses: 5660
% SAT solver unit clauses: 272
% SAT solver binary clauses: 2430
% 
% Sat splits: 52
% Sat splitting refutations: 1376
% 
% Memory used [KB]: 10874
% Time elapsed: 0.904 s
% ------------------------------
% ------------------------------
% Success in time 2.211 s
