Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.ring
% (45293)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 612
% Passive clauses: 6486
% Generated clauses: 23876
% Final active clauses: 606
% Final passive clauses: 5628
% Input formulas: 260
% Initial clauses: 380
% 
% Duplicate literals: 82
% Trivial inequalities: 81
% Fw subsumption resolutions: 762
% Fw demodulations: 4136
% Bw demodulations: 6
% 
% Simple tautologies: 197
% Equational tautologies: 27
% Forward subsumptions: 12070
% 
% Binary resolution: 3612
% Factoring: 350
% Forward superposition: 6325
% Backward superposition: 8245
% Self superposition: 78
% Equality factoring: 2
% Equality resolution: 21
% 
% Split clauses: 163
% Split components: 327
% SAT solver clauses: 281
% SAT solver unit clauses: 19
% SAT solver binary clauses: 61
% 
% Sat splits: 163
% 
% TWLsolver clauses: 281
% TWLsolver calls for satisfiability: 231
% 
% Memory used [KB]: 19573
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.ring
% (45307)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 338
% Passive clauses: 2459
% Generated clauses: 16557
% Final active clauses: 335
% Final passive clauses: 2060
% Input formulas: 260
% Initial clauses: 380
% 
% Duplicate literals: 59
% Trivial inequalities: 8
% Fw subsumption resolutions: 361
% Fw demodulations: 1932
% Bw demodulations: 64
% Condensations: 91
% Global subsumptions: 18
% 
% Simple tautologies: 158
% Equational tautologies: 11
% Forward subsumptions: 8662
% Fw demodulations to eq. taut.: 4
% 
% Binary resolution: 1583
% Factoring: 216
% Forward superposition: 4795
% Backward superposition: 4472
% Self superposition: 94
% Equality factoring: 2
% Equality resolution: 5
% 
% Split clauses: 18
% Split components: 36
% SAT solver clauses: 3187
% SAT solver unit clauses: 579
% SAT solver binary clauses: 78
% SAT solver learnt clauses: 36
% SAT solver learnt literals: 10
% 
% Sat splits: 18
% 
% TWLsolver clauses: 3115
% TWLsolver calls for satisfiability: 1364
% 
% Memory used [KB]: 26609
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.ring
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.ring
% SZS output start Proof for thm.ring
fof(f77051,plain,(
  $false),
  inference(subsumption_resolution,[],[f77050,f1033])).
fof(f1033,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),X3),s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1))))) = s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),X3),s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f462])).
fof(f462,plain,(
  ! [X0,X1,X2,X3] : s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),X3),s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1))))) = s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),X3),s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f461])).
fof(f461,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),X3),s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1))))) = s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),X3),s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f179])).
fof(f179,axiom,(
  ! [X6] : ! [X19,X7,X8] : s('type.ring.ring'(X6),'const.ring.ring_R|49|_fupd_2'(s(fun(X6,X6),X8),s('type.ring.ring'(X6),'const.ring.ring_R|48|_fupd_2'(s(fun(X6,X6),X7),s('type.ring.ring'(X6),X19))))) = s('type.ring.ring'(X6),'const.ring.ring_R|48|_fupd_2'(s(fun(X6,X6),X7),s('type.ring.ring'(X6),'const.ring.ring_R|49|_fupd_2'(s(fun(X6,X6),X8),s('type.ring.ring'(X6),X19)))))),
  file('thm.ring.ring_fupdcanon_comp|split|0','thm.ring.ring_5|49|_prove_ltb')).
fof(f77050,plain,(
  s('type.ring.ring'(sK6),'const.ring.ring_R|49|_fupd_2'(s(fun(sK6,sK6),sK8),s('type.ring.ring'(sK6),'const.ring.ring_R|48|_fupd_2'(s(fun(sK6,sK6),sK7),s('type.ring.ring'(sK6),sK23('const.combin.o_2'(s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK6,sK6),sK8))),s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK6,sK6),sK7)))),'const.combin.o_2'(s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK6,sK6),sK7))),s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK6,sK6),sK8)))),'type.ring.ring'(sK6),'type.ring.ring'(sK6))))))) != s('type.ring.ring'(sK6),'const.ring.ring_R|48|_fupd_2'(s(fun(sK6,sK6),sK7),s('type.ring.ring'(sK6),'const.ring.ring_R|49|_fupd_2'(s(fun(sK6,sK6),sK8),s('type.ring.ring'(sK6),sK23('const.combin.o_2'(s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK6,sK6),sK8))),s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK6,sK6),sK7)))),'const.combin.o_2'(s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK6,sK6),sK7))),s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK6,sK6),sK8)))),'type.ring.ring'(sK6),'type.ring.ring'(sK6)))))))),
  inference(forward_demodulation,[],[f77049,f989])).
fof(f989,plain,(
  ( ! [X2,X0,X1] : (s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),X1),s('type.ring.ring'(X0),X2))) = s('type.ring.ring'(X0),happ(s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_R|48|_fupd_1'(s(fun(X0,X0),X1))),s('type.ring.ring'(X0),X2)))) )),
  inference(cnf_transformation,[],[f382])).
fof(f382,plain,(
  ! [X0,X1,X2] : s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),X1),s('type.ring.ring'(X0),X2))) = s('type.ring.ring'(X0),happ(s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_R|48|_fupd_1'(s(fun(X0,X0),X1))),s('type.ring.ring'(X0),X2)))),
  inference(flattening,[],[f381])).
fof(f381,plain,(
  ! [X0] : ! [X1,X2] : s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),X1),s('type.ring.ring'(X0),X2))) = s('type.ring.ring'(X0),happ(s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_R|48|_fupd_1'(s(fun(X0,X0),X1))),s('type.ring.ring'(X0),X2)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X6] : ! [X11,X12] : s('type.ring.ring'(X6),'const.ring.ring_R|48|_fupd_2'(s(fun(X6,X6),X11),s('type.ring.ring'(X6),X12))) = s('type.ring.ring'(X6),happ(s(fun('type.ring.ring'(X6),'type.ring.ring'(X6)),'const.ring.ring_R|48|_fupd_1'(s(fun(X6,X6),X11))),s('type.ring.ring'(X6),X12)))),
  file('thm.ring.ring_fupdcanon_comp|split|0','ARITY_1_2_const.ring.ring_R|48|_fupd')).
fof(f77049,plain,(
  s('type.ring.ring'(sK6),'const.ring.ring_R|48|_fupd_2'(s(fun(sK6,sK6),sK7),s('type.ring.ring'(sK6),'const.ring.ring_R|49|_fupd_2'(s(fun(sK6,sK6),sK8),s('type.ring.ring'(sK6),sK23('const.combin.o_2'(s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK6,sK6),sK8))),s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK6,sK6),sK7)))),'const.combin.o_2'(s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK6,sK6),sK7))),s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK6,sK6),sK8)))),'type.ring.ring'(sK6),'type.ring.ring'(sK6))))))) != s('type.ring.ring'(sK6),'const.ring.ring_R|49|_fupd_2'(s(fun(sK6,sK6),sK8),s('type.ring.ring'(sK6),happ(s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK6,sK6),sK7))),s('type.ring.ring'(sK6),sK23('const.combin.o_2'(s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK6,sK6),sK8))),s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK6,sK6),sK7)))),'const.combin.o_2'(s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK6,sK6),sK7))),s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK6,sK6),sK8)))),'type.ring.ring'(sK6),'type.ring.ring'(sK6)))))))),
  inference(forward_demodulation,[],[f77048,f988])).
fof(f988,plain,(
  ( ! [X2,X0,X1] : (s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),X1),s('type.ring.ring'(X0),X2))) = s('type.ring.ring'(X0),happ(s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_R|49|_fupd_1'(s(fun(X0,X0),X1))),s('type.ring.ring'(X0),X2)))) )),
  inference(cnf_transformation,[],[f380])).
fof(f380,plain,(
  ! [X0,X1,X2] : s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),X1),s('type.ring.ring'(X0),X2))) = s('type.ring.ring'(X0),happ(s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_R|49|_fupd_1'(s(fun(X0,X0),X1))),s('type.ring.ring'(X0),X2)))),
  inference(flattening,[],[f379])).
fof(f379,plain,(
  ! [X0] : ! [X1,X2] : s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),X1),s('type.ring.ring'(X0),X2))) = s('type.ring.ring'(X0),happ(s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_R|49|_fupd_1'(s(fun(X0,X0),X1))),s('type.ring.ring'(X0),X2)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X6] : ! [X11,X12] : s('type.ring.ring'(X6),'const.ring.ring_R|49|_fupd_2'(s(fun(X6,X6),X11),s('type.ring.ring'(X6),X12))) = s('type.ring.ring'(X6),happ(s(fun('type.ring.ring'(X6),'type.ring.ring'(X6)),'const.ring.ring_R|49|_fupd_1'(s(fun(X6,X6),X11))),s('type.ring.ring'(X6),X12)))),
  file('thm.ring.ring_fupdcanon_comp|split|0','ARITY_1_2_const.ring.ring_R|49|_fupd')).
fof(f77048,plain,(
  s('type.ring.ring'(sK6),'const.ring.ring_R|49|_fupd_2'(s(fun(sK6,sK6),sK8),s('type.ring.ring'(sK6),happ(s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK6,sK6),sK7))),s('type.ring.ring'(sK6),sK23('const.combin.o_2'(s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK6,sK6),sK8))),s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK6,sK6),sK7)))),'const.combin.o_2'(s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK6,sK6),sK7))),s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK6,sK6),sK8)))),'type.ring.ring'(sK6),'type.ring.ring'(sK6))))))) != s('type.ring.ring'(sK6),'const.ring.ring_R|48|_fupd_2'(s(fun(sK6,sK6),sK7),s('type.ring.ring'(sK6),happ(s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK6,sK6),sK8))),s('type.ring.ring'(sK6),sK23('const.combin.o_2'(s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK6,sK6),sK8))),s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK6,sK6),sK7)))),'const.combin.o_2'(s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK6,sK6),sK7))),s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK6,sK6),sK8)))),'type.ring.ring'(sK6),'type.ring.ring'(sK6)))))))),
  inference(forward_demodulation,[],[f77047,f16138])).
fof(f16138,plain,(
  ( ! [X61,X64,X62,X65,X63] : (s('type.ring.ring'(X61),'const.ring.ring_R|49|_fupd_2'(s(fun(X61,X61),X62),s('type.ring.ring'(X61),happ(s(fun(X63,'type.ring.ring'(X61)),X64),s(X63,X65))))) = s('type.ring.ring'(X61),'const.combin.o_3'(s(fun('type.ring.ring'(X61),'type.ring.ring'(X61)),'const.ring.ring_R|49|_fupd_1'(s(fun(X61,X61),X62))),s(fun(X63,'type.ring.ring'(X61)),X64),s(X63,X65)))) )),
  inference(superposition,[],[f1208,f988])).
fof(f1208,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5)))))) )),
  inference(cnf_transformation,[],[f673])).
fof(f673,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5)))))),
  inference(flattening,[],[f672])).
fof(f672,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5)))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X9,X6,X10] : ! [X8,X7,X14] : s(X9,'const.combin.o_3'(s(fun(X6,X9),X8),s(fun(X10,X6),X7),s(X10,X14))) = s(X9,happ(s(fun(X6,X9),X8),s(X6,happ(s(fun(X10,X6),X7),s(X10,X14)))))),
  file('thm.ring.ring_fupdcanon_comp|split|0','thm.combin.o_THM')).
fof(f77047,plain,(
  s('type.ring.ring'(sK6),'const.ring.ring_R|48|_fupd_2'(s(fun(sK6,sK6),sK7),s('type.ring.ring'(sK6),happ(s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK6,sK6),sK8))),s('type.ring.ring'(sK6),sK23('const.combin.o_2'(s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK6,sK6),sK8))),s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK6,sK6),sK7)))),'const.combin.o_2'(s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK6,sK6),sK7))),s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK6,sK6),sK8)))),'type.ring.ring'(sK6),'type.ring.ring'(sK6))))))) != s('type.ring.ring'(sK6),'const.combin.o_3'(s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK6,sK6),sK8))),s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK6,sK6),sK7))),s('type.ring.ring'(sK6),sK23('const.combin.o_2'(s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK6,sK6),sK8))),s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK6,sK6),sK7)))),'const.combin.o_2'(s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK6,sK6),sK7))),s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK6,sK6),sK8)))),'type.ring.ring'(sK6),'type.ring.ring'(sK6)))))),
  inference(forward_demodulation,[],[f77046,f16139])).
fof(f16139,plain,(
  ( ! [X70,X68,X66,X69,X67] : (s('type.ring.ring'(X66),'const.ring.ring_R|48|_fupd_2'(s(fun(X66,X66),X67),s('type.ring.ring'(X66),happ(s(fun(X68,'type.ring.ring'(X66)),X69),s(X68,X70))))) = s('type.ring.ring'(X66),'const.combin.o_3'(s(fun('type.ring.ring'(X66),'type.ring.ring'(X66)),'const.ring.ring_R|48|_fupd_1'(s(fun(X66,X66),X67))),s(fun(X68,'type.ring.ring'(X66)),X69),s(X68,X70)))) )),
  inference(superposition,[],[f1208,f989])).
fof(f77046,plain,(
  s('type.ring.ring'(sK6),'const.combin.o_3'(s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK6,sK6),sK8))),s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK6,sK6),sK7))),s('type.ring.ring'(sK6),sK23('const.combin.o_2'(s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK6,sK6),sK8))),s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK6,sK6),sK7)))),'const.combin.o_2'(s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK6,sK6),sK7))),s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK6,sK6),sK8)))),'type.ring.ring'(sK6),'type.ring.ring'(sK6))))) != s('type.ring.ring'(sK6),'const.combin.o_3'(s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK6,sK6),sK7))),s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK6,sK6),sK8))),s('type.ring.ring'(sK6),sK23('const.combin.o_2'(s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK6,sK6),sK8))),s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK6,sK6),sK7)))),'const.combin.o_2'(s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK6,sK6),sK7))),s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK6,sK6),sK8)))),'type.ring.ring'(sK6),'type.ring.ring'(sK6)))))),
  inference(forward_demodulation,[],[f77045,f1212])).
fof(f1212,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))) )),
  inference(cnf_transformation,[],[f681])).
fof(f681,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(flattening,[],[f680])).
fof(f680,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X9,X10,X6] : ! [X11,X12,X13] : s(X9,'const.combin.o_3'(s(fun(X10,X9),X11),s(fun(X6,X10),X12),s(X6,X13))) = s(X9,happ(s(fun(X6,X9),'const.combin.o_2'(s(fun(X10,X9),X11),s(fun(X6,X10),X12))),s(X6,X13)))),
  file('thm.ring.ring_fupdcanon_comp|split|0','ARITY_2_3_const.combin.o')).
fof(f77045,plain,(
  s('type.ring.ring'(sK6),'const.combin.o_3'(s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK6,sK6),sK7))),s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK6,sK6),sK8))),s('type.ring.ring'(sK6),sK23('const.combin.o_2'(s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK6,sK6),sK8))),s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK6,sK6),sK7)))),'const.combin.o_2'(s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK6,sK6),sK7))),s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK6,sK6),sK8)))),'type.ring.ring'(sK6),'type.ring.ring'(sK6))))) != s('type.ring.ring'(sK6),happ(s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.combin.o_2'(s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK6,sK6),sK8))),s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK6,sK6),sK7))))),s('type.ring.ring'(sK6),sK23('const.combin.o_2'(s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK6,sK6),sK8))),s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK6,sK6),sK7)))),'const.combin.o_2'(s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK6,sK6),sK7))),s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK6,sK6),sK8)))),'type.ring.ring'(sK6),'type.ring.ring'(sK6)))))),
  inference(forward_demodulation,[],[f76852,f1212])).
fof(f76852,plain,(
  s('type.ring.ring'(sK6),happ(s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.combin.o_2'(s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK6,sK6),sK8))),s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK6,sK6),sK7))))),s('type.ring.ring'(sK6),sK23('const.combin.o_2'(s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK6,sK6),sK8))),s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK6,sK6),sK7)))),'const.combin.o_2'(s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK6,sK6),sK7))),s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK6,sK6),sK8)))),'type.ring.ring'(sK6),'type.ring.ring'(sK6))))) != s('type.ring.ring'(sK6),happ(s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.combin.o_2'(s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK6,sK6),sK7))),s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK6,sK6),sK8))))),s('type.ring.ring'(sK6),sK23('const.combin.o_2'(s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK6,sK6),sK8))),s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK6,sK6),sK7)))),'const.combin.o_2'(s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK6,sK6),sK7))),s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK6,sK6),sK8)))),'type.ring.ring'(sK6),'type.ring.ring'(sK6)))))),
  inference(unit_resulting_resolution,[],[f879,f1139])).
fof(f1139,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,sK23(X3,X2,X1,X0)))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,sK23(X3,X2,X1,X0)))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)) )),
  inference(cnf_transformation,[],[f855])).
fof(f855,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X5] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) & (s(X0,happ(s(fun(X1,X0),X2),s(X1,sK23(X3,X2,X1,X0)))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,sK23(X3,X2,X1,X0)))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK23])],[f854])).
fof(f854,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X5] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) & (? [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(rectify,[],[f853])).
fof(f853,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4)))) & (? [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(nnf_transformation,[],[f603])).
fof(f603,plain,(
  ! [X0,X1,X2,X3] : (s(fun(X1,X0),X2) = s(fun(X1,X0),X3) <=> ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))))),
  inference(flattening,[],[f602])).
fof(f602,plain,(
  ! [X0,X1] : ! [X2,X3] : (s(fun(X1,X0),X2) = s(fun(X1,X0),X3) <=> ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))))),
  inference(rectify,[],[f61])).
fof(f61,axiom,(
  ! [X9,X6] : ! [X8,X7] : (s(fun(X6,X9),X7) = s(fun(X6,X9),X8) <=> ! [X14] : s(X9,happ(s(fun(X6,X9),X7),s(X6,X14))) = s(X9,happ(s(fun(X6,X9),X8),s(X6,X14))))),
  file('thm.ring.ring_fupdcanon_comp|split|0','thm.bool.FUN_EQ_THM')).
fof(f879,plain,(
  s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.combin.o_2'(s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK6,sK6),sK8))),s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK6,sK6),sK7))))) != s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.combin.o_2'(s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK6,sK6),sK7))),s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK6,sK6),sK8)))))),
  inference(cnf_transformation,[],[f787])).
fof(f787,plain,(
  s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.combin.o_2'(s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK6,sK6),sK8))),s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK6,sK6),sK7))))) != s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.combin.o_2'(s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK6,sK6),sK7))),s(fun('type.ring.ring'(sK6),'type.ring.ring'(sK6)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK6,sK6),sK8)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK6,sK7,sK8])],[f744])).
fof(f744,plain,(
  ? [X0,X1,X2] : s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.combin.o_2'(s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_R|49|_fupd_1'(s(fun(X0,X0),X2))),s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_R|48|_fupd_1'(s(fun(X0,X0),X1))))) != s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.combin.o_2'(s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_R|48|_fupd_1'(s(fun(X0,X0),X1))),s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_R|49|_fupd_1'(s(fun(X0,X0),X2)))))),
  inference(ennf_transformation,[],[f263])).
fof(f263,plain,(
  ~! [X0,X1,X2] : s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.combin.o_2'(s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_R|49|_fupd_1'(s(fun(X0,X0),X2))),s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_R|48|_fupd_1'(s(fun(X0,X0),X1))))) = s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.combin.o_2'(s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_R|48|_fupd_1'(s(fun(X0,X0),X1))),s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_R|49|_fupd_1'(s(fun(X0,X0),X2)))))),
  inference(flattening,[],[f262])).
fof(f262,plain,(
  ~! [X0] : ! [X1,X2] : s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.combin.o_2'(s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_R|49|_fupd_1'(s(fun(X0,X0),X2))),s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_R|48|_fupd_1'(s(fun(X0,X0),X1))))) = s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.combin.o_2'(s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_R|48|_fupd_1'(s(fun(X0,X0),X1))),s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_R|49|_fupd_1'(s(fun(X0,X0),X2)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : s(fun('type.ring.ring'(X6),'type.ring.ring'(X6)),'const.combin.o_2'(s(fun('type.ring.ring'(X6),'type.ring.ring'(X6)),'const.ring.ring_R|49|_fupd_1'(s(fun(X6,X6),X8))),s(fun('type.ring.ring'(X6),'type.ring.ring'(X6)),'const.ring.ring_R|48|_fupd_1'(s(fun(X6,X6),X7))))) = s(fun('type.ring.ring'(X6),'type.ring.ring'(X6)),'const.combin.o_2'(s(fun('type.ring.ring'(X6),'type.ring.ring'(X6)),'const.ring.ring_R|48|_fupd_1'(s(fun(X6,X6),X7))),s(fun('type.ring.ring'(X6),'type.ring.ring'(X6)),'const.ring.ring_R|49|_fupd_1'(s(fun(X6,X6),X8)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : s(fun('type.ring.ring'(X6),'type.ring.ring'(X6)),'const.combin.o_2'(s(fun('type.ring.ring'(X6),'type.ring.ring'(X6)),'const.ring.ring_R|49|_fupd_1'(s(fun(X6,X6),X8))),s(fun('type.ring.ring'(X6),'type.ring.ring'(X6)),'const.ring.ring_R|48|_fupd_1'(s(fun(X6,X6),X7))))) = s(fun('type.ring.ring'(X6),'type.ring.ring'(X6)),'const.combin.o_2'(s(fun('type.ring.ring'(X6),'type.ring.ring'(X6)),'const.ring.ring_R|48|_fupd_1'(s(fun(X6,X6),X7))),s(fun('type.ring.ring'(X6),'type.ring.ring'(X6)),'const.ring.ring_R|49|_fupd_1'(s(fun(X6,X6),X8)))))),
  file('thm.ring.ring_fupdcanon_comp|split|0','thm.ring.ring_fupdcanon_comp|split|0')).
% SZS output end Proof for thm.ring
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 419
% Passive clauses: 4064
% Generated clauses: 76154
% Final active clauses: 412
% Final passive clauses: 3456
% Input formulas: 260
% Initial clauses: 380
% 
% Duplicate literals: 3
% Trivial inequalities: 11
% Fw subsumption resolutions: 100
% Fw demodulations: 29762
% Bw demodulations: 136
% Condensations: 8
% 
% Simple tautologies: 92
% Equational tautologies: 40
% Forward subsumptions: 15821
% Fw demodulations to eq. taut.: 25888
% 
% Unit resulting resolution: 29285
% Factoring: 1
% Forward superposition: 7991
% Backward superposition: 8327
% Self superposition: 128
% Equality factoring: 2
% Equality resolution: 7
% 
% Split clauses: 12
% Split components: 24
% SAT solver clauses: 4250
% SAT solver unit clauses: 2128
% SAT solver binary clauses: 1786
% 
% Sat splits: 12
% Sat splitting refutations: 160
% 
% Memory used [KB]: 69082
% Time elapsed: 2.639 s
% ------------------------------
% ------------------------------
% Success in time 3.96 s
