Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.transc
% (78108)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 164
% Passive clauses: 2113
% Generated clauses: 21767
% Final active clauses: 153
% Final passive clauses: 1760
% Input formulas: 39
% Initial clauses: 54
% 
% Duplicate literals: 1468
% Trivial inequalities: 100
% Fw subsumption resolutions: 2871
% Fw demodulations: 3207
% Bw demodulations: 9
% 
% Simple tautologies: 342
% Equational tautologies: 20
% Forward subsumptions: 10266
% 
% Binary resolution: 362
% Forward superposition: 8330
% Backward superposition: 4990
% Self superposition: 150
% Equality factoring: 34
% Equality resolution: 26
% 
% Split clauses: 823
% Split components: 1652
% SAT solver clauses: 1534
% SAT solver unit clauses: 6
% SAT solver binary clauses: 56
% 
% Sat splits: 823
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 1533
% TWLsolver calls for satisfiability: 201
% 
% Memory used [KB]: 14200
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.transc
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.transc
% SZS output start Proof for thm.transc
fof(f3413,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f3400,f118,f176,f120,f178,f169,f212,f229,f231,f116,f174,f131,f188,f293,f296,f117,f172,f135,f190,f373,f381,f127,f185])).
fof(f185,plain,(
  ( ! [X0] : (p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'const.transc.cos_0'),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.sin_1'(s('type.realax.real',X0))))),s('type.realax.real',X0))))) ) | $spl20),
  inference(cnf_transformation,[],[f185_D])).
fof(f185_D,plain,(
  ( ! [X0] : (p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'const.transc.cos_0'),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.sin_1'(s('type.realax.real',X0))))),s('type.realax.real',X0))))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f127,plain,(
  ( ! [X0] : (p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'const.transc.cos_0'),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.sin_1'(s('type.realax.real',X0))))),s('type.realax.real',X0))))) )),
  inference(cnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0] : p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'const.transc.cos_0'),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.sin_1'(s('type.realax.real',X0))))),s('type.realax.real',X0))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X6] : p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'const.transc.cos_0'),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.sin_1'(s('type.realax.real',X6))))),s('type.realax.real',X6))))),
  file('thm.transc.DIFF_COMPOSITE|split|9','thm.transc.DIFF_COS')).
fof(f381,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'const.transc.cos_0'),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.sin_1'(s('type.realax.real',X2))))),s('type.realax.real',X2)))) ) | $spl106),
  inference(cnf_transformation,[],[f381_D])).
fof(f381_D,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'const.transc.cos_0'),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.sin_1'(s('type.realax.real',X2))))),s('type.realax.real',X2)))) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f373,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'const.transc.cos_0'),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.sin_1'(s('type.realax.real',X2))))),s('type.realax.real',X2)))) ) | ($spl20 | $spl26)),
  inference(resolution,[],[f185,f188])).
fof(f190,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) | $spl30),
  inference(cnf_transformation,[],[f190_D])).
fof(f190_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f135,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f53])).
fof(f53,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X9] : (s(bool,'F_0') = s(bool,X9) <=> ~p(s(bool,X9)))),
  file('thm.transc.DIFF_COMPOSITE|split|9','thm.bool.EQ_CLAUSES|split|2')).
fof(f172,plain,(
  ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.transc.cos_0'),s(fun('type.realax.real','type.realax.real'),sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.sin_1'(s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK2),s('type.realax.real',sK0))))))),s('type.realax.real',sK1))),s('type.realax.real',sK0)))) | $spl1),
  inference(cnf_transformation,[],[f172_D])).
fof(f172_D,plain,(
  ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.transc.cos_0'),s(fun('type.realax.real','type.realax.real'),sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.sin_1'(s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK2),s('type.realax.real',sK0))))))),s('type.realax.real',sK1))),s('type.realax.real',sK0)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f117,plain,(
  ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.transc.cos_0'),s(fun('type.realax.real','type.realax.real'),sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.sin_1'(s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK2),s('type.realax.real',sK0))))))),s('type.realax.real',sK1))),s('type.realax.real',sK0))))),
  inference(cnf_transformation,[],[f100])).
fof(f100,plain,(
  p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),sK2),s('type.realax.real',sK1),s('type.realax.real',sK0)))) & ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.transc.cos_0'),s(fun('type.realax.real','type.realax.real'),sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.sin_1'(s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK2),s('type.realax.real',sK0))))))),s('type.realax.real',sK1))),s('type.realax.real',sK0))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f85])).
fof(f85,plain,(
  ? [X0,X1,X2] : (p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X2),s('type.realax.real',X1),s('type.realax.real',X0)))) & ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.transc.cos_0'),s(fun('type.realax.real','type.realax.real'),X2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.sin_1'(s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X2),s('type.realax.real',X0))))))),s('type.realax.real',X1))),s('type.realax.real',X0)))))),
  inference(ennf_transformation,[],[f41])).
fof(f41,plain,(
  ~! [X0,X1,X2] : (p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X2),s('type.realax.real',X1),s('type.realax.real',X0)))) => p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.transc.cos_0'),s(fun('type.realax.real','type.realax.real'),X2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.sin_1'(s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X2),s('type.realax.real',X0))))))),s('type.realax.real',X1))),s('type.realax.real',X0)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : (p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X8),s('type.realax.real',X7),s('type.realax.real',X6)))) => p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.transc.cos_0'),s(fun('type.realax.real','type.realax.real'),X8))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.sin_1'(s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X8),s('type.realax.real',X6))))))),s('type.realax.real',X7))),s('type.realax.real',X6)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : (p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X8),s('type.realax.real',X7),s('type.realax.real',X6)))) => p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.transc.cos_0'),s(fun('type.realax.real','type.realax.real'),X8))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.sin_1'(s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X8),s('type.realax.real',X6))))))),s('type.realax.real',X7))),s('type.realax.real',X6)))))),
  file('thm.transc.DIFF_COMPOSITE|split|9','thm.transc.DIFF_COMPOSITE|split|9')).
fof(f296,plain,(
  s(bool,'F_0') = s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.transc.cos_0'),s(fun('type.realax.real','type.realax.real'),sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.sin_1'(s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK2),s('type.realax.real',sK0))))))),s('type.realax.real',sK1))),s('type.realax.real',sK0))) | $spl94),
  inference(cnf_transformation,[],[f296_D])).
fof(f296_D,plain,(
  s(bool,'F_0') = s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.transc.cos_0'),s(fun('type.realax.real','type.realax.real'),sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.sin_1'(s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK2),s('type.realax.real',sK0))))))),s('type.realax.real',sK1))),s('type.realax.real',sK0))) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f293,plain,(
  s(bool,'F_0') = s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.transc.cos_0'),s(fun('type.realax.real','type.realax.real'),sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.sin_1'(s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK2),s('type.realax.real',sK0))))))),s('type.realax.real',sK1))),s('type.realax.real',sK0))) | ($spl1 | $spl30)),
  inference(resolution,[],[f172,f190])).
fof(f188,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl26),
  inference(cnf_transformation,[],[f188_D])).
fof(f188_D,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f131,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X9] : (s(bool,'T_0') = s(bool,X9) <=> p(s(bool,X9)))),
  file('thm.transc.DIFF_COMPOSITE|split|9','thm.bool.EQ_CLAUSES|split|0')).
fof(f174,plain,(
  p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),sK2),s('type.realax.real',sK1),s('type.realax.real',sK0)))) | $spl2),
  inference(cnf_transformation,[],[f174_D])).
fof(f174_D,plain,(
  p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),sK2),s('type.realax.real',sK1),s('type.realax.real',sK0)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f116,plain,(
  p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),sK2),s('type.realax.real',sK1),s('type.realax.real',sK0))))),
  inference(cnf_transformation,[],[f100])).
fof(f231,plain,(
  s(bool,'T_0') = s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),sK2),s('type.realax.real',sK1),s('type.realax.real',sK0))) | $spl82),
  inference(cnf_transformation,[],[f231_D])).
fof(f231_D,plain,(
  s(bool,'T_0') = s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),sK2),s('type.realax.real',sK1),s('type.realax.real',sK0))) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f229,plain,(
  s(bool,'T_0') = s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),sK2),s('type.realax.real',sK1),s('type.realax.real',sK0))) | ($spl2 | $spl26)),
  inference(resolution,[],[f188,f174])).
fof(f212,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),X0),s(fun('type.realax.real','type.realax.real'),X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3))),s('type.realax.real',X4)))) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X3),s('type.realax.real',X4)))) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',X2),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X4))))))) ) | $spl66),
  inference(cnf_transformation,[],[f212_D])).
fof(f212_D,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),X0),s(fun('type.realax.real','type.realax.real'),X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3))),s('type.realax.real',X4)))) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X3),s('type.realax.real',X4)))) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',X2),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X4))))))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f169,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),X0),s(fun('type.realax.real','type.realax.real'),X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3))),s('type.realax.real',X4)))) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X3),s('type.realax.real',X4)))) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',X2),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X4))))))) )),
  inference(cnf_transformation,[],[f99])).
fof(f99,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',X2),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X4)))))) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X3),s('type.realax.real',X4)))) | p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),X0),s(fun('type.realax.real','type.realax.real'),X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3))),s('type.realax.real',X4)))))),
  inference(flattening,[],[f98])).
fof(f98,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',X2),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X4)))))) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X3),s('type.realax.real',X4))))) | p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),X0),s(fun('type.realax.real','type.realax.real'),X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3))),s('type.realax.real',X4)))))),
  inference(ennf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0,X1,X2,X3,X4] : ((p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',X2),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X4)))))) & p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X3),s('type.realax.real',X4))))) => p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),X0),s(fun('type.realax.real','type.realax.real'),X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3))),s('type.realax.real',X4)))))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X10,X8,X11,X7,X6] : ((p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X10),s('type.realax.real',X11),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X8),s('type.realax.real',X6)))))) & p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X8),s('type.realax.real',X7),s('type.realax.real',X6))))) => p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),X10),s(fun('type.realax.real','type.realax.real'),X8))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X11),s('type.realax.real',X7))),s('type.realax.real',X6)))))),
  file('thm.transc.DIFF_COMPOSITE|split|9','thm.lim.DIFF_CHAIN')).
fof(f178,plain,(
  ~p(s(bool,'F_0')) | $spl7),
  inference(cnf_transformation,[],[f178_D])).
fof(f178_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f120,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f42])).
fof(f42,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.transc.DIFF_COMPOSITE|split|9',aHL_FALSITY)).
fof(f176,plain,(
  p(s(bool,'T_0')) | $spl4),
  inference(cnf_transformation,[],[f176_D])).
fof(f176_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f118,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.transc.DIFF_COMPOSITE|split|9',aHL_TRUTH)).
fof(f3400,plain,(
  $false | ($spl4 | $spl7 | $spl66 | $spl82 | $spl94 | $spl106)),
  inference(subsumption_resolution,[],[f3399,f176])).
fof(f3399,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl7 | $spl66 | $spl82 | $spl94 | $spl106)),
  inference(forward_demodulation,[],[f3398,f381])).
fof(f3398,plain,(
  ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'const.transc.cos_0'),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.sin_1'(s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK2),s('type.realax.real',sK0))))))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK2),s('type.realax.real',sK0)))))) | ($spl4 | $spl7 | $spl66 | $spl82 | $spl94)),
  inference(subsumption_resolution,[],[f3397,f176])).
fof(f3397,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'const.transc.cos_0'),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.sin_1'(s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK2),s('type.realax.real',sK0))))))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK2),s('type.realax.real',sK0)))))) | ($spl7 | $spl66 | $spl82 | $spl94)),
  inference(forward_demodulation,[],[f3396,f231])).
fof(f3396,plain,(
  ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),sK2),s('type.realax.real',sK1),s('type.realax.real',sK0)))) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'const.transc.cos_0'),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.sin_1'(s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK2),s('type.realax.real',sK0))))))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK2),s('type.realax.real',sK0)))))) | ($spl7 | $spl66 | $spl94)),
  inference(subsumption_resolution,[],[f3318,f178])).
fof(f3318,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),sK2),s('type.realax.real',sK1),s('type.realax.real',sK0)))) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'const.transc.cos_0'),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.sin_1'(s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK2),s('type.realax.real',sK0))))))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK2),s('type.realax.real',sK0)))))) | ($spl66 | $spl94)),
  inference(superposition,[],[f212,f296])).
% SZS output end Proof for thm.transc
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 68
% Passive clauses: 361
% Generated clauses: 3255
% Final active clauses: 56
% Final passive clauses: 253
% Input formulas: 39
% Initial clauses: 54
% 
% Duplicate literals: 154
% Trivial inequalities: 38
% Fw subsumption resolutions: 165
% Fw demodulations: 315
% Bw demodulations: 48
% Condensations: 153
% Global subsumptions: 4
% 
% Simple tautologies: 83
% Equational tautologies: 18
% Forward subsumptions: 1427
% 
% Binary resolution: 42
% Forward superposition: 1193
% Backward superposition: 618
% Self superposition: 68
% Equality factoring: 30
% Equality resolution: 9
% 
% Split clauses: 33
% Split components: 71
% SAT solver clauses: 590
% SAT solver unit clauses: 77
% SAT solver binary clauses: 30
% SAT solver learnt clauses: 9
% 
% Sat splits: 33
% Sat splitting refutations: 8
% 
% TWLsolver clauses: 573
% TWLsolver calls for satisfiability: 234
% 
% Memory used [KB]: 5500
% Time elapsed: 0.154 s
% ------------------------------
% ------------------------------
% Success in time 0.772 s
