Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.relation
% (91873)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 98
% Passive clauses: 2990
% Generated clauses: 18828
% Final active clauses: 84
% Final passive clauses: 2746
% Input formulas: 6
% Initial clauses: 16
% 
% Duplicate literals: 1130
% Trivial inequalities: 197
% Fw subsumption resolutions: 2273
% Fw demodulations: 1558
% Bw demodulations: 9
% 
% Simple tautologies: 335
% Equational tautologies: 98
% Forward subsumptions: 7541
% 
% Binary resolution: 84
% Forward superposition: 7575
% Backward superposition: 5143
% Self superposition: 383
% Equality factoring: 274
% Equality resolution: 11
% 
% Split clauses: 900
% Split components: 1810
% SAT solver clauses: 2186
% SAT solver unit clauses: 329
% SAT solver binary clauses: 597
% 
% Sat splits: 900
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 2178
% TWLsolver calls for satisfiability: 97
% 
% Memory used [KB]: 9978
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.relation
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.relation
% SZS output start Proof for thm.relation
fof(f1978,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f1977,f41,f62,f1633,f34,f53,f35,f55,f74,f76,f32,f49,f36,f56,f503,f510,f508,f384,f385,f33,f47,f64,f72,f598,f1730,f1732,f39,f58,f38,f59,f37,f60,f142,f147,f31,f50,f40,f57,f1837,f1848,f30,f51,f1745,f1598,f1747,f42,f61])).
fof(f61,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP0(X0,X2,X1) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK9(X2,X1,X0)))),s(X0,sK10(X2,X1,X0)))))) ) | $spl22),
  inference(cnf_transformation,[],[f61_D])).
fof(f61_D,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP0(X0,X2,X1) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK9(X2,X1,X0)))),s(X0,sK10(X2,X1,X0)))))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f42,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | sP0(X0,X2,X1) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK9(X2,X1,X0)))),s(X0,sK10(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f28])).
fof(f28,plain,(
  ! [X0,X1,X2] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK9(X2,X1,X0)))),s(X0,sK10(X2,X1,X0))))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK9(X2,X1,X0)))),s(X0,sK10(X2,X1,X0)))))) | sP0(X0,X2,X1) | ! [X3,X4] : (~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK9,sK10])],[f27])).
fof(f27,plain,(
  ! [X0,X1,X2] : (? [X5,X6] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X6)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X5))),s(X0,X6))))) | sP0(X0,X2,X1) | ! [X3,X4] : (~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))))))),
  inference(rectify,[],[f21])).
fof(f21,plain,(
  ! [X0,X1,X2] : (? [X6,X7] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))),s(X0,X7)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X6))),s(X0,X7))))) | sP0(X0,X2,X1) | ! [X8,X9] : (~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X8),s(X0,X9)))) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X8))),s(X0,X9)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X8),s(X0,X9)))))))),
  inference(definition_folding,[],[f18,f20])).
fof(f20,plain,(
  ! [X0,X2,X1] : (~sP0(X0,X2,X1) | ? [X3,X4,X5] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X4)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X4))),s(X0,X5)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X5)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X5))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f18,plain,(
  ! [X0,X1,X2] : (? [X6,X7] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))),s(X0,X7)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X6))),s(X0,X7))))) | ? [X3,X4,X5] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X4)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X4))),s(X0,X5)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X5)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X5))))) | ! [X8,X9] : (~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X8),s(X0,X9)))) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X8))),s(X0,X9)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X8),s(X0,X9)))))))),
  inference(flattening,[],[f17])).
fof(f17,plain,(
  ! [X0,X1,X2] : ((? [X6,X7] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))),s(X0,X7)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X6))),s(X0,X7))))) | ? [X3,X4,X5] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X4)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X4))),s(X0,X5)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X5))))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X5)))))) | ! [X8,X9] : (~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X8),s(X0,X9)))) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X8))),s(X0,X9)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X8),s(X0,X9)))))))),
  inference(ennf_transformation,[],[f12])).
fof(f12,plain,(
  ! [X0,X1,X2] : ((! [X6,X7] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))),s(X0,X7)))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X6))),s(X0,X7))))) & ! [X3,X4,X5] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X4)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X4))),s(X0,X5)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X5))))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X5)))))) => ! [X8,X9] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X8),s(X0,X9)))) => (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X8))),s(X0,X9)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X8),s(X0,X9)))))))),
  inference(flattening,[],[f11])).
fof(f11,plain,(
  ! [X0] : ! [X1,X2] : ((! [X6,X7] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))),s(X0,X7)))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X6))),s(X0,X7))))) & ! [X3,X4,X5] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X4)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X4))),s(X0,X5)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X5))))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X5)))))) => ! [X8,X9] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X8),s(X0,X9)))) => (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X8))),s(X0,X9)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X8),s(X0,X9)))))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6] : ! [X7,X8] : ((! [X9,X10] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X9))),s(X6,X10)))) => p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,X9))),s(X6,X10))))) & ! [X9,X10,X11] : ((p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X9))),s(X6,X10)))) & p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,X10))),s(X6,X11)))) & p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X10),s(X6,X11))))) => p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,X9))),s(X6,X11)))))) => ! [X12,X13] : (p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X12),s(X6,X13)))) => (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,X12))),s(X6,X13)))) & p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X12),s(X6,X13)))))))),
  file('thm.relation.TC_STRONG_INDUCT_LEFT|49|','thm.relation.relation_49_prove_ltb')).
fof(f1747,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1))))) | $spl353),
  inference(cnf_transformation,[],[f1747_D])).
fof(f1747_D,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1))))) <=> ~$spl353),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl353])])).
fof(f1598,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,X1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,X1))))) ) | $spl308),
  inference(cnf_transformation,[],[f1598_D])).
fof(f1598_D,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,X1)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,X1))))) ) <=> ~$spl308),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl308])])).
fof(f1745,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0),s(sK1,X1)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,X1)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))))) ) | ($spl22 | $spl347)),
  inference(resolution,[],[f1732,f61])).
fof(f51,plain,(
  ( ! [X8,X9] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X8))),s(sK1,X9)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X8))),s(sK1,X9))))) ) | $spl6),
  inference(cnf_transformation,[],[f51_D])).
fof(f51_D,plain,(
  ( ! [X8,X9] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X8))),s(sK1,X9)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X8))),s(sK1,X9))))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f30,plain,(
  ( ! [X8,X9] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X8))),s(sK1,X9)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X8))),s(sK1,X9))))) )),
  inference(cnf_transformation,[],[f23])).
fof(f23,plain,(
  ! [X8,X9] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X8))),s(sK1,X9)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X8))),s(sK1,X9))))) & ! [X5,X6,X7] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X5))),s(sK1,X6)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X6))),s(sK1,X7)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X6),s(sK1,X7)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X5))),s(sK1,X7))))) & (p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK4),s(sK1,sK5)))) & ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK4))),s(sK1,sK5)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1,sK2,sK3,sK4,sK5])],[f22])).
fof(f22,plain,(
  ? [X0,X1,X2] : (! [X8,X9] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X8))),s(X0,X9)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X8))),s(X0,X9))))) & ! [X5,X6,X7] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X6)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X6))),s(X0,X7)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X6),s(X0,X7)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X5))),s(X0,X7))))) & ? [X3,X4] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4))))))),
  inference(rectify,[],[f16])).
fof(f16,plain,(
  ? [X0,X1,X2] : (! [X6,X7] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))),s(X0,X7)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X6))),s(X0,X7))))) & ! [X3,X4,X5] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X4))),s(X0,X5)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X5)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X5))))) & ? [X8,X9] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X8),s(X0,X9)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X8))),s(X0,X9))))))),
  inference(flattening,[],[f15])).
fof(f15,plain,(
  ? [X0,X1,X2] : ((! [X6,X7] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))),s(X0,X7)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X6))),s(X0,X7))))) & ! [X3,X4,X5] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X4))),s(X0,X5)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X5))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X5)))))) & ? [X8,X9] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X8),s(X0,X9)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X8))),s(X0,X9))))))),
  inference(ennf_transformation,[],[f9])).
fof(f9,plain,(
  ~! [X0,X1,X2] : ((! [X6,X7] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))),s(X0,X7)))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X6))),s(X0,X7))))) & ! [X3,X4,X5] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X4)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X4))),s(X0,X5)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X5))))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X5)))))) => ! [X8,X9] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X8),s(X0,X9)))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X8))),s(X0,X9))))))),
  inference(flattening,[],[f8])).
fof(f8,plain,(
  ~! [X0] : ! [X1,X2] : ((! [X6,X7] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))),s(X0,X7)))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X6))),s(X0,X7))))) & ! [X3,X4,X5] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X4)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X4))),s(X0,X5)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X5))))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X5)))))) => ! [X8,X9] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X8),s(X0,X9)))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X8))),s(X0,X9))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : ((! [X9,X10] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X9))),s(X6,X10)))) => p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,X9))),s(X6,X10))))) & ! [X9,X10,X11] : ((p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X9))),s(X6,X10)))) & p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,X10))),s(X6,X11)))) & p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X10),s(X6,X11))))) => p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,X9))),s(X6,X11)))))) => ! [X12,X13] : (p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X12),s(X6,X13)))) => p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,X12))),s(X6,X13))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : ((! [X9,X10] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X9))),s(X6,X10)))) => p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,X9))),s(X6,X10))))) & ! [X9,X10,X11] : ((p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X9))),s(X6,X10)))) & p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,X10))),s(X6,X11)))) & p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X10),s(X6,X11))))) => p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,X9))),s(X6,X11)))))) => ! [X12,X13] : (p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X12),s(X6,X13)))) => p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,X12))),s(X6,X13))))))),
  file('thm.relation.TC_STRONG_INDUCT_LEFT|49|','thm.relation.TC_STRONG_INDUCT_LEFT|49|')).
fof(f1848,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1))))) | $spl361),
  inference(cnf_transformation,[],[f1848_D])).
fof(f1848_D,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1))))) <=> ~$spl361),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl361])])).
fof(f1837,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1))))) | ($spl6 | $spl353)),
  inference(resolution,[],[f1747,f51])).
fof(f57,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK6(X2,X1,X0)))),s(X0,sK8(X2,X1,X0))))) | ~sP0(X0,X1,X2)) ) | $spl14),
  inference(cnf_transformation,[],[f57_D])).
fof(f57_D,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK6(X2,X1,X0)))),s(X0,sK8(X2,X1,X0))))) | ~sP0(X0,X1,X2)) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f40,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK6(X2,X1,X0)))),s(X0,sK8(X2,X1,X0))))) | ~sP0(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f26])).
fof(f26,plain,(
  ! [X0,X1,X2] : (~sP0(X0,X1,X2) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK6(X2,X1,X0)))),s(X0,sK7(X2,X1,X0))))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK7(X2,X1,X0)))),s(X0,sK8(X2,X1,X0))))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X2),s(X0,sK7(X2,X1,X0)),s(X0,sK8(X2,X1,X0))))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK6(X2,X1,X0)))),s(X0,sK8(X2,X1,X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK6,sK7,sK8])],[f25])).
fof(f25,plain,(
  ! [X0,X1,X2] : (~sP0(X0,X1,X2) | ? [X3,X4,X5] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X5)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X2),s(X0,X4),s(X0,X5)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X5))))))),
  inference(rectify,[],[f24])).
fof(f24,plain,(
  ! [X0,X2,X1] : (~sP0(X0,X2,X1) | ? [X3,X4,X5] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X4)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X4))),s(X0,X5)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X5)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X5))))))),
  inference(nnf_transformation,[],[f20])).
fof(f50,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X5))),s(sK1,X7)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X6),s(sK1,X7)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X6))),s(sK1,X7)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X5))),s(sK1,X6))))) ) | $spl4),
  inference(cnf_transformation,[],[f50_D])).
fof(f50_D,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X5))),s(sK1,X7)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X6),s(sK1,X7)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X6))),s(sK1,X7)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X5))),s(sK1,X6))))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f31,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X5))),s(sK1,X7)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X6),s(sK1,X7)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X6))),s(sK1,X7)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X5))),s(sK1,X6))))) )),
  inference(cnf_transformation,[],[f23])).
fof(f147,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X1))),s(sK1,sK8(X0,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(X0,sK3,sK1)))),s(sK1,X1)))) | ~sP0(sK1,sK3,X0) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X1),s(sK1,sK8(X0,sK3,sK1)))))) ) | $spl58),
  inference(cnf_transformation,[],[f147_D])).
fof(f147_D,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X1))),s(sK1,sK8(X0,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(X0,sK3,sK1)))),s(sK1,X1)))) | ~sP0(sK1,sK3,X0) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X1),s(sK1,sK8(X0,sK3,sK1)))))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f142,plain,(
  ( ! [X0,X1] : (~sP0(sK1,sK3,X0) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X1),s(sK1,sK8(X0,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X1))),s(sK1,sK8(X0,sK3,sK1))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6(X0,sK3,sK1)))),s(sK1,X1))))) ) | ($spl4 | $spl14)),
  inference(resolution,[],[f57,f50])).
fof(f60,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK6(X2,X1,X0)))),s(X0,sK7(X2,X1,X0))))) | ~sP0(X0,X1,X2)) ) | $spl20),
  inference(cnf_transformation,[],[f60_D])).
fof(f60_D,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK6(X2,X1,X0)))),s(X0,sK7(X2,X1,X0))))) | ~sP0(X0,X1,X2)) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f37,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,sK6(X2,X1,X0)))),s(X0,sK7(X2,X1,X0))))) | ~sP0(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f26])).
fof(f59,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK7(X2,X1,X0)))),s(X0,sK8(X2,X1,X0))))) | ~sP0(X0,X1,X2)) ) | $spl18),
  inference(cnf_transformation,[],[f59_D])).
fof(f59_D,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK7(X2,X1,X0)))),s(X0,sK8(X2,X1,X0))))) | ~sP0(X0,X1,X2)) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f38,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK7(X2,X1,X0)))),s(X0,sK8(X2,X1,X0))))) | ~sP0(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f26])).
fof(f58,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X2),s(X0,sK7(X2,X1,X0)),s(X0,sK8(X2,X1,X0))))) | ~sP0(X0,X1,X2)) ) | $spl16),
  inference(cnf_transformation,[],[f58_D])).
fof(f58_D,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X2),s(X0,sK7(X2,X1,X0)),s(X0,sK8(X2,X1,X0))))) | ~sP0(X0,X1,X2)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f39,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X2),s(X0,sK7(X2,X1,X0)),s(X0,sK8(X2,X1,X0))))) | ~sP0(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f26])).
fof(f1732,plain,(
  ~sP0(sK1,sK3,sK2) | $spl347),
  inference(cnf_transformation,[],[f1732_D])).
fof(f1732_D,plain,(
  ~sP0(sK1,sK3,sK2) <=> ~$spl347),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl347])])).
fof(f1730,plain,(
  ~sP0(sK1,sK3,sK2) | ($spl16 | $spl18 | $spl20 | $spl58)),
  inference(subsumption_resolution,[],[f1729,f58])).
fof(f1729,plain,(
  ~sP0(sK1,sK3,sK2) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1))))) | ($spl18 | $spl20 | $spl58)),
  inference(subsumption_resolution,[],[f1724,f59])).
fof(f1724,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1))))) | ~sP0(sK1,sK3,sK2) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1))))) | ($spl20 | $spl58)),
  inference(duplicate_literal_removal,[],[f1707])).
fof(f1707,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK7(sK2,sK3,sK1)))),s(sK1,sK8(sK2,sK3,sK1))))) | ~sP0(sK1,sK3,sK2) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7(sK2,sK3,sK1)),s(sK1,sK8(sK2,sK3,sK1))))) | ~sP0(sK1,sK3,sK2) | ($spl20 | $spl58)),
  inference(resolution,[],[f147,f60])).
fof(f598,plain,(
  $false | ($spl8 | $spl11 | $spl118)),
  inference(subsumption_resolution,[],[f574,f53])).
fof(f574,plain,(
  ~p(s(bool,'T_0')) | ($spl11 | $spl118)),
  inference(backward_demodulation,[],[f508,f55])).
fof(f72,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl28),
  inference(cnf_transformation,[],[f72_D])).
fof(f72_D,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f64,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl12),
  inference(superposition,[],[f56,f56])).
fof(f47,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK4))),s(sK1,sK5)))) | $spl1),
  inference(cnf_transformation,[],[f47_D])).
fof(f47_D,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK4))),s(sK1,sK5)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f33,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK4))),s(sK1,sK5))))),
  inference(cnf_transformation,[],[f23])).
fof(f385,plain,(
  ( ! [X9] : (s(bool,X9) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK4))),s(sK1,sK5))) | s(bool,'T_0') = s(bool,X9)) ) | $spl78),
  inference(cnf_transformation,[],[f385_D])).
fof(f385_D,plain,(
  ( ! [X9] : (s(bool,X9) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK4))),s(sK1,sK5))) | s(bool,'T_0') = s(bool,X9)) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f384,plain,(
  ( ! [X9] : (s(bool,X9) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK4))),s(sK1,sK5))) | s(bool,'T_0') = s(bool,X9)) ) | ($spl1 | $spl8 | $spl28)),
  inference(subsumption_resolution,[],[f242,f53])).
fof(f242,plain,(
  ( ! [X9] : (~p(s(bool,'T_0')) | s(bool,X9) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK4))),s(sK1,sK5))) | s(bool,'T_0') = s(bool,X9)) ) | ($spl1 | $spl28)),
  inference(superposition,[],[f47,f72])).
fof(f508,plain,(
  ( ! [X9] : (s(bool,'T_0') = s(bool,X9)) ) | $spl118),
  inference(cnf_transformation,[],[f508_D])).
fof(f508_D,plain,(
  ( ! [X9] : (s(bool,'T_0') = s(bool,X9)) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f510,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK4))),s(sK1,sK5))) | $spl120),
  inference(cnf_transformation,[],[f510_D])).
fof(f510_D,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK4))),s(sK1,sK5))) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f503,plain,(
  ( ! [X9] : (s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK4))),s(sK1,sK5))) | s(bool,'T_0') = s(bool,X9)) ) | ($spl12 | $spl78)),
  inference(duplicate_literal_removal,[],[f446])).
fof(f446,plain,(
  ( ! [X9] : (s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK4))),s(sK1,sK5))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) ) | ($spl12 | $spl78)),
  inference(superposition,[],[f385,f56])).
fof(f56,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl12),
  inference(cnf_transformation,[],[f56_D])).
fof(f56_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f36,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.relation.TC_STRONG_INDUCT_LEFT|49|',aHL_BOOL_CASES)).
fof(f49,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK4),s(sK1,sK5)))) | $spl2),
  inference(cnf_transformation,[],[f49_D])).
fof(f49_D,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK4),s(sK1,sK5)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f32,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK4),s(sK1,sK5))))),
  inference(cnf_transformation,[],[f23])).
fof(f76,plain,(
  s(bool,'T_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK4),s(sK1,sK5))) | $spl32),
  inference(cnf_transformation,[],[f76_D])).
fof(f76_D,plain,(
  s(bool,'T_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK4),s(sK1,sK5))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f74,plain,(
  s(bool,'T_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK4),s(sK1,sK5))) | ($spl2 | $spl11 | $spl12)),
  inference(subsumption_resolution,[],[f69,f55])).
fof(f69,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK4),s(sK1,sK5))) | ($spl2 | $spl12)),
  inference(superposition,[],[f49,f56])).
fof(f55,plain,(
  ~p(s(bool,'F_0')) | $spl11),
  inference(cnf_transformation,[],[f55_D])).
fof(f55_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f35,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f10])).
fof(f10,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.relation.TC_STRONG_INDUCT_LEFT|49|',aHL_FALSITY)).
fof(f53,plain,(
  p(s(bool,'T_0')) | $spl8),
  inference(cnf_transformation,[],[f53_D])).
fof(f53_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f34,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.relation.TC_STRONG_INDUCT_LEFT|49|',aHL_TRUTH)).
fof(f1633,plain,(
  $false | ($spl8 | $spl11 | $spl32 | $spl120 | $spl308)),
  inference(subsumption_resolution,[],[f1632,f53])).
fof(f1632,plain,(
  ~p(s(bool,'T_0')) | ($spl11 | $spl32 | $spl120 | $spl308)),
  inference(forward_demodulation,[],[f1631,f76])).
fof(f1631,plain,(
  ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK4),s(sK1,sK5)))) | ($spl11 | $spl120 | $spl308)),
  inference(subsumption_resolution,[],[f1617,f55])).
fof(f1617,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK4),s(sK1,sK5)))) | ($spl120 | $spl308)),
  inference(superposition,[],[f1598,f510])).
fof(f62,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP0(X0,X2,X1) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK9(X2,X1,X0)))),s(X0,sK10(X2,X1,X0)))))) ) | $spl24),
  inference(cnf_transformation,[],[f62_D])).
fof(f62_D,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP0(X0,X2,X1) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK9(X2,X1,X0)))),s(X0,sK10(X2,X1,X0)))))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f41,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X3),s(X0,X4)))) | sP0(X0,X2,X1) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK9(X2,X1,X0)))),s(X0,sK10(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f28])).
fof(f1977,plain,(
  ( ! [X26,X25] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X25),s(sK1,X26)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X25))),s(sK1,X26))))) ) | ($spl24 | $spl347 | $spl361)),
  inference(subsumption_resolution,[],[f1974,f1848])).
fof(f1974,plain,(
  ( ! [X26,X25] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X25),s(sK1,X26)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X25))),s(sK1,X26)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK9(sK3,sK2,sK1)))),s(sK1,sK10(sK3,sK2,sK1)))))) ) | ($spl24 | $spl347)),
  inference(resolution,[],[f62,f1732])).
% SZS output end Proof for thm.relation
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 77
% Passive clauses: 214
% Generated clauses: 2034
% Final active clauses: 44
% Final passive clauses: 87
% Input formulas: 6
% Initial clauses: 16
% 
% Duplicate literals: 105
% Trivial inequalities: 8
% Fw subsumption resolutions: 183
% Fw demodulations: 158
% Bw demodulations: 65
% Condensations: 123
% 
% Simple tautologies: 30
% Equational tautologies: 17
% Forward subsumptions: 805
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 71
% Forward superposition: 674
% Backward superposition: 263
% Self superposition: 30
% Equality factoring: 32
% 
% Split clauses: 37
% Split components: 74
% SAT solver clauses: 378
% SAT solver unit clauses: 18
% SAT solver binary clauses: 13
% SAT solver learnt clauses: 1
% 
% Sat splits: 37
% Sat splitting refutations: 12
% 
% TWLsolver clauses: 377
% TWLsolver calls for satisfiability: 158
% 
% Memory used [KB]: 2430
% Time elapsed: 0.079 s
% ------------------------------
% ------------------------------
% Success in time 0.693 s
