Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.relation
% (68314)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 55
% Passive clauses: 1340
% Generated clauses: 11401
% Final active clauses: 51
% Final passive clauses: 1135
% Input formulas: 24
% Initial clauses: 37
% 
% Duplicate literals: 915
% Trivial inequalities: 132
% Fw subsumption resolutions: 1557
% Fw demodulations: 83
% 
% Simple tautologies: 269
% Equational tautologies: 118
% Forward subsumptions: 4506
% 
% Binary resolution: 25
% Forward superposition: 4293
% Backward superposition: 3501
% Self superposition: 425
% Equality factoring: 362
% Equality resolution: 11
% 
% Split clauses: 733
% Split components: 1513
% SAT solver clauses: 1884
% SAT solver unit clauses: 243
% SAT solver binary clauses: 588
% 
% Sat splits: 733
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1875
% TWLsolver calls for satisfiability: 47
% 
% Memory used [KB]: 5756
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.relation
% (68327)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 34
% Passive clauses: 97
% Generated clauses: 5265
% Final active clauses: 31
% Final passive clauses: 59
% Input formulas: 24
% Initial clauses: 37
% 
% Duplicate literals: 1147
% Trivial inequalities: 61
% Fw subsumption resolutions: 137
% Fw demodulations: 3
% Bw demodulations: 3
% Condensations: 1600
% Global subsumptions: 4
% 
% Simple tautologies: 135
% Equational tautologies: 78
% Forward subsumptions: 1305
% 
% Binary resolution: 16
% Forward superposition: 989
% Backward superposition: 928
% Self superposition: 246
% Equality factoring: 222
% Equality resolution: 5
% 
% Split clauses: 27
% Split components: 54
% SAT solver clauses: 305
% SAT solver unit clauses: 48
% SAT solver binary clauses: 70
% SAT solver learnt clauses: 8
% 
% Sat splits: 27
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 289
% TWLsolver calls for satisfiability: 69
% 
% Memory used [KB]: 3454
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.relation
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.relation
% SZS output start Proof for thm.relation
fof(f38058,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f38024])).
fof(f38024,plain,(
  s(fun(sK0,fun(sK0,bool)),'const.relation.inv_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.RTC_1'(s(fun(sK0,fun(sK0,bool)),sK1))))) != s(fun(sK0,fun(sK0,bool)),'const.relation.inv_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.RTC_1'(s(fun(sK0,fun(sK0,bool)),sK1)))))),
  inference(superposition,[],[f79,f27140])).
fof(f27140,plain,(
  ( ! [X2,X3] : (s(fun(X2,fun(X2,bool)),'const.relation.inv_1'(s(fun(X2,fun(X2,bool)),'const.relation.RTC_1'(s(fun(X2,fun(X2,bool)),X3))))) = s(fun(X2,fun(X2,bool)),'const.relation.RTC_1'(s(fun(X2,fun(X2,bool)),'const.relation.inv_1'(s(fun(X2,fun(X2,bool)),X3)))))) )),
  inference(forward_demodulation,[],[f27105,f101])).
fof(f101,plain,(
  ( ! [X0,X1] : (s(fun(X0,fun(X0,bool)),'const.relation.RTC_1'(s(fun(X0,fun(X0,bool)),X1))) = s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),X1)))))) )),
  inference(cnf_transformation,[],[f39])).
fof(f39,plain,(
  ! [X0,X1] : s(fun(X0,fun(X0,bool)),'const.relation.RTC_1'(s(fun(X0,fun(X0,bool)),X1))) = s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),X1)))))),
  inference(flattening,[],[f38])).
fof(f38,plain,(
  ! [X0] : ! [X1] : s(fun(X0,fun(X0,bool)),'const.relation.RTC_1'(s(fun(X0,fun(X0,bool)),X1))) = s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),X1)))))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X6] : ! [X7] : s(fun(X6,fun(X6,bool)),'const.relation.RTC_1'(s(fun(X6,fun(X6,bool)),X7))) = s(fun(X6,fun(X6,bool)),'const.relation.TC_1'(s(fun(X6,fun(X6,bool)),'const.relation.RC_1'(s(fun(X6,fun(X6,bool)),X7)))))),
  file('thm.relation.inv_MOVES_OUT|split|4','thm.relation.TC_RC_EQNS|split|1')).
fof(f27105,plain,(
  ( ! [X2,X3] : (s(fun(X2,fun(X2,bool)),'const.relation.RTC_1'(s(fun(X2,fun(X2,bool)),'const.relation.inv_1'(s(fun(X2,fun(X2,bool)),X3))))) = s(fun(X2,fun(X2,bool)),'const.relation.inv_1'(s(fun(X2,fun(X2,bool)),'const.relation.TC_1'(s(fun(X2,fun(X2,bool)),'const.relation.RC_1'(s(fun(X2,fun(X2,bool)),X3)))))))) )),
  inference(superposition,[],[f4008,f106])).
fof(f106,plain,(
  ( ! [X0,X1] : (s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),X1))))) = s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),X1)))))) )),
  inference(cnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0,X1] : s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),X1))))) = s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),X1)))))),
  inference(flattening,[],[f48])).
fof(f48,plain,(
  ! [X0] : ! [X1] : s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),X1))))) = s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),X1)))))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X6] : ! [X7] : s(fun(X6,fun(X6,bool)),'const.relation.inv_1'(s(fun(X6,fun(X6,bool)),'const.relation.TC_1'(s(fun(X6,fun(X6,bool)),X7))))) = s(fun(X6,fun(X6,bool)),'const.relation.TC_1'(s(fun(X6,fun(X6,bool)),'const.relation.inv_1'(s(fun(X6,fun(X6,bool)),X7)))))),
  file('thm.relation.inv_MOVES_OUT|split|4','thm.relation.inv_TC')).
fof(f4008,plain,(
  ( ! [X0,X1] : (s(fun(X0,fun(X0,bool)),'const.relation.RTC_1'(s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),X1))))) = s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),X1)))))))) )),
  inference(superposition,[],[f101,f105])).
fof(f105,plain,(
  ( ! [X0,X1] : (s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),X1))))) = s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),X1)))))) )),
  inference(cnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0,X1] : s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),X1))))) = s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),X1)))))),
  inference(flattening,[],[f46])).
fof(f46,plain,(
  ! [X0] : ! [X1] : s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),X1))))) = s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),X1)))))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X6] : ! [X7] : s(fun(X6,fun(X6,bool)),'const.relation.inv_1'(s(fun(X6,fun(X6,bool)),'const.relation.RC_1'(s(fun(X6,fun(X6,bool)),X7))))) = s(fun(X6,fun(X6,bool)),'const.relation.RC_1'(s(fun(X6,fun(X6,bool)),'const.relation.inv_1'(s(fun(X6,fun(X6,bool)),X7)))))),
  file('thm.relation.inv_MOVES_OUT|split|4','thm.relation.inv_RC')).
fof(f79,plain,(
  s(fun(sK0,fun(sK0,bool)),'const.relation.inv_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.RTC_1'(s(fun(sK0,fun(sK0,bool)),sK1))))) != s(fun(sK0,fun(sK0,bool)),'const.relation.RTC_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.inv_1'(s(fun(sK0,fun(sK0,bool)),sK1)))))),
  inference(cnf_transformation,[],[f63])).
fof(f63,plain,(
  s(fun(sK0,fun(sK0,bool)),'const.relation.inv_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.RTC_1'(s(fun(sK0,fun(sK0,bool)),sK1))))) != s(fun(sK0,fun(sK0,bool)),'const.relation.RTC_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.inv_1'(s(fun(sK0,fun(sK0,bool)),sK1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f61])).
fof(f61,plain,(
  ? [X0,X1] : s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),'const.relation.RTC_1'(s(fun(X0,fun(X0,bool)),X1))))) != s(fun(X0,fun(X0,bool)),'const.relation.RTC_1'(s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),X1)))))),
  inference(ennf_transformation,[],[f27])).
fof(f27,plain,(
  ~! [X0,X1] : s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),'const.relation.RTC_1'(s(fun(X0,fun(X0,bool)),X1))))) = s(fun(X0,fun(X0,bool)),'const.relation.RTC_1'(s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),X1)))))),
  inference(flattening,[],[f26])).
fof(f26,plain,(
  ~! [X0] : ! [X1] : s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),'const.relation.RTC_1'(s(fun(X0,fun(X0,bool)),X1))))) = s(fun(X0,fun(X0,bool)),'const.relation.RTC_1'(s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),X1)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7] : s(fun(X6,fun(X6,bool)),'const.relation.inv_1'(s(fun(X6,fun(X6,bool)),'const.relation.RTC_1'(s(fun(X6,fun(X6,bool)),X7))))) = s(fun(X6,fun(X6,bool)),'const.relation.RTC_1'(s(fun(X6,fun(X6,bool)),'const.relation.inv_1'(s(fun(X6,fun(X6,bool)),X7)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7] : s(fun(X6,fun(X6,bool)),'const.relation.inv_1'(s(fun(X6,fun(X6,bool)),'const.relation.RTC_1'(s(fun(X6,fun(X6,bool)),X7))))) = s(fun(X6,fun(X6,bool)),'const.relation.RTC_1'(s(fun(X6,fun(X6,bool)),'const.relation.inv_1'(s(fun(X6,fun(X6,bool)),X7)))))),
  file('thm.relation.inv_MOVES_OUT|split|4','thm.relation.inv_MOVES_OUT|split|4')).
% SZS output end Proof for thm.relation
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 128
% Passive clauses: 1255
% Generated clauses: 37927
% Final active clauses: 95
% Final passive clauses: 497
% Input formulas: 24
% Initial clauses: 37
% 
% Duplicate literals: 4143
% Trivial inequalities: 314
% Fw subsumption resolutions: 2153
% Fw demodulations: 1065
% Bw demodulations: 476
% Condensations: 7947
% Global subsumptions: 1
% 
% Simple tautologies: 518
% Equational tautologies: 981
% Forward subsumptions: 19665
% Fw demodulations to eq. taut.: 10
% Bw demodulations to eq. taut.: 36
% 
% Unit resulting resolution: 2149
% Forward superposition: 11629
% Backward superposition: 6853
% Self superposition: 442
% Equality factoring: 342
% Equality resolution: 18
% 
% Split clauses: 122
% Split components: 245
% SAT solver clauses: 1514
% SAT solver unit clauses: 95
% SAT solver binary clauses: 85
% 
% Sat splits: 122
% Sat splitting refutations: 12
% 
% Memory used [KB]: 16502
% Time elapsed: 5.274 s
% ------------------------------
% ------------------------------
% Success in time 6.588 s
