Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.extreal
% (94042)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 258
% Passive clauses: 3334
% Generated clauses: 29519
% Final active clauses: 210
% Final passive clauses: 2913
% Input formulas: 69
% Initial clauses: 114
% 
% Duplicate literals: 742
% Trivial inequalities: 1573
% Fw subsumption resolutions: 1281
% Fw demodulations: 5591
% Bw demodulations: 45
% 
% Simple tautologies: 295
% Equational tautologies: 33
% Forward subsumptions: 16355
% 
% Binary resolution: 125
% Factoring: 1
% Forward superposition: 8641
% Backward superposition: 11105
% Self superposition: 111
% Equality factoring: 58
% Equality resolution: 52
% 
% Split clauses: 64
% Split components: 128
% SAT solver clauses: 169
% SAT solver unit clauses: 15
% SAT solver binary clauses: 85
% SAT solver learnt clauses: 1
% SAT solver learnt literals: 1
% 
% Sat splits: 64
% Sat splitting refutations: 7
% 
% TWLsolver clauses: 168
% TWLsolver calls for satisfiability: 34
% 
% Memory used [KB]: 17398
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 144 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.extreal
% (94069)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 211
% Passive clauses: 1437
% Generated clauses: 17041
% Final active clauses: 182
% Final passive clauses: 1080
% Input formulas: 69
% Initial clauses: 114
% 
% Duplicate literals: 327
% Trivial inequalities: 440
% Fw subsumption resolutions: 816
% Fw demodulations: 2979
% Bw demodulations: 173
% Condensations: 323
% Global subsumptions: 11
% 
% Simple tautologies: 216
% Equational tautologies: 18
% Forward subsumptions: 8658
% 
% Binary resolution: 138
% Factoring: 1
% Forward superposition: 4866
% Backward superposition: 5279
% Self superposition: 79
% Equality factoring: 28
% Equality resolution: 21
% 
% Split clauses: 45
% Split components: 90
% SAT solver clauses: 1787
% SAT solver unit clauses: 155
% SAT solver binary clauses: 58
% SAT solver learnt clauses: 23
% SAT solver learnt literals: 1
% 
% Sat splits: 45
% Sat splitting refutations: 8
% 
% TWLsolver clauses: 1742
% TWLsolver calls for satisfiability: 741
% 
% Memory used [KB]: 18805
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.extreal
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.extreal
% SZS output start Proof for thm.extreal
fof(f16089,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f216,f336,f334,f215,f335,f333,f343,f813,f819,f1960,f1964,f1962,f2576,f2585,f2584,f2835,f2839,f2837,f2915,f2918,f2916,f2919,f2923,f2925,f2927,f2929,f2931,f2933,f2935,f2937,f2939,f2941,f2943,f2945,f2947,f2949,f2951,f2953,f2957,f2959,f2961,f2963,f2965,f2967,f2969,f2971,f2973,f2975,f2977,f2979,f2981,f2983,f2985,f2987,f2989,f2991,f2993,f2995,f2998,f3002,f3004,f3051,f3052,f3053,f3054,f3055,f3056,f3057,f3058,f3059,f3060,f3061,f3062,f3063,f3064,f3065,f3066,f3067,f3068,f3069,f3070,f3071,f3072,f3073,f3216,f3968,f3972,f3970,f4303,f4309,f4424,f4425,f4427,f4430,f4431,f4432,f4434,f4435,f4436,f4507,f4527,f4529,f4531,f4533,f4534,f4536,f4537,f4546,f1922,f4590,f1923,f4591,f1924,f4592,f1925,f4593,f1926,f4594,f1927,f4595,f1928,f4596,f1929,f4597,f1930,f821,f4598,f7143,f7147,f7145,f7158,f3220,f3218,f7204,f8155,f8159,f8157,f8362,f8377,f8530,f8534,f8532,f8600,f8615,f8789,f8793,f8791,f8914,f8929,f9103,f9107,f9105,f9173,f9188,f9362,f9366,f9364,f9710,f9725,f16075])).
fof(f16075,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f16044,f426])).
fof(f426,plain,(
  s(bool,'T_0') != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))),
  inference(subsumption_resolution,[],[f423,f337])).
fof(f337,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f222,f217])).
fof(f217,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.extreal.le_infty|split|2',aHL_TRUTH)).
fof(f222,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f175])).
fof(f175,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f72])).
fof(f72,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f55])).
fof(f55,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.extreal.le_infty|split|2','thm.bool.NOT_CLAUSES|split|1')).
fof(f423,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))),
  inference(superposition,[],[f247,f219])).
fof(f219,plain,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))),
  inference(cnf_transformation,[],[f18])).
fof(f18,axiom,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))),
  file('thm.extreal.le_infty|split|2','thm.extreal.extreal_le_def|split|4')).
fof(f247,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f178])).
fof(f178,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f95])).
fof(f95,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X24] : (s(bool,'T_0') = s(bool,X24) <=> p(s(bool,X24)))),
  file('thm.extreal.le_infty|split|2','thm.bool.EQ_CLAUSES|split|1')).
fof(f16044,plain,(
  s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | ($spl0 | $spl3)),
  inference(backward_demodulation,[],[f16043,f364])).
fof(f364,plain,(
  s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal','const.extreal.NegInf_0'))) | $spl0),
  inference(unit_resulting_resolution,[],[f333,f248])).
fof(f248,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f178])).
fof(f16043,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK0) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f16042,f336])).
fof(f16042,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK0) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK0) | $spl0),
  inference(subsumption_resolution,[],[f15980,f354])).
fof(f354,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f337,f247])).
fof(f15980,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK0) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK0) | $spl0),
  inference(superposition,[],[f2633,f364])).
fof(f2633,plain,(
  ( ! [X3] : (s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X3) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X3)) )),
  inference(superposition,[],[f237,f243])).
fof(f243,plain,(
  ( ! [X0] : (s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK1(X0)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0)) )),
  inference(cnf_transformation,[],[f177])).
fof(f177,plain,(
  ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK1(X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1])],[f91])).
fof(f91,plain,(
  ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | ? [X1] : s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1))))),
  inference(rectify,[],[f49])).
fof(f49,axiom,(
  ! [X6] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X6) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X6) | ? [X37] : s('type.extreal.extreal',X6) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X37))))),
  file('thm.extreal.le_infty|split|2','thm.extreal.extreal_cases')).
fof(f237,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.NegInf_0')))) )),
  inference(cnf_transformation,[],[f86])).
fof(f86,plain,(
  ! [X0] : s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.NegInf_0')))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X19] : s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X19))),s('type.extreal.extreal','const.extreal.NegInf_0')))),
  file('thm.extreal.le_infty|split|2','thm.extreal.extreal_le_def|split|5')).
fof(f9725,plain,(
  $false | ($spl16 | $spl19 | $spl64 | $spl70)),
  inference(subsumption_resolution,[],[f9724,f5015])).
fof(f5015,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.PosInf_0') | ($spl16 | $spl19)),
  inference(superposition,[],[f2838,f2837])).
fof(f2838,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',sK2('const.combin.I_1'(s('type.extreal.extreal','const.extreal.PosInf_0')),'type.extreal.extreal')) | $spl19),
  inference(cnf_transformation,[],[f2838_D])).
fof(f2838_D,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',sK2('const.combin.I_1'(s('type.extreal.extreal','const.extreal.PosInf_0')),'type.extreal.extreal')) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f9724,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.PosInf_0') | ($spl64 | $spl70)),
  inference(forward_demodulation,[],[f9680,f8793])).
fof(f9680,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK2(sK2(sK2('const.extreal.NegInf_0','type.extreal.extreal'),'type.extreal.extreal'),'type.extreal.extreal')) | $spl70),
  inference(superposition,[],[f290,f9364])).
fof(f290,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,sK2(X1,X0))) )),
  inference(cnf_transformation,[],[f204])).
fof(f204,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,sK2(X1,X0))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2])],[f122])).
fof(f122,plain,(
  ! [X0,X1] : ? [X2] : s(X0,X1) = s(X0,X2)),
  inference(flattening,[],[f121])).
fof(f121,plain,(
  ! [X0] : ! [X1] : ? [X2] : s(X0,X1) = s(X0,X2)),
  inference(rectify,[],[f67])).
fof(f67,axiom,(
  ! [X7] : ! [X11] : ? [X6] : s(X7,X6) = s(X7,X11)),
  file('thm.extreal.le_infty|split|2','thm.bool.EXISTS_REFL')).
fof(f9710,plain,(
  $false | ($spl16 | $spl19 | $spl64 | $spl70)),
  inference(subsumption_resolution,[],[f9709,f5015])).
fof(f9709,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.PosInf_0') | ($spl64 | $spl70)),
  inference(forward_demodulation,[],[f9659,f8793])).
fof(f9659,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK2(sK2(sK2('const.extreal.NegInf_0','type.extreal.extreal'),'type.extreal.extreal'),'type.extreal.extreal')) | $spl70),
  inference(superposition,[],[f9364,f290])).
fof(f9364,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK2(sK2(sK2(sK2('const.extreal.NegInf_0','type.extreal.extreal'),'type.extreal.extreal'),'type.extreal.extreal'),'type.extreal.extreal')) | $spl70),
  inference(cnf_transformation,[],[f9364_D])).
fof(f9364_D,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK2(sK2(sK2(sK2('const.extreal.NegInf_0','type.extreal.extreal'),'type.extreal.extreal'),'type.extreal.extreal'),'type.extreal.extreal')) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f9366,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK2(sK2(sK2(sK2('const.extreal.NegInf_0','type.extreal.extreal'),'type.extreal.extreal'),'type.extreal.extreal'),'type.extreal.extreal')) | $spl72),
  inference(cnf_transformation,[],[f9366_D])).
fof(f9366_D,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK2(sK2(sK2(sK2('const.extreal.NegInf_0','type.extreal.extreal'),'type.extreal.extreal'),'type.extreal.extreal'),'type.extreal.extreal')) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f9362,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK2(sK2(sK2(sK2('const.extreal.NegInf_0','type.extreal.extreal'),'type.extreal.extreal'),'type.extreal.extreal'),'type.extreal.extreal')) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK2(sK2(sK2(sK2('const.extreal.NegInf_0','type.extreal.extreal'),'type.extreal.extreal'),'type.extreal.extreal'),'type.extreal.extreal')) | $spl64),
  inference(subsumption_resolution,[],[f9324,f217])).
fof(f9324,plain,(
  ~p(s(bool,'T_0')) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK2(sK2(sK2(sK2('const.extreal.NegInf_0','type.extreal.extreal'),'type.extreal.extreal'),'type.extreal.extreal'),'type.extreal.extreal')) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK2(sK2(sK2(sK2('const.extreal.NegInf_0','type.extreal.extreal'),'type.extreal.extreal'),'type.extreal.extreal'),'type.extreal.extreal')) | $spl64),
  inference(superposition,[],[f311,f8935])).
fof(f8935,plain,(
  s(bool,'T_0') = s(bool,sK5(sK2(sK2(sK2(sK2('const.extreal.NegInf_0','type.extreal.extreal'),'type.extreal.extreal'),'type.extreal.extreal'),'type.extreal.extreal'))) | $spl64),
  inference(unit_resulting_resolution,[],[f8793,f1022])).
fof(f1022,plain,(
  ( ! [X5] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X5) | s(bool,'T_0') = s(bool,sK5(sK2(X5,'type.extreal.extreal')))) )),
  inference(superposition,[],[f543,f290])).
fof(f543,plain,(
  ( ! [X6] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X6) | s(bool,'T_0') = s(bool,sK5(X6))) )),
  inference(subsumption_resolution,[],[f530,f337])).
fof(f530,plain,(
  ( ! [X6] : (p(s(bool,'F_0')) | s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X6) | s(bool,'T_0') = s(bool,sK5(X6))) )),
  inference(superposition,[],[f312,f242])).
fof(f242,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.extreal.le_infty|split|2',aHL_BOOL_CASES)).
fof(f312,plain,(
  ( ! [X0] : (p(s(bool,sK5(X0))) | s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0)) )),
  inference(cnf_transformation,[],[f210])).
fof(f210,plain,(
  ! [X0] : ((~p(s(bool,sK5(X0))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)) & ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0)) | p(s(bool,sK5(X0)))) & s('type.realax.real','const.extreal.real_1'(s('type.extreal.extreal',X0))) = s('type.realax.real','const.bool.COND_3'(s(bool,sK5(X0)),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.min.@_1'(s(fun('type.realax.real',bool),'_dst_x0x1_2'(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal',X0))),s(fun('type.realax.real','type.extreal.extreal'),'const.extreal.Normal_0'))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5])],[f209])).
fof(f209,plain,(
  ! [X0] : ? [X1] : ((~p(s(bool,X1)) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)) & ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0)) | p(s(bool,X1))) & s('type.realax.real','const.extreal.real_1'(s('type.extreal.extreal',X0))) = s('type.realax.real','const.bool.COND_3'(s(bool,X1),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.min.@_1'(s(fun('type.realax.real',bool),'_dst_x0x1_2'(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal',X0))),s(fun('type.realax.real','type.extreal.extreal'),'const.extreal.Normal_0'))))))))),
  inference(flattening,[],[f208])).
fof(f208,plain,(
  ! [X0] : ? [X1] : (((~p(s(bool,X1)) | (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0))) & ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0)) | p(s(bool,X1)))) & s('type.realax.real','const.extreal.real_1'(s('type.extreal.extreal',X0))) = s('type.realax.real','const.bool.COND_3'(s(bool,X1),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.min.@_1'(s(fun('type.realax.real',bool),'_dst_x0x1_2'(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal',X0))),s(fun('type.realax.real','type.extreal.extreal'),'const.extreal.Normal_0'))))))))),
  inference(nnf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0] : ? [X1] : ((p(s(bool,X1)) <=> (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0))) & s('type.realax.real','const.extreal.real_1'(s('type.extreal.extreal',X0))) = s('type.realax.real','const.bool.COND_3'(s(bool,X1),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.min.@_1'(s(fun('type.realax.real',bool),'_dst_x0x1_2'(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal',X0))),s(fun('type.realax.real','type.extreal.extreal'),'const.extreal.Normal_0'))))))))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X6] : ? [X38] : ((p(s(bool,X38)) <=> (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X6) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X6))) & s('type.realax.real','const.extreal.real_1'(s('type.extreal.extreal',X6))) = s('type.realax.real','const.bool.COND_3'(s(bool,X38),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.min.@_1'(s(fun('type.realax.real',bool),'_dst_x0x1_2'(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal',X6))),s(fun('type.realax.real','type.extreal.extreal'),'const.extreal.Normal_0'))))))))),
  file('thm.extreal.le_infty|split|2','thm.extreal.real_def')).
fof(f311,plain,(
  ( ! [X0] : (~p(s(bool,sK5(X0))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)) )),
  inference(cnf_transformation,[],[f210])).
fof(f9188,plain,(
  $false | ($spl16 | $spl19 | $spl60 | $spl66)),
  inference(subsumption_resolution,[],[f9187,f5015])).
fof(f9187,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.PosInf_0') | ($spl60 | $spl66)),
  inference(forward_demodulation,[],[f9144,f8534])).
fof(f9144,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK2('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0')),'type.extreal.extreal')) | $spl66),
  inference(superposition,[],[f290,f9105])).
fof(f9173,plain,(
  $false | ($spl16 | $spl19 | $spl60 | $spl66)),
  inference(subsumption_resolution,[],[f9172,f5015])).
fof(f9172,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.PosInf_0') | ($spl60 | $spl66)),
  inference(forward_demodulation,[],[f9123,f8534])).
fof(f9123,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK2('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0')),'type.extreal.extreal')) | $spl66),
  inference(superposition,[],[f9105,f290])).
fof(f9105,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK2(sK2('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0')),'type.extreal.extreal'),'type.extreal.extreal')) | $spl66),
  inference(cnf_transformation,[],[f9105_D])).
fof(f9105_D,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK2(sK2('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0')),'type.extreal.extreal'),'type.extreal.extreal')) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f9107,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK2(sK2('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0')),'type.extreal.extreal'),'type.extreal.extreal')) | $spl68),
  inference(cnf_transformation,[],[f9107_D])).
fof(f9107_D,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK2(sK2('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0')),'type.extreal.extreal'),'type.extreal.extreal')) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f9103,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK2(sK2('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0')),'type.extreal.extreal'),'type.extreal.extreal')) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK2(sK2('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0')),'type.extreal.extreal'),'type.extreal.extreal')) | $spl60),
  inference(subsumption_resolution,[],[f9065,f217])).
fof(f9065,plain,(
  ~p(s(bool,'T_0')) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK2(sK2('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0')),'type.extreal.extreal'),'type.extreal.extreal')) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK2(sK2('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0')),'type.extreal.extreal'),'type.extreal.extreal')) | $spl60),
  inference(superposition,[],[f311,f8676])).
fof(f8676,plain,(
  s(bool,'T_0') = s(bool,sK5(sK2(sK2('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0')),'type.extreal.extreal'),'type.extreal.extreal'))) | $spl60),
  inference(unit_resulting_resolution,[],[f8534,f1022])).
fof(f8929,plain,(
  $false | ($spl16 | $spl19 | $spl56 | $spl62)),
  inference(subsumption_resolution,[],[f8928,f5015])).
fof(f8928,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.PosInf_0') | ($spl56 | $spl62)),
  inference(forward_demodulation,[],[f8885,f8159])).
fof(f8885,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK2(sK2('const.extreal.NegInf_0','type.extreal.extreal'),'type.extreal.extreal')) | $spl62),
  inference(superposition,[],[f290,f8791])).
fof(f8914,plain,(
  $false | ($spl16 | $spl19 | $spl56 | $spl62)),
  inference(subsumption_resolution,[],[f8913,f5015])).
fof(f8913,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.PosInf_0') | ($spl56 | $spl62)),
  inference(forward_demodulation,[],[f8864,f8159])).
fof(f8864,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK2(sK2('const.extreal.NegInf_0','type.extreal.extreal'),'type.extreal.extreal')) | $spl62),
  inference(superposition,[],[f8791,f290])).
fof(f8791,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK2(sK2(sK2('const.extreal.NegInf_0','type.extreal.extreal'),'type.extreal.extreal'),'type.extreal.extreal')) | $spl62),
  inference(cnf_transformation,[],[f8791_D])).
fof(f8791_D,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK2(sK2(sK2('const.extreal.NegInf_0','type.extreal.extreal'),'type.extreal.extreal'),'type.extreal.extreal')) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f8793,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK2(sK2(sK2('const.extreal.NegInf_0','type.extreal.extreal'),'type.extreal.extreal'),'type.extreal.extreal')) | $spl64),
  inference(cnf_transformation,[],[f8793_D])).
fof(f8793_D,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK2(sK2(sK2('const.extreal.NegInf_0','type.extreal.extreal'),'type.extreal.extreal'),'type.extreal.extreal')) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f8789,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK2(sK2(sK2('const.extreal.NegInf_0','type.extreal.extreal'),'type.extreal.extreal'),'type.extreal.extreal')) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK2(sK2(sK2('const.extreal.NegInf_0','type.extreal.extreal'),'type.extreal.extreal'),'type.extreal.extreal')) | $spl56),
  inference(subsumption_resolution,[],[f8751,f217])).
fof(f8751,plain,(
  ~p(s(bool,'T_0')) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK2(sK2(sK2('const.extreal.NegInf_0','type.extreal.extreal'),'type.extreal.extreal'),'type.extreal.extreal')) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK2(sK2(sK2('const.extreal.NegInf_0','type.extreal.extreal'),'type.extreal.extreal'),'type.extreal.extreal')) | $spl56),
  inference(superposition,[],[f311,f8383])).
fof(f8383,plain,(
  s(bool,'T_0') = s(bool,sK5(sK2(sK2(sK2('const.extreal.NegInf_0','type.extreal.extreal'),'type.extreal.extreal'),'type.extreal.extreal'))) | $spl56),
  inference(unit_resulting_resolution,[],[f8159,f1022])).
fof(f8615,plain,(
  $false | ($spl16 | $spl19 | $spl58)),
  inference(subsumption_resolution,[],[f8614,f5015])).
fof(f8614,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.PosInf_0') | $spl58),
  inference(forward_demodulation,[],[f8571,f282])).
fof(f282,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))) )),
  inference(cnf_transformation,[],[f114])).
fof(f114,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))),
  inference(flattening,[],[f113])).
fof(f113,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X7] : ! [X6] : s(X7,X6) = s(X7,'const.combin.I_1'(s(X7,X6)))),
  file('thm.extreal.le_infty|split|2','thm.combin.I_THM')).
fof(f8571,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))) | $spl58),
  inference(superposition,[],[f290,f8532])).
fof(f8600,plain,(
  $false | ($spl16 | $spl19 | $spl58)),
  inference(subsumption_resolution,[],[f8599,f5015])).
fof(f8599,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.PosInf_0') | $spl58),
  inference(forward_demodulation,[],[f8550,f282])).
fof(f8550,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))) | $spl58),
  inference(superposition,[],[f8532,f290])).
fof(f8532,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK2('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0')),'type.extreal.extreal')) | $spl58),
  inference(cnf_transformation,[],[f8532_D])).
fof(f8532_D,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK2('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0')),'type.extreal.extreal')) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f8534,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK2('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0')),'type.extreal.extreal')) | $spl60),
  inference(cnf_transformation,[],[f8534_D])).
fof(f8534_D,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK2('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0')),'type.extreal.extreal')) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f8530,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK2('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0')),'type.extreal.extreal')) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK2('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0')),'type.extreal.extreal'))),
  inference(subsumption_resolution,[],[f8492,f217])).
fof(f8492,plain,(
  ~p(s(bool,'T_0')) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK2('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0')),'type.extreal.extreal')) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK2('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0')),'type.extreal.extreal'))),
  inference(superposition,[],[f311,f8071])).
fof(f8071,plain,(
  s(bool,'T_0') = s(bool,sK5(sK2('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0')),'type.extreal.extreal')))),
  inference(unit_resulting_resolution,[],[f282,f1022])).
fof(f8377,plain,(
  $false | ($spl16 | $spl19 | $spl54)),
  inference(subsumption_resolution,[],[f8376,f5015])).
fof(f8376,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.PosInf_0') | $spl54),
  inference(forward_demodulation,[],[f8333,f290])).
fof(f8333,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK2('const.extreal.NegInf_0','type.extreal.extreal')) | $spl54),
  inference(superposition,[],[f290,f8157])).
fof(f8362,plain,(
  $false | ($spl16 | $spl19 | $spl54)),
  inference(subsumption_resolution,[],[f8361,f5015])).
fof(f8361,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.PosInf_0') | $spl54),
  inference(forward_demodulation,[],[f8312,f290])).
fof(f8312,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK2('const.extreal.NegInf_0','type.extreal.extreal')) | $spl54),
  inference(superposition,[],[f8157,f290])).
fof(f8157,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK2(sK2('const.extreal.NegInf_0','type.extreal.extreal'),'type.extreal.extreal')) | $spl54),
  inference(cnf_transformation,[],[f8157_D])).
fof(f8157_D,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK2(sK2('const.extreal.NegInf_0','type.extreal.extreal'),'type.extreal.extreal')) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f8159,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK2(sK2('const.extreal.NegInf_0','type.extreal.extreal'),'type.extreal.extreal')) | $spl56),
  inference(cnf_transformation,[],[f8159_D])).
fof(f8159_D,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK2(sK2('const.extreal.NegInf_0','type.extreal.extreal'),'type.extreal.extreal')) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f8155,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK2(sK2('const.extreal.NegInf_0','type.extreal.extreal'),'type.extreal.extreal')) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK2(sK2('const.extreal.NegInf_0','type.extreal.extreal'),'type.extreal.extreal'))),
  inference(subsumption_resolution,[],[f8117,f217])).
fof(f8117,plain,(
  ~p(s(bool,'T_0')) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK2(sK2('const.extreal.NegInf_0','type.extreal.extreal'),'type.extreal.extreal')) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK2(sK2('const.extreal.NegInf_0','type.extreal.extreal'),'type.extreal.extreal'))),
  inference(superposition,[],[f311,f8073])).
fof(f8073,plain,(
  s(bool,'T_0') = s(bool,sK5(sK2(sK2('const.extreal.NegInf_0','type.extreal.extreal'),'type.extreal.extreal')))),
  inference(unit_resulting_resolution,[],[f290,f1022])).
fof(f7204,plain,(
  $false | ($spl8 | $spl16 | $spl19 | $spl26)),
  inference(subsumption_resolution,[],[f7203,f5015])).
fof(f7203,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.PosInf_0') | ($spl8 | $spl26)),
  inference(backward_demodulation,[],[f7166,f1962])).
fof(f7166,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK2(sK2('const.extreal.PosInf_0','type.extreal.extreal'),'type.extreal.extreal')) | $spl26),
  inference(superposition,[],[f3220,f290])).
fof(f3218,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK2(sK2(sK2('const.extreal.PosInf_0','type.extreal.extreal'),'type.extreal.extreal'),'type.extreal.extreal')) | $spl24),
  inference(cnf_transformation,[],[f3218_D])).
fof(f3218_D,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK2(sK2(sK2('const.extreal.PosInf_0','type.extreal.extreal'),'type.extreal.extreal'),'type.extreal.extreal')) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f3220,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK2(sK2(sK2('const.extreal.PosInf_0','type.extreal.extreal'),'type.extreal.extreal'),'type.extreal.extreal')) | $spl26),
  inference(cnf_transformation,[],[f3220_D])).
fof(f3220_D,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK2(sK2(sK2('const.extreal.PosInf_0','type.extreal.extreal'),'type.extreal.extreal'),'type.extreal.extreal')) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f7158,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK2(sK2(sK2('const.extreal.PosInf_0','type.extreal.extreal'),'type.extreal.extreal'),'type.extreal.extreal')) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK2(sK2(sK2('const.extreal.PosInf_0','type.extreal.extreal'),'type.extreal.extreal'),'type.extreal.extreal')) | $spl8),
  inference(subsumption_resolution,[],[f5611,f217])).
fof(f5611,plain,(
  ~p(s(bool,'T_0')) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK2(sK2(sK2('const.extreal.PosInf_0','type.extreal.extreal'),'type.extreal.extreal'),'type.extreal.extreal')) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK2(sK2(sK2('const.extreal.PosInf_0','type.extreal.extreal'),'type.extreal.extreal'),'type.extreal.extreal')) | $spl8),
  inference(superposition,[],[f311,f2673])).
fof(f2673,plain,(
  s(bool,'T_0') = s(bool,sK5(sK2(sK2(sK2('const.extreal.PosInf_0','type.extreal.extreal'),'type.extreal.extreal'),'type.extreal.extreal'))) | $spl8),
  inference(unit_resulting_resolution,[],[f1962,f804])).
fof(f804,plain,(
  ( ! [X2] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X2) | s(bool,'T_0') = s(bool,sK5(sK2(X2,'type.extreal.extreal')))) )),
  inference(superposition,[],[f542,f290])).
fof(f542,plain,(
  ( ! [X5] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X5) | s(bool,'T_0') = s(bool,sK5(X5))) )),
  inference(subsumption_resolution,[],[f529,f337])).
fof(f529,plain,(
  ( ! [X5] : (p(s(bool,'F_0')) | s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X5) | s(bool,'T_0') = s(bool,sK5(X5))) )),
  inference(superposition,[],[f313,f242])).
fof(f313,plain,(
  ( ! [X0] : (p(s(bool,sK5(X0))) | s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0)) )),
  inference(cnf_transformation,[],[f210])).
fof(f7145,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK2(sK2(sK2(sK2('const.extreal.PosInf_0','type.extreal.extreal'),'type.extreal.extreal'),'type.extreal.extreal'),'type.extreal.extreal')) | $spl50),
  inference(cnf_transformation,[],[f7145_D])).
fof(f7145_D,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK2(sK2(sK2(sK2('const.extreal.PosInf_0','type.extreal.extreal'),'type.extreal.extreal'),'type.extreal.extreal'),'type.extreal.extreal')) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f7147,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK2(sK2(sK2(sK2('const.extreal.PosInf_0','type.extreal.extreal'),'type.extreal.extreal'),'type.extreal.extreal'),'type.extreal.extreal')) | $spl52),
  inference(cnf_transformation,[],[f7147_D])).
fof(f7147_D,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK2(sK2(sK2(sK2('const.extreal.PosInf_0','type.extreal.extreal'),'type.extreal.extreal'),'type.extreal.extreal'),'type.extreal.extreal')) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f7143,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK2(sK2(sK2(sK2('const.extreal.PosInf_0','type.extreal.extreal'),'type.extreal.extreal'),'type.extreal.extreal'),'type.extreal.extreal')) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK2(sK2(sK2(sK2('const.extreal.PosInf_0','type.extreal.extreal'),'type.extreal.extreal'),'type.extreal.extreal'),'type.extreal.extreal')) | $spl24),
  inference(subsumption_resolution,[],[f7106,f217])).
fof(f7106,plain,(
  ~p(s(bool,'T_0')) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK2(sK2(sK2(sK2('const.extreal.PosInf_0','type.extreal.extreal'),'type.extreal.extreal'),'type.extreal.extreal'),'type.extreal.extreal')) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK2(sK2(sK2(sK2('const.extreal.PosInf_0','type.extreal.extreal'),'type.extreal.extreal'),'type.extreal.extreal'),'type.extreal.extreal')) | $spl24),
  inference(superposition,[],[f311,f3229])).
fof(f3229,plain,(
  s(bool,'T_0') = s(bool,sK5(sK2(sK2(sK2(sK2('const.extreal.PosInf_0','type.extreal.extreal'),'type.extreal.extreal'),'type.extreal.extreal'),'type.extreal.extreal'))) | $spl24),
  inference(unit_resulting_resolution,[],[f3218,f804])).
fof(f4598,plain,(
  ( ! [X10] : (s(bool,'T_0') = s(bool,sK5(sK2('const.extreal.extreal_logr_2'(s('type.realax.real',X10),s('type.extreal.extreal','const.extreal.NegInf_0')),'type.extreal.extreal')))) ) | $spl48),
  inference(cnf_transformation,[],[f4598_D])).
fof(f4598_D,plain,(
  ( ! [X10] : (s(bool,'T_0') = s(bool,sK5(sK2('const.extreal.extreal_logr_2'(s('type.realax.real',X10),s('type.extreal.extreal','const.extreal.NegInf_0')),'type.extreal.extreal')))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f821,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.PosInf_0') | $spl7),
  inference(cnf_transformation,[],[f821_D])).
fof(f821_D,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.PosInf_0') <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f1930,plain,(
  ( ! [X10] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.PosInf_0') | s(bool,'T_0') = s(bool,sK5(sK2('const.extreal.extreal_logr_2'(s('type.realax.real',X10),s('type.extreal.extreal','const.extreal.NegInf_0')),'type.extreal.extreal')))) )),
  inference(superposition,[],[f804,f232])).
fof(f232,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_logr_2'(s('type.realax.real',X0),s('type.extreal.extreal','const.extreal.NegInf_0')))) )),
  inference(cnf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0] : s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_logr_2'(s('type.realax.real',X0),s('type.extreal.extreal','const.extreal.NegInf_0')))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X39] : s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_logr_2'(s('type.realax.real',X39),s('type.extreal.extreal','const.extreal.NegInf_0')))),
  file('thm.extreal.le_infty|split|2','thm.extreal.extreal_logr_def|split|2')).
fof(f4597,plain,(
  ( ! [X9] : (s(bool,'T_0') = s(bool,sK5(sK2('const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X9))),s('type.extreal.extreal','const.extreal.NegInf_0')),'type.extreal.extreal')))) ) | $spl46),
  inference(cnf_transformation,[],[f4597_D])).
fof(f4597_D,plain,(
  ( ! [X9] : (s(bool,'T_0') = s(bool,sK5(sK2('const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X9))),s('type.extreal.extreal','const.extreal.NegInf_0')),'type.extreal.extreal')))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f1929,plain,(
  ( ! [X9] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.PosInf_0') | s(bool,'T_0') = s(bool,sK5(sK2('const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X9))),s('type.extreal.extreal','const.extreal.NegInf_0')),'type.extreal.extreal')))) )),
  inference(superposition,[],[f804,f846])).
fof(f846,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0))),s('type.extreal.extreal','const.extreal.NegInf_0')))) )),
  inference(superposition,[],[f235,f233])).
fof(f233,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0] : s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f65])).
fof(f65,axiom,(
  ! [X31] : s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X31))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X31)))))),
  file('thm.extreal.le_infty|split|2','thm.extreal.extreal_of_num_def')).
fof(f235,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.NegInf_0')))) )),
  inference(cnf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0] : s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.NegInf_0')))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X19] : s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X19))),s('type.extreal.extreal','const.extreal.NegInf_0')))),
  file('thm.extreal.le_infty|split|2','thm.extreal.extreal_add_def|split|6')).
fof(f4596,plain,(
  ( ! [X8] : (s(bool,'T_0') = s(bool,sK5(sK2('const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X8))),s('type.extreal.extreal','const.extreal.NegInf_0')),'type.extreal.extreal')))) ) | $spl44),
  inference(cnf_transformation,[],[f4596_D])).
fof(f4596_D,plain,(
  ( ! [X8] : (s(bool,'T_0') = s(bool,sK5(sK2('const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X8))),s('type.extreal.extreal','const.extreal.NegInf_0')),'type.extreal.extreal')))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f1928,plain,(
  ( ! [X8] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.PosInf_0') | s(bool,'T_0') = s(bool,sK5(sK2('const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X8))),s('type.extreal.extreal','const.extreal.NegInf_0')),'type.extreal.extreal')))) )),
  inference(superposition,[],[f804,f235])).
fof(f4595,plain,(
  ( ! [X7] : (s(bool,'T_0') = s(bool,sK5(sK2('const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X7)))),'type.extreal.extreal')))) ) | $spl42),
  inference(cnf_transformation,[],[f4595_D])).
fof(f4595_D,plain,(
  ( ! [X7] : (s(bool,'T_0') = s(bool,sK5(sK2('const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X7)))),'type.extreal.extreal')))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f1927,plain,(
  ( ! [X7] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.PosInf_0') | s(bool,'T_0') = s(bool,sK5(sK2('const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X7)))),'type.extreal.extreal')))) )),
  inference(superposition,[],[f804,f812])).
fof(f812,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0)))))) )),
  inference(superposition,[],[f234,f233])).
fof(f234,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))))) )),
  inference(cnf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0] : s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X21] : s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X21)))))),
  file('thm.extreal.le_infty|split|2','thm.extreal.extreal_add_def|split|5')).
fof(f4594,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,sK5(sK2('const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X6)))),'type.extreal.extreal')))) ) | $spl40),
  inference(cnf_transformation,[],[f4594_D])).
fof(f4594_D,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,sK5(sK2('const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X6)))),'type.extreal.extreal')))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f1926,plain,(
  ( ! [X6] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.PosInf_0') | s(bool,'T_0') = s(bool,sK5(sK2('const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X6)))),'type.extreal.extreal')))) )),
  inference(superposition,[],[f804,f234])).
fof(f4593,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,sK5(sK2('const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X5))),s('type.extreal.extreal','const.extreal.PosInf_0')),'type.extreal.extreal')))) ) | $spl38),
  inference(cnf_transformation,[],[f4593_D])).
fof(f4593_D,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,sK5(sK2('const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X5))),s('type.extreal.extreal','const.extreal.PosInf_0')),'type.extreal.extreal')))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f1925,plain,(
  ( ! [X5] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.PosInf_0') | s(bool,'T_0') = s(bool,sK5(sK2('const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X5))),s('type.extreal.extreal','const.extreal.PosInf_0')),'type.extreal.extreal')))) )),
  inference(superposition,[],[f804,f1055])).
fof(f1055,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0))),s('type.extreal.extreal','const.extreal.PosInf_0')))) )),
  inference(superposition,[],[f239,f233])).
fof(f239,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0')))) )),
  inference(cnf_transformation,[],[f88])).
fof(f88,plain,(
  ! [X0] : s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0')))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X20] : s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X20))),s('type.extreal.extreal','const.extreal.PosInf_0')))),
  file('thm.extreal.le_infty|split|2','thm.extreal.extreal_sub_def|split|3')).
fof(f4592,plain,(
  ( ! [X4] : (s(bool,'T_0') = s(bool,sK5(sK2('const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4))),s('type.extreal.extreal','const.extreal.PosInf_0')),'type.extreal.extreal')))) ) | $spl36),
  inference(cnf_transformation,[],[f4592_D])).
fof(f4592_D,plain,(
  ( ! [X4] : (s(bool,'T_0') = s(bool,sK5(sK2('const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4))),s('type.extreal.extreal','const.extreal.PosInf_0')),'type.extreal.extreal')))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f1924,plain,(
  ( ! [X4] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.PosInf_0') | s(bool,'T_0') = s(bool,sK5(sK2('const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4))),s('type.extreal.extreal','const.extreal.PosInf_0')),'type.extreal.extreal')))) )),
  inference(superposition,[],[f804,f239])).
fof(f4591,plain,(
  ( ! [X3] : (s(bool,'T_0') = s(bool,sK5(sK2('const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X3)))),'type.extreal.extreal')))) ) | $spl34),
  inference(cnf_transformation,[],[f4591_D])).
fof(f4591_D,plain,(
  ( ! [X3] : (s(bool,'T_0') = s(bool,sK5(sK2('const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X3)))),'type.extreal.extreal')))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f1923,plain,(
  ( ! [X3] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.PosInf_0') | s(bool,'T_0') = s(bool,sK5(sK2('const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X3)))),'type.extreal.extreal')))) )),
  inference(superposition,[],[f804,f893])).
fof(f893,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0)))))) )),
  inference(superposition,[],[f236,f233])).
fof(f236,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))))) )),
  inference(cnf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0] : s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X21] : s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X21)))))),
  file('thm.extreal.le_infty|split|2','thm.extreal.extreal_sub_def|split|5')).
fof(f4590,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,sK5(sK2('const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))),'type.extreal.extreal')))) ) | $spl32),
  inference(cnf_transformation,[],[f4590_D])).
fof(f4590_D,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,sK5(sK2('const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))),'type.extreal.extreal')))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f1922,plain,(
  ( ! [X2] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.PosInf_0') | s(bool,'T_0') = s(bool,sK5(sK2('const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))),'type.extreal.extreal')))) )),
  inference(superposition,[],[f804,f236])).
fof(f4546,plain,(
  $false | ($spl17 | $spl18)),
  inference(subsumption_resolution,[],[f3983,f4280])).
fof(f4280,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.PosInf_0') | $spl18),
  inference(superposition,[],[f282,f3988])).
fof(f3988,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.combin.I_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))) | $spl18),
  inference(superposition,[],[f2839,f290])).
fof(f3983,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.PosInf_0') | ($spl17 | $spl18)),
  inference(forward_demodulation,[],[f2836,f2839])).
fof(f2836,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',sK2('const.combin.I_1'(s('type.extreal.extreal','const.extreal.PosInf_0')),'type.extreal.extreal')) | $spl17),
  inference(cnf_transformation,[],[f2836_D])).
fof(f2836_D,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',sK2('const.combin.I_1'(s('type.extreal.extreal','const.extreal.PosInf_0')),'type.extreal.extreal')) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f4537,plain,(
  $false | ($spl3 | $spl18)),
  inference(subsumption_resolution,[],[f4501,f3444])).
fof(f3444,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,X1))),s(X0,X1)))) )),
  inference(unit_resulting_resolution,[],[f2389,f248])).
fof(f2389,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,X1))),s(X0,X1))))) )),
  inference(forward_demodulation,[],[f2231,f290])).
fof(f2231,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,sK2(X1,X0)))),s(X0,X1))))) )),
  inference(unit_resulting_resolution,[],[f290,f326])).
fof(f326,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,X1))),s(X0,X2)))) | s(X0,X1) != s(X0,X2)) )),
  inference(cnf_transformation,[],[f214])).
fof(f214,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,X1))),s(X0,X2))))))),
  inference(nnf_transformation,[],[f156])).
fof(f156,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,X1))),s(X0,X2)))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f63])).
fof(f63,axiom,(
  ! [X44,X45,X46] : (p(s(bool,happ(s(fun(X44,bool),'=_1'(s(X44,X45))),s(X44,X46)))) <=> s(X44,X45) = s(X44,X46))),
  file('thm.extreal.le_infty|split|2',arity_1_HL_EQ)).
fof(f4501,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.NegInf_0'))) | ($spl3 | $spl18)),
  inference(backward_demodulation,[],[f4445,f2175])).
fof(f2175,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',sK0))) | $spl3),
  inference(unit_resulting_resolution,[],[f1969,f247])).
fof(f1969,plain,(
  ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',sK0)))) | $spl3),
  inference(unit_resulting_resolution,[],[f336,f325])).
fof(f325,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) )),
  inference(cnf_transformation,[],[f214])).
fof(f4445,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0)) ) | $spl18),
  inference(global_subsumption,[],[f216,f215,f217,f219,f220,f337,f226,f227,f228,f229,f230,f231,f232,f233,f234,f235,f236,f237,f238,f239,f240,f241,f242,f243,f248,f247,f252,f251,f281,f282,f339,f290,f291,f292,f297,f298,f299,f300,f301,f302,f310,f309,f308,f307,f306,f305,f304,f303,f314,f313,f312,f311,f315,f319,f316,f321,f320,f322,f323,f324,f326,f325,f327,f328,f340,f341,f354,f368,f367,f366,f365,f404,f403,f406,f410,f409,f412,f416,f426,f428,f495,f502,f511,f512,f542,f543,f545,f597,f707,f706,f726,f756,f765,f803,f804,f805,f806,f807,f808,f812,f813,f842,f846,f858,f862,f876,f880,f889,f893,f905,f909,f923,f939,f950,f960,f977,f1022,f1023,f1024,f1025,f1051,f1050,f1055,f1069,f1080,f1099,f1101,f1105,f1106,f1145,f1144,f1143,f1142,f1184,f1195,f1196,f1258,f1268,f1284,f1283,f1323,f1506,f1533,f1532,f1531,f1530,f1621,f1642,f1643,f1743,f1744,f1745,f1747,f1755,f1754,f1753,f1752,f1848,f1909,f1873,f1874,f1919,f1918,f1917,f1916,f1915,f1914,f1920,f1921,f1931,f1932,f1933,f1934,f1935,f1936,f1960,f2076,f1979,f2081,f1970,f1993,f1994,f2103,f2037,f2038,f2145,f2355,f2364,f2365,f2366,f2378,f2389,f2272,f2273,f2452,f2316,f2317,f2494,f2519,f2576,f2630,f2631,f2632,f2633,f2635,f2636,f2637,f2638,f2640,f2641,f2642,f2643,f2655,f2656,f2657,f2658,f2722,f2723,f2730,f2731,f2733,f2764,f2835,f2915,f2574,f1266,f2582,f1265,f2581,f1264,f2580,f1263,f2579,f1863,f2578,f1633,f2577,f294,f293,f3343,f3419,f3420,f3366,f3378,f3379,f3380,f3433,f3427,f3444,f3438,f3576,f3577,f3700,f3781,f3782,f3724,f3737,f3738,f3739,f3806,f3807,f3814,f3813,f3825,f3824,f3832,f3833,f3844,f3845,f3846,f3847,f3860,f3861,f3862,f3863,f3864,f3865,f3866,f3867,f2839,f3988,f4032,f4031,f4033,f4034,f4035,f4036,f4118,f4129,f4130,f4131,f4132,f4134,f4135,f4136,f4138,f4139,f4140,f4141,f4142,f4144,f4145,f4146,f4147,f4148,f4155,f4156,f4157,f4158,f4159,f4162,f4163,f4164,f4202,f4203,f4204,f4205,f4207,f4208,f4210,f4211,f4212,f4213,f4214,f4216,f4217,f4218,f4219,f4220,f4227,f4228,f4229,f4230,f4231,f4234,f4235,f4236,f4280,f4317,f4422,f4320,f4421,f4420,f4419,f4327,f4336,f4337,f4343,f4347,f4349,f4350,f4437,f4438,f4439,f4353,f4354,f4355,f4356,f4357,f4358,f4360,f4361,f4440,f4363,f4441,f4364,f4442,f4365,f4443,f4366,f4444,f4368,f4418])).
fof(f4418,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)))) ) | $spl18),
  inference(duplicate_literal_removal,[],[f4369])).
fof(f4369,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0)) ) | $spl18),
  inference(backward_demodulation,[],[f4280,f2630])).
fof(f4368,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~p(s(bool,happ(s(fun(X5,bool),X7),s(X5,sK3(X9,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(fun(X5,bool),X6),s(fun(X5,bool),X7),s(fun('type.realax.real',fun(X5,bool)),X8)),X5))))) | p(s(bool,'|dollar|exists_1'(s(fun(X5,bool),X7))))) ) | $spl18),
  inference(backward_demodulation,[],[f4280,f2578])).
fof(f4444,plain,(
  ( ! [X21] : (s(bool,'T_0') = s(bool,sK5(sK2(sK2(X21,'type.extreal.extreal'),'type.extreal.extreal')))) ) | $spl18),
  inference(global_subsumption,[],[f216,f215,f217,f219,f220,f337,f226,f227,f228,f229,f230,f231,f232,f233,f234,f235,f236,f237,f238,f239,f240,f241,f242,f243,f248,f247,f252,f251,f281,f282,f339,f290,f291,f292,f297,f298,f299,f300,f301,f302,f310,f309,f308,f307,f306,f305,f304,f303,f314,f313,f312,f311,f315,f319,f316,f321,f320,f322,f323,f324,f326,f325,f327,f328,f340,f341,f354,f368,f367,f366,f365,f404,f403,f406,f410,f409,f412,f416,f426,f428,f495,f502,f511,f512,f542,f543,f545,f597,f707,f706,f726,f756,f765,f803,f804,f805,f806,f807,f808,f812,f813,f842,f846,f858,f862,f876,f880,f889,f893,f905,f909,f923,f939,f950,f960,f977,f1022,f1023,f1024,f1025,f1051,f1050,f1055,f1069,f1080,f1099,f1101,f1105,f1106,f1145,f1144,f1143,f1142,f1184,f1195,f1196,f1258,f1268,f1284,f1283,f1323,f1506,f1533,f1532,f1531,f1530,f1621,f1642,f1643,f1743,f1744,f1745,f1747,f1755,f1754,f1753,f1752,f1848,f1909,f1873,f1874,f1919,f1918,f1917,f1916,f1915,f1914,f1920,f1921,f1931,f1932,f1933,f1934,f1935,f1936,f1960,f2076,f1979,f2081,f1970,f1993,f1994,f2103,f2037,f2038,f2145,f2355,f2364,f2365,f2366,f2378,f2389,f2272,f2273,f2452,f2316,f2317,f2494,f2519,f2576,f2630,f2631,f2632,f2633,f2635,f2636,f2637,f2638,f2640,f2641,f2642,f2643,f2655,f2656,f2657,f2658,f2722,f2723,f2730,f2731,f2733,f2764,f2835,f2915,f2574,f1266,f2582,f1265,f2581,f1264,f2580,f1263,f2579,f1863,f2578,f1633,f2577,f294,f293,f3343,f3419,f3420,f3366,f3378,f3379,f3380,f3433,f3427,f3444,f3438,f3576,f3577,f3700,f3781,f3782,f3724,f3737,f3738,f3739,f3806,f3807,f3814,f3813,f3825,f3824,f3832,f3833,f3844,f3845,f3846,f3847,f3860,f3861,f3862,f3863,f3864,f3865,f3866,f3867,f2839,f3988,f4032,f4031,f4033,f4034,f4035,f4036,f4118,f4129,f4130,f4131,f4132,f4134,f4135,f4136,f4138,f4139,f4140,f4141,f4142,f4144,f4145,f4146,f4147,f4148,f4155,f4156,f4157,f4158,f4159,f4162,f4163,f4164,f4202,f4203,f4204,f4205,f4207,f4208,f4210,f4211,f4212,f4213,f4214,f4216,f4217,f4218,f4219,f4220,f4227,f4228,f4229,f4230,f4231,f4234,f4235,f4236,f4280,f4317,f4422,f4320,f4421,f4420,f4419,f4327,f4336,f4337,f4343,f4347,f4349,f4350,f4437,f4438,f4439,f4353,f4354,f4355,f4356,f4357,f4358,f4360,f4361,f4440,f4363,f4441,f4364,f4442,f4365,f4443,f4366])).
fof(f4366,plain,(
  ( ! [X21] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X21) | s(bool,'T_0') = s(bool,sK5(sK2(sK2(X21,'type.extreal.extreal'),'type.extreal.extreal')))) ) | $spl18),
  inference(backward_demodulation,[],[f4280,f1936])).
fof(f4443,plain,(
  ( ! [X19,X20] : (s(bool,'T_0') = s(bool,sK5(sK2('const.ind_type.FCONS_3'(s('type.extreal.extreal',X19),s(fun('type.num.num','type.extreal.extreal'),X20),s('type.num.num','const.num.|48|_0')),'type.extreal.extreal')))) ) | $spl18),
  inference(global_subsumption,[],[f216,f215,f217,f219,f220,f337,f226,f227,f228,f229,f230,f231,f232,f233,f234,f235,f236,f237,f238,f239,f240,f241,f242,f243,f248,f247,f252,f251,f281,f282,f339,f290,f291,f292,f297,f298,f299,f300,f301,f302,f310,f309,f308,f307,f306,f305,f304,f303,f314,f313,f312,f311,f315,f319,f316,f321,f320,f322,f323,f324,f326,f325,f327,f328,f340,f341,f354,f368,f367,f366,f365,f404,f403,f406,f410,f409,f412,f416,f426,f428,f495,f502,f511,f512,f542,f543,f545,f597,f707,f706,f726,f756,f765,f803,f804,f805,f806,f807,f808,f812,f813,f842,f846,f858,f862,f876,f880,f889,f893,f905,f909,f923,f939,f950,f960,f977,f1022,f1023,f1024,f1025,f1051,f1050,f1055,f1069,f1080,f1099,f1101,f1105,f1106,f1145,f1144,f1143,f1142,f1184,f1195,f1196,f1258,f1268,f1284,f1283,f1323,f1506,f1533,f1532,f1531,f1530,f1621,f1642,f1643,f1743,f1744,f1745,f1747,f1755,f1754,f1753,f1752,f1848,f1909,f1873,f1874,f1919,f1918,f1917,f1916,f1915,f1914,f1920,f1921,f1931,f1932,f1933,f1934,f1935,f1936,f1960,f2076,f1979,f2081,f1970,f1993,f1994,f2103,f2037,f2038,f2145,f2355,f2364,f2365,f2366,f2378,f2389,f2272,f2273,f2452,f2316,f2317,f2494,f2519,f2576,f2630,f2631,f2632,f2633,f2635,f2636,f2637,f2638,f2640,f2641,f2642,f2643,f2655,f2656,f2657,f2658,f2722,f2723,f2730,f2731,f2733,f2764,f2835,f2915,f2574,f1266,f2582,f1265,f2581,f1264,f2580,f1263,f2579,f1863,f2578,f1633,f2577,f294,f293,f3343,f3419,f3420,f3366,f3378,f3379,f3380,f3433,f3427,f3444,f3438,f3576,f3577,f3700,f3781,f3782,f3724,f3737,f3738,f3739,f3806,f3807,f3814,f3813,f3825,f3824,f3832,f3833,f3844,f3845,f3846,f3847,f3860,f3861,f3862,f3863,f3864,f3865,f3866,f3867,f2839,f3988,f4032,f4031,f4033,f4034,f4035,f4036,f4118,f4129,f4130,f4131,f4132,f4134,f4135,f4136,f4138,f4139,f4140,f4141,f4142,f4144,f4145,f4146,f4147,f4148,f4155,f4156,f4157,f4158,f4159,f4162,f4163,f4164,f4202,f4203,f4204,f4205,f4207,f4208,f4210,f4211,f4212,f4213,f4214,f4216,f4217,f4218,f4219,f4220,f4227,f4228,f4229,f4230,f4231,f4234,f4235,f4236,f4280,f4317,f4422,f4320,f4421,f4420,f4419,f4327,f4336,f4337,f4343,f4347,f4349,f4350,f4437,f4438,f4439,f4353,f4354,f4355,f4356,f4357,f4358,f4360,f4361,f4440,f4363,f4441,f4364,f4442,f4365])).
fof(f4365,plain,(
  ( ! [X19,X20] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X19) | s(bool,'T_0') = s(bool,sK5(sK2('const.ind_type.FCONS_3'(s('type.extreal.extreal',X19),s(fun('type.num.num','type.extreal.extreal'),X20),s('type.num.num','const.num.|48|_0')),'type.extreal.extreal')))) ) | $spl18),
  inference(backward_demodulation,[],[f4280,f1935])).
fof(f4442,plain,(
  ( ! [X18] : (s(bool,'T_0') = s(bool,sK5(sK2('const.combin.I_1'(s('type.extreal.extreal',X18)),'type.extreal.extreal')))) ) | $spl18),
  inference(global_subsumption,[],[f216,f215,f217,f219,f220,f337,f226,f227,f228,f229,f230,f231,f232,f233,f234,f235,f236,f237,f238,f239,f240,f241,f242,f243,f248,f247,f252,f251,f281,f282,f339,f290,f291,f292,f297,f298,f299,f300,f301,f302,f310,f309,f308,f307,f306,f305,f304,f303,f314,f313,f312,f311,f315,f319,f316,f321,f320,f322,f323,f324,f326,f325,f327,f328,f340,f341,f354,f368,f367,f366,f365,f404,f403,f406,f410,f409,f412,f416,f426,f428,f495,f502,f511,f512,f542,f543,f545,f597,f707,f706,f726,f756,f765,f803,f804,f805,f806,f807,f808,f812,f813,f842,f846,f858,f862,f876,f880,f889,f893,f905,f909,f923,f939,f950,f960,f977,f1022,f1023,f1024,f1025,f1051,f1050,f1055,f1069,f1080,f1099,f1101,f1105,f1106,f1145,f1144,f1143,f1142,f1184,f1195,f1196,f1258,f1268,f1284,f1283,f1323,f1506,f1533,f1532,f1531,f1530,f1621,f1642,f1643,f1743,f1744,f1745,f1747,f1755,f1754,f1753,f1752,f1848,f1909,f1873,f1874,f1919,f1918,f1917,f1916,f1915,f1914,f1920,f1921,f1931,f1932,f1933,f1934,f1935,f1936,f1960,f2076,f1979,f2081,f1970,f1993,f1994,f2103,f2037,f2038,f2145,f2355,f2364,f2365,f2366,f2378,f2389,f2272,f2273,f2452,f2316,f2317,f2494,f2519,f2576,f2630,f2631,f2632,f2633,f2635,f2636,f2637,f2638,f2640,f2641,f2642,f2643,f2655,f2656,f2657,f2658,f2722,f2723,f2730,f2731,f2733,f2764,f2835,f2915,f2574,f1266,f2582,f1265,f2581,f1264,f2580,f1263,f2579,f1863,f2578,f1633,f2577,f294,f293,f3343,f3419,f3420,f3366,f3378,f3379,f3380,f3433,f3427,f3444,f3438,f3576,f3577,f3700,f3781,f3782,f3724,f3737,f3738,f3739,f3806,f3807,f3814,f3813,f3825,f3824,f3832,f3833,f3844,f3845,f3846,f3847,f3860,f3861,f3862,f3863,f3864,f3865,f3866,f3867,f2839,f3988,f4032,f4031,f4033,f4034,f4035,f4036,f4118,f4129,f4130,f4131,f4132,f4134,f4135,f4136,f4138,f4139,f4140,f4141,f4142,f4144,f4145,f4146,f4147,f4148,f4155,f4156,f4157,f4158,f4159,f4162,f4163,f4164,f4202,f4203,f4204,f4205,f4207,f4208,f4210,f4211,f4212,f4213,f4214,f4216,f4217,f4218,f4219,f4220,f4227,f4228,f4229,f4230,f4231,f4234,f4235,f4236,f4280,f4317,f4422,f4320,f4421,f4420,f4419,f4327,f4336,f4337,f4343,f4347,f4349,f4350,f4437,f4438,f4439,f4353,f4354,f4355,f4356,f4357,f4358,f4360,f4361,f4440,f4363,f4441,f4364])).
fof(f4364,plain,(
  ( ! [X18] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X18) | s(bool,'T_0') = s(bool,sK5(sK2('const.combin.I_1'(s('type.extreal.extreal',X18)),'type.extreal.extreal')))) ) | $spl18),
  inference(backward_demodulation,[],[f4280,f1934])).
fof(f4441,plain,(
  ( ! [X17] : (s(bool,'T_0') = s(bool,sK5(sK2(happ(s(fun('type.extreal.extreal','type.extreal.extreal'),'const.combin.I_0'),s('type.extreal.extreal',X17)),'type.extreal.extreal')))) ) | $spl18),
  inference(global_subsumption,[],[f216,f215,f217,f219,f220,f337,f226,f227,f228,f229,f230,f231,f232,f233,f234,f235,f236,f237,f238,f239,f240,f241,f242,f243,f248,f247,f252,f251,f281,f282,f339,f290,f291,f292,f297,f298,f299,f300,f301,f302,f310,f309,f308,f307,f306,f305,f304,f303,f314,f313,f312,f311,f315,f319,f316,f321,f320,f322,f323,f324,f326,f325,f327,f328,f340,f341,f354,f368,f367,f366,f365,f404,f403,f406,f410,f409,f412,f416,f426,f428,f495,f502,f511,f512,f542,f543,f545,f597,f707,f706,f726,f756,f765,f803,f804,f805,f806,f807,f808,f812,f813,f842,f846,f858,f862,f876,f880,f889,f893,f905,f909,f923,f939,f950,f960,f977,f1022,f1023,f1024,f1025,f1051,f1050,f1055,f1069,f1080,f1099,f1101,f1105,f1106,f1145,f1144,f1143,f1142,f1184,f1195,f1196,f1258,f1268,f1284,f1283,f1323,f1506,f1533,f1532,f1531,f1530,f1621,f1642,f1643,f1743,f1744,f1745,f1747,f1755,f1754,f1753,f1752,f1848,f1909,f1873,f1874,f1919,f1918,f1917,f1916,f1915,f1914,f1920,f1921,f1931,f1932,f1933,f1934,f1935,f1936,f1960,f2076,f1979,f2081,f1970,f1993,f1994,f2103,f2037,f2038,f2145,f2355,f2364,f2365,f2366,f2378,f2389,f2272,f2273,f2452,f2316,f2317,f2494,f2519,f2576,f2630,f2631,f2632,f2633,f2635,f2636,f2637,f2638,f2640,f2641,f2642,f2643,f2655,f2656,f2657,f2658,f2722,f2723,f2730,f2731,f2733,f2764,f2835,f2915,f2574,f1266,f2582,f1265,f2581,f1264,f2580,f1263,f2579,f1863,f2578,f1633,f2577,f294,f293,f3343,f3419,f3420,f3366,f3378,f3379,f3380,f3433,f3427,f3444,f3438,f3576,f3577,f3700,f3781,f3782,f3724,f3737,f3738,f3739,f3806,f3807,f3814,f3813,f3825,f3824,f3832,f3833,f3844,f3845,f3846,f3847,f3860,f3861,f3862,f3863,f3864,f3865,f3866,f3867,f2839,f3988,f4032,f4031,f4033,f4034,f4035,f4036,f4118,f4129,f4130,f4131,f4132,f4134,f4135,f4136,f4138,f4139,f4140,f4141,f4142,f4144,f4145,f4146,f4147,f4148,f4155,f4156,f4157,f4158,f4159,f4162,f4163,f4164,f4202,f4203,f4204,f4205,f4207,f4208,f4210,f4211,f4212,f4213,f4214,f4216,f4217,f4218,f4219,f4220,f4227,f4228,f4229,f4230,f4231,f4234,f4235,f4236,f4280,f4317,f4422,f4320,f4421,f4420,f4419,f4327,f4336,f4337,f4343,f4347,f4349,f4350,f4437,f4438,f4439,f4353,f4354,f4355,f4356,f4357,f4358,f4360,f4361,f4440,f4363])).
fof(f4363,plain,(
  ( ! [X17] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X17) | s(bool,'T_0') = s(bool,sK5(sK2(happ(s(fun('type.extreal.extreal','type.extreal.extreal'),'const.combin.I_0'),s('type.extreal.extreal',X17)),'type.extreal.extreal')))) ) | $spl18),
  inference(backward_demodulation,[],[f4280,f1933])).
fof(f4440,plain,(
  ( ! [X12,X13,X11] : (s(bool,'T_0') = s(bool,sK5(sK2('const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X11),s('type.extreal.extreal',X12),s(fun('type.realax.real','type.extreal.extreal'),X13)),'type.extreal.extreal')))) ) | $spl18),
  inference(global_subsumption,[],[f216,f215,f217,f219,f220,f337,f226,f227,f228,f229,f230,f231,f232,f233,f234,f235,f236,f237,f238,f239,f240,f241,f242,f243,f248,f247,f252,f251,f281,f282,f339,f290,f291,f292,f297,f298,f299,f300,f301,f302,f310,f309,f308,f307,f306,f305,f304,f303,f314,f313,f312,f311,f315,f319,f316,f321,f320,f322,f323,f324,f326,f325,f327,f328,f340,f341,f354,f368,f367,f366,f365,f404,f403,f406,f410,f409,f412,f416,f426,f428,f495,f502,f511,f512,f542,f543,f545,f597,f707,f706,f726,f756,f765,f803,f804,f805,f806,f807,f808,f812,f813,f842,f846,f858,f862,f876,f880,f889,f893,f905,f909,f923,f939,f950,f960,f977,f1022,f1023,f1024,f1025,f1051,f1050,f1055,f1069,f1080,f1099,f1101,f1105,f1106,f1145,f1144,f1143,f1142,f1184,f1195,f1196,f1258,f1268,f1284,f1283,f1323,f1506,f1533,f1532,f1531,f1530,f1621,f1642,f1643,f1743,f1744,f1745,f1747,f1755,f1754,f1753,f1752,f1848,f1909,f1873,f1874,f1919,f1918,f1917,f1916,f1915,f1914,f1920,f1921,f1931,f1932,f1933,f1934,f1935,f1936,f1960,f2076,f1979,f2081,f1970,f1993,f1994,f2103,f2037,f2038,f2145,f2355,f2364,f2365,f2366,f2378,f2389,f2272,f2273,f2452,f2316,f2317,f2494,f2519,f2576,f2630,f2631,f2632,f2633,f2635,f2636,f2637,f2638,f2640,f2641,f2642,f2643,f2655,f2656,f2657,f2658,f2722,f2723,f2730,f2731,f2733,f2764,f2835,f2915,f2574,f1266,f2582,f1265,f2581,f1264,f2580,f1263,f2579,f1863,f2578,f1633,f2577,f294,f293,f3343,f3419,f3420,f3366,f3378,f3379,f3380,f3433,f3427,f3444,f3438,f3576,f3577,f3700,f3781,f3782,f3724,f3737,f3738,f3739,f3806,f3807,f3814,f3813,f3825,f3824,f3832,f3833,f3844,f3845,f3846,f3847,f3860,f3861,f3862,f3863,f3864,f3865,f3866,f3867,f2839,f3988,f4032,f4031,f4033,f4034,f4035,f4036,f4118,f4129,f4130,f4131,f4132,f4134,f4135,f4136,f4138,f4139,f4140,f4141,f4142,f4144,f4145,f4146,f4147,f4148,f4155,f4156,f4157,f4158,f4159,f4162,f4163,f4164,f4202,f4203,f4204,f4205,f4207,f4208,f4210,f4211,f4212,f4213,f4214,f4216,f4217,f4218,f4219,f4220,f4227,f4228,f4229,f4230,f4231,f4234,f4235,f4236,f4280,f4317,f4422,f4320,f4421,f4420,f4419,f4327,f4336,f4337,f4343,f4347,f4349,f4350,f4437,f4438,f4439,f4353,f4354,f4355,f4356,f4357,f4358,f4360,f4361])).
fof(f4361,plain,(
  ( ! [X12,X13,X11] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X11) | s(bool,'T_0') = s(bool,sK5(sK2('const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X11),s('type.extreal.extreal',X12),s(fun('type.realax.real','type.extreal.extreal'),X13)),'type.extreal.extreal')))) ) | $spl18),
  inference(backward_demodulation,[],[f4280,f1931])).
fof(f4360,plain,(
  ( ! [X1] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1))) | s(bool,'T_0') = s(bool,sK5(sK2(happ(s(fun('type.realax.real','type.extreal.extreal'),'const.extreal.Normal_0'),s('type.realax.real',X1)),'type.extreal.extreal')))) ) | $spl18),
  inference(backward_demodulation,[],[f4280,f1921])).
fof(f4358,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,sK5(sK2('const.ind_type.FCONS_3'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(fun('type.num.num','type.extreal.extreal'),X0),s('type.num.num','const.num.|48|_0')),'type.extreal.extreal')))) ) | $spl18),
  inference(backward_demodulation,[],[f4280,f1918])).
fof(f4357,plain,(
  s(bool,'T_0') = s(bool,sK5(sK2('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0')),'type.extreal.extreal'))) | $spl18),
  inference(backward_demodulation,[],[f4280,f1917])).
fof(f4356,plain,(
  s(bool,'T_0') = s(bool,sK5(sK2(happ(s(fun('type.extreal.extreal','type.extreal.extreal'),'const.combin.I_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),'type.extreal.extreal'))) | $spl18),
  inference(backward_demodulation,[],[f4280,f1916])).
fof(f4355,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK5(sK2('const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0'),s(fun('type.realax.real','type.extreal.extreal'),X1)),'type.extreal.extreal')))) ) | $spl18),
  inference(backward_demodulation,[],[f4280,f1915])).
fof(f4354,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK5(sK2('const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0),s(fun('type.realax.real','type.extreal.extreal'),X1)),'type.extreal.extreal')))) ) | $spl18),
  inference(backward_demodulation,[],[f4280,f1914])).
fof(f4353,plain,(
  ( ! [X54,X52,X53] : (s(bool,'F_0') = s(bool,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(bool,X52),s(bool,X53),s(fun('type.realax.real',bool),X54))) | s(bool,'T_0') = s(bool,X53)) ) | $spl18),
  inference(backward_demodulation,[],[f4280,f1909])).
fof(f4439,plain,(
  ( ! [X12] : (s('type.realax.real','const.extreal.real_1'(s('type.extreal.extreal',X12))) = s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.min.@_1'(s(fun('type.realax.real',bool),'_dst_x0x1_2'(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal',X12))),s(fun('type.realax.real','type.extreal.extreal'),'const.extreal.Normal_0')))))))) ) | $spl18),
  inference(backward_demodulation,[],[f4438,f3860])).
fof(f4438,plain,(
  ( ! [X97,X98,X96] : (s(bool,'T_0') = s(bool,sK5('const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X96),s('type.extreal.extreal',X97),s(fun('type.realax.real','type.extreal.extreal'),X98))))) ) | $spl18),
  inference(global_subsumption,[],[f216,f215,f217,f219,f220,f337,f226,f227,f228,f229,f230,f231,f232,f233,f234,f235,f236,f237,f238,f239,f240,f241,f242,f243,f248,f247,f252,f251,f281,f282,f339,f290,f291,f292,f297,f298,f299,f300,f301,f302,f310,f309,f308,f307,f306,f305,f304,f303,f314,f313,f312,f311,f315,f319,f316,f321,f320,f322,f323,f324,f326,f325,f327,f328,f340,f341,f354,f368,f367,f366,f365,f404,f403,f406,f410,f409,f412,f416,f426,f428,f495,f502,f511,f512,f542,f543,f545,f597,f707,f706,f726,f756,f765,f803,f804,f805,f806,f807,f808,f812,f813,f842,f846,f858,f862,f876,f880,f889,f893,f905,f909,f923,f939,f950,f960,f977,f1022,f1023,f1024,f1025,f1051,f1050,f1055,f1069,f1080,f1099,f1101,f1105,f1106,f1145,f1144,f1143,f1142,f1184,f1195,f1196,f1258,f1268,f1284,f1283,f1323,f1506,f1533,f1532,f1531,f1530,f1621,f1642,f1643,f1743,f1744,f1745,f1747,f1755,f1754,f1753,f1752,f1848,f1909,f1873,f1874,f1919,f1918,f1917,f1916,f1915,f1914,f1920,f1921,f1931,f1932,f1933,f1934,f1935,f1936,f1960,f2076,f1979,f2081,f1970,f1993,f1994,f2103,f2037,f2038,f2145,f2355,f2364,f2365,f2366,f2378,f2389,f2272,f2273,f2452,f2316,f2317,f2494,f2519,f2576,f2630,f2631,f2632,f2633,f2635,f2636,f2637,f2638,f2640,f2641,f2642,f2643,f2655,f2656,f2657,f2658,f2722,f2723,f2730,f2731,f2733,f2764,f2835,f2915,f2574,f1266,f2582,f1265,f2581,f1264,f2580,f1263,f2579,f1863,f2578,f1633,f2577,f294,f293,f3343,f3419,f3420,f3366,f3378,f3379,f3380,f3433,f3427,f3444,f3438,f3576,f3577,f3700,f3781,f3782,f3724,f3737,f3738,f3739,f3806,f3807,f3814,f3813,f3825,f3824,f3832,f3833,f3844,f3845,f3846,f3847,f3860,f3861,f3862,f3863,f3864,f3865,f3866,f3867,f2839,f3988,f4032,f4031,f4033,f4034,f4035,f4036,f4118,f4129,f4130,f4131,f4132,f4134,f4135,f4136,f4138,f4139,f4140,f4141,f4142,f4144,f4145,f4146,f4147,f4148,f4155,f4156,f4157,f4158,f4159,f4162,f4163,f4164,f4202,f4203,f4204,f4205,f4207,f4208,f4210,f4211,f4212,f4213,f4214,f4216,f4217,f4218,f4219,f4220,f4227,f4228,f4229,f4230,f4231,f4234,f4235,f4236,f4280,f4317,f4422,f4320,f4421,f4420,f4419,f4327,f4336,f4337,f4343,f4347,f4349,f4350,f4437])).
fof(f4437,plain,(
  ( ! [X97,X98,X96] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X97) | s(bool,'T_0') = s(bool,sK5('const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X96),s('type.extreal.extreal',X97),s(fun('type.realax.real','type.extreal.extreal'),X98))))) ) | $spl18),
  inference(forward_demodulation,[],[f4351,f4280])).
fof(f4351,plain,(
  ( ! [X97,X98,X96] : (s(bool,'T_0') = s(bool,sK5('const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X96),s('type.extreal.extreal',X97),s(fun('type.realax.real','type.extreal.extreal'),X98)))) | s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X97)) ) | $spl18),
  inference(backward_demodulation,[],[f4280,f1873])).
fof(f4350,plain,(
  ( ! [X59,X57,X58,X56,X55] : (s(X55,X59) = s(X55,sK3(X59,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(fun(X55,bool),X56),s(fun(X55,bool),X57),s(fun('type.realax.real',fun(X55,bool)),X58)),X55)) | p(s(bool,'|dollar|exists_1'(s(fun(X55,bool),X57))))) ) | $spl18),
  inference(backward_demodulation,[],[f4280,f1863])).
fof(f4349,plain,(
  ( ! [X6,X4,X5] : (s(bool,'T_0') = s(bool,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(bool,X4),s(bool,X5),s(fun('type.realax.real',bool),X6))) | s(bool,'F_0') = s(bool,X5)) ) | $spl18),
  inference(backward_demodulation,[],[f4280,f1848])).
fof(f4347,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK5('const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0'),s(fun('type.realax.real','type.extreal.extreal'),X1))))) ) | $spl18),
  inference(backward_demodulation,[],[f4280,f1752])).
fof(f4343,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,sK5('const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0))),s('type.extreal.extreal','const.extreal.NegInf_0'))))) ) | $spl18),
  inference(backward_demodulation,[],[f4280,f1283])).
fof(f4337,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0))),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl18),
  inference(backward_demodulation,[],[f4280,f1055])).
fof(f4336,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,sK5('const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.NegInf_0'))))) ) | $spl18),
  inference(backward_demodulation,[],[f4280,f1050])).
fof(f4327,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl18),
  inference(backward_demodulation,[],[f4280,f365])).
fof(f4419,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | ~p(s(bool,sK5(X0)))) ) | $spl18),
  inference(duplicate_literal_removal,[],[f4325])).
fof(f4325,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | ~p(s(bool,sK5(X0))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0)) ) | $spl18),
  inference(backward_demodulation,[],[f4280,f311])).
fof(f4420,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X5) | s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal',X4),s(X0,X6),s(X0,X7),s(fun('type.realax.real',X0),X8))) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal',X5),s(X0,X2),s(X0,X1),s(fun('type.realax.real',X0),X3))) | s(X0,happ(s(fun('type.realax.real',X0),X3),s('type.realax.real',sK4(X8,X7,X6,X5,X4,X3,X2,X1,X0)))) != s(X0,happ(s(fun('type.realax.real',X0),X8),s('type.realax.real',sK4(X8,X7,X6,X5,X4,X3,X2,X1,X0)))) | s('type.extreal.extreal',X4) != s('type.extreal.extreal',X5)) ) | $spl18),
  inference(duplicate_literal_removal,[],[f4323])).
fof(f4323,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X5) | s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal',X4),s(X0,X6),s(X0,X7),s(fun('type.realax.real',X0),X8))) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal',X5),s(X0,X2),s(X0,X1),s(fun('type.realax.real',X0),X3))) | s(X0,happ(s(fun('type.realax.real',X0),X3),s('type.realax.real',sK4(X8,X7,X6,X5,X4,X3,X2,X1,X0)))) != s(X0,happ(s(fun('type.realax.real',X0),X8),s('type.realax.real',sK4(X8,X7,X6,X5,X4,X3,X2,X1,X0)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X5) | s('type.extreal.extreal',X4) != s('type.extreal.extreal',X5)) ) | $spl18),
  inference(backward_demodulation,[],[f4280,f307])).
fof(f4421,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X5) | s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal',X4),s(X0,X6),s(X0,X7),s(fun('type.realax.real',X0),X8))) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal',X5),s(X0,X2),s(X0,X1),s(fun('type.realax.real',X0),X3))) | s('type.extreal.extreal',X5) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK4(X8,X7,X6,X5,X4,X3,X2,X1,X0)))) | s('type.extreal.extreal',X4) != s('type.extreal.extreal',X5)) ) | $spl18),
  inference(duplicate_literal_removal,[],[f4321])).
fof(f4321,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X5) | s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal',X4),s(X0,X6),s(X0,X7),s(fun('type.realax.real',X0),X8))) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal',X5),s(X0,X2),s(X0,X1),s(fun('type.realax.real',X0),X3))) | s('type.extreal.extreal',X5) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK4(X8,X7,X6,X5,X4,X3,X2,X1,X0)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X5) | s('type.extreal.extreal',X4) != s('type.extreal.extreal',X5)) ) | $spl18),
  inference(backward_demodulation,[],[f4280,f303])).
fof(f4320,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X2) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(X0,X1),s(X0,X2),s(fun('type.realax.real',X0),X3)))) ) | $spl18),
  inference(backward_demodulation,[],[f4280,f298])).
fof(f4422,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK1(X0))))) ) | $spl18),
  inference(duplicate_literal_removal,[],[f4319])).
fof(f4319,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK1(X0)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0)) ) | $spl18),
  inference(backward_demodulation,[],[f4280,f243])).
fof(f4317,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl18),
  inference(backward_demodulation,[],[f4280,f239])).
fof(f4236,plain,(
  ( ! [X288,X285,X287,X289,X286] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',happ(s(fun(X286,'type.extreal.extreal'),X287),s(X286,happ(s(fun(X285,X286),X288),s(X285,X289))))) | s(bool,'T_0') = s(bool,sK5(sK2(happ(s(fun(X285,'type.extreal.extreal'),'_dst_x0x1_2'(s(fun(X286,'type.extreal.extreal'),X287),s(fun(X285,X286),X288))),s(X285,X289)),'type.extreal.extreal')))) )),
  inference(superposition,[],[f804,f327])).
fof(f4235,plain,(
  ( ! [X280,X282,X281,X283,X284] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',happ(s(fun(X281,'type.extreal.extreal'),X282),s(X281,happ(s(fun(X280,X281),X283),s(X280,X284))))) | s(bool,'T_0') = s(bool,sK5(happ(s(fun(X280,'type.extreal.extreal'),'_dst_x0x1_2'(s(fun(X281,'type.extreal.extreal'),X282),s(fun(X280,X281),X283))),s(X280,X284))))) )),
  inference(superposition,[],[f543,f327])).
fof(f4234,plain,(
  ( ! [X277,X279,X275,X276,X278] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',happ(s(fun(X276,'type.extreal.extreal'),X277),s(X276,happ(s(fun(X275,X276),X278),s(X275,X279))))) | s(bool,'T_0') = s(bool,sK5(happ(s(fun(X275,'type.extreal.extreal'),'_dst_x0x1_2'(s(fun(X276,'type.extreal.extreal'),X277),s(fun(X275,X276),X278))),s(X275,X279))))) )),
  inference(superposition,[],[f542,f327])).
fof(f4231,plain,(
  ( ! [X261,X263,X259,X260,X262,X264] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X264),s('type.extreal.extreal',happ(s(fun(X260,'type.extreal.extreal'),X261),s(X260,happ(s(fun(X259,X260),X262),s(X259,X263)))))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',happ(s(fun(X259,'type.extreal.extreal'),'_dst_x0x1_2'(s(fun(X260,'type.extreal.extreal'),X261),s(fun(X259,X260),X262))),s(X259,X263))),s('type.extreal.extreal',X264))))) )),
  inference(superposition,[],[f321,f327])).
fof(f4230,plain,(
  ( ! [X255,X253,X257,X254,X256,X258] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',happ(s(fun(X254,'type.extreal.extreal'),X255),s(X254,happ(s(fun(X253,X254),X256),s(X253,X257))))),s('type.extreal.extreal',X258)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X258),s('type.extreal.extreal',happ(s(fun(X253,'type.extreal.extreal'),'_dst_x0x1_2'(s(fun(X254,'type.extreal.extreal'),X255),s(fun(X253,X254),X256))),s(X253,X257))))))) )),
  inference(superposition,[],[f321,f327])).
fof(f4229,plain,(
  ( ! [X251,X249,X252,X250,X248,X247] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X252),s('type.extreal.extreal',happ(s(fun(X248,'type.extreal.extreal'),X249),s(X248,happ(s(fun(X247,X248),X250),s(X247,X251)))))))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',happ(s(fun(X247,'type.extreal.extreal'),'_dst_x0x1_2'(s(fun(X248,'type.extreal.extreal'),X249),s(fun(X247,X248),X250))),s(X247,X251))),s('type.extreal.extreal',X252))))) )),
  inference(superposition,[],[f320,f327])).
fof(f4228,plain,(
  ( ! [X243,X241,X246,X244,X242,X245] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',happ(s(fun(X242,'type.extreal.extreal'),X243),s(X242,happ(s(fun(X241,X242),X244),s(X241,X245))))),s('type.extreal.extreal',X246)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X246),s('type.extreal.extreal',happ(s(fun(X241,'type.extreal.extreal'),'_dst_x0x1_2'(s(fun(X242,'type.extreal.extreal'),X243),s(fun(X241,X242),X244))),s(X241,X245))))))) )),
  inference(superposition,[],[f320,f327])).
fof(f4227,plain,(
  ( ! [X239,X237,X240,X238,X236] : (s('type.realax.real','const.extreal.real_1'(s('type.extreal.extreal',happ(s(fun(X237,'type.extreal.extreal'),X238),s(X237,happ(s(fun(X236,X237),X239),s(X236,X240))))))) = s('type.realax.real','const.bool.COND_3'(s(bool,sK5(happ(s(fun(X236,'type.extreal.extreal'),'_dst_x0x1_2'(s(fun(X237,'type.extreal.extreal'),X238),s(fun(X236,X237),X239))),s(X236,X240)))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.min.@_1'(s(fun('type.realax.real',bool),'_dst_x0x1_2'(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal',happ(s(fun(X237,'type.extreal.extreal'),X238),s(X237,happ(s(fun(X236,X237),X239),s(X236,X240))))))),s(fun('type.realax.real','type.extreal.extreal'),'const.extreal.Normal_0')))))))) )),
  inference(superposition,[],[f314,f327])).
fof(f4220,plain,(
  ( ! [X189,X187,X185,X190,X188,X186,X184] : (~p(s(bool,happ(s(fun(X184,bool),happ(s(fun(X186,fun(X184,bool)),X187),s(X186,happ(s(fun(X185,X186),X188),s(X185,X189))))),s(X184,X190)))) | p(s(bool,happ(s(fun(X184,bool),happ(s(fun(X185,fun(X184,bool)),'_dst_x0x1_2'(s(fun(X186,fun(X184,bool)),X187),s(fun(X185,X186),X188))),s(X185,X189))),s(X184,'const.min.@_1'(s(fun(X184,bool),happ(s(fun(X185,fun(X184,bool)),'_dst_x0x1_2'(s(fun(X186,fun(X184,bool)),X187),s(fun(X185,X186),X188))),s(X185,X189))))))))) )),
  inference(superposition,[],[f292,f327])).
fof(f4219,plain,(
  ( ! [X177,X182,X180,X178,X183,X181,X179] : (p(s(bool,'|dollar|exists_1'(s(fun(X177,bool),happ(s(fun(X179,fun(X177,bool)),X180),s(X179,happ(s(fun(X178,X179),X181),s(X178,X182)))))))) | s(X177,X183) = s(X177,sK3(X183,happ(s(fun(X178,fun(X177,bool)),'_dst_x0x1_2'(s(fun(X179,fun(X177,bool)),X180),s(fun(X178,X179),X181))),s(X178,X182)),X177))) )),
  inference(superposition,[],[f293,f327])).
fof(f4218,plain,(
  ( ! [X175,X173,X171,X176,X174,X172,X170] : (~p(s(bool,happ(s(fun(X170,bool),happ(s(fun(X172,fun(X170,bool)),X173),s(X172,happ(s(fun(X171,X172),X174),s(X171,X175))))),s(X170,sK3(X176,happ(s(fun(X171,fun(X170,bool)),'_dst_x0x1_2'(s(fun(X172,fun(X170,bool)),X173),s(fun(X171,X172),X174))),s(X171,X175)),X170))))) | p(s(bool,'|dollar|exists_1'(s(fun(X170,bool),happ(s(fun(X171,fun(X170,bool)),'_dst_x0x1_2'(s(fun(X172,fun(X170,bool)),X173),s(fun(X171,X172),X174))),s(X171,X175))))))) )),
  inference(superposition,[],[f294,f327])).
fof(f4217,plain,(
  ( ! [X167,X165,X169,X166,X168] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,happ(s(fun(X166,bool),X167),s(X166,happ(s(fun(X165,X166),X168),s(X165,X169))))))) | s(bool,'F_0') = s(bool,happ(s(fun(X165,bool),'_dst_x0x1_2'(s(fun(X166,bool),X167),s(fun(X165,X166),X168))),s(X165,X169)))) )),
  inference(superposition,[],[f512,f327])).
fof(f4216,plain,(
  ( ! [X163,X161,X159,X164,X162,X160] : (~p(s(bool,happ(s(fun(X160,bool),X161),s(X160,happ(s(fun(X159,X160),X162),s(X159,X163)))))) | ~p(s(bool,X164)) | s(bool,X164) = s(bool,happ(s(fun(X159,bool),'_dst_x0x1_2'(s(fun(X160,bool),X161),s(fun(X159,X160),X162))),s(X159,X163)))) )),
  inference(superposition,[],[f319,f327])).
fof(f4214,plain,(
  ( ! [X152,X151,X149,X147,X150,X148] : (p(s(bool,happ(s(fun(X148,bool),X149),s(X148,happ(s(fun(X147,X148),X150),s(X147,X151)))))) | p(s(bool,X152)) | s(bool,X152) = s(bool,happ(s(fun(X147,bool),'_dst_x0x1_2'(s(fun(X148,bool),X149),s(fun(X147,X148),X150))),s(X147,X151)))) )),
  inference(superposition,[],[f316,f327])).
fof(f4213,plain,(
  ( ! [X146,X144,X142,X145,X143] : (p(s(bool,happ(s(fun(X143,bool),X144),s(X143,happ(s(fun(X142,X143),X145),s(X142,X146)))))) | s(bool,'F_0') = s(bool,happ(s(fun(X142,bool),'_dst_x0x1_2'(s(fun(X143,bool),X144),s(fun(X142,X143),X145))),s(X142,X146)))) )),
  inference(superposition,[],[f252,f327])).
fof(f4212,plain,(
  ( ! [X140,X138,X141,X139,X137] : (~p(s(bool,happ(s(fun(X138,bool),X139),s(X138,happ(s(fun(X137,X138),X140),s(X137,X141)))))) | s(bool,'F_0') != s(bool,happ(s(fun(X137,bool),'_dst_x0x1_2'(s(fun(X138,bool),X139),s(fun(X137,X138),X140))),s(X137,X141)))) )),
  inference(superposition,[],[f251,f327])).
fof(f4211,plain,(
  ( ! [X134,X132,X136,X135,X133] : (~p(s(bool,happ(s(fun(X133,bool),X134),s(X133,happ(s(fun(X132,X133),X135),s(X132,X136)))))) | s(bool,'T_0') = s(bool,happ(s(fun(X132,bool),'_dst_x0x1_2'(s(fun(X133,bool),X134),s(fun(X132,X133),X135))),s(X132,X136)))) )),
  inference(superposition,[],[f248,f327])).
fof(f4210,plain,(
  ( ! [X127,X130,X128,X131,X129] : (p(s(bool,happ(s(fun(X128,bool),X129),s(X128,happ(s(fun(X127,X128),X130),s(X127,X131)))))) | s(bool,'T_0') != s(bool,happ(s(fun(X127,bool),'_dst_x0x1_2'(s(fun(X128,bool),X129),s(fun(X127,X128),X130))),s(X127,X131)))) )),
  inference(superposition,[],[f247,f327])).
fof(f4208,plain,(
  ( ! [X121,X118,X120,X119,X117] : (~p(s(bool,happ(s(fun(X118,bool),X119),s(X118,happ(s(fun(X117,X118),X120),s(X117,sK3(X121,'_dst_x0x1_2'(s(fun(X118,bool),X119),s(fun(X117,X118),X120)),X117))))))) | p(s(bool,'|dollar|exists_1'(s(fun(X117,bool),'_dst_x0x1_2'(s(fun(X118,bool),X119),s(fun(X117,X118),X120))))))) )),
  inference(superposition,[],[f294,f327])).
fof(f4207,plain,(
  ( ! [X116,X114,X112,X115,X113] : (~p(s(bool,happ(s(fun(X113,bool),X114),s(X113,happ(s(fun(X112,X113),X115),s(X112,X116)))))) | p(s(bool,happ(s(fun(X112,bool),'_dst_x0x1_2'(s(fun(X113,bool),X114),s(fun(X112,X113),X115))),s(X112,'const.min.@_1'(s(fun(X112,bool),'_dst_x0x1_2'(s(fun(X113,bool),X114),s(fun(X112,X113),X115))))))))) )),
  inference(superposition,[],[f292,f327])).
fof(f4205,plain,(
  ( ! [X103,X101,X99,X105,X102,X100,X104] : (p(s(bool,happ(s(fun(X99,bool),'=_1'(s(X99,X105))),s(X99,happ(s(fun(X101,X99),X102),s(X101,happ(s(fun(X100,X101),X103),s(X100,X104)))))))) | s(X99,X105) != s(X99,happ(s(fun(X100,X99),'_dst_x0x1_2'(s(fun(X101,X99),X102),s(fun(X100,X101),X103))),s(X100,X104)))) )),
  inference(superposition,[],[f326,f327])).
fof(f4204,plain,(
  ( ! [X94,X92,X97,X95,X93,X98,X96] : (p(s(bool,happ(s(fun(X92,bool),'=_1'(s(X92,happ(s(fun(X94,X92),X95),s(X94,happ(s(fun(X93,X94),X96),s(X93,X97))))))),s(X92,X98)))) | s(X92,X98) != s(X92,happ(s(fun(X93,X92),'_dst_x0x1_2'(s(fun(X94,X92),X95),s(fun(X93,X94),X96))),s(X93,X97)))) )),
  inference(superposition,[],[f326,f327])).
fof(f4203,plain,(
  ( ! [X90,X88,X87,X85,X91,X89,X86] : (~p(s(bool,happ(s(fun(X85,bool),'=_1'(s(X85,X91))),s(X85,happ(s(fun(X87,X85),X88),s(X87,happ(s(fun(X86,X87),X89),s(X86,X90)))))))) | s(X85,X91) = s(X85,happ(s(fun(X86,X85),'_dst_x0x1_2'(s(fun(X87,X85),X88),s(fun(X86,X87),X89))),s(X86,X90)))) )),
  inference(superposition,[],[f325,f327])).
fof(f4202,plain,(
  ( ! [X80,X78,X83,X81,X79,X84,X82] : (~p(s(bool,happ(s(fun(X78,bool),'=_1'(s(X78,happ(s(fun(X80,X78),X81),s(X80,happ(s(fun(X79,X80),X82),s(X79,X83))))))),s(X78,X84)))) | s(X78,X84) = s(X78,happ(s(fun(X79,X78),'_dst_x0x1_2'(s(fun(X80,X78),X81),s(fun(X79,X80),X82))),s(X79,X83)))) )),
  inference(superposition,[],[f325,f327])).
fof(f4164,plain,(
  ( ! [X300,X298,X301,X297,X299] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',happ(s(fun(X299,'type.extreal.extreal'),'_dst_x0x1_2'(s(fun(X297,'type.extreal.extreal'),X298),s(fun(X299,X297),X300))),s(X299,X301))) | s(bool,'T_0') = s(bool,sK5(sK2(happ(s(fun(X297,'type.extreal.extreal'),X298),s(X297,happ(s(fun(X299,X297),X300),s(X299,X301)))),'type.extreal.extreal')))) )),
  inference(superposition,[],[f804,f327])).
fof(f4163,plain,(
  ( ! [X292,X294,X296,X293,X295] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',happ(s(fun(X294,'type.extreal.extreal'),'_dst_x0x1_2'(s(fun(X292,'type.extreal.extreal'),X293),s(fun(X294,X292),X295))),s(X294,X296))) | s(bool,'T_0') = s(bool,sK5(happ(s(fun(X292,'type.extreal.extreal'),X293),s(X292,happ(s(fun(X294,X292),X295),s(X294,X296))))))) )),
  inference(superposition,[],[f543,f327])).
fof(f4162,plain,(
  ( ! [X288,X290,X287,X289,X291] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',happ(s(fun(X289,'type.extreal.extreal'),'_dst_x0x1_2'(s(fun(X287,'type.extreal.extreal'),X288),s(fun(X289,X287),X290))),s(X289,X291))) | s(bool,'T_0') = s(bool,sK5(happ(s(fun(X287,'type.extreal.extreal'),X288),s(X287,happ(s(fun(X289,X287),X290),s(X289,X291))))))) )),
  inference(superposition,[],[f542,f327])).
fof(f4159,plain,(
  ( ! [X271,X273,X275,X276,X272,X274] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X276),s('type.extreal.extreal',happ(s(fun(X273,'type.extreal.extreal'),'_dst_x0x1_2'(s(fun(X271,'type.extreal.extreal'),X272),s(fun(X273,X271),X274))),s(X273,X275)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',happ(s(fun(X271,'type.extreal.extreal'),X272),s(X271,happ(s(fun(X273,X271),X274),s(X273,X275))))),s('type.extreal.extreal',X276))))) )),
  inference(superposition,[],[f321,f327])).
fof(f4158,plain,(
  ( ! [X269,X265,X267,X268,X270,X266] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',happ(s(fun(X267,'type.extreal.extreal'),'_dst_x0x1_2'(s(fun(X265,'type.extreal.extreal'),X266),s(fun(X267,X265),X268))),s(X267,X269))),s('type.extreal.extreal',X270)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X270),s('type.extreal.extreal',happ(s(fun(X265,'type.extreal.extreal'),X266),s(X265,happ(s(fun(X267,X265),X268),s(X267,X269))))))))) )),
  inference(superposition,[],[f321,f327])).
fof(f4157,plain,(
  ( ! [X261,X263,X259,X260,X262,X264] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X264),s('type.extreal.extreal',happ(s(fun(X261,'type.extreal.extreal'),'_dst_x0x1_2'(s(fun(X259,'type.extreal.extreal'),X260),s(fun(X261,X259),X262))),s(X261,X263)))))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',happ(s(fun(X259,'type.extreal.extreal'),X260),s(X259,happ(s(fun(X261,X259),X262),s(X261,X263))))),s('type.extreal.extreal',X264))))) )),
  inference(superposition,[],[f320,f327])).
fof(f4156,plain,(
  ( ! [X255,X253,X257,X254,X256,X258] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',happ(s(fun(X255,'type.extreal.extreal'),'_dst_x0x1_2'(s(fun(X253,'type.extreal.extreal'),X254),s(fun(X255,X253),X256))),s(X255,X257))),s('type.extreal.extreal',X258)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X258),s('type.extreal.extreal',happ(s(fun(X253,'type.extreal.extreal'),X254),s(X253,happ(s(fun(X255,X253),X256),s(X255,X257))))))))) )),
  inference(superposition,[],[f320,f327])).
fof(f4155,plain,(
  ( ! [X251,X249,X252,X250,X248] : (s('type.realax.real','const.extreal.real_1'(s('type.extreal.extreal',happ(s(fun(X250,'type.extreal.extreal'),'_dst_x0x1_2'(s(fun(X248,'type.extreal.extreal'),X249),s(fun(X250,X248),X251))),s(X250,X252))))) = s('type.realax.real','const.bool.COND_3'(s(bool,sK5(happ(s(fun(X248,'type.extreal.extreal'),X249),s(X248,happ(s(fun(X250,X248),X251),s(X250,X252)))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.min.@_1'(s(fun('type.realax.real',bool),'_dst_x0x1_2'(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal',happ(s(fun(X250,'type.extreal.extreal'),'_dst_x0x1_2'(s(fun(X248,'type.extreal.extreal'),X249),s(fun(X250,X248),X251))),s(X250,X252))))),s(fun('type.realax.real','type.extreal.extreal'),'const.extreal.Normal_0')))))))) )),
  inference(superposition,[],[f314,f327])).
fof(f4148,plain,(
  ( ! [X198,X196,X194,X199,X197,X195,X193] : (~p(s(bool,happ(s(fun(X193,bool),happ(s(fun(X196,fun(X193,bool)),'_dst_x0x1_2'(s(fun(X194,fun(X193,bool)),X195),s(fun(X196,X194),X197))),s(X196,X198))),s(X193,X199)))) | p(s(bool,happ(s(fun(X193,bool),happ(s(fun(X194,fun(X193,bool)),X195),s(X194,happ(s(fun(X196,X194),X197),s(X196,X198))))),s(X193,'const.min.@_1'(s(fun(X193,bool),happ(s(fun(X194,fun(X193,bool)),X195),s(X194,happ(s(fun(X196,X194),X197),s(X196,X198))))))))))) )),
  inference(superposition,[],[f292,f327])).
fof(f4147,plain,(
  ( ! [X191,X189,X187,X192,X190,X188,X186] : (p(s(bool,'|dollar|exists_1'(s(fun(X186,bool),happ(s(fun(X189,fun(X186,bool)),'_dst_x0x1_2'(s(fun(X187,fun(X186,bool)),X188),s(fun(X189,X187),X190))),s(X189,X191)))))) | s(X186,X192) = s(X186,sK3(X192,happ(s(fun(X187,fun(X186,bool)),X188),s(X187,happ(s(fun(X189,X187),X190),s(X189,X191)))),X186))) )),
  inference(superposition,[],[f293,f327])).
fof(f4146,plain,(
  ( ! [X185,X182,X180,X184,X183,X181,X179] : (~p(s(bool,happ(s(fun(X179,bool),happ(s(fun(X182,fun(X179,bool)),'_dst_x0x1_2'(s(fun(X180,fun(X179,bool)),X181),s(fun(X182,X180),X183))),s(X182,X184))),s(X179,sK3(X185,happ(s(fun(X180,fun(X179,bool)),X181),s(X180,happ(s(fun(X182,X180),X183),s(X182,X184)))),X179))))) | p(s(bool,'|dollar|exists_1'(s(fun(X179,bool),happ(s(fun(X180,fun(X179,bool)),X181),s(X180,happ(s(fun(X182,X180),X183),s(X182,X184))))))))) )),
  inference(superposition,[],[f294,f327])).
fof(f4145,plain,(
  ( ! [X177,X175,X178,X176,X174] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,happ(s(fun(X176,bool),'_dst_x0x1_2'(s(fun(X174,bool),X175),s(fun(X176,X174),X177))),s(X176,X178))))) | s(bool,'F_0') = s(bool,happ(s(fun(X174,bool),X175),s(X174,happ(s(fun(X176,X174),X177),s(X176,X178)))))) )),
  inference(superposition,[],[f512,f327])).
fof(f4144,plain,(
  ( ! [X173,X171,X169,X172,X170,X168] : (~p(s(bool,happ(s(fun(X170,bool),'_dst_x0x1_2'(s(fun(X168,bool),X169),s(fun(X170,X168),X171))),s(X170,X172)))) | ~p(s(bool,X173)) | s(bool,X173) = s(bool,happ(s(fun(X168,bool),X169),s(X168,happ(s(fun(X170,X168),X171),s(X170,X172)))))) )),
  inference(superposition,[],[f319,f327])).
fof(f4142,plain,(
  ( ! [X158,X156,X161,X159,X157,X160] : (p(s(bool,happ(s(fun(X158,bool),'_dst_x0x1_2'(s(fun(X156,bool),X157),s(fun(X158,X156),X159))),s(X158,X160)))) | p(s(bool,X161)) | s(bool,X161) = s(bool,happ(s(fun(X156,bool),X157),s(X156,happ(s(fun(X158,X156),X159),s(X158,X160)))))) )),
  inference(superposition,[],[f316,f327])).
fof(f4141,plain,(
  ( ! [X154,X152,X151,X155,X153] : (p(s(bool,happ(s(fun(X153,bool),'_dst_x0x1_2'(s(fun(X151,bool),X152),s(fun(X153,X151),X154))),s(X153,X155)))) | s(bool,'F_0') = s(bool,happ(s(fun(X151,bool),X152),s(X151,happ(s(fun(X153,X151),X154),s(X153,X155)))))) )),
  inference(superposition,[],[f252,f327])).
fof(f4140,plain,(
  ( ! [X146,X149,X147,X150,X148] : (~p(s(bool,happ(s(fun(X148,bool),'_dst_x0x1_2'(s(fun(X146,bool),X147),s(fun(X148,X146),X149))),s(X148,X150)))) | s(bool,'F_0') != s(bool,happ(s(fun(X146,bool),X147),s(X146,happ(s(fun(X148,X146),X149),s(X148,X150)))))) )),
  inference(superposition,[],[f251,f327])).
fof(f4139,plain,(
  ( ! [X144,X142,X145,X143,X141] : (~p(s(bool,happ(s(fun(X143,bool),'_dst_x0x1_2'(s(fun(X141,bool),X142),s(fun(X143,X141),X144))),s(X143,X145)))) | s(bool,'T_0') = s(bool,happ(s(fun(X141,bool),X142),s(X141,happ(s(fun(X143,X141),X144),s(X143,X145)))))) )),
  inference(superposition,[],[f248,f327])).
fof(f4138,plain,(
  ( ! [X140,X138,X136,X139,X137] : (p(s(bool,happ(s(fun(X138,bool),'_dst_x0x1_2'(s(fun(X136,bool),X137),s(fun(X138,X136),X139))),s(X138,X140)))) | s(bool,'T_0') != s(bool,happ(s(fun(X136,bool),X137),s(X136,happ(s(fun(X138,X136),X139),s(X138,X140)))))) )),
  inference(superposition,[],[f247,f327])).
fof(f4136,plain,(
  ( ! [X127,X130,X128,X126,X129] : (~p(s(bool,happ(s(fun(X128,bool),'_dst_x0x1_2'(s(fun(X126,bool),X127),s(fun(X128,X126),X129))),s(X128,X130)))) | p(s(bool,happ(s(fun(X126,bool),X127),s(X126,'const.min.@_1'(s(fun(X126,bool),X127))))))) )),
  inference(superposition,[],[f292,f327])).
fof(f4135,plain,(
  ( ! [X125,X123,X121,X124,X122] : (~p(s(bool,happ(s(fun(X123,bool),'_dst_x0x1_2'(s(fun(X121,bool),'=_1'(s(X121,X122))),s(fun(X123,X121),X124))),s(X123,X125)))) | s(X121,X122) = s(X121,happ(s(fun(X123,X121),X124),s(X123,X125)))) )),
  inference(superposition,[],[f325,f327])).
fof(f4134,plain,(
  ( ! [X118,X116,X120,X119,X117] : (p(s(bool,happ(s(fun(X118,bool),'_dst_x0x1_2'(s(fun(X116,bool),'=_1'(s(X116,X117))),s(fun(X118,X116),X119))),s(X118,X120)))) | s(X116,X117) != s(X116,happ(s(fun(X118,X116),X119),s(X118,X120)))) )),
  inference(superposition,[],[f326,f327])).
fof(f4132,plain,(
  ( ! [X103,X109,X107,X105,X108,X106,X104] : (p(s(bool,happ(s(fun(X103,bool),'=_1'(s(X103,X109))),s(X103,happ(s(fun(X106,X103),'_dst_x0x1_2'(s(fun(X104,X103),X105),s(fun(X106,X104),X107))),s(X106,X108)))))) | s(X103,X109) != s(X103,happ(s(fun(X104,X103),X105),s(X104,happ(s(fun(X106,X104),X107),s(X106,X108)))))) )),
  inference(superposition,[],[f326,f327])).
fof(f4131,plain,(
  ( ! [X101,X99,X97,X102,X100,X98,X96] : (p(s(bool,happ(s(fun(X96,bool),'=_1'(s(X96,happ(s(fun(X99,X96),'_dst_x0x1_2'(s(fun(X97,X96),X98),s(fun(X99,X97),X100))),s(X99,X101))))),s(X96,X102)))) | s(X96,X102) != s(X96,happ(s(fun(X97,X96),X98),s(X97,happ(s(fun(X99,X97),X100),s(X99,X101)))))) )),
  inference(superposition,[],[f326,f327])).
fof(f4130,plain,(
  ( ! [X94,X92,X90,X95,X93,X91,X89] : (~p(s(bool,happ(s(fun(X89,bool),'=_1'(s(X89,X95))),s(X89,happ(s(fun(X92,X89),'_dst_x0x1_2'(s(fun(X90,X89),X91),s(fun(X92,X90),X93))),s(X92,X94)))))) | s(X89,X95) = s(X89,happ(s(fun(X90,X89),X91),s(X90,happ(s(fun(X92,X90),X93),s(X92,X94)))))) )),
  inference(superposition,[],[f325,f327])).
fof(f4129,plain,(
  ( ! [X88,X87,X85,X83,X86,X84,X82] : (~p(s(bool,happ(s(fun(X82,bool),'=_1'(s(X82,happ(s(fun(X85,X82),'_dst_x0x1_2'(s(fun(X83,X82),X84),s(fun(X85,X83),X86))),s(X85,X87))))),s(X82,X88)))) | s(X82,X88) = s(X82,happ(s(fun(X83,X82),X84),s(X83,happ(s(fun(X85,X83),X86),s(X85,X87)))))) )),
  inference(superposition,[],[f325,f327])).
fof(f4118,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s(bool,'F_0') = s(bool,happ(s(fun(X15,bool),X16),s(X15,happ(s(fun(X14,X15),X17),s(X14,X18))))) | s(bool,'T_0') = s(bool,happ(s(fun(X14,bool),'_dst_x0x1_2'(s(fun(X15,bool),X16),s(fun(X14,X15),X17))),s(X14,X18)))) )),
  inference(superposition,[],[f327,f242])).
fof(f4036,plain,(
  ( ! [X21,X19,X22,X20,X18] : (s(bool,'F_0') = s(bool,happ(s(fun(X20,bool),'_dst_x0x1_2'(s(fun(X18,bool),X19),s(fun(X20,X18),X21))),s(X20,X22))) | s(bool,'T_0') = s(bool,happ(s(fun(X18,bool),X19),s(X18,happ(s(fun(X20,X18),X21),s(X20,X22)))))) )),
  inference(superposition,[],[f327,f242])).
fof(f4035,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X16] : (s(X10,happ(s(fun(X12,X10),X13),s(X12,happ(s(fun(X11,X12),X14),s(X11,happ(s(fun(X15,X11),X16),s(X15,X17))))))) = s(X10,happ(s(fun(X15,X10),'_dst_x0x1_2'(s(fun(X11,X10),'_dst_x0x1_2'(s(fun(X12,X10),X13),s(fun(X11,X12),X14))),s(fun(X15,X11),X16))),s(X15,X17)))) )),
  inference(superposition,[],[f327,f327])).
fof(f4034,plain,(
  ( ! [X14,X12,X19,X17,X15,X13,X18,X16] : (s(X18,happ(s(fun(X12,X18),X19),s(X12,happ(s(fun(X14,X12),X15),s(X14,happ(s(fun(X13,X14),X16),s(X13,X17))))))) = s(X18,happ(s(fun(X13,X18),'_dst_x0x1_2'(s(fun(X12,X18),X19),s(fun(X13,X12),'_dst_x0x1_2'(s(fun(X14,X12),X15),s(fun(X13,X14),X16))))),s(X13,X17)))) )),
  inference(superposition,[],[f327,f327])).
fof(f4033,plain,(
  ( ! [X6,X4,X10,X8,X7,X5,X11,X9] : (s(X10,happ(s(fun(X4,X10),X11),s(X4,happ(s(fun(X7,X4),'_dst_x0x1_2'(s(fun(X5,X4),X6),s(fun(X7,X5),X8))),s(X7,X9))))) = s(X10,happ(s(fun(X5,X10),'_dst_x0x1_2'(s(fun(X4,X10),X11),s(fun(X5,X4),X6))),s(X5,happ(s(fun(X7,X5),X8),s(X7,X9)))))) )),
  inference(superposition,[],[f327,f327])).
fof(f4031,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X3,X1),X4),s(X3,X5))))))),s(X0,happ(s(fun(X3,X0),'_dst_x0x1_2'(s(fun(X1,X0),X2),s(fun(X3,X1),X4))),s(X3,X5))))))) )),
  inference(unit_resulting_resolution,[],[f327,f326])).
fof(f4032,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,happ(s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X3),s(fun(X1,X2),X4))),s(X1,X5))))),s(X0,happ(s(fun(X2,X0),X3),s(X2,happ(s(fun(X1,X2),X4),s(X1,X5))))))))) )),
  inference(unit_resulting_resolution,[],[f327,f326])).
fof(f3867,plain,(
  ( ! [X29] : (s('type.realax.real','const.extreal.real_1'(s('type.extreal.extreal',X29))) = s('type.realax.real','const.bool.COND_3'(s(bool,sK5(sK2(X29,'type.extreal.extreal'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.min.@_1'(s(fun('type.realax.real',bool),'_dst_x0x1_2'(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal',X29))),s(fun('type.realax.real','type.extreal.extreal'),'const.extreal.Normal_0')))))))) )),
  inference(superposition,[],[f314,f290])).
fof(f3866,plain,(
  ( ! [X28,X26,X27] : (s('type.realax.real','const.extreal.real_1'(s('type.extreal.extreal',happ(s(fun('type.num.num','type.extreal.extreal'),X27),s('type.num.num',X28))))) = s('type.realax.real','const.bool.COND_3'(s(bool,sK5('const.ind_type.FCONS_3'(s('type.extreal.extreal',X26),s(fun('type.num.num','type.extreal.extreal'),X27),s('type.num.num','const.num.SUC_1'(s('type.num.num',X28)))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.min.@_1'(s(fun('type.realax.real',bool),'_dst_x0x1_2'(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal',happ(s(fun('type.num.num','type.extreal.extreal'),X27),s('type.num.num',X28))))),s(fun('type.realax.real','type.extreal.extreal'),'const.extreal.Normal_0')))))))) )),
  inference(superposition,[],[f314,f299])).
fof(f3865,plain,(
  ( ! [X24,X25] : (s('type.realax.real','const.extreal.real_1'(s('type.extreal.extreal',X24))) = s('type.realax.real','const.bool.COND_3'(s(bool,sK5('const.ind_type.FCONS_3'(s('type.extreal.extreal',X24),s(fun('type.num.num','type.extreal.extreal'),X25),s('type.num.num','const.num.|48|_0')))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.min.@_1'(s(fun('type.realax.real',bool),'_dst_x0x1_2'(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal',X24))),s(fun('type.realax.real','type.extreal.extreal'),'const.extreal.Normal_0')))))))) )),
  inference(superposition,[],[f314,f291])).
fof(f3864,plain,(
  ( ! [X23] : (s('type.realax.real','const.extreal.real_1'(s('type.extreal.extreal',X23))) = s('type.realax.real','const.bool.COND_3'(s(bool,sK5('const.combin.I_1'(s('type.extreal.extreal',X23)))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.min.@_1'(s(fun('type.realax.real',bool),'_dst_x0x1_2'(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal',X23))),s(fun('type.realax.real','type.extreal.extreal'),'const.extreal.Normal_0')))))))) )),
  inference(superposition,[],[f314,f282])).
fof(f3863,plain,(
  ( ! [X22] : (s('type.realax.real','const.extreal.real_1'(s('type.extreal.extreal',X22))) = s('type.realax.real','const.bool.COND_3'(s(bool,sK5(happ(s(fun('type.extreal.extreal','type.extreal.extreal'),'const.combin.I_0'),s('type.extreal.extreal',X22)))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.min.@_1'(s(fun('type.realax.real',bool),'_dst_x0x1_2'(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal',X22))),s(fun('type.realax.real','type.extreal.extreal'),'const.extreal.Normal_0')))))))) )),
  inference(superposition,[],[f314,f339])).
fof(f3862,plain,(
  ( ! [X21,X19,X20,X18] : (s('type.realax.real','const.extreal.real_1'(s('type.extreal.extreal',happ(s(fun('type.realax.real','type.extreal.extreal'),X21),s('type.realax.real',X18))))) = s('type.realax.real','const.bool.COND_3'(s(bool,sK5('const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X18))),s('type.extreal.extreal',X19),s('type.extreal.extreal',X20),s(fun('type.realax.real','type.extreal.extreal'),X21)))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.min.@_1'(s(fun('type.realax.real',bool),'_dst_x0x1_2'(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal',happ(s(fun('type.realax.real','type.extreal.extreal'),X21),s('type.realax.real',X18))))),s(fun('type.realax.real','type.extreal.extreal'),'const.extreal.Normal_0')))))))) )),
  inference(superposition,[],[f314,f302])).
fof(f3861,plain,(
  ( ! [X17,X15,X16] : (s('type.realax.real','const.extreal.real_1'(s('type.extreal.extreal',X16))) = s('type.realax.real','const.bool.COND_3'(s(bool,sK5('const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X15),s('type.extreal.extreal',X16),s(fun('type.realax.real','type.extreal.extreal'),X17)))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.min.@_1'(s(fun('type.realax.real',bool),'_dst_x0x1_2'(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal',X16))),s(fun('type.realax.real','type.extreal.extreal'),'const.extreal.Normal_0')))))))) )),
  inference(superposition,[],[f314,f298])).
fof(f3860,plain,(
  ( ! [X14,X12,X13] : (s('type.realax.real','const.extreal.real_1'(s('type.extreal.extreal',X12))) = s('type.realax.real','const.bool.COND_3'(s(bool,sK5('const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X12),s('type.extreal.extreal',X13),s(fun('type.realax.real','type.extreal.extreal'),X14)))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.min.@_1'(s(fun('type.realax.real',bool),'_dst_x0x1_2'(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal',X12))),s(fun('type.realax.real','type.extreal.extreal'),'const.extreal.Normal_0')))))))) )),
  inference(superposition,[],[f314,f297])).
fof(f3847,plain,(
  ( ! [X2] : (s('type.realax.real','const.extreal.real_1'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))))) = s('type.realax.real','const.bool.COND_3'(s(bool,sK5(happ(s(fun('type.realax.real','type.extreal.extreal'),'const.extreal.Normal_0'),s('type.realax.real',X2)))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.min.@_1'(s(fun('type.realax.real',bool),'_dst_x0x1_2'(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))))),s(fun('type.realax.real','type.extreal.extreal'),'const.extreal.Normal_0')))))))) )),
  inference(superposition,[],[f314,f241])).
fof(f3846,plain,(
  ( ! [X1] : (s('type.realax.real','const.extreal.real_1'(s('type.extreal.extreal',X1))) = s('type.realax.real','const.bool.COND_3'(s(bool,sK5('const.extreal.Normal_1'(s('type.realax.real',sK1(X1))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.min.@_1'(s(fun('type.realax.real',bool),'_dst_x0x1_2'(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal',X1))),s(fun('type.realax.real','type.extreal.extreal'),'const.extreal.Normal_0'))))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X1) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X1)) )),
  inference(superposition,[],[f314,f243])).
fof(f3845,plain,(
  ( ! [X0] : (s('type.realax.real','const.extreal.real_1'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0))))) = s('type.realax.real','const.bool.COND_3'(s(bool,sK5('const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0)))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.min.@_1'(s(fun('type.realax.real',bool),'_dst_x0x1_2'(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0))))),s(fun('type.realax.real','type.extreal.extreal'),'const.extreal.Normal_0')))))))) )),
  inference(superposition,[],[f314,f233])).
fof(f3844,plain,(
  ( ! [X1] : (s('type.realax.real','const.extreal.real_1'(s('type.extreal.extreal',X1))) = s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.min.@_1'(s(fun('type.realax.real',bool),'_dst_x0x1_2'(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal',X1))),s(fun('type.realax.real','type.extreal.extreal'),'const.extreal.Normal_0'))))))) | s(bool,'T_0') = s(bool,sK5(X1))) )),
  inference(superposition,[],[f314,f242])).
fof(f3833,plain,(
  s('type.realax.real','const.extreal.real_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))) = s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.min.@_1'(s(fun('type.realax.real',bool),'_dst_x0x1_2'(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))),s(fun('type.realax.real','type.extreal.extreal'),'const.extreal.Normal_0')))))))),
  inference(superposition,[],[f314,f807])).
fof(f3832,plain,(
  s('type.realax.real','const.extreal.real_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))) = s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.min.@_1'(s(fun('type.realax.real',bool),'_dst_x0x1_2'(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),s(fun('type.realax.real','type.extreal.extreal'),'const.extreal.Normal_0')))))))),
  inference(superposition,[],[f314,f1025])).
fof(f3824,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'=_1'(s(bool,'T_0'))),s(bool,'F_0')))),
  inference(unit_resulting_resolution,[],[f2081,f252])).
fof(f3825,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'=_1'(s(bool,'T_0'))),s(bool,'F_0')))),
  inference(unit_resulting_resolution,[],[f2081,f247])).
fof(f3813,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'=_1'(s(bool,'F_0'))),s(bool,'T_0')))),
  inference(unit_resulting_resolution,[],[f2076,f252])).
fof(f3814,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'=_1'(s(bool,'F_0'))),s(bool,'T_0')))),
  inference(unit_resulting_resolution,[],[f2076,f247])).
fof(f3807,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK7(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK7(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f323,f242])).
fof(f3806,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK7(X0,X1,X2,bool))) )),
  inference(superposition,[],[f323,f242])).
fof(f3739,plain,(
  ( ! [X208,X206,X207,X205] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',happ(s(fun('type.realax.real','type.extreal.extreal'),X208),s('type.realax.real',X205))) | s(bool,'T_0') = s(bool,sK5(sK2('const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X205))),s('type.extreal.extreal',X206),s('type.extreal.extreal',X207),s(fun('type.realax.real','type.extreal.extreal'),X208)),'type.extreal.extreal')))) )),
  inference(superposition,[],[f804,f302])).
fof(f3738,plain,(
  ( ! [X204,X202,X203,X201] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',happ(s(fun('type.realax.real','type.extreal.extreal'),X204),s('type.realax.real',X201))) | s(bool,'T_0') = s(bool,sK5('const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X201))),s('type.extreal.extreal',X202),s('type.extreal.extreal',X203),s(fun('type.realax.real','type.extreal.extreal'),X204))))) )),
  inference(superposition,[],[f543,f302])).
fof(f3737,plain,(
  ( ! [X198,X200,X199,X197] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',happ(s(fun('type.realax.real','type.extreal.extreal'),X200),s('type.realax.real',X197))) | s(bool,'T_0') = s(bool,sK5('const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X197))),s('type.extreal.extreal',X198),s('type.extreal.extreal',X199),s(fun('type.realax.real','type.extreal.extreal'),X200))))) )),
  inference(superposition,[],[f542,f302])).
fof(f3724,plain,(
  ( ! [X127,X130,X128,X126,X131,X129] : (p(s(bool,'|dollar|exists_1'(s(fun(X126,bool),happ(s(fun('type.realax.real',fun(X126,bool)),X130),s('type.realax.real',X127)))))) | s(X126,X131) = s(X126,sK3(X131,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X127))),s(fun(X126,bool),X128),s(fun(X126,bool),X129),s(fun('type.realax.real',fun(X126,bool)),X130)),X126))) )),
  inference(superposition,[],[f293,f302])).
fof(f3782,plain,(
  ( ! [X125,X123,X121,X124,X122,X120] : (p(s(bool,'|dollar|exists_1'(s(fun(X120,bool),happ(s(fun('type.realax.real',fun(X120,bool)),X124),s('type.realax.real',X121)))))) | ~p(s(bool,happ(s(fun(X120,bool),happ(s(fun('type.realax.real',fun(X120,bool)),X124),s('type.realax.real',X121))),s(X120,sK3(X125,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X121))),s(fun(X120,bool),X122),s(fun(X120,bool),X123),s(fun('type.realax.real',fun(X120,bool)),X124)),X120)))))) )),
  inference(forward_demodulation,[],[f3723,f302])).
fof(f3723,plain,(
  ( ! [X125,X123,X121,X124,X122,X120] : (~p(s(bool,happ(s(fun(X120,bool),happ(s(fun('type.realax.real',fun(X120,bool)),X124),s('type.realax.real',X121))),s(X120,sK3(X125,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X121))),s(fun(X120,bool),X122),s(fun(X120,bool),X123),s(fun('type.realax.real',fun(X120,bool)),X124)),X120))))) | p(s(bool,'|dollar|exists_1'(s(fun(X120,bool),'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X121))),s(fun(X120,bool),X122),s(fun(X120,bool),X123),s(fun('type.realax.real',fun(X120,bool)),X124))))))) )),
  inference(superposition,[],[f294,f302])).
fof(f3781,plain,(
  ( ! [X118,X116,X119,X117] : (s(bool,'T_0') = s(bool,happ(s(fun('type.realax.real',bool),X119),s('type.realax.real',X116))) | s(bool,'F_0') = s(bool,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X116))),s(bool,X117),s(bool,X118),s(fun('type.realax.real',bool),X119)))) )),
  inference(forward_demodulation,[],[f3722,f282])).
fof(f3722,plain,(
  ( ! [X118,X116,X119,X117] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,happ(s(fun('type.realax.real',bool),X119),s('type.realax.real',X116))))) | s(bool,'F_0') = s(bool,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X116))),s(bool,X117),s(bool,X118),s(fun('type.realax.real',bool),X119)))) )),
  inference(superposition,[],[f512,f302])).
fof(f3700,plain,(
  ( ! [X6,X8,X7,X5] : (s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),X8),s('type.realax.real',X5))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5))),s(bool,X6),s(bool,X7),s(fun('type.realax.real',bool),X8)))) )),
  inference(superposition,[],[f302,f242])).
fof(f3577,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(X6,happ(s(fun('type.realax.real',X6),X7),s('type.realax.real',sK1(X5)))) = s(X6,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal',X5),s(X6,X8),s(X6,X9),s(fun('type.realax.real',X6),X7))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X5) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X5)) )),
  inference(superposition,[],[f302,f243])).
fof(f3576,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X1,happ(s(fun('type.realax.real',X1),X2),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))))) = s(X1,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0))),s(X1,X3),s(X1,X4),s(fun('type.realax.real',X1),X2)))) )),
  inference(superposition,[],[f302,f233])).
fof(f3438,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,X1))),s(X0,'const.min.@_1'(s(fun(X0,bool),'=_1'(s(X0,X1))))))))) )),
  inference(unit_resulting_resolution,[],[f2389,f292])).
fof(f3427,plain,(
  p(s(bool,happ(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.min.@_1'(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))))))),
  inference(unit_resulting_resolution,[],[f2355,f292])).
fof(f3433,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))),
  inference(unit_resulting_resolution,[],[f2355,f248])).
fof(f3380,plain,(
  ( ! [X156,X157,X155] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',happ(s(fun('type.num.num','type.extreal.extreal'),X156),s('type.num.num',X157))) | s(bool,'T_0') = s(bool,sK5(sK2('const.ind_type.FCONS_3'(s('type.extreal.extreal',X155),s(fun('type.num.num','type.extreal.extreal'),X156),s('type.num.num','const.num.SUC_1'(s('type.num.num',X157)))),'type.extreal.extreal')))) )),
  inference(superposition,[],[f804,f299])).
fof(f3379,plain,(
  ( ! [X154,X152,X153] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',happ(s(fun('type.num.num','type.extreal.extreal'),X153),s('type.num.num',X154))) | s(bool,'T_0') = s(bool,sK5('const.ind_type.FCONS_3'(s('type.extreal.extreal',X152),s(fun('type.num.num','type.extreal.extreal'),X153),s('type.num.num','const.num.SUC_1'(s('type.num.num',X154))))))) )),
  inference(superposition,[],[f543,f299])).
fof(f3378,plain,(
  ( ! [X151,X149,X150] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',happ(s(fun('type.num.num','type.extreal.extreal'),X150),s('type.num.num',X151))) | s(bool,'T_0') = s(bool,sK5('const.ind_type.FCONS_3'(s('type.extreal.extreal',X149),s(fun('type.num.num','type.extreal.extreal'),X150),s('type.num.num','const.num.SUC_1'(s('type.num.num',X151))))))) )),
  inference(superposition,[],[f542,f299])).
fof(f3366,plain,(
  ( ! [X101,X99,X97,X100,X98] : (p(s(bool,'|dollar|exists_1'(s(fun(X97,bool),happ(s(fun('type.num.num',fun(X97,bool)),X99),s('type.num.num',X100)))))) | s(X97,X101) = s(X97,sK3(X101,'const.ind_type.FCONS_3'(s(fun(X97,bool),X98),s(fun('type.num.num',fun(X97,bool)),X99),s('type.num.num','const.num.SUC_1'(s('type.num.num',X100)))),X97))) )),
  inference(superposition,[],[f293,f299])).
fof(f3420,plain,(
  ( ! [X94,X92,X95,X93,X96] : (p(s(bool,'|dollar|exists_1'(s(fun(X92,bool),happ(s(fun('type.num.num',fun(X92,bool)),X94),s('type.num.num',X95)))))) | ~p(s(bool,happ(s(fun(X92,bool),happ(s(fun('type.num.num',fun(X92,bool)),X94),s('type.num.num',X95))),s(X92,sK3(X96,'const.ind_type.FCONS_3'(s(fun(X92,bool),X93),s(fun('type.num.num',fun(X92,bool)),X94),s('type.num.num','const.num.SUC_1'(s('type.num.num',X95)))),X92)))))) )),
  inference(forward_demodulation,[],[f3365,f299])).
fof(f3365,plain,(
  ( ! [X94,X92,X95,X93,X96] : (~p(s(bool,happ(s(fun(X92,bool),happ(s(fun('type.num.num',fun(X92,bool)),X94),s('type.num.num',X95))),s(X92,sK3(X96,'const.ind_type.FCONS_3'(s(fun(X92,bool),X93),s(fun('type.num.num',fun(X92,bool)),X94),s('type.num.num','const.num.SUC_1'(s('type.num.num',X95)))),X92))))) | p(s(bool,'|dollar|exists_1'(s(fun(X92,bool),'const.ind_type.FCONS_3'(s(fun(X92,bool),X93),s(fun('type.num.num',fun(X92,bool)),X94),s('type.num.num','const.num.SUC_1'(s('type.num.num',X95))))))))) )),
  inference(superposition,[],[f294,f299])).
fof(f3419,plain,(
  ( ! [X90,X91,X89] : (s(bool,'T_0') = s(bool,happ(s(fun('type.num.num',bool),X90),s('type.num.num',X91))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X89),s(fun('type.num.num',bool),X90),s('type.num.num','const.num.SUC_1'(s('type.num.num',X91)))))) )),
  inference(forward_demodulation,[],[f3364,f282])).
fof(f3364,plain,(
  ( ! [X90,X91,X89] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,happ(s(fun('type.num.num',bool),X90),s('type.num.num',X91))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X89),s(fun('type.num.num',bool),X90),s('type.num.num','const.num.SUC_1'(s('type.num.num',X91)))))) )),
  inference(superposition,[],[f512,f299])).
fof(f3343,plain,(
  ( ! [X6,X4,X5] : (s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),X5),s('type.num.num',X6))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X4),s(fun('type.num.num',bool),X5),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))) )),
  inference(superposition,[],[f299,f242])).
fof(f293,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(X0,X2) = s(X0,sK3(X2,X1,X0))) )),
  inference(cnf_transformation,[],[f205])).
fof(f205,plain,(
  ! [X0,X1,X2] : ((s(X0,X2) = s(X0,sK3(X2,X1,X0)) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK3(X2,X1,X0)))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3])],[f165])).
fof(f165,plain,(
  ! [X0,X1,X2] : (? [X3] : (s(X0,X2) = s(X0,X3) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(ennf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0,X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f127])).
fof(f127,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X7] : ! [X15,X24] : (! [X6] : (s(X7,X6) = s(X7,X24) => p(s(bool,happ(s(fun(X7,bool),X15),s(X7,X6))))) => p(s(bool,'|dollar|exists_1'(s(fun(X7,bool),X15)))))),
  file('thm.extreal.le_infty|split|2','thm.bool.JRH_INDUCT_UTIL')).
fof(f294,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK3(X2,X1,X0))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f205])).
fof(f2577,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK3(X4,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2),s(fun('type.realax.real',fun(X0,bool)),X3)),X0)))))) )),
  inference(forward_demodulation,[],[f2568,f297])).
fof(f2568,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK3(X4,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2),s(fun('type.realax.real',fun(X0,bool)),X3)),X0))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2),s(fun('type.realax.real',fun(X0,bool)),X3))))))) )),
  inference(superposition,[],[f294,f297])).
fof(f1633,plain,(
  ( ! [X43,X41,X44,X42,X40] : (p(s(bool,'|dollar|exists_1'(s(fun(X40,bool),X41)))) | s(X40,X44) = s(X40,sK3(X44,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(fun(X40,bool),X41),s(fun(X40,bool),X42),s(fun('type.realax.real',fun(X40,bool)),X43)),X40))) )),
  inference(superposition,[],[f293,f297])).
fof(f2578,plain,(
  ( ! [X6,X8,X7,X5,X9] : (p(s(bool,'|dollar|exists_1'(s(fun(X5,bool),X7)))) | ~p(s(bool,happ(s(fun(X5,bool),X7),s(X5,sK3(X9,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.PosInf_0'),s(fun(X5,bool),X6),s(fun(X5,bool),X7),s(fun('type.realax.real',fun(X5,bool)),X8)),X5)))))) )),
  inference(forward_demodulation,[],[f2569,f298])).
fof(f2569,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~p(s(bool,happ(s(fun(X5,bool),X7),s(X5,sK3(X9,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.PosInf_0'),s(fun(X5,bool),X6),s(fun(X5,bool),X7),s(fun('type.realax.real',fun(X5,bool)),X8)),X5))))) | p(s(bool,'|dollar|exists_1'(s(fun(X5,bool),'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.PosInf_0'),s(fun(X5,bool),X6),s(fun(X5,bool),X7),s(fun('type.realax.real',fun(X5,bool)),X8))))))) )),
  inference(superposition,[],[f294,f298])).
fof(f1863,plain,(
  ( ! [X59,X57,X58,X56,X55] : (p(s(bool,'|dollar|exists_1'(s(fun(X55,bool),X57)))) | s(X55,X59) = s(X55,sK3(X59,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.PosInf_0'),s(fun(X55,bool),X56),s(fun(X55,bool),X57),s(fun('type.realax.real',fun(X55,bool)),X58)),X55))) )),
  inference(superposition,[],[f293,f298])).
fof(f2579,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,'|dollar|exists_1'(s(fun(X10,bool),X11)))) | ~p(s(bool,happ(s(fun(X10,bool),X11),s(X10,sK3(X12,happ(s(fun(fun(X10,bool),fun(X10,bool)),'const.combin.I_0'),s(fun(X10,bool),X11)),X10)))))) )),
  inference(forward_demodulation,[],[f2570,f339])).
fof(f2570,plain,(
  ( ! [X12,X10,X11] : (~p(s(bool,happ(s(fun(X10,bool),X11),s(X10,sK3(X12,happ(s(fun(fun(X10,bool),fun(X10,bool)),'const.combin.I_0'),s(fun(X10,bool),X11)),X10))))) | p(s(bool,'|dollar|exists_1'(s(fun(X10,bool),happ(s(fun(fun(X10,bool),fun(X10,bool)),'const.combin.I_0'),s(fun(X10,bool),X11))))))) )),
  inference(superposition,[],[f294,f339])).
fof(f1263,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(X0,X2) = s(X0,sK3(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'const.combin.I_0'),s(fun(X0,bool),X1)),X0))) )),
  inference(superposition,[],[f293,f339])).
fof(f2580,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,'|dollar|exists_1'(s(fun(X13,bool),X14)))) | ~p(s(bool,happ(s(fun(X13,bool),X14),s(X13,sK3(X15,'const.combin.I_1'(s(fun(X13,bool),X14)),X13)))))) )),
  inference(forward_demodulation,[],[f2571,f282])).
fof(f2571,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,happ(s(fun(X13,bool),X14),s(X13,sK3(X15,'const.combin.I_1'(s(fun(X13,bool),X14)),X13))))) | p(s(bool,'|dollar|exists_1'(s(fun(X13,bool),'const.combin.I_1'(s(fun(X13,bool),X14))))))) )),
  inference(superposition,[],[f294,f282])).
fof(f1264,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'|dollar|exists_1'(s(fun(X3,bool),X4)))) | s(X3,X5) = s(X3,sK3(X5,'const.combin.I_1'(s(fun(X3,bool),X4)),X3))) )),
  inference(superposition,[],[f293,f282])).
fof(f2581,plain,(
  ( ! [X19,X17,X18,X16] : (p(s(bool,'|dollar|exists_1'(s(fun(X16,bool),X17)))) | ~p(s(bool,happ(s(fun(X16,bool),X17),s(X16,sK3(X19,'const.ind_type.FCONS_3'(s(fun(X16,bool),X17),s(fun('type.num.num',fun(X16,bool)),X18),s('type.num.num','const.num.|48|_0')),X16)))))) )),
  inference(forward_demodulation,[],[f2572,f291])).
fof(f2572,plain,(
  ( ! [X19,X17,X18,X16] : (~p(s(bool,happ(s(fun(X16,bool),X17),s(X16,sK3(X19,'const.ind_type.FCONS_3'(s(fun(X16,bool),X17),s(fun('type.num.num',fun(X16,bool)),X18),s('type.num.num','const.num.|48|_0')),X16))))) | p(s(bool,'|dollar|exists_1'(s(fun(X16,bool),'const.ind_type.FCONS_3'(s(fun(X16,bool),X17),s(fun('type.num.num',fun(X16,bool)),X18),s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f294,f291])).
fof(f1265,plain,(
  ( ! [X6,X8,X7,X9] : (p(s(bool,'|dollar|exists_1'(s(fun(X6,bool),X7)))) | s(X6,X9) = s(X6,sK3(X9,'const.ind_type.FCONS_3'(s(fun(X6,bool),X7),s(fun('type.num.num',fun(X6,bool)),X8),s('type.num.num','const.num.|48|_0')),X6))) )),
  inference(superposition,[],[f293,f291])).
fof(f2582,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,'|dollar|exists_1'(s(fun(X20,bool),X21)))) | ~p(s(bool,happ(s(fun(X20,bool),X21),s(X20,sK3(X22,sK2(X21,fun(X20,bool)),X20)))))) )),
  inference(forward_demodulation,[],[f2573,f290])).
fof(f2573,plain,(
  ( ! [X21,X22,X20] : (~p(s(bool,happ(s(fun(X20,bool),X21),s(X20,sK3(X22,sK2(X21,fun(X20,bool)),X20))))) | p(s(bool,'|dollar|exists_1'(s(fun(X20,bool),sK2(X21,fun(X20,bool))))))) )),
  inference(superposition,[],[f294,f290])).
fof(f1266,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,'|dollar|exists_1'(s(fun(X10,bool),X11)))) | s(X10,X12) = s(X10,sK3(X12,sK2(X11,fun(X10,bool)),X10))) )),
  inference(superposition,[],[f293,f290])).
fof(f2574,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X1)))) | s(bool,'T_0') = s(bool,sK3(X0,X1,bool))) )),
  inference(superposition,[],[f294,f242])).
fof(f2764,plain,(
  ( ! [X21,X20] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X20),s('type.realax.real',X21))) | s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X20))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X21)))))) )),
  inference(forward_demodulation,[],[f2743,f282])).
fof(f2743,plain,(
  ( ! [X21,X20] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'const.real.real_lte_2'(s('type.realax.real',X20),s('type.realax.real',X21))))) | s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X20))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X21)))))) )),
  inference(superposition,[],[f512,f315])).
fof(f2733,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',X4))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4)))))) )),
  inference(superposition,[],[f315,f242])).
fof(f2731,plain,(
  ( ! [X2,X3] : (s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',sK1(X2)))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))),s('type.extreal.extreal',X2))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X2) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X2)) )),
  inference(superposition,[],[f315,f243])).
fof(f2730,plain,(
  ( ! [X0,X1] : (s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0)))))) )),
  inference(superposition,[],[f315,f233])).
fof(f2723,plain,(
  ( ! [X2,X3] : (s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1(X2)),s('type.realax.real',X3))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X2) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X2)) )),
  inference(superposition,[],[f315,f243])).
fof(f2722,plain,(
  ( ! [X0,X1] : (s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))),s('type.realax.real',X1))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))))) )),
  inference(superposition,[],[f315,f233])).
fof(f2658,plain,(
  ( ! [X45,X44] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal',X45))),s('type.extreal.extreal',X44)))) | s('type.extreal.extreal',X45) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK1(X44)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X44) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X44)) )),
  inference(superposition,[],[f326,f243])).
fof(f2657,plain,(
  ( ! [X43,X42] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal',X42))),s('type.extreal.extreal',X43)))) | s('type.extreal.extreal',X43) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK1(X42)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X42) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X42)) )),
  inference(superposition,[],[f326,f243])).
fof(f2656,plain,(
  ( ! [X41,X40] : (~p(s(bool,happ(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal',X41))),s('type.extreal.extreal',X40)))) | s('type.extreal.extreal',X41) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK1(X40)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X40) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X40)) )),
  inference(superposition,[],[f325,f243])).
fof(f2655,plain,(
  ( ! [X39,X38] : (~p(s(bool,happ(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal',X38))),s('type.extreal.extreal',X39)))) | s('type.extreal.extreal',X39) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK1(X38)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X38) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X38)) )),
  inference(superposition,[],[f325,f243])).
fof(f2643,plain,(
  ( ! [X17,X16] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X17),s('type.extreal.extreal',X16)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK1(X16)))),s('type.extreal.extreal',X17)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X16) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X16)) )),
  inference(superposition,[],[f321,f243])).
fof(f2642,plain,(
  ( ! [X14,X15] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X14),s('type.extreal.extreal',X15)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X15),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK1(X14))))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X14) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X14)) )),
  inference(superposition,[],[f321,f243])).
fof(f2641,plain,(
  ( ! [X12,X13] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X13),s('type.extreal.extreal',X12)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK1(X12)))),s('type.extreal.extreal',X13)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X12) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X12)) )),
  inference(superposition,[],[f320,f243])).
fof(f2640,plain,(
  ( ! [X10,X11] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X10),s('type.extreal.extreal',X11)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X11),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK1(X10))))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X10) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X10)) )),
  inference(superposition,[],[f320,f243])).
fof(f2638,plain,(
  ( ! [X8] : (s(bool,'T_0') != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X8))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X8) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X8)) )),
  inference(superposition,[],[f1080,f243])).
fof(f2637,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X7) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X7)) )),
  inference(superposition,[],[f950,f243])).
fof(f2636,plain,(
  ( ! [X6] : (s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X6))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X6) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X6)) )),
  inference(superposition,[],[f240,f243])).
fof(f2635,plain,(
  ( ! [X5] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X5),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X5) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X5)) )),
  inference(superposition,[],[f239,f243])).
fof(f2632,plain,(
  ( ! [X2] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X2))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X2) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X2)) )),
  inference(superposition,[],[f236,f243])).
fof(f2631,plain,(
  ( ! [X1] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X1) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X1)) )),
  inference(superposition,[],[f235,f243])).
fof(f2630,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0)) )),
  inference(superposition,[],[f234,f243])).
fof(f2519,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,X1))),s(X0,X2))) | s(X0,X1) != s(X0,X2)) )),
  inference(subsumption_resolution,[],[f2354,f337])).
fof(f2354,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | s(X0,X1) != s(X0,X2) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,X1))),s(X0,X2)))) )),
  inference(superposition,[],[f326,f242])).
fof(f2494,plain,(
  ( ! [X47,X46] : (s(bool,X46) != s(bool,X47) | p(s(bool,happ(s(fun(bool,bool),'=_1'(s(bool,X47))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X46)) )),
  inference(forward_demodulation,[],[f2329,f282])).
fof(f2329,plain,(
  ( ! [X47,X46] : (p(s(bool,happ(s(fun(bool,bool),'=_1'(s(bool,X47))),s(bool,'T_0')))) | s(bool,X47) != s(bool,'const.combin.I_1'(s(bool,X46))) | s(bool,'F_0') = s(bool,X46)) )),
  inference(superposition,[],[f326,f512])).
fof(f2317,plain,(
  ( ! [X26,X25] : (p(s(bool,happ(s(fun(bool,bool),'=_1'(s(bool,X26))),s(bool,X25)))) | s(bool,'F_0') != s(bool,X26) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f326,f242])).
fof(f2316,plain,(
  ( ! [X24,X23] : (p(s(bool,happ(s(fun(bool,bool),'=_1'(s(bool,X24))),s(bool,'F_0')))) | s(bool,X23) != s(bool,X24) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f326,f242])).
fof(f2452,plain,(
  ( ! [X47,X46] : (s(bool,X46) != s(bool,X47) | p(s(bool,happ(s(fun(bool,bool),'=_1'(s(bool,'T_0'))),s(bool,X47)))) | s(bool,'F_0') = s(bool,X46)) )),
  inference(forward_demodulation,[],[f2285,f282])).
fof(f2285,plain,(
  ( ! [X47,X46] : (p(s(bool,happ(s(fun(bool,bool),'=_1'(s(bool,'T_0'))),s(bool,X47)))) | s(bool,X47) != s(bool,'const.combin.I_1'(s(bool,X46))) | s(bool,'F_0') = s(bool,X46)) )),
  inference(superposition,[],[f326,f512])).
fof(f2273,plain,(
  ( ! [X26,X25] : (p(s(bool,happ(s(fun(bool,bool),'=_1'(s(bool,X25))),s(bool,X26)))) | s(bool,'F_0') != s(bool,X26) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f326,f242])).
fof(f2272,plain,(
  ( ! [X24,X23] : (p(s(bool,happ(s(fun(bool,bool),'=_1'(s(bool,'F_0'))),s(bool,X24)))) | s(bool,X23) != s(bool,X24) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f326,f242])).
fof(f2378,plain,(
  p(s(bool,happ(s(fun(bool,bool),'=_1'(s(bool,'F_0'))),s(bool,'F_0'))))),
  inference(forward_demodulation,[],[f2242,f1069])).
fof(f2242,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),'=_1'(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0))))))),s(bool,'F_0'))))) )),
  inference(unit_resulting_resolution,[],[f1069,f326])).
fof(f2366,plain,(
  p(s(bool,happ(s(fun(bool,bool),'=_1'(s(bool,'T_0'))),s(bool,'T_0'))))),
  inference(forward_demodulation,[],[f2254,f1919])).
fof(f2254,plain,(
  p(s(bool,happ(s(fun(bool,bool),'=_1'(s(bool,sK5(sK2(sK2('const.extreal.PosInf_0','type.extreal.extreal'),'type.extreal.extreal'))))),s(bool,'T_0'))))),
  inference(unit_resulting_resolution,[],[f1919,f326])).
fof(f2365,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0))))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0))))))) )),
  inference(forward_demodulation,[],[f2255,f233])).
fof(f2255,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))))))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0))))))) )),
  inference(unit_resulting_resolution,[],[f233,f326])).
fof(f2364,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))))))) )),
  inference(forward_demodulation,[],[f2256,f241])).
fof(f2256,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal',happ(s(fun('type.realax.real','type.extreal.extreal'),'const.extreal.Normal_0'),s('type.realax.real',X0))))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))))))) )),
  inference(unit_resulting_resolution,[],[f241,f326])).
fof(f2355,plain,(
  p(s(bool,happ(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.NegInf_0'))))),
  inference(forward_demodulation,[],[f2265,f232])).
fof(f2265,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal','const.extreal.extreal_logr_2'(s('type.realax.real',X0),s('type.extreal.extreal','const.extreal.NegInf_0'))))),s('type.extreal.extreal','const.extreal.NegInf_0'))))) )),
  inference(unit_resulting_resolution,[],[f232,f326])).
fof(f2145,plain,(
  ( ! [X47,X46] : (s(bool,X46) = s(bool,X47) | ~p(s(bool,happ(s(fun(bool,bool),'=_1'(s(bool,X47))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X46)) )),
  inference(forward_demodulation,[],[f2050,f282])).
fof(f2050,plain,(
  ( ! [X47,X46] : (~p(s(bool,happ(s(fun(bool,bool),'=_1'(s(bool,X47))),s(bool,'T_0')))) | s(bool,X47) = s(bool,'const.combin.I_1'(s(bool,X46))) | s(bool,'F_0') = s(bool,X46)) )),
  inference(superposition,[],[f325,f512])).
fof(f2038,plain,(
  ( ! [X26,X25] : (~p(s(bool,happ(s(fun(bool,bool),'=_1'(s(bool,X26))),s(bool,X25)))) | s(bool,'F_0') = s(bool,X26) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f325,f242])).
fof(f2037,plain,(
  ( ! [X24,X23] : (~p(s(bool,happ(s(fun(bool,bool),'=_1'(s(bool,X24))),s(bool,'F_0')))) | s(bool,X23) = s(bool,X24) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f325,f242])).
fof(f2103,plain,(
  ( ! [X47,X46] : (s(bool,X46) = s(bool,X47) | ~p(s(bool,happ(s(fun(bool,bool),'=_1'(s(bool,'T_0'))),s(bool,X47)))) | s(bool,'F_0') = s(bool,X46)) )),
  inference(forward_demodulation,[],[f2006,f282])).
fof(f2006,plain,(
  ( ! [X47,X46] : (~p(s(bool,happ(s(fun(bool,bool),'=_1'(s(bool,'T_0'))),s(bool,X47)))) | s(bool,X47) = s(bool,'const.combin.I_1'(s(bool,X46))) | s(bool,'F_0') = s(bool,X46)) )),
  inference(superposition,[],[f325,f512])).
fof(f1994,plain,(
  ( ! [X26,X25] : (~p(s(bool,happ(s(fun(bool,bool),'=_1'(s(bool,X25))),s(bool,X26)))) | s(bool,'F_0') = s(bool,X26) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f325,f242])).
fof(f1993,plain,(
  ( ! [X24,X23] : (~p(s(bool,happ(s(fun(bool,bool),'=_1'(s(bool,'F_0'))),s(bool,X24)))) | s(bool,X23) = s(bool,X24) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f325,f242])).
fof(f1970,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0))))))) )),
  inference(unit_resulting_resolution,[],[f281,f325])).
fof(f2081,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),'=_1'(s(bool,'T_0'))),s(bool,'F_0'))))),
  inference(forward_demodulation,[],[f1977,f1069])).
fof(f1977,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),'=_1'(s(bool,'T_0'))),s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0))))))))) )),
  inference(unit_resulting_resolution,[],[f1323,f325])).
fof(f1979,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0))))),s('type.extreal.extreal','const.extreal.NegInf_0'))))) )),
  inference(unit_resulting_resolution,[],[f281,f325])).
fof(f2076,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),'=_1'(s(bool,'F_0'))),s(bool,'T_0'))))),
  inference(forward_demodulation,[],[f1986,f1069])).
fof(f1986,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),'=_1'(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0))))))),s(bool,'T_0'))))) )),
  inference(unit_resulting_resolution,[],[f1323,f325])).
fof(f1936,plain,(
  ( ! [X21] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X21) | s(bool,'T_0') = s(bool,sK5(sK2(sK2(X21,'type.extreal.extreal'),'type.extreal.extreal')))) )),
  inference(superposition,[],[f804,f290])).
fof(f1935,plain,(
  ( ! [X19,X20] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X19) | s(bool,'T_0') = s(bool,sK5(sK2('const.ind_type.FCONS_3'(s('type.extreal.extreal',X19),s(fun('type.num.num','type.extreal.extreal'),X20),s('type.num.num','const.num.|48|_0')),'type.extreal.extreal')))) )),
  inference(superposition,[],[f804,f291])).
fof(f1934,plain,(
  ( ! [X18] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X18) | s(bool,'T_0') = s(bool,sK5(sK2('const.combin.I_1'(s('type.extreal.extreal',X18)),'type.extreal.extreal')))) )),
  inference(superposition,[],[f804,f282])).
fof(f1933,plain,(
  ( ! [X17] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X17) | s(bool,'T_0') = s(bool,sK5(sK2(happ(s(fun('type.extreal.extreal','type.extreal.extreal'),'const.combin.I_0'),s('type.extreal.extreal',X17)),'type.extreal.extreal')))) )),
  inference(superposition,[],[f804,f339])).
fof(f1932,plain,(
  ( ! [X14,X15,X16] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X15) | s(bool,'T_0') = s(bool,sK5(sK2('const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X14),s('type.extreal.extreal',X15),s(fun('type.realax.real','type.extreal.extreal'),X16)),'type.extreal.extreal')))) )),
  inference(superposition,[],[f804,f298])).
fof(f1931,plain,(
  ( ! [X12,X13,X11] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X11) | s(bool,'T_0') = s(bool,sK5(sK2('const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X11),s('type.extreal.extreal',X12),s(fun('type.realax.real','type.extreal.extreal'),X13)),'type.extreal.extreal')))) )),
  inference(superposition,[],[f804,f297])).
fof(f1921,plain,(
  ( ! [X1] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1))) | s(bool,'T_0') = s(bool,sK5(sK2(happ(s(fun('type.realax.real','type.extreal.extreal'),'const.extreal.Normal_0'),s('type.realax.real',X1)),'type.extreal.extreal')))) )),
  inference(superposition,[],[f804,f241])).
fof(f1920,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0))) | s(bool,'T_0') = s(bool,sK5(sK2('const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0)))),'type.extreal.extreal')))) )),
  inference(superposition,[],[f804,f233])).
fof(f1914,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK5(sK2('const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X0),s(fun('type.realax.real','type.extreal.extreal'),X1)),'type.extreal.extreal')))) )),
  inference(unit_resulting_resolution,[],[f297,f804])).
fof(f1915,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK5(sK2('const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.PosInf_0'),s(fun('type.realax.real','type.extreal.extreal'),X1)),'type.extreal.extreal')))) )),
  inference(unit_resulting_resolution,[],[f298,f804])).
fof(f1916,plain,(
  s(bool,'T_0') = s(bool,sK5(sK2(happ(s(fun('type.extreal.extreal','type.extreal.extreal'),'const.combin.I_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),'type.extreal.extreal')))),
  inference(unit_resulting_resolution,[],[f339,f804])).
fof(f1917,plain,(
  s(bool,'T_0') = s(bool,sK5(sK2('const.combin.I_1'(s('type.extreal.extreal','const.extreal.PosInf_0')),'type.extreal.extreal')))),
  inference(unit_resulting_resolution,[],[f282,f804])).
fof(f1918,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,sK5(sK2('const.ind_type.FCONS_3'(s('type.extreal.extreal','const.extreal.PosInf_0'),s(fun('type.num.num','type.extreal.extreal'),X0),s('type.num.num','const.num.|48|_0')),'type.extreal.extreal')))) )),
  inference(unit_resulting_resolution,[],[f291,f804])).
fof(f1919,plain,(
  s(bool,'T_0') = s(bool,sK5(sK2(sK2('const.extreal.PosInf_0','type.extreal.extreal'),'type.extreal.extreal')))),
  inference(unit_resulting_resolution,[],[f290,f804])).
fof(f1874,plain,(
  ( ! [X101,X99,X100] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X100) | s(bool,'T_0') = s(bool,sK5('const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X99),s('type.extreal.extreal',X100),s(fun('type.realax.real','type.extreal.extreal'),X101))))) )),
  inference(superposition,[],[f543,f298])).
fof(f1873,plain,(
  ( ! [X97,X98,X96] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X97) | s(bool,'T_0') = s(bool,sK5('const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X96),s('type.extreal.extreal',X97),s(fun('type.realax.real','type.extreal.extreal'),X98))))) )),
  inference(superposition,[],[f542,f298])).
fof(f1909,plain,(
  ( ! [X54,X52,X53] : (s(bool,'T_0') = s(bool,X53) | s(bool,'F_0') = s(bool,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.PosInf_0'),s(bool,X52),s(bool,X53),s(fun('type.realax.real',bool),X54)))) )),
  inference(forward_demodulation,[],[f1862,f282])).
fof(f1862,plain,(
  ( ! [X54,X52,X53] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X53))) | s(bool,'F_0') = s(bool,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.PosInf_0'),s(bool,X52),s(bool,X53),s(fun('type.realax.real',bool),X54)))) )),
  inference(superposition,[],[f512,f298])).
fof(f1848,plain,(
  ( ! [X6,X4,X5] : (s(bool,'F_0') = s(bool,X5) | s(bool,'T_0') = s(bool,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.PosInf_0'),s(bool,X4),s(bool,X5),s(fun('type.realax.real',bool),X6)))) )),
  inference(superposition,[],[f298,f242])).
fof(f1752,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK5('const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0'),s(fun('type.realax.real','type.extreal.extreal'),X1))))) )),
  inference(unit_resulting_resolution,[],[f298,f543])).
fof(f1753,plain,(
  ( ! [X0,X1] : (p(s(bool,sK5('const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0'),s(fun('type.realax.real','type.extreal.extreal'),X1)))))) )),
  inference(unit_resulting_resolution,[],[f298,f312])).
fof(f1754,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK5('const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.PosInf_0'),s(fun('type.realax.real','type.extreal.extreal'),X1))))) )),
  inference(unit_resulting_resolution,[],[f298,f542])).
fof(f1755,plain,(
  ( ! [X0,X1] : (p(s(bool,sK5('const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.PosInf_0'),s(fun('type.realax.real','type.extreal.extreal'),X1)))))) )),
  inference(unit_resulting_resolution,[],[f298,f313])).
fof(f1747,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,X3))) | s(bool,'F_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f1732,f217])).
fof(f1732,plain,(
  ( ! [X3] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,X3))) | s(bool,'F_0') = s(bool,X3)) )),
  inference(superposition,[],[f251,f512])).
fof(f1745,plain,(
  ( ! [X21,X20] : (s(bool,'T_0') = s(bool,X20) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X20),s(fun('type.num.num',bool),X21),s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f1725,f282])).
fof(f1725,plain,(
  ( ! [X21,X20] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X20))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X20),s(fun('type.num.num',bool),X21),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f512,f291])).
fof(f1744,plain,(
  ( ! [X18] : (s(bool,'T_0') = s(bool,X18) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X18)))) )),
  inference(forward_demodulation,[],[f1723,f282])).
fof(f1723,plain,(
  ( ! [X18] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X18))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X18)))) )),
  inference(superposition,[],[f512,f339])).
fof(f1743,plain,(
  ( ! [X17,X15,X16] : (s(bool,'T_0') = s(bool,X15) | s(bool,'F_0') = s(bool,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(bool,X15),s(bool,X16),s(fun('type.realax.real',bool),X17)))) )),
  inference(forward_demodulation,[],[f1722,f282])).
fof(f1722,plain,(
  ( ! [X17,X15,X16] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X15))) | s(bool,'F_0') = s(bool,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(bool,X15),s(bool,X16),s(fun('type.realax.real',bool),X17)))) )),
  inference(superposition,[],[f512,f297])).
fof(f1643,plain,(
  ( ! [X80,X78,X79] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X78) | s(bool,'T_0') = s(bool,sK5('const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X78),s('type.extreal.extreal',X79),s(fun('type.realax.real','type.extreal.extreal'),X80))))) )),
  inference(superposition,[],[f543,f297])).
fof(f1642,plain,(
  ( ! [X76,X77,X75] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X75) | s(bool,'T_0') = s(bool,sK5('const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X75),s('type.extreal.extreal',X76),s(fun('type.realax.real','type.extreal.extreal'),X77))))) )),
  inference(superposition,[],[f542,f297])).
fof(f1621,plain,(
  ( ! [X6,X4,X5] : (s(bool,'F_0') = s(bool,X4) | s(bool,'T_0') = s(bool,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(bool,X4),s(bool,X5),s(fun('type.realax.real',bool),X6)))) )),
  inference(superposition,[],[f297,f242])).
fof(f1530,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK5('const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0),s(fun('type.realax.real','type.extreal.extreal'),X1))))) )),
  inference(unit_resulting_resolution,[],[f297,f543])).
fof(f1531,plain,(
  ( ! [X0,X1] : (p(s(bool,sK5('const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0),s(fun('type.realax.real','type.extreal.extreal'),X1)))))) )),
  inference(unit_resulting_resolution,[],[f297,f312])).
fof(f1532,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK5('const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X0),s(fun('type.realax.real','type.extreal.extreal'),X1))))) )),
  inference(unit_resulting_resolution,[],[f297,f542])).
fof(f1533,plain,(
  ( ! [X0,X1] : (p(s(bool,sK5('const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X0),s(fun('type.realax.real','type.extreal.extreal'),X1)))))) )),
  inference(unit_resulting_resolution,[],[f297,f313])).
fof(f1506,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) )),
  inference(subsumption_resolution,[],[f1475,f337])).
fof(f1475,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) )),
  inference(superposition,[],[f321,f242])).
fof(f1323,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X1)))))) )),
  inference(subsumption_resolution,[],[f1312,f337])).
fof(f1312,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X1)))))) )),
  inference(superposition,[],[f247,f1069])).
fof(f1283,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,sK5('const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0))),s('type.extreal.extreal','const.extreal.PosInf_0'))))) )),
  inference(unit_resulting_resolution,[],[f1055,f543])).
fof(f1284,plain,(
  ( ! [X0] : (p(s(bool,sK5('const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0))),s('type.extreal.extreal','const.extreal.PosInf_0')))))) )),
  inference(unit_resulting_resolution,[],[f1055,f312])).
fof(f1268,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,sK3(X2,X1,X0)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) )),
  inference(subsumption_resolution,[],[f1267,f337])).
fof(f1267,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | s(X0,X2) = s(X0,sK3(X2,X1,X0)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) )),
  inference(superposition,[],[f293,f242])).
fof(f1258,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X1))),s('type.extreal.extreal','const.extreal.NegInf_0')))) )),
  inference(subsumption_resolution,[],[f1247,f337])).
fof(f1247,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X1))),s('type.extreal.extreal','const.extreal.NegInf_0')))) )),
  inference(superposition,[],[f247,f939])).
fof(f1196,plain,(
  ( ! [X28,X27] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X27) | s(bool,'T_0') = s(bool,sK5('const.ind_type.FCONS_3'(s('type.extreal.extreal',X27),s(fun('type.num.num','type.extreal.extreal'),X28),s('type.num.num','const.num.|48|_0'))))) )),
  inference(superposition,[],[f542,f291])).
fof(f1195,plain,(
  ( ! [X26,X25] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X25) | s(bool,'T_0') = s(bool,sK5('const.ind_type.FCONS_3'(s('type.extreal.extreal',X25),s(fun('type.num.num','type.extreal.extreal'),X26),s('type.num.num','const.num.|48|_0'))))) )),
  inference(superposition,[],[f543,f291])).
fof(f1184,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X3),s(fun('type.num.num',bool),X4),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X3)) )),
  inference(superposition,[],[f291,f242])).
fof(f1142,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,sK5('const.ind_type.FCONS_3'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(fun('type.num.num','type.extreal.extreal'),X0),s('type.num.num','const.num.|48|_0'))))) )),
  inference(unit_resulting_resolution,[],[f291,f543])).
fof(f1143,plain,(
  ( ! [X0] : (p(s(bool,sK5('const.ind_type.FCONS_3'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(fun('type.num.num','type.extreal.extreal'),X0),s('type.num.num','const.num.|48|_0')))))) )),
  inference(unit_resulting_resolution,[],[f291,f312])).
fof(f1144,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,sK5('const.ind_type.FCONS_3'(s('type.extreal.extreal','const.extreal.PosInf_0'),s(fun('type.num.num','type.extreal.extreal'),X0),s('type.num.num','const.num.|48|_0'))))) )),
  inference(unit_resulting_resolution,[],[f291,f542])).
fof(f1145,plain,(
  ( ! [X0] : (p(s(bool,sK5('const.ind_type.FCONS_3'(s('type.extreal.extreal','const.extreal.PosInf_0'),s(fun('type.num.num','type.extreal.extreal'),X0),s('type.num.num','const.num.|48|_0')))))) )),
  inference(unit_resulting_resolution,[],[f291,f313])).
fof(f1106,plain,(
  ( ! [X1] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1))) | s(bool,'T_0') = s(bool,sK5(happ(s(fun('type.realax.real','type.extreal.extreal'),'const.extreal.Normal_0'),s('type.realax.real',X1))))) )),
  inference(superposition,[],[f542,f241])).
fof(f1105,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))) | s(bool,'T_0') = s(bool,sK5(happ(s(fun('type.realax.real','type.extreal.extreal'),'const.extreal.Normal_0'),s('type.realax.real',X0))))) )),
  inference(superposition,[],[f543,f241])).
fof(f1101,plain,(
  ( ! [X10,X9] : (s(bool,X9) = s(bool,X10) | ~p(s(bool,X10)) | s(bool,'F_0') = s(bool,X9)) )),
  inference(forward_demodulation,[],[f1100,f290])).
fof(f1100,plain,(
  ( ! [X10,X9] : (~p(s(bool,X10)) | s(bool,X10) = s(bool,sK2(X9,bool)) | s(bool,'F_0') = s(bool,X9)) )),
  inference(subsumption_resolution,[],[f1095,f217])).
fof(f1095,plain,(
  ( ! [X10,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,X10)) | s(bool,X10) = s(bool,sK2(X9,bool)) | s(bool,'F_0') = s(bool,X9)) )),
  inference(superposition,[],[f319,f511])).
fof(f1099,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,sK2(X3,bool)) | s(bool,'F_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f1091,f217])).
fof(f1091,plain,(
  ( ! [X3] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,sK2(X3,bool)) | s(bool,'F_0') = s(bool,X3)) )),
  inference(superposition,[],[f251,f511])).
fof(f1080,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))))) )),
  inference(subsumption_resolution,[],[f1070,f337])).
fof(f1070,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))))) )),
  inference(superposition,[],[f247,f240])).
fof(f1069,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0)))))) )),
  inference(superposition,[],[f240,f233])).
fof(f1050,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,sK5('const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0'))))) )),
  inference(unit_resulting_resolution,[],[f239,f543])).
fof(f1051,plain,(
  ( ! [X0] : (p(s(bool,sK5('const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0')))))) )),
  inference(unit_resulting_resolution,[],[f239,f312])).
fof(f1025,plain,(
  s(bool,'T_0') = s(bool,sK5('const.extreal.NegInf_0'))),
  inference(equality_resolution,[],[f543])).
fof(f1024,plain,(
  ( ! [X7] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X7) | s(bool,'T_0') = s(bool,sK5(happ(s(fun('type.extreal.extreal','type.extreal.extreal'),'const.combin.I_0'),s('type.extreal.extreal',X7))))) )),
  inference(superposition,[],[f543,f339])).
fof(f1023,plain,(
  ( ! [X6] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X6) | s(bool,'T_0') = s(bool,sK5('const.combin.I_1'(s('type.extreal.extreal',X6))))) )),
  inference(superposition,[],[f543,f282])).
fof(f977,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,sK5('const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0))))))) )),
  inference(unit_resulting_resolution,[],[f909,f248])).
fof(f960,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,sK5('const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0))),s('type.extreal.extreal','const.extreal.NegInf_0'))))) )),
  inference(unit_resulting_resolution,[],[f880,f248])).
fof(f950,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1))),s('type.extreal.extreal','const.extreal.NegInf_0')))) )),
  inference(subsumption_resolution,[],[f940,f337])).
fof(f940,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1))),s('type.extreal.extreal','const.extreal.NegInf_0')))) )),
  inference(superposition,[],[f247,f237])).
fof(f939,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0))),s('type.extreal.extreal','const.extreal.NegInf_0')))) )),
  inference(superposition,[],[f237,f233])).
fof(f923,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,sK5('const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0))))))) )),
  inference(unit_resulting_resolution,[],[f862,f248])).
fof(f909,plain,(
  ( ! [X0] : (p(s(bool,sK5('const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0)))))))) )),
  inference(superposition,[],[f889,f233])).
fof(f905,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,sK5('const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))))))) )),
  inference(unit_resulting_resolution,[],[f889,f248])).
fof(f889,plain,(
  ( ! [X0] : (p(s(bool,sK5('const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))))))) )),
  inference(unit_resulting_resolution,[],[f236,f312])).
fof(f880,plain,(
  ( ! [X0] : (p(s(bool,sK5('const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0))),s('type.extreal.extreal','const.extreal.NegInf_0')))))) )),
  inference(superposition,[],[f842,f233])).
fof(f876,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,sK5('const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.NegInf_0'))))) )),
  inference(unit_resulting_resolution,[],[f842,f248])).
fof(f862,plain,(
  ( ! [X0] : (p(s(bool,sK5('const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0)))))))) )),
  inference(superposition,[],[f808,f233])).
fof(f858,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,sK5('const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))))))) )),
  inference(unit_resulting_resolution,[],[f808,f248])).
fof(f842,plain,(
  ( ! [X0] : (p(s(bool,sK5('const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.NegInf_0')))))) )),
  inference(unit_resulting_resolution,[],[f235,f312])).
fof(f808,plain,(
  ( ! [X0] : (p(s(bool,sK5('const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))))))) )),
  inference(unit_resulting_resolution,[],[f234,f312])).
fof(f807,plain,(
  s(bool,'T_0') = s(bool,sK5('const.extreal.PosInf_0'))),
  inference(equality_resolution,[],[f542])).
fof(f806,plain,(
  ( ! [X4] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X4) | s(bool,'T_0') = s(bool,sK5(happ(s(fun('type.extreal.extreal','type.extreal.extreal'),'const.combin.I_0'),s('type.extreal.extreal',X4))))) )),
  inference(superposition,[],[f542,f339])).
fof(f805,plain,(
  ( ! [X3] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X3) | s(bool,'T_0') = s(bool,sK5('const.combin.I_1'(s('type.extreal.extreal',X3))))) )),
  inference(superposition,[],[f542,f282])).
fof(f803,plain,(
  ( ! [X1] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X1))) | s(bool,'T_0') = s(bool,sK5('const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1))))))) )),
  inference(superposition,[],[f542,f233])).
fof(f765,plain,(
  s(bool,'T_0') = s(bool,sK5(happ(s(fun('type.extreal.extreal','type.extreal.extreal'),'const.combin.I_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))))),
  inference(unit_resulting_resolution,[],[f707,f248])).
fof(f756,plain,(
  s(bool,'T_0') = s(bool,sK5(happ(s(fun('type.extreal.extreal','type.extreal.extreal'),'const.combin.I_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))))),
  inference(unit_resulting_resolution,[],[f706,f248])).
fof(f726,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X2))) | s(bool,'F_0') = s(bool,X2)) )),
  inference(superposition,[],[f339,f242])).
fof(f706,plain,(
  p(s(bool,sK5(happ(s(fun('type.extreal.extreal','type.extreal.extreal'),'const.combin.I_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))))),
  inference(unit_resulting_resolution,[],[f339,f312])).
fof(f707,plain,(
  p(s(bool,sK5(happ(s(fun('type.extreal.extreal','type.extreal.extreal'),'const.combin.I_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))))),
  inference(unit_resulting_resolution,[],[f339,f313])).
fof(f597,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f571,f337])).
fof(f571,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f316,f242])).
fof(f545,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,sK5('const.extreal.extreal_logr_2'(s('type.realax.real',X0),s('type.extreal.extreal','const.extreal.NegInf_0'))))) )),
  inference(unit_resulting_resolution,[],[f495,f248])).
fof(f512,plain,(
  ( ! [X4] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X4))) | s(bool,'F_0') = s(bool,X4)) )),
  inference(superposition,[],[f242,f282])).
fof(f511,plain,(
  ( ! [X3] : (s(bool,'T_0') = s(bool,sK2(X3,bool)) | s(bool,'F_0') = s(bool,X3)) )),
  inference(superposition,[],[f242,f290])).
fof(f502,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f242,f242])).
fof(f495,plain,(
  ( ! [X0] : (p(s(bool,sK5('const.extreal.extreal_logr_2'(s('type.realax.real',X0),s('type.extreal.extreal','const.extreal.NegInf_0')))))) )),
  inference(unit_resulting_resolution,[],[f232,f312])).
fof(f428,plain,(
  s(bool,'T_0') = s(bool,sK5('const.combin.I_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))))),
  inference(unit_resulting_resolution,[],[f410,f248])).
fof(f416,plain,(
  s(bool,'T_0') = s(bool,sK5('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))),
  inference(unit_resulting_resolution,[],[f404,f248])).
fof(f412,plain,(
  s(bool,'T_0') = s(bool,sK5(sK2('const.extreal.PosInf_0','type.extreal.extreal')))),
  inference(unit_resulting_resolution,[],[f409,f248])).
fof(f409,plain,(
  p(s(bool,sK5(sK2('const.extreal.PosInf_0','type.extreal.extreal'))))),
  inference(unit_resulting_resolution,[],[f290,f313])).
fof(f410,plain,(
  p(s(bool,sK5('const.combin.I_1'(s('type.extreal.extreal','const.extreal.PosInf_0')))))),
  inference(unit_resulting_resolution,[],[f282,f313])).
fof(f406,plain,(
  s(bool,'T_0') = s(bool,sK5(sK2('const.extreal.NegInf_0','type.extreal.extreal')))),
  inference(unit_resulting_resolution,[],[f403,f248])).
fof(f403,plain,(
  p(s(bool,sK5(sK2('const.extreal.NegInf_0','type.extreal.extreal'))))),
  inference(unit_resulting_resolution,[],[f290,f312])).
fof(f404,plain,(
  p(s(bool,sK5('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0')))))),
  inference(unit_resulting_resolution,[],[f282,f312])).
fof(f365,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.PosInf_0')))) )),
  inference(unit_resulting_resolution,[],[f226,f248])).
fof(f366,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)))) )),
  inference(unit_resulting_resolution,[],[f227,f248])).
fof(f367,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0')))) )),
  inference(unit_resulting_resolution,[],[f230,f248])).
fof(f368,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)))) )),
  inference(unit_resulting_resolution,[],[f228,f248])).
fof(f341,plain,(
  ( ! [X6,X4,X2,X0,X5,X1] : (s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),sK6(X2,X1,X0)),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),'const.pair.pair_CASE_2'(s('type.pair.prod'('type.num.num',X1),'const.pair.,_2'(s('type.num.num',X4),s(X1,X5))),s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),sK6(X2,X1,X0)),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))) )),
  inference(backward_demodulation,[],[f328,f322])).
fof(f340,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),'const.pair.pair_CASE_2'(s('type.pair.prod'(X0,fun('type.num.num',X0)),'const.pair.,_2'(s(X0,X1),s(fun('type.num.num',X0),X2))),s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'))),s('type.num.num',X3)))) )),
  inference(backward_demodulation,[],[f328,f301])).
fof(f328,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X5),s(X1,X4))),s(X2,X3))) = s(X0,'const.pair.pair_CASE_2'(s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X3))),s(fun(X1,fun(X2,X0)),X5)))) )),
  inference(cnf_transformation,[],[f160])).
fof(f160,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X5),s(X1,X4))),s(X2,X3))) = s(X0,'const.pair.pair_CASE_2'(s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X3))),s(fun(X1,fun(X2,X0)),X5)))),
  inference(flattening,[],[f159])).
fof(f159,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X5),s(X1,X4))),s(X2,X3))) = s(X0,'const.pair.pair_CASE_2'(s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X3))),s(fun(X1,fun(X2,X0)),X5)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X7,X13,X14] : ! [X12,X6,X10] : s(X7,happ(s(fun(X14,X7),happ(s(fun(X13,fun(X14,X7)),X10),s(X13,X6))),s(X14,X12))) = s(X7,'const.pair.pair_CASE_2'(s('type.pair.prod'(X13,X14),'const.pair.,_2'(s(X13,X6),s(X14,X12))),s(fun(X13,fun(X14,X7)),X10)))),
  file('thm.extreal.le_infty|split|2','thm.pair.pair_case_thm')).
fof(f327,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))) )),
  inference(cnf_transformation,[],[f158])).
fof(f158,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(flattening,[],[f157])).
fof(f157,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X25,X26,X27] : ! [X28,X29,X30] : s(X25,happ(s(fun(X26,X25),X28),s(X26,happ(s(fun(X27,X26),X29),s(X27,X30))))) = s(X25,happ(s(fun(X27,X25),'_dst_x0x1_2'(s(fun(X26,X25),X28),s(fun(X27,X26),X29))),s(X27,X30)))),
  file('thm.extreal.le_infty|split|2','_dst_x0x1_ax')).
fof(f324,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X3)))) | s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(fun(X1,X0),fun(X1,X0)),X2),s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X4),s(fun(X1,fun(X1,bool)),X3),s(X1,X5))))),s(X1,X5))) | s(fun(X1,X0),X4) != s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X2)))) )),
  inference(cnf_transformation,[],[f172])).
fof(f172,plain,(
  ! [X0,X1,X2,X3,X4] : (s(fun(X1,X0),X4) != s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X2))) | ~p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X3)))) | ! [X5] : s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(fun(X1,X0),fun(X1,X0)),X2),s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X4),s(fun(X1,fun(X1,bool)),X3),s(X1,X5))))),s(X1,X5))))),
  inference(flattening,[],[f171])).
fof(f171,plain,(
  ! [X0,X1,X2,X3,X4] : (s(fun(X1,X0),X4) != s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X2))) | (~p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X3)))) | ! [X5] : s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(fun(X1,X0),fun(X1,X0)),X2),s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X4),s(fun(X1,fun(X1,bool)),X3),s(X1,X5))))),s(X1,X5)))))),
  inference(ennf_transformation,[],[f155])).
fof(f155,plain,(
  ! [X0,X1,X2,X3,X4] : (s(fun(X1,X0),X4) = s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X2))) => (p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X3)))) => ! [X5] : s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(fun(X1,X0),fun(X1,X0)),X2),s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X4),s(fun(X1,fun(X1,bool)),X3),s(X1,X5))))),s(X1,X5)))))),
  inference(flattening,[],[f154])).
fof(f154,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (s(fun(X1,X0),X4) = s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X2))) => (p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X3)))) => ! [X5] : s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(fun(X1,X0),fun(X1,X0)),X2),s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X4),s(fun(X1,fun(X1,bool)),X3),s(X1,X5))))),s(X1,X5)))))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X13,X7] : ! [X17,X18,X10] : (s(fun(X7,X13),X10) = s(fun(X7,X13),'const.relation.WFREC_2'(s(fun(X7,fun(X7,bool)),X18),s(fun(fun(X7,X13),fun(X7,X13)),X17))) => (p(s(bool,'const.relation.WF_1'(s(fun(X7,fun(X7,bool)),X18)))) => ! [X6] : s(X13,happ(s(fun(X7,X13),X10),s(X7,X6))) = s(X13,happ(s(fun(X7,X13),happ(s(fun(fun(X7,X13),fun(X7,X13)),X17),s(fun(X7,X13),'const.relation.RESTRICT_3'(s(fun(X7,X13),X10),s(fun(X7,fun(X7,bool)),X18),s(X7,X6))))),s(X7,X6)))))),
  file('thm.extreal.le_infty|split|2','thm.relation.WFREC_COROLLARY')).
fof(f323,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK7(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK7(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f213])).
fof(f213,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK7(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK7(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7])],[f170])).
fof(f170,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f153])).
fof(f153,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f152])).
fof(f152,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.extreal.le_infty|split|2',aHL_EXT)).
fof(f322,plain,(
  ( ! [X6,X4,X2,X0,X5,X1] : (s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),sK6(X2,X1,X0)),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),sK6(X2,X1,X0)),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))) )),
  inference(cnf_transformation,[],[f212])).
fof(f212,plain,(
  ! [X0,X1,X2] : ! [X4,X5,X6] : s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),sK6(X2,X1,X0)),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),sK6(X2,X1,X0)),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK6])],[f151])).
fof(f151,plain,(
  ! [X0,X1,X2] : ? [X3] : ! [X4,X5,X6] : s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),X3),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),X3),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))),
  inference(flattening,[],[f150])).
fof(f150,plain,(
  ! [X0,X1] : ! [X2] : ? [X3] : ! [X4,X5,X6] : s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),X3),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),X3),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X13,X7] : ! [X34] : ? [X10] : ! [X35,X36,X37] : s(X13,happ(s(fun('type.ind_type.recspace'(X7),X13),X10),s('type.ind_type.recspace'(X7),'const.ind_type.CONSTR_3'(s('type.num.num',X35),s(X7,X36),s(fun('type.num.num','type.ind_type.recspace'(X7)),X37))))) = s(X13,happ(s(fun(fun('type.num.num',X13),X13),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X7)),fun(fun('type.num.num',X13),X13)),happ(s(fun(X7,fun(fun('type.num.num','type.ind_type.recspace'(X7)),fun(fun('type.num.num',X13),X13))),happ(s(fun('type.num.num',fun(X7,fun(fun('type.num.num','type.ind_type.recspace'(X7)),fun(fun('type.num.num',X13),X13)))),X34),s('type.num.num',X35))),s(X7,X36))),s(fun('type.num.num','type.ind_type.recspace'(X7)),X37))),s(fun('type.num.num',X13),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X7),X13),X10),s(fun('type.num.num','type.ind_type.recspace'(X7)),X37)))))),
  file('thm.extreal.le_infty|split|2','thm.ind_type.CONSTR_REC')).
fof(f320,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0))))) )),
  inference(cnf_transformation,[],[f211])).
fof(f211,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0))))) & (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))))),
  inference(nnf_transformation,[],[f149])).
fof(f149,plain,(
  ! [X0,X1] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) <=> ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))))),
  inference(flattening,[],[f148])).
fof(f148,plain,(
  ! [X0,X1] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) <=> ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X6,X12] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X12)))) <=> ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X12),s('type.extreal.extreal',X6)))))),
  file('thm.extreal.le_infty|split|2','thm.extreal.extreal_lt_def')).
fof(f321,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0))))) )),
  inference(cnf_transformation,[],[f211])).
fof(f316,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f169])).
fof(f169,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f168])).
fof(f168,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X22,X23] : ((p(s(bool,X22)) => p(s(bool,X23))) => ((p(s(bool,X23)) => p(s(bool,X22))) => s(bool,X22) = s(bool,X23)))),
  file('thm.extreal.le_infty|split|2','thm.bool.IMP_ANTISYM_AX')).
fof(f319,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f169])).
fof(f315,plain,(
  ( ! [X0,X1] : (s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))))) )),
  inference(cnf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0,X1] : s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X12,X6] : s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X12))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X6))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X12)))))),
  file('thm.extreal.le_infty|split|2','thm.extreal.extreal_le_def|split|0')).
fof(f314,plain,(
  ( ! [X0] : (s('type.realax.real','const.extreal.real_1'(s('type.extreal.extreal',X0))) = s('type.realax.real','const.bool.COND_3'(s(bool,sK5(X0)),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.min.@_1'(s(fun('type.realax.real',bool),'_dst_x0x1_2'(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal',X0))),s(fun('type.realax.real','type.extreal.extreal'),'const.extreal.Normal_0')))))))) )),
  inference(cnf_transformation,[],[f210])).
fof(f303,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s('type.extreal.extreal',X4) != s('type.extreal.extreal',X5) | s('type.extreal.extreal',X5) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK4(X8,X7,X6,X5,X4,X3,X2,X1,X0)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X5) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X5) | s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal',X4),s(X0,X6),s(X0,X7),s(fun('type.realax.real',X0),X8))) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal',X5),s(X0,X2),s(X0,X1),s(fun('type.realax.real',X0),X3)))) )),
  inference(cnf_transformation,[],[f207])).
fof(f207,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : (s('type.extreal.extreal',X4) != s('type.extreal.extreal',X5) | (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X5) & s(X0,X2) != s(X0,X6)) | (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X5) & s(X0,X1) != s(X0,X7)) | (s('type.extreal.extreal',X5) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK4(X8,X7,X6,X5,X4,X3,X2,X1,X0)))) & s(X0,happ(s(fun('type.realax.real',X0),X3),s('type.realax.real',sK4(X8,X7,X6,X5,X4,X3,X2,X1,X0)))) != s(X0,happ(s(fun('type.realax.real',X0),X8),s('type.realax.real',sK4(X8,X7,X6,X5,X4,X3,X2,X1,X0))))) | s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal',X4),s(X0,X6),s(X0,X7),s(fun('type.realax.real',X0),X8))) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal',X5),s(X0,X2),s(X0,X1),s(fun('type.realax.real',X0),X3))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f167])).
fof(f167,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : (s('type.extreal.extreal',X4) != s('type.extreal.extreal',X5) | (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X5) & s(X0,X2) != s(X0,X6)) | (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X5) & s(X0,X1) != s(X0,X7)) | ? [X9] : (s('type.extreal.extreal',X5) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X9))) & s(X0,happ(s(fun('type.realax.real',X0),X3),s('type.realax.real',X9))) != s(X0,happ(s(fun('type.realax.real',X0),X8),s('type.realax.real',X9)))) | s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal',X4),s(X0,X6),s(X0,X7),s(fun('type.realax.real',X0),X8))) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal',X5),s(X0,X2),s(X0,X1),s(fun('type.realax.real',X0),X3))))),
  inference(flattening,[],[f166])).
fof(f166,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : ((s('type.extreal.extreal',X4) != s('type.extreal.extreal',X5) | (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X5) & s(X0,X2) != s(X0,X6)) | (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X5) & s(X0,X1) != s(X0,X7)) | ? [X9] : (s('type.extreal.extreal',X5) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X9))) & s(X0,happ(s(fun('type.realax.real',X0),X3),s('type.realax.real',X9))) != s(X0,happ(s(fun('type.realax.real',X0),X8),s('type.realax.real',X9))))) | s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal',X4),s(X0,X6),s(X0,X7),s(fun('type.realax.real',X0),X8))) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal',X5),s(X0,X2),s(X0,X1),s(fun('type.realax.real',X0),X3))))),
  inference(ennf_transformation,[],[f144])).
fof(f144,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : ((s('type.extreal.extreal',X4) = s('type.extreal.extreal',X5) & (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X5) => s(X0,X2) = s(X0,X6)) & (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X5) => s(X0,X1) = s(X0,X7)) & ! [X9] : (s('type.extreal.extreal',X5) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X9))) => s(X0,happ(s(fun('type.realax.real',X0),X3),s('type.realax.real',X9))) = s(X0,happ(s(fun('type.realax.real',X0),X8),s('type.realax.real',X9))))) => s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal',X4),s(X0,X6),s(X0,X7),s(fun('type.realax.real',X0),X8))) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal',X5),s(X0,X2),s(X0,X1),s(fun('type.realax.real',X0),X3))))),
  inference(flattening,[],[f143])).
fof(f143,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8] : ((s('type.extreal.extreal',X4) = s('type.extreal.extreal',X5) & (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X5) => s(X0,X2) = s(X0,X6)) & (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X5) => s(X0,X1) = s(X0,X7)) & ! [X9] : (s('type.extreal.extreal',X5) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X9))) => s(X0,happ(s(fun('type.realax.real',X0),X3),s('type.realax.real',X9))) = s(X0,happ(s(fun('type.realax.real',X0),X8),s('type.realax.real',X9))))) => s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal',X4),s(X0,X6),s(X0,X7),s(fun('type.realax.real',X0),X8))) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal',X5),s(X0,X2),s(X0,X1),s(fun('type.realax.real',X0),X3))))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X7] : ! [X40,X41,X42,X17,X43,X8,X9,X10] : ((s('type.extreal.extreal',X17) = s('type.extreal.extreal',X43) & (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X43) => s(X7,X8) = s(X7,X41)) & (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X43) => s(X7,X9) = s(X7,X40)) & ! [X11] : (s('type.extreal.extreal',X43) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X11))) => s(X7,happ(s(fun('type.realax.real',X7),X10),s('type.realax.real',X11))) = s(X7,happ(s(fun('type.realax.real',X7),X42),s('type.realax.real',X11))))) => s(X7,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal',X17),s(X7,X8),s(X7,X9),s(fun('type.realax.real',X7),X10))) = s(X7,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal',X43),s(X7,X41),s(X7,X40),s(fun('type.realax.real',X7),X42))))),
  file('thm.extreal.le_infty|split|2','thm.extreal.extreal_case_cong')).
fof(f304,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s('type.extreal.extreal',X4) != s('type.extreal.extreal',X5) | s('type.extreal.extreal',X5) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK4(X8,X7,X6,X5,X4,X3,X2,X1,X0)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X5) | s(X0,X2) != s(X0,X6) | s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal',X4),s(X0,X6),s(X0,X7),s(fun('type.realax.real',X0),X8))) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal',X5),s(X0,X2),s(X0,X1),s(fun('type.realax.real',X0),X3)))) )),
  inference(cnf_transformation,[],[f207])).
fof(f305,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s('type.extreal.extreal',X4) != s('type.extreal.extreal',X5) | s('type.extreal.extreal',X5) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK4(X8,X7,X6,X5,X4,X3,X2,X1,X0)))) | s(X0,X1) != s(X0,X7) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X5) | s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal',X4),s(X0,X6),s(X0,X7),s(fun('type.realax.real',X0),X8))) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal',X5),s(X0,X2),s(X0,X1),s(fun('type.realax.real',X0),X3)))) )),
  inference(cnf_transformation,[],[f207])).
fof(f306,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s('type.extreal.extreal',X4) != s('type.extreal.extreal',X5) | s('type.extreal.extreal',X5) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK4(X8,X7,X6,X5,X4,X3,X2,X1,X0)))) | s(X0,X1) != s(X0,X7) | s(X0,X2) != s(X0,X6) | s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal',X4),s(X0,X6),s(X0,X7),s(fun('type.realax.real',X0),X8))) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal',X5),s(X0,X2),s(X0,X1),s(fun('type.realax.real',X0),X3)))) )),
  inference(cnf_transformation,[],[f207])).
fof(f307,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(X0,happ(s(fun('type.realax.real',X0),X3),s('type.realax.real',sK4(X8,X7,X6,X5,X4,X3,X2,X1,X0)))) != s(X0,happ(s(fun('type.realax.real',X0),X8),s('type.realax.real',sK4(X8,X7,X6,X5,X4,X3,X2,X1,X0)))) | s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal',X4),s(X0,X6),s(X0,X7),s(fun('type.realax.real',X0),X8))) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal',X5),s(X0,X2),s(X0,X1),s(fun('type.realax.real',X0),X3))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X5) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X5) | s('type.extreal.extreal',X4) != s('type.extreal.extreal',X5)) )),
  inference(cnf_transformation,[],[f207])).
fof(f308,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(X0,happ(s(fun('type.realax.real',X0),X3),s('type.realax.real',sK4(X8,X7,X6,X5,X4,X3,X2,X1,X0)))) != s(X0,happ(s(fun('type.realax.real',X0),X8),s('type.realax.real',sK4(X8,X7,X6,X5,X4,X3,X2,X1,X0)))) | s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal',X4),s(X0,X6),s(X0,X7),s(fun('type.realax.real',X0),X8))) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal',X5),s(X0,X2),s(X0,X1),s(fun('type.realax.real',X0),X3))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X5) | s(X0,X2) != s(X0,X6) | s('type.extreal.extreal',X4) != s('type.extreal.extreal',X5)) )),
  inference(cnf_transformation,[],[f207])).
fof(f309,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(X0,happ(s(fun('type.realax.real',X0),X3),s('type.realax.real',sK4(X8,X7,X6,X5,X4,X3,X2,X1,X0)))) != s(X0,happ(s(fun('type.realax.real',X0),X8),s('type.realax.real',sK4(X8,X7,X6,X5,X4,X3,X2,X1,X0)))) | s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal',X4),s(X0,X6),s(X0,X7),s(fun('type.realax.real',X0),X8))) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal',X5),s(X0,X2),s(X0,X1),s(fun('type.realax.real',X0),X3))) | s(X0,X1) != s(X0,X7) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X5) | s('type.extreal.extreal',X4) != s('type.extreal.extreal',X5)) )),
  inference(cnf_transformation,[],[f207])).
fof(f310,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(X0,happ(s(fun('type.realax.real',X0),X3),s('type.realax.real',sK4(X8,X7,X6,X5,X4,X3,X2,X1,X0)))) != s(X0,happ(s(fun('type.realax.real',X0),X8),s('type.realax.real',sK4(X8,X7,X6,X5,X4,X3,X2,X1,X0)))) | s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal',X4),s(X0,X6),s(X0,X7),s(fun('type.realax.real',X0),X8))) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal',X5),s(X0,X2),s(X0,X1),s(fun('type.realax.real',X0),X3))) | s(X0,X1) != s(X0,X7) | s(X0,X2) != s(X0,X6) | s('type.extreal.extreal',X4) != s('type.extreal.extreal',X5)) )),
  inference(cnf_transformation,[],[f207])).
fof(f302,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,happ(s(fun('type.realax.real',X0),X4),s('type.realax.real',X1))) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1))),s(X0,X2),s(X0,X3),s(fun('type.realax.real',X0),X4)))) )),
  inference(cnf_transformation,[],[f142])).
fof(f142,plain,(
  ! [X0,X1,X2,X3,X4] : s(X0,happ(s(fun('type.realax.real',X0),X4),s('type.realax.real',X1))) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1))),s(X0,X2),s(X0,X3),s(fun('type.realax.real',X0),X4)))),
  inference(flattening,[],[f141])).
fof(f141,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s(X0,happ(s(fun('type.realax.real',X0),X4),s('type.realax.real',X1))) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1))),s(X0,X2),s(X0,X3),s(fun('type.realax.real',X0),X4)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X7] : ! [X11,X8,X9,X10] : s(X7,happ(s(fun('type.realax.real',X7),X10),s('type.realax.real',X11))) = s(X7,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X11))),s(X7,X8),s(X7,X9),s(fun('type.realax.real',X7),X10)))),
  file('thm.extreal.le_infty|split|2','thm.extreal.extreal_case_def|split|2')).
fof(f301,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))) )),
  inference(cnf_transformation,[],[f140])).
fof(f140,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))),
  inference(flattening,[],[f139])).
fof(f139,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X7] : ! [X16,X32,X33] : s(X7,'const.ind_type.FCONS_3'(s(X7,X16),s(fun('type.num.num',X7),X32),s('type.num.num',X33))) = s(X7,happ(s(fun('type.num.num',X7),happ(s(fun(fun('type.num.num',X7),fun('type.num.num',X7)),happ(s(fun(X7,fun(fun('type.num.num',X7),fun('type.num.num',X7))),'const.ind_type.FCONS_0'),s(X7,X16))),s(fun('type.num.num',X7),X32))),s('type.num.num',X33)))),
  file('thm.extreal.le_infty|split|2','ARITY_0_3_const.ind_type.FCONS')).
fof(f300,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))) )),
  inference(cnf_transformation,[],[f138])).
fof(f138,plain,(
  ! [X0,X1,X2,X3] : s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))),
  inference(flattening,[],[f137])).
fof(f137,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X7] : ! [X16,X32,X33] : s('type.ind_type.recspace'(X7),'const.ind_type.CONSTR_3'(s('type.num.num',X16),s(X7,X32),s(fun('type.num.num','type.ind_type.recspace'(X7)),X33))) = s('type.ind_type.recspace'(X7),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X7)),'type.ind_type.recspace'(X7)),happ(s(fun(X7,fun(fun('type.num.num','type.ind_type.recspace'(X7)),'type.ind_type.recspace'(X7))),'const.ind_type.CONSTR_1'(s('type.num.num',X16))),s(X7,X32))),s(fun('type.num.num','type.ind_type.recspace'(X7)),X33)))),
  file('thm.extreal.le_infty|split|2','ARITY_1_3_const.ind_type.CONSTR')).
fof(f299,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))) )),
  inference(cnf_transformation,[],[f136])).
fof(f136,plain,(
  ! [X0,X1,X2,X3] : s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))),
  inference(flattening,[],[f135])).
fof(f135,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X7] : ! [X11,X10,X31] : s(X7,happ(s(fun('type.num.num',X7),X10),s('type.num.num',X31))) = s(X7,'const.ind_type.FCONS_3'(s(X7,X11),s(fun('type.num.num',X7),X10),s('type.num.num','const.num.SUC_1'(s('type.num.num',X31)))))),
  file('thm.extreal.le_infty|split|2','thm.ind_type.FCONS|split|1')).
fof(f298,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X2) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.PosInf_0'),s(X0,X1),s(X0,X2),s(fun('type.realax.real',X0),X3)))) )),
  inference(cnf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0,X1,X2,X3] : s(X0,X2) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.PosInf_0'),s(X0,X1),s(X0,X2),s(fun('type.realax.real',X0),X3)))),
  inference(flattening,[],[f133])).
fof(f133,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,X2) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.PosInf_0'),s(X0,X1),s(X0,X2),s(fun('type.realax.real',X0),X3)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X7] : ! [X8,X9,X10] : s(X7,X9) = s(X7,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.PosInf_0'),s(X7,X8),s(X7,X9),s(fun('type.realax.real',X7),X10)))),
  file('thm.extreal.le_infty|split|2','thm.extreal.extreal_case_def|split|1')).
fof(f297,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X1) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(X0,X1),s(X0,X2),s(fun('type.realax.real',X0),X3)))) )),
  inference(cnf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0,X1,X2,X3] : s(X0,X1) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(X0,X1),s(X0,X2),s(fun('type.realax.real',X0),X3)))),
  inference(flattening,[],[f131])).
fof(f131,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,X1) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(X0,X1),s(X0,X2),s(fun('type.realax.real',X0),X3)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X7] : ! [X8,X9,X10] : s(X7,X8) = s(X7,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(X7,X8),s(X7,X9),s(fun('type.realax.real',X7),X10)))),
  file('thm.extreal.le_infty|split|2','thm.extreal.extreal_case_def|split|0')).
fof(f292,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1))))))) )),
  inference(cnf_transformation,[],[f164])).
fof(f164,plain,(
  ! [X0,X1,X2] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(ennf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(flattening,[],[f125])).
fof(f125,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X7] : ! [X15,X6] : (p(s(bool,happ(s(fun(X7,bool),X15),s(X7,X6)))) => p(s(bool,happ(s(fun(X7,bool),X15),s(X7,'const.min.@_1'(s(fun(X7,bool),X15)))))))),
  file('thm.extreal.le_infty|split|2','thm.bool.SELECT_AX')).
fof(f291,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(flattening,[],[f123])).
fof(f123,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X7] : ! [X11,X10] : s(X7,X11) = s(X7,'const.ind_type.FCONS_3'(s(X7,X11),s(fun('type.num.num',X7),X10),s('type.num.num','const.num.|48|_0')))),
  file('thm.extreal.le_infty|split|2','thm.ind_type.FCONS|split|0')).
fof(f339,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))) )),
  inference(forward_demodulation,[],[f283,f282])).
fof(f283,plain,(
  ( ! [X0,X1] : (s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f116])).
fof(f116,plain,(
  ! [X0,X1] : s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))),
  inference(flattening,[],[f115])).
fof(f115,plain,(
  ! [X0] : ! [X1] : s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X7] : ! [X16] : s(X7,'const.combin.I_1'(s(X7,X16))) = s(X7,happ(s(fun(X7,X7),'const.combin.I_0'),s(X7,X16)))),
  file('thm.extreal.le_infty|split|2','ARITY_0_1_const.combin.I')).
fof(f281,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f112])).
fof(f112,plain,(
  ! [X0] : s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0)))),
  inference(flattening,[],[f111])).
fof(f111,plain,(
  ! [X0] : ~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0)))),
  inference(rectify,[],[f48])).
fof(f48,axiom,(
  ! [X31] : ~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X31)))),
  file('thm.extreal.le_infty|split|2','thm.extreal.num_not_infty|split|0')).
fof(f251,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f180])).
fof(f180,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f98])).
fof(f98,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f97])).
fof(f97,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X24] : (s(bool,'F_0') = s(bool,X24) <=> ~p(s(bool,X24)))),
  file('thm.extreal.le_infty|split|2','thm.bool.EQ_CLAUSES|split|2')).
fof(f252,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f180])).
fof(f241,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))) = s('type.extreal.extreal',happ(s(fun('type.realax.real','type.extreal.extreal'),'const.extreal.Normal_0'),s('type.realax.real',X0)))) )),
  inference(cnf_transformation,[],[f90])).
fof(f90,plain,(
  ! [X0] : s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))) = s('type.extreal.extreal',happ(s(fun('type.realax.real','type.extreal.extreal'),'const.extreal.Normal_0'),s('type.realax.real',X0)))),
  inference(rectify,[],[f62])).
fof(f62,axiom,(
  ! [X16] : s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X16))) = s('type.extreal.extreal',happ(s(fun('type.realax.real','type.extreal.extreal'),'const.extreal.Normal_0'),s('type.realax.real',X16)))),
  file('thm.extreal.le_infty|split|2','ARITY_0_1_const.extreal.Normal')).
fof(f240,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))))) )),
  inference(cnf_transformation,[],[f89])).
fof(f89,plain,(
  ! [X0] : s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X21] : s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X21)))))),
  file('thm.extreal.le_infty|split|2','thm.extreal.extreal_le_def|split|6')).
fof(f238,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0')))) )),
  inference(cnf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0')))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X20] : s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X20))),s('type.extreal.extreal','const.extreal.PosInf_0')))),
  file('thm.extreal.le_infty|split|2','thm.extreal.extreal_le_def|split|3')).
fof(f231,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)))) )),
  inference(cnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X11] : s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X11)))),
  file('thm.extreal.le_infty|split|2','thm.extreal.extreal_le_def|split|1')).
fof(f230,plain,(
  ( ! [X0] : (p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'))))) )),
  inference(cnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0] : p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'))))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X7] : p(s(bool,'const.relation.WF_1'(s(fun(X7,fun(X7,bool)),'const.relation.EMPTY_REL_0'))))),
  file('thm.extreal.le_infty|split|2','thm.relation.WF_EMPTY_REL')).
fof(f229,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0))))) )),
  inference(cnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0] : p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0))))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X6] : p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X6))))),
  file('thm.extreal.le_infty|split|2','thm.extreal.le_infty|split|0')).
fof(f228,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0))))) )),
  inference(cnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0] : p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0))))),
  inference(rectify,[],[f64])).
fof(f64,axiom,(
  ! [X6] : p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X6))))),
  file('thm.extreal.le_infty|split|2','thm.real.REAL_LE_REFL')).
fof(f227,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0))))) )),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0] : p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0))))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X6] : p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X6))))),
  file('thm.extreal.le_infty|split|2','thm.extreal.le_refl')).
fof(f226,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.PosInf_0'))))) )),
  inference(cnf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0] : p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.PosInf_0'))))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X6] : p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal','const.extreal.PosInf_0'))))),
  file('thm.extreal.le_infty|split|2','thm.extreal.le_infty|split|1')).
fof(f220,plain,(
  s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))),
  inference(cnf_transformation,[],[f27])).
fof(f27,axiom,(
  s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))),
  file('thm.extreal.le_infty|split|2','thm.extreal.extreal_le_def|split|2')).
fof(f4536,plain,(
  $false | ($spl3 | $spl18)),
  inference(subsumption_resolution,[],[f4535,f354])).
fof(f4535,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl3 | $spl18)),
  inference(forward_demodulation,[],[f4500,f3433])).
fof(f4500,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.NegInf_0'))) | ($spl3 | $spl18)),
  inference(backward_demodulation,[],[f4445,f2174])).
fof(f2174,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',sK0))) | $spl3),
  inference(unit_resulting_resolution,[],[f1969,f252])).
fof(f4534,plain,(
  $false | ($spl3 | $spl18)),
  inference(subsumption_resolution,[],[f4499,f3444])).
fof(f4499,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.NegInf_0'))) | ($spl3 | $spl18)),
  inference(backward_demodulation,[],[f4445,f2183])).
fof(f2183,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal',sK0))),s('type.extreal.extreal','const.extreal.NegInf_0'))) | $spl3),
  inference(unit_resulting_resolution,[],[f1978,f247])).
fof(f1978,plain,(
  ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal',sK0))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl3),
  inference(unit_resulting_resolution,[],[f336,f325])).
fof(f4533,plain,(
  $false | ($spl3 | $spl18)),
  inference(subsumption_resolution,[],[f4532,f354])).
fof(f4532,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl3 | $spl18)),
  inference(forward_demodulation,[],[f4498,f3433])).
fof(f4498,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.NegInf_0'))) | ($spl3 | $spl18)),
  inference(backward_demodulation,[],[f4445,f2182])).
fof(f2182,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal',sK0))),s('type.extreal.extreal','const.extreal.NegInf_0'))) | $spl3),
  inference(unit_resulting_resolution,[],[f1978,f252])).
fof(f4531,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f4530,f217])).
fof(f4530,plain,(
  ~p(s(bool,'T_0')) | $spl18),
  inference(forward_demodulation,[],[f4493,f3433])).
fof(f4493,plain,(
  ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl18),
  inference(backward_demodulation,[],[f4445,f1979])).
fof(f4529,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f4528,f217])).
fof(f4528,plain,(
  ~p(s(bool,'T_0')) | $spl18),
  inference(forward_demodulation,[],[f4492,f3433])).
fof(f4492,plain,(
  ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl18),
  inference(backward_demodulation,[],[f4445,f1970])).
fof(f4527,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f4491,f231])).
fof(f4491,plain,(
  s(bool,'T_0') != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | $spl18),
  inference(backward_demodulation,[],[f4445,f1258])).
fof(f4507,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f4455,f231])).
fof(f4455,plain,(
  s(bool,'T_0') != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | $spl18),
  inference(backward_demodulation,[],[f4445,f950])).
fof(f4436,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f4344,f231])).
fof(f4344,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X1)))))) ) | $spl18),
  inference(backward_demodulation,[],[f4280,f1323])).
fof(f4435,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f4339,f231])).
fof(f4339,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))))) ) | $spl18),
  inference(backward_demodulation,[],[f4280,f1080])).
fof(f4434,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f4433,f354])).
fof(f4433,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl18),
  inference(forward_demodulation,[],[f4338,f231])).
fof(f4338,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0)))))) ) | $spl18),
  inference(backward_demodulation,[],[f4280,f1069])).
fof(f4432,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f4328,f231])).
fof(f4328,plain,(
  s(bool,'T_0') != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | $spl18),
  inference(backward_demodulation,[],[f4280,f426])).
fof(f4431,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f4429,f354])).
fof(f4429,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl18),
  inference(backward_demodulation,[],[f4327,f939])).
fof(f4430,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f4428,f354])).
fof(f4428,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl18),
  inference(backward_demodulation,[],[f4327,f237])).
fof(f4427,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f4426,f354])).
fof(f4426,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl18),
  inference(forward_demodulation,[],[f4318,f231])).
fof(f4318,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))))) ) | $spl18),
  inference(backward_demodulation,[],[f4280,f240])).
fof(f4425,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f4316,f950])).
fof(f4316,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl18),
  inference(backward_demodulation,[],[f4280,f238])).
fof(f4424,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f4423,f354])).
fof(f4423,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl18),
  inference(forward_demodulation,[],[f4314,f231])).
fof(f4314,plain,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | $spl18),
  inference(backward_demodulation,[],[f4280,f219])).
fof(f4309,plain,(
  $false | ($spl7 | $spl18)),
  inference(subsumption_resolution,[],[f4280,f821])).
fof(f4303,plain,(
  $false | ($spl7 | $spl18)),
  inference(subsumption_resolution,[],[f4268,f821])).
fof(f4268,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.PosInf_0') | $spl18),
  inference(superposition,[],[f3988,f282])).
fof(f3970,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK2(sK2('const.combin.I_1'(s('type.extreal.extreal','const.extreal.PosInf_0')),'type.extreal.extreal'),'type.extreal.extreal')) | $spl28),
  inference(cnf_transformation,[],[f3970_D])).
fof(f3970_D,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK2(sK2('const.combin.I_1'(s('type.extreal.extreal','const.extreal.PosInf_0')),'type.extreal.extreal'),'type.extreal.extreal')) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f3972,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK2(sK2('const.combin.I_1'(s('type.extreal.extreal','const.extreal.PosInf_0')),'type.extreal.extreal'),'type.extreal.extreal')) | $spl30),
  inference(cnf_transformation,[],[f3972_D])).
fof(f3972_D,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK2(sK2('const.combin.I_1'(s('type.extreal.extreal','const.extreal.PosInf_0')),'type.extreal.extreal'),'type.extreal.extreal')) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f3968,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK2(sK2('const.combin.I_1'(s('type.extreal.extreal','const.extreal.PosInf_0')),'type.extreal.extreal'),'type.extreal.extreal')) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK2(sK2('const.combin.I_1'(s('type.extreal.extreal','const.extreal.PosInf_0')),'type.extreal.extreal'),'type.extreal.extreal')) | $spl16),
  inference(subsumption_resolution,[],[f3934,f217])).
fof(f3934,plain,(
  ~p(s(bool,'T_0')) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK2(sK2('const.combin.I_1'(s('type.extreal.extreal','const.extreal.PosInf_0')),'type.extreal.extreal'),'type.extreal.extreal')) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK2(sK2('const.combin.I_1'(s('type.extreal.extreal','const.extreal.PosInf_0')),'type.extreal.extreal'),'type.extreal.extreal')) | $spl16),
  inference(superposition,[],[f311,f3127])).
fof(f3127,plain,(
  s(bool,'T_0') = s(bool,sK5(sK2(sK2('const.combin.I_1'(s('type.extreal.extreal','const.extreal.PosInf_0')),'type.extreal.extreal'),'type.extreal.extreal'))) | $spl16),
  inference(unit_resulting_resolution,[],[f2837,f804])).
fof(f3216,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK2(sK2(sK2('const.extreal.PosInf_0','type.extreal.extreal'),'type.extreal.extreal'),'type.extreal.extreal')) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK2(sK2(sK2('const.extreal.PosInf_0','type.extreal.extreal'),'type.extreal.extreal'),'type.extreal.extreal')) | $spl8),
  inference(subsumption_resolution,[],[f3189,f217])).
fof(f3189,plain,(
  ~p(s(bool,'T_0')) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK2(sK2(sK2('const.extreal.PosInf_0','type.extreal.extreal'),'type.extreal.extreal'),'type.extreal.extreal')) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK2(sK2(sK2('const.extreal.PosInf_0','type.extreal.extreal'),'type.extreal.extreal'),'type.extreal.extreal')) | $spl8),
  inference(superposition,[],[f311,f2673])).
fof(f3073,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f2920,f2994])).
fof(f2994,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) != s(X0,X2)) ) | $spl20),
  inference(unit_resulting_resolution,[],[f2916,f326])).
fof(f2920,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl20),
  inference(unit_resulting_resolution,[],[f337,f2916,f316])).
fof(f3072,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f2921,f2994])).
fof(f2921,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1)) ) | $spl20),
  inference(unit_resulting_resolution,[],[f2916,f2916,f316])).
fof(f3071,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f2938,f337])).
fof(f2938,plain,(
  p(s(bool,'F_0')) | $spl20),
  inference(unit_resulting_resolution,[],[f354,f2916,f316])).
fof(f3070,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f2940,f2916])).
fof(f2940,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl20),
  inference(unit_resulting_resolution,[],[f426,f2916,f316])).
fof(f3069,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f2942,f2916])).
fof(f2942,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.NegInf_0'))))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f950,f2916,f316])).
fof(f3068,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f2944,f2916])).
fof(f2944,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0))),s('type.extreal.extreal','const.extreal.NegInf_0'))))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f1258,f2916,f316])).
fof(f3067,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f2946,f2916])).
fof(f2946,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))))))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f1080,f2916,f316])).
fof(f3066,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f2948,f2916])).
fof(f2948,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0))))))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f1323,f2916,f316])).
fof(f3065,plain,(
  $false | ($spl3 | $spl20)),
  inference(subsumption_resolution,[],[f2950,f2916])).
fof(f2950,plain,(
  p(s(bool,happ(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',sK0)))) | ($spl3 | $spl20)),
  inference(unit_resulting_resolution,[],[f2175,f2916,f316])).
fof(f3064,plain,(
  $false | ($spl3 | $spl20)),
  inference(subsumption_resolution,[],[f2952,f2916])).
fof(f2952,plain,(
  p(s(bool,happ(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal',sK0))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | ($spl3 | $spl20)),
  inference(unit_resulting_resolution,[],[f2183,f2916,f316])).
fof(f3063,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f2954,f2994])).
fof(f2954,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl20),
  inference(unit_resulting_resolution,[],[f337,f2916,f316])).
fof(f3062,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f2955,f2994])).
fof(f2955,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1)) ) | $spl20),
  inference(unit_resulting_resolution,[],[f2916,f2916,f316])).
fof(f3061,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f2972,f337])).
fof(f2972,plain,(
  p(s(bool,'F_0')) | $spl20),
  inference(unit_resulting_resolution,[],[f354,f2916,f316])).
fof(f3060,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f2974,f2916])).
fof(f2974,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl20),
  inference(unit_resulting_resolution,[],[f426,f2916,f316])).
fof(f3059,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f2976,f2916])).
fof(f2976,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.NegInf_0'))))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f950,f2916,f316])).
fof(f3058,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f2978,f2916])).
fof(f2978,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0))),s('type.extreal.extreal','const.extreal.NegInf_0'))))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f1258,f2916,f316])).
fof(f3057,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f2980,f2916])).
fof(f2980,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))))))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f1080,f2916,f316])).
fof(f3056,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f2982,f2916])).
fof(f2982,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0))))))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f1323,f2916,f316])).
fof(f3055,plain,(
  $false | ($spl3 | $spl20)),
  inference(subsumption_resolution,[],[f2984,f2916])).
fof(f2984,plain,(
  p(s(bool,happ(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',sK0)))) | ($spl3 | $spl20)),
  inference(unit_resulting_resolution,[],[f2175,f2916,f316])).
fof(f3054,plain,(
  $false | ($spl3 | $spl20)),
  inference(subsumption_resolution,[],[f2986,f2916])).
fof(f2986,plain,(
  p(s(bool,happ(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal',sK0))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | ($spl3 | $spl20)),
  inference(unit_resulting_resolution,[],[f2183,f2916,f316])).
fof(f3053,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f2988,f2994])).
fof(f2988,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl20),
  inference(unit_resulting_resolution,[],[f2916,f252])).
fof(f3052,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f2992,f2916])).
fof(f2992,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f2916,f321])).
fof(f3051,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f2997,f2916])).
fof(f2997,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f2916,f321])).
fof(f3004,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f236,f2916,f312])).
fof(f3002,plain,(
  $false | ($spl8 | $spl20)),
  inference(unit_resulting_resolution,[],[f1962,f2916,f313])).
fof(f2998,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f2916,f2916,f321])).
fof(f2995,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f297,f2916,f326])).
fof(f2993,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f2916,f2916,f321])).
fof(f2991,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f231,f2916,f247])).
fof(f2989,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f354,f2916,f252])).
fof(f2987,plain,(
  $false | ($spl3 | $spl20)),
  inference(unit_resulting_resolution,[],[f2916,f2183,f2916,f316])).
fof(f2985,plain,(
  $false | ($spl3 | $spl20)),
  inference(unit_resulting_resolution,[],[f2916,f2175,f2916,f316])).
fof(f2983,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f2916,f1323,f2916,f316])).
fof(f2981,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f2916,f1080,f2916,f316])).
fof(f2979,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f2916,f1258,f2916,f316])).
fof(f2977,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f2916,f950,f2916,f316])).
fof(f2975,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f2916,f426,f2916,f316])).
fof(f2973,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f337,f354,f2916,f316])).
fof(f2971,plain,(
  $false | ($spl3 | $spl20)),
  inference(unit_resulting_resolution,[],[f2916,f2183,f2916,f316])).
fof(f2969,plain,(
  $false | ($spl3 | $spl20)),
  inference(unit_resulting_resolution,[],[f2916,f2175,f2916,f316])).
fof(f2967,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f2916,f1323,f2916,f316])).
fof(f2965,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f2916,f1080,f2916,f316])).
fof(f2963,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f2916,f1258,f2916,f316])).
fof(f2961,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f2916,f950,f2916,f316])).
fof(f2959,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f2916,f426,f2916,f316])).
fof(f2957,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f2916,f354,f2916,f316])).
fof(f2953,plain,(
  $false | ($spl3 | $spl20)),
  inference(unit_resulting_resolution,[],[f2916,f2183,f2916,f316])).
fof(f2951,plain,(
  $false | ($spl3 | $spl20)),
  inference(unit_resulting_resolution,[],[f2916,f2175,f2916,f316])).
fof(f2949,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f2916,f1323,f2916,f316])).
fof(f2947,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f2916,f1080,f2916,f316])).
fof(f2945,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f2916,f1258,f2916,f316])).
fof(f2943,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f2916,f950,f2916,f316])).
fof(f2941,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f2916,f426,f2916,f316])).
fof(f2939,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f337,f354,f2916,f316])).
fof(f2937,plain,(
  $false | ($spl3 | $spl20)),
  inference(unit_resulting_resolution,[],[f2916,f2183,f2916,f316])).
fof(f2935,plain,(
  $false | ($spl3 | $spl20)),
  inference(unit_resulting_resolution,[],[f2916,f2175,f2916,f316])).
fof(f2933,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f2916,f1323,f2916,f316])).
fof(f2931,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f2916,f1080,f2916,f316])).
fof(f2929,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f2916,f1258,f2916,f316])).
fof(f2927,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f2916,f950,f2916,f316])).
fof(f2925,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f2916,f426,f2916,f316])).
fof(f2923,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f2916,f354,f2916,f316])).
fof(f2919,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f217,f2916])).
fof(f2916,plain,(
  ( ! [X3] : (~p(s(bool,X3))) ) | $spl20),
  inference(cnf_transformation,[],[f2916_D])).
fof(f2916_D,plain,(
  ( ! [X3] : (~p(s(bool,X3))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f2918,plain,(
  p(s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0')))) | $spl22),
  inference(cnf_transformation,[],[f2918_D])).
fof(f2918_D,plain,(
  p(s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0')))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f2915,plain,(
  ( ! [X3] : (p(s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0')))) | ~p(s(bool,X3))) )),
  inference(forward_demodulation,[],[f2908,f339])).
fof(f2908,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | p(s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))))))) )),
  inference(superposition,[],[f292,f339])).
fof(f2837,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK2('const.combin.I_1'(s('type.extreal.extreal','const.extreal.PosInf_0')),'type.extreal.extreal')) | $spl16),
  inference(cnf_transformation,[],[f2837_D])).
fof(f2837_D,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK2('const.combin.I_1'(s('type.extreal.extreal','const.extreal.PosInf_0')),'type.extreal.extreal')) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f2839,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK2('const.combin.I_1'(s('type.extreal.extreal','const.extreal.PosInf_0')),'type.extreal.extreal')) | $spl18),
  inference(cnf_transformation,[],[f2839_D])).
fof(f2839_D,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK2('const.combin.I_1'(s('type.extreal.extreal','const.extreal.PosInf_0')),'type.extreal.extreal')) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f2835,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK2('const.combin.I_1'(s('type.extreal.extreal','const.extreal.PosInf_0')),'type.extreal.extreal')) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK2('const.combin.I_1'(s('type.extreal.extreal','const.extreal.PosInf_0')),'type.extreal.extreal'))),
  inference(subsumption_resolution,[],[f2809,f217])).
fof(f2809,plain,(
  ~p(s(bool,'T_0')) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK2('const.combin.I_1'(s('type.extreal.extreal','const.extreal.PosInf_0')),'type.extreal.extreal')) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK2('const.combin.I_1'(s('type.extreal.extreal','const.extreal.PosInf_0')),'type.extreal.extreal'))),
  inference(superposition,[],[f311,f1917])).
fof(f2584,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'const.combin.I_0')))) | $spl12),
  inference(cnf_transformation,[],[f2584_D])).
fof(f2584_D,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'const.combin.I_0')))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f2585,plain,(
  ( ! [X3] : (~p(s(bool,sK3(X3,'const.combin.I_0',bool)))) ) | $spl14),
  inference(cnf_transformation,[],[f2585_D])).
fof(f2585_D,plain,(
  ( ! [X3] : (~p(s(bool,sK3(X3,'const.combin.I_0',bool)))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f2576,plain,(
  ( ! [X3] : (~p(s(bool,sK3(X3,'const.combin.I_0',bool))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'const.combin.I_0'))))) )),
  inference(superposition,[],[f294,f339])).
fof(f1962,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK2(sK2('const.extreal.PosInf_0','type.extreal.extreal'),'type.extreal.extreal')) | $spl8),
  inference(cnf_transformation,[],[f1962_D])).
fof(f1962_D,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK2(sK2('const.extreal.PosInf_0','type.extreal.extreal'),'type.extreal.extreal')) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f1964,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK2(sK2('const.extreal.PosInf_0','type.extreal.extreal'),'type.extreal.extreal')) | $spl10),
  inference(cnf_transformation,[],[f1964_D])).
fof(f1964_D,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK2(sK2('const.extreal.PosInf_0','type.extreal.extreal'),'type.extreal.extreal')) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f1960,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK2(sK2('const.extreal.PosInf_0','type.extreal.extreal'),'type.extreal.extreal')) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK2(sK2('const.extreal.PosInf_0','type.extreal.extreal'),'type.extreal.extreal'))),
  inference(subsumption_resolution,[],[f1940,f217])).
fof(f1940,plain,(
  ~p(s(bool,'T_0')) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK2(sK2('const.extreal.PosInf_0','type.extreal.extreal'),'type.extreal.extreal')) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK2(sK2('const.extreal.PosInf_0','type.extreal.extreal'),'type.extreal.extreal'))),
  inference(superposition,[],[f311,f1919])).
fof(f819,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,sK5('const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))))))) ) | $spl4),
  inference(cnf_transformation,[],[f819_D])).
fof(f819_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,sK5('const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))))))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f813,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.PosInf_0') | s(bool,'T_0') = s(bool,sK5('const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))))))) )),
  inference(superposition,[],[f542,f234])).
fof(f343,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f342,f227])).
fof(f342,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f334,f335])).
fof(f333,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl0),
  inference(cnf_transformation,[],[f333_D])).
fof(f333_D,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal','const.extreal.NegInf_0')))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f335,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK0) | $spl2),
  inference(cnf_transformation,[],[f335_D])).
fof(f335_D,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK0) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f215,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK0) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal','const.extreal.NegInf_0'))))),
  inference(cnf_transformation,[],[f174])).
fof(f174,plain,(
  (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK0)) & (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',sK0))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f173])).
fof(f173,plain,(
  ? [X0] : ((p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0)) & (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0)))),
  inference(nnf_transformation,[],[f161])).
fof(f161,plain,(
  ? [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0')))) <~> s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0')))) <=> s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal','const.extreal.NegInf_0')))) <=> s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X6))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal','const.extreal.NegInf_0')))) <=> s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X6))),
  file('thm.extreal.le_infty|split|2','thm.extreal.le_infty|split|2')).
fof(f334,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl1),
  inference(cnf_transformation,[],[f334_D])).
fof(f334_D,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal','const.extreal.NegInf_0')))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f336,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',sK0) | $spl3),
  inference(cnf_transformation,[],[f336_D])).
fof(f336_D,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',sK0) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f216,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',sK0) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal','const.extreal.NegInf_0'))))),
  inference(cnf_transformation,[],[f174])).
% SZS output end Proof for thm.extreal
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 254
% Passive clauses: 1667
% Generated clauses: 16029
% Final active clauses: 155
% Final passive clauses: 1170
% Input formulas: 69
% Initial clauses: 114
% 
% Duplicate literals: 356
% Trivial inequalities: 102
% Fw subsumption resolutions: 433
% Fw demodulations: 3237
% Bw demodulations: 245
% Condensations: 391
% Global subsumptions: 14
% 
% Simple tautologies: 133
% Equational tautologies: 80
% Forward subsumptions: 9495
% Bw demodulations to eq. taut.: 7
% 
% Unit resulting resolution: 1283
% Factoring: 1
% Forward superposition: 4132
% Backward superposition: 5343
% Self superposition: 88
% Equality factoring: 30
% Equality resolution: 14
% 
% Split clauses: 25
% Split components: 50
% SAT solver clauses: 1829
% SAT solver unit clauses: 313
% SAT solver binary clauses: 612
% 
% Sat splits: 25
% Sat splitting refutations: 96
% 
% Memory used [KB]: 16758
% Time elapsed: 0.704 s
% ------------------------------
% ------------------------------
% Success in time 2.015 s
