Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.rich_list
% (50434)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1143
% Passive clauses: 9159
% Generated clauses: 30429
% Final active clauses: 1040
% Final passive clauses: 7586
% Input formulas: 261
% Initial clauses: 579
% 
% Duplicate literals: 209
% Trivial inequalities: 56
% Fw subsumption resolutions: 2090
% Fw demodulations: 3456
% Bw demodulations: 102
% 
% Simple tautologies: 322
% Equational tautologies: 25
% Forward subsumptions: 14677
% 
% Binary resolution: 8247
% Factoring: 907
% Forward superposition: 6499
% Backward superposition: 7985
% Self superposition: 46
% Equality factoring: 2
% Equality resolution: 14
% 
% Split clauses: 348
% Split components: 698
% SAT solver clauses: 609
% SAT solver unit clauses: 17
% SAT solver binary clauses: 114
% 
% Sat splits: 348
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 605
% TWLsolver calls for satisfiability: 486
% 
% Memory used [KB]: 16119
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.rich_list
% (50455)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 585
% Passive clauses: 4564
% Generated clauses: 24340
% Final active clauses: 532
% Final passive clauses: 3878
% Input formulas: 261
% Initial clauses: 579
% 
% Duplicate literals: 105
% Trivial inequalities: 46
% Fw subsumption resolutions: 887
% Fw demodulations: 3985
% Bw demodulations: 147
% Condensations: 112
% Global subsumptions: 24
% 
% Simple tautologies: 300
% Equational tautologies: 12
% Forward subsumptions: 9355
% Bw demodulations to eq. taut.: 4
% 
% Binary resolution: 3039
% Factoring: 133
% Forward superposition: 6152
% Backward superposition: 4449
% Self superposition: 54
% Equality factoring: 2
% Equality resolution: 8
% 
% Split clauses: 65
% Split components: 130
% SAT solver clauses: 5502
% SAT solver unit clauses: 647
% SAT solver binary clauses: 188
% SAT solver learnt clauses: 48
% SAT solver learnt literals: 12
% 
% Sat splits: 65
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 5406
% TWLsolver calls for satisfiability: 3655
% 
% Memory used [KB]: 31598
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.rich_list
% (50472)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1028
% Passive clauses: 8616
% Generated clauses: 123742
% Final active clauses: 954
% Final passive clauses: 5795
% Input formulas: 261
% Initial clauses: 579
% 
% Duplicate literals: 9
% Trivial inequalities: 60
% Fw subsumption resolutions: 1185
% Fw demodulations: 56102
% Bw demodulations: 1250
% Condensations: 38
% Global subsumptions: 4
% 
% Simple tautologies: 163
% Equational tautologies: 50
% Forward subsumptions: 38738
% Fw demodulations to eq. taut.: 18027
% Bw demodulations to eq. taut.: 20
% 
% Unit resulting resolution: 35441
% Factoring: 1
% Forward superposition: 17623
% Backward superposition: 10157
% Self superposition: 117
% Equality factoring: 2
% Equality resolution: 13
% 
% Split clauses: 61
% Split components: 122
% SAT solver clauses: 9433
% SAT solver unit clauses: 1311
% SAT solver binary clauses: 5228
% 
% Sat splits: 61
% Sat splitting refutations: 689
% 
% Memory used [KB]: 220294
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.rich_list
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.rich_list
% SZS output start Proof for thm.rich_list
fof(f15072,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f854,f855,f856,f857,f858,f859,f865,f860,f861,f866,f881,f883,f884,f886,f888,f976,f1091,f1094,f1097,f1098,f1093,f1221,f1362,f1387,f1389,f1390,f1391,f1394,f1483,f1512,f1590,f1591,f1592,f1593,f1092,f1751,f1793,f1810,f1812,f1813,f1815,f1817,f1905,f1934,f1936,f1937,f1939,f1941,f2029,f2058,f2060,f2061,f2063,f2065,f2153,f2182,f2229,f1752,f3285,f3287,f3288,f3289,f3292,f3381,f3401,f3407,f3410,f3412,f3413,f3415,f3417,f3505,f3535,f3537,f3538,f3540,f3542,f3630,f3941,f3943,f3945,f3947,f3948,f4130,f4132,f4133,f4135,f4137,f4225,f4318,f4422,f4424,f4425,f4427,f4429,f4517,f4559,f4561,f4562,f4564,f4566,f4654,f4708,f4710,f4711,f4712,f4715,f4804,f4903,f5189,f5212,f5195,f5213,f5374,f5331,f5441,f5442,f5476,f4542,f5477,f4679,f5478,f5479,f5480,f5481,f5482,f5532,f5533,f5534,f5535,f5536,f5537,f5547,f5563,f5565,f5566,f5568,f5570,f5658,f5710,f5712,f5713,f5715,f5717,f5805,f6436,f5693,f6652,f5840,f7044,f7046,f7047,f7048,f7051,f7140,f7391,f7393,f7394,f7395,f7398,f7487,f7632,f7637,f7647,f7655,f7664,f7673,f7682,f7951,f7962,f7970,f7981,f7982,f5543,f7985,f7986,f7987,f7988,f8011,f8013,f8014,f8015,f8018,f8107,f8140,f8155,f8160,f8170,f8178,f5545,f8305,f8308,f8309,f864,f8310,f8311,f863,f8483,f8485,f8486,f8488,f8490,f8578,f8800,f8802,f8803,f8805,f8807,f8895,f8924,f8926,f8927,f8929,f8931,f9019,f9047,f9049,f9050,f9052,f9054,f9142,f9669,f9672,f9673,f9674,f9675,f9676,f9677,f9678,f9671,f10209,f3524,f10212,f10214,f10215,f10217,f10219,f10307,f10335,f10337,f10338,f10340,f10342,f10430,f10678,f10680,f10681,f10683,f10685,f10773,f10891,f10893,f10894,f10895,f10898,f10987,f11008,f11042,f11043,f11044,f11045,f11046,f11047,f11065,f1003,f11066,f2048,f11067,f2172,f11068,f8914,f11069,f3525,f11070,f3649,f11127,f11134,f10865,f11182,f11198,f11185,f11199,f11186,f11200,f11187,f11201,f11188,f11202,f11189,f11203,f11190,f11204,f11245,f11573,f11575,f11576,f11577,f11580,f11669,f11766,f11998,f12000,f12001,f12002,f12005,f12094,f7241,f12210,f12212,f12213,f12215,f12217,f12305,f12336,f12371,f12373,f12374,f12376,f12378,f12466,f12497,f12573,f12575,f12576,f12577,f12580,f12669,f7588,f12788,f12790,f12791,f12793,f12795,f12883,f12923,f12997,f12999,f13000,f13002,f13004,f13092,f13131,f13218,f13264,f13314,f7640,f13357,f13378,f13379,f13380,f13381,f13406,f13411,f8163,f13417,f13429,f13448,f13458,f13465,f13468,f13510,f13511,f7994,f13514,f13515,f13525,f13526,f13553,f13555,f13556,f13557,f13560,f13649,f13678,f13811,f13814,f13815,f13816,f13813,f13971,f14023,f14025,f14080,f14082,f14445,f7996,f14446,f14450,f14459,f14461,f14462,f14464,f14466,f14554,f1924,f14601,f14603,f14604,f1004,f14605,f14608,f1005,f14697,f1007,f14718,f14851,f13223,f14852,f13422,f14866,f13978,f1006,f15071])).
fof(f15071,plain,(
  $false | ($spl1 | $spl2)),
  inference(trivial_inequality_removal,[],[f15070])).
fof(f15070,plain,(
  s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))) != s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f15067,f8326])).
fof(f8326,plain,(
  s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.FILTER_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) | $spl2),
  inference(resolution,[],[f1081,f663])).
fof(f663,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'(X0),X2) = s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))))) )),
  inference(cnf_transformation,[],[f483])).
fof(f483,plain,(
  ! [X0,X1,X2] : ((s('type.list.list'(X0),X2) != s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))))) & (~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) | s('type.list.list'(X0),X2) = s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))))),
  inference(nnf_transformation,[],[f327])).
fof(f327,plain,(
  ! [X0,X1,X2] : (s('type.list.list'(X0),X2) = s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))) <=> p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))))),
  inference(flattening,[],[f326])).
fof(f326,plain,(
  ! [X0] : ! [X1,X2] : (s('type.list.list'(X0),X2) = s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))) <=> p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))))),
  inference(rectify,[],[f63])).
fof(f63,axiom,(
  ! [X6] : ! [X7,X8] : (s('type.list.list'(X6),X8) = s('type.list.list'(X6),'const.list.FILTER_2'(s(fun(X6,bool),X7),s('type.list.list'(X6),X8))) <=> p(s(bool,'const.list.EVERY_2'(s(fun(X6,bool),X7),s('type.list.list'(X6),X8)))))),
  file('thm.rich_list.EVERY_REVERSE','thm.list.FILTER_EQ_ID')).
fof(f1081,plain,(
  p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)))) | $spl2),
  inference(equality_resolution,[],[f865])).
fof(f15067,plain,(
  s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))) != s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),'const.list.FILTER_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))))) | $spl1),
  inference(superposition,[],[f8478,f635])).
fof(f635,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'(X0),'const.list.REVERSE_1'(s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X2),s('type.list.list'(X0),X1))))) = s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X2),s('type.list.list'(X0),'const.list.REVERSE_1'(s('type.list.list'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f295])).
fof(f295,plain,(
  ! [X0,X1,X2] : s('type.list.list'(X0),'const.list.REVERSE_1'(s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X2),s('type.list.list'(X0),X1))))) = s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X2),s('type.list.list'(X0),'const.list.REVERSE_1'(s('type.list.list'(X0),X1)))))),
  inference(flattening,[],[f294])).
fof(f294,plain,(
  ! [X0] : ! [X1,X2] : s('type.list.list'(X0),'const.list.REVERSE_1'(s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X2),s('type.list.list'(X0),X1))))) = s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X2),s('type.list.list'(X0),'const.list.REVERSE_1'(s('type.list.list'(X0),X1)))))),
  inference(rectify,[],[f242])).
fof(f242,axiom,(
  ! [X6] : ! [X8,X7] : s('type.list.list'(X6),'const.list.REVERSE_1'(s('type.list.list'(X6),'const.list.FILTER_2'(s(fun(X6,bool),X7),s('type.list.list'(X6),X8))))) = s('type.list.list'(X6),'const.list.FILTER_2'(s(fun(X6,bool),X7),s('type.list.list'(X6),'const.list.REVERSE_1'(s('type.list.list'(X6),X8)))))),
  file('thm.rich_list.EVERY_REVERSE','thm.list.FILTER_REVERSE')).
fof(f8478,plain,(
  s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))) != s('type.list.list'(sK12),'const.list.FILTER_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))) | $spl1),
  inference(resolution,[],[f863,f662])).
fof(f662,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) | s('type.list.list'(X0),X2) != s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) )),
  inference(cnf_transformation,[],[f483])).
fof(f1006,plain,(
  ( ! [X28] : (~p(s(bool,X28))) ) | $spl12),
  inference(cnf_transformation,[],[f1006_D])).
fof(f1006_D,plain,(
  ( ! [X28] : (~p(s(bool,X28))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f13978,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20(sK14,sK13,sK12))))) ) | $spl124),
  inference(cnf_transformation,[],[f13978_D])).
fof(f13978_D,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20(sK14,sK13,sK12))))) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f14866,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20(sK14,sK13,sK12))))) ) | ($spl2 | $spl22)),
  inference(resolution,[],[f13837,f11162])).
fof(f11162,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | s(bool,X1) = s(bool,X2)) ) | $spl2),
  inference(duplicate_literal_removal,[],[f11150])).
fof(f11150,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | s(bool,X1) = s(bool,X2) | ~p(s(bool,X0))) ) | $spl2),
  inference(resolution,[],[f11144,f608])).
fof(f608,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X2)) | ~sP3(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f465])).
fof(f465,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP3(X2,X1,X0)) & (~sP3(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f437])).
fof(f437,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP3(X2,X1,X0))),
  inference(definition_folding,[],[f267,f436,f435,f434,f433])).
fof(f433,plain,(
  ! [X0,X1,X2] : (sP0(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f434,plain,(
  ! [X1,X0,X2] : (sP1(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f435,plain,(
  ! [X2,X0,X1] : (sP2(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f436,plain,(
  ! [X2,X1,X0] : (sP3(X2,X1,X0) <=> (sP0(X0,X1,X2) & sP1(X1,X0,X2) & sP2(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f267,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f266])).
fof(f266,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f145])).
fof(f145,axiom,(
  ! [X57,X63,X64] : ((p(s(bool,X64)) <=> s(bool,X57) = s(bool,X63)) <=> ((p(s(bool,X64)) | p(s(bool,X63)) | p(s(bool,X57))) & (p(s(bool,X64)) | ~p(s(bool,X57)) | ~p(s(bool,X63))) & (p(s(bool,X63)) | ~p(s(bool,X57)) | ~p(s(bool,X64))) & (p(s(bool,X57)) | ~p(s(bool,X63)) | ~p(s(bool,X64)))))),
  file('thm.rich_list.EVERY_REVERSE','thm.sat.sat_2|49|_prove_ltb')).
fof(f11144,plain,(
  ( ! [X4,X5,X3] : (sP3(X4,X3,X5) | ~sP2(X4,X5,X3) | p(s(bool,X3)) | ~p(s(bool,X4))) ) | $spl2),
  inference(resolution,[],[f11141,f603])).
fof(f603,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f464])).
fof(f464,plain,(
  ! [X0,X1,X2] : ((~sP0(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP0(X0,X1,X2)))),
  inference(flattening,[],[f463])).
fof(f463,plain,(
  ! [X0,X1,X2] : ((~sP0(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP0(X0,X1,X2)))),
  inference(nnf_transformation,[],[f433])).
fof(f11141,plain,(
  ( ! [X14,X12,X13] : (~sP0(X14,X12,X13) | p(s(bool,X12)) | ~sP2(X13,X14,X12) | sP3(X13,X12,X14)) ) | $spl2),
  inference(subsumption_resolution,[],[f11138,f865])).
fof(f11138,plain,(
  ( ! [X14,X12,X13] : (s(bool,X12) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) | sP3(X13,X12,X14) | p(s(bool,X12)) | ~sP2(X13,X14,X12) | ~sP0(X14,X12,X13)) ) | $spl2),
  inference(resolution,[],[f10694,f592])).
fof(f592,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X1)) | ~sP2(X0,X2,X1) | ~sP1(X1,X2,X0) | ~sP0(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f456])).
fof(f456,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | (sP0(X2,X1,X0) & sP1(X1,X2,X0) & sP2(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP0(X2,X1,X0) | ~sP1(X1,X2,X0) | ~sP2(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f455])).
fof(f455,plain,(
  ! [X2,X1,X0] : ((~sP3(X2,X1,X0) | (sP0(X0,X1,X2) & sP1(X1,X0,X2) & sP2(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP0(X0,X1,X2) | ~sP1(X1,X0,X2) | ~sP2(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP3(X2,X1,X0)))),
  inference(flattening,[],[f454])).
fof(f454,plain,(
  ! [X2,X1,X0] : ((~sP3(X2,X1,X0) | (sP0(X0,X1,X2) & sP1(X1,X0,X2) & sP2(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP0(X0,X1,X2) | ~sP1(X1,X0,X2) | ~sP2(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP3(X2,X1,X0)))),
  inference(nnf_transformation,[],[f436])).
fof(f10694,plain,(
  ( ! [X57,X58,X56] : (sP1(X56,X57,X58) | s(bool,X56) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)))) ) | $spl2),
  inference(resolution,[],[f8354,f601])).
fof(f601,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f462])).
fof(f462,plain,(
  ! [X0,X1,X2] : ((~sP1(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1,X2)))),
  inference(rectify,[],[f461])).
fof(f461,plain,(
  ! [X1,X0,X2] : ((~sP1(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0,X2)))),
  inference(flattening,[],[f460])).
fof(f460,plain,(
  ! [X1,X0,X2] : ((~sP1(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0,X2)))),
  inference(nnf_transformation,[],[f434])).
fof(f8354,plain,(
  ( ! [X67] : (~p(s(bool,X67)) | s(bool,X67) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)))) ) | $spl2),
  inference(resolution,[],[f1081,f618])).
fof(f618,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f468])).
fof(f468,plain,(
  ! [X0,X1] : ((s(bool,X0) != s(bool,X1) | ((~p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | p(s(bool,X0))))) & ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(flattening,[],[f467])).
fof(f467,plain,(
  ! [X0,X1] : ((s(bool,X0) != s(bool,X1) | ((~p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | p(s(bool,X0))))) & (((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0)))) | s(bool,X0) = s(bool,X1)))),
  inference(nnf_transformation,[],[f394])).
fof(f394,plain,(
  ! [X0,X1] : (s(bool,X0) = s(bool,X1) <=> ((~p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | p(s(bool,X0)))))),
  inference(ennf_transformation,[],[f274])).
fof(f274,plain,(
  ! [X0,X1] : (s(bool,X0) = s(bool,X1) <=> ((p(s(bool,X0)) => p(s(bool,X1))) & (p(s(bool,X1)) => p(s(bool,X0)))))),
  inference(rectify,[],[f151])).
fof(f151,axiom,(
  ! [X25,X26] : (s(bool,X25) = s(bool,X26) <=> ((p(s(bool,X25)) => p(s(bool,X26))) & (p(s(bool,X26)) => p(s(bool,X25)))))),
  file('thm.rich_list.EVERY_REVERSE','thm.bool.EQ_IMP_THM')).
fof(f13837,plain,(
  ( ! [X39,X38] : (sP2(X38,happ(s(fun(sK12,bool),sK13),s(sK12,sK20(sK14,sK13,sK12))),X39)) ) | $spl22),
  inference(resolution,[],[f13801,f596])).
fof(f596,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f459])).
fof(f459,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(rectify,[],[f458])).
fof(f458,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(flattening,[],[f457])).
fof(f457,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(nnf_transformation,[],[f435])).
fof(f13801,plain,(
  ~p(s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20(sK14,sK13,sK12))))) | $spl22),
  inference(equality_resolution,[],[f1924])).
fof(f13422,plain,(
  s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12)))) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))) | $spl118),
  inference(cnf_transformation,[],[f13422_D])).
fof(f13422_D,plain,(
  s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12)))) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f14852,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12)))) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) ) | $spl1),
  inference(resolution,[],[f8546,f12638])).
fof(f12638,plain,(
  ( ! [X198,X199,X197] : (sP8(X198,X199,X197) | s(bool,X197) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12))))) ) | $spl1),
  inference(resolution,[],[f9071,f730])).
fof(f730,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f511])).
fof(f511,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(flattening,[],[f510])).
fof(f510,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(nnf_transformation,[],[f444])).
fof(f444,plain,(
  ! [X0,X1,X2] : (sP8(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f9071,plain,(
  ( ! [X60] : (p(s(bool,X60)) | s(bool,X60) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12))))) ) | $spl1),
  inference(resolution,[],[f8799,f615])).
fof(f615,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f468])).
fof(f8799,plain,(
  ~p(s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12))))) | $spl1),
  inference(resolution,[],[f8480,f828])).
fof(f828,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK44(X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f572])).
fof(f572,plain,(
  ! [X0,X1] : ((! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK43(X1,X0)))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK44(X1,X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK43,sK44])],[f571])).
fof(f571,plain,(
  ! [X0,X1] : ((! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(rectify,[],[f570])).
fof(f570,plain,(
  ! [X0,X1] : ((! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(nnf_transformation,[],[f431])).
fof(f431,plain,(
  ! [X0,X1] : (? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(ennf_transformation,[],[f382])).
fof(f382,plain,(
  ! [X0,X1] : (~! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(flattening,[],[f381])).
fof(f381,plain,(
  ! [X0] : ! [X1] : (~! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(rectify,[],[f236])).
fof(f236,axiom,(
  ! [X6] : ! [X7] : (~! [X12] : p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X12)))) <=> ? [X12] : ~p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X12)))))),
  file('thm.rich_list.EVERY_REVERSE','thm.bool.NOT_FORALL_THM')).
fof(f8480,plain,(
  ~p(s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20('const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)),sK13,sK12))))) | $spl1),
  inference(resolution,[],[f863,f666])).
fof(f666,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK20(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f486])).
fof(f486,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))))) & ((p(s(bool,'const.bool.IN_2'(s(X0,sK20(X2,X1,X0)),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK20(X2,X1,X0)))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK20])],[f485])).
fof(f485,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))))) & (? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))))))),
  inference(rectify,[],[f484])).
fof(f484,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) | ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))) & (? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))))))),
  inference(nnf_transformation,[],[f406])).
fof(f406,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) <=> ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))))),
  inference(ennf_transformation,[],[f329])).
fof(f329,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))))),
  inference(flattening,[],[f328])).
fof(f328,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))))),
  inference(rectify,[],[f81])).
fof(f81,axiom,(
  ! [X6] : ! [X7,X8] : (p(s(bool,'const.list.EVERY_2'(s(fun(X6,bool),X7),s('type.list.list'(X6),X8)))) <=> ! [X43] : (p(s(bool,'const.bool.IN_2'(s(X6,X43),s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),X8)))))) => p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X43))))))),
  file('thm.rich_list.EVERY_REVERSE','thm.list.EVERY_MEM')).
fof(f8546,plain,(
  ( ! [X140,X139] : (~sP8(X139,X140,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))) | ~p(s(bool,X140)) | ~p(s(bool,X139))) ) | $spl1),
  inference(resolution,[],[f863,f729])).
fof(f729,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f511])).
fof(f13223,plain,(
  s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))) | $spl116),
  inference(cnf_transformation,[],[f13223_D])).
fof(f13223_D,plain,(
  s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f14851,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) ) | $spl1),
  inference(resolution,[],[f8546,f12063])).
fof(f12063,plain,(
  ( ! [X198,X199,X197] : (sP8(X198,X199,X197) | s(bool,X197) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12))))) ) | $spl1),
  inference(resolution,[],[f8948,f730])).
fof(f8948,plain,(
  ( ! [X60] : (p(s(bool,X60)) | s(bool,X60) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12))))) ) | $spl1),
  inference(resolution,[],[f8798,f615])).
fof(f8798,plain,(
  ~p(s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12))))) | $spl1),
  inference(resolution,[],[f8480,f827])).
fof(f827,plain,(
  ( ! [X0,X5,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK43(X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))) )),
  inference(cnf_transformation,[],[f572])).
fof(f14718,plain,(
  ( ! [X20] : (p(s(bool,X20)) | s(bool,X20) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20(sK14,sK13,sK12))))) ) | $spl128),
  inference(cnf_transformation,[],[f14718_D])).
fof(f14718_D,plain,(
  ( ! [X20] : (p(s(bool,X20)) | s(bool,X20) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20(sK14,sK13,sK12))))) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f1007,plain,(
  ( ! [X26] : (p(s(bool,X26))) ) | $spl14),
  inference(cnf_transformation,[],[f1007_D])).
fof(f1007_D,plain,(
  ( ! [X26] : (p(s(bool,X26))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f14697,plain,(
  ( ! [X212,X211] : (p(s(bool,X211)) | p(s(bool,X212)) | s(bool,X212) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20(sK14,sK13,sK12))))) ) | $spl122),
  inference(resolution,[],[f13812,f839])).
fof(f839,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f583])).
fof(f583,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f453])).
fof(f453,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f452])).
fof(f452,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f393])).
fof(f393,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f392])).
fof(f392,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f265])).
fof(f265,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X12,X29,X27,X30] : ((s(bool,X12) = s(bool,X29) & (p(s(bool,X29)) => s(bool,X27) = s(bool,X30))) => ((p(s(bool,X12)) => p(s(bool,X27))) <=> (p(s(bool,X29)) => p(s(bool,X30)))))),
  file('thm.rich_list.EVERY_REVERSE','thm.bool.IMP_CONG')).
fof(f13812,plain,(
  p(s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20(sK14,sK13,sK12))))) | $spl122),
  inference(cnf_transformation,[],[f13812_D])).
fof(f13812_D,plain,(
  p(s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20(sK14,sK13,sK12))))) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f1005,plain,(
  ( ! [X19,X17] : (~p(s(bool,X17)) | s(bool,X17) != s(bool,X19)) ) | $spl10),
  inference(cnf_transformation,[],[f1005_D])).
fof(f1005_D,plain,(
  ( ! [X19,X17] : (~p(s(bool,X17)) | s(bool,X17) != s(bool,X19)) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f14608,plain,(
  ( ! [X41,X42,X40] : (p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20(sK14,sK13,sK12)))) | s(bool,X41) != s(bool,X42)) ) | $spl122),
  inference(resolution,[],[f13812,f586])).
fof(f586,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f453])).
fof(f14605,plain,(
  ( ! [X33,X31,X32] : (p(s(bool,X31)) | ~p(s(bool,X32)) | p(s(bool,X33)) | s(bool,X33) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20(sK14,sK13,sK12))))) ) | $spl122),
  inference(resolution,[],[f13812,f585])).
fof(f585,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f453])).
fof(f1004,plain,(
  ( ! [X13,X11] : (p(s(bool,X11)) | s(bool,X11) != s(bool,X13)) ) | $spl8),
  inference(cnf_transformation,[],[f1004_D])).
fof(f1004_D,plain,(
  ( ! [X13,X11] : (p(s(bool,X11)) | s(bool,X11) != s(bool,X13)) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f14604,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20(sK14,sK13,sK12))))) ) | $spl122),
  inference(resolution,[],[f13812,f584])).
fof(f584,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f453])).
fof(f14603,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20(sK14,sK13,sK12)))) | s(bool,X26) != s(bool,X27)) ) | $spl122),
  inference(resolution,[],[f13812,f582])).
fof(f582,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f453])).
fof(f14601,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | p(s(bool,X20)) | s(bool,X19) != s(bool,X21) | s(bool,X20) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20(sK14,sK13,sK12))))) ) | $spl122),
  inference(resolution,[],[f13812,f580])).
fof(f580,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f453])).
fof(f1924,plain,(
  ( ! [X7] : (s(bool,X7) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20(sK14,sK13,sK12)))) | ~p(s(bool,X7))) ) | $spl22),
  inference(cnf_transformation,[],[f1924_D])).
fof(f1924_D,plain,(
  ( ! [X7] : (s(bool,X7) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20(sK14,sK13,sK12)))) | ~p(s(bool,X7))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f14554,plain,(
  ( ! [X196,X197] : (p(s(bool,X196)) | ~p(s(bool,X197)) | s(bool,X197) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20(sK14,sK13,sK12))))) ) | $spl123),
  inference(resolution,[],[f13813,f839])).
fof(f14466,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X28) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20(sK14,sK13,sK12)))) | s(bool,X27) != s(bool,X29)) ) | $spl123),
  inference(resolution,[],[f13813,f586])).
fof(f14464,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20(sK14,sK13,sK12))))) ) | $spl123),
  inference(resolution,[],[f13813,f585])).
fof(f14462,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20(sK14,sK13,sK12))))) ) | $spl123),
  inference(resolution,[],[f13813,f584])).
fof(f14461,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X13) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20(sK14,sK13,sK12)))) | s(bool,X12) != s(bool,X14)) ) | $spl123),
  inference(resolution,[],[f13813,f582])).
fof(f14459,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20(sK14,sK13,sK12))))) ) | $spl123),
  inference(resolution,[],[f13813,f580])).
fof(f14450,plain,(
  s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12)))) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20(sK14,sK13,sK12)))) | $spl126),
  inference(cnf_transformation,[],[f14450_D])).
fof(f14450_D,plain,(
  s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12)))) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20(sK14,sK13,sK12)))) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f14446,plain,(
  ( ! [X10] : (p(s(bool,X10)) | s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12)))) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20(sK14,sK13,sK12))))) ) | ($spl1 | $spl22)),
  inference(resolution,[],[f13898,f12652])).
fof(f12652,plain,(
  ( ! [X241,X239,X240] : (sP10(X240,X239,X241) | s(bool,X239) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12))))) ) | $spl1),
  inference(resolution,[],[f9071,f751])).
fof(f751,plain,(
  ( ! [X2,X0,X1] : (sP10(X2,X0,X1) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f520])).
fof(f520,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X0)) & (~p(s(bool,X1)) | ~p(s(bool,X2)))) | sP10(X2,X0,X1)) & (~sP10(X2,X0,X1) | p(s(bool,X0)) | (p(s(bool,X1)) & p(s(bool,X2)))))),
  inference(flattening,[],[f519])).
fof(f519,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X0)) & (~p(s(bool,X1)) | ~p(s(bool,X2)))) | sP10(X2,X0,X1)) & (~sP10(X2,X0,X1) | (p(s(bool,X0)) | (p(s(bool,X1)) & p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f448])).
fof(f448,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X0)) | (p(s(bool,X1)) & p(s(bool,X2)))) <=> sP10(X2,X0,X1))),
  inference(definition_folding,[],[f345,f447])).
fof(f447,plain,(
  ! [X2,X0,X1] : (sP10(X2,X0,X1) <=> ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f345,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X0)) | (p(s(bool,X1)) & p(s(bool,X2)))) <=> ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(rectify,[],[f230])).
fof(f230,axiom,(
  ! [X6,X11,X46] : ((p(s(bool,X6)) | (p(s(bool,X11)) & p(s(bool,X46)))) <=> ((p(s(bool,X6)) | p(s(bool,X11))) & (p(s(bool,X6)) | p(s(bool,X46)))))),
  file('thm.rich_list.EVERY_REVERSE','thm.bool.LEFT_OR_OVER_AND')).
fof(f13898,plain,(
  ( ! [X154,X155] : (~sP10(X155,happ(s(fun(sK12,bool),sK13),s(sK12,sK20(sK14,sK13,sK12))),X154) | p(s(bool,X154))) ) | $spl22),
  inference(resolution,[],[f13801,f745])).
fof(f745,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | p(s(bool,X1)) | ~sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f518])).
fof(f518,plain,(
  ! [X0,X1,X2] : ((~sP10(X0,X1,X2) | ((p(s(bool,X1)) | p(s(bool,X2))) & (p(s(bool,X1)) | p(s(bool,X0))))) & ((~p(s(bool,X1)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) & ~p(s(bool,X0))) | sP10(X0,X1,X2)))),
  inference(rectify,[],[f517])).
fof(f517,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(flattening,[],[f516])).
fof(f516,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2))))) & (((~p(s(bool,X0)) & ~p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X2)))) | sP10(X2,X0,X1)))),
  inference(nnf_transformation,[],[f447])).
fof(f7996,plain,(
  s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20(sK14,sK13,sK12)))) | $spl66),
  inference(cnf_transformation,[],[f7996_D])).
fof(f7996_D,plain,(
  s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20(sK14,sK13,sK12)))) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f14445,plain,(
  ( ! [X9] : (p(s(bool,X9)) | s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20(sK14,sK13,sK12))))) ) | ($spl1 | $spl22)),
  inference(resolution,[],[f13898,f12077])).
fof(f12077,plain,(
  ( ! [X241,X239,X240] : (sP10(X240,X239,X241) | s(bool,X239) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12))))) ) | $spl1),
  inference(resolution,[],[f8948,f751])).
fof(f14082,plain,(
  ( ! [X14,X15] : (p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12)))) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) ) | $spl1),
  inference(resolution,[],[f8533,f12621])).
fof(f12621,plain,(
  ( ! [X146,X147,X148] : (sP6(X147,X148,X146) | s(bool,X146) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12))))) ) | $spl1),
  inference(resolution,[],[f9071,f713])).
fof(f713,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f504])).
fof(f504,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(rectify,[],[f503])).
fof(f503,plain,(
  ! [X2,X0,X1] : ((~sP6(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP6(X2,X0,X1)))),
  inference(flattening,[],[f502])).
fof(f502,plain,(
  ! [X2,X0,X1] : ((~sP6(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP6(X2,X0,X1)))),
  inference(nnf_transformation,[],[f441])).
fof(f441,plain,(
  ! [X2,X0,X1] : (sP6(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f8533,plain,(
  ( ! [X114,X113] : (~sP6(X113,X114,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))) | p(s(bool,X114)) | ~p(s(bool,X113))) ) | $spl1),
  inference(resolution,[],[f863,f712])).
fof(f712,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f504])).
fof(f14080,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) ) | $spl1),
  inference(resolution,[],[f8533,f12046])).
fof(f12046,plain,(
  ( ! [X146,X147,X148] : (sP6(X147,X148,X146) | s(bool,X146) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12))))) ) | $spl1),
  inference(resolution,[],[f8948,f713])).
fof(f14025,plain,(
  ( ! [X17,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12)))) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) ) | $spl1),
  inference(resolution,[],[f8532,f12622])).
fof(f12622,plain,(
  ( ! [X151,X149,X150] : (sP6(X150,X149,X151) | s(bool,X149) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12))))) ) | $spl1),
  inference(resolution,[],[f9071,f714])).
fof(f714,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f504])).
fof(f8532,plain,(
  ( ! [X111,X112] : (~sP6(X111,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))),X112) | p(s(bool,X112)) | ~p(s(bool,X111))) ) | $spl1),
  inference(resolution,[],[f863,f712])).
fof(f14023,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) ) | $spl1),
  inference(resolution,[],[f8532,f12047])).
fof(f12047,plain,(
  ( ! [X151,X149,X150] : (sP6(X150,X149,X151) | s(bool,X149) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12))))) ) | $spl1),
  inference(resolution,[],[f8948,f714])).
fof(f13971,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20(sK14,sK13,sK12))))) ) | ($spl2 | $spl22)),
  inference(resolution,[],[f13837,f11162])).
fof(f13813,plain,(
  ~p(s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20(sK14,sK13,sK12))))) | $spl123),
  inference(cnf_transformation,[],[f13813_D])).
fof(f13813_D,plain,(
  ~p(s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20(sK14,sK13,sK12))))) <=> ~$spl123),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl123])])).
fof(f13816,plain,(
  ( ! [X26,X27] : (~p(s(bool,X27)) | ~p(s(bool,X26)) | ~p(s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20(sK14,sK13,sK12)))))) ) | $spl22),
  inference(subsumption_resolution,[],[f13800,f618])).
fof(f13800,plain,(
  ( ! [X26,X27] : (s(bool,X26) != s(bool,X27) | ~p(s(bool,X27)) | ~p(s(bool,X26)) | ~p(s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20(sK14,sK13,sK12)))))) ) | $spl22),
  inference(superposition,[],[f1924,f622])).
fof(f622,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f396])).
fof(f396,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f395])).
fof(f395,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f275])).
fof(f275,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X25,X26] : ((p(s(bool,X25)) => p(s(bool,X26))) => ((p(s(bool,X26)) => p(s(bool,X25))) => s(bool,X25) = s(bool,X26)))),
  file('thm.rich_list.EVERY_REVERSE','thm.bool.IMP_ANTISYM_AX')).
fof(f13815,plain,(
  ( ! [X24,X25] : (~p(s(bool,X25)) | ~p(s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20(sK14,sK13,sK12))))) | ~p(s(bool,X24))) ) | $spl22),
  inference(subsumption_resolution,[],[f13799,f618])).
fof(f13799,plain,(
  ( ! [X24,X25] : (s(bool,X24) != s(bool,X25) | ~p(s(bool,X25)) | ~p(s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20(sK14,sK13,sK12))))) | ~p(s(bool,X24))) ) | $spl22),
  inference(superposition,[],[f1924,f622])).
fof(f13814,plain,(
  ( ! [X19,X18] : (~p(s(bool,X19)) | ~p(s(bool,X18)) | ~p(s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20(sK14,sK13,sK12)))))) ) | $spl22),
  inference(subsumption_resolution,[],[f13796,f618])).
fof(f13796,plain,(
  ( ! [X19,X18] : (s(bool,X18) != s(bool,X19) | ~p(s(bool,X19)) | ~p(s(bool,X18)) | ~p(s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20(sK14,sK13,sK12)))))) ) | $spl22),
  inference(superposition,[],[f1924,f618])).
fof(f13811,plain,(
  ( ! [X17,X16] : (~p(s(bool,X17)) | ~p(s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20(sK14,sK13,sK12))))) | ~p(s(bool,X16))) ) | $spl22),
  inference(subsumption_resolution,[],[f13795,f618])).
fof(f13795,plain,(
  ( ! [X17,X16] : (s(bool,X16) != s(bool,X17) | ~p(s(bool,X17)) | ~p(s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20(sK14,sK13,sK12))))) | ~p(s(bool,X16))) ) | $spl22),
  inference(superposition,[],[f1924,f618])).
fof(f13678,plain,(
  ( ! [X30,X28] : (p(s(bool,X28)) | s(bool,X28) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.FILTER_2'(s(fun(sK12,bool),X30),s('type.list.list'(sK12),sK14)))))) ) | $spl120),
  inference(cnf_transformation,[],[f13678_D])).
fof(f13678_D,plain,(
  ( ! [X30,X28] : (p(s(bool,X28)) | s(bool,X28) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.FILTER_2'(s(fun(sK12,bool),X30),s('type.list.list'(sK12),sK14)))))) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f13649,plain,(
  ( ! [X315,X316,X314] : (p(s(bool,X314)) | p(s(bool,X315)) | s(bool,X315) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.FILTER_2'(s(fun(sK12,bool),X316),s('type.list.list'(sK12),sK14)))))) ) | $spl2),
  inference(resolution,[],[f8325,f839])).
fof(f8325,plain,(
  ( ! [X9] : (p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.FILTER_2'(s(fun(sK12,bool),X9),s('type.list.list'(sK12),sK14))))))) ) | $spl2),
  inference(resolution,[],[f1081,f661])).
fof(f661,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X2),s('type.list.list'(X0),X3)))))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X3))))) )),
  inference(cnf_transformation,[],[f405])).
fof(f405,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X3)))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X2),s('type.list.list'(X0),X3)))))))),
  inference(ennf_transformation,[],[f325])).
fof(f325,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X3)))) => p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X2),s('type.list.list'(X0),X3)))))))),
  inference(flattening,[],[f324])).
fof(f324,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X3)))) => p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X2),s('type.list.list'(X0),X3)))))))),
  inference(rectify,[],[f75])).
fof(f75,axiom,(
  ! [X6] : ! [X17,X18,X8] : (p(s(bool,'const.list.EVERY_2'(s(fun(X6,bool),X17),s('type.list.list'(X6),X8)))) => p(s(bool,'const.list.EVERY_2'(s(fun(X6,bool),X17),s('type.list.list'(X6),'const.list.FILTER_2'(s(fun(X6,bool),X18),s('type.list.list'(X6),X8)))))))),
  file('thm.rich_list.EVERY_REVERSE','thm.list.EVERY_FILTER_IMP')).
fof(f13560,plain,(
  ( ! [X57,X54,X56,X55] : (p(s(bool,X54)) | ~p(s(bool,X55)) | s(bool,X54) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.FILTER_2'(s(fun(sK12,bool),X56),s('type.list.list'(sK12),sK14))))) | s(bool,X55) != s(bool,X57)) ) | $spl2),
  inference(resolution,[],[f8325,f586])).
fof(f13557,plain,(
  ( ! [X45,X43,X46,X44] : (p(s(bool,X43)) | ~p(s(bool,X44)) | p(s(bool,X45)) | s(bool,X45) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.FILTER_2'(s(fun(sK12,bool),X46),s('type.list.list'(sK12),sK14)))))) ) | $spl2),
  inference(resolution,[],[f8325,f585])).
fof(f13556,plain,(
  ( ! [X39,X41,X42,X40] : (p(s(bool,X39)) | p(s(bool,X40)) | s(bool,X39) != s(bool,X41) | s(bool,X40) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.FILTER_2'(s(fun(sK12,bool),X42),s('type.list.list'(sK12),sK14)))))) ) | $spl2),
  inference(resolution,[],[f8325,f584])).
fof(f13555,plain,(
  ( ! [X37,X35,X38,X36] : (p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.FILTER_2'(s(fun(sK12,bool),X37),s('type.list.list'(sK12),sK14))))) | s(bool,X36) != s(bool,X38)) ) | $spl2),
  inference(resolution,[],[f8325,f582])).
fof(f13553,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,X27)) | p(s(bool,X28)) | s(bool,X27) != s(bool,X29) | s(bool,X28) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.FILTER_2'(s(fun(sK12,bool),X30),s('type.list.list'(sK12),sK14)))))) ) | $spl2),
  inference(resolution,[],[f8325,f580])).
fof(f13526,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12)))) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) ) | $spl1),
  inference(resolution,[],[f8516,f12608])).
fof(f12608,plain,(
  ( ! [X109,X107,X108] : (sP4(X108,X107,X109) | s(bool,X107) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12))))) ) | $spl1),
  inference(resolution,[],[f9071,f697])).
fof(f697,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f496])).
fof(f496,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f495])).
fof(f495,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(flattening,[],[f494])).
fof(f494,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(nnf_transformation,[],[f438])).
fof(f438,plain,(
  ! [X2,X0,X1] : (sP4(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f8516,plain,(
  ( ! [X80,X79] : (~sP4(X79,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))),X80) | ~p(s(bool,X80)) | ~p(s(bool,X79))) ) | $spl1),
  inference(resolution,[],[f863,f695])).
fof(f695,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f496])).
fof(f13525,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) ) | $spl1),
  inference(resolution,[],[f8516,f12033])).
fof(f12033,plain,(
  ( ! [X109,X107,X108] : (sP4(X108,X107,X109) | s(bool,X107) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12))))) ) | $spl1),
  inference(resolution,[],[f8948,f697])).
fof(f13515,plain,(
  ( ! [X16] : (s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12)))) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))) | p(s(bool,X16))) ) | $spl1),
  inference(resolution,[],[f12652,f8556])).
fof(f8556,plain,(
  ( ! [X159,X160] : (~sP10(X160,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))),X159) | p(s(bool,X159))) ) | $spl1),
  inference(resolution,[],[f863,f745])).
fof(f13514,plain,(
  ( ! [X15] : (s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12)))) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))) | p(s(bool,X15))) ) | $spl1),
  inference(resolution,[],[f12652,f8558])).
fof(f8558,plain,(
  ( ! [X163,X164] : (~sP10(X163,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))),X164) | p(s(bool,X163))) ) | $spl1),
  inference(resolution,[],[f863,f746])).
fof(f746,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | ~sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f518])).
fof(f7994,plain,(
  s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12)))) | $spl64),
  inference(cnf_transformation,[],[f7994_D])).
fof(f7994_D,plain,(
  s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12)))) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f13511,plain,(
  ( ! [X12] : (s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12)))) | p(s(bool,X12))) ) | $spl1),
  inference(resolution,[],[f12652,f8997])).
fof(f8997,plain,(
  ( ! [X154,X155] : (~sP10(X155,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12))),X154) | p(s(bool,X154))) ) | $spl1),
  inference(resolution,[],[f8798,f745])).
fof(f13510,plain,(
  ( ! [X11] : (s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12)))) | p(s(bool,X11))) ) | $spl1),
  inference(resolution,[],[f12652,f8999])).
fof(f8999,plain,(
  ( ! [X158,X159] : (~sP10(X158,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12))),X159) | p(s(bool,X158))) ) | $spl1),
  inference(resolution,[],[f8798,f746])).
fof(f13468,plain,(
  ( ! [X14,X15] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12)))) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) ) | $spl1),
  inference(resolution,[],[f8502,f12590])).
fof(f12590,plain,(
  ( ! [X61,X59,X60] : (sP0(X60,X61,X59) | s(bool,X59) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12))))) ) | $spl1),
  inference(resolution,[],[f9071,f603])).
fof(f8502,plain,(
  ( ! [X57,X58] : (~sP0(X57,X58,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))) | p(s(bool,X58)) | p(s(bool,X57))) ) | $spl1),
  inference(resolution,[],[f863,f602])).
fof(f602,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP0(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f464])).
fof(f13465,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | p(s(bool,X9)) | s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) ) | $spl1),
  inference(resolution,[],[f8502,f12015])).
fof(f12015,plain,(
  ( ! [X61,X59,X60] : (sP0(X60,X61,X59) | s(bool,X59) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12))))) ) | $spl1),
  inference(resolution,[],[f8948,f603])).
fof(f13458,plain,(
  ( ! [X21,X20] : (s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12)))) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))) | p(s(bool,X20)) | p(s(bool,X21))) ) | $spl1),
  inference(resolution,[],[f12592,f8501])).
fof(f8501,plain,(
  ( ! [X56,X55] : (~sP0('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))),X55,X56) | p(s(bool,X56)) | p(s(bool,X55))) ) | $spl1),
  inference(resolution,[],[f863,f602])).
fof(f12592,plain,(
  ( ! [X66,X67,X65] : (sP0(X65,X66,X67) | s(bool,X65) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12))))) ) | $spl1),
  inference(resolution,[],[f9071,f605])).
fof(f605,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f464])).
fof(f13448,plain,(
  ( ! [X19,X18] : (s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12)))) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))) | p(s(bool,X18)) | p(s(bool,X19))) ) | $spl1),
  inference(resolution,[],[f12591,f8500])).
fof(f8500,plain,(
  ( ! [X54,X53] : (~sP0(X53,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))),X54) | p(s(bool,X54)) | p(s(bool,X53))) ) | $spl1),
  inference(resolution,[],[f863,f602])).
fof(f12591,plain,(
  ( ! [X64,X62,X63] : (sP0(X63,X62,X64) | s(bool,X62) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12))))) ) | $spl1),
  inference(resolution,[],[f9071,f604])).
fof(f604,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f464])).
fof(f13429,plain,(
  ( ! [X19,X18] : (s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12)))) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))) | ~p(s(bool,X18)) | ~p(s(bool,X19))) ) | $spl1),
  inference(resolution,[],[f12589,f8497])).
fof(f8497,plain,(
  ( ! [X47,X48] : (~sP1(X47,X48,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))) | ~p(s(bool,X48)) | ~p(s(bool,X47))) ) | $spl1),
  inference(resolution,[],[f863,f598])).
fof(f598,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP1(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f462])).
fof(f12589,plain,(
  ( ! [X57,X58,X56] : (sP1(X57,X58,X56) | s(bool,X56) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12))))) ) | $spl1),
  inference(resolution,[],[f9071,f599])).
fof(f599,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f462])).
fof(f13417,plain,(
  ( ! [X19,X18] : (s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12)))) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))) | ~p(s(bool,X18)) | ~p(s(bool,X19))) ) | $spl1),
  inference(resolution,[],[f12586,f8494])).
fof(f8494,plain,(
  ( ! [X41,X42] : (~sP2(X41,X42,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))) | ~p(s(bool,X42)) | ~p(s(bool,X41))) ) | $spl1),
  inference(resolution,[],[f863,f594])).
fof(f594,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f459])).
fof(f12586,plain,(
  ( ! [X47,X48,X49] : (sP2(X48,X49,X47) | s(bool,X47) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12))))) ) | $spl1),
  inference(resolution,[],[f9071,f595])).
fof(f595,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f459])).
fof(f8163,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12)))) | s(bool,X0) = s(bool,X1) | ~p(s(bool,X1))) ) | $spl70),
  inference(cnf_transformation,[],[f8163_D])).
fof(f8163_D,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12)))) | s(bool,X0) = s(bool,X1) | ~p(s(bool,X1))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f13411,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12)))) | ~p(s(bool,X1)) | s(bool,X0) = s(bool,X1) | ~p(s(bool,X2))) ) | $spl1),
  inference(resolution,[],[f12586,f13298])).
fof(f13298,plain,(
  ( ! [X2,X0,X1] : (~sP2(X1,X0,X2) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X2) | ~p(s(bool,X1))) ) | $spl1),
  inference(subsumption_resolution,[],[f13277,f599])).
fof(f13277,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP2(X1,X0,X2) | ~sP1(X2,X0,X1) | s(bool,X0) = s(bool,X2) | ~p(s(bool,X1))) ) | $spl1),
  inference(resolution,[],[f13255,f608])).
fof(f13255,plain,(
  ( ! [X10,X11,X9] : (sP3(X10,X11,X9) | ~p(s(bool,X9)) | ~sP2(X10,X9,X11) | ~sP1(X11,X9,X10)) ) | $spl1),
  inference(subsumption_resolution,[],[f13252,f8946])).
fof(f8946,plain,(
  ( ! [X58] : (s(bool,X58) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) | ~p(s(bool,X58))) ) | $spl1),
  inference(resolution,[],[f8798,f613])).
fof(f613,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f468])).
fof(f13252,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) | sP3(X10,X11,X9) | ~p(s(bool,X9)) | ~sP2(X10,X9,X11) | ~sP1(X11,X9,X10)) ) | $spl1),
  inference(resolution,[],[f12017,f591])).
fof(f591,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2)) | ~sP2(X0,X2,X1) | ~sP1(X1,X2,X0) | ~sP0(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f456])).
fof(f12017,plain,(
  ( ! [X66,X67,X65] : (sP0(X65,X66,X67) | s(bool,X65) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12))))) ) | $spl1),
  inference(resolution,[],[f8948,f605])).
fof(f13406,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) ) | $spl1),
  inference(resolution,[],[f8501,f12017])).
fof(f13381,plain,(
  ( ! [X16] : (s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))) | p(s(bool,X16))) ) | $spl1),
  inference(resolution,[],[f12077,f8556])).
fof(f13380,plain,(
  ( ! [X15] : (s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))) | p(s(bool,X15))) ) | $spl1),
  inference(resolution,[],[f12077,f8558])).
fof(f13379,plain,(
  ( ! [X14] : (s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12)))) | p(s(bool,X14))) ) | $spl1),
  inference(resolution,[],[f12077,f9120])).
fof(f9120,plain,(
  ( ! [X154,X155] : (~sP10(X155,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12))),X154) | p(s(bool,X154))) ) | $spl1),
  inference(resolution,[],[f8799,f745])).
fof(f13378,plain,(
  ( ! [X13] : (s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12)))) | p(s(bool,X13))) ) | $spl1),
  inference(resolution,[],[f12077,f9122])).
fof(f9122,plain,(
  ( ! [X158,X159] : (~sP10(X158,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12))),X159) | p(s(bool,X158))) ) | $spl1),
  inference(resolution,[],[f8799,f746])).
fof(f13357,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) ) | $spl1),
  inference(resolution,[],[f8500,f12016])).
fof(f12016,plain,(
  ( ! [X64,X62,X63] : (sP0(X63,X62,X64) | s(bool,X62) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12))))) ) | $spl1),
  inference(resolution,[],[f8948,f604])).
fof(f7640,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) | s(bool,X0) = s(bool,X1) | ~p(s(bool,X1))) ) | $spl62),
  inference(cnf_transformation,[],[f7640_D])).
fof(f7640_D,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) | s(bool,X0) = s(bool,X1) | ~p(s(bool,X1))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f13314,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,X18)) | s(bool,X18) = s(bool,X19) | ~p(s(bool,X20)) | s(bool,X19) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12))))) ) | $spl1),
  inference(resolution,[],[f13298,f12011])).
fof(f12011,plain,(
  ( ! [X47,X48,X49] : (sP2(X48,X49,X47) | s(bool,X47) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12))))) ) | $spl1),
  inference(resolution,[],[f8948,f595])).
fof(f13264,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) ) | $spl1),
  inference(resolution,[],[f8497,f12014])).
fof(f12014,plain,(
  ( ! [X57,X58,X56] : (sP1(X57,X58,X56) | s(bool,X56) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12))))) ) | $spl1),
  inference(resolution,[],[f8948,f599])).
fof(f13218,plain,(
  ( ! [X15,X16] : (s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))) | ~p(s(bool,X15)) | ~p(s(bool,X16))) ) | $spl1),
  inference(resolution,[],[f12011,f8494])).
fof(f13131,plain,(
  ( ! [X29,X27] : (~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),X29),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))))) ) | $spl114),
  inference(cnf_transformation,[],[f13131_D])).
fof(f13131_D,plain,(
  ( ! [X29,X27] : (~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),X29),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))))) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f13092,plain,(
  ( ! [X310,X312,X311] : (p(s(bool,X310)) | ~p(s(bool,X311)) | s(bool,X311) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),X312),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))))) ) | $spl1),
  inference(resolution,[],[f8470,f839])).
fof(f8470,plain,(
  ( ! [X1] : (~p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),X1),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))))))) ) | $spl1),
  inference(resolution,[],[f863,f648])).
fof(f648,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X3)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X3))))))) )),
  inference(cnf_transformation,[],[f476])).
fof(f476,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X3)))))) | (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) & p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X3)))))) & (~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X3)))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X3))))))))),
  inference(flattening,[],[f475])).
fof(f475,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X3)))))) | (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) & p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X3)))))) & ((~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X3))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X3))))))))),
  inference(nnf_transformation,[],[f315])).
fof(f315,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X3)))))) <=> (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) & p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X3))))))),
  inference(flattening,[],[f314])).
fof(f314,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X3)))))) <=> (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) & p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X3))))))),
  inference(rectify,[],[f76])).
fof(f76,axiom,(
  ! [X6] : ! [X7,X13,X14] : (p(s(bool,'const.list.EVERY_2'(s(fun(X6,bool),X7),s('type.list.list'(X6),'const.list.APPEND_2'(s('type.list.list'(X6),X13),s('type.list.list'(X6),X14)))))) <=> (p(s(bool,'const.list.EVERY_2'(s(fun(X6,bool),X7),s('type.list.list'(X6),X13)))) & p(s(bool,'const.list.EVERY_2'(s(fun(X6,bool),X7),s('type.list.list'(X6),X14))))))),
  file('thm.rich_list.EVERY_REVERSE','thm.list.EVERY_APPEND')).
fof(f13004,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X54) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),X55),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))))) | s(bool,X53) != s(bool,X56)) ) | $spl1),
  inference(resolution,[],[f8470,f586])).
fof(f13002,plain,(
  ( ! [X47,X48,X46,X49] : (p(s(bool,X46)) | ~p(s(bool,X47)) | ~p(s(bool,X48)) | s(bool,X47) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),X49),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))))) ) | $spl1),
  inference(resolution,[],[f8470,f585])).
fof(f13000,plain,(
  ( ! [X39,X41,X38,X40] : (p(s(bool,X38)) | ~p(s(bool,X39)) | s(bool,X38) != s(bool,X40) | s(bool,X39) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),X41),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))))) ) | $spl1),
  inference(resolution,[],[f8470,f584])).
fof(f12999,plain,(
  ( ! [X37,X35,X36,X34] : (~p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X35) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),X36),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))))) | s(bool,X34) != s(bool,X37)) ) | $spl1),
  inference(resolution,[],[f8470,f582])).
fof(f12997,plain,(
  ( ! [X28,X26,X29,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,X28) | s(bool,X27) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),X29),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))))) ) | $spl1),
  inference(resolution,[],[f8470,f580])).
fof(f12923,plain,(
  ( ! [X29,X27] : (~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))),s('type.list.list'(sK12),X29)))))) ) | $spl112),
  inference(cnf_transformation,[],[f12923_D])).
fof(f12923_D,plain,(
  ( ! [X29,X27] : (~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))),s('type.list.list'(sK12),X29)))))) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f12883,plain,(
  ( ! [X310,X312,X311] : (p(s(bool,X310)) | ~p(s(bool,X311)) | s(bool,X311) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))),s('type.list.list'(sK12),X312)))))) ) | $spl1),
  inference(resolution,[],[f8469,f839])).
fof(f8469,plain,(
  ( ! [X0] : (~p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))),s('type.list.list'(sK12),X0))))))) ) | $spl1),
  inference(resolution,[],[f863,f647])).
fof(f647,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X3))))))) )),
  inference(cnf_transformation,[],[f476])).
fof(f12795,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X54) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))),s('type.list.list'(sK12),X55))))) | s(bool,X53) != s(bool,X56)) ) | $spl1),
  inference(resolution,[],[f8469,f586])).
fof(f12793,plain,(
  ( ! [X47,X48,X46,X49] : (p(s(bool,X46)) | ~p(s(bool,X47)) | ~p(s(bool,X48)) | s(bool,X47) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))),s('type.list.list'(sK12),X49)))))) ) | $spl1),
  inference(resolution,[],[f8469,f585])).
fof(f12791,plain,(
  ( ! [X39,X41,X38,X40] : (p(s(bool,X38)) | ~p(s(bool,X39)) | s(bool,X38) != s(bool,X40) | s(bool,X39) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))),s('type.list.list'(sK12),X41)))))) ) | $spl1),
  inference(resolution,[],[f8469,f584])).
fof(f12790,plain,(
  ( ! [X37,X35,X36,X34] : (~p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X35) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))),s('type.list.list'(sK12),X36))))) | s(bool,X34) != s(bool,X37)) ) | $spl1),
  inference(resolution,[],[f8469,f582])).
fof(f12788,plain,(
  ( ! [X28,X26,X29,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,X28) | s(bool,X27) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))),s('type.list.list'(sK12),X29)))))) ) | $spl1),
  inference(resolution,[],[f8469,f580])).
fof(f7588,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12)))) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) | $spl60),
  inference(cnf_transformation,[],[f7588_D])).
fof(f7588_D,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12)))) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f12669,plain,(
  ( ! [X288,X290,X289] : (s(bool,X288) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12)))) | p(s(bool,X289)) | p(s(bool,X290)) | s(bool,X288) != s(bool,X290)) ) | $spl1),
  inference(resolution,[],[f9071,f839])).
fof(f12580,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12)))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) ) | $spl1),
  inference(resolution,[],[f9071,f586])).
fof(f12577,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,X16) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12)))) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) ) | $spl1),
  inference(resolution,[],[f9071,f585])).
fof(f12576,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12)))) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl1),
  inference(resolution,[],[f9071,f584])).
fof(f12575,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12)))) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) ) | $spl1),
  inference(resolution,[],[f9071,f582])).
fof(f12573,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12)))) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl1),
  inference(resolution,[],[f9071,f580])).
fof(f12497,plain,(
  ( ! [X28,X26] : (~p(s(bool,X26)) | s(bool,X26) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X28),s('type.list.list'(sK12),sK14)))))))) ) | $spl110),
  inference(cnf_transformation,[],[f12497_D])).
fof(f12497_D,plain,(
  ( ! [X28,X26] : (~p(s(bool,X26)) | s(bool,X26) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X28),s('type.list.list'(sK12),sK14)))))))) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f12466,plain,(
  ( ! [X310,X309,X311] : (p(s(bool,X309)) | ~p(s(bool,X310)) | s(bool,X310) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X311),s('type.list.list'(sK12),sK14)))))))) ) | $spl1),
  inference(resolution,[],[f8602,f839])).
fof(f8602,plain,(
  ( ! [X3] : (~p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X3),s('type.list.list'(sK12),sK14))))))))) ) | $spl1),
  inference(forward_demodulation,[],[f8472,f637])).
fof(f637,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'(X0),'const.list.REVERSE_1'(s('type.list.list'(X0),'const.list.SNOC_2'(s(X0,X1),s('type.list.list'(X0),X2))))) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),'const.list.REVERSE_1'(s('type.list.list'(X0),X2)))))) )),
  inference(cnf_transformation,[],[f299])).
fof(f299,plain,(
  ! [X0,X1,X2] : s('type.list.list'(X0),'const.list.REVERSE_1'(s('type.list.list'(X0),'const.list.SNOC_2'(s(X0,X1),s('type.list.list'(X0),X2))))) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),'const.list.REVERSE_1'(s('type.list.list'(X0),X2)))))),
  inference(flattening,[],[f298])).
fof(f298,plain,(
  ! [X0] : ! [X1,X2] : s('type.list.list'(X0),'const.list.REVERSE_1'(s('type.list.list'(X0),'const.list.SNOC_2'(s(X0,X1),s('type.list.list'(X0),X2))))) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),'const.list.REVERSE_1'(s('type.list.list'(X0),X2)))))),
  inference(rectify,[],[f224])).
fof(f224,axiom,(
  ! [X6] : ! [X12,X8] : s('type.list.list'(X6),'const.list.REVERSE_1'(s('type.list.list'(X6),'const.list.SNOC_2'(s(X6,X12),s('type.list.list'(X6),X8))))) = s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X12),s('type.list.list'(X6),'const.list.REVERSE_1'(s('type.list.list'(X6),X8)))))),
  file('thm.rich_list.EVERY_REVERSE','thm.list.REVERSE_SNOC')).
fof(f8472,plain,(
  ( ! [X3] : (~p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X3),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))))))) ) | $spl1),
  inference(resolution,[],[f863,f654])).
fof(f654,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X3)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3))))))) )),
  inference(cnf_transformation,[],[f480])).
fof(f480,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3)))))) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) & p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X3)))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X3)))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3))))))))),
  inference(flattening,[],[f479])).
fof(f479,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3)))))) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) & p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X3)))))) & ((~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X3))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3))))))))),
  inference(nnf_transformation,[],[f319])).
fof(f319,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3)))))) <=> (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) & p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X3))))))),
  inference(flattening,[],[f318])).
fof(f318,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3)))))) <=> (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) & p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X3))))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6] : ! [X7,X9,X10] : (p(s(bool,'const.list.EVERY_2'(s(fun(X6,bool),X7),s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X9),s('type.list.list'(X6),X10)))))) <=> (p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X9)))) & p(s(bool,'const.list.EVERY_2'(s(fun(X6,bool),X7),s('type.list.list'(X6),X10))))))),
  file('thm.rich_list.EVERY_REVERSE','thm.list.EVERY_DEF|split|1')).
fof(f12378,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,X52)) | ~p(s(bool,X53)) | s(bool,X53) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X54),s('type.list.list'(sK12),sK14))))))) | s(bool,X52) != s(bool,X55)) ) | $spl1),
  inference(resolution,[],[f8602,f586])).
fof(f12376,plain,(
  ( ! [X47,X45,X48,X46] : (p(s(bool,X45)) | ~p(s(bool,X46)) | ~p(s(bool,X47)) | s(bool,X46) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X48),s('type.list.list'(sK12),sK14)))))))) ) | $spl1),
  inference(resolution,[],[f8602,f585])).
fof(f12374,plain,(
  ( ! [X39,X37,X38,X40] : (p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,X39) | s(bool,X38) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK14)))))))) ) | $spl1),
  inference(resolution,[],[f8602,f584])).
fof(f12373,plain,(
  ( ! [X35,X33,X36,X34] : (~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X35),s('type.list.list'(sK12),sK14))))))) | s(bool,X33) != s(bool,X36)) ) | $spl1),
  inference(resolution,[],[f8602,f582])).
fof(f12371,plain,(
  ( ! [X28,X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X28),s('type.list.list'(sK12),sK14)))))))) ) | $spl1),
  inference(resolution,[],[f8602,f580])).
fof(f12336,plain,(
  ( ! [X28,X26] : (~p(s(bool,X26)) | s(bool,X26) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X28),s('type.list.list'(sK12),sK14)))))))) ) | $spl108),
  inference(cnf_transformation,[],[f12336_D])).
fof(f12336_D,plain,(
  ( ! [X28,X26] : (~p(s(bool,X26)) | s(bool,X26) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X28),s('type.list.list'(sK12),sK14)))))))) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f12305,plain,(
  ( ! [X310,X309,X311] : (p(s(bool,X309)) | ~p(s(bool,X310)) | s(bool,X310) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X311),s('type.list.list'(sK12),sK14)))))))) ) | $spl1),
  inference(resolution,[],[f8601,f839])).
fof(f8601,plain,(
  ( ! [X2] : (~p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X2),s('type.list.list'(sK12),sK14))))))))) ) | $spl1),
  inference(forward_demodulation,[],[f8471,f638])).
fof(f638,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'(X0),'const.list.REVERSE_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2))))) = s('type.list.list'(X0),'const.list.SNOC_2'(s(X0,X1),s('type.list.list'(X0),'const.list.REVERSE_1'(s('type.list.list'(X0),X2)))))) )),
  inference(cnf_transformation,[],[f301])).
fof(f301,plain,(
  ! [X0,X1,X2] : s('type.list.list'(X0),'const.list.REVERSE_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2))))) = s('type.list.list'(X0),'const.list.SNOC_2'(s(X0,X1),s('type.list.list'(X0),'const.list.REVERSE_1'(s('type.list.list'(X0),X2)))))),
  inference(flattening,[],[f300])).
fof(f300,plain,(
  ! [X0] : ! [X1,X2] : s('type.list.list'(X0),'const.list.REVERSE_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2))))) = s('type.list.list'(X0),'const.list.SNOC_2'(s(X0,X1),s('type.list.list'(X0),'const.list.REVERSE_1'(s('type.list.list'(X0),X2)))))),
  inference(rectify,[],[f225])).
fof(f225,axiom,(
  ! [X6] : ! [X12,X8] : s('type.list.list'(X6),'const.list.REVERSE_1'(s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X12),s('type.list.list'(X6),X8))))) = s('type.list.list'(X6),'const.list.SNOC_2'(s(X6,X12),s('type.list.list'(X6),'const.list.REVERSE_1'(s('type.list.list'(X6),X8)))))),
  file('thm.rich_list.EVERY_REVERSE','thm.list.REVERSE_SNOC_DEF|split|1')).
fof(f8471,plain,(
  ( ! [X2] : (~p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X2),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))))))) ) | $spl1),
  inference(resolution,[],[f863,f650])).
fof(f650,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X3)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.SNOC_2'(s(X0,X2),s('type.list.list'(X0),X3))))))) )),
  inference(cnf_transformation,[],[f478])).
fof(f478,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.SNOC_2'(s(X0,X2),s('type.list.list'(X0),X3)))))) | (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X3)))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))) & (~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.SNOC_2'(s(X0,X2),s('type.list.list'(X0),X3))))))))),
  inference(flattening,[],[f477])).
fof(f477,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.SNOC_2'(s(X0,X2),s('type.list.list'(X0),X3)))))) | (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X3)))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))) & ((~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.SNOC_2'(s(X0,X2),s('type.list.list'(X0),X3))))))))),
  inference(nnf_transformation,[],[f317])).
fof(f317,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.SNOC_2'(s(X0,X2),s('type.list.list'(X0),X3)))))) <=> (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X3)))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(flattening,[],[f316])).
fof(f316,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.SNOC_2'(s(X0,X2),s('type.list.list'(X0),X3)))))) <=> (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X3)))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X6] : ! [X7,X12,X8] : (p(s(bool,'const.list.EVERY_2'(s(fun(X6,bool),X7),s('type.list.list'(X6),'const.list.SNOC_2'(s(X6,X12),s('type.list.list'(X6),X8)))))) <=> (p(s(bool,'const.list.EVERY_2'(s(fun(X6,bool),X7),s('type.list.list'(X6),X8)))) & p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X12))))))),
  file('thm.rich_list.EVERY_REVERSE','thm.list.EVERY_SNOC')).
fof(f12217,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,X52)) | ~p(s(bool,X53)) | s(bool,X53) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X54),s('type.list.list'(sK12),sK14))))))) | s(bool,X52) != s(bool,X55)) ) | $spl1),
  inference(resolution,[],[f8601,f586])).
fof(f12215,plain,(
  ( ! [X47,X45,X48,X46] : (p(s(bool,X45)) | ~p(s(bool,X46)) | ~p(s(bool,X47)) | s(bool,X46) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X48),s('type.list.list'(sK12),sK14)))))))) ) | $spl1),
  inference(resolution,[],[f8601,f585])).
fof(f12213,plain,(
  ( ! [X39,X37,X38,X40] : (p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,X39) | s(bool,X38) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X40),s('type.list.list'(sK12),sK14)))))))) ) | $spl1),
  inference(resolution,[],[f8601,f584])).
fof(f12212,plain,(
  ( ! [X35,X33,X36,X34] : (~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X35),s('type.list.list'(sK12),sK14))))))) | s(bool,X33) != s(bool,X36)) ) | $spl1),
  inference(resolution,[],[f8601,f582])).
fof(f12210,plain,(
  ( ! [X28,X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X28),s('type.list.list'(sK12),sK14)))))))) ) | $spl1),
  inference(resolution,[],[f8601,f580])).
fof(f7241,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) | $spl58),
  inference(cnf_transformation,[],[f7241_D])).
fof(f7241_D,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f12094,plain,(
  ( ! [X288,X290,X289] : (s(bool,X288) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) | p(s(bool,X289)) | p(s(bool,X290)) | s(bool,X288) != s(bool,X290)) ) | $spl1),
  inference(resolution,[],[f8948,f839])).
fof(f12005,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) ) | $spl1),
  inference(resolution,[],[f8948,f586])).
fof(f12002,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,X16) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) ) | $spl1),
  inference(resolution,[],[f8948,f585])).
fof(f12001,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl1),
  inference(resolution,[],[f8948,f584])).
fof(f12000,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) ) | $spl1),
  inference(resolution,[],[f8948,f582])).
fof(f11998,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl1),
  inference(resolution,[],[f8948,f580])).
fof(f11766,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) | $spl106),
  inference(cnf_transformation,[],[f11766_D])).
fof(f11766_D,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f11669,plain,(
  ( ! [X288,X290,X289] : (s(bool,X288) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))) | p(s(bool,X289)) | p(s(bool,X290)) | s(bool,X288) != s(bool,X290)) ) | $spl1),
  inference(resolution,[],[f8507,f839])).
fof(f8507,plain,(
  ( ! [X65] : (p(s(bool,X65)) | s(bool,X65) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) ) | $spl1),
  inference(resolution,[],[f863,f615])).
fof(f11580,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) ) | $spl1),
  inference(resolution,[],[f8507,f586])).
fof(f11577,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,X16) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) ) | $spl1),
  inference(resolution,[],[f8507,f585])).
fof(f11576,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl1),
  inference(resolution,[],[f8507,f584])).
fof(f11575,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) ) | $spl1),
  inference(resolution,[],[f8507,f582])).
fof(f11573,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl1),
  inference(resolution,[],[f8507,f580])).
fof(f11245,plain,(
  $false | ($spl6 | $spl18)),
  inference(equality_resolution,[],[f11177])).
fof(f11177,plain,(
  ( ! [X14] : (s(bool,X14) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) ) | ($spl6 | $spl18)),
  inference(subsumption_resolution,[],[f1512,f1003])).
fof(f11204,plain,(
  ( ! [X33,X31] : (p(s(bool,X31)) | p(s(bool,X33)) | s(bool,X31) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK30(sK13,X33,sK12))))) ) | $spl104),
  inference(cnf_transformation,[],[f11204_D])).
fof(f11204_D,plain,(
  ( ! [X33,X31] : (p(s(bool,X31)) | p(s(bool,X33)) | s(bool,X31) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK30(sK13,X33,sK12))))) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f11190,plain,(
  ( ! [X33,X31,X32] : (p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK30(sK13,X33,sK12)))) | p(s(bool,X33))) ) | ($spl1 | $spl2)),
  inference(resolution,[],[f11162,f10347])).
fof(f10347,plain,(
  ( ! [X59,X60,X58] : (sP2(X59,happ(s(fun(sK12,bool),sK13),s(sK12,sK30(sK13,X58,sK12))),X60) | p(s(bool,X58))) ) | $spl1),
  inference(resolution,[],[f8797,f596])).
fof(f8797,plain,(
  ( ! [X5] : (~p(s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK30(sK13,X5,sK12))))) | p(s(bool,X5))) ) | $spl1),
  inference(resolution,[],[f8480,f803])).
fof(f803,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK30(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f546])).
fof(f546,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK29(X2,X1,X0))))) & ~p(s(bool,X1))) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) | p(s(bool,X1))) & ((~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK30(X2,X1,X0))))) & ~p(s(bool,X1))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK29,sK30])],[f545])).
fof(f545,plain,(
  ! [X0,X1,X2] : ((? [X6] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) & ~p(s(bool,X1))) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) | p(s(bool,X1))) & ((? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) & ~p(s(bool,X1))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1)))))),
  inference(rectify,[],[f544])).
fof(f544,plain,(
  ! [X0,X1,X2] : ((? [X3] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) & ~p(s(bool,X1))) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,X1))) & ((? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) & ~p(s(bool,X1))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1)))))),
  inference(flattening,[],[f543])).
fof(f543,plain,(
  ! [X0,X1,X2] : ((? [X3] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) & ~p(s(bool,X1))) | (! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,X1)))) & ((? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) & ~p(s(bool,X1))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1)))))),
  inference(nnf_transformation,[],[f367])).
fof(f367,plain,(
  ! [X0,X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1))) <=> (! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,X1))))),
  inference(flattening,[],[f366])).
fof(f366,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1))) <=> (! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,X1))))),
  inference(rectify,[],[f252])).
fof(f252,axiom,(
  ! [X6] : ! [X39,X7] : (! [X12] : (p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X12)))) | p(s(bool,X39))) <=> (! [X12] : p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X12)))) | p(s(bool,X39))))),
  file('thm.rich_list.EVERY_REVERSE','thm.bool.LEFT_FORALL_OR_THM')).
fof(f11203,plain,(
  ( ! [X30,X28] : (p(s(bool,X28)) | p(s(bool,X30)) | s(bool,X28) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK29(sK13,X30,sK12))))) ) | $spl102),
  inference(cnf_transformation,[],[f11203_D])).
fof(f11203_D,plain,(
  ( ! [X30,X28] : (p(s(bool,X28)) | p(s(bool,X30)) | s(bool,X28) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK29(sK13,X30,sK12))))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f11189,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X28) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK29(sK13,X30,sK12)))) | p(s(bool,X30))) ) | ($spl1 | $spl2)),
  inference(resolution,[],[f11162,f10224])).
fof(f10224,plain,(
  ( ! [X59,X60,X58] : (sP2(X59,happ(s(fun(sK12,bool),sK13),s(sK12,sK29(sK13,X58,sK12))),X60) | p(s(bool,X58))) ) | $spl1),
  inference(resolution,[],[f8796,f596])).
fof(f8796,plain,(
  ( ! [X4] : (~p(s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK29(sK13,X4,sK12))))) | p(s(bool,X4))) ) | $spl1),
  inference(resolution,[],[f8480,f801])).
fof(f801,plain,(
  ( ! [X2,X0,X5,X1] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK29(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f546])).
fof(f11202,plain,(
  ( ! [X26] : (p(s(bool,X26)) | s(bool,X26) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20('const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)),sK13,sK12))))) ) | $spl100),
  inference(cnf_transformation,[],[f11202_D])).
fof(f11202_D,plain,(
  ( ! [X26] : (p(s(bool,X26)) | s(bool,X26) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20('const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)),sK13,sK12))))) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f11188,plain,(
  ( ! [X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20('const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)),sK13,sK12))))) ) | ($spl1 | $spl2)),
  inference(resolution,[],[f11162,f8812])).
fof(f8812,plain,(
  ( ! [X39,X38] : (sP2(X38,happ(s(fun(sK12,bool),sK13),s(sK12,sK20('const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)),sK13,sK12))),X39)) ) | $spl1),
  inference(resolution,[],[f8480,f596])).
fof(f11201,plain,(
  ( ! [X24] : (p(s(bool,X24)) | s(bool,X24) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12))))) ) | $spl98),
  inference(cnf_transformation,[],[f11201_D])).
fof(f11201_D,plain,(
  ( ! [X24] : (p(s(bool,X24)) | s(bool,X24) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12))))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f11187,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12))))) ) | ($spl1 | $spl2)),
  inference(resolution,[],[f11162,f9059])).
fof(f9059,plain,(
  ( ! [X39,X38] : (sP2(X38,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12))),X39)) ) | $spl1),
  inference(resolution,[],[f8799,f596])).
fof(f11200,plain,(
  ( ! [X22] : (p(s(bool,X22)) | s(bool,X22) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12))))) ) | $spl96),
  inference(cnf_transformation,[],[f11200_D])).
fof(f11200_D,plain,(
  ( ! [X22] : (p(s(bool,X22)) | s(bool,X22) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12))))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f11186,plain,(
  ( ! [X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12))))) ) | ($spl1 | $spl2)),
  inference(resolution,[],[f11162,f8936])).
fof(f8936,plain,(
  ( ! [X39,X38] : (sP2(X38,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12))),X39)) ) | $spl1),
  inference(resolution,[],[f8798,f596])).
fof(f11199,plain,(
  ( ! [X20] : (p(s(bool,X20)) | s(bool,X20) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) ) | $spl94),
  inference(cnf_transformation,[],[f11199_D])).
fof(f11199_D,plain,(
  ( ! [X20] : (p(s(bool,X20)) | s(bool,X20) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f11185,plain,(
  ( ! [X21,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X20) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) ) | ($spl1 | $spl2)),
  inference(resolution,[],[f11162,f893])).
fof(f893,plain,(
  ( ! [X43,X44] : (sP2(X43,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))),X44)) ) | $spl1),
  inference(resolution,[],[f863,f596])).
fof(f11198,plain,(
  ( ! [X14,X12] : (p(s(bool,X12)) | s(bool,X14) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) | s(bool,X12) = s(bool,X14)) ) | $spl92),
  inference(cnf_transformation,[],[f11198_D])).
fof(f11198_D,plain,(
  ( ! [X14,X12] : (p(s(bool,X12)) | s(bool,X14) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) | s(bool,X12) = s(bool,X14)) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f11182,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) = s(bool,X14) | s(bool,X14) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)))) ) | $spl2),
  inference(resolution,[],[f11162,f10690])).
fof(f10690,plain,(
  ( ! [X45,X46,X44] : (sP2(X45,X44,X46) | s(bool,X44) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)))) ) | $spl2),
  inference(resolution,[],[f8354,f596])).
fof(f10865,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) | $spl88),
  inference(cnf_transformation,[],[f10865_D])).
fof(f10865_D,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f11134,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X0) != s(bool,X2)) ) | $spl2),
  inference(resolution,[],[f10694,f11036])).
fof(f11036,plain,(
  ( ! [X4,X5,X3] : (~sP1(X3,X4,X5) | p(s(bool,X5)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X4)) ) | $spl2),
  inference(duplicate_literal_removal,[],[f11028])).
fof(f11028,plain,(
  ( ! [X4,X5,X3] : (~sP1(X3,X4,X5) | p(s(bool,X5)) | ~p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X3) != s(bool,X4)) ) | $spl2),
  inference(resolution,[],[f11012,f609])).
fof(f609,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | s(bool,X0) != s(bool,X1) | ~sP3(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f465])).
fof(f11012,plain,(
  ( ! [X10,X11,X9] : (sP3(X9,X10,X11) | ~sP1(X10,X11,X9) | p(s(bool,X9)) | ~p(s(bool,X11))) ) | $spl2),
  inference(resolution,[],[f10889,f605])).
fof(f10889,plain,(
  ( ! [X10,X11,X9] : (~sP0(X11,X10,X9) | p(s(bool,X9)) | ~sP1(X10,X11,X9) | sP3(X9,X10,X11)) ) | $spl2),
  inference(subsumption_resolution,[],[f10886,f865])).
fof(f10886,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) | sP3(X9,X10,X11) | p(s(bool,X9)) | ~sP1(X10,X11,X9) | ~sP0(X11,X10,X9)) ) | $spl2),
  inference(resolution,[],[f10691,f593])).
fof(f593,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X0)) | ~sP2(X0,X2,X1) | ~sP1(X1,X2,X0) | ~sP0(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f456])).
fof(f10691,plain,(
  ( ! [X47,X48,X49] : (sP2(X47,X48,X49) | s(bool,X47) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)))) ) | $spl2),
  inference(resolution,[],[f8354,f597])).
fof(f597,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f459])).
fof(f11127,plain,(
  ( ! [X4,X5,X3] : (s(bool,X3) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) | p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X3) != s(bool,X5)) ) | $spl2),
  inference(resolution,[],[f10693,f11026])).
fof(f11026,plain,(
  ( ! [X4,X5,X3] : (~sP1(X3,X4,X5) | p(s(bool,X5)) | ~p(s(bool,X3)) | s(bool,X3) != s(bool,X4)) ) | $spl2),
  inference(duplicate_literal_removal,[],[f11017])).
fof(f11017,plain,(
  ( ! [X4,X5,X3] : (~sP1(X3,X4,X5) | p(s(bool,X5)) | ~p(s(bool,X3)) | p(s(bool,X5)) | s(bool,X3) != s(bool,X4)) ) | $spl2),
  inference(resolution,[],[f11011,f609])).
fof(f11011,plain,(
  ( ! [X6,X8,X7] : (sP3(X6,X7,X8) | ~sP1(X7,X8,X6) | p(s(bool,X6)) | ~p(s(bool,X7))) ) | $spl2),
  inference(resolution,[],[f10889,f604])).
fof(f10693,plain,(
  ( ! [X54,X55,X53] : (sP1(X54,X53,X55) | s(bool,X53) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)))) ) | $spl2),
  inference(resolution,[],[f8354,f600])).
fof(f600,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f462])).
fof(f3649,plain,(
  ( ! [X14,X16] : (p(s(bool,X14)) | ~p(s(bool,X16)) | s(bool,X16) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK30(sK13,X14,sK12))))) ) | $spl34),
  inference(cnf_transformation,[],[f3649_D])).
fof(f3649_D,plain,(
  ( ! [X14,X16] : (p(s(bool,X14)) | ~p(s(bool,X16)) | s(bool,X16) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK30(sK13,X14,sK12))))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f11070,plain,(
  ( ! [X39,X41,X40] : (p(s(bool,X39)) | ~p(s(bool,X40)) | s(bool,X40) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK30(sK13,X41,sK12)))) | p(s(bool,X41))) ) | ($spl1 | $spl2)),
  inference(resolution,[],[f11036,f10351])).
fof(f10351,plain,(
  ( ! [X70,X72,X71] : (sP1(happ(s(fun(sK12,bool),sK13),s(sK12,sK30(sK13,X70,sK12))),X71,X72) | p(s(bool,X70))) ) | $spl1),
  inference(resolution,[],[f8797,f601])).
fof(f3525,plain,(
  ( ! [X14,X16] : (p(s(bool,X14)) | ~p(s(bool,X16)) | s(bool,X16) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK29(sK13,X14,sK12))))) ) | $spl32),
  inference(cnf_transformation,[],[f3525_D])).
fof(f3525_D,plain,(
  ( ! [X14,X16] : (p(s(bool,X14)) | ~p(s(bool,X16)) | s(bool,X16) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK29(sK13,X14,sK12))))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f11069,plain,(
  ( ! [X37,X38,X36] : (p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X37) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK29(sK13,X38,sK12)))) | p(s(bool,X38))) ) | ($spl1 | $spl2)),
  inference(resolution,[],[f11036,f10228])).
fof(f10228,plain,(
  ( ! [X70,X72,X71] : (sP1(happ(s(fun(sK12,bool),sK13),s(sK12,sK29(sK13,X70,sK12))),X71,X72) | p(s(bool,X70))) ) | $spl1),
  inference(resolution,[],[f8796,f601])).
fof(f8914,plain,(
  ( ! [X7] : (~p(s(bool,X7)) | s(bool,X7) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20('const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)),sK13,sK12))))) ) | $spl76),
  inference(cnf_transformation,[],[f8914_D])).
fof(f8914_D,plain,(
  ( ! [X7] : (~p(s(bool,X7)) | s(bool,X7) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20('const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)),sK13,sK12))))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f11068,plain,(
  ( ! [X35,X34] : (p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X35) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20('const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)),sK13,sK12))))) ) | ($spl1 | $spl2)),
  inference(resolution,[],[f11036,f8816])).
fof(f8816,plain,(
  ( ! [X47,X46] : (sP1(happ(s(fun(sK12,bool),sK13),s(sK12,sK20('const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)),sK13,sK12))),X46,X47)) ) | $spl1),
  inference(resolution,[],[f8480,f601])).
fof(f2172,plain,(
  ( ! [X7] : (~p(s(bool,X7)) | s(bool,X7) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12))))) ) | $spl26),
  inference(cnf_transformation,[],[f2172_D])).
fof(f2172_D,plain,(
  ( ! [X7] : (~p(s(bool,X7)) | s(bool,X7) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12))))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f11067,plain,(
  ( ! [X33,X32] : (p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X33) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12))))) ) | ($spl1 | $spl2)),
  inference(resolution,[],[f11036,f9063])).
fof(f9063,plain,(
  ( ! [X47,X46] : (sP1(happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12))),X46,X47)) ) | $spl1),
  inference(resolution,[],[f8799,f601])).
fof(f2048,plain,(
  ( ! [X7] : (~p(s(bool,X7)) | s(bool,X7) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12))))) ) | $spl24),
  inference(cnf_transformation,[],[f2048_D])).
fof(f2048_D,plain,(
  ( ! [X7] : (~p(s(bool,X7)) | s(bool,X7) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12))))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f11066,plain,(
  ( ! [X30,X31] : (p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X31) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12))))) ) | ($spl1 | $spl2)),
  inference(resolution,[],[f11036,f8940])).
fof(f8940,plain,(
  ( ! [X47,X46] : (sP1(happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12))),X46,X47)) ) | $spl1),
  inference(resolution,[],[f8798,f601])).
fof(f1003,plain,(
  ( ! [X12] : (s(bool,X12) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))) | ~p(s(bool,X12))) ) | $spl6),
  inference(cnf_transformation,[],[f1003_D])).
fof(f1003_D,plain,(
  ( ! [X12] : (s(bool,X12) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))) | ~p(s(bool,X12))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f11065,plain,(
  ( ! [X28,X29] : (p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X29) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) ) | ($spl1 | $spl2)),
  inference(resolution,[],[f11036,f897])).
fof(f897,plain,(
  ( ! [X52,X51] : (sP1('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))),X51,X52)) ) | $spl1),
  inference(resolution,[],[f863,f601])).
fof(f11047,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X26) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK30(sK13,X27,sK12)))) | p(s(bool,X27))) ) | ($spl1 | $spl2)),
  inference(resolution,[],[f11026,f10350])).
fof(f10350,plain,(
  ( ! [X68,X69,X67] : (sP1(X68,happ(s(fun(sK12,bool),sK13),s(sK12,sK30(sK13,X67,sK12))),X69) | p(s(bool,X67))) ) | $spl1),
  inference(resolution,[],[f8797,f600])).
fof(f11046,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK29(sK13,X24,sK12)))) | p(s(bool,X24))) ) | ($spl1 | $spl2)),
  inference(resolution,[],[f11026,f10227])).
fof(f10227,plain,(
  ( ! [X68,X69,X67] : (sP1(X68,happ(s(fun(sK12,bool),sK13),s(sK12,sK29(sK13,X67,sK12))),X69) | p(s(bool,X67))) ) | $spl1),
  inference(resolution,[],[f8796,f600])).
fof(f11045,plain,(
  ( ! [X21,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X21) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20('const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)),sK13,sK12))))) ) | ($spl1 | $spl2)),
  inference(resolution,[],[f11026,f8815])).
fof(f8815,plain,(
  ( ! [X45,X44] : (sP1(X44,happ(s(fun(sK12,bool),sK13),s(sK12,sK20('const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)),sK13,sK12))),X45)) ) | $spl1),
  inference(resolution,[],[f8480,f600])).
fof(f11044,plain,(
  ( ! [X19,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12))))) ) | ($spl1 | $spl2)),
  inference(resolution,[],[f11026,f9062])).
fof(f9062,plain,(
  ( ! [X45,X44] : (sP1(X44,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12))),X45)) ) | $spl1),
  inference(resolution,[],[f8799,f600])).
fof(f11043,plain,(
  ( ! [X17,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X17) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12))))) ) | ($spl1 | $spl2)),
  inference(resolution,[],[f11026,f8939])).
fof(f8939,plain,(
  ( ! [X45,X44] : (sP1(X44,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12))),X45)) ) | $spl1),
  inference(resolution,[],[f8798,f600])).
fof(f11042,plain,(
  ( ! [X14,X15] : (p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X15) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) ) | ($spl1 | $spl2)),
  inference(resolution,[],[f11026,f896])).
fof(f896,plain,(
  ( ! [X50,X49] : (sP1(X49,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))),X50)) ) | $spl1),
  inference(resolution,[],[f863,f600])).
fof(f11008,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK12,sK20('const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)),sK13,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK14)))))) ) | $spl90),
  inference(cnf_transformation,[],[f11008_D])).
fof(f11008_D,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK12,sK20('const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)),sK13,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK14)))))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f10987,plain,(
  ( ! [X194,X193] : (p(s(bool,X193)) | p(s(bool,X194)) | s(bool,X194) != s(bool,'const.bool.IN_2'(s(sK12,sK20('const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)),sK13,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK14)))))) ) | $spl1),
  inference(resolution,[],[f8604,f839])).
fof(f8604,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK12,sK20('const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)),sK13,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK14)))))) | $spl1),
  inference(forward_demodulation,[],[f8479,f636])).
fof(f636,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1))))) = s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.REVERSE_1'(s('type.list.list'(X0),X1)))))))) )),
  inference(cnf_transformation,[],[f297])).
fof(f297,plain,(
  ! [X0,X1,X2] : s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1))))) = s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.REVERSE_1'(s('type.list.list'(X0),X1)))))))),
  inference(flattening,[],[f296])).
fof(f296,plain,(
  ! [X0] : ! [X1,X2] : s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1))))) = s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.REVERSE_1'(s('type.list.list'(X0),X1)))))))),
  inference(rectify,[],[f88])).
fof(f88,axiom,(
  ! [X6] : ! [X8,X12] : s(bool,'const.bool.IN_2'(s(X6,X12),s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),X8))))) = s(bool,'const.bool.IN_2'(s(X6,X12),s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),'const.list.REVERSE_1'(s('type.list.list'(X6),X8)))))))),
  file('thm.rich_list.EVERY_REVERSE','thm.list.MEM_REVERSE')).
fof(f8479,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK12,sK20('const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)),sK13,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))))) | $spl1),
  inference(resolution,[],[f863,f665])).
fof(f665,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK20(X2,X1,X0)),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2))))))) )),
  inference(cnf_transformation,[],[f486])).
fof(f10898,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.bool.IN_2'(s(sK12,sK20('const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)),sK13,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK14))))) | s(bool,X23) != s(bool,X24)) ) | $spl1),
  inference(resolution,[],[f8604,f586])).
fof(f10895,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.bool.IN_2'(s(sK12,sK20('const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)),sK13,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK14)))))) ) | $spl1),
  inference(resolution,[],[f8604,f585])).
fof(f10894,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.bool.IN_2'(s(sK12,sK20('const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)),sK13,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK14)))))) ) | $spl1),
  inference(resolution,[],[f8604,f584])).
fof(f10893,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.bool.IN_2'(s(sK12,sK20('const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)),sK13,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK14))))) | s(bool,X8) != s(bool,X9)) ) | $spl1),
  inference(resolution,[],[f8604,f582])).
fof(f10891,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK12,sK20('const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)),sK13,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK14)))))) ) | $spl1),
  inference(resolution,[],[f8604,f580])).
fof(f10773,plain,(
  ( ! [X285,X287,X286] : (s(bool,X285) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) | p(s(bool,X286)) | ~p(s(bool,X287)) | s(bool,X285) != s(bool,X287)) ) | $spl2),
  inference(resolution,[],[f8354,f839])).
fof(f10685,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) ) | $spl2),
  inference(resolution,[],[f8354,f586])).
fof(f10683,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,X20) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X20) != s(bool,X22)) ) | $spl2),
  inference(resolution,[],[f8354,f585])).
fof(f10681,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl2),
  inference(resolution,[],[f8354,f584])).
fof(f10680,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) ) | $spl2),
  inference(resolution,[],[f8354,f582])).
fof(f10678,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl2),
  inference(resolution,[],[f8354,f580])).
fof(f10430,plain,(
  ( ! [X300,X301,X299] : (p(s(bool,X299)) | p(s(bool,X300)) | ~p(s(bool,X301)) | s(bool,X301) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK30(sK13,X299,sK12))))) ) | $spl1),
  inference(resolution,[],[f8797,f839])).
fof(f10342,plain,(
  ( ! [X45,X43,X44,X42] : (p(s(bool,X42)) | ~p(s(bool,X43)) | ~p(s(bool,X44)) | s(bool,X44) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK30(sK13,X42,sK12)))) | s(bool,X43) != s(bool,X45)) ) | $spl1),
  inference(resolution,[],[f8797,f586])).
fof(f10340,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,X34)) | p(s(bool,X35)) | ~p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK30(sK13,X34,sK12))))) ) | $spl1),
  inference(resolution,[],[f8797,f585])).
fof(f10338,plain,(
  ( ! [X28,X26,X29,X27] : (p(s(bool,X26)) | p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,X29) | s(bool,X28) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK30(sK13,X26,sK12))))) ) | $spl1),
  inference(resolution,[],[f8797,f584])).
fof(f10337,plain,(
  ( ! [X24,X23,X25,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK30(sK13,X22,sK12)))) | s(bool,X23) != s(bool,X25)) ) | $spl1),
  inference(resolution,[],[f8797,f582])).
fof(f10335,plain,(
  ( ! [X14,X17,X15,X16] : (p(s(bool,X14)) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK30(sK13,X14,sK12))))) ) | $spl1),
  inference(resolution,[],[f8797,f580])).
fof(f10307,plain,(
  ( ! [X300,X301,X299] : (p(s(bool,X299)) | p(s(bool,X300)) | ~p(s(bool,X301)) | s(bool,X301) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK29(sK13,X299,sK12))))) ) | $spl1),
  inference(resolution,[],[f8796,f839])).
fof(f10219,plain,(
  ( ! [X45,X43,X44,X42] : (p(s(bool,X42)) | ~p(s(bool,X43)) | ~p(s(bool,X44)) | s(bool,X44) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK29(sK13,X42,sK12)))) | s(bool,X43) != s(bool,X45)) ) | $spl1),
  inference(resolution,[],[f8796,f586])).
fof(f10217,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,X34)) | p(s(bool,X35)) | ~p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK29(sK13,X34,sK12))))) ) | $spl1),
  inference(resolution,[],[f8796,f585])).
fof(f10215,plain,(
  ( ! [X28,X26,X29,X27] : (p(s(bool,X26)) | p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,X29) | s(bool,X28) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK29(sK13,X26,sK12))))) ) | $spl1),
  inference(resolution,[],[f8796,f584])).
fof(f10214,plain,(
  ( ! [X24,X23,X25,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK29(sK13,X22,sK12)))) | s(bool,X23) != s(bool,X25)) ) | $spl1),
  inference(resolution,[],[f8796,f582])).
fof(f10212,plain,(
  ( ! [X14,X17,X15,X16] : (p(s(bool,X14)) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK29(sK13,X14,sK12))))) ) | $spl1),
  inference(resolution,[],[f8796,f580])).
fof(f3524,plain,(
  ( ! [X11] : (p(s(bool,X11)) | ~p(s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK30(sK13,X11,sK12)))))) ) | $spl30),
  inference(cnf_transformation,[],[f3524_D])).
fof(f3524_D,plain,(
  ( ! [X11] : (p(s(bool,X11)) | ~p(s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK30(sK13,X11,sK12)))))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f10209,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | ~p(s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK30(sK13,X11,sK12)))))) ) | $spl1),
  inference(resolution,[],[f8796,f803])).
fof(f9671,plain,(
  s('type.list.list'(sK12),sK14) != s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))) | $spl79),
  inference(cnf_transformation,[],[f9671_D])).
fof(f9671_D,plain,(
  s('type.list.list'(sK12),sK14) != s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))) <=> ~$spl79),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl79])])).
fof(f9678,plain,(
  ( ! [X3] : (s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),X3),s('type.list.list'(sK12),sK14))) | p(s(bool,'const.bool.IN_2'(s(sK12,sK17(sK13,X3,'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)),'const.list.FILTER_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)),sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK14))))))) ) | $spl86),
  inference(cnf_transformation,[],[f9678_D])).
fof(f9678_D,plain,(
  ( ! [X3] : (s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),X3),s('type.list.list'(sK12),sK14))) | p(s(bool,'const.bool.IN_2'(s(sK12,sK17(sK13,X3,'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)),'const.list.FILTER_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)),sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK14))))))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f9677,plain,(
  ( ! [X3] : (s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),X3),s('type.list.list'(sK12),sK14))) | s('type.list.list'(sK12),sK14) != s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))) | p(s(bool,'const.bool.IN_2'(s(sK12,sK17(sK13,X3,'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)),'const.list.FILTER_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)),sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK14))))))) ) | $spl2),
  inference(forward_demodulation,[],[f9546,f8326])).
fof(f9546,plain,(
  ( ! [X3] : (s('type.list.list'(sK12),sK14) != s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))) | s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),X3),s('type.list.list'(sK12),'const.list.FILTER_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))))) | p(s(bool,'const.bool.IN_2'(s(sK12,sK17(sK13,X3,'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)),'const.list.FILTER_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)),sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK14))))))) ) | $spl2),
  inference(superposition,[],[f862,f8326])).
fof(f862,plain,(
  ( ! [X0,X1] : (s('type.list.list'(sK12),X1) != s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))) | s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),X0),s('type.list.list'(sK12),X1))) | p(s(bool,'const.bool.IN_2'(s(sK12,sK17(sK13,X0,'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)),X1,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK14))))))) )),
  inference(forward_demodulation,[],[f846,f636])).
fof(f846,plain,(
  ( ! [X0,X1] : (s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),X0),s('type.list.list'(sK12),X1))) | p(s(bool,'const.bool.IN_2'(s(sK12,sK17(sK13,X0,'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)),X1,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))))) | s('type.list.list'(sK12),X1) != s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))) )),
  inference(superposition,[],[f578,f645])).
fof(f645,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X3),s('type.list.list'(X0),X1))) = s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X4),s('type.list.list'(X0),X2))) | p(s(bool,'const.bool.IN_2'(s(X0,sK17(X4,X3,X2,X1,X0)),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | s('type.list.list'(X0),X1) != s('type.list.list'(X0),X2)) )),
  inference(cnf_transformation,[],[f474])).
fof(f474,plain,(
  ! [X0,X1,X2,X3,X4] : (s('type.list.list'(X0),X1) != s('type.list.list'(X0),X2) | (p(s(bool,'const.bool.IN_2'(s(X0,sK17(X4,X3,X2,X1,X0)),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & s(bool,happ(s(fun(X0,bool),X3),s(X0,sK17(X4,X3,X2,X1,X0)))) != s(bool,happ(s(fun(X0,bool),X4),s(X0,sK17(X4,X3,X2,X1,X0))))) | s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X3),s('type.list.list'(X0),X1))) = s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X4),s('type.list.list'(X0),X2))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK17])],[f400])).
fof(f400,plain,(
  ! [X0,X1,X2,X3,X4] : (s('type.list.list'(X0),X1) != s('type.list.list'(X0),X2) | ? [X5] : (p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & s(bool,happ(s(fun(X0,bool),X3),s(X0,X5))) != s(bool,happ(s(fun(X0,bool),X4),s(X0,X5)))) | s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X3),s('type.list.list'(X0),X1))) = s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X4),s('type.list.list'(X0),X2))))),
  inference(flattening,[],[f399])).
fof(f399,plain,(
  ! [X0,X1,X2,X3,X4] : ((s('type.list.list'(X0),X1) != s('type.list.list'(X0),X2) | ? [X5] : (p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & s(bool,happ(s(fun(X0,bool),X3),s(X0,X5))) != s(bool,happ(s(fun(X0,bool),X4),s(X0,X5))))) | s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X3),s('type.list.list'(X0),X1))) = s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X4),s('type.list.list'(X0),X2))))),
  inference(ennf_transformation,[],[f313])).
fof(f313,plain,(
  ! [X0,X1,X2,X3,X4] : ((s('type.list.list'(X0),X1) = s('type.list.list'(X0),X2) & ! [X5] : (p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) => s(bool,happ(s(fun(X0,bool),X3),s(X0,X5))) = s(bool,happ(s(fun(X0,bool),X4),s(X0,X5))))) => s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X3),s('type.list.list'(X0),X1))) = s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X4),s('type.list.list'(X0),X2))))),
  inference(flattening,[],[f312])).
fof(f312,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : ((s('type.list.list'(X0),X1) = s('type.list.list'(X0),X2) & ! [X5] : (p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) => s(bool,happ(s(fun(X0,bool),X3),s(X0,X5))) = s(bool,happ(s(fun(X0,bool),X4),s(X0,X5))))) => s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X3),s('type.list.list'(X0),X1))) = s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X4),s('type.list.list'(X0),X2))))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X6] : ! [X13,X14,X7,X15] : ((s('type.list.list'(X6),X13) = s('type.list.list'(X6),X14) & ! [X12] : (p(s(bool,'const.bool.IN_2'(s(X6,X12),s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),X14)))))) => s(bool,happ(s(fun(X6,bool),X7),s(X6,X12))) = s(bool,happ(s(fun(X6,bool),X15),s(X6,X12))))) => s(bool,'const.list.EVERY_2'(s(fun(X6,bool),X7),s('type.list.list'(X6),X13))) = s(bool,'const.list.EVERY_2'(s(fun(X6,bool),X15),s('type.list.list'(X6),X14))))),
  file('thm.rich_list.EVERY_REVERSE','thm.list.EVERY_CONG')).
fof(f578,plain,(
  s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))),
  inference(cnf_transformation,[],[f451])).
fof(f451,plain,(
  s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13,sK14])],[f391])).
fof(f391,plain,(
  ? [X0,X1,X2] : s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))) != s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.REVERSE_1'(s('type.list.list'(X0),X2)))))),
  inference(ennf_transformation,[],[f264])).
fof(f264,plain,(
  ~! [X0,X1,X2] : s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))) = s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.REVERSE_1'(s('type.list.list'(X0),X2)))))),
  inference(flattening,[],[f263])).
fof(f263,plain,(
  ~! [X0] : ! [X1,X2] : s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))) = s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.REVERSE_1'(s('type.list.list'(X0),X2)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : s(bool,'const.list.EVERY_2'(s(fun(X6,bool),X7),s('type.list.list'(X6),X8))) = s(bool,'const.list.EVERY_2'(s(fun(X6,bool),X7),s('type.list.list'(X6),'const.list.REVERSE_1'(s('type.list.list'(X6),X8)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : s(bool,'const.list.EVERY_2'(s(fun(X6,bool),X7),s('type.list.list'(X6),X8))) = s(bool,'const.list.EVERY_2'(s(fun(X6,bool),X7),s('type.list.list'(X6),'const.list.REVERSE_1'(s('type.list.list'(X6),X8)))))),
  file('thm.rich_list.EVERY_REVERSE','thm.rich_list.EVERY_REVERSE')).
fof(f9676,plain,(
  ( ! [X2] : (s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),X2),s('type.list.list'(sK12),sK14))) | s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK17(X2,sK13,'const.list.FILTER_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)),sK12)))) != s(bool,happ(s(fun(sK12,bool),X2),s(sK12,sK17(X2,sK13,'const.list.FILTER_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)),sK12))))) ) | $spl84),
  inference(cnf_transformation,[],[f9676_D])).
fof(f9676_D,plain,(
  ( ! [X2] : (s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),X2),s('type.list.list'(sK12),sK14))) | s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK17(X2,sK13,'const.list.FILTER_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)),sK12)))) != s(bool,happ(s(fun(sK12,bool),X2),s(sK12,sK17(X2,sK13,'const.list.FILTER_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)),sK12))))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f9675,plain,(
  ( ! [X2] : (s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),X2),s('type.list.list'(sK12),sK14))) | s('type.list.list'(sK12),sK14) != s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))) | s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK17(X2,sK13,'const.list.FILTER_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)),sK12)))) != s(bool,happ(s(fun(sK12,bool),X2),s(sK12,sK17(X2,sK13,'const.list.FILTER_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)),sK12))))) ) | $spl2),
  inference(forward_demodulation,[],[f9545,f8326])).
fof(f9545,plain,(
  ( ! [X2] : (s('type.list.list'(sK12),sK14) != s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))) | s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK17(X2,sK13,'const.list.FILTER_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)),sK12)))) != s(bool,happ(s(fun(sK12,bool),X2),s(sK12,sK17(X2,sK13,'const.list.FILTER_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)),sK12)))) | s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),X2),s('type.list.list'(sK12),'const.list.FILTER_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)))))) ) | $spl2),
  inference(superposition,[],[f849,f8326])).
fof(f849,plain,(
  ( ! [X6,X7] : (s('type.list.list'(sK12),X7) != s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))) | s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK17(X6,sK13,X7,'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)),sK12)))) != s(bool,happ(s(fun(sK12,bool),X6),s(sK12,sK17(X6,sK13,X7,'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)),sK12)))) | s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),X6),s('type.list.list'(sK12),X7)))) )),
  inference(superposition,[],[f578,f646])).
fof(f646,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X3),s('type.list.list'(X0),X1))) = s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X4),s('type.list.list'(X0),X2))) | s(bool,happ(s(fun(X0,bool),X3),s(X0,sK17(X4,X3,X2,X1,X0)))) != s(bool,happ(s(fun(X0,bool),X4),s(X0,sK17(X4,X3,X2,X1,X0)))) | s('type.list.list'(X0),X1) != s('type.list.list'(X0),X2)) )),
  inference(cnf_transformation,[],[f474])).
fof(f9674,plain,(
  ( ! [X1] : (s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),X1),s('type.list.list'(sK12),sK14))) | s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK17(sK13,X1,'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)),'const.list.FILTER_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)),sK12)))) != s(bool,happ(s(fun(sK12,bool),X1),s(sK12,sK17(sK13,X1,'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)),'const.list.FILTER_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)),sK12))))) ) | $spl82),
  inference(cnf_transformation,[],[f9674_D])).
fof(f9674_D,plain,(
  ( ! [X1] : (s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),X1),s('type.list.list'(sK12),sK14))) | s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK17(sK13,X1,'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)),'const.list.FILTER_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)),sK12)))) != s(bool,happ(s(fun(sK12,bool),X1),s(sK12,sK17(sK13,X1,'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)),'const.list.FILTER_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)),sK12))))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f9673,plain,(
  ( ! [X1] : (s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),X1),s('type.list.list'(sK12),sK14))) | s('type.list.list'(sK12),sK14) != s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))) | s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK17(sK13,X1,'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)),'const.list.FILTER_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)),sK12)))) != s(bool,happ(s(fun(sK12,bool),X1),s(sK12,sK17(sK13,X1,'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)),'const.list.FILTER_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)),sK12))))) ) | $spl2),
  inference(forward_demodulation,[],[f9544,f8326])).
fof(f9544,plain,(
  ( ! [X1] : (s('type.list.list'(sK12),sK14) != s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))) | s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK17(sK13,X1,'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)),'const.list.FILTER_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)),sK12)))) != s(bool,happ(s(fun(sK12,bool),X1),s(sK12,sK17(sK13,X1,'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)),'const.list.FILTER_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)),sK12)))) | s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),X1),s('type.list.list'(sK12),'const.list.FILTER_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)))))) ) | $spl2),
  inference(superposition,[],[f848,f8326])).
fof(f848,plain,(
  ( ! [X4,X5] : (s('type.list.list'(sK12),X5) != s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))) | s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK17(sK13,X4,'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)),X5,sK12)))) != s(bool,happ(s(fun(sK12,bool),X4),s(sK12,sK17(sK13,X4,'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)),X5,sK12)))) | s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),X4),s('type.list.list'(sK12),X5)))) )),
  inference(superposition,[],[f578,f646])).
fof(f9672,plain,(
  ( ! [X0] : (s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),X0),s('type.list.list'(sK12),sK14))) | p(s(bool,'const.bool.IN_2'(s(sK12,sK17(X0,sK13,'const.list.FILTER_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)),sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK14))))))) ) | $spl80),
  inference(cnf_transformation,[],[f9672_D])).
fof(f9672_D,plain,(
  ( ! [X0] : (s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),X0),s('type.list.list'(sK12),sK14))) | p(s(bool,'const.bool.IN_2'(s(sK12,sK17(X0,sK13,'const.list.FILTER_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)),sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK14))))))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f9669,plain,(
  ( ! [X0] : (s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),X0),s('type.list.list'(sK12),sK14))) | p(s(bool,'const.bool.IN_2'(s(sK12,sK17(X0,sK13,'const.list.FILTER_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)),sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK14)))))) | s('type.list.list'(sK12),sK14) != s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))) ) | $spl2),
  inference(forward_demodulation,[],[f9668,f8326])).
fof(f9668,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK17(X0,sK13,'const.list.FILTER_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)),sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK14)))))) | s('type.list.list'(sK12),sK14) != s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))) | s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),X0),s('type.list.list'(sK12),'const.list.FILTER_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)))))) ) | $spl2),
  inference(forward_demodulation,[],[f9543,f8326])).
fof(f9543,plain,(
  ( ! [X0] : (s('type.list.list'(sK12),sK14) != s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))) | p(s(bool,'const.bool.IN_2'(s(sK12,sK17(X0,sK13,'const.list.FILTER_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)),sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.FILTER_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)))))))) | s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),X0),s('type.list.list'(sK12),'const.list.FILTER_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)))))) ) | $spl2),
  inference(superposition,[],[f847,f8326])).
fof(f847,plain,(
  ( ! [X2,X3] : (s('type.list.list'(sK12),X3) != s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))) | p(s(bool,'const.bool.IN_2'(s(sK12,sK17(X2,sK13,X3,'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)),sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),X3)))))) | s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),X2),s('type.list.list'(sK12),X3)))) )),
  inference(superposition,[],[f578,f645])).
fof(f9142,plain,(
  ( ! [X196,X197] : (p(s(bool,X196)) | ~p(s(bool,X197)) | s(bool,X197) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12))))) ) | $spl1),
  inference(resolution,[],[f8799,f839])).
fof(f9054,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X28) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12)))) | s(bool,X27) != s(bool,X29)) ) | $spl1),
  inference(resolution,[],[f8799,f586])).
fof(f9052,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12))))) ) | $spl1),
  inference(resolution,[],[f8799,f585])).
fof(f9050,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12))))) ) | $spl1),
  inference(resolution,[],[f8799,f584])).
fof(f9049,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X13) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12)))) | s(bool,X12) != s(bool,X14)) ) | $spl1),
  inference(resolution,[],[f8799,f582])).
fof(f9047,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12))))) ) | $spl1),
  inference(resolution,[],[f8799,f580])).
fof(f9019,plain,(
  ( ! [X196,X197] : (p(s(bool,X196)) | ~p(s(bool,X197)) | s(bool,X197) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12))))) ) | $spl1),
  inference(resolution,[],[f8798,f839])).
fof(f8931,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X28) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) | s(bool,X27) != s(bool,X29)) ) | $spl1),
  inference(resolution,[],[f8798,f586])).
fof(f8929,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12))))) ) | $spl1),
  inference(resolution,[],[f8798,f585])).
fof(f8927,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12))))) ) | $spl1),
  inference(resolution,[],[f8798,f584])).
fof(f8926,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X13) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) | s(bool,X12) != s(bool,X14)) ) | $spl1),
  inference(resolution,[],[f8798,f582])).
fof(f8924,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12))))) ) | $spl1),
  inference(resolution,[],[f8798,f580])).
fof(f8895,plain,(
  ( ! [X196,X197] : (p(s(bool,X196)) | ~p(s(bool,X197)) | s(bool,X197) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20('const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)),sK13,sK12))))) ) | $spl1),
  inference(resolution,[],[f8480,f839])).
fof(f8807,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X28) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20('const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)),sK13,sK12)))) | s(bool,X27) != s(bool,X29)) ) | $spl1),
  inference(resolution,[],[f8480,f586])).
fof(f8805,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20('const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)),sK13,sK12))))) ) | $spl1),
  inference(resolution,[],[f8480,f585])).
fof(f8803,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20('const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)),sK13,sK12))))) ) | $spl1),
  inference(resolution,[],[f8480,f584])).
fof(f8802,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X13) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20('const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)),sK13,sK12)))) | s(bool,X12) != s(bool,X14)) ) | $spl1),
  inference(resolution,[],[f8480,f582])).
fof(f8800,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20('const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)),sK13,sK12))))) ) | $spl1),
  inference(resolution,[],[f8480,f580])).
fof(f8578,plain,(
  ( ! [X202,X201] : (p(s(bool,X201)) | ~p(s(bool,X202)) | s(bool,X202) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) ) | $spl1),
  inference(resolution,[],[f863,f839])).
fof(f8490,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X33) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))) | s(bool,X32) != s(bool,X34)) ) | $spl1),
  inference(resolution,[],[f863,f586])).
fof(f8488,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) ) | $spl1),
  inference(resolution,[],[f863,f585])).
fof(f8486,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X20) != s(bool,X22) | s(bool,X21) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) ) | $spl1),
  inference(resolution,[],[f863,f584])).
fof(f8485,plain,(
  ( ! [X19,X17,X18] : (~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X18) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))) | s(bool,X17) != s(bool,X19)) ) | $spl1),
  inference(resolution,[],[f863,f582])).
fof(f8483,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) ) | $spl1),
  inference(resolution,[],[f863,f580])).
fof(f863,plain,(
  ~p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) | $spl1),
  inference(cnf_transformation,[],[f863_D])).
fof(f863_D,plain,(
  ~p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f8311,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | ~sP2(X6,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))),X7) | s(bool,X7) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)))) ) | $spl74),
  inference(cnf_transformation,[],[f8311_D])).
fof(f8311_D,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | ~sP2(X6,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))),X7) | s(bool,X7) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f8310,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | s(bool,X7) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) | ~sP2(X6,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))),X7)) )),
  inference(subsumption_resolution,[],[f2184,f605])).
fof(f2184,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | s(bool,X7) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) | ~sP2(X6,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))),X7) | ~sP0('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))),X7,X6)) )),
  inference(subsumption_resolution,[],[f2177,f599])).
fof(f2177,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | s(bool,X7) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) | ~sP2(X6,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))),X7) | ~sP1(X7,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))),X6) | ~sP0('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))),X7,X6)) )),
  inference(resolution,[],[f852,f591])).
fof(f852,plain,(
  ( ! [X12,X13] : (~sP3(X13,X12,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)))) )),
  inference(superposition,[],[f578,f608])).
fof(f864,plain,(
  p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) | $spl0),
  inference(cnf_transformation,[],[f864_D])).
fof(f864_D,plain,(
  p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f8309,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | ~sP2(X8,X9,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))) | s(bool,X9) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)))) ) | $spl72),
  inference(cnf_transformation,[],[f8309_D])).
fof(f8309_D,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | ~sP2(X8,X9,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))) | s(bool,X9) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f8308,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | s(bool,X9) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) | p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) | ~sP2(X8,X9,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) )),
  inference(subsumption_resolution,[],[f2231,f603])).
fof(f2231,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | s(bool,X9) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) | p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) | ~sP2(X8,X9,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))) | ~sP0(X9,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))),X8)) )),
  inference(subsumption_resolution,[],[f2225,f601])).
fof(f2225,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | s(bool,X9) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) | p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) | ~sP2(X8,X9,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))) | ~sP1('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))),X9,X8) | ~sP0(X9,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))),X8)) )),
  inference(resolution,[],[f853,f592])).
fof(f853,plain,(
  ( ! [X14,X15] : (~sP3(X15,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))),X14) | ~p(s(bool,X15)) | s(bool,X14) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)))) )),
  inference(superposition,[],[f578,f608])).
fof(f8305,plain,(
  $false | ($spl0 | $spl17)),
  inference(subsumption_resolution,[],[f8182,f1605])).
fof(f1605,plain,(
  ~p(s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20(sK14,sK13,sK12))))) | $spl17),
  inference(resolution,[],[f1092,f666])).
fof(f8182,plain,(
  p(s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20(sK14,sK13,sK12))))) | ($spl0 | $spl17)),
  inference(resolution,[],[f1511,f1604])).
fof(f1604,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK12,sK20(sK14,sK13,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK14)))))) | $spl17),
  inference(resolution,[],[f1092,f665])).
fof(f1511,plain,(
  ( ! [X10] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X10),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK14)))))) | p(s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,X10))))) ) | $spl0),
  inference(forward_demodulation,[],[f1384,f636])).
fof(f1384,plain,(
  ( ! [X10] : (p(s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,X10)))) | ~p(s(bool,'const.bool.IN_2'(s(sK12,X10),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))))))) ) | $spl0),
  inference(resolution,[],[f864,f664])).
fof(f664,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))))) )),
  inference(cnf_transformation,[],[f486])).
fof(f5545,plain,(
  s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12)))) | $spl50),
  inference(cnf_transformation,[],[f5545_D])).
fof(f5545_D,plain,(
  s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12)))) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f8178,plain,(
  ( ! [X12,X13] : (s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12)))) | p(s(bool,X12)) | p(s(bool,X13))) ) | $spl17),
  inference(resolution,[],[f7408,f1627])).
fof(f1627,plain,(
  ( ! [X57,X58] : (~sP0(X57,X58,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) | p(s(bool,X58)) | p(s(bool,X57))) ) | $spl17),
  inference(resolution,[],[f1092,f602])).
fof(f7408,plain,(
  ( ! [X61,X59,X60] : (sP0(X60,X61,X59) | s(bool,X59) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12))))) ) | $spl17),
  inference(resolution,[],[f2082,f603])).
fof(f2082,plain,(
  ( ! [X60] : (p(s(bool,X60)) | s(bool,X60) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12))))) ) | $spl17),
  inference(resolution,[],[f1809,f615])).
fof(f1809,plain,(
  ~p(s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12))))) | $spl17),
  inference(resolution,[],[f1605,f828])).
fof(f8170,plain,(
  ( ! [X15,X16] : (s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12)))) | ~p(s(bool,X15)) | ~p(s(bool,X16))) ) | $spl17),
  inference(resolution,[],[f7407,f1622])).
fof(f1622,plain,(
  ( ! [X47,X48] : (~sP1(X47,X48,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) | ~p(s(bool,X48)) | ~p(s(bool,X47))) ) | $spl17),
  inference(resolution,[],[f1092,f598])).
fof(f7407,plain,(
  ( ! [X57,X58,X56] : (sP1(X57,X58,X56) | s(bool,X56) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12))))) ) | $spl17),
  inference(resolution,[],[f2082,f599])).
fof(f8160,plain,(
  ( ! [X15,X16] : (s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12)))) | ~p(s(bool,X15)) | ~p(s(bool,X16))) ) | $spl17),
  inference(resolution,[],[f7404,f1619])).
fof(f1619,plain,(
  ( ! [X41,X42] : (~sP2(X41,X42,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) | ~p(s(bool,X42)) | ~p(s(bool,X41))) ) | $spl17),
  inference(resolution,[],[f1092,f594])).
fof(f7404,plain,(
  ( ! [X47,X48,X49] : (sP2(X48,X49,X47) | s(bool,X47) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12))))) ) | $spl17),
  inference(resolution,[],[f2082,f595])).
fof(f8155,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12)))) | ~p(s(bool,X1)) | s(bool,X0) = s(bool,X1) | ~p(s(bool,X2))) ) | ($spl4 | $spl17)),
  inference(resolution,[],[f7404,f5182])).
fof(f5182,plain,(
  ( ! [X2,X0,X1] : (~sP2(X1,X0,X2) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X2) | ~p(s(bool,X1))) ) | ($spl4 | $spl17)),
  inference(subsumption_resolution,[],[f5170,f599])).
fof(f5170,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP2(X1,X0,X2) | ~sP1(X2,X0,X1) | s(bool,X0) = s(bool,X2) | ~p(s(bool,X1))) ) | ($spl4 | $spl17)),
  inference(resolution,[],[f5167,f608])).
fof(f5167,plain,(
  ( ! [X4,X5,X3] : (sP3(X4,X5,X3) | ~p(s(bool,X3)) | ~sP2(X4,X3,X5) | ~sP1(X5,X3,X4)) ) | ($spl4 | $spl17)),
  inference(subsumption_resolution,[],[f5164,f866])).
fof(f5164,plain,(
  ( ! [X4,X5,X3] : (s(bool,X3) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) | sP3(X4,X5,X3) | ~p(s(bool,X3)) | ~sP2(X4,X3,X5) | ~sP1(X5,X3,X4)) ) | $spl17),
  inference(resolution,[],[f4727,f591])).
fof(f4727,plain,(
  ( ! [X66,X67,X65] : (sP0(X65,X66,X67) | s(bool,X65) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)))) ) | $spl17),
  inference(resolution,[],[f1632,f605])).
fof(f1632,plain,(
  ( ! [X65] : (p(s(bool,X65)) | s(bool,X65) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)))) ) | $spl17),
  inference(resolution,[],[f1092,f615])).
fof(f8140,plain,(
  ( ! [X30,X28] : (p(s(bool,X28)) | s(bool,X28) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),'const.list.FILTER_2'(s(fun(sK12,bool),X30),s('type.list.list'(sK12),sK14)))))))) ) | $spl68),
  inference(cnf_transformation,[],[f8140_D])).
fof(f8140_D,plain,(
  ( ! [X30,X28] : (p(s(bool,X28)) | s(bool,X28) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),'const.list.FILTER_2'(s(fun(sK12,bool),X30),s('type.list.list'(sK12),sK14)))))))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f8107,plain,(
  ( ! [X315,X316,X314] : (p(s(bool,X314)) | p(s(bool,X315)) | s(bool,X315) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),'const.list.FILTER_2'(s(fun(sK12,bool),X316),s('type.list.list'(sK12),sK14)))))))) ) | $spl0),
  inference(resolution,[],[f1510,f839])).
fof(f1510,plain,(
  ( ! [X9] : (p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),'const.list.FILTER_2'(s(fun(sK12,bool),X9),s('type.list.list'(sK12),sK14))))))))) ) | $spl0),
  inference(forward_demodulation,[],[f1382,f635])).
fof(f1382,plain,(
  ( ! [X9] : (p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.FILTER_2'(s(fun(sK12,bool),X9),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))))))) ) | $spl0),
  inference(resolution,[],[f864,f661])).
fof(f8018,plain,(
  ( ! [X57,X54,X56,X55] : (p(s(bool,X54)) | ~p(s(bool,X55)) | s(bool,X54) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),'const.list.FILTER_2'(s(fun(sK12,bool),X56),s('type.list.list'(sK12),sK14))))))) | s(bool,X55) != s(bool,X57)) ) | $spl0),
  inference(resolution,[],[f1510,f586])).
fof(f8015,plain,(
  ( ! [X45,X43,X46,X44] : (p(s(bool,X43)) | ~p(s(bool,X44)) | p(s(bool,X45)) | s(bool,X45) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),'const.list.FILTER_2'(s(fun(sK12,bool),X46),s('type.list.list'(sK12),sK14)))))))) ) | $spl0),
  inference(resolution,[],[f1510,f585])).
fof(f8014,plain,(
  ( ! [X39,X41,X42,X40] : (p(s(bool,X39)) | p(s(bool,X40)) | s(bool,X39) != s(bool,X41) | s(bool,X40) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),'const.list.FILTER_2'(s(fun(sK12,bool),X42),s('type.list.list'(sK12),sK14)))))))) ) | $spl0),
  inference(resolution,[],[f1510,f584])).
fof(f8013,plain,(
  ( ! [X37,X35,X38,X36] : (p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),'const.list.FILTER_2'(s(fun(sK12,bool),X37),s('type.list.list'(sK12),sK14))))))) | s(bool,X36) != s(bool,X38)) ) | $spl0),
  inference(resolution,[],[f1510,f582])).
fof(f8011,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,X27)) | p(s(bool,X28)) | s(bool,X27) != s(bool,X29) | s(bool,X28) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),'const.list.FILTER_2'(s(fun(sK12,bool),X30),s('type.list.list'(sK12),sK14)))))))) ) | $spl0),
  inference(resolution,[],[f1510,f580])).
fof(f7988,plain,(
  ( ! [X19] : (s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20(sK14,sK13,sK12)))) | p(s(bool,X19))) ) | $spl17),
  inference(resolution,[],[f7123,f1883])).
fof(f1883,plain,(
  ( ! [X154,X155] : (~sP10(X155,happ(s(fun(sK12,bool),sK13),s(sK12,sK20(sK14,sK13,sK12))),X154) | p(s(bool,X154))) ) | $spl17),
  inference(resolution,[],[f1605,f745])).
fof(f7123,plain,(
  ( ! [X241,X239,X240] : (sP10(X240,X239,X241) | s(bool,X239) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12))))) ) | $spl17),
  inference(resolution,[],[f1958,f751])).
fof(f1958,plain,(
  ( ! [X60] : (p(s(bool,X60)) | s(bool,X60) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12))))) ) | $spl17),
  inference(resolution,[],[f1808,f615])).
fof(f1808,plain,(
  ~p(s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12))))) | $spl17),
  inference(resolution,[],[f1605,f827])).
fof(f7987,plain,(
  ( ! [X18] : (s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20(sK14,sK13,sK12)))) | p(s(bool,X18))) ) | $spl17),
  inference(resolution,[],[f7123,f1885])).
fof(f1885,plain,(
  ( ! [X158,X159] : (~sP10(X158,happ(s(fun(sK12,bool),sK13),s(sK12,sK20(sK14,sK13,sK12))),X159) | p(s(bool,X158))) ) | $spl17),
  inference(resolution,[],[f1605,f746])).
fof(f7986,plain,(
  ( ! [X17] : (s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12)))) | p(s(bool,X17))) ) | $spl17),
  inference(resolution,[],[f7123,f2131])).
fof(f2131,plain,(
  ( ! [X154,X155] : (~sP10(X155,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12))),X154) | p(s(bool,X154))) ) | $spl17),
  inference(resolution,[],[f1809,f745])).
fof(f7985,plain,(
  ( ! [X16] : (s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12)))) | p(s(bool,X16))) ) | $spl17),
  inference(resolution,[],[f7123,f2133])).
fof(f2133,plain,(
  ( ! [X158,X159] : (~sP10(X158,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12))),X159) | p(s(bool,X158))) ) | $spl17),
  inference(resolution,[],[f1809,f746])).
fof(f5543,plain,(
  s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) | $spl48),
  inference(cnf_transformation,[],[f5543_D])).
fof(f5543_D,plain,(
  s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f7982,plain,(
  ( ! [X13] : (s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) | p(s(bool,X13))) ) | $spl17),
  inference(resolution,[],[f7123,f1681])).
fof(f1681,plain,(
  ( ! [X159,X160] : (~sP10(X160,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)),X159) | p(s(bool,X159))) ) | $spl17),
  inference(resolution,[],[f1092,f745])).
fof(f7981,plain,(
  ( ! [X12] : (s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) | p(s(bool,X12))) ) | $spl17),
  inference(resolution,[],[f7123,f1683])).
fof(f1683,plain,(
  ( ! [X163,X164] : (~sP10(X163,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)),X164) | p(s(bool,X163))) ) | $spl17),
  inference(resolution,[],[f1092,f746])).
fof(f7970,plain,(
  ( ! [X15,X16] : (s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) | ~p(s(bool,X15)) | ~p(s(bool,X16))) ) | $spl17),
  inference(resolution,[],[f7109,f1671])).
fof(f1671,plain,(
  ( ! [X140,X139] : (~sP8(X139,X140,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) | ~p(s(bool,X140)) | ~p(s(bool,X139))) ) | $spl17),
  inference(resolution,[],[f1092,f729])).
fof(f7109,plain,(
  ( ! [X198,X199,X197] : (sP8(X198,X199,X197) | s(bool,X197) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12))))) ) | $spl17),
  inference(resolution,[],[f1958,f730])).
fof(f7962,plain,(
  ( ! [X17,X18] : (s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) | p(s(bool,X17)) | ~p(s(bool,X18))) ) | $spl17),
  inference(resolution,[],[f7093,f1657])).
fof(f1657,plain,(
  ( ! [X111,X112] : (~sP6(X111,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)),X112) | p(s(bool,X112)) | ~p(s(bool,X111))) ) | $spl17),
  inference(resolution,[],[f1092,f712])).
fof(f7093,plain,(
  ( ! [X151,X149,X150] : (sP6(X150,X149,X151) | s(bool,X149) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12))))) ) | $spl17),
  inference(resolution,[],[f1958,f714])).
fof(f7951,plain,(
  ( ! [X15,X16] : (s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) | p(s(bool,X15)) | ~p(s(bool,X16))) ) | $spl17),
  inference(resolution,[],[f7092,f1658])).
fof(f1658,plain,(
  ( ! [X114,X113] : (~sP6(X113,X114,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) | p(s(bool,X114)) | ~p(s(bool,X113))) ) | $spl17),
  inference(resolution,[],[f1092,f712])).
fof(f7092,plain,(
  ( ! [X146,X147,X148] : (sP6(X147,X148,X146) | s(bool,X146) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12))))) ) | $spl17),
  inference(resolution,[],[f1958,f713])).
fof(f7682,plain,(
  ( ! [X17,X18] : (s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) | ~p(s(bool,X17)) | ~p(s(bool,X18))) ) | $spl17),
  inference(resolution,[],[f7079,f1641])).
fof(f1641,plain,(
  ( ! [X80,X79] : (~sP4(X79,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)),X80) | ~p(s(bool,X80)) | ~p(s(bool,X79))) ) | $spl17),
  inference(resolution,[],[f1092,f695])).
fof(f7079,plain,(
  ( ! [X109,X107,X108] : (sP4(X108,X107,X109) | s(bool,X107) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12))))) ) | $spl17),
  inference(resolution,[],[f1958,f697])).
fof(f7673,plain,(
  ( ! [X17,X16] : (s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) | p(s(bool,X16)) | p(s(bool,X17))) ) | $spl17),
  inference(resolution,[],[f7063,f1626])).
fof(f1626,plain,(
  ( ! [X56,X55] : (~sP0('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)),X55,X56) | p(s(bool,X56)) | p(s(bool,X55))) ) | $spl17),
  inference(resolution,[],[f1092,f602])).
fof(f7063,plain,(
  ( ! [X66,X67,X65] : (sP0(X65,X66,X67) | s(bool,X65) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12))))) ) | $spl17),
  inference(resolution,[],[f1958,f605])).
fof(f7664,plain,(
  ( ! [X14,X15] : (s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) | p(s(bool,X14)) | p(s(bool,X15))) ) | $spl17),
  inference(resolution,[],[f7062,f1625])).
fof(f1625,plain,(
  ( ! [X54,X53] : (~sP0(X53,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)),X54) | p(s(bool,X54)) | p(s(bool,X53))) ) | $spl17),
  inference(resolution,[],[f1092,f602])).
fof(f7062,plain,(
  ( ! [X64,X62,X63] : (sP0(X63,X62,X64) | s(bool,X62) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12))))) ) | $spl17),
  inference(resolution,[],[f1958,f604])).
fof(f7655,plain,(
  ( ! [X12,X13] : (s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) | p(s(bool,X12)) | p(s(bool,X13))) ) | $spl17),
  inference(resolution,[],[f7061,f1627])).
fof(f7061,plain,(
  ( ! [X61,X59,X60] : (sP0(X60,X61,X59) | s(bool,X59) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12))))) ) | $spl17),
  inference(resolution,[],[f1958,f603])).
fof(f7647,plain,(
  ( ! [X15,X16] : (s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) | ~p(s(bool,X15)) | ~p(s(bool,X16))) ) | $spl17),
  inference(resolution,[],[f7060,f1622])).
fof(f7060,plain,(
  ( ! [X57,X58,X56] : (sP1(X57,X58,X56) | s(bool,X56) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12))))) ) | $spl17),
  inference(resolution,[],[f1958,f599])).
fof(f7637,plain,(
  ( ! [X15,X16] : (s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) | ~p(s(bool,X15)) | ~p(s(bool,X16))) ) | $spl17),
  inference(resolution,[],[f7057,f1619])).
fof(f7057,plain,(
  ( ! [X47,X48,X49] : (sP2(X48,X49,X47) | s(bool,X47) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12))))) ) | $spl17),
  inference(resolution,[],[f1958,f595])).
fof(f7632,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) | ~p(s(bool,X1)) | s(bool,X0) = s(bool,X1) | ~p(s(bool,X2))) ) | ($spl4 | $spl17)),
  inference(resolution,[],[f7057,f5182])).
fof(f7487,plain,(
  ( ! [X288,X290,X289] : (s(bool,X288) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12)))) | p(s(bool,X289)) | p(s(bool,X290)) | s(bool,X288) != s(bool,X290)) ) | $spl17),
  inference(resolution,[],[f2082,f839])).
fof(f7398,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12)))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) ) | $spl17),
  inference(resolution,[],[f2082,f586])).
fof(f7395,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,X16) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12)))) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) ) | $spl17),
  inference(resolution,[],[f2082,f585])).
fof(f7394,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12)))) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl17),
  inference(resolution,[],[f2082,f584])).
fof(f7393,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12)))) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) ) | $spl17),
  inference(resolution,[],[f2082,f582])).
fof(f7391,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12)))) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl17),
  inference(resolution,[],[f2082,f580])).
fof(f7140,plain,(
  ( ! [X288,X290,X289] : (s(bool,X288) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) | p(s(bool,X289)) | p(s(bool,X290)) | s(bool,X288) != s(bool,X290)) ) | $spl17),
  inference(resolution,[],[f1958,f839])).
fof(f7051,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) ) | $spl17),
  inference(resolution,[],[f1958,f586])).
fof(f7048,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,X16) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) ) | $spl17),
  inference(resolution,[],[f1958,f585])).
fof(f7047,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl17),
  inference(resolution,[],[f1958,f584])).
fof(f7046,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) ) | $spl17),
  inference(resolution,[],[f1958,f582])).
fof(f7044,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl17),
  inference(resolution,[],[f1958,f580])).
fof(f5840,plain,(
  ( ! [X29,X27] : (~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),X29),s('type.list.list'(sK12),sK14)))))) ) | $spl56),
  inference(cnf_transformation,[],[f5840_D])).
fof(f5840_D,plain,(
  ( ! [X29,X27] : (~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),X29),s('type.list.list'(sK12),sK14)))))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f6652,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),X2),s('type.list.list'(sK12),sK14)))))) ) | ($spl4 | $spl17)),
  inference(resolution,[],[f5726,f5183])).
fof(f5183,plain,(
  ( ! [X4,X5,X3] : (~sP1(X5,X3,X4) | ~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5)) ) | ($spl4 | $spl17)),
  inference(subsumption_resolution,[],[f5171,f597])).
fof(f5171,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | ~sP2(X4,X3,X5) | ~sP1(X5,X3,X4) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5)) ) | ($spl4 | $spl17)),
  inference(resolution,[],[f5167,f609])).
fof(f5726,plain,(
  ( ! [X83,X81,X82] : (sP1('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),X81),s('type.list.list'(sK12),sK14)))),X82,X83)) ) | $spl17),
  inference(resolution,[],[f1595,f601])).
fof(f1595,plain,(
  ( ! [X1] : (~p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),X1),s('type.list.list'(sK12),sK14))))))) ) | $spl17),
  inference(resolution,[],[f1092,f648])).
fof(f5693,plain,(
  ( ! [X29,X27] : (~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),X29)))))) ) | $spl54),
  inference(cnf_transformation,[],[f5693_D])).
fof(f5693_D,plain,(
  ( ! [X29,X27] : (~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),X29)))))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f6436,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),X2)))))) ) | ($spl4 | $spl17)),
  inference(resolution,[],[f5579,f5183])).
fof(f5579,plain,(
  ( ! [X83,X81,X82] : (sP1('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),X81)))),X82,X83)) ) | $spl17),
  inference(resolution,[],[f1594,f601])).
fof(f1594,plain,(
  ( ! [X0] : (~p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),X0))))))) ) | $spl17),
  inference(resolution,[],[f1092,f647])).
fof(f5805,plain,(
  ( ! [X310,X312,X311] : (p(s(bool,X310)) | ~p(s(bool,X311)) | s(bool,X311) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),X312),s('type.list.list'(sK12),sK14)))))) ) | $spl17),
  inference(resolution,[],[f1595,f839])).
fof(f5717,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X54) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),X55),s('type.list.list'(sK12),sK14))))) | s(bool,X53) != s(bool,X56)) ) | $spl17),
  inference(resolution,[],[f1595,f586])).
fof(f5715,plain,(
  ( ! [X47,X48,X46,X49] : (p(s(bool,X46)) | ~p(s(bool,X47)) | ~p(s(bool,X48)) | s(bool,X47) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),X49),s('type.list.list'(sK12),sK14)))))) ) | $spl17),
  inference(resolution,[],[f1595,f585])).
fof(f5713,plain,(
  ( ! [X39,X41,X38,X40] : (p(s(bool,X38)) | ~p(s(bool,X39)) | s(bool,X38) != s(bool,X40) | s(bool,X39) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),X41),s('type.list.list'(sK12),sK14)))))) ) | $spl17),
  inference(resolution,[],[f1595,f584])).
fof(f5712,plain,(
  ( ! [X37,X35,X36,X34] : (~p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X35) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),X36),s('type.list.list'(sK12),sK14))))) | s(bool,X34) != s(bool,X37)) ) | $spl17),
  inference(resolution,[],[f1595,f582])).
fof(f5710,plain,(
  ( ! [X28,X26,X29,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,X28) | s(bool,X27) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),X29),s('type.list.list'(sK12),sK14)))))) ) | $spl17),
  inference(resolution,[],[f1595,f580])).
fof(f5658,plain,(
  ( ! [X310,X312,X311] : (p(s(bool,X310)) | ~p(s(bool,X311)) | s(bool,X311) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),X312)))))) ) | $spl17),
  inference(resolution,[],[f1594,f839])).
fof(f5570,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X54) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),X55))))) | s(bool,X53) != s(bool,X56)) ) | $spl17),
  inference(resolution,[],[f1594,f586])).
fof(f5568,plain,(
  ( ! [X47,X48,X46,X49] : (p(s(bool,X46)) | ~p(s(bool,X47)) | ~p(s(bool,X48)) | s(bool,X47) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),X49)))))) ) | $spl17),
  inference(resolution,[],[f1594,f585])).
fof(f5566,plain,(
  ( ! [X39,X41,X38,X40] : (p(s(bool,X38)) | ~p(s(bool,X39)) | s(bool,X38) != s(bool,X40) | s(bool,X39) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),X41)))))) ) | $spl17),
  inference(resolution,[],[f1594,f584])).
fof(f5565,plain,(
  ( ! [X37,X35,X36,X34] : (~p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X35) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),X36))))) | s(bool,X34) != s(bool,X37)) ) | $spl17),
  inference(resolution,[],[f1594,f582])).
fof(f5563,plain,(
  ( ! [X28,X26,X29,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,X28) | s(bool,X27) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),X29)))))) ) | $spl17),
  inference(resolution,[],[f1594,f580])).
fof(f5547,plain,(
  s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20(sK14,sK13,sK12)))) | $spl52),
  inference(cnf_transformation,[],[f5547_D])).
fof(f5547_D,plain,(
  s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20(sK14,sK13,sK12)))) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f5537,plain,(
  ( ! [X19] : (s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20(sK14,sK13,sK12)))) | p(s(bool,X19))) ) | $spl17),
  inference(resolution,[],[f4787,f1883])).
fof(f4787,plain,(
  ( ! [X241,X239,X240] : (sP10(X240,X239,X241) | s(bool,X239) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)))) ) | $spl17),
  inference(resolution,[],[f1632,f751])).
fof(f5536,plain,(
  ( ! [X18] : (s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20(sK14,sK13,sK12)))) | p(s(bool,X18))) ) | $spl17),
  inference(resolution,[],[f4787,f1885])).
fof(f5535,plain,(
  ( ! [X17] : (s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12)))) | p(s(bool,X17))) ) | $spl17),
  inference(resolution,[],[f4787,f2131])).
fof(f5534,plain,(
  ( ! [X16] : (s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12)))) | p(s(bool,X16))) ) | $spl17),
  inference(resolution,[],[f4787,f2133])).
fof(f5533,plain,(
  ( ! [X15] : (s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) | p(s(bool,X15))) ) | $spl17),
  inference(resolution,[],[f4787,f2007])).
fof(f2007,plain,(
  ( ! [X154,X155] : (~sP10(X155,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12))),X154) | p(s(bool,X154))) ) | $spl17),
  inference(resolution,[],[f1808,f745])).
fof(f5532,plain,(
  ( ! [X14] : (s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) | p(s(bool,X14))) ) | $spl17),
  inference(resolution,[],[f4787,f2009])).
fof(f2009,plain,(
  ( ! [X158,X159] : (~sP10(X158,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12))),X159) | p(s(bool,X158))) ) | $spl17),
  inference(resolution,[],[f1808,f746])).
fof(f5482,plain,(
  ( ! [X59,X60,X58] : (~p(s(bool,X58)) | p(s(bool,X59)) | s(bool,X58) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK30(sK13,X60,sK12)))) | p(s(bool,X60))) ) | ($spl4 | $spl17)),
  inference(resolution,[],[f5183,f3551])).
fof(f3551,plain,(
  ( ! [X70,X72,X71] : (sP1(happ(s(fun(sK12,bool),sK13),s(sK12,sK30(sK13,X70,sK12))),X71,X72) | p(s(bool,X70))) ) | $spl17),
  inference(resolution,[],[f1807,f601])).
fof(f1807,plain,(
  ( ! [X5] : (~p(s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK30(sK13,X5,sK12))))) | p(s(bool,X5))) ) | $spl17),
  inference(resolution,[],[f1605,f803])).
fof(f5481,plain,(
  ( ! [X57,X56,X55] : (~p(s(bool,X55)) | p(s(bool,X56)) | s(bool,X55) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK29(sK13,X57,sK12)))) | p(s(bool,X57))) ) | ($spl4 | $spl17)),
  inference(resolution,[],[f5183,f3426])).
fof(f3426,plain,(
  ( ! [X70,X72,X71] : (sP1(happ(s(fun(sK12,bool),sK13),s(sK12,sK29(sK13,X70,sK12))),X71,X72) | p(s(bool,X70))) ) | $spl17),
  inference(resolution,[],[f1806,f601])).
fof(f1806,plain,(
  ( ! [X4] : (~p(s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK29(sK13,X4,sK12))))) | p(s(bool,X4))) ) | $spl17),
  inference(resolution,[],[f1605,f801])).
fof(f5480,plain,(
  ( ! [X54,X53] : (~p(s(bool,X53)) | p(s(bool,X54)) | s(bool,X53) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20(sK14,sK13,sK12))))) ) | ($spl4 | $spl17)),
  inference(resolution,[],[f5183,f1826])).
fof(f1826,plain,(
  ( ! [X47,X46] : (sP1(happ(s(fun(sK12,bool),sK13),s(sK12,sK20(sK14,sK13,sK12))),X46,X47)) ) | $spl17),
  inference(resolution,[],[f1605,f601])).
fof(f5479,plain,(
  ( ! [X52,X51] : (~p(s(bool,X51)) | p(s(bool,X52)) | s(bool,X51) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12))))) ) | ($spl4 | $spl17)),
  inference(resolution,[],[f5183,f2074])).
fof(f2074,plain,(
  ( ! [X47,X46] : (sP1(happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12))),X46,X47)) ) | $spl17),
  inference(resolution,[],[f1809,f601])).
fof(f5478,plain,(
  ( ! [X50,X49] : (~p(s(bool,X49)) | p(s(bool,X50)) | s(bool,X49) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12))))) ) | ($spl4 | $spl17)),
  inference(resolution,[],[f5183,f1950])).
fof(f1950,plain,(
  ( ! [X47,X46] : (sP1(happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12))),X46,X47)) ) | $spl17),
  inference(resolution,[],[f1808,f601])).
fof(f4679,plain,(
  ( ! [X29,X27] : (~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X29),s('type.list.list'(sK12),sK14)))))) ) | $spl40),
  inference(cnf_transformation,[],[f4679_D])).
fof(f4679_D,plain,(
  ( ! [X29,X27] : (~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X29),s('type.list.list'(sK12),sK14)))))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f5477,plain,(
  ( ! [X47,X48,X46] : (~p(s(bool,X46)) | p(s(bool,X47)) | s(bool,X46) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X48),s('type.list.list'(sK12),sK14)))))) ) | ($spl4 | $spl17)),
  inference(resolution,[],[f5183,f4575])).
fof(f4575,plain,(
  ( ! [X83,X81,X82] : (sP1('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X81),s('type.list.list'(sK12),sK14)))),X82,X83)) ) | $spl17),
  inference(resolution,[],[f1597,f601])).
fof(f1597,plain,(
  ( ! [X3] : (~p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X3),s('type.list.list'(sK12),sK14))))))) ) | $spl17),
  inference(resolution,[],[f1092,f654])).
fof(f4542,plain,(
  ( ! [X29,X27] : (~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X29),s('type.list.list'(sK12),sK14)))))) ) | $spl38),
  inference(cnf_transformation,[],[f4542_D])).
fof(f4542_D,plain,(
  ( ! [X29,X27] : (~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X29),s('type.list.list'(sK12),sK14)))))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f5476,plain,(
  ( ! [X45,X43,X44] : (~p(s(bool,X43)) | p(s(bool,X44)) | s(bool,X43) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X45),s('type.list.list'(sK12),sK14)))))) ) | ($spl4 | $spl17)),
  inference(resolution,[],[f5183,f4438])).
fof(f4438,plain,(
  ( ! [X83,X81,X82] : (sP1('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X81),s('type.list.list'(sK12),sK14)))),X82,X83)) ) | $spl17),
  inference(resolution,[],[f1596,f601])).
fof(f1596,plain,(
  ( ! [X2] : (~p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X2),s('type.list.list'(sK12),sK14))))))) ) | $spl17),
  inference(resolution,[],[f1092,f650])).
fof(f5442,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f5373,f1006])).
fof(f5373,plain,(
  ( ! [X512,X514,X513] : (p(s(bool,'const.list.EVERY_2'(s(fun(X512,bool),X513),s('type.list.list'(X512),X514))))) ) | $spl12),
  inference(resolution,[],[f1006,f665])).
fof(f5441,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f5323,f1006])).
fof(f5323,plain,(
  ( ! [X333,X332,X334] : (p(s(bool,'const.bool.IN_2'(s(X332,sK20(X333,X334,X332)),s(fun(X332,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X332),X333))))))) ) | $spl12),
  inference(resolution,[],[f1006,f665])).
fof(f5331,plain,(
  $false | ($spl0 | $spl12)),
  inference(resolution,[],[f1006,f864])).
fof(f5374,plain,(
  $false | ($spl12 | $spl17)),
  inference(resolution,[],[f1006,f1604])).
fof(f5213,plain,(
  ( ! [X17] : (~p(s(bool,X17)) | s(bool,X17) = s(bool,'const.bool.IN_2'(s(sK12,sK20(sK14,sK13,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK14)))))) ) | $spl46),
  inference(cnf_transformation,[],[f5213_D])).
fof(f5213_D,plain,(
  ( ! [X17] : (~p(s(bool,X17)) | s(bool,X17) = s(bool,'const.bool.IN_2'(s(sK12,sK20(sK14,sK13,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK14)))))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f5195,plain,(
  ( ! [X17,X18] : (~p(s(bool,X17)) | s(bool,X17) = s(bool,'const.bool.IN_2'(s(sK12,sK20(sK14,sK13,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK14))))) | ~p(s(bool,X18))) ) | ($spl4 | $spl17)),
  inference(resolution,[],[f5182,f3298])).
fof(f3298,plain,(
  ( ! [X35,X36] : (sP2(X35,X36,'const.bool.IN_2'(s(sK12,sK20(sK14,sK13,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK14)))))) ) | $spl17),
  inference(resolution,[],[f1604,f595])).
fof(f5212,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,X1) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) | s(bool,X0) = s(bool,X1)) ) | $spl44),
  inference(cnf_transformation,[],[f5212_D])).
fof(f5212_D,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,X1) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) | s(bool,X0) = s(bool,X1)) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f5189,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,X1) | ~p(s(bool,X2)) | s(bool,X1) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)))) ) | ($spl4 | $spl17)),
  inference(resolution,[],[f5182,f4721])).
fof(f4721,plain,(
  ( ! [X47,X48,X49] : (sP2(X48,X49,X47) | s(bool,X47) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)))) ) | $spl17),
  inference(resolution,[],[f1632,f595])).
fof(f4903,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) | $spl42),
  inference(cnf_transformation,[],[f4903_D])).
fof(f4903_D,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f4804,plain,(
  ( ! [X288,X290,X289] : (s(bool,X288) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) | p(s(bool,X289)) | p(s(bool,X290)) | s(bool,X288) != s(bool,X290)) ) | $spl17),
  inference(resolution,[],[f1632,f839])).
fof(f4715,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) ) | $spl17),
  inference(resolution,[],[f1632,f586])).
fof(f4712,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,X16) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) ) | $spl17),
  inference(resolution,[],[f1632,f585])).
fof(f4711,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl17),
  inference(resolution,[],[f1632,f584])).
fof(f4710,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) ) | $spl17),
  inference(resolution,[],[f1632,f582])).
fof(f4708,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl17),
  inference(resolution,[],[f1632,f580])).
fof(f4654,plain,(
  ( ! [X310,X312,X311] : (p(s(bool,X310)) | ~p(s(bool,X311)) | s(bool,X311) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X312),s('type.list.list'(sK12),sK14)))))) ) | $spl17),
  inference(resolution,[],[f1597,f839])).
fof(f4566,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X54) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X55),s('type.list.list'(sK12),sK14))))) | s(bool,X53) != s(bool,X56)) ) | $spl17),
  inference(resolution,[],[f1597,f586])).
fof(f4564,plain,(
  ( ! [X47,X48,X46,X49] : (p(s(bool,X46)) | ~p(s(bool,X47)) | ~p(s(bool,X48)) | s(bool,X47) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X49),s('type.list.list'(sK12),sK14)))))) ) | $spl17),
  inference(resolution,[],[f1597,f585])).
fof(f4562,plain,(
  ( ! [X39,X41,X38,X40] : (p(s(bool,X38)) | ~p(s(bool,X39)) | s(bool,X38) != s(bool,X40) | s(bool,X39) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X41),s('type.list.list'(sK12),sK14)))))) ) | $spl17),
  inference(resolution,[],[f1597,f584])).
fof(f4561,plain,(
  ( ! [X37,X35,X36,X34] : (~p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X35) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X36),s('type.list.list'(sK12),sK14))))) | s(bool,X34) != s(bool,X37)) ) | $spl17),
  inference(resolution,[],[f1597,f582])).
fof(f4559,plain,(
  ( ! [X28,X26,X29,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,X28) | s(bool,X27) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X29),s('type.list.list'(sK12),sK14)))))) ) | $spl17),
  inference(resolution,[],[f1597,f580])).
fof(f4517,plain,(
  ( ! [X310,X312,X311] : (p(s(bool,X310)) | ~p(s(bool,X311)) | s(bool,X311) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X312),s('type.list.list'(sK12),sK14)))))) ) | $spl17),
  inference(resolution,[],[f1596,f839])).
fof(f4429,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X54) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X55),s('type.list.list'(sK12),sK14))))) | s(bool,X53) != s(bool,X56)) ) | $spl17),
  inference(resolution,[],[f1596,f586])).
fof(f4427,plain,(
  ( ! [X47,X48,X46,X49] : (p(s(bool,X46)) | ~p(s(bool,X47)) | ~p(s(bool,X48)) | s(bool,X47) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X49),s('type.list.list'(sK12),sK14)))))) ) | $spl17),
  inference(resolution,[],[f1596,f585])).
fof(f4425,plain,(
  ( ! [X39,X41,X38,X40] : (p(s(bool,X38)) | ~p(s(bool,X39)) | s(bool,X38) != s(bool,X40) | s(bool,X39) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X41),s('type.list.list'(sK12),sK14)))))) ) | $spl17),
  inference(resolution,[],[f1596,f584])).
fof(f4424,plain,(
  ( ! [X37,X35,X36,X34] : (~p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X35) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X36),s('type.list.list'(sK12),sK14))))) | s(bool,X34) != s(bool,X37)) ) | $spl17),
  inference(resolution,[],[f1596,f582])).
fof(f4422,plain,(
  ( ! [X28,X26,X29,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,X28) | s(bool,X27) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X29),s('type.list.list'(sK12),sK14)))))) ) | $spl17),
  inference(resolution,[],[f1596,f580])).
fof(f4318,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) | $spl36),
  inference(cnf_transformation,[],[f4318_D])).
fof(f4318_D,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f4225,plain,(
  ( ! [X285,X287,X286] : (s(bool,X285) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))) | p(s(bool,X286)) | ~p(s(bool,X287)) | s(bool,X285) != s(bool,X287)) ) | $spl0),
  inference(resolution,[],[f1411,f839])).
fof(f1411,plain,(
  ( ! [X67] : (~p(s(bool,X67)) | s(bool,X67) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) ) | $spl0),
  inference(resolution,[],[f864,f618])).
fof(f4137,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) ) | $spl0),
  inference(resolution,[],[f1411,f586])).
fof(f4135,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,X20) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X20) != s(bool,X22)) ) | $spl0),
  inference(resolution,[],[f1411,f585])).
fof(f4133,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl0),
  inference(resolution,[],[f1411,f584])).
fof(f4132,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) ) | $spl0),
  inference(resolution,[],[f1411,f582])).
fof(f4130,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl0),
  inference(resolution,[],[f1411,f580])).
fof(f3948,plain,(
  $false | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f3929,f864])).
fof(f3929,plain,(
  ~p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) | $spl6),
  inference(equality_resolution,[],[f1003])).
fof(f3947,plain,(
  ( ! [X39,X38] : (~p(s(bool,X39)) | ~p(s(bool,X38))) ) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f3946,f864])).
fof(f3946,plain,(
  ( ! [X39,X38] : (~p(s(bool,X39)) | ~p(s(bool,X38)) | ~p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))))) ) | $spl6),
  inference(subsumption_resolution,[],[f3928,f618])).
fof(f3928,plain,(
  ( ! [X39,X38] : (s(bool,X38) != s(bool,X39) | ~p(s(bool,X39)) | ~p(s(bool,X38)) | ~p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))))) ) | $spl6),
  inference(superposition,[],[f1003,f622])).
fof(f3945,plain,(
  ( ! [X37,X36] : (~p(s(bool,X37)) | ~p(s(bool,X36))) ) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f3944,f864])).
fof(f3944,plain,(
  ( ! [X37,X36] : (~p(s(bool,X37)) | ~p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) | ~p(s(bool,X36))) ) | $spl6),
  inference(subsumption_resolution,[],[f3927,f618])).
fof(f3927,plain,(
  ( ! [X37,X36] : (s(bool,X36) != s(bool,X37) | ~p(s(bool,X37)) | ~p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) | ~p(s(bool,X36))) ) | $spl6),
  inference(superposition,[],[f1003,f622])).
fof(f3943,plain,(
  ( ! [X30,X31] : (~p(s(bool,X31)) | ~p(s(bool,X30))) ) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f3942,f864])).
fof(f3942,plain,(
  ( ! [X30,X31] : (~p(s(bool,X31)) | ~p(s(bool,X30)) | ~p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))))) ) | $spl6),
  inference(subsumption_resolution,[],[f3924,f618])).
fof(f3924,plain,(
  ( ! [X30,X31] : (s(bool,X30) != s(bool,X31) | ~p(s(bool,X31)) | ~p(s(bool,X30)) | ~p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))))) ) | $spl6),
  inference(superposition,[],[f1003,f618])).
fof(f3941,plain,(
  ( ! [X28,X29] : (~p(s(bool,X29)) | ~p(s(bool,X28))) ) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f3940,f864])).
fof(f3940,plain,(
  ( ! [X28,X29] : (~p(s(bool,X29)) | ~p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) | ~p(s(bool,X28))) ) | $spl6),
  inference(subsumption_resolution,[],[f3923,f618])).
fof(f3923,plain,(
  ( ! [X28,X29] : (s(bool,X28) != s(bool,X29) | ~p(s(bool,X29)) | ~p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) | ~p(s(bool,X28))) ) | $spl6),
  inference(superposition,[],[f1003,f618])).
fof(f3630,plain,(
  ( ! [X300,X301,X299] : (p(s(bool,X299)) | p(s(bool,X300)) | ~p(s(bool,X301)) | s(bool,X301) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK30(sK13,X299,sK12))))) ) | $spl17),
  inference(resolution,[],[f1807,f839])).
fof(f3542,plain,(
  ( ! [X45,X43,X44,X42] : (p(s(bool,X42)) | ~p(s(bool,X43)) | ~p(s(bool,X44)) | s(bool,X44) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK30(sK13,X42,sK12)))) | s(bool,X43) != s(bool,X45)) ) | $spl17),
  inference(resolution,[],[f1807,f586])).
fof(f3540,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,X34)) | p(s(bool,X35)) | ~p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK30(sK13,X34,sK12))))) ) | $spl17),
  inference(resolution,[],[f1807,f585])).
fof(f3538,plain,(
  ( ! [X28,X26,X29,X27] : (p(s(bool,X26)) | p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,X29) | s(bool,X28) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK30(sK13,X26,sK12))))) ) | $spl17),
  inference(resolution,[],[f1807,f584])).
fof(f3537,plain,(
  ( ! [X24,X23,X25,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK30(sK13,X22,sK12)))) | s(bool,X23) != s(bool,X25)) ) | $spl17),
  inference(resolution,[],[f1807,f582])).
fof(f3535,plain,(
  ( ! [X14,X17,X15,X16] : (p(s(bool,X14)) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK30(sK13,X14,sK12))))) ) | $spl17),
  inference(resolution,[],[f1807,f580])).
fof(f3505,plain,(
  ( ! [X300,X301,X299] : (p(s(bool,X299)) | p(s(bool,X300)) | ~p(s(bool,X301)) | s(bool,X301) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK29(sK13,X299,sK12))))) ) | $spl17),
  inference(resolution,[],[f1806,f839])).
fof(f3417,plain,(
  ( ! [X45,X43,X44,X42] : (p(s(bool,X42)) | ~p(s(bool,X43)) | ~p(s(bool,X44)) | s(bool,X44) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK29(sK13,X42,sK12)))) | s(bool,X43) != s(bool,X45)) ) | $spl17),
  inference(resolution,[],[f1806,f586])).
fof(f3415,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,X34)) | p(s(bool,X35)) | ~p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK29(sK13,X34,sK12))))) ) | $spl17),
  inference(resolution,[],[f1806,f585])).
fof(f3413,plain,(
  ( ! [X28,X26,X29,X27] : (p(s(bool,X26)) | p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,X29) | s(bool,X28) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK29(sK13,X26,sK12))))) ) | $spl17),
  inference(resolution,[],[f1806,f584])).
fof(f3412,plain,(
  ( ! [X24,X23,X25,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK29(sK13,X22,sK12)))) | s(bool,X23) != s(bool,X25)) ) | $spl17),
  inference(resolution,[],[f1806,f582])).
fof(f3410,plain,(
  ( ! [X14,X17,X15,X16] : (p(s(bool,X14)) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK29(sK13,X14,sK12))))) ) | $spl17),
  inference(resolution,[],[f1806,f580])).
fof(f3407,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | ~p(s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK30(sK13,X11,sK12)))))) ) | $spl17),
  inference(resolution,[],[f1806,f803])).
fof(f3401,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK12,sK20(sK14,sK13,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK14)))))) ) | $spl28),
  inference(cnf_transformation,[],[f3401_D])).
fof(f3401_D,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK12,sK20(sK14,sK13,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK14)))))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f3381,plain,(
  ( ! [X194,X193] : (p(s(bool,X193)) | p(s(bool,X194)) | s(bool,X194) != s(bool,'const.bool.IN_2'(s(sK12,sK20(sK14,sK13,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK14)))))) ) | $spl17),
  inference(resolution,[],[f1604,f839])).
fof(f3292,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.bool.IN_2'(s(sK12,sK20(sK14,sK13,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK14))))) | s(bool,X23) != s(bool,X24)) ) | $spl17),
  inference(resolution,[],[f1604,f586])).
fof(f3289,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.bool.IN_2'(s(sK12,sK20(sK14,sK13,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK14)))))) ) | $spl17),
  inference(resolution,[],[f1604,f585])).
fof(f3288,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.bool.IN_2'(s(sK12,sK20(sK14,sK13,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK14)))))) ) | $spl17),
  inference(resolution,[],[f1604,f584])).
fof(f3287,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.bool.IN_2'(s(sK12,sK20(sK14,sK13,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK14))))) | s(bool,X8) != s(bool,X9)) ) | $spl17),
  inference(resolution,[],[f1604,f582])).
fof(f3285,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK12,sK20(sK14,sK13,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK14)))))) ) | $spl17),
  inference(resolution,[],[f1604,f580])).
fof(f1752,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) | s(bool,X2) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) ) | $spl20),
  inference(cnf_transformation,[],[f1752_D])).
fof(f1752_D,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) | s(bool,X2) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f2229,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | s(bool,X5) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) | s(bool,X5) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) )),
  inference(duplicate_literal_removal,[],[f2223])).
fof(f2223,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | s(bool,X5) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) | s(bool,X5) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))) | ~p(s(bool,X4))) )),
  inference(resolution,[],[f853,f607])).
fof(f607,plain,(
  ( ! [X2,X0,X1] : (sP3(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f465])).
fof(f2182,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | s(bool,X5) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) | s(bool,X5) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) )),
  inference(duplicate_literal_removal,[],[f2176])).
fof(f2176,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | s(bool,X5) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) | s(bool,X5) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))) | ~p(s(bool,X4))) )),
  inference(resolution,[],[f852,f607])).
fof(f2153,plain,(
  ( ! [X196,X197] : (p(s(bool,X196)) | ~p(s(bool,X197)) | s(bool,X197) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12))))) ) | $spl17),
  inference(resolution,[],[f1809,f839])).
fof(f2065,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X28) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12)))) | s(bool,X27) != s(bool,X29)) ) | $spl17),
  inference(resolution,[],[f1809,f586])).
fof(f2063,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12))))) ) | $spl17),
  inference(resolution,[],[f1809,f585])).
fof(f2061,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12))))) ) | $spl17),
  inference(resolution,[],[f1809,f584])).
fof(f2060,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X13) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12)))) | s(bool,X12) != s(bool,X14)) ) | $spl17),
  inference(resolution,[],[f1809,f582])).
fof(f2058,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK44(sK13,sK12))))) ) | $spl17),
  inference(resolution,[],[f1809,f580])).
fof(f2029,plain,(
  ( ! [X196,X197] : (p(s(bool,X196)) | ~p(s(bool,X197)) | s(bool,X197) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12))))) ) | $spl17),
  inference(resolution,[],[f1808,f839])).
fof(f1941,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X28) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) | s(bool,X27) != s(bool,X29)) ) | $spl17),
  inference(resolution,[],[f1808,f586])).
fof(f1939,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12))))) ) | $spl17),
  inference(resolution,[],[f1808,f585])).
fof(f1937,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12))))) ) | $spl17),
  inference(resolution,[],[f1808,f584])).
fof(f1936,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X13) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12)))) | s(bool,X12) != s(bool,X14)) ) | $spl17),
  inference(resolution,[],[f1808,f582])).
fof(f1934,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK43(sK13,sK12))))) ) | $spl17),
  inference(resolution,[],[f1808,f580])).
fof(f1905,plain,(
  ( ! [X196,X197] : (p(s(bool,X196)) | ~p(s(bool,X197)) | s(bool,X197) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20(sK14,sK13,sK12))))) ) | $spl17),
  inference(resolution,[],[f1605,f839])).
fof(f1817,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X28) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20(sK14,sK13,sK12)))) | s(bool,X27) != s(bool,X29)) ) | $spl17),
  inference(resolution,[],[f1605,f586])).
fof(f1815,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20(sK14,sK13,sK12))))) ) | $spl17),
  inference(resolution,[],[f1605,f585])).
fof(f1813,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20(sK14,sK13,sK12))))) ) | $spl17),
  inference(resolution,[],[f1605,f584])).
fof(f1812,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X13) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20(sK14,sK13,sK12)))) | s(bool,X12) != s(bool,X14)) ) | $spl17),
  inference(resolution,[],[f1605,f582])).
fof(f1810,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,sK20(sK14,sK13,sK12))))) ) | $spl17),
  inference(resolution,[],[f1605,f580])).
fof(f1793,plain,(
  ( ! [X2,X3] : (s(bool,X2) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) )),
  inference(duplicate_literal_removal,[],[f1788])).
fof(f1788,plain,(
  ( ! [X2,X3] : (s(bool,X2) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) | p(s(bool,X3)) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) )),
  inference(resolution,[],[f851,f609])).
fof(f851,plain,(
  ( ! [X10,X11] : (sP3(X11,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))),X10) | s(bool,X10) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) | p(s(bool,X11))) )),
  inference(superposition,[],[f578,f606])).
fof(f606,plain,(
  ( ! [X2,X0,X1] : (sP3(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f465])).
fof(f1751,plain,(
  ( ! [X2,X3] : (s(bool,X2) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) )),
  inference(duplicate_literal_removal,[],[f1746])).
fof(f1746,plain,(
  ( ! [X2,X3] : (s(bool,X2) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) | p(s(bool,X3)) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) )),
  inference(resolution,[],[f850,f609])).
fof(f850,plain,(
  ( ! [X8,X9] : (sP3(X9,X8,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))) | s(bool,X8) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) | p(s(bool,X9))) )),
  inference(superposition,[],[f578,f606])).
fof(f1092,plain,(
  ~p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)))) | $spl17),
  inference(cnf_transformation,[],[f1092_D])).
fof(f1092_D,plain,(
  ~p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)))) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f1593,plain,(
  ( ! [X39,X38] : (~p(s(bool,X39)) | ~p(s(bool,X38)) | ~p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))))) ) | $spl4),
  inference(subsumption_resolution,[],[f1579,f618])).
fof(f1579,plain,(
  ( ! [X39,X38] : (s(bool,X38) != s(bool,X39) | ~p(s(bool,X39)) | ~p(s(bool,X38)) | ~p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))))) ) | $spl4),
  inference(superposition,[],[f866,f622])).
fof(f1592,plain,(
  ( ! [X37,X36] : (~p(s(bool,X37)) | ~p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)))) | ~p(s(bool,X36))) ) | $spl4),
  inference(subsumption_resolution,[],[f1578,f618])).
fof(f1578,plain,(
  ( ! [X37,X36] : (s(bool,X36) != s(bool,X37) | ~p(s(bool,X37)) | ~p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)))) | ~p(s(bool,X36))) ) | $spl4),
  inference(superposition,[],[f866,f622])).
fof(f1591,plain,(
  ( ! [X30,X31] : (~p(s(bool,X31)) | ~p(s(bool,X30)) | ~p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))))) ) | $spl4),
  inference(subsumption_resolution,[],[f1575,f618])).
fof(f1575,plain,(
  ( ! [X30,X31] : (s(bool,X30) != s(bool,X31) | ~p(s(bool,X31)) | ~p(s(bool,X30)) | ~p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))))) ) | $spl4),
  inference(superposition,[],[f866,f618])).
fof(f1590,plain,(
  ( ! [X28,X29] : (~p(s(bool,X29)) | ~p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)))) | ~p(s(bool,X28))) ) | $spl4),
  inference(subsumption_resolution,[],[f1574,f618])).
fof(f1574,plain,(
  ( ! [X28,X29] : (s(bool,X28) != s(bool,X29) | ~p(s(bool,X29)) | ~p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)))) | ~p(s(bool,X28))) ) | $spl4),
  inference(superposition,[],[f866,f618])).
fof(f1512,plain,(
  ( ! [X14] : (p(s(bool,X14)) | s(bool,X14) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) ) | $spl18),
  inference(cnf_transformation,[],[f1512_D])).
fof(f1512_D,plain,(
  ( ! [X14] : (p(s(bool,X14)) | s(bool,X14) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f1483,plain,(
  ( ! [X206,X205] : (p(s(bool,X205)) | p(s(bool,X206)) | s(bool,X206) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) ) | $spl0),
  inference(resolution,[],[f864,f839])).
fof(f1394,plain,(
  ( ! [X35,X36,X34] : (p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X34) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))) | s(bool,X35) != s(bool,X36)) ) | $spl0),
  inference(resolution,[],[f864,f586])).
fof(f1391,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | p(s(bool,X27)) | s(bool,X27) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) ) | $spl0),
  inference(resolution,[],[f864,f585])).
fof(f1390,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | p(s(bool,X23)) | s(bool,X22) != s(bool,X24) | s(bool,X23) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) ) | $spl0),
  inference(resolution,[],[f864,f584])).
fof(f1389,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))) | s(bool,X20) != s(bool,X21)) ) | $spl0),
  inference(resolution,[],[f864,f582])).
fof(f1387,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) ) | $spl0),
  inference(resolution,[],[f864,f580])).
fof(f1362,plain,(
  $false | ($spl4 | $spl14)),
  inference(equality_resolution,[],[f1317])).
fof(f1317,plain,(
  ( ! [X18] : (s(bool,X18) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)))) ) | ($spl4 | $spl14)),
  inference(subsumption_resolution,[],[f866,f1007])).
fof(f1221,plain,(
  $false | ($spl1 | $spl14)),
  inference(resolution,[],[f1007,f863])).
fof(f1093,plain,(
  p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)))) | $spl16),
  inference(cnf_transformation,[],[f1093_D])).
fof(f1093_D,plain,(
  p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f1098,plain,(
  ( ! [X35,X34] : (p(s(bool,X35)) | p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)))) | p(s(bool,X34))) ) | $spl2),
  inference(subsumption_resolution,[],[f1078,f615])).
fof(f1078,plain,(
  ( ! [X35,X34] : (s(bool,X34) != s(bool,X35) | p(s(bool,X35)) | p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)))) | p(s(bool,X34))) ) | $spl2),
  inference(superposition,[],[f865,f619])).
fof(f619,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f396])).
fof(f1097,plain,(
  ( ! [X33,X32] : (p(s(bool,X33)) | p(s(bool,X32)) | p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))))) ) | $spl2),
  inference(subsumption_resolution,[],[f1077,f615])).
fof(f1077,plain,(
  ( ! [X33,X32] : (s(bool,X32) != s(bool,X33) | p(s(bool,X33)) | p(s(bool,X32)) | p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))))) ) | $spl2),
  inference(superposition,[],[f865,f619])).
fof(f1094,plain,(
  ( ! [X26,X27] : (p(s(bool,X27)) | p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)))) | p(s(bool,X26))) ) | $spl2),
  inference(subsumption_resolution,[],[f1074,f615])).
fof(f1074,plain,(
  ( ! [X26,X27] : (s(bool,X26) != s(bool,X27) | p(s(bool,X27)) | p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14)))) | p(s(bool,X26))) ) | $spl2),
  inference(superposition,[],[f865,f615])).
fof(f1091,plain,(
  ( ! [X24,X25] : (p(s(bool,X25)) | p(s(bool,X24)) | p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))))) ) | $spl2),
  inference(subsumption_resolution,[],[f1073,f615])).
fof(f1073,plain,(
  ( ! [X24,X25] : (s(bool,X24) != s(bool,X25) | p(s(bool,X25)) | p(s(bool,X24)) | p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))))) ) | $spl2),
  inference(superposition,[],[f865,f615])).
fof(f976,plain,(
  ( ! [X202,X201] : (p(s(bool,X201)) | ~p(s(bool,X202)) | s(bool,X202) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) ) | $spl1),
  inference(resolution,[],[f863,f839])).
fof(f888,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X33) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))) | s(bool,X32) != s(bool,X34)) ) | $spl1),
  inference(resolution,[],[f863,f586])).
fof(f886,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) ) | $spl1),
  inference(resolution,[],[f863,f585])).
fof(f884,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X20) != s(bool,X22) | s(bool,X21) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) ) | $spl1),
  inference(resolution,[],[f863,f584])).
fof(f883,plain,(
  ( ! [X19,X17,X18] : (~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X18) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))) | s(bool,X17) != s(bool,X19)) ) | $spl1),
  inference(resolution,[],[f863,f582])).
fof(f881,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) ) | $spl1),
  inference(resolution,[],[f863,f580])).
fof(f866,plain,(
  ( ! [X18] : (s(bool,X18) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) | ~p(s(bool,X18))) ) | $spl4),
  inference(cnf_transformation,[],[f866_D])).
fof(f866_D,plain,(
  ( ! [X18] : (s(bool,X18) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) | ~p(s(bool,X18))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f861,plain,(
  ( ! [X23] : (s(bool,X23) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) | ~p(s(bool,X23)) | ~p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))))) )),
  inference(superposition,[],[f578,f622])).
fof(f860,plain,(
  ( ! [X22] : (s(bool,X22) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) | ~p(s(bool,X22))) )),
  inference(superposition,[],[f578,f622])).
fof(f865,plain,(
  ( ! [X16] : (s(bool,X16) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) | p(s(bool,X16))) ) | $spl2),
  inference(cnf_transformation,[],[f865_D])).
fof(f865_D,plain,(
  ( ! [X16] : (s(bool,X16) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) | p(s(bool,X16))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f859,plain,(
  ( ! [X21] : (s(bool,X21) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) | p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) | p(s(bool,X21))) )),
  inference(superposition,[],[f578,f619])).
fof(f858,plain,(
  ( ! [X20] : (s(bool,X20) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) | p(s(bool,X20)) | p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))))) )),
  inference(superposition,[],[f578,f619])).
fof(f857,plain,(
  ( ! [X19] : (s(bool,X19) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) | ~p(s(bool,X19)) | ~p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))))) )),
  inference(superposition,[],[f578,f618])).
fof(f856,plain,(
  ( ! [X18] : (s(bool,X18) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) | ~p(s(bool,X18))) )),
  inference(superposition,[],[f578,f618])).
fof(f855,plain,(
  ( ! [X17] : (s(bool,X17) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) | p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14)))))) | p(s(bool,X17))) )),
  inference(superposition,[],[f578,f615])).
fof(f854,plain,(
  ( ! [X16] : (s(bool,X16) != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK14))) | p(s(bool,X16)) | p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.REVERSE_1'(s('type.list.list'(sK12),sK14))))))) )),
  inference(superposition,[],[f578,f615])).
% SZS output end Proof for thm.rich_list
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 832
% Passive clauses: 5593
% Generated clauses: 14377
% Final active clauses: 466
% Final passive clauses: 2154
% Input formulas: 261
% Initial clauses: 261
% 
% Selected by SInE selection: 72
% SInE iterations: 1
% 
% Duplicate literals: 906
% Trivial inequalities: 13
% Fw subsumption resolutions: 855
% Fw demodulations: 372
% 
% Simple tautologies: 1675
% Equational tautologies: 38
% Forward subsumptions: 4793
% 
% Binary resolution: 10206
% Forward superposition: 1548
% Backward superposition: 269
% Equality resolution: 21
% 
% Split clauses: 350
% Split components: 779
% SAT solver clauses: 401
% SAT solver unit clauses: 2
% SAT solver binary clauses: 22
% 
% Sat splits: 350
% Sat splitting refutations: 10
% 
% Memory used [KB]: 6780
% Time elapsed: 0.323 s
% ------------------------------
% ------------------------------
% Success in time 12.081 s
