Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.pred_set
% (84205)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 104
% Passive clauses: 2102
% Generated clauses: 20245
% Final active clauses: 90
% Final passive clauses: 1756
% Input formulas: 23
% Initial clauses: 51
% 
% Duplicate literals: 1048
% Trivial inequalities: 288
% Fw subsumption resolutions: 2623
% Fw demodulations: 2484
% Bw demodulations: 7
% 
% Simple tautologies: 351
% Equational tautologies: 73
% Forward subsumptions: 8141
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 288
% Factoring: 5
% Forward superposition: 5904
% Backward superposition: 6859
% Self superposition: 303
% Equality factoring: 219
% Equality resolution: 20
% 
% Split clauses: 627
% Split components: 1267
% SAT solver clauses: 1341
% SAT solver unit clauses: 7
% SAT solver binary clauses: 103
% 
% Sat splits: 627
% Sat splitting refutations: 11
% 
% TWLsolver clauses: 1339
% TWLsolver calls for satisfiability: 144
% 
% Memory used [KB]: 8827
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pred_set
% (84218)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 67
% Passive clauses: 479
% Generated clauses: 8428
% Final active clauses: 60
% Final passive clauses: 403
% Input formulas: 23
% Initial clauses: 51
% 
% Duplicate literals: 727
% Trivial inequalities: 111
% Fw subsumption resolutions: 670
% Fw demodulations: 439
% Bw demodulations: 10
% Condensations: 1146
% Global subsumptions: 9
% 
% Simple tautologies: 181
% Equational tautologies: 51
% Forward subsumptions: 3201
% 
% Binary resolution: 162
% Factoring: 3
% Forward superposition: 2340
% Backward superposition: 1981
% Self superposition: 169
% Equality factoring: 131
% Equality resolution: 10
% 
% Split clauses: 68
% Split components: 137
% SAT solver clauses: 875
% SAT solver unit clauses: 57
% SAT solver binary clauses: 49
% SAT solver learnt clauses: 18
% 
% Sat splits: 68
% Sat splitting refutations: 10
% 
% TWLsolver clauses: 839
% TWLsolver calls for satisfiability: 183
% 
% Memory used [KB]: 6268
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f7772,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f148,f151,f149,f174,f178,f176,f175,f187,f230,f234,f232,f423,f427,f425,f1056,f1057,f1058,f1059,f1070,f1060,f1069,f1067,f1071,f1072,f1066,f1064,f1154,f1155,f1156,f1157,f1168,f1158,f1167,f1165,f1169,f1170,f1164,f1162,f1260,f1269,f1261,f1262,f1263,f1273,f1266,f1264,f1272,f1267,f1270,f1526,f1529,f1547,f1545,f2972,f3416,f4070,f4071,f4072,f4076,f4079,f4082,f4085,f4088,f4091,f4124,f4125,f4126,f4130,f4133,f4136,f4139,f4142,f4145,f7545,f7547,f7551,f7556,f7559,f7570,f7575,f7584,f7590,f7597,f7604,f7609,f7613,f7616,f7623,f7628,f7633,f7637,f7643,f7647,f7651,f7654,f7660,f7665,f7670,f7675,f7679,f7684,f7689,f7694,f7698,f7701,f7704,f7708,f7711,f7714,f7716,f7721,f7725,f7729,f7732,f7735,f7737,f7739,f7741,f7743,f7746,f7748,f7750,f7752,f7754,f7756,f7758,f7760,f7762,f7764,f7766,f7768,f7771])).
fof(f7771,plain,(
  $false | ($spl17 | $spl19)),
  inference(subsumption_resolution,[],[f7770,f87])).
fof(f87,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.UNION_EMPTY|split|1',aHL_TRUTH)).
fof(f7770,plain,(
  ~p(s(bool,'T_0')) | ($spl17 | $spl19)),
  inference(forward_demodulation,[],[f7453,f4996])).
fof(f4996,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl19),
  inference(unit_resulting_resolution,[],[f89,f1066,f177])).
fof(f177,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f169,f89])).
fof(f169,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f132,f90])).
fof(f90,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.pred_set.UNION_EMPTY|split|1',aHL_BOOL_CASES)).
fof(f132,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f53])).
fof(f53,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f48])).
fof(f48,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X8,X9] : ((p(s(bool,X8)) => p(s(bool,X9))) => ((p(s(bool,X9)) => p(s(bool,X8))) => s(bool,X8) = s(bool,X9)))),
  file('thm.pred_set.UNION_EMPTY|split|1','thm.bool.IMP_ANTISYM_AX')).
fof(f89,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f27])).
fof(f27,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.pred_set.UNION_EMPTY|split|1',aHL_FALSITY)).
fof(f7453,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1)))) | $spl17),
  inference(unit_resulting_resolution,[],[f1063,f283])).
fof(f283,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))))) )),
  inference(subsumption_resolution,[],[f280,f89])).
fof(f280,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))))) )),
  inference(superposition,[],[f130,f90])).
fof(f130,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) & ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))))),
  inference(flattening,[],[f83])).
fof(f83,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))))),
  inference(nnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(flattening,[],[f46])).
fof(f46,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X6] : ! [X7,X10,X11] : (p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X10)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X7)))) | p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X10))))))),
  file('thm.pred_set.UNION_EMPTY|split|1','thm.pred_set.IN_UNION')).
fof(f1063,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | $spl17),
  inference(cnf_transformation,[],[f1063_D])).
fof(f1063_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f7768,plain,(
  $false | ($spl3 | $spl17 | $spl19)),
  inference(subsumption_resolution,[],[f7767,f151])).
fof(f7767,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl17 | $spl19)),
  inference(forward_demodulation,[],[f7455,f7439])).
fof(f7439,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X1)))))) ) | $spl19),
  inference(subsumption_resolution,[],[f7374,f87])).
fof(f7374,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X1)))))) ) | $spl19),
  inference(superposition,[],[f283,f4996])).
fof(f7455,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | $spl17),
  inference(unit_resulting_resolution,[],[f1063,f90])).
fof(f7766,plain,(
  $false | ($spl3 | $spl17 | $spl19)),
  inference(subsumption_resolution,[],[f7765,f151])).
fof(f7765,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl3 | $spl17 | $spl19)),
  inference(forward_demodulation,[],[f7456,f7439])).
fof(f7456,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | ($spl3 | $spl17)),
  inference(unit_resulting_resolution,[],[f151,f1063,f140])).
fof(f140,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f90,f90])).
fof(f7764,plain,(
  $false | ($spl17 | $spl19 | $spl45)),
  inference(subsumption_resolution,[],[f7763,f1272])).
fof(f7763,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | ($spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7457,f7439])).
fof(f7457,plain,(
  s(bool,f(s(bool,'F_0'))) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | ($spl17 | $spl45)),
  inference(unit_resulting_resolution,[],[f1272,f1063,f140])).
fof(f7762,plain,(
  $false | ($spl11 | $spl17 | $spl19)),
  inference(subsumption_resolution,[],[f7761,f245])).
fof(f245,plain,(
  s(bool,'T_0') != s(bool,g(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))) | $spl11),
  inference(backward_demodulation,[],[f239,f226])).
fof(f226,plain,(
  s(bool,f(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))) != s(bool,g(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0))))),
  inference(unit_resulting_resolution,[],[f86,f136])).
fof(f136,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK3(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK3(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK3(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK3(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3])],[f55])).
fof(f55,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f49])).
fof(f49,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.pred_set.UNION_EMPTY|split|1',aHL_EXT)).
fof(f86,plain,(
  s(fun(sK0,bool),sK1) != s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))),
  inference(cnf_transformation,[],[f56])).
fof(f56,plain,(
  s(fun(sK0,bool),sK1) != s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f51])).
fof(f51,plain,(
  ? [X0,X1] : s(fun(X0,bool),X1) != s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))),
  inference(ennf_transformation,[],[f26])).
fof(f26,plain,(
  ~! [X0,X1] : s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))),
  inference(flattening,[],[f25])).
fof(f25,plain,(
  ~! [X0] : ! [X1] : s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7] : s(fun(X6,bool),X7) = s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),X7),s(fun(X6,bool),'const.pred_set.EMPTY_0')))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7] : s(fun(X6,bool),X7) = s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),X7),s(fun(X6,bool),'const.pred_set.EMPTY_0')))),
  file('thm.pred_set.UNION_EMPTY|split|1','thm.pred_set.UNION_EMPTY|split|1')).
fof(f239,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))) | $spl11),
  inference(unit_resulting_resolution,[],[f89,f234,f177])).
fof(f7761,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))) | ($spl11 | $spl17 | $spl19)),
  inference(forward_demodulation,[],[f7458,f7439])).
fof(f7458,plain,(
  s(bool,g(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | ($spl11 | $spl17)),
  inference(unit_resulting_resolution,[],[f245,f1063,f140])).
fof(f7760,plain,(
  $false | ($spl15 | $spl17 | $spl19)),
  inference(subsumption_resolution,[],[f7759,f438])).
fof(f438,plain,(
  s(bool,'T_0') != s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))) | $spl15),
  inference(backward_demodulation,[],[f432,f227])).
fof(f227,plain,(
  s(bool,f(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))) != s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0))))),
  inference(unit_resulting_resolution,[],[f86,f136])).
fof(f432,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))) | $spl15),
  inference(unit_resulting_resolution,[],[f89,f427,f177])).
fof(f7759,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))) | ($spl15 | $spl17 | $spl19)),
  inference(forward_demodulation,[],[f7459,f7439])).
fof(f7459,plain,(
  s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | ($spl15 | $spl17)),
  inference(unit_resulting_resolution,[],[f438,f1063,f140])).
fof(f7758,plain,(
  $false | ($spl3 | $spl17 | $spl19)),
  inference(subsumption_resolution,[],[f7757,f151])).
fof(f7757,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl3 | $spl17 | $spl19)),
  inference(forward_demodulation,[],[f7462,f7439])).
fof(f7462,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | ($spl3 | $spl17)),
  inference(unit_resulting_resolution,[],[f151,f1063,f140])).
fof(f7756,plain,(
  $false | ($spl17 | $spl19 | $spl45)),
  inference(subsumption_resolution,[],[f7755,f1272])).
fof(f7755,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | ($spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7463,f7439])).
fof(f7463,plain,(
  s(bool,f(s(bool,'F_0'))) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | ($spl17 | $spl45)),
  inference(unit_resulting_resolution,[],[f1272,f1063,f140])).
fof(f7754,plain,(
  $false | ($spl11 | $spl17 | $spl19)),
  inference(subsumption_resolution,[],[f7753,f245])).
fof(f7753,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))) | ($spl11 | $spl17 | $spl19)),
  inference(forward_demodulation,[],[f7464,f7439])).
fof(f7464,plain,(
  s(bool,g(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | ($spl11 | $spl17)),
  inference(unit_resulting_resolution,[],[f245,f1063,f140])).
fof(f7752,plain,(
  $false | ($spl15 | $spl17 | $spl19)),
  inference(subsumption_resolution,[],[f7751,f438])).
fof(f7751,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))) | ($spl15 | $spl17 | $spl19)),
  inference(forward_demodulation,[],[f7465,f7439])).
fof(f7465,plain,(
  s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | ($spl15 | $spl17)),
  inference(unit_resulting_resolution,[],[f438,f1063,f140])).
fof(f7750,plain,(
  $false | ($spl17 | $spl19)),
  inference(subsumption_resolution,[],[f7749,f87])).
fof(f7749,plain,(
  ~p(s(bool,'T_0')) | ($spl17 | $spl19)),
  inference(forward_demodulation,[],[f7468,f7439])).
fof(f7468,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))))) | $spl17),
  inference(unit_resulting_resolution,[],[f1063,f142])).
fof(f142,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f89,f90])).
fof(f7748,plain,(
  $false | ($spl3 | $spl17 | $spl19)),
  inference(subsumption_resolution,[],[f7747,f151])).
fof(f7747,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl17 | $spl19)),
  inference(forward_demodulation,[],[f7469,f7439])).
fof(f7469,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | $spl17),
  inference(unit_resulting_resolution,[],[f89,f1063,f177])).
fof(f7746,plain,(
  $false | ($spl3 | $spl17 | $spl19)),
  inference(subsumption_resolution,[],[f7745,f151])).
fof(f7745,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl17 | $spl19)),
  inference(forward_demodulation,[],[f7744,f7439])).
fof(f7744,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | $spl17),
  inference(forward_demodulation,[],[f7471,f163])).
fof(f163,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f89,f126,f132])).
fof(f126,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(cnf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0,X1] : ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),
  inference(flattening,[],[f42])).
fof(f42,plain,(
  ! [X0] : ! [X1] : ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X6] : ! [X11] : ~p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),'const.pred_set.EMPTY_0'))))),
  file('thm.pred_set.UNION_EMPTY|split|1','thm.pred_set.NOT_IN_EMPTY')).
fof(f7471,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) = s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))),s(fun(bool,bool),'const.pred_set.EMPTY_0'))) | $spl17),
  inference(unit_resulting_resolution,[],[f1063,f207])).
fof(f207,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f163,f90])).
fof(f7743,plain,(
  $false | ($spl11 | $spl17 | $spl19)),
  inference(subsumption_resolution,[],[f7742,f245])).
fof(f7742,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))) | ($spl11 | $spl17 | $spl19)),
  inference(forward_demodulation,[],[f7472,f7439])).
fof(f7472,plain,(
  s(bool,g(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | ($spl11 | $spl17)),
  inference(unit_resulting_resolution,[],[f1063,f273])).
fof(f273,plain,(
  ( ! [X3] : (s(bool,X3) = s(bool,g(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,X3)) ) | $spl11),
  inference(subsumption_resolution,[],[f266,f89])).
fof(f266,plain,(
  ( ! [X3] : (p(s(bool,'F_0')) | s(bool,X3) = s(bool,g(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,X3)) ) | $spl11),
  inference(superposition,[],[f177,f253])).
fof(f253,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))) | $spl11),
  inference(unit_resulting_resolution,[],[f245,f90])).
fof(f7741,plain,(
  $false | ($spl11 | $spl17 | $spl19)),
  inference(subsumption_resolution,[],[f7740,f239])).
fof(f7740,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))) | ($spl11 | $spl17 | $spl19)),
  inference(forward_demodulation,[],[f7473,f7439])).
fof(f7473,plain,(
  s(bool,f(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | ($spl11 | $spl17)),
  inference(unit_resulting_resolution,[],[f1063,f349])).
fof(f349,plain,(
  ( ! [X13] : (s(bool,X13) != s(bool,f(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,X13)) ) | $spl11),
  inference(subsumption_resolution,[],[f330,f86])).
fof(f330,plain,(
  ( ! [X13] : (s(bool,X13) != s(bool,f(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))) | s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,X13)) ) | $spl11),
  inference(superposition,[],[f136,f273])).
fof(f7739,plain,(
  $false | ($spl15 | $spl17 | $spl19)),
  inference(subsumption_resolution,[],[f7738,f438])).
fof(f7738,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))) | ($spl15 | $spl17 | $spl19)),
  inference(forward_demodulation,[],[f7474,f7439])).
fof(f7474,plain,(
  s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | ($spl15 | $spl17)),
  inference(unit_resulting_resolution,[],[f1063,f496])).
fof(f496,plain,(
  ( ! [X3] : (s(bool,X3) = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))) | s(bool,'T_0') = s(bool,X3)) ) | $spl15),
  inference(subsumption_resolution,[],[f484,f89])).
fof(f484,plain,(
  ( ! [X3] : (p(s(bool,'F_0')) | s(bool,X3) = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))) | s(bool,'T_0') = s(bool,X3)) ) | $spl15),
  inference(superposition,[],[f177,f468])).
fof(f468,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))) | $spl15),
  inference(unit_resulting_resolution,[],[f438,f90])).
fof(f7737,plain,(
  $false | ($spl17 | $spl19 | $spl24)),
  inference(subsumption_resolution,[],[f7736,f4996])).
fof(f7736,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl17 | $spl19 | $spl24)),
  inference(forward_demodulation,[],[f7475,f7439])).
fof(f7475,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | ($spl17 | $spl24)),
  inference(unit_resulting_resolution,[],[f1063,f1070])).
fof(f7735,plain,(
  $false | ($spl17 | $spl19 | $spl29 | $spl34)),
  inference(subsumption_resolution,[],[f7734,f5835])).
fof(f5835,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) | $spl29),
  inference(unit_resulting_resolution,[],[f89,f1164,f177])).
fof(f7734,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) | ($spl17 | $spl19 | $spl34)),
  inference(forward_demodulation,[],[f7476,f7439])).
fof(f7476,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | ($spl17 | $spl34)),
  inference(unit_resulting_resolution,[],[f1063,f1168])).
fof(f7732,plain,(
  $false | ($spl3 | $spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(subsumption_resolution,[],[f7731,f151])).
fof(f7731,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl3 | $spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(forward_demodulation,[],[f7730,f1637])).
fof(f1637,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | $spl51),
  inference(unit_resulting_resolution,[],[f89,f1547,f177])).
fof(f7730,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | ($spl3 | $spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7478,f7439])).
fof(f7478,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))))) | ($spl3 | $spl15 | $spl17 | $spl45)),
  inference(unit_resulting_resolution,[],[f151,f1063,f1506])).
fof(f1506,plain,(
  ( ! [X2,X3] : (s(bool,X3) = s(bool,f(s(bool,X2))) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X3)) ) | ($spl15 | $spl45)),
  inference(duplicate_literal_removal,[],[f1476])).
fof(f1476,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X3) = s(bool,f(s(bool,X2))) | s(bool,'T_0') = s(bool,X3)) ) | ($spl15 | $spl45)),
  inference(superposition,[],[f1410,f140])).
fof(f1410,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl15 | $spl45)),
  inference(superposition,[],[f1352,f90])).
fof(f1352,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | ($spl15 | $spl45)),
  inference(forward_demodulation,[],[f1333,f468])).
fof(f1333,plain,(
  s(bool,f(s(bool,'F_0'))) = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))) | ($spl15 | $spl45)),
  inference(unit_resulting_resolution,[],[f1272,f496])).
fof(f7729,plain,(
  $false | ($spl3 | $spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(subsumption_resolution,[],[f7728,f151])).
fof(f7728,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(forward_demodulation,[],[f7727,f1637])).
fof(f7727,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | ($spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7726,f1352])).
fof(f7726,plain,(
  s(bool,f(s(bool,'T_0'))) = s(bool,f(s(bool,'F_0'))) | ($spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7479,f7439])).
fof(f7479,plain,(
  s(bool,f(s(bool,'F_0'))) = s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))))) | ($spl15 | $spl17 | $spl45)),
  inference(unit_resulting_resolution,[],[f1272,f1063,f1506])).
fof(f7725,plain,(
  $false | ($spl3 | $spl11 | $spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(subsumption_resolution,[],[f7724,f151])).
fof(f7724,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl11 | $spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(forward_demodulation,[],[f7723,f1637])).
fof(f7723,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | ($spl11 | $spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7722,f253])).
fof(f7722,plain,(
  s(bool,f(s(bool,'T_0'))) = s(bool,g(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))) | ($spl11 | $spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7480,f7439])).
fof(f7480,plain,(
  s(bool,g(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))) = s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))))) | ($spl11 | $spl15 | $spl17 | $spl45)),
  inference(unit_resulting_resolution,[],[f245,f1063,f1506])).
fof(f7721,plain,(
  $false | ($spl3 | $spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(subsumption_resolution,[],[f7720,f151])).
fof(f7720,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(forward_demodulation,[],[f7719,f1637])).
fof(f7719,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | ($spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7718,f468])).
fof(f7718,plain,(
  s(bool,f(s(bool,'T_0'))) = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))) | ($spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7481,f7439])).
fof(f7481,plain,(
  s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))) = s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))))) | ($spl15 | $spl17 | $spl45)),
  inference(unit_resulting_resolution,[],[f438,f1063,f1506])).
fof(f7716,plain,(
  $false | ($spl3 | $spl15 | $spl17 | $spl19 | $spl45)),
  inference(subsumption_resolution,[],[f7715,f1272])).
fof(f7715,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | ($spl3 | $spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7483,f7439])).
fof(f7483,plain,(
  s(bool,f(s(bool,'F_0'))) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | ($spl3 | $spl15 | $spl17 | $spl45)),
  inference(unit_resulting_resolution,[],[f151,f1063,f1506])).
fof(f7714,plain,(
  $false | ($spl15 | $spl17 | $spl19 | $spl45)),
  inference(subsumption_resolution,[],[f7713,f1272])).
fof(f7713,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | ($spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7712,f1352])).
fof(f7712,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,f(s(bool,'F_0'))))) | ($spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7484,f7439])).
fof(f7484,plain,(
  s(bool,f(s(bool,f(s(bool,'F_0'))))) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | ($spl15 | $spl17 | $spl45)),
  inference(unit_resulting_resolution,[],[f1272,f1063,f1506])).
fof(f7711,plain,(
  $false | ($spl11 | $spl15 | $spl17 | $spl19 | $spl45)),
  inference(subsumption_resolution,[],[f7710,f1272])).
fof(f7710,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | ($spl11 | $spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7709,f253])).
fof(f7709,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,g(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))))) | ($spl11 | $spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7485,f7439])).
fof(f7485,plain,(
  s(bool,f(s(bool,g(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))))) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | ($spl11 | $spl15 | $spl17 | $spl45)),
  inference(unit_resulting_resolution,[],[f245,f1063,f1506])).
fof(f7708,plain,(
  $false | ($spl15 | $spl17 | $spl19 | $spl45)),
  inference(subsumption_resolution,[],[f7707,f1272])).
fof(f7707,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | ($spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7706,f468])).
fof(f7706,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))))) | ($spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7486,f7439])).
fof(f7486,plain,(
  s(bool,f(s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))))) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | ($spl15 | $spl17 | $spl45)),
  inference(unit_resulting_resolution,[],[f438,f1063,f1506])).
fof(f7704,plain,(
  $false | ($spl3 | $spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(subsumption_resolution,[],[f7703,f151])).
fof(f7703,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(forward_demodulation,[],[f7702,f1637])).
fof(f7702,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | ($spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7488,f7439])).
fof(f7488,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))))) | ($spl15 | $spl17 | $spl45)),
  inference(unit_resulting_resolution,[],[f89,f1063,f1521])).
fof(f1521,plain,(
  ( ! [X2,X1] : (p(s(bool,X2)) | s(bool,X2) = s(bool,f(s(bool,X1))) | s(bool,'T_0') = s(bool,X1)) ) | ($spl15 | $spl45)),
  inference(subsumption_resolution,[],[f1484,f142])).
fof(f1484,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X2) = s(bool,f(s(bool,X1))) | s(bool,'T_0') = s(bool,X1)) ) | ($spl15 | $spl45)),
  inference(superposition,[],[f132,f1410])).
fof(f7701,plain,(
  $false | ($spl3 | $spl6 | $spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(subsumption_resolution,[],[f7700,f151])).
fof(f7700,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl6 | $spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(forward_demodulation,[],[f7699,f1637])).
fof(f7699,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | ($spl6 | $spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7489,f7439])).
fof(f7489,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))))) | ($spl6 | $spl15 | $spl17 | $spl45)),
  inference(unit_resulting_resolution,[],[f1063,f1522])).
fof(f1522,plain,(
  ( ! [X16] : (s(bool,'F_0') = s(bool,f(s(bool,X16))) | s(bool,'T_0') = s(bool,X16)) ) | ($spl6 | $spl15 | $spl45)),
  inference(subsumption_resolution,[],[f1492,f142])).
fof(f1492,plain,(
  ( ! [X16] : (p(s(bool,X16)) | s(bool,'F_0') = s(bool,f(s(bool,X16))) | s(bool,'T_0') = s(bool,X16)) ) | ($spl6 | $spl15 | $spl45)),
  inference(superposition,[],[f176,f1410])).
fof(f7698,plain,(
  $false | ($spl3 | $spl6 | $spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(subsumption_resolution,[],[f7697,f151])).
fof(f7697,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl3 | $spl6 | $spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(forward_demodulation,[],[f7696,f1637])).
fof(f7696,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | ($spl3 | $spl6 | $spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7695,f1352])).
fof(f7695,plain,(
  s(bool,f(s(bool,'T_0'))) = s(bool,f(s(bool,'F_0'))) | ($spl3 | $spl6 | $spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7490,f7439])).
fof(f7490,plain,(
  s(bool,f(s(bool,'F_0'))) = s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))))) | ($spl3 | $spl6 | $spl15 | $spl17 | $spl45)),
  inference(unit_resulting_resolution,[],[f151,f1063,f1795])).
fof(f1795,plain,(
  ( ! [X26,X27] : (s(bool,f(s(bool,X26))) = s(bool,f(s(bool,X27))) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,X26)) ) | ($spl6 | $spl15 | $spl45)),
  inference(subsumption_resolution,[],[f1776,f89])).
fof(f1776,plain,(
  ( ! [X26,X27] : (p(s(bool,'F_0')) | s(bool,f(s(bool,X26))) = s(bool,f(s(bool,X27))) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,X26)) ) | ($spl6 | $spl15 | $spl45)),
  inference(superposition,[],[f1521,f1522])).
fof(f7694,plain,(
  $false | ($spl3 | $spl6 | $spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(subsumption_resolution,[],[f7693,f151])).
fof(f7693,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl6 | $spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(forward_demodulation,[],[f7692,f1637])).
fof(f7692,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | ($spl6 | $spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7691,f1352])).
fof(f7691,plain,(
  s(bool,f(s(bool,'T_0'))) = s(bool,f(s(bool,'F_0'))) | ($spl6 | $spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7690,f1352])).
fof(f7690,plain,(
  s(bool,f(s(bool,'T_0'))) = s(bool,f(s(bool,f(s(bool,'F_0'))))) | ($spl6 | $spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7491,f7439])).
fof(f7491,plain,(
  s(bool,f(s(bool,f(s(bool,'F_0'))))) = s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))))) | ($spl6 | $spl15 | $spl17 | $spl45)),
  inference(unit_resulting_resolution,[],[f1272,f1063,f1795])).
fof(f7689,plain,(
  $false | ($spl3 | $spl6 | $spl11 | $spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(subsumption_resolution,[],[f7688,f151])).
fof(f7688,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl6 | $spl11 | $spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(forward_demodulation,[],[f7687,f1637])).
fof(f7687,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | ($spl6 | $spl11 | $spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7686,f1352])).
fof(f7686,plain,(
  s(bool,f(s(bool,'T_0'))) = s(bool,f(s(bool,'F_0'))) | ($spl6 | $spl11 | $spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7685,f253])).
fof(f7685,plain,(
  s(bool,f(s(bool,'T_0'))) = s(bool,f(s(bool,g(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))))) | ($spl6 | $spl11 | $spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7492,f7439])).
fof(f7492,plain,(
  s(bool,f(s(bool,g(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))))) = s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))))) | ($spl6 | $spl11 | $spl15 | $spl17 | $spl45)),
  inference(unit_resulting_resolution,[],[f245,f1063,f1795])).
fof(f7684,plain,(
  $false | ($spl3 | $spl6 | $spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(subsumption_resolution,[],[f7683,f151])).
fof(f7683,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl6 | $spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(forward_demodulation,[],[f7682,f1637])).
fof(f7682,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | ($spl6 | $spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7681,f1352])).
fof(f7681,plain,(
  s(bool,f(s(bool,'T_0'))) = s(bool,f(s(bool,'F_0'))) | ($spl6 | $spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7680,f468])).
fof(f7680,plain,(
  s(bool,f(s(bool,'T_0'))) = s(bool,f(s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))))) | ($spl6 | $spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7493,f7439])).
fof(f7493,plain,(
  s(bool,f(s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))))) = s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))))) | ($spl6 | $spl15 | $spl17 | $spl45)),
  inference(unit_resulting_resolution,[],[f438,f1063,f1795])).
fof(f7679,plain,(
  $false | ($spl3 | $spl6 | $spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(subsumption_resolution,[],[f7678,f151])).
fof(f7678,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl3 | $spl6 | $spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(forward_demodulation,[],[f7677,f1637])).
fof(f7677,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | ($spl3 | $spl6 | $spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7676,f1352])).
fof(f7676,plain,(
  s(bool,f(s(bool,'T_0'))) = s(bool,f(s(bool,'F_0'))) | ($spl3 | $spl6 | $spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7495,f7439])).
fof(f7495,plain,(
  s(bool,f(s(bool,'F_0'))) = s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))))) | ($spl3 | $spl6 | $spl15 | $spl17 | $spl45)),
  inference(unit_resulting_resolution,[],[f151,f1063,f1795])).
fof(f7675,plain,(
  $false | ($spl3 | $spl6 | $spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(subsumption_resolution,[],[f7674,f151])).
fof(f7674,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl6 | $spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(forward_demodulation,[],[f7673,f1637])).
fof(f7673,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | ($spl6 | $spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7672,f1352])).
fof(f7672,plain,(
  s(bool,f(s(bool,'T_0'))) = s(bool,f(s(bool,'F_0'))) | ($spl6 | $spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7671,f1352])).
fof(f7671,plain,(
  s(bool,f(s(bool,'T_0'))) = s(bool,f(s(bool,f(s(bool,'F_0'))))) | ($spl6 | $spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7496,f7439])).
fof(f7496,plain,(
  s(bool,f(s(bool,f(s(bool,'F_0'))))) = s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))))) | ($spl6 | $spl15 | $spl17 | $spl45)),
  inference(unit_resulting_resolution,[],[f1272,f1063,f1795])).
fof(f7670,plain,(
  $false | ($spl3 | $spl6 | $spl11 | $spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(subsumption_resolution,[],[f7669,f151])).
fof(f7669,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl6 | $spl11 | $spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(forward_demodulation,[],[f7668,f1637])).
fof(f7668,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | ($spl6 | $spl11 | $spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7667,f1352])).
fof(f7667,plain,(
  s(bool,f(s(bool,'T_0'))) = s(bool,f(s(bool,'F_0'))) | ($spl6 | $spl11 | $spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7666,f253])).
fof(f7666,plain,(
  s(bool,f(s(bool,'T_0'))) = s(bool,f(s(bool,g(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))))) | ($spl6 | $spl11 | $spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7497,f7439])).
fof(f7497,plain,(
  s(bool,f(s(bool,g(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))))) = s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))))) | ($spl6 | $spl11 | $spl15 | $spl17 | $spl45)),
  inference(unit_resulting_resolution,[],[f245,f1063,f1795])).
fof(f7665,plain,(
  $false | ($spl3 | $spl6 | $spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(subsumption_resolution,[],[f7664,f151])).
fof(f7664,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl6 | $spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(forward_demodulation,[],[f7663,f1637])).
fof(f7663,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | ($spl6 | $spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7662,f1352])).
fof(f7662,plain,(
  s(bool,f(s(bool,'T_0'))) = s(bool,f(s(bool,'F_0'))) | ($spl6 | $spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7661,f468])).
fof(f7661,plain,(
  s(bool,f(s(bool,'T_0'))) = s(bool,f(s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))))) | ($spl6 | $spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7498,f7439])).
fof(f7498,plain,(
  s(bool,f(s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))))) = s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))))) | ($spl6 | $spl15 | $spl17 | $spl45)),
  inference(unit_resulting_resolution,[],[f438,f1063,f1795])).
fof(f7660,plain,(
  $false | ($spl3 | $spl6 | $spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(subsumption_resolution,[],[f7659,f151])).
fof(f7659,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl6 | $spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(forward_demodulation,[],[f7658,f1637])).
fof(f7658,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | ($spl6 | $spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(forward_demodulation,[],[f7657,f1637])).
fof(f7657,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,'T_0'))))) | ($spl6 | $spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7500,f7439])).
fof(f7500,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))))))) | ($spl6 | $spl15 | $spl17 | $spl45)),
  inference(unit_resulting_resolution,[],[f1063,f2031])).
fof(f2031,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,X0))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl6 | $spl15 | $spl45)),
  inference(condensation,[],[f2030])).
fof(f2030,plain,(
  ( ! [X45,X46] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,X46))))) | s(bool,'T_0') = s(bool,X45) | s(bool,'T_0') = s(bool,X46)) ) | ($spl6 | $spl15 | $spl45)),
  inference(duplicate_literal_removal,[],[f1947])).
fof(f1947,plain,(
  ( ! [X45,X46] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,X46))))) | s(bool,'T_0') = s(bool,X45) | s(bool,'T_0') = s(bool,X46) | s(bool,'T_0') = s(bool,X45)) ) | ($spl6 | $spl15 | $spl45)),
  inference(superposition,[],[f1522,f1506])).
fof(f7654,plain,(
  $false | ($spl3 | $spl15 | $spl17 | $spl19 | $spl45)),
  inference(subsumption_resolution,[],[f7653,f1272])).
fof(f7653,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | ($spl3 | $spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7652,f1352])).
fof(f7652,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,f(s(bool,'F_0'))))) | ($spl3 | $spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7502,f7439])).
fof(f7502,plain,(
  s(bool,f(s(bool,f(s(bool,'F_0'))))) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | ($spl3 | $spl15 | $spl17 | $spl45)),
  inference(unit_resulting_resolution,[],[f151,f1063,f2060])).
fof(f2060,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,f(s(bool,f(s(bool,X1))))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1)) ) | ($spl15 | $spl45)),
  inference(condensation,[],[f2059])).
fof(f2059,plain,(
  ( ! [X19,X20,X18] : (s(bool,X20) = s(bool,f(s(bool,f(s(bool,X19))))) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,X20) | s(bool,'T_0') = s(bool,X19)) ) | ($spl15 | $spl45)),
  inference(duplicate_literal_removal,[],[f1898])).
fof(f1898,plain,(
  ( ! [X19,X20,X18] : (s(bool,X20) = s(bool,f(s(bool,f(s(bool,X19))))) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,X20) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,X18)) ) | ($spl15 | $spl45)),
  inference(superposition,[],[f1506,f1506])).
fof(f7651,plain,(
  $false | ($spl15 | $spl17 | $spl19 | $spl45)),
  inference(subsumption_resolution,[],[f7650,f1272])).
fof(f7650,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | ($spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7649,f1352])).
fof(f7649,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,f(s(bool,'F_0'))))) | ($spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7648,f1352])).
fof(f7648,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,'F_0'))))))) | ($spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7503,f7439])).
fof(f7503,plain,(
  s(bool,f(s(bool,f(s(bool,f(s(bool,'F_0'))))))) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | ($spl15 | $spl17 | $spl45)),
  inference(unit_resulting_resolution,[],[f1272,f1063,f2060])).
fof(f7647,plain,(
  $false | ($spl11 | $spl15 | $spl17 | $spl19 | $spl45)),
  inference(subsumption_resolution,[],[f7646,f1272])).
fof(f7646,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | ($spl11 | $spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7645,f1352])).
fof(f7645,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,f(s(bool,'F_0'))))) | ($spl11 | $spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7644,f253])).
fof(f7644,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,f(s(bool,g(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))))))) | ($spl11 | $spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7504,f7439])).
fof(f7504,plain,(
  s(bool,f(s(bool,f(s(bool,g(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))))))) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | ($spl11 | $spl15 | $spl17 | $spl45)),
  inference(unit_resulting_resolution,[],[f245,f1063,f2060])).
fof(f7643,plain,(
  $false | ($spl15 | $spl17 | $spl19 | $spl45)),
  inference(subsumption_resolution,[],[f7642,f1272])).
fof(f7642,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | ($spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7641,f1352])).
fof(f7641,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,f(s(bool,'F_0'))))) | ($spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7640,f468])).
fof(f7640,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,f(s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))))))) | ($spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7505,f7439])).
fof(f7505,plain,(
  s(bool,f(s(bool,f(s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))))))) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | ($spl15 | $spl17 | $spl45)),
  inference(unit_resulting_resolution,[],[f438,f1063,f2060])).
fof(f7637,plain,(
  $false | ($spl3 | $spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(subsumption_resolution,[],[f7636,f151])).
fof(f7636,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl3 | $spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(forward_demodulation,[],[f7635,f1637])).
fof(f7635,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | ($spl3 | $spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(forward_demodulation,[],[f7634,f1637])).
fof(f7634,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,'T_0'))))) | ($spl3 | $spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7507,f7439])).
fof(f7507,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))))))) | ($spl3 | $spl15 | $spl17 | $spl45)),
  inference(unit_resulting_resolution,[],[f151,f1063,f2060])).
fof(f7633,plain,(
  $false | ($spl3 | $spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(subsumption_resolution,[],[f7632,f151])).
fof(f7632,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(forward_demodulation,[],[f7631,f1637])).
fof(f7631,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | ($spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(forward_demodulation,[],[f7630,f1352])).
fof(f7630,plain,(
  s(bool,f(s(bool,'T_0'))) = s(bool,f(s(bool,'F_0'))) | ($spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(forward_demodulation,[],[f7629,f1637])).
fof(f7629,plain,(
  s(bool,f(s(bool,'F_0'))) = s(bool,f(s(bool,f(s(bool,'T_0'))))) | ($spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7508,f7439])).
fof(f7508,plain,(
  s(bool,f(s(bool,'F_0'))) = s(bool,f(s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))))))) | ($spl15 | $spl17 | $spl45)),
  inference(unit_resulting_resolution,[],[f1272,f1063,f2060])).
fof(f7628,plain,(
  $false | ($spl3 | $spl11 | $spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(subsumption_resolution,[],[f7627,f151])).
fof(f7627,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl11 | $spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(forward_demodulation,[],[f7626,f1637])).
fof(f7626,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | ($spl11 | $spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(forward_demodulation,[],[f7625,f1637])).
fof(f7625,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,'T_0'))))) | ($spl11 | $spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7624,f253])).
fof(f7624,plain,(
  s(bool,f(s(bool,f(s(bool,'T_0'))))) = s(bool,g(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))) | ($spl11 | $spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7509,f7439])).
fof(f7509,plain,(
  s(bool,g(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))) = s(bool,f(s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))))))) | ($spl11 | $spl15 | $spl17 | $spl45)),
  inference(unit_resulting_resolution,[],[f245,f1063,f2060])).
fof(f7623,plain,(
  $false | ($spl3 | $spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(subsumption_resolution,[],[f7622,f151])).
fof(f7622,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(forward_demodulation,[],[f7621,f1637])).
fof(f7621,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | ($spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(forward_demodulation,[],[f7620,f1637])).
fof(f7620,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,'T_0'))))) | ($spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7619,f468])).
fof(f7619,plain,(
  s(bool,f(s(bool,f(s(bool,'T_0'))))) = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))) | ($spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7510,f7439])).
fof(f7510,plain,(
  s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))) = s(bool,f(s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))))))) | ($spl15 | $spl17 | $spl45)),
  inference(unit_resulting_resolution,[],[f438,f1063,f2060])).
fof(f7616,plain,(
  $false | ($spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(subsumption_resolution,[],[f7615,f87])).
fof(f7615,plain,(
  ~p(s(bool,'T_0')) | ($spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(forward_demodulation,[],[f7614,f1637])).
fof(f7614,plain,(
  ~p(s(bool,f(s(bool,'T_0')))) | ($spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7512,f7439])).
fof(f7512,plain,(
  ~p(s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))))))) | ($spl15 | $spl17 | $spl45)),
  inference(unit_resulting_resolution,[],[f1063,f2143])).
fof(f2143,plain,(
  ( ! [X1] : (~p(s(bool,f(s(bool,X1)))) | s(bool,'T_0') = s(bool,X1)) ) | ($spl15 | $spl45)),
  inference(duplicate_literal_removal,[],[f2142])).
fof(f2142,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,X1) | ~p(s(bool,f(s(bool,X1)))) | s(bool,'T_0') = s(bool,X1)) ) | ($spl15 | $spl45)),
  inference(condensation,[],[f2141])).
fof(f2141,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,f(s(bool,X9)))) | s(bool,X8) = s(bool,X10) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,X8)) ) | ($spl15 | $spl45)),
  inference(subsumption_resolution,[],[f1933,f177])).
fof(f1933,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,f(s(bool,X9)))) | ~p(s(bool,X10)) | s(bool,X8) = s(bool,X10) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,X8)) ) | ($spl15 | $spl45)),
  inference(superposition,[],[f135,f1506])).
fof(f135,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f54])).
fof(f7613,plain,(
  $false | ($spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(subsumption_resolution,[],[f7612,f87])).
fof(f7612,plain,(
  ~p(s(bool,'T_0')) | ($spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(forward_demodulation,[],[f7611,f1637])).
fof(f7611,plain,(
  ~p(s(bool,f(s(bool,'T_0')))) | ($spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(forward_demodulation,[],[f7610,f1637])).
fof(f7610,plain,(
  ~p(s(bool,f(s(bool,f(s(bool,'T_0')))))) | ($spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7513,f7439])).
fof(f7513,plain,(
  ~p(s(bool,f(s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))))))))) | ($spl15 | $spl17 | $spl45)),
  inference(unit_resulting_resolution,[],[f1063,f2214])).
fof(f2214,plain,(
  ( ! [X0] : (~p(s(bool,f(s(bool,f(s(bool,X0)))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl15 | $spl45)),
  inference(condensation,[],[f2213])).
fof(f2213,plain,(
  ( ! [X12,X11] : (~p(s(bool,f(s(bool,f(s(bool,X12)))))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X12)) ) | ($spl15 | $spl45)),
  inference(duplicate_literal_removal,[],[f2177])).
fof(f2177,plain,(
  ( ! [X12,X11] : (~p(s(bool,f(s(bool,f(s(bool,X12)))))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X12) | s(bool,'T_0') = s(bool,X11)) ) | ($spl15 | $spl45)),
  inference(superposition,[],[f2143,f1506])).
fof(f7609,plain,(
  $false | ($spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(subsumption_resolution,[],[f7608,f87])).
fof(f7608,plain,(
  ~p(s(bool,'T_0')) | ($spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(forward_demodulation,[],[f7607,f1637])).
fof(f7607,plain,(
  ~p(s(bool,f(s(bool,'T_0')))) | ($spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(forward_demodulation,[],[f7606,f1637])).
fof(f7606,plain,(
  ~p(s(bool,f(s(bool,f(s(bool,'T_0')))))) | ($spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(forward_demodulation,[],[f7605,f1637])).
fof(f7605,plain,(
  ~p(s(bool,f(s(bool,f(s(bool,f(s(bool,'T_0')))))))) | ($spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7514,f7439])).
fof(f7514,plain,(
  ~p(s(bool,f(s(bool,f(s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))))))))))) | ($spl15 | $spl17 | $spl45)),
  inference(unit_resulting_resolution,[],[f1063,f2332])).
fof(f2332,plain,(
  ( ! [X0] : (~p(s(bool,f(s(bool,f(s(bool,f(s(bool,X0)))))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl15 | $spl45)),
  inference(condensation,[],[f2331])).
fof(f2331,plain,(
  ( ! [X12,X11] : (~p(s(bool,f(s(bool,f(s(bool,f(s(bool,X12)))))))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X12)) ) | ($spl15 | $spl45)),
  inference(duplicate_literal_removal,[],[f2284])).
fof(f2284,plain,(
  ( ! [X12,X11] : (~p(s(bool,f(s(bool,f(s(bool,f(s(bool,X12)))))))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X12) | s(bool,'T_0') = s(bool,X11)) ) | ($spl15 | $spl45)),
  inference(superposition,[],[f2214,f1506])).
fof(f7604,plain,(
  $false | ($spl3 | $spl6 | $spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(subsumption_resolution,[],[f7603,f151])).
fof(f7603,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl6 | $spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(forward_demodulation,[],[f7602,f1637])).
fof(f7602,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | ($spl6 | $spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(forward_demodulation,[],[f7601,f1637])).
fof(f7601,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,'T_0'))))) | ($spl6 | $spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(forward_demodulation,[],[f7600,f1637])).
fof(f7600,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,'T_0'))))))) | ($spl6 | $spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7515,f7439])).
fof(f7515,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))))))))) | ($spl6 | $spl15 | $spl17 | $spl45)),
  inference(unit_resulting_resolution,[],[f1063,f2487])).
fof(f2487,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,X0))))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl6 | $spl15 | $spl45)),
  inference(condensation,[],[f2486])).
fof(f2486,plain,(
  ( ! [X12,X11] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,X12))))))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X12)) ) | ($spl6 | $spl15 | $spl45)),
  inference(duplicate_literal_removal,[],[f2406])).
fof(f2406,plain,(
  ( ! [X12,X11] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,X12))))))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X12) | s(bool,'T_0') = s(bool,X11)) ) | ($spl6 | $spl15 | $spl45)),
  inference(superposition,[],[f2031,f1506])).
fof(f7597,plain,(
  $false | ($spl3 | $spl6 | $spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(subsumption_resolution,[],[f7596,f151])).
fof(f7596,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl6 | $spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(forward_demodulation,[],[f7595,f1637])).
fof(f7595,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | ($spl6 | $spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(forward_demodulation,[],[f7594,f1637])).
fof(f7594,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,'T_0'))))) | ($spl6 | $spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7517,f7439])).
fof(f7517,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))))))) | ($spl6 | $spl15 | $spl17 | $spl45)),
  inference(unit_resulting_resolution,[],[f89,f1063,f2530])).
fof(f2530,plain,(
  ( ! [X8,X9] : (p(s(bool,X9)) | s(bool,X9) = s(bool,f(s(bool,f(s(bool,X8))))) | s(bool,'T_0') = s(bool,X8)) ) | ($spl6 | $spl15 | $spl45)),
  inference(subsumption_resolution,[],[f2454,f89])).
fof(f2454,plain,(
  ( ! [X8,X9] : (p(s(bool,'F_0')) | p(s(bool,X9)) | s(bool,X9) = s(bool,f(s(bool,f(s(bool,X8))))) | s(bool,'T_0') = s(bool,X8)) ) | ($spl6 | $spl15 | $spl45)),
  inference(superposition,[],[f132,f2031])).
fof(f7590,plain,(
  $false | ($spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(subsumption_resolution,[],[f7589,f87])).
fof(f7589,plain,(
  ~p(s(bool,'T_0')) | ($spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(forward_demodulation,[],[f7588,f1637])).
fof(f7588,plain,(
  ~p(s(bool,f(s(bool,'T_0')))) | ($spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(forward_demodulation,[],[f7587,f1637])).
fof(f7587,plain,(
  ~p(s(bool,f(s(bool,f(s(bool,'T_0')))))) | ($spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(forward_demodulation,[],[f7586,f1637])).
fof(f7586,plain,(
  ~p(s(bool,f(s(bool,f(s(bool,f(s(bool,'T_0')))))))) | ($spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(forward_demodulation,[],[f7585,f1637])).
fof(f7585,plain,(
  ~p(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,'T_0')))))))))) | ($spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7519,f7439])).
fof(f7519,plain,(
  ~p(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))))))))))))) | ($spl15 | $spl17 | $spl45)),
  inference(unit_resulting_resolution,[],[f1063,f2821])).
fof(f2821,plain,(
  ( ! [X0] : (~p(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X0)))))))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl15 | $spl45)),
  inference(condensation,[],[f2820])).
fof(f2820,plain,(
  ( ! [X12,X11] : (~p(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X12)))))))))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X12)) ) | ($spl15 | $spl45)),
  inference(duplicate_literal_removal,[],[f2752])).
fof(f2752,plain,(
  ( ! [X12,X11] : (~p(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X12)))))))))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X12) | s(bool,'T_0') = s(bool,X11)) ) | ($spl15 | $spl45)),
  inference(superposition,[],[f2332,f1506])).
fof(f7584,plain,(
  $false | ($spl3 | $spl6 | $spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(subsumption_resolution,[],[f7583,f151])).
fof(f7583,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl6 | $spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(forward_demodulation,[],[f7582,f1637])).
fof(f7582,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | ($spl6 | $spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(forward_demodulation,[],[f7581,f1637])).
fof(f7581,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,'T_0'))))) | ($spl6 | $spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(forward_demodulation,[],[f7580,f1637])).
fof(f7580,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,'T_0'))))))) | ($spl6 | $spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(forward_demodulation,[],[f7579,f1637])).
fof(f7579,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,'T_0'))))))))) | ($spl6 | $spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7520,f7439])).
fof(f7520,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))))))))))) | ($spl6 | $spl15 | $spl17 | $spl45)),
  inference(unit_resulting_resolution,[],[f1063,f3278])).
fof(f3278,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X0))))))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl6 | $spl15 | $spl45)),
  inference(condensation,[],[f3277])).
fof(f3277,plain,(
  ( ! [X12,X11] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X12))))))))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X12)) ) | ($spl6 | $spl15 | $spl45)),
  inference(duplicate_literal_removal,[],[f3158])).
fof(f3158,plain,(
  ( ! [X12,X11] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X12))))))))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X12) | s(bool,'T_0') = s(bool,X11)) ) | ($spl6 | $spl15 | $spl45)),
  inference(superposition,[],[f2487,f1506])).
fof(f7575,plain,(
  $false | ($spl3 | $spl6 | $spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(subsumption_resolution,[],[f7574,f151])).
fof(f7574,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl6 | $spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(forward_demodulation,[],[f7573,f1637])).
fof(f7573,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | ($spl6 | $spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(forward_demodulation,[],[f7572,f1637])).
fof(f7572,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,'T_0'))))) | ($spl6 | $spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(forward_demodulation,[],[f7571,f1637])).
fof(f7571,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,'T_0'))))))) | ($spl6 | $spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7522,f7439])).
fof(f7522,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))))))))) | ($spl6 | $spl15 | $spl17 | $spl45)),
  inference(unit_resulting_resolution,[],[f89,f1063,f3346])).
fof(f3346,plain,(
  ( ! [X17,X18] : (p(s(bool,X18)) | s(bool,X18) = s(bool,f(s(bool,f(s(bool,f(s(bool,X17))))))) | s(bool,'T_0') = s(bool,X17)) ) | ($spl6 | $spl15 | $spl45)),
  inference(subsumption_resolution,[],[f3237,f89])).
fof(f3237,plain,(
  ( ! [X17,X18] : (p(s(bool,'F_0')) | p(s(bool,X18)) | s(bool,X18) = s(bool,f(s(bool,f(s(bool,f(s(bool,X17))))))) | s(bool,'T_0') = s(bool,X17)) ) | ($spl6 | $spl15 | $spl45)),
  inference(superposition,[],[f132,f2487])).
fof(f7570,plain,(
  $false | ($spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(subsumption_resolution,[],[f7569,f87])).
fof(f7569,plain,(
  ~p(s(bool,'T_0')) | ($spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(forward_demodulation,[],[f7568,f1637])).
fof(f7568,plain,(
  ~p(s(bool,f(s(bool,'T_0')))) | ($spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(forward_demodulation,[],[f7567,f1637])).
fof(f7567,plain,(
  ~p(s(bool,f(s(bool,f(s(bool,'T_0')))))) | ($spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(forward_demodulation,[],[f7566,f1637])).
fof(f7566,plain,(
  ~p(s(bool,f(s(bool,f(s(bool,f(s(bool,'T_0')))))))) | ($spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(forward_demodulation,[],[f7565,f1637])).
fof(f7565,plain,(
  ~p(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,'T_0')))))))))) | ($spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(forward_demodulation,[],[f7564,f1637])).
fof(f7564,plain,(
  ~p(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,'T_0')))))))))))) | ($spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7523,f7439])).
fof(f7523,plain,(
  ~p(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))))))))))))))) | ($spl15 | $spl17 | $spl45)),
  inference(unit_resulting_resolution,[],[f1063,f3748])).
fof(f3748,plain,(
  ( ! [X0] : (~p(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X0)))))))))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl15 | $spl45)),
  inference(condensation,[],[f3747])).
fof(f3747,plain,(
  ( ! [X12,X11] : (~p(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X12)))))))))))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X12)) ) | ($spl15 | $spl45)),
  inference(duplicate_literal_removal,[],[f3648])).
fof(f3648,plain,(
  ( ! [X12,X11] : (~p(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X12)))))))))))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X12) | s(bool,'T_0') = s(bool,X11)) ) | ($spl15 | $spl45)),
  inference(superposition,[],[f2821,f1506])).
fof(f7559,plain,(
  $false | ($spl11 | $spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(subsumption_resolution,[],[f7558,f1637])).
fof(f7558,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl11 | $spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(forward_demodulation,[],[f7557,f1637])).
fof(f7557,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,'T_0'))))) | ($spl11 | $spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7525,f7439])).
fof(f7525,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))))))) | ($spl11 | $spl15 | $spl17 | $spl45)),
  inference(unit_resulting_resolution,[],[f1063,f6384])).
fof(f6384,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,X0))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl11 | $spl15 | $spl45)),
  inference(forward_demodulation,[],[f6262,f239])).
fof(f6262,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,f(s(bool,f(s(bool,X0))))) != s(bool,f(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0))))) ) | ($spl11 | $spl15 | $spl45)),
  inference(condensation,[],[f6261])).
fof(f6261,plain,(
  ( ! [X39,X38] : (s(bool,f(s(bool,f(s(bool,X39))))) != s(bool,f(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,X38) | s(bool,'T_0') = s(bool,X39)) ) | ($spl11 | $spl15 | $spl45)),
  inference(duplicate_literal_removal,[],[f6005])).
fof(f6005,plain,(
  ( ! [X39,X38] : (s(bool,f(s(bool,f(s(bool,X39))))) != s(bool,f(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,X38) | s(bool,'T_0') = s(bool,X38) | s(bool,'T_0') = s(bool,X39)) ) | ($spl11 | $spl15 | $spl45)),
  inference(superposition,[],[f349,f2060])).
fof(f7556,plain,(
  $false | ($spl11 | $spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(subsumption_resolution,[],[f7555,f1637])).
fof(f7555,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl11 | $spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(forward_demodulation,[],[f7554,f1637])).
fof(f7554,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,'T_0'))))) | ($spl11 | $spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(forward_demodulation,[],[f7553,f1637])).
fof(f7553,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,'T_0'))))))) | ($spl11 | $spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(forward_demodulation,[],[f7552,f1637])).
fof(f7552,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,'T_0'))))))))) | ($spl11 | $spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7526,f7439])).
fof(f7526,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))))))))))) | ($spl11 | $spl15 | $spl17 | $spl45)),
  inference(unit_resulting_resolution,[],[f1063,f6503])).
fof(f6503,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X0))))))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl11 | $spl15 | $spl45)),
  inference(condensation,[],[f6502])).
fof(f6502,plain,(
  ( ! [X14,X13] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X14))))))))) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X14)) ) | ($spl11 | $spl15 | $spl45)),
  inference(duplicate_literal_removal,[],[f6411])).
fof(f6411,plain,(
  ( ! [X14,X13] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X14))))))))) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X14)) ) | ($spl11 | $spl15 | $spl45)),
  inference(superposition,[],[f6384,f2060])).
fof(f7551,plain,(
  $false | ($spl11 | $spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(subsumption_resolution,[],[f7550,f1637])).
fof(f7550,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl11 | $spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(forward_demodulation,[],[f7549,f1637])).
fof(f7549,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,'T_0'))))) | ($spl11 | $spl15 | $spl17 | $spl19 | $spl45 | $spl51)),
  inference(forward_demodulation,[],[f7548,f1637])).
fof(f7548,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,'T_0'))))))) | ($spl11 | $spl15 | $spl17 | $spl19 | $spl45)),
  inference(forward_demodulation,[],[f7527,f7439])).
fof(f7527,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))))))))) | ($spl11 | $spl15 | $spl17 | $spl45)),
  inference(unit_resulting_resolution,[],[f1063,f6505])).
fof(f6505,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,X0))))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl11 | $spl15 | $spl45)),
  inference(condensation,[],[f6504])).
fof(f6504,plain,(
  ( ! [X12,X11] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,X12))))))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X12)) ) | ($spl11 | $spl15 | $spl45)),
  inference(duplicate_literal_removal,[],[f6410])).
fof(f6410,plain,(
  ( ! [X12,X11] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,X12))))))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X12) | s(bool,'T_0') = s(bool,X11)) ) | ($spl11 | $spl15 | $spl45)),
  inference(superposition,[],[f6384,f1506])).
fof(f7547,plain,(
  $false | ($spl17 | $spl19)),
  inference(subsumption_resolution,[],[f7546,f87])).
fof(f7546,plain,(
  ~p(s(bool,'T_0')) | ($spl17 | $spl19)),
  inference(forward_demodulation,[],[f7529,f7439])).
fof(f7529,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))))) | $spl17),
  inference(unit_resulting_resolution,[],[f87,f1063,f135])).
fof(f7545,plain,(
  $false | ($spl17 | $spl19)),
  inference(subsumption_resolution,[],[f7544,f87])).
fof(f7544,plain,(
  ~p(s(bool,'T_0')) | ($spl17 | $spl19)),
  inference(forward_demodulation,[],[f7532,f7439])).
fof(f7532,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))))) | $spl17),
  inference(unit_resulting_resolution,[],[f87,f1063,f135])).
fof(f4145,plain,(
  $false | ($spl15 | $spl26 | $spl45)),
  inference(subsumption_resolution,[],[f4144,f3431])).
fof(f3431,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) | $spl26),
  inference(subsumption_resolution,[],[f3364,f86])).
fof(f3364,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) | s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | $spl26),
  inference(superposition,[],[f128,f1162])).
fof(f128,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.bool.IN_2'(s(X0,sK2(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK2(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) )),
  inference(cnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (s(bool,'const.bool.IN_2'(s(X0,sK2(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK2(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2])],[f81])).
fof(f81,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(rectify,[],[f80])).
fof(f80,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(nnf_transformation,[],[f45])).
fof(f45,plain,(
  ! [X0,X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(flattening,[],[f44])).
fof(f44,plain,(
  ! [X0] : ! [X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X6] : ! [X7,X10] : (s(fun(X6,bool),X7) = s(fun(X6,bool),X10) <=> ! [X11] : s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X7))) = s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X10))))),
  file('thm.pred_set.UNION_EMPTY|split|1','thm.pred_set.EXTENSION')).
fof(f4144,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) | ($spl15 | $spl26 | $spl45)),
  inference(condensation,[],[f4143])).
fof(f4143,plain,(
  ( ! [X4] : (s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1)))) ) | ($spl15 | $spl26 | $spl45)),
  inference(subsumption_resolution,[],[f4123,f2143])).
fof(f4123,plain,(
  ( ! [X4] : (p(s(bool,f(s(bool,X4)))) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1)))) ) | ($spl15 | $spl26 | $spl45)),
  inference(superposition,[],[f3430,f1506])).
fof(f3430,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1)))) | $spl26),
  inference(subsumption_resolution,[],[f3429,f89])).
fof(f3429,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1)))) | $spl26),
  inference(forward_demodulation,[],[f3412,f163])).
fof(f3412,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1)))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | $spl26),
  inference(subsumption_resolution,[],[f3365,f87])).
fof(f3365,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1)))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | $spl26),
  inference(superposition,[],[f129,f1162])).
fof(f129,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f84])).
fof(f4142,plain,(
  $false | ($spl15 | $spl26)),
  inference(subsumption_resolution,[],[f4141,f89])).
fof(f4141,plain,(
  p(s(bool,'F_0')) | ($spl15 | $spl26)),
  inference(forward_demodulation,[],[f4140,f468])).
fof(f4140,plain,(
  p(s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0))))) | ($spl15 | $spl26)),
  inference(subsumption_resolution,[],[f4122,f3431])).
fof(f4122,plain,(
  p(s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) | ($spl15 | $spl26)),
  inference(superposition,[],[f3430,f496])).
fof(f4139,plain,(
  $false | ($spl11 | $spl26)),
  inference(subsumption_resolution,[],[f4138,f89])).
fof(f4138,plain,(
  p(s(bool,'F_0')) | ($spl11 | $spl26)),
  inference(forward_demodulation,[],[f4137,f253])).
fof(f4137,plain,(
  p(s(bool,g(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0))))) | ($spl11 | $spl26)),
  inference(subsumption_resolution,[],[f4121,f3431])).
fof(f4121,plain,(
  p(s(bool,g(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) | ($spl11 | $spl26)),
  inference(superposition,[],[f3430,f273])).
fof(f4136,plain,(
  $false | $spl26),
  inference(subsumption_resolution,[],[f4135,f3431])).
fof(f4135,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) | $spl26),
  inference(condensation,[],[f4134])).
fof(f4134,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X2)) ) | $spl26),
  inference(subsumption_resolution,[],[f4119,f142])).
fof(f4119,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X2)) ) | $spl26),
  inference(superposition,[],[f3430,f140])).
fof(f4133,plain,(
  $false | $spl26),
  inference(subsumption_resolution,[],[f4132,f3431])).
fof(f4132,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) | $spl26),
  inference(condensation,[],[f4131])).
fof(f4131,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1)))) ) | $spl26),
  inference(subsumption_resolution,[],[f4118,f142])).
fof(f4118,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1)))) ) | $spl26),
  inference(superposition,[],[f3430,f140])).
fof(f4130,plain,(
  $false | $spl26),
  inference(subsumption_resolution,[],[f4129,f3431])).
fof(f4129,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) | $spl26),
  inference(subsumption_resolution,[],[f4116,f89])).
fof(f4116,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) | $spl26),
  inference(superposition,[],[f3430,f90])).
fof(f4126,plain,(
  $false | $spl26),
  inference(subsumption_resolution,[],[f4111,f3431])).
fof(f4111,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) | $spl26),
  inference(unit_resulting_resolution,[],[f87,f3430,f135])).
fof(f4125,plain,(
  $false | $spl26),
  inference(subsumption_resolution,[],[f4113,f3431])).
fof(f4113,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) | $spl26),
  inference(unit_resulting_resolution,[],[f87,f3430,f135])).
fof(f4124,plain,(
  $false | $spl26),
  inference(subsumption_resolution,[],[f4115,f3431])).
fof(f4115,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) | $spl26),
  inference(unit_resulting_resolution,[],[f3430,f142])).
fof(f4091,plain,(
  $false | ($spl15 | $spl16 | $spl45)),
  inference(subsumption_resolution,[],[f4090,f2987])).
fof(f2987,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl16),
  inference(subsumption_resolution,[],[f2922,f86])).
fof(f2922,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | $spl16),
  inference(superposition,[],[f128,f1064])).
fof(f4090,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl15 | $spl16 | $spl45)),
  inference(condensation,[],[f4089])).
fof(f4089,plain,(
  ( ! [X4] : (s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1)))) ) | ($spl15 | $spl16 | $spl45)),
  inference(subsumption_resolution,[],[f4069,f2143])).
fof(f4069,plain,(
  ( ! [X4] : (p(s(bool,f(s(bool,X4)))) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1)))) ) | ($spl15 | $spl16 | $spl45)),
  inference(superposition,[],[f2986,f1506])).
fof(f2986,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1)))) | $spl16),
  inference(subsumption_resolution,[],[f2985,f89])).
fof(f2985,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1)))) | $spl16),
  inference(forward_demodulation,[],[f2968,f163])).
fof(f2968,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1)))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | $spl16),
  inference(subsumption_resolution,[],[f2923,f87])).
fof(f2923,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1)))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | $spl16),
  inference(superposition,[],[f129,f1064])).
fof(f4088,plain,(
  $false | ($spl15 | $spl16)),
  inference(subsumption_resolution,[],[f4087,f89])).
fof(f4087,plain,(
  p(s(bool,'F_0')) | ($spl15 | $spl16)),
  inference(forward_demodulation,[],[f4086,f468])).
fof(f4086,plain,(
  p(s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0))))) | ($spl15 | $spl16)),
  inference(subsumption_resolution,[],[f4068,f2987])).
fof(f4068,plain,(
  p(s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl15 | $spl16)),
  inference(superposition,[],[f2986,f496])).
fof(f4085,plain,(
  $false | ($spl11 | $spl16)),
  inference(subsumption_resolution,[],[f4084,f89])).
fof(f4084,plain,(
  p(s(bool,'F_0')) | ($spl11 | $spl16)),
  inference(forward_demodulation,[],[f4083,f253])).
fof(f4083,plain,(
  p(s(bool,g(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0))))) | ($spl11 | $spl16)),
  inference(subsumption_resolution,[],[f4067,f2987])).
fof(f4067,plain,(
  p(s(bool,g(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl11 | $spl16)),
  inference(superposition,[],[f2986,f273])).
fof(f4082,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f4081,f2987])).
fof(f4081,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl16),
  inference(condensation,[],[f4080])).
fof(f4080,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X2)) ) | $spl16),
  inference(subsumption_resolution,[],[f4065,f142])).
fof(f4065,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X2)) ) | $spl16),
  inference(superposition,[],[f2986,f140])).
fof(f4079,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f4078,f2987])).
fof(f4078,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl16),
  inference(condensation,[],[f4077])).
fof(f4077,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1)))) ) | $spl16),
  inference(subsumption_resolution,[],[f4064,f142])).
fof(f4064,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1)))) ) | $spl16),
  inference(superposition,[],[f2986,f140])).
fof(f4076,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f4075,f2987])).
fof(f4075,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl16),
  inference(subsumption_resolution,[],[f4062,f89])).
fof(f4062,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl16),
  inference(superposition,[],[f2986,f90])).
fof(f4072,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f4057,f2987])).
fof(f4057,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl16),
  inference(unit_resulting_resolution,[],[f87,f2986,f135])).
fof(f4071,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f4059,f2987])).
fof(f4059,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl16),
  inference(unit_resulting_resolution,[],[f87,f2986,f135])).
fof(f4070,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f4061,f2987])).
fof(f4061,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl16),
  inference(unit_resulting_resolution,[],[f2986,f142])).
fof(f3416,plain,(
  $false | ($spl15 | $spl26 | $spl33)),
  inference(subsumption_resolution,[],[f3415,f89])).
fof(f3415,plain,(
  p(s(bool,'F_0')) | ($spl15 | $spl26 | $spl33)),
  inference(forward_demodulation,[],[f3414,f163])).
fof(f3414,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | ($spl15 | $spl26 | $spl33)),
  inference(subsumption_resolution,[],[f3413,f89])).
fof(f3413,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | ($spl15 | $spl26 | $spl33)),
  inference(forward_demodulation,[],[f3412,f1205])).
fof(f1205,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) | ($spl15 | $spl33)),
  inference(forward_demodulation,[],[f1190,f468])).
fof(f1190,plain,(
  s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) | ($spl15 | $spl33)),
  inference(unit_resulting_resolution,[],[f1167,f496])).
fof(f2972,plain,(
  $false | ($spl15 | $spl16 | $spl23)),
  inference(subsumption_resolution,[],[f2971,f89])).
fof(f2971,plain,(
  p(s(bool,'F_0')) | ($spl15 | $spl16 | $spl23)),
  inference(forward_demodulation,[],[f2970,f163])).
fof(f2970,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | ($spl15 | $spl16 | $spl23)),
  inference(subsumption_resolution,[],[f2969,f89])).
fof(f2969,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | ($spl15 | $spl16 | $spl23)),
  inference(forward_demodulation,[],[f2968,f1105])).
fof(f1105,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl15 | $spl23)),
  inference(forward_demodulation,[],[f1090,f468])).
fof(f1090,plain,(
  s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl15 | $spl23)),
  inference(unit_resulting_resolution,[],[f1069,f496])).
fof(f1545,plain,(
  ( ! [X4,X2,X3] : (s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,X4) | s(bool,X4) = s(bool,sK3(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X2,X3,bool,bool))))) ) | $spl48),
  inference(cnf_transformation,[],[f1545_D])).
fof(f1545_D,plain,(
  ( ! [X4,X2,X3] : (s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,X4) | s(bool,X4) = s(bool,sK3(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X2,X3,bool,bool))))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f1547,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | $spl51),
  inference(cnf_transformation,[],[f1547_D])).
fof(f1547_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl51),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl51])])).
fof(f1529,plain,(
  ( ! [X12,X13,X11] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X11) = s(fun(bool,bool),X12) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X11,X12,bool,bool)))) | s(bool,X13) = s(bool,sK3(X11,X12,bool,bool)) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f229,f140])).
fof(f229,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK3(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK3(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f136,f90])).
fof(f1526,plain,(
  ( ! [X4,X2,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X2,X3,bool,bool)))) | s(bool,X4) = s(bool,sK3(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f229,f140])).
fof(f1270,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) ) | $spl42),
  inference(cnf_transformation,[],[f1270_D])).
fof(f1270_D,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f1267,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK3(X0,X1,bool,bool))) ) | $spl38),
  inference(cnf_transformation,[],[f1267_D])).
fof(f1267_D,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK3(X0,X1,bool,bool))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f1272,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl45),
  inference(cnf_transformation,[],[f1272_D])).
fof(f1272_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl45),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl45])])).
fof(f1264,plain,(
  ( ! [X12,X13,X11] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X11) = s(fun(bool,bool),X12) | s(bool,'T_0') = s(bool,sK3(X11,X12,bool,bool)) | s(bool,X13) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f228,f140])).
fof(f228,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK3(X0,X1,X2,bool))) )),
  inference(superposition,[],[f136,f90])).
fof(f1266,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl36),
  inference(cnf_transformation,[],[f1266_D])).
fof(f1266_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f1273,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) ) | $spl46),
  inference(cnf_transformation,[],[f1273_D])).
fof(f1273_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f1263,plain,(
  ( ! [X10,X8,X9] : (s(bool,X8) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X9) = s(fun(bool,bool),X10) | s(bool,'T_0') = s(bool,sK3(X9,X10,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f228,f140])).
fof(f1262,plain,(
  ( ! [X6,X7,X5] : (s(bool,X5) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X6) = s(fun(bool,bool),X7) | s(bool,'T_0') = s(bool,sK3(X6,X7,bool,bool)) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f228,f140])).
fof(f1261,plain,(
  ( ! [X4,X2,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,sK3(X2,X3,bool,bool)) | s(bool,X4) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f228,f140])).
fof(f1269,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl41),
  inference(cnf_transformation,[],[f1269_D])).
fof(f1269_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl41),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl41])])).
fof(f1260,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK3(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f228,f90])).
fof(f1162,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | $spl26),
  inference(cnf_transformation,[],[f1162_D])).
fof(f1162_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f1164,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) | $spl29),
  inference(cnf_transformation,[],[f1164_D])).
fof(f1164_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f1170,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | $spl15),
  inference(forward_demodulation,[],[f1160,f468])).
fof(f1160,plain,(
  s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))) != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | $spl15),
  inference(superposition,[],[f410,f496])).
fof(f410,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))))),
  inference(unit_resulting_resolution,[],[f86,f128])).
fof(f1169,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | $spl11),
  inference(forward_demodulation,[],[f1159,f253])).
fof(f1159,plain,(
  s(bool,g(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))) != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | $spl11),
  inference(superposition,[],[f410,f273])).
fof(f1165,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl30),
  inference(cnf_transformation,[],[f1165_D])).
fof(f1165_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f1167,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) | $spl33),
  inference(cnf_transformation,[],[f1167_D])).
fof(f1167_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f1158,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) | s(bool,X3) = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f410,f140])).
fof(f1168,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X1)) ) | $spl34),
  inference(cnf_transformation,[],[f1168_D])).
fof(f1168_D,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X1)) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f1157,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f410,f140])).
fof(f1156,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))))) )),
  inference(superposition,[],[f410,f140])).
fof(f1155,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f410,f140])).
fof(f1154,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))))),
  inference(superposition,[],[f410,f90])).
fof(f1064,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | $spl16),
  inference(cnf_transformation,[],[f1064_D])).
fof(f1064_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f1066,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl19),
  inference(cnf_transformation,[],[f1066_D])).
fof(f1066_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f1072,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | $spl15),
  inference(forward_demodulation,[],[f1062,f468])).
fof(f1062,plain,(
  s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | $spl15),
  inference(superposition,[],[f409,f496])).
fof(f409,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))))),
  inference(unit_resulting_resolution,[],[f86,f128])).
fof(f1071,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | $spl11),
  inference(forward_demodulation,[],[f1061,f253])).
fof(f1061,plain,(
  s(bool,g(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | $spl11),
  inference(superposition,[],[f409,f273])).
fof(f1067,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl20),
  inference(cnf_transformation,[],[f1067_D])).
fof(f1067_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f1069,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl23),
  inference(cnf_transformation,[],[f1069_D])).
fof(f1069_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f1060,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,X3) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f409,f140])).
fof(f1070,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X1)) ) | $spl24),
  inference(cnf_transformation,[],[f1070_D])).
fof(f1070_D,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X1)) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f1059,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f409,f140])).
fof(f1058,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))))) )),
  inference(superposition,[],[f409,f140])).
fof(f1057,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f409,f140])).
fof(f1056,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))))),
  inference(superposition,[],[f409,f90])).
fof(f425,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))) | $spl12),
  inference(cnf_transformation,[],[f425_D])).
fof(f425_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f427,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))) | $spl15),
  inference(cnf_transformation,[],[f427_D])).
fof(f427_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f423,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0))))),
  inference(superposition,[],[f227,f90])).
fof(f232,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))) | $spl8),
  inference(cnf_transformation,[],[f232_D])).
fof(f232_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f234,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))) | $spl11),
  inference(cnf_transformation,[],[f234_D])).
fof(f234_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f230,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0))))),
  inference(superposition,[],[f226,f90])).
fof(f187,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f180,f87])).
fof(f180,plain,(
  ~p(s(bool,'T_0')) | $spl4),
  inference(backward_demodulation,[],[f175,f89])).
fof(f175,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl4),
  inference(cnf_transformation,[],[f175_D])).
fof(f175_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f176,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1)) ) | $spl6),
  inference(cnf_transformation,[],[f176_D])).
fof(f176_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1)) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f178,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f170,f142])).
fof(f170,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f132,f90])).
fof(f174,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f168,f142])).
fof(f168,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f132,f90])).
fof(f149,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl0),
  inference(cnf_transformation,[],[f149_D])).
fof(f149_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f151,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl3),
  inference(cnf_transformation,[],[f151_D])).
fof(f151_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f148,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f90])).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 89
% Passive clauses: 300
% Generated clauses: 7586
% Final active clauses: 68
% Final passive clauses: 174
% Input formulas: 23
% Initial clauses: 51
% 
% Duplicate literals: 553
% Trivial inequalities: 30
% Fw subsumption resolutions: 526
% Fw demodulations: 1647
% Bw demodulations: 36
% Condensations: 468
% 
% Simple tautologies: 105
% Equational tautologies: 83
% Forward subsumptions: 3758
% Fw demodulations to eq. taut.: 10
% 
% Unit resulting resolution: 469
% Factoring: 1
% Forward superposition: 2196
% Backward superposition: 1467
% Self superposition: 68
% Equality factoring: 34
% Equality resolution: 5
% 
% Split clauses: 26
% Split components: 57
% SAT solver clauses: 434
% SAT solver unit clauses: 27
% SAT solver binary clauses: 94
% 
% Sat splits: 26
% Sat splitting refutations: 80
% 
% Memory used [KB]: 4349
% Time elapsed: 0.263 s
% ------------------------------
% ------------------------------
% Success in time 1.59 s
