Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.integer
% (75717)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 770
% Passive clauses: 9044
% Generated clauses: 22730
% Final active clauses: 721
% Final passive clauses: 7682
% Input formulas: 62
% Initial clauses: 234
% 
% Duplicate literals: 420
% Trivial inequalities: 22
% Fw subsumption resolutions: 2225
% Fw demodulations: 1328
% Bw demodulations: 47
% 
% Simple tautologies: 280
% Equational tautologies: 82
% Forward subsumptions: 8830
% 
% Binary resolution: 10068
% Factoring: 1525
% Forward superposition: 3300
% Backward superposition: 3295
% Self superposition: 42
% Equality factoring: 28
% Equality resolution: 14
% 
% Split clauses: 383
% Split components: 803
% SAT solver clauses: 807
% SAT solver unit clauses: 20
% SAT solver binary clauses: 122
% 
% Sat splits: 383
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 805
% TWLsolver calls for satisfiability: 428
% 
% Memory used [KB]: 9722
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.integer
% (75724)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 473
% Passive clauses: 2949
% Generated clauses: 12862
% Final active clauses: 471
% Final passive clauses: 2475
% Input formulas: 62
% Initial clauses: 234
% 
% Duplicate literals: 247
% Trivial inequalities: 29
% Fw subsumption resolutions: 968
% Fw demodulations: 404
% Bw demodulations: 3
% Condensations: 444
% Global subsumptions: 22
% 
% Simple tautologies: 283
% Equational tautologies: 17
% Forward subsumptions: 4123
% 
% Binary resolution: 4020
% Factoring: 403
% Forward superposition: 1573
% Backward superposition: 1479
% Self superposition: 38
% Equality factoring: 26
% Equality resolution: 9
% 
% Split clauses: 40
% Split components: 80
% SAT solver clauses: 3577
% SAT solver unit clauses: 231
% SAT solver binary clauses: 90
% SAT solver learnt clauses: 44
% SAT solver learnt literals: 10
% 
% Sat splits: 40
% 
% TWLsolver clauses: 3489
% TWLsolver calls for satisfiability: 2096
% 
% Memory used [KB]: 13688
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.integer
% (75738)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 2230
% Passive clauses: 13703
% Generated clauses: 273782
% Final active clauses: 1031
% Final passive clauses: 2003
% Input formulas: 62
% Initial clauses: 234
% 
% Duplicate literals: 4322
% Trivial inequalities: 152
% Fw subsumption resolutions: 6960
% Fw demodulations: 99393
% Bw demodulations: 5460
% Condensations: 8555
% Global subsumptions: 1612
% 
% Simple tautologies: 825
% Equational tautologies: 602
% Forward subsumptions: 110797
% Fw demodulations to eq. taut.: 20874
% Bw demodulations to eq. taut.: 587
% 
% Unit resulting resolution: 61850
% Factoring: 9
% Forward superposition: 38789
% Backward superposition: 38675
% Self superposition: 90
% Equality factoring: 58
% Equality resolution: 93
% 
% Split clauses: 933
% Split components: 1909
% SAT solver clauses: 22237
% SAT solver unit clauses: 3061
% SAT solver binary clauses: 5545
% 
% Sat splits: 933
% Sat splitting refutations: 5054
% 
% Memory used [KB]: 113089
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.integer
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.integer
% SZS output start Proof for thm.integer
fof(f9643,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f543,f547,f550,f554,f552,f539,f557,f541,f558,f570,f565,f577,f567,f578,f699,f580,f677,f679,f838,f839,f846,f574,f848,f815,f849,f817,f850,f863,f872,f865,f873,f882,f890,f884,f891,f1013,f1015,f1014,f893,f1018,f1016,f1021,f572,f991,f1022,f556,f993,f1023,f555,f1148,f701,f700,f1029,f704,f702,f1126,f707,f1128,f708,f1154,f1250,f1252,f1284,f1406,f1383,f1411,f1385,f1412,f1546,f1552,f1667,f1680,f1681,f1684,f1685,f544,f1822,f1823,f1835,f1665,f1839,f1840,f1666,f1892,f2019,f1907,f2024,f2022,f2335,f2350,f2351,f2356,f2355,f2353,f2357,f2354,f2443,f2445,f2444,f2791,f2792,f3021,f3022,f3034,f3038,f3035,f3039,f3171,f3172,f2856,f2857,f2858,f2859,f3181,f3182,f3184,f3183,f3185,f3186,f3180,f3188,f3191,f3192,f3195,f3322,f3320,f3319,f3314,f3305,f3304,f3303,f3398,f3399,f3400,f3401,f3402,f3403,f3552,f3685,f4113,f3027,f4114,f3029,f4115,f4122,f4116,f4153,f3176,f4154,f3178,f4155,f4162,f4156,f4168,f1838,f4294,f1274,f1273,f4390,f1277,f576,f4392,f1278,f575,f4667,f3175,f4669,f4962,f5458,f5459,f5460,f5461,f5462,f5463,f6001,f6002,f6003,f6004,f6005,f6006,f6016,f6017,f6018,f6019,f6020,f6021,f6154,f6152,f6155,f6283,f6319,f6449,f6521,f6646,f6645,f6642,f6641,f6640,f6631,f6713,f6714,f6715,f6716,f6735,f6737,f3193,f6742,f6740,f6743,f1905,f6747,f3686,f6749,f6750,f6776,f6785,f6804,f4164,f6805,f4124,f6817,f6826,f6874,f6879,f7088,f7089,f7090,f7104,f4673,f7164,f7176,f7207,f7214,f7217,f7215,f7222,f7223,f1903,f7225,f6450,f1902,f7353,f7359,f7539,f4163,f7602,f4672,f7603,f7611,f7987,f867,f8183,f8161,f549,f545,f8298,f8299,f8185,f8302,f8300,f8615,f2021,f2020,f8754,f1668,f8870,f8922,f1682,f8923,f9205,f9206,f9221,f3025,f9222,f9223,f4169,f6434,f9417,f2337,f9432,f2336,f9418,f9558,f9555,f9554,f9553,f9544,f9626,f9627,f9628,f9629])).
fof(f9629,plain,(
  $false | ($spl107 | $spl238)),
  inference(subsumption_resolution,[],[f9546,f9438])).
fof(f9438,plain,(
  ( ! [X1] : (p(s(bool,X1))) ) | ($spl107 | $spl238)),
  inference(subsumption_resolution,[],[f9437,f519])).
fof(f519,plain,(
  ( ! [X1] : (p(s(bool,X1)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f296])).
fof(f296,plain,(
  (~p(s(bool,'F_0')) | ! [X1] : p(s(bool,X1))) & (~p(s(bool,sK15)) | p(s(bool,'F_0')))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK15])],[f295])).
fof(f295,plain,(
  (~p(s(bool,'F_0')) | ! [X1] : p(s(bool,X1))) & (? [X0] : ~p(s(bool,X0)) | p(s(bool,'F_0')))),
  inference(rectify,[],[f294])).
fof(f294,plain,(
  (~p(s(bool,'F_0')) | ! [X0] : p(s(bool,X0))) & (? [X0] : ~p(s(bool,X0)) | p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f143])).
fof(f143,plain,(
  p(s(bool,'F_0')) <=> ! [X0] : p(s(bool,X0))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'F_0')) <=> ! [X8] : p(s(bool,X8))),
  file('thm.integer.INT_MOD_COMMON_FACTOR','thm.bool.F_DEF')).
fof(f9437,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1))) ) | ($spl107 | $spl238)),
  inference(forward_demodulation,[],[f9434,f7382])).
fof(f7382,plain,(
  s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl107),
  inference(resolution,[],[f1904,f343])).
fof(f343,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f213])).
fof(f213,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f80])).
fof(f80,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X8] : (s(bool,'F_0') = s(bool,X8) <=> ~p(s(bool,X8)))),
  file('thm.integer.INT_MOD_COMMON_FACTOR','thm.bool.EQ_CLAUSES|split|3')).
fof(f1904,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl107),
  inference(cnf_transformation,[],[f1904_D])).
fof(f1904_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl107),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl107])])).
fof(f9434,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X1))) ) | $spl238),
  inference(resolution,[],[f9418,f299])).
fof(f299,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | p(s(bool,X2)) | ~sP0(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f196])).
fof(f196,plain,(
  ! [X0,X1,X2] : (~sP0(X0,X1,X2) | ((p(s(bool,X2)) | (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) & (~p(s(bool,X2)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))))),
  inference(rectify,[],[f195])).
fof(f195,plain,(
  ! [X2,X1,X4] : (~sP0(X2,X1,X4) | ((p(s(bool,X4)) | (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X2)))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) & (~p(s(bool,X4)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X2)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))))),
  inference(flattening,[],[f194])).
fof(f194,plain,(
  ! [X2,X1,X4] : (~sP0(X2,X1,X4) | ((p(s(bool,X4)) | (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X2)))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) & (~p(s(bool,X4)) | (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X2)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))))),
  inference(nnf_transformation,[],[f173])).
fof(f173,plain,(
  ! [X2,X1,X4] : (~sP0(X2,X1,X4) | (p(s(bool,X4)) <~> (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X2)))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f9546,plain,(
  ~p(s(bool,'T_0')) | ($spl107 | $spl238)),
  inference(resolution,[],[f9438,f518])).
fof(f518,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f293])).
fof(f293,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f142])).
fof(f142,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f31])).
fof(f31,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.integer.INT_MOD_COMMON_FACTOR','thm.bool.NOT_CLAUSES|split|1')).
fof(f9628,plain,(
  $false | ($spl107 | $spl238)),
  inference(subsumption_resolution,[],[f9545,f9438])).
fof(f9545,plain,(
  ~p(s(bool,'T_0')) | ($spl107 | $spl238)),
  inference(resolution,[],[f9438,f516])).
fof(f516,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f292])).
fof(f292,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f141])).
fof(f141,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f32])).
fof(f32,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.integer.INT_MOD_COMMON_FACTOR','thm.bool.NOT_CLAUSES|split|2')).
fof(f9627,plain,(
  $false | ($spl107 | $spl238)),
  inference(subsumption_resolution,[],[f9542,f9438])).
fof(f9542,plain,(
  ( ! [X242] : (~p(s(bool,X242))) ) | ($spl107 | $spl238)),
  inference(resolution,[],[f9438,f510])).
fof(f510,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f169])).
fof(f169,plain,(
  ! [X0] : ((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) => ~p(s(bool,X0)))),
  inference(flattening,[],[f133])).
fof(f133,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) => ~p(s(bool,X0)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X8] : ((p(s(bool,X8)) => p(s(bool,'F_0'))) => ~p(s(bool,X8)))),
  file('thm.integer.INT_MOD_COMMON_FACTOR','thm.bool.IMP_F')).
fof(f9626,plain,(
  $false | ($spl107 | $spl238)),
  inference(subsumption_resolution,[],[f9536,f9438])).
fof(f9536,plain,(
  ( ! [X236] : (~p(s(bool,X236))) ) | ($spl107 | $spl238)),
  inference(resolution,[],[f9438,f460])).
fof(f460,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f263])).
fof(f263,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f262])).
fof(f262,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f160])).
fof(f160,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f114])).
fof(f114,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f113])).
fof(f113,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X8] : ((p(s(bool,X8)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X8)))),
  file('thm.integer.INT_MOD_COMMON_FACTOR','thm.bool.IMP_CLAUSES|split|4')).
fof(f9544,plain,(
  $false | ($spl107 | $spl238)),
  inference(resolution,[],[f9438,f514])).
fof(f514,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f140])).
fof(f140,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.integer.INT_MOD_COMMON_FACTOR',aHL_FALSITY)).
fof(f9553,plain,(
  $false | ($spl87 | $spl107 | $spl238)),
  inference(resolution,[],[f9438,f1666])).
fof(f9554,plain,(
  $false | ($spl11 | $spl107 | $spl238)),
  inference(resolution,[],[f9438,f555])).
fof(f9555,plain,(
  $false | ($spl107 | $spl121 | $spl238)),
  inference(resolution,[],[f9438,f2354])).
fof(f9558,plain,(
  $false | ($spl5 | $spl107 | $spl238)),
  inference(resolution,[],[f9438,f548])).
fof(f548,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl5),
  inference(cnf_transformation,[],[f548_D])).
fof(f548_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f9418,plain,(
  ( ! [X29] : (sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X29)) ) | $spl238),
  inference(cnf_transformation,[],[f9418_D])).
fof(f9418_D,plain,(
  ( ! [X29] : (sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X29)) ) <=> ~$spl238),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl238])])).
fof(f2336,plain,(
  ( ! [X34] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X34),s('type.integer.int',sK13)))) ) | $spl116),
  inference(cnf_transformation,[],[f2336_D])).
fof(f2336_D,plain,(
  ( ! [X34] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X34),s('type.integer.int',sK13)))) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f9432,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',sK13))) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X1)) ) | ($spl104 | $spl234)),
  inference(subsumption_resolution,[],[f9431,f524])).
fof(f524,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f340])).
fof(f340,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f212])).
fof(f212,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f78])).
fof(f78,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X11] : ! [X12] : (s(X11,X12) = s(X11,X12) <=> p(s(bool,'T_0')))),
  file('thm.integer.INT_MOD_COMMON_FACTOR','thm.bool.REFL_CLAUSE')).
fof(f9431,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',sK13))) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X1)) ) | ($spl104 | $spl234)),
  inference(forward_demodulation,[],[f9430,f7263])).
fof(f7263,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))) | $spl104),
  inference(resolution,[],[f1903,f347])).
fof(f347,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f215])).
fof(f215,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X8] : (s(bool,'T_0') = s(bool,X8) <=> p(s(bool,X8)))),
  file('thm.integer.INT_MOD_COMMON_FACTOR','thm.bool.EQ_CLAUSES|split|1')).
fof(f9430,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',sK13))) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X1) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) ) | $spl234),
  inference(forward_demodulation,[],[f9429,f354])).
fof(f354,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(cnf_transformation,[],[f89])).
fof(f89,plain,(
  ! [X0] : s('type.integer.int',X0) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(rectify,[],[f49])).
fof(f49,axiom,(
  ! [X12] : s('type.integer.int',X12) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.integer.INT_MOD_COMMON_FACTOR','thm.integer.INT_ADD_RID')).
fof(f9429,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X1) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) ) | $spl234),
  inference(subsumption_resolution,[],[f9408,f523])).
fof(f523,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X0))))) )),
  inference(cnf_transformation,[],[f144])).
fof(f144,plain,(
  ! [X0] : p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X0))))),
  inference(rectify,[],[f50])).
fof(f50,axiom,(
  ! [X12] : p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X12),s('type.integer.int',X12))))),
  file('thm.integer.INT_MOD_COMMON_FACTOR','thm.integer.INT_LE_REFL')).
fof(f9408,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X1) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) ) | $spl234),
  inference(equality_resolution,[],[f9206])).
fof(f2337,plain,(
  ( ! [X31] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X31),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X31))))) ) | $spl118),
  inference(cnf_transformation,[],[f2337_D])).
fof(f2337_D,plain,(
  ( ! [X31] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X31),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X31))))) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f9417,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X27),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X28),s('type.integer.int',sK13))) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X29) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X27))))) ) | ($spl104 | $spl234)),
  inference(forward_demodulation,[],[f9416,f353])).
fof(f353,plain,(
  ( ! [X0] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))) )),
  inference(cnf_transformation,[],[f88])).
fof(f88,plain,(
  ! [X0] : s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))),
  inference(rectify,[],[f51])).
fof(f51,axiom,(
  ! [X12] : s(bool,'const.integer.int_lt_2'(s('type.integer.int',X12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X12)))))),
  file('thm.integer.INT_MOD_COMMON_FACTOR','thm.integer.INT_NEG_GT|48|')).
fof(f9416,plain,(
  ( ! [X28,X29,X27] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X28),s('type.integer.int',sK13))) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X29) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X27)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X27))))) ) | ($spl104 | $spl234)),
  inference(subsumption_resolution,[],[f9415,f524])).
fof(f9415,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,'T_0')) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X28),s('type.integer.int',sK13))) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X29) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X27)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X27))))) ) | ($spl104 | $spl234)),
  inference(forward_demodulation,[],[f9414,f7263])).
fof(f9414,plain,(
  ( ! [X28,X29,X27] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X28),s('type.integer.int',sK13))) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X29) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X27)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X27))))) ) | $spl234),
  inference(forward_demodulation,[],[f9413,f354])).
fof(f9413,plain,(
  ( ! [X28,X29,X27] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X28),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X29) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X27)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X27))))) ) | $spl234),
  inference(subsumption_resolution,[],[f9409,f523])).
fof(f9409,plain,(
  ( ! [X28,X29,X27] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X28),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X29) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X27)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X27))))) ) | $spl234),
  inference(trivial_inequality_removal,[],[f9393])).
fof(f9393,plain,(
  ( ! [X28,X29,X27] : (s('type.integer.int',X27) != s('type.integer.int',X27) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X28),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X29) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X27)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X27))))) ) | $spl234),
  inference(superposition,[],[f9206,f352])).
fof(f352,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X12] : (s('type.integer.int',X12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X12)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X12)))))))),
  file('thm.integer.INT_MOD_COMMON_FACTOR','thm.integer.INT_LT_NEGTOTAL')).
fof(f6434,plain,(
  ( ! [X2,X1] : (~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | sP0(sK13,'const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),X1) | p(s(bool,X2))) ) | $spl188),
  inference(cnf_transformation,[],[f6434_D])).
fof(f6434_D,plain,(
  ( ! [X2,X1] : (~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | sP0(sK13,'const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),X1) | p(s(bool,X2))) ) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f4169,plain,(
  ( ! [X3,X1] : (s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int',X1) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',sK13)))) ) | $spl164),
  inference(cnf_transformation,[],[f4169_D])).
fof(f4169_D,plain,(
  ( ! [X3,X1] : (s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int',X1) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',sK13)))) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f9223,plain,(
  ( ! [X17,X15,X18,X16] : (s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X15),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X16),s(bool,X17)))) | p(s(bool,X17)) | sP0(sK13,'const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),X16) | s('type.integer.int',X15) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X18),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',sK13)))) ) | $spl107),
  inference(subsumption_resolution,[],[f8945,f481])).
fof(f481,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f277])).
fof(f277,plain,(
  ! [X0] : (((~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | p(s(bool,X0))) & (~p(s(bool,X0)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f276])).
fof(f276,plain,(
  ! [X0] : (((~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X0] : ((p(s(bool,X0)) | p(s(bool,'F_0'))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X8] : ((p(s(bool,X8)) | p(s(bool,'F_0'))) <=> p(s(bool,X8)))),
  file('thm.integer.INT_MOD_COMMON_FACTOR','thm.bool.OR_CLAUSES|split|3')).
fof(f8945,plain,(
  ( ! [X17,X15,X18,X16] : (p(s(bool,'F_0')) | s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X15),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X16),s(bool,X17)))) | p(s(bool,X17)) | sP0(sK13,'const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),X16) | s('type.integer.int',X15) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X18),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',sK13)))) ) | $spl107),
  inference(superposition,[],[f303,f7382])).
fof(f303,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.integer.int',X2) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int',X1))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X4),s(bool,X5)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X2),s('type.integer.int',X1)))) | p(s(bool,X5)) | sP0(X2,X1,X4) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int',X1))),s('type.integer.int',X2)))) )),
  inference(cnf_transformation,[],[f198])).
fof(f198,plain,(
  ! [X0,X1,X2] : (! [X3] : (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int',X1))),s('type.integer.int',X2))) | ! [X4] : (sP0(X2,X1,X4) | ! [X5] : (((p(s(bool,X5)) | (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X2)))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X2),s('type.integer.int',X1)))))) & (~p(s(bool,X5)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X2)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X2),s('type.integer.int',X1)))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X4),s(bool,X5))))))) | s('type.integer.int',X2) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int',X1))))),
  inference(flattening,[],[f197])).
fof(f197,plain,(
  ! [X0,X1,X2] : (! [X3] : (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int',X1))),s('type.integer.int',X2))) | ! [X4] : (sP0(X2,X1,X4) | ! [X5] : (((p(s(bool,X5)) | (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X2)))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X2),s('type.integer.int',X1)))))) & (~p(s(bool,X5)) | (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X2)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X2),s('type.integer.int',X1))))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X4),s(bool,X5))))))) | s('type.integer.int',X2) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int',X1))))),
  inference(nnf_transformation,[],[f174])).
fof(f174,plain,(
  ! [X0,X1,X2] : (! [X3] : (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int',X1))),s('type.integer.int',X2))) | ! [X4] : (sP0(X2,X1,X4) | ! [X5] : ((p(s(bool,X5)) <~> (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X2)))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X2),s('type.integer.int',X1)))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X4),s(bool,X5))))))) | s('type.integer.int',X2) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int',X1))))),
  inference(definition_folding,[],[f146,f173])).
fof(f146,plain,(
  ! [X0,X1,X2] : (! [X3] : (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int',X1))),s('type.integer.int',X2))) | ! [X4] : ((p(s(bool,X4)) <~> (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X2)))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | ! [X5] : ((p(s(bool,X5)) <~> (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X2)))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X2),s('type.integer.int',X1)))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X4),s(bool,X5))))))) | s('type.integer.int',X2) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int',X1))))),
  inference(ennf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0,X1,X2] : (? [X3] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int',X1))),s('type.integer.int',X2))) & ? [X4] : ((p(s(bool,X4)) <=> (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X2)))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) & ? [X5] : ((p(s(bool,X5)) <=> (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X2)))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X2),s('type.integer.int',X1)))))) & p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X4),s(bool,X5))))))) => s('type.integer.int',X2) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int',X1))))),
  inference(rectify,[],[f53])).
fof(f53,axiom,(
  ! [X21,X22,X23] : (? [X7] : (s('type.integer.int',X21) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X7),s('type.integer.int',X22))),s('type.integer.int',X23))) & ? [X24] : ((p(s(bool,X24)) <=> (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X22),s('type.integer.int',X23)))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X23),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) & ? [X25] : ((p(s(bool,X25)) <=> (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X23)))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X23),s('type.integer.int',X22)))))) & p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X22),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X24),s(bool,X25))))))) => s('type.integer.int',X23) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X21),s('type.integer.int',X22))))),
  file('thm.integer.INT_MOD_COMMON_FACTOR','thm.integer.INT_MOD_UNIQUE')).
fof(f9222,plain,(
  ( ! [X14,X12,X10,X13] : (s('type.integer.int',X12) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X13),s('type.integer.int',sK13))) | s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X14),s('type.integer.int',sK13))),s('type.integer.int',X12))) | sP0(X12,sK13,X10) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X12)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X12),s('type.integer.int',sK13))))) ) | $spl236),
  inference(cnf_transformation,[],[f9222_D])).
fof(f9222_D,plain,(
  ( ! [X14,X12,X10,X13] : (s('type.integer.int',X12) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X13),s('type.integer.int',sK13))) | s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X14),s('type.integer.int',sK13))),s('type.integer.int',X12))) | sP0(X12,sK13,X10) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X12)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X12),s('type.integer.int',sK13))))) ) <=> ~$spl236),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl236])])).
fof(f3025,plain,(
  ( ! [X8] : (~p(s(bool,X8))) ) | $spl132),
  inference(cnf_transformation,[],[f3025_D])).
fof(f3025_D,plain,(
  ( ! [X8] : (~p(s(bool,X8))) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f9221,plain,(
  ( ! [X14,X12,X10,X13,X11] : (~p(s(bool,X11)) | s('type.integer.int',X12) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X13),s('type.integer.int',sK13))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X12),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X12)))) | sP0(X12,sK13,X10) | s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X14),s('type.integer.int',sK13))),s('type.integer.int',X12)))) ) | $spl107),
  inference(duplicate_literal_removal,[],[f9220])).
fof(f9220,plain,(
  ( ! [X14,X12,X10,X13,X11] : (~p(s(bool,X11)) | s('type.integer.int',X12) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X13),s('type.integer.int',sK13))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X12),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X12)))) | ~p(s(bool,X11)) | sP0(X12,sK13,X10) | s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X14),s('type.integer.int',sK13))),s('type.integer.int',X12)))) ) | $spl107),
  inference(forward_demodulation,[],[f8944,f338])).
fof(f338,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0,X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f74])).
fof(f74,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X11] : ! [X9,X10] : s(X11,X10) = s(X11,'const.bool.COND_3'(s(bool,'F_0'),s(X11,X9),s(X11,X10)))),
  file('thm.integer.INT_MOD_COMMON_FACTOR','thm.bool.COND_CLAUSES|split|1')).
fof(f8944,plain,(
  ( ! [X14,X12,X10,X13,X11] : (~p(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X10),s(bool,X11)))) | s('type.integer.int',X12) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X13),s('type.integer.int',sK13))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X12),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X12)))) | ~p(s(bool,X11)) | sP0(X12,sK13,X10) | s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X14),s('type.integer.int',sK13))),s('type.integer.int',X12)))) ) | $spl107),
  inference(superposition,[],[f304,f7382])).
fof(f304,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.integer.int',X2) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int',X1))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X4),s(bool,X5)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X2),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X2)))) | ~p(s(bool,X5)) | sP0(X2,X1,X4) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int',X1))),s('type.integer.int',X2)))) )),
  inference(cnf_transformation,[],[f198])).
fof(f9206,plain,(
  ( ! [X8,X11,X9] : (s('type.integer.int',X8) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int',sK13))),s('type.integer.int',X8))) | sP0(X8,sK13,X9) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X8)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X8),s('type.integer.int',sK13))))) ) | $spl234),
  inference(cnf_transformation,[],[f9206_D])).
fof(f9206_D,plain,(
  ( ! [X8,X11,X9] : (s('type.integer.int',X8) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int',sK13))),s('type.integer.int',X8))) | sP0(X8,sK13,X9) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X8)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X8),s('type.integer.int',sK13))))) ) <=> ~$spl234),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl234])])).
fof(f9205,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,X10)) | s('type.integer.int',X8) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X8),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X8)))) | sP0(X8,sK13,X9) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int',sK13))),s('type.integer.int',X8)))) ) | $spl107),
  inference(duplicate_literal_removal,[],[f9204])).
fof(f9204,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,X10)) | s('type.integer.int',X8) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X8),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X8)))) | ~p(s(bool,X10)) | sP0(X8,sK13,X9) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int',sK13))),s('type.integer.int',X8)))) ) | $spl107),
  inference(forward_demodulation,[],[f8925,f338])).
fof(f8925,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X9),s(bool,X10)))) | s('type.integer.int',X8) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X8),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X8)))) | ~p(s(bool,X10)) | sP0(X8,sK13,X9) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int',sK13))),s('type.integer.int',X8)))) ) | $spl107),
  inference(backward_demodulation,[],[f7382,f537])).
fof(f537,plain,(
  ( ! [X10,X8,X11,X9] : (s('type.integer.int',X8) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X9),s(bool,X10)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X8),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X8)))) | ~p(s(bool,X10)) | sP0(X8,sK13,X9) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int',sK13))),s('type.integer.int',X8)))) )),
  inference(superposition,[],[f298,f304])).
fof(f298,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13)))),
  inference(cnf_transformation,[],[f193])).
fof(f193,plain,(
  s('type.integer.int',sK13) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) & s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13,sK14])],[f145])).
fof(f145,plain,(
  ? [X0] : (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) & ? [X1] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X0))),s('type.integer.int',X0))))),
  inference(ennf_transformation,[],[f65])).
fof(f65,plain,(
  ~! [X0] : (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) => ! [X1] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X0))),s('type.integer.int',X0))))),
  inference(flattening,[],[f64])).
fof(f64,plain,(
  ~! [X0] : (~s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) => ! [X1] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X0))),s('type.integer.int',X0))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : (~s('type.integer.int',X6) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) => ! [X7] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X7),s('type.integer.int',X6))),s('type.integer.int',X6))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : (~s('type.integer.int',X6) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) => ! [X7] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X7),s('type.integer.int',X6))),s('type.integer.int',X6))))),
  file('thm.integer.INT_MOD_COMMON_FACTOR','thm.integer.INT_MOD_COMMON_FACTOR')).
fof(f8923,plain,(
  ( ! [X10,X8] : (p(s(bool,X8)) | ~p(s(bool,X10)) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X8),s(bool,X10))))) ) | $spl232),
  inference(cnf_transformation,[],[f8923_D])).
fof(f8923_D,plain,(
  ( ! [X10,X8] : (p(s(bool,X8)) | ~p(s(bool,X10)) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X8),s(bool,X10))))) ) <=> ~$spl232),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl232])])).
fof(f1682,plain,(
  ( ! [X7,X5] : (s('type.integer.int',X5) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X5),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))))) ) | $spl94),
  inference(cnf_transformation,[],[f1682_D])).
fof(f1682_D,plain,(
  ( ! [X7,X5] : (s('type.integer.int',X5) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X5),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f8922,plain,(
  ( ! [X10,X8,X11,X9] : (s('type.integer.int',X9) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))) | p(s(bool,X8)) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X8),s(bool,X10)))) | ~p(s(bool,X10))) ) | ($spl0 | $spl5)),
  inference(forward_demodulation,[],[f8921,f354])).
fof(f8921,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X8)) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X8),s(bool,X10)))) | ~p(s(bool,X10)) | s('type.integer.int',X9) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl0 | $spl5)),
  inference(subsumption_resolution,[],[f8920,f523])).
fof(f8920,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X8)) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X8),s(bool,X10)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,X10)) | s('type.integer.int',X9) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl0 | $spl5)),
  inference(subsumption_resolution,[],[f8919,f545])).
fof(f8919,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X8)) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X8),s(bool,X10)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,X10)) | s('type.integer.int',X9) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl5),
  inference(resolution,[],[f8755,f304])).
fof(f8755,plain,(
  ( ! [X4] : (~sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),'const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),X4) | p(s(bool,X4))) ) | $spl5),
  inference(resolution,[],[f548,f299])).
fof(f8870,plain,(
  ( ! [X2,X1] : (~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | sP0('const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),'const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),X1) | p(s(bool,X2))) ) | $spl230),
  inference(cnf_transformation,[],[f8870_D])).
fof(f8870_D,plain,(
  ( ! [X2,X1] : (~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | sP0('const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),'const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),X1) | p(s(bool,X2))) ) <=> ~$spl230),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl230])])).
fof(f1668,plain,(
  ( ! [X0,X3] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))))) ) | $spl90),
  inference(cnf_transformation,[],[f1668_D])).
fof(f1668_D,plain,(
  ( ! [X0,X3] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f8754,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | p(s(bool,X2)) | sP0('const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),'const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),X1) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))))) ) | $spl5),
  inference(resolution,[],[f548,f303])).
fof(f2020,plain,(
  ( ! [X2,X1] : (~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X1) | ~p(s(bool,X2))) ) | $spl108),
  inference(cnf_transformation,[],[f2020_D])).
fof(f2020_D,plain,(
  ( ! [X2,X1] : (~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X1) | ~p(s(bool,X2))) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f2021,plain,(
  ( ! [X0,X3] : (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int',sK13))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int',sK13)))) ) | $spl110),
  inference(cnf_transformation,[],[f2021_D])).
fof(f2021_D,plain,(
  ( ! [X0,X3] : (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int',sK13))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int',sK13)))) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f8615,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.integer.int',X4) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X7),s('type.integer.int',sK13))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X4),s('type.integer.int',sK13))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X5),s(bool,X6)))) | ~p(s(bool,X6)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X5)) ) | $spl104),
  inference(forward_demodulation,[],[f8614,f354])).
fof(f8614,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X4),s('type.integer.int',sK13))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X5),s(bool,X6)))) | ~p(s(bool,X6)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X5) | s('type.integer.int',X4) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X7),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl104),
  inference(subsumption_resolution,[],[f8613,f523])).
fof(f8613,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X4),s('type.integer.int',sK13))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X5),s(bool,X6)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,X6)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X5) | s('type.integer.int',X4) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X7),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl104),
  inference(subsumption_resolution,[],[f8350,f524])).
fof(f8350,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'T_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X4),s('type.integer.int',sK13))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X5),s(bool,X6)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,X6)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X5) | s('type.integer.int',X4) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X7),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl104),
  inference(superposition,[],[f304,f7263])).
fof(f8300,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | ~sP0('const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),'const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),X4)) ) | $spl226),
  inference(cnf_transformation,[],[f8300_D])).
fof(f8300_D,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | ~sP0('const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),'const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),X4)) ) <=> ~$spl226),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl226])])).
fof(f8302,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl229),
  inference(cnf_transformation,[],[f8302_D])).
fof(f8302_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl229),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl229])])).
fof(f8185,plain,(
  ( ! [X4] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,X4)) | ~sP0('const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),'const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),X4)) ) | $spl0),
  inference(resolution,[],[f545,f301])).
fof(f301,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | ~p(s(bool,X2)) | ~sP0(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f196])).
fof(f8299,plain,(
  ( ! [X2,X1] : (~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),'const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),X1) | ~p(s(bool,X2))) ) | $spl224),
  inference(cnf_transformation,[],[f8299_D])).
fof(f8299_D,plain,(
  ( ! [X2,X1] : (~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),'const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),X1) | ~p(s(bool,X2))) ) <=> ~$spl224),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl224])])).
fof(f8298,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | ~p(s(bool,X2)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),'const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),X1)) ) | $spl0),
  inference(forward_demodulation,[],[f8297,f354])).
fof(f8297,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | ~p(s(bool,X2)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),'const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),X1) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl0),
  inference(subsumption_resolution,[],[f8184,f523])).
fof(f8184,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,X2)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),'const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),X1) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl0),
  inference(resolution,[],[f545,f304])).
fof(f545,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))))) | $spl0),
  inference(cnf_transformation,[],[f545_D])).
fof(f545_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f549,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl4),
  inference(cnf_transformation,[],[f549_D])).
fof(f549_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f8161,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))))) | $spl87),
  inference(forward_demodulation,[],[f3160,f353])).
fof(f3160,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))))) | $spl87),
  inference(subsumption_resolution,[],[f3136,f523])).
fof(f3136,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))))) | $spl87),
  inference(superposition,[],[f1666,f352])).
fof(f8183,plain,(
  ( ! [X10,X8] : (~p(s(bool,X8)) | ~p(s(bool,X10))) ) | ($spl4 | $spl96)),
  inference(duplicate_literal_removal,[],[f8182])).
fof(f8182,plain,(
  ( ! [X10,X8] : (~p(s(bool,X8)) | ~p(s(bool,X8)) | ~p(s(bool,X10))) ) | ($spl4 | $spl96)),
  inference(forward_demodulation,[],[f8181,f339])).
fof(f339,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f76])).
fof(f76,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X11] : ! [X9,X10] : s(X11,X9) = s(X11,'const.bool.COND_3'(s(bool,'T_0'),s(X11,X9),s(X11,X10)))),
  file('thm.integer.INT_MOD_COMMON_FACTOR','thm.bool.COND_CLAUSES|split|0')).
fof(f8181,plain,(
  ( ! [X10,X8] : (~p(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X8),s(bool,X10)))) | ~p(s(bool,X8)) | ~p(s(bool,X10))) ) | ($spl4 | $spl96)),
  inference(forward_demodulation,[],[f1685,f1575])).
fof(f1575,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl4),
  inference(resolution,[],[f549,f347])).
fof(f867,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl46),
  inference(cnf_transformation,[],[f867_D])).
fof(f867_D,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f7987,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))) | ~p(s(bool,X1)) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))) | p(s(bool,X2))) ) | ($spl1 | $spl4)),
  inference(forward_demodulation,[],[f7986,f354])).
fof(f7986,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,X1)) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))) | p(s(bool,X2)) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f7985,f1669])).
fof(f1669,plain,(
  ( ! [X4] : (~sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),'const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),X4) | ~p(s(bool,X4))) ) | $spl4),
  inference(subsumption_resolution,[],[f1553,f523])).
fof(f1553,plain,(
  ( ! [X4] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,X4)) | ~sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),'const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),X4)) ) | $spl4),
  inference(resolution,[],[f549,f301])).
fof(f7985,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,X1)) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))) | p(s(bool,X2)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),'const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),X1) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl1 | $spl4)),
  inference(forward_demodulation,[],[f7984,f339])).
fof(f7984,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))) | p(s(bool,X2)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),'const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),X1) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl1 | $spl4)),
  inference(forward_demodulation,[],[f7983,f1575])).
fof(f7983,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | p(s(bool,X2)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),'const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),X1) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl1),
  inference(subsumption_resolution,[],[f7710,f481])).
fof(f7710,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'F_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | p(s(bool,X2)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),'const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),X1) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl1),
  inference(superposition,[],[f303,f1728])).
fof(f1728,plain,(
  s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))) | $spl1),
  inference(resolution,[],[f544,f343])).
fof(f7611,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl223),
  inference(cnf_transformation,[],[f7611_D])).
fof(f7611_D,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl223),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl223])])).
fof(f7603,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,X9))) ) | ($spl107 | $spl130)),
  inference(resolution,[],[f6329,f4108])).
fof(f4108,plain,(
  ( ! [X14,X15] : (~sP3(X14,X15,'F_0') | s(bool,'T_0') != s(bool,X15) | ~p(s(bool,X14))) ) | $spl130),
  inference(subsumption_resolution,[],[f4107,f321])).
fof(f321,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f209])).
fof(f209,plain,(
  ! [X0,X1,X2] : ((~sP1(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP1(X0,X1,X2)))),
  inference(flattening,[],[f208])).
fof(f208,plain,(
  ! [X0,X1,X2] : ((~sP1(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP1(X0,X1,X2)))),
  inference(nnf_transformation,[],[f175])).
fof(f175,plain,(
  ! [X0,X1,X2] : (sP1(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f4107,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | ~sP3(X14,X15,'F_0') | ~sP1(X15,'F_0',X14)) ) | $spl130),
  inference(subsumption_resolution,[],[f4106,f460])).
fof(f4106,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | p(s(bool,'F_0')) | ~sP3(X14,X15,'F_0') | ~sP1(X15,'F_0',X14)) ) | $spl130),
  inference(subsumption_resolution,[],[f4099,f319])).
fof(f319,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f207])).
fof(f207,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(rectify,[],[f206])).
fof(f206,plain,(
  ! [X1,X0,X2] : ((~sP2(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP2(X1,X0,X2)))),
  inference(flattening,[],[f205])).
fof(f205,plain,(
  ! [X1,X0,X2] : ((~sP2(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP2(X1,X0,X2)))),
  inference(nnf_transformation,[],[f176])).
fof(f176,plain,(
  ! [X1,X0,X2] : (sP2(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f4099,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | p(s(bool,'F_0')) | ~sP3(X14,X15,'F_0') | ~sP2('F_0',X15,X14) | ~sP1(X15,'F_0',X14)) ) | $spl130),
  inference(resolution,[],[f3419,f310])).
fof(f310,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1)) | ~sP3(X0,X2,X1) | ~sP2(X1,X2,X0) | ~sP1(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f201])).
fof(f201,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | (sP1(X2,X1,X0) & sP2(X1,X2,X0) & sP3(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP1(X2,X1,X0) | ~sP2(X1,X2,X0) | ~sP3(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f200])).
fof(f200,plain,(
  ! [X2,X1,X0] : ((~sP4(X2,X1,X0) | (sP1(X0,X1,X2) & sP2(X1,X0,X2) & sP3(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP1(X0,X1,X2) | ~sP2(X1,X0,X2) | ~sP3(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP4(X2,X1,X0)))),
  inference(flattening,[],[f199])).
fof(f199,plain,(
  ! [X2,X1,X0] : ((~sP4(X2,X1,X0) | (sP1(X0,X1,X2) & sP2(X1,X0,X2) & sP3(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP1(X0,X1,X2) | ~sP2(X1,X0,X2) | ~sP3(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP4(X2,X1,X0)))),
  inference(nnf_transformation,[],[f178])).
fof(f178,plain,(
  ! [X2,X1,X0] : (sP4(X2,X1,X0) <=> (sP1(X0,X1,X2) & sP2(X1,X0,X2) & sP3(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f3419,plain,(
  ( ! [X10,X11] : (~sP4(X11,'F_0',X10) | ~p(s(bool,X11)) | s(bool,'T_0') != s(bool,X10)) ) | $spl130),
  inference(superposition,[],[f2855,f326])).
fof(f326,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X2)) | ~sP4(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f210])).
fof(f210,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP4(X2,X1,X0)) & (~sP4(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f179])).
fof(f179,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP4(X2,X1,X0))),
  inference(definition_folding,[],[f68,f178,f177,f176,f175])).
fof(f177,plain,(
  ! [X2,X0,X1] : (sP3(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f68,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f67])).
fof(f67,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f59])).
fof(f59,axiom,(
  ! [X26,X7,X6] : ((p(s(bool,X6)) <=> s(bool,X7) = s(bool,X26)) <=> ((p(s(bool,X6)) | p(s(bool,X7)) | p(s(bool,X26))) & (p(s(bool,X6)) | ~p(s(bool,X26)) | ~p(s(bool,X7))) & (p(s(bool,X7)) | ~p(s(bool,X26)) | ~p(s(bool,X6))) & (p(s(bool,X26)) | ~p(s(bool,X7)) | ~p(s(bool,X6)))))),
  file('thm.integer.INT_MOD_COMMON_FACTOR','thm.sat.sat_2|49|_prove_ltb')).
fof(f2855,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl130),
  inference(equality_resolution,[],[f2792])).
fof(f6329,plain,(
  ( ! [X28,X27] : (sP3(X27,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X28)) ) | $spl107),
  inference(resolution,[],[f1904,f314])).
fof(f314,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f204])).
fof(f204,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f203])).
fof(f203,plain,(
  ! [X2,X0,X1] : ((~sP3(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP3(X2,X0,X1)))),
  inference(flattening,[],[f202])).
fof(f202,plain,(
  ! [X2,X0,X1] : ((~sP3(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP3(X2,X0,X1)))),
  inference(nnf_transformation,[],[f177])).
fof(f4672,plain,(
  s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl166),
  inference(cnf_transformation,[],[f4672_D])).
fof(f4672_D,plain,(
  s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f7602,plain,(
  ( ! [X8] : (s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,X8))) ) | ($spl107 | $spl130)),
  inference(resolution,[],[f6329,f4053])).
fof(f4053,plain,(
  ( ! [X2,X3] : (~sP3(X2,X3,'F_0') | s(bool,'F_0') = s(bool,X3) | ~p(s(bool,X2))) ) | $spl130),
  inference(subsumption_resolution,[],[f4047,f321])).
fof(f4047,plain,(
  ( ! [X2,X3] : (~sP3(X2,X3,'F_0') | ~sP1(X3,'F_0',X2) | s(bool,'F_0') = s(bool,X3) | ~p(s(bool,X2))) ) | $spl130),
  inference(resolution,[],[f4008,f326])).
fof(f4008,plain,(
  ( ! [X4,X5] : (sP4(X4,'F_0',X5) | ~sP3(X4,X5,'F_0') | ~sP1(X5,'F_0',X4)) ) | $spl130),
  inference(subsumption_resolution,[],[f4006,f514])).
fof(f4006,plain,(
  ( ! [X4,X5] : (sP4(X4,'F_0',X5) | p(s(bool,'F_0')) | ~sP3(X4,X5,'F_0') | ~sP1(X5,'F_0',X4)) ) | $spl130),
  inference(resolution,[],[f4003,f310])).
fof(f4003,plain,(
  ( ! [X4,X5] : (sP2('F_0',X4,X5)) ) | $spl130),
  inference(subsumption_resolution,[],[f3993,f317])).
fof(f317,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f207])).
fof(f3993,plain,(
  ( ! [X4,X5] : (sP2('F_0',X4,X5) | p(s(bool,X5))) ) | $spl130),
  inference(resolution,[],[f3990,f315])).
fof(f315,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f204])).
fof(f3990,plain,(
  ( ! [X6,X7] : (~sP3(X7,X6,'F_0') | sP2('F_0',X6,X7)) ) | $spl130),
  inference(subsumption_resolution,[],[f3986,f318])).
fof(f318,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f207])).
fof(f3986,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | ~sP3(X7,X6,'F_0') | sP2('F_0',X6,X7)) ) | $spl130),
  inference(resolution,[],[f3980,f306])).
fof(f306,plain,(
  ( ! [X2,X0,X1] : (sP2(X1,X2,X0) | ~sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f201])).
fof(f3980,plain,(
  ( ! [X2,X3] : (sP4(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP3(X3,X2,'F_0')) ) | $spl130),
  inference(subsumption_resolution,[],[f3979,f347])).
fof(f3979,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP4(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP3(X3,X2,'F_0')) ) | $spl130),
  inference(subsumption_resolution,[],[f3976,f3954])).
fof(f3954,plain,(
  ( ! [X4,X5] : (sP1(X4,'F_0',X5) | s(bool,'T_0') != s(bool,X4)) ) | $spl130),
  inference(subsumption_resolution,[],[f3949,f321])).
fof(f3949,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | p(s(bool,X5)) | sP1(X4,'F_0',X5)) ) | $spl130),
  inference(resolution,[],[f3417,f305])).
fof(f305,plain,(
  ( ! [X2,X0,X1] : (sP1(X2,X1,X0) | ~sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f201])).
fof(f3417,plain,(
  ( ! [X6,X7] : (sP4(X7,'F_0',X6) | s(bool,'T_0') != s(bool,X6) | p(s(bool,X7))) ) | $spl130),
  inference(superposition,[],[f2855,f324])).
fof(f324,plain,(
  ( ! [X2,X0,X1] : (sP4(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f210])).
fof(f3976,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP4(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP3(X3,X2,'F_0') | ~sP1(X2,'F_0',X3)) ) | $spl130),
  inference(resolution,[],[f3955,f309])).
fof(f309,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2)) | ~sP3(X0,X2,X1) | ~sP2(X1,X2,X0) | ~sP1(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f201])).
fof(f3955,plain,(
  ( ! [X6,X7] : (sP2('F_0',X6,X7) | s(bool,'T_0') != s(bool,X6)) ) | $spl130),
  inference(subsumption_resolution,[],[f3950,f317])).
fof(f3950,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X6) | p(s(bool,X7)) | sP2('F_0',X6,X7)) ) | $spl130),
  inference(resolution,[],[f3417,f306])).
fof(f4163,plain,(
  s(bool,'F_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))) | $spl163),
  inference(cnf_transformation,[],[f4163_D])).
fof(f4163_D,plain,(
  s(bool,'F_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))) <=> ~$spl163),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl163])])).
fof(f7539,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,'F_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) ) | ($spl104 | $spl130)),
  inference(resolution,[],[f1921,f4054])).
fof(f4054,plain,(
  ( ! [X4,X5] : (~sP1(X5,'F_0',X4) | p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5)) ) | $spl130),
  inference(subsumption_resolution,[],[f4048,f315])).
fof(f4048,plain,(
  ( ! [X4,X5] : (~sP3(X4,X5,'F_0') | ~sP1(X5,'F_0',X4) | p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5)) ) | $spl130),
  inference(resolution,[],[f4008,f327])).
fof(f327,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | s(bool,X0) != s(bool,X1) | ~sP4(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f210])).
fof(f1921,plain,(
  ( ! [X35,X34] : (sP1('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X34,X35)) ) | $spl104),
  inference(resolution,[],[f1903,f323])).
fof(f323,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f209])).
fof(f7359,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | p(s(bool,X2)) | sP0(sK13,'const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),X1) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',sK13)))) ) | $spl107),
  inference(resolution,[],[f1904,f303])).
fof(f7353,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int',sK13))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int',sK13))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | ~p(s(bool,X2)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X1)) ) | $spl104),
  inference(forward_demodulation,[],[f7352,f354])).
fof(f7352,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int',sK13))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | ~p(s(bool,X2)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X1) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl104),
  inference(subsumption_resolution,[],[f7240,f523])).
fof(f7240,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int',sK13))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,X2)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X1) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl104),
  inference(resolution,[],[f1903,f304])).
fof(f1902,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | $spl105),
  inference(cnf_transformation,[],[f1902_D])).
fof(f1902_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) <=> ~$spl105),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl105])])).
fof(f6450,plain,(
  ( ! [X10,X8] : (p(s(bool,X8)) | ~p(s(bool,X10)) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X8),s(bool,X10))))) ) | $spl190),
  inference(cnf_transformation,[],[f6450_D])).
fof(f6450_D,plain,(
  ( ! [X10,X8] : (p(s(bool,X8)) | ~p(s(bool,X10)) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X8),s(bool,X10))))) ) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f7225,plain,(
  ( ! [X12,X10,X13,X11] : (s('type.integer.int',X11) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X13),s('type.integer.int',sK13))) | p(s(bool,X10)) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X11),s('type.integer.int',sK13))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X10),s(bool,X12)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | ~p(s(bool,X12))) ) | $spl107),
  inference(forward_demodulation,[],[f7224,f354])).
fof(f7224,plain,(
  ( ! [X12,X10,X13,X11] : (p(s(bool,X10)) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X11),s('type.integer.int',sK13))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X10),s(bool,X12)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | ~p(s(bool,X12)) | s('type.integer.int',X11) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X13),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl107),
  inference(subsumption_resolution,[],[f7221,f523])).
fof(f7221,plain,(
  ( ! [X12,X10,X13,X11] : (p(s(bool,X10)) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X11),s('type.integer.int',sK13))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X10),s(bool,X12)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,X12)) | s('type.integer.int',X11) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X13),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl107),
  inference(resolution,[],[f6320,f304])).
fof(f6320,plain,(
  ( ! [X4] : (~sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X4) | p(s(bool,X4))) ) | $spl107),
  inference(resolution,[],[f1904,f299])).
fof(f1903,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | $spl104),
  inference(cnf_transformation,[],[f1903_D])).
fof(f1903_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f7223,plain,(
  ( ! [X6,X8] : (p(s(bool,X6)) | p(s(bool,X8)) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X6),s(bool,X8))))) ) | $spl220),
  inference(cnf_transformation,[],[f7223_D])).
fof(f7223_D,plain,(
  ( ! [X6,X8] : (p(s(bool,X6)) | p(s(bool,X8)) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X6),s(bool,X8))))) ) <=> ~$spl220),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl220])])).
fof(f7222,plain,(
  ( ! [X6,X8,X7,X9] : (s('type.integer.int',X7) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X9),s('type.integer.int',sK13))) | p(s(bool,X6)) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X7),s('type.integer.int',sK13))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X6),s(bool,X8)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | p(s(bool,X8))) ) | $spl107),
  inference(forward_demodulation,[],[f7220,f354])).
fof(f7220,plain,(
  ( ! [X6,X8,X7,X9] : (p(s(bool,X6)) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X7),s('type.integer.int',sK13))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X6),s(bool,X8)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | p(s(bool,X8)) | s('type.integer.int',X7) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X9),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl107),
  inference(resolution,[],[f6320,f303])).
fof(f7215,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X1)) ) | $spl218),
  inference(cnf_transformation,[],[f7215_D])).
fof(f7215_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X1)) ) <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f7217,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X27),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X28),s('type.integer.int',sK13))) | ~p(s(bool,X29)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X29) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X27))))) ) | $spl178),
  inference(forward_demodulation,[],[f7216,f353])).
fof(f7216,plain,(
  ( ! [X28,X29,X27] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X28),s('type.integer.int',sK13))) | ~p(s(bool,X29)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X29) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X27)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X27))))) ) | $spl178),
  inference(forward_demodulation,[],[f6771,f354])).
fof(f6771,plain,(
  ( ! [X28,X29,X27] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X28),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,X29)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X29) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X27)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X27))))) ) | $spl178),
  inference(trivial_inequality_removal,[],[f6755])).
fof(f6755,plain,(
  ( ! [X28,X29,X27] : (s('type.integer.int',X27) != s('type.integer.int',X27) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X28),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,X29)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X29) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X27)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X27))))) ) | $spl178),
  inference(superposition,[],[f6006,f352])).
fof(f7214,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',sK13))) | ~p(s(bool,X1)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X1) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) ) | $spl178),
  inference(forward_demodulation,[],[f6770,f354])).
fof(f6770,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,X1)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X1) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) ) | $spl178),
  inference(equality_resolution,[],[f6006])).
fof(f7207,plain,(
  $false | $spl116),
  inference(equality_resolution,[],[f2336])).
fof(f7176,plain,(
  ( ! [X2,X0,X1] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',sK13))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X0),s(bool,X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | ~p(s(bool,X1)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X0)) )),
  inference(forward_demodulation,[],[f7175,f354])).
fof(f7175,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X0),s(bool,X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | ~p(s(bool,X1)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X0) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(subsumption_resolution,[],[f7157,f523])).
fof(f7157,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X0),s(bool,X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,X1)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X0) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(equality_resolution,[],[f537])).
fof(f7164,plain,(
  ( ! [X33,X31,X34,X32] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X31),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X34),s('type.integer.int',sK13))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X32),s(bool,X33)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | ~p(s(bool,X33)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X32) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X31))))) )),
  inference(forward_demodulation,[],[f7163,f353])).
fof(f7163,plain,(
  ( ! [X33,X31,X34,X32] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X34),s('type.integer.int',sK13))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X32),s(bool,X33)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | ~p(s(bool,X33)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X32) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X31)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X31))))) )),
  inference(forward_demodulation,[],[f7162,f354])).
fof(f7162,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X32),s(bool,X33)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | ~p(s(bool,X33)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X32) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X34),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X31)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X31))))) )),
  inference(subsumption_resolution,[],[f7158,f523])).
fof(f7158,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X32),s(bool,X33)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,X33)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X32) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X34),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X31)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X31))))) )),
  inference(trivial_inequality_removal,[],[f7142])).
fof(f7142,plain,(
  ( ! [X33,X31,X34,X32] : (s('type.integer.int',X31) != s('type.integer.int',X31) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X32),s(bool,X33)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,X33)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X32) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X34),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X31)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X31))))) )),
  inference(superposition,[],[f537,f352])).
fof(f4673,plain,(
  s(bool,'F_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl167),
  inference(cnf_transformation,[],[f4673_D])).
fof(f4673_D,plain,(
  s(bool,'F_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl167),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl167])])).
fof(f7104,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,'F_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl106 | $spl130)),
  inference(resolution,[],[f3577,f4054])).
fof(f3577,plain,(
  ( ! [X35,X34] : (sP1('const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X34,X35)) ) | $spl106),
  inference(resolution,[],[f1905,f323])).
fof(f7090,plain,(
  ( ! [X10,X9] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X9),s('type.integer.int',sK13))) | ~p(s(bool,X10))) ) | ($spl106 | $spl174 | $spl176 | $spl178)),
  inference(forward_demodulation,[],[f7087,f354])).
fof(f7087,plain,(
  ( ! [X10,X9] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X9),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,X10))) ) | ($spl106 | $spl174 | $spl176 | $spl178)),
  inference(trivial_inequality_removal,[],[f7086])).
fof(f7086,plain,(
  ( ! [X10,X9] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X9),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X10))) ) | ($spl106 | $spl174 | $spl176 | $spl178)),
  inference(duplicate_literal_removal,[],[f7085])).
fof(f7085,plain,(
  ( ! [X10,X9] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X9),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X10)) | ~p(s(bool,X10))) ) | ($spl106 | $spl174 | $spl176 | $spl178)),
  inference(resolution,[],[f6927,f3674])).
fof(f3674,plain,(
  ( ! [X4] : (~sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X4) | ~p(s(bool,X4))) ) | $spl106),
  inference(subsumption_resolution,[],[f3563,f523])).
fof(f3563,plain,(
  ( ! [X4] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,X4)) | ~sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X4)) ) | $spl106),
  inference(resolution,[],[f1905,f301])).
fof(f6927,plain,(
  ( ! [X8,X11,X9] : (sP0(X8,sK13,X9) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int',sK13))),s('type.integer.int',X8))) | s('type.integer.int',X8) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X9))) ) | ($spl174 | $spl176 | $spl178)),
  inference(subsumption_resolution,[],[f6926,f6004])).
fof(f6926,plain,(
  ( ! [X8,X11,X9] : (~p(s(bool,X9)) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int',sK13))),s('type.integer.int',X8))) | s('type.integer.int',X8) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | sP0(X8,sK13,X9) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X8))))) ) | ($spl174 | $spl178)),
  inference(subsumption_resolution,[],[f6002,f6006])).
fof(f7089,plain,(
  ( ! [X6,X7] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X6),s('type.integer.int',sK13))),s('type.integer.int',X7))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',X7)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X7) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl216),
  inference(cnf_transformation,[],[f7089_D])).
fof(f7089_D,plain,(
  ( ! [X6,X7] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X6),s('type.integer.int',sK13))),s('type.integer.int',X7))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',X7)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X7) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl216),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl216])])).
fof(f7088,plain,(
  ( ! [X6,X8,X7] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X6),s('type.integer.int',sK13))),s('type.integer.int',X7))) | s('type.integer.int',X7) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X8)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',X7))))) ) | ($spl174 | $spl176 | $spl178)),
  inference(duplicate_literal_removal,[],[f7084])).
fof(f7084,plain,(
  ( ! [X6,X8,X7] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X6),s('type.integer.int',sK13))),s('type.integer.int',X7))) | s('type.integer.int',X7) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X8)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',X7)))) | ~p(s(bool,X8))) ) | ($spl174 | $spl176 | $spl178)),
  inference(resolution,[],[f6927,f301])).
fof(f6879,plain,(
  ( ! [X4,X5] : (sP4(X4,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X5) | ~sP1(X5,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X4) | ~sP3(X4,X5,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) ) | $spl204),
  inference(cnf_transformation,[],[f6879_D])).
fof(f6879_D,plain,(
  ( ! [X4,X5] : (sP4(X4,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X5) | ~sP1(X5,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X4) | ~sP3(X4,X5,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) ) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f6874,plain,(
  ( ! [X4,X5] : (sP4(X4,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X5) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | ~sP3(X4,X5,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))) | ~sP1(X5,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X4)) ) | $spl105),
  inference(resolution,[],[f3449,f310])).
fof(f3449,plain,(
  ( ! [X35,X36] : (sP2('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X35,X36)) ) | $spl105),
  inference(resolution,[],[f1902,f319])).
fof(f6826,plain,(
  ( ! [X6,X7] : (sP4('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X6,X7) | ~sP1(X7,X6,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))) | ~sP2(X6,X7,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) ) | $spl198),
  inference(cnf_transformation,[],[f6826_D])).
fof(f6826_D,plain,(
  ( ! [X6,X7] : (sP4('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X6,X7) | ~sP1(X7,X6,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))) | ~sP2(X6,X7,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) ) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f6817,plain,(
  ( ! [X6,X7] : (sP4('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X6,X7) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | ~sP2(X6,X7,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))) | ~sP1(X7,X6,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) ) | $spl105),
  inference(resolution,[],[f3446,f311])).
fof(f311,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0)) | ~sP3(X0,X2,X1) | ~sP2(X1,X2,X0) | ~sP1(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f201])).
fof(f3446,plain,(
  ( ! [X30,X29] : (sP3('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X29,X30)) ) | $spl105),
  inference(resolution,[],[f1902,f315])).
fof(f4124,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))) | $spl159),
  inference(cnf_transformation,[],[f4124_D])).
fof(f4124_D,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))) <=> ~$spl159),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl159])])).
fof(f6805,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))) | ~p(s(bool,X9))) ) | ($spl105 | $spl130)),
  inference(resolution,[],[f3445,f4108])).
fof(f3445,plain,(
  ( ! [X28,X27] : (sP3(X27,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X28)) ) | $spl105),
  inference(resolution,[],[f1902,f314])).
fof(f4164,plain,(
  s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))) | $spl162),
  inference(cnf_transformation,[],[f4164_D])).
fof(f4164_D,plain,(
  s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f6804,plain,(
  ( ! [X8] : (s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))) | ~p(s(bool,X8))) ) | ($spl105 | $spl130)),
  inference(resolution,[],[f3445,f4053])).
fof(f6785,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',sK13))) | ~p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) ) | ($spl106 | $spl178)),
  inference(forward_demodulation,[],[f6784,f354])).
fof(f6784,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) ) | ($spl106 | $spl178)),
  inference(subsumption_resolution,[],[f6770,f3674])).
fof(f6776,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X27),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X28),s('type.integer.int',sK13))) | ~p(s(bool,X29)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X27))))) ) | ($spl106 | $spl178)),
  inference(forward_demodulation,[],[f6775,f353])).
fof(f6775,plain,(
  ( ! [X28,X29,X27] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X28),s('type.integer.int',sK13))) | ~p(s(bool,X29)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X27)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X27))))) ) | ($spl106 | $spl178)),
  inference(forward_demodulation,[],[f6774,f354])).
fof(f6774,plain,(
  ( ! [X28,X29,X27] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X28),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,X29)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X27)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X27))))) ) | ($spl106 | $spl178)),
  inference(subsumption_resolution,[],[f6771,f3674])).
fof(f6750,plain,(
  ( ! [X10,X8] : (~p(s(bool,X8)) | ~p(s(bool,X10)) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X8),s(bool,X10))))) ) | $spl194),
  inference(cnf_transformation,[],[f6750_D])).
fof(f6750_D,plain,(
  ( ! [X10,X8] : (~p(s(bool,X8)) | ~p(s(bool,X10)) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X8),s(bool,X10))))) ) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f6749,plain,(
  ( ! [X10,X8,X11,X9] : (s('type.integer.int',X9) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int',sK13))) | ~p(s(bool,X8)) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X9),s('type.integer.int',sK13))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X8),s(bool,X10)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | ~p(s(bool,X10))) ) | $spl106),
  inference(forward_demodulation,[],[f6748,f354])).
fof(f6748,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,X8)) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X9),s('type.integer.int',sK13))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X8),s(bool,X10)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | ~p(s(bool,X10)) | s('type.integer.int',X9) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl106),
  inference(subsumption_resolution,[],[f6746,f523])).
fof(f6746,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,X8)) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X9),s('type.integer.int',sK13))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X8),s(bool,X10)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,X10)) | s('type.integer.int',X9) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl106),
  inference(resolution,[],[f3674,f304])).
fof(f3686,plain,(
  ( ! [X6,X4] : (~p(s(bool,X4)) | p(s(bool,X6)) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X4),s(bool,X6))))) ) | $spl154),
  inference(cnf_transformation,[],[f3686_D])).
fof(f3686_D,plain,(
  ( ! [X6,X4] : (~p(s(bool,X4)) | p(s(bool,X6)) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X4),s(bool,X6))))) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f6747,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.integer.int',X5) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X7),s('type.integer.int',sK13))) | ~p(s(bool,X4)) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X5),s('type.integer.int',sK13))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X4),s(bool,X6)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | p(s(bool,X6))) ) | $spl106),
  inference(forward_demodulation,[],[f6745,f354])).
fof(f6745,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,X4)) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X5),s('type.integer.int',sK13))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X4),s(bool,X6)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | p(s(bool,X6)) | s('type.integer.int',X5) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X7),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl106),
  inference(resolution,[],[f3674,f303])).
fof(f1905,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl106),
  inference(cnf_transformation,[],[f1905_D])).
fof(f1905_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f6743,plain,(
  ( ! [X2,X1] : (~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | p(s(bool,X2))) ) | $spl192),
  inference(cnf_transformation,[],[f6743_D])).
fof(f6743_D,plain,(
  ( ! [X2,X1] : (~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | p(s(bool,X2))) ) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f6740,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X4),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X5),s(bool,X6)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X6)) | s('type.integer.int',X4) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',sK13)))) ) | ($spl105 | $spl112)),
  inference(resolution,[],[f6738,f303])).
fof(f6738,plain,(
  ( ! [X4] : (~sP0(sK13,'const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),X4)) ) | ($spl105 | $spl112)),
  inference(subsumption_resolution,[],[f3436,f2022])).
fof(f3436,plain,(
  ( ! [X4] : (~sP0(sK13,'const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),X4) | p(s(bool,X4))) ) | $spl105),
  inference(resolution,[],[f1902,f299])).
fof(f6742,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | p(s(bool,X2)) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',sK13)))) ) | ($spl105 | $spl112 | $spl121)),
  inference(subsumption_resolution,[],[f6739,f2354])).
fof(f6739,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | p(s(bool,X2)) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',sK13)))) ) | ($spl105 | $spl112)),
  inference(resolution,[],[f6738,f302])).
fof(f302,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.integer.int',X2) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int',X1))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X4),s(bool,X5)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X2)))) | p(s(bool,X5)) | sP0(X2,X1,X4) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int',X1))),s('type.integer.int',X2)))) )),
  inference(cnf_transformation,[],[f198])).
fof(f3193,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X0),s(bool,X1)))) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X0) | p(s(bool,X1))) ) | $spl152),
  inference(cnf_transformation,[],[f3193_D])).
fof(f3193_D,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X0),s(bool,X1)))) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X0) | p(s(bool,X1))) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f6737,plain,(
  ( ! [X33,X31,X34,X32] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X31),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X34),s('type.integer.int',sK13))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X32),s(bool,X33)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | p(s(bool,X33)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X32) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X31))))) )),
  inference(forward_demodulation,[],[f6736,f353])).
fof(f6736,plain,(
  ( ! [X33,X31,X34,X32] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X34),s('type.integer.int',sK13))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X32),s(bool,X33)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | p(s(bool,X33)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X32) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X31)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X31))))) )),
  inference(forward_demodulation,[],[f6309,f354])).
fof(f6309,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X32),s(bool,X33)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | p(s(bool,X33)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X32) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X34),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X31)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X31))))) )),
  inference(trivial_inequality_removal,[],[f6293])).
fof(f6293,plain,(
  ( ! [X33,X31,X34,X32] : (s('type.integer.int',X31) != s('type.integer.int',X31) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X32),s(bool,X33)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | p(s(bool,X33)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X32) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X34),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X31)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X31))))) )),
  inference(superposition,[],[f536,f352])).
fof(f536,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.integer.int',X4) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X5),s(bool,X6)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X4),s('type.integer.int',sK13)))) | p(s(bool,X6)) | sP0(X4,sK13,X5) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X7),s('type.integer.int',sK13))),s('type.integer.int',X4)))) )),
  inference(superposition,[],[f298,f303])).
fof(f6735,plain,(
  ( ! [X2,X0,X1] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',sK13))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X0),s(bool,X1)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | p(s(bool,X1)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X0)) )),
  inference(forward_demodulation,[],[f6308,f354])).
fof(f6308,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X0),s(bool,X1)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | p(s(bool,X1)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X0) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(equality_resolution,[],[f536])).
fof(f6716,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f6633,f867])).
fof(f6633,plain,(
  ~p(s(bool,'T_0')) | $spl46),
  inference(resolution,[],[f867,f518])).
fof(f6715,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f6632,f867])).
fof(f6632,plain,(
  ~p(s(bool,'T_0')) | $spl46),
  inference(resolution,[],[f867,f516])).
fof(f6714,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f6629,f867])).
fof(f6629,plain,(
  ( ! [X242] : (~p(s(bool,X242))) ) | $spl46),
  inference(resolution,[],[f867,f510])).
fof(f6713,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f6623,f867])).
fof(f6623,plain,(
  ( ! [X236] : (~p(s(bool,X236))) ) | $spl46),
  inference(resolution,[],[f867,f460])).
fof(f6631,plain,(
  $false | $spl46),
  inference(resolution,[],[f867,f514])).
fof(f6640,plain,(
  $false | ($spl46 | $spl87)),
  inference(resolution,[],[f867,f1666])).
fof(f6641,plain,(
  $false | ($spl11 | $spl46)),
  inference(resolution,[],[f867,f555])).
fof(f6642,plain,(
  $false | ($spl46 | $spl121)),
  inference(resolution,[],[f867,f2354])).
fof(f6645,plain,(
  $false | ($spl1 | $spl46)),
  inference(resolution,[],[f867,f544])).
fof(f6646,plain,(
  $false | ($spl46 | $spl107)),
  inference(resolution,[],[f867,f1904])).
fof(f6521,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,'F_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) ) | ($spl104 | $spl130)),
  inference(resolution,[],[f1921,f4054])).
fof(f6449,plain,(
  ( ! [X10,X8,X11,X9] : (s('type.integer.int',X9) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int',sK13))) | p(s(bool,X8)) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X9),s('type.integer.int',sK13))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X8),s(bool,X10)))) | ~p(s(bool,X10))) ) | ($spl104 | $spl107)),
  inference(forward_demodulation,[],[f6448,f354])).
fof(f6448,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X8)) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X9),s('type.integer.int',sK13))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X8),s(bool,X10)))) | ~p(s(bool,X10)) | s('type.integer.int',X9) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl104 | $spl107)),
  inference(subsumption_resolution,[],[f6447,f523])).
fof(f6447,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X8)) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X9),s('type.integer.int',sK13))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X8),s(bool,X10)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,X10)) | s('type.integer.int',X9) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl104 | $spl107)),
  inference(subsumption_resolution,[],[f6446,f1903])).
fof(f6446,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X8)) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X9),s('type.integer.int',sK13))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X8),s(bool,X10)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,X10)) | s('type.integer.int',X9) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl107),
  inference(resolution,[],[f6320,f304])).
fof(f6319,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | p(s(bool,X2)) | sP0(sK13,'const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),X1) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',sK13)))) ) | $spl107),
  inference(resolution,[],[f1904,f303])).
fof(f6283,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int',sK13))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int',sK13))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | ~p(s(bool,X2)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X1)) ) | $spl104),
  inference(forward_demodulation,[],[f6282,f354])).
fof(f6282,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int',sK13))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | ~p(s(bool,X2)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X1) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl104),
  inference(subsumption_resolution,[],[f6170,f523])).
fof(f6170,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int',sK13))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,X2)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X1) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl104),
  inference(resolution,[],[f1903,f304])).
fof(f6155,plain,(
  ( ! [X2,X0] : (~p(s(bool,X0)) | p(s(bool,X2)) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X0),s(bool,X2))))) ) | $spl186),
  inference(cnf_transformation,[],[f6155_D])).
fof(f6155_D,plain,(
  ( ! [X2,X0] : (~p(s(bool,X0)) | p(s(bool,X2)) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X0),s(bool,X2))))) ) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f6152,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,X4)) | s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X5),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X4),s(bool,X6)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X6)) | s('type.integer.int',X5) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',sK13)))) ) | $spl112),
  inference(resolution,[],[f2022,f303])).
fof(f6154,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,X0)) | s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X0),s(bool,X2)))) | p(s(bool,X2)) | s('type.integer.int',X1) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',sK13)))) ) | ($spl112 | $spl121)),
  inference(subsumption_resolution,[],[f6151,f2354])).
fof(f6151,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,X0)) | s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X0),s(bool,X2)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | p(s(bool,X2)) | s('type.integer.int',X1) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',sK13)))) ) | $spl112),
  inference(resolution,[],[f2022,f302])).
fof(f6021,plain,(
  ( ! [X14,X12,X10,X13] : (~p(s(bool,X10)) | s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X14),s('type.integer.int',sK13))),s('type.integer.int',X12))) | s('type.integer.int',X12) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X13),s('type.integer.int',sK13))) | sP0(X12,sK13,X10) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X12)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X12),s('type.integer.int',sK13))))) ) | $spl184),
  inference(cnf_transformation,[],[f6021_D])).
fof(f6021_D,plain,(
  ( ! [X14,X12,X10,X13] : (~p(s(bool,X10)) | s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X14),s('type.integer.int',sK13))),s('type.integer.int',X12))) | s('type.integer.int',X12) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X13),s('type.integer.int',sK13))) | sP0(X12,sK13,X10) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X12)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X12),s('type.integer.int',sK13))))) ) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f6020,plain,(
  ( ! [X14,X12,X10,X13,X11] : (~p(s(bool,X10)) | s('type.integer.int',X12) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X13),s('type.integer.int',sK13))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X12),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X12)))) | ~p(s(bool,X11)) | sP0(X12,sK13,X10) | s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X14),s('type.integer.int',sK13))),s('type.integer.int',X12)))) ) | $spl106),
  inference(forward_demodulation,[],[f5745,f339])).
fof(f5745,plain,(
  ( ! [X14,X12,X10,X13,X11] : (~p(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X10),s(bool,X11)))) | s('type.integer.int',X12) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X13),s('type.integer.int',sK13))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X12),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X12)))) | ~p(s(bool,X11)) | sP0(X12,sK13,X10) | s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X14),s('type.integer.int',sK13))),s('type.integer.int',X12)))) ) | $spl106),
  inference(superposition,[],[f304,f3585])).
fof(f3585,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl106),
  inference(resolution,[],[f1905,f347])).
fof(f6019,plain,(
  ( ! [X8,X7,X5,X9] : (~p(s(bool,X5)) | s('type.integer.int',X8) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X9),s('type.integer.int',sK13))),s('type.integer.int',X7))) | s('type.integer.int',X7) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X8),s('type.integer.int',sK13))) | sP0(X7,sK13,X5) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X7),s('type.integer.int',sK13))))) ) | $spl182),
  inference(cnf_transformation,[],[f6019_D])).
fof(f6019_D,plain,(
  ( ! [X8,X7,X5,X9] : (~p(s(bool,X5)) | s('type.integer.int',X8) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X9),s('type.integer.int',sK13))),s('type.integer.int',X7))) | s('type.integer.int',X7) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X8),s('type.integer.int',sK13))) | sP0(X7,sK13,X5) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X7),s('type.integer.int',sK13))))) ) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f6018,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~p(s(bool,X5)) | s('type.integer.int',X7) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X8),s('type.integer.int',sK13))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X7),s('type.integer.int',sK13)))) | p(s(bool,X6)) | sP0(X7,sK13,X5) | s('type.integer.int',X8) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X9),s('type.integer.int',sK13))),s('type.integer.int',X7)))) ) | $spl106),
  inference(forward_demodulation,[],[f5744,f339])).
fof(f5744,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~p(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X5),s(bool,X6)))) | s('type.integer.int',X7) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X8),s('type.integer.int',sK13))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X7),s('type.integer.int',sK13)))) | p(s(bool,X6)) | sP0(X7,sK13,X5) | s('type.integer.int',X8) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X9),s('type.integer.int',sK13))),s('type.integer.int',X7)))) ) | $spl106),
  inference(superposition,[],[f303,f3585])).
fof(f6017,plain,(
  ( ! [X4,X2,X0,X3] : (~p(s(bool,X0)) | s('type.integer.int',X3) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X4),s('type.integer.int',sK13))),s('type.integer.int',X2))) | s('type.integer.int',X2) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X3),s('type.integer.int',sK13))) | sP0(X2,sK13,X0) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X2))))) ) | $spl180),
  inference(cnf_transformation,[],[f6017_D])).
fof(f6017_D,plain,(
  ( ! [X4,X2,X0,X3] : (~p(s(bool,X0)) | s('type.integer.int',X3) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X4),s('type.integer.int',sK13))),s('type.integer.int',X2))) | s('type.integer.int',X2) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X3),s('type.integer.int',sK13))) | sP0(X2,sK13,X0) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X2))))) ) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f6016,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,X0)) | s('type.integer.int',X2) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X3),s('type.integer.int',sK13))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X2)))) | p(s(bool,X1)) | sP0(X2,sK13,X0) | s('type.integer.int',X3) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X4),s('type.integer.int',sK13))),s('type.integer.int',X2)))) ) | $spl106),
  inference(forward_demodulation,[],[f5743,f339])).
fof(f5743,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)))) | s('type.integer.int',X2) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X3),s('type.integer.int',sK13))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X2)))) | p(s(bool,X1)) | sP0(X2,sK13,X0) | s('type.integer.int',X3) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X4),s('type.integer.int',sK13))),s('type.integer.int',X2)))) ) | $spl106),
  inference(superposition,[],[f302,f3585])).
fof(f6006,plain,(
  ( ! [X4,X7,X5] : (s('type.integer.int',X4) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X7),s('type.integer.int',sK13))),s('type.integer.int',X4))) | ~p(s(bool,X5)) | sP0(X4,sK13,X5) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X4),s('type.integer.int',sK13))))) ) | $spl178),
  inference(cnf_transformation,[],[f6006_D])).
fof(f6006_D,plain,(
  ( ! [X4,X7,X5] : (s('type.integer.int',X4) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X7),s('type.integer.int',sK13))),s('type.integer.int',X4))) | ~p(s(bool,X5)) | sP0(X4,sK13,X5) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X4),s('type.integer.int',sK13))))) ) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f6005,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,X5)) | s('type.integer.int',X4) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X4),s('type.integer.int',sK13)))) | p(s(bool,X6)) | sP0(X4,sK13,X5) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X7),s('type.integer.int',sK13))),s('type.integer.int',X4)))) ) | $spl106),
  inference(forward_demodulation,[],[f5728,f339])).
fof(f5728,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X5),s(bool,X6)))) | s('type.integer.int',X4) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X4),s('type.integer.int',sK13)))) | p(s(bool,X6)) | sP0(X4,sK13,X5) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X7),s('type.integer.int',sK13))),s('type.integer.int',X4)))) ) | $spl106),
  inference(backward_demodulation,[],[f3585,f536])).
fof(f6004,plain,(
  ( ! [X0,X3,X1] : (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int',sK13))),s('type.integer.int',X0))) | ~p(s(bool,X1)) | sP0(X0,sK13,X1) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))))) ) | $spl176),
  inference(cnf_transformation,[],[f6004_D])).
fof(f6004_D,plain,(
  ( ! [X0,X3,X1] : (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int',sK13))),s('type.integer.int',X0))) | ~p(s(bool,X1)) | sP0(X0,sK13,X1) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))))) ) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f6003,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,X1)) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | p(s(bool,X2)) | sP0(X0,sK13,X1) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int',sK13))),s('type.integer.int',X0)))) ) | $spl106),
  inference(forward_demodulation,[],[f5727,f339])).
fof(f5727,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2)))) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | p(s(bool,X2)) | sP0(X0,sK13,X1) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int',sK13))),s('type.integer.int',X0)))) ) | $spl106),
  inference(backward_demodulation,[],[f3585,f535])).
fof(f535,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | p(s(bool,X2)) | sP0(X0,sK13,X1) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int',sK13))),s('type.integer.int',X0)))) )),
  inference(superposition,[],[f298,f302])).
fof(f6002,plain,(
  ( ! [X8,X11,X9] : (~p(s(bool,X9)) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int',sK13))),s('type.integer.int',X8))) | s('type.integer.int',X8) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | sP0(X8,sK13,X9) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X8)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X8),s('type.integer.int',sK13))))) ) | $spl174),
  inference(cnf_transformation,[],[f6002_D])).
fof(f6002_D,plain,(
  ( ! [X8,X11,X9] : (~p(s(bool,X9)) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int',sK13))),s('type.integer.int',X8))) | s('type.integer.int',X8) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | sP0(X8,sK13,X9) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X8)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X8),s('type.integer.int',sK13))))) ) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f6001,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,X9)) | s('type.integer.int',X8) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X8),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X8)))) | ~p(s(bool,X10)) | sP0(X8,sK13,X9) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int',sK13))),s('type.integer.int',X8)))) ) | $spl106),
  inference(forward_demodulation,[],[f5726,f339])).
fof(f5726,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X9),s(bool,X10)))) | s('type.integer.int',X8) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X8),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X8)))) | ~p(s(bool,X10)) | sP0(X8,sK13,X9) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int',sK13))),s('type.integer.int',X8)))) ) | $spl106),
  inference(backward_demodulation,[],[f3585,f537])).
fof(f5463,plain,(
  ( ! [X14,X12,X10,X13] : (~p(s(bool,X10)) | s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X14),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))),s('type.integer.int',X12))) | s('type.integer.int',X12) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))) | sP0(X12,'const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),X10) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X12)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X12),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))))) ) | $spl172),
  inference(cnf_transformation,[],[f5463_D])).
fof(f5463_D,plain,(
  ( ! [X14,X12,X10,X13] : (~p(s(bool,X10)) | s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X14),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))),s('type.integer.int',X12))) | s('type.integer.int',X12) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))) | sP0(X12,'const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),X10) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X12)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X12),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))))) ) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f5462,plain,(
  ( ! [X14,X12,X10,X13,X11] : (~p(s(bool,X10)) | s('type.integer.int',X12) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X12),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X12)))) | ~p(s(bool,X11)) | sP0(X12,'const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),X10) | s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X14),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))),s('type.integer.int',X12)))) ) | $spl4),
  inference(forward_demodulation,[],[f5190,f339])).
fof(f5190,plain,(
  ( ! [X14,X12,X10,X13,X11] : (~p(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X10),s(bool,X11)))) | s('type.integer.int',X12) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X12),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X12)))) | ~p(s(bool,X11)) | sP0(X12,'const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),X10) | s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X14),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))),s('type.integer.int',X12)))) ) | $spl4),
  inference(superposition,[],[f304,f1575])).
fof(f5461,plain,(
  ( ! [X8,X7,X5,X9] : (~p(s(bool,X5)) | s('type.integer.int',X8) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))),s('type.integer.int',X7))) | s('type.integer.int',X7) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X8),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))) | sP0(X7,'const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),X5) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))))) ) | $spl170),
  inference(cnf_transformation,[],[f5461_D])).
fof(f5461_D,plain,(
  ( ! [X8,X7,X5,X9] : (~p(s(bool,X5)) | s('type.integer.int',X8) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))),s('type.integer.int',X7))) | s('type.integer.int',X7) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X8),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))) | sP0(X7,'const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),X5) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))))) ) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f5460,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~p(s(bool,X5)) | s('type.integer.int',X7) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X8),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))))) | p(s(bool,X6)) | sP0(X7,'const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),X5) | s('type.integer.int',X8) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))),s('type.integer.int',X7)))) ) | $spl4),
  inference(forward_demodulation,[],[f5189,f339])).
fof(f5189,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~p(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X5),s(bool,X6)))) | s('type.integer.int',X7) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X8),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))))) | p(s(bool,X6)) | sP0(X7,'const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),X5) | s('type.integer.int',X8) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))),s('type.integer.int',X7)))) ) | $spl4),
  inference(superposition,[],[f303,f1575])).
fof(f5459,plain,(
  ( ! [X4,X2,X0,X3] : (~p(s(bool,X0)) | s('type.integer.int',X3) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X4),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))),s('type.integer.int',X2))) | s('type.integer.int',X2) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))) | sP0(X2,'const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),X0) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X2))))) ) | $spl168),
  inference(cnf_transformation,[],[f5459_D])).
fof(f5459_D,plain,(
  ( ! [X4,X2,X0,X3] : (~p(s(bool,X0)) | s('type.integer.int',X3) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X4),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))),s('type.integer.int',X2))) | s('type.integer.int',X2) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))) | sP0(X2,'const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),X0) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X2))))) ) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f5458,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,X0)) | s('type.integer.int',X2) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X2)))) | p(s(bool,X1)) | sP0(X2,'const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),X0) | s('type.integer.int',X3) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X4),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))),s('type.integer.int',X2)))) ) | $spl4),
  inference(forward_demodulation,[],[f5188,f339])).
fof(f5188,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)))) | s('type.integer.int',X2) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X2)))) | p(s(bool,X1)) | sP0(X2,'const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),X0) | s('type.integer.int',X3) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X4),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))),s('type.integer.int',X2)))) ) | $spl4),
  inference(superposition,[],[f302,f1575])).
fof(f4962,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int',sK13))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int',sK13))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | p(s(bool,X2)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X1)) ) | $spl105),
  inference(forward_demodulation,[],[f4961,f354])).
fof(f4961,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int',sK13))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | p(s(bool,X2)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X1) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl105),
  inference(subsumption_resolution,[],[f4690,f481])).
fof(f4690,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'F_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int',sK13))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | p(s(bool,X2)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X1) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl105),
  inference(superposition,[],[f303,f3458])).
fof(f3458,plain,(
  s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))) | $spl105),
  inference(resolution,[],[f1902,f343])).
fof(f4669,plain,(
  ( ! [X16] : (p(s(bool,X16)) | s(bool,'F_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl106 | $spl130)),
  inference(resolution,[],[f4054,f3577])).
fof(f3175,plain,(
  s(bool,'F_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl143),
  inference(cnf_transformation,[],[f3175_D])).
fof(f3175_D,plain,(
  s(bool,'F_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl143),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl143])])).
fof(f4667,plain,(
  ( ! [X14] : (p(s(bool,X14)) | s(bool,'F_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl114 | $spl130)),
  inference(resolution,[],[f4054,f4184])).
fof(f4184,plain,(
  ( ! [X31,X32] : (sP1('const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X31,X32)) ) | $spl114),
  inference(resolution,[],[f2023,f323])).
fof(f2023,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl114),
  inference(cnf_transformation,[],[f2023_D])).
fof(f2023_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f575,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) | $spl21),
  inference(cnf_transformation,[],[f575_D])).
fof(f575_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f1278,plain,(
  ( ! [X26,X24,X23,X25,X22] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X22),s('type.integer.int',X23))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X22) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X26),s('type.integer.int',X23))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X23,X24) | ~p(s(bool,X25)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))),s('type.integer.int',X23)))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X23),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X24),s(bool,X25))))) ) | $spl78),
  inference(cnf_transformation,[],[f1278_D])).
fof(f1278_D,plain,(
  ( ! [X26,X24,X23,X25,X22] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X22),s('type.integer.int',X23))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X22) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X26),s('type.integer.int',X23))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X23,X24) | ~p(s(bool,X25)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))),s('type.integer.int',X23)))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X23),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X24),s(bool,X25))))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f4392,plain,(
  ( ! [X26,X24,X23,X25,X22] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X22),s('type.integer.int',X23))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X23),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X24),s(bool,X25)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))),s('type.integer.int',X23)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) | ~p(s(bool,X25)) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X23,X24) | s('type.integer.int',X22) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X26),s('type.integer.int',X23))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) ) | $spl18),
  inference(superposition,[],[f574,f304])).
fof(f576,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) | $spl20),
  inference(cnf_transformation,[],[f576_D])).
fof(f576_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f1277,plain,(
  ( ! [X14,X12,X15,X13,X16] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X12),s('type.integer.int',X13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X16),s('type.integer.int',X13))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X13,X14) | p(s(bool,X15)) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X14),s(bool,X15))))) ) | $spl76),
  inference(cnf_transformation,[],[f1277_D])).
fof(f1277_D,plain,(
  ( ! [X14,X12,X15,X13,X16] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X12),s('type.integer.int',X13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X16),s('type.integer.int',X13))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X13,X14) | p(s(bool,X15)) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X14),s(bool,X15))))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f4390,plain,(
  ( ! [X14,X12,X15,X13,X16] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X12),s('type.integer.int',X13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X14),s(bool,X15)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) | p(s(bool,X15)) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X13,X14) | s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X16),s('type.integer.int',X13))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) ) | $spl18),
  inference(superposition,[],[f574,f302])).
fof(f1273,plain,(
  ( ! [X2,X1] : (~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),'const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),X1) | ~p(s(bool,X2))) ) | $spl72),
  inference(cnf_transformation,[],[f1273_D])).
fof(f1273_D,plain,(
  ( ! [X2,X1] : (~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),'const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),X1) | ~p(s(bool,X2))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f1274,plain,(
  ( ! [X0,X3] : (s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) ) | $spl74),
  inference(cnf_transformation,[],[f1274_D])).
fof(f1274_D,plain,(
  ( ! [X0,X3] : (s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f4294,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) | ~p(s(bool,X2)) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),'const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),X1) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) ) | $spl18),
  inference(resolution,[],[f574,f304])).
fof(f1838,plain,(
  ( ! [X2,X0] : (p(s(bool,X0)) | p(s(bool,X2)) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X0),s(bool,X2))))) ) | $spl100),
  inference(cnf_transformation,[],[f1838_D])).
fof(f1838_D,plain,(
  ( ! [X2,X0] : (p(s(bool,X0)) | p(s(bool,X2)) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X0),s(bool,X2))))) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f4168,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X0)) | s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X0),s(bool,X2)))) | p(s(bool,X2)) | s('type.integer.int',X1) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',sK13)))) ) | ($spl105 | $spl121)),
  inference(subsumption_resolution,[],[f4165,f2354])).
fof(f4165,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X0)) | s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X0),s(bool,X2)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | p(s(bool,X2)) | s('type.integer.int',X1) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',sK13)))) ) | $spl105),
  inference(resolution,[],[f3436,f302])).
fof(f4156,plain,(
  ( ! [X11] : (s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))) | ~p(s(bool,X11))) ) | ($spl105 | $spl130)),
  inference(resolution,[],[f4053,f3445])).
fof(f4162,plain,(
  s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))) | $spl160),
  inference(cnf_transformation,[],[f4162_D])).
fof(f4162_D,plain,(
  s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f4155,plain,(
  ( ! [X10] : (s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))) | ~p(s(bool,X10))) ) | ($spl1 | $spl130)),
  inference(resolution,[],[f4053,f1715])).
fof(f1715,plain,(
  ( ! [X28,X27] : (sP3(X27,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))),X28)) ) | $spl1),
  inference(resolution,[],[f544,f314])).
fof(f3178,plain,(
  s(bool,'F_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))) | $spl144),
  inference(cnf_transformation,[],[f3178_D])).
fof(f3178_D,plain,(
  s(bool,'F_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f4154,plain,(
  ( ! [X9] : (s(bool,'F_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))) | ~p(s(bool,X9))) ) | ($spl121 | $spl130)),
  inference(resolution,[],[f4053,f2505])).
fof(f2505,plain,(
  ( ! [X28,X27] : (sP3(X27,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X28)) ) | $spl121),
  inference(resolution,[],[f2354,f314])).
fof(f3176,plain,(
  s(bool,'F_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl142),
  inference(cnf_transformation,[],[f3176_D])).
fof(f3176_D,plain,(
  s(bool,'F_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f4153,plain,(
  ( ! [X8] : (s(bool,'F_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,X8))) ) | ($spl115 | $spl130)),
  inference(resolution,[],[f4053,f2194])).
fof(f2194,plain,(
  ( ! [X24,X25] : (sP3(X24,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X25)) ) | $spl115),
  inference(resolution,[],[f2024,f314])).
fof(f4116,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))) | ~p(s(bool,X11))) ) | ($spl105 | $spl130)),
  inference(resolution,[],[f4108,f3445])).
fof(f4122,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))) | $spl157),
  inference(cnf_transformation,[],[f4122_D])).
fof(f4122_D,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))) <=> ~$spl157),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl157])])).
fof(f4115,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))) | ~p(s(bool,X10))) ) | ($spl1 | $spl130)),
  inference(resolution,[],[f4108,f1715])).
fof(f3029,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))) | $spl137),
  inference(cnf_transformation,[],[f3029_D])).
fof(f3029_D,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))) <=> ~$spl137),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl137])])).
fof(f4114,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))) | ~p(s(bool,X9))) ) | ($spl121 | $spl130)),
  inference(resolution,[],[f4108,f2505])).
fof(f3027,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl135),
  inference(cnf_transformation,[],[f3027_D])).
fof(f3027_D,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl135),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl135])])).
fof(f4113,plain,(
  ( ! [X8] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,X8))) ) | ($spl115 | $spl130)),
  inference(resolution,[],[f4108,f2194])).
fof(f3685,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.integer.int',X5) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X7),s('type.integer.int',sK13))) | ~p(s(bool,X4)) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X5),s('type.integer.int',sK13))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X4),s(bool,X6)))) | p(s(bool,X6))) ) | ($spl105 | $spl106)),
  inference(forward_demodulation,[],[f3684,f354])).
fof(f3684,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,X4)) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X5),s('type.integer.int',sK13))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X4),s(bool,X6)))) | p(s(bool,X6)) | s('type.integer.int',X5) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X7),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl105 | $spl106)),
  inference(subsumption_resolution,[],[f3682,f1902])).
fof(f3682,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,X4)) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X5),s('type.integer.int',sK13))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X4),s(bool,X6)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | p(s(bool,X6)) | s('type.integer.int',X5) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X7),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl106),
  inference(resolution,[],[f3674,f303])).
fof(f3552,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int',sK13))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int',sK13))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | p(s(bool,X2)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X1)) ) | $spl105),
  inference(forward_demodulation,[],[f3435,f354])).
fof(f3435,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int',sK13))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | p(s(bool,X2)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X1) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl105),
  inference(resolution,[],[f1902,f303])).
fof(f3403,plain,(
  $false | $spl132),
  inference(subsumption_resolution,[],[f3302,f3025])).
fof(f3302,plain,(
  p(s(bool,'F_0')) | $spl132),
  inference(resolution,[],[f3025,f517])).
fof(f517,plain,(
  p(s(bool,'F_0')) | p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f293])).
fof(f3402,plain,(
  $false | $spl132),
  inference(subsumption_resolution,[],[f3301,f3025])).
fof(f3301,plain,(
  p(s(bool,'F_0')) | $spl132),
  inference(resolution,[],[f3025,f515])).
fof(f515,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f292])).
fof(f3401,plain,(
  $false | $spl132),
  inference(subsumption_resolution,[],[f3296,f3025])).
fof(f3296,plain,(
  ( ! [X241] : (p(s(bool,X241))) ) | $spl132),
  inference(resolution,[],[f3025,f471])).
fof(f471,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f271])).
fof(f271,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,'T_0'))))),
  inference(flattening,[],[f270])).
fof(f270,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f164])).
fof(f164,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X8] : ((p(s(bool,X8)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  file('thm.integer.INT_MOD_COMMON_FACTOR','thm.bool.IMP_CLAUSES|split|1')).
fof(f3400,plain,(
  $false | $spl132),
  inference(subsumption_resolution,[],[f3294,f3025])).
fof(f3294,plain,(
  p(s(bool,'F_0')) | $spl132),
  inference(resolution,[],[f3025,f468])).
fof(f468,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f269])).
fof(f269,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f268])).
fof(f268,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f163])).
fof(f163,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f117])).
fof(f117,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X8] : ((p(s(bool,'F_0')) => p(s(bool,X8))) <=> p(s(bool,'T_0')))),
  file('thm.integer.INT_MOD_COMMON_FACTOR','thm.bool.IMP_CLAUSES|split|2')).
fof(f3399,plain,(
  $false | $spl132),
  inference(subsumption_resolution,[],[f3292,f3025])).
fof(f3292,plain,(
  ( ! [X238] : (p(s(bool,X238))) ) | $spl132),
  inference(resolution,[],[f3025,f465])).
fof(f465,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f267])).
fof(f267,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))))),
  inference(flattening,[],[f266])).
fof(f266,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f162])).
fof(f162,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f116])).
fof(f116,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X8] : ((p(s(bool,X8)) => p(s(bool,X8))) <=> p(s(bool,'T_0')))),
  file('thm.integer.INT_MOD_COMMON_FACTOR','thm.bool.IMP_CLAUSES|split|3')).
fof(f3398,plain,(
  $false | $spl132),
  inference(subsumption_resolution,[],[f3291,f3025])).
fof(f3291,plain,(
  ( ! [X237] : (p(s(bool,X237))) ) | $spl132),
  inference(resolution,[],[f3025,f462])).
fof(f462,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f265])).
fof(f265,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f264])).
fof(f264,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (~p(s(bool,'T_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f161])).
fof(f161,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(ennf_transformation,[],[f115])).
fof(f115,plain,(
  ! [X0] : ((p(s(bool,'T_0')) => p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X8] : ((p(s(bool,'T_0')) => p(s(bool,X8))) <=> p(s(bool,X8)))),
  file('thm.integer.INT_MOD_COMMON_FACTOR','thm.bool.IMP_CLAUSES|split|0')).
fof(f3303,plain,(
  $false | $spl132),
  inference(resolution,[],[f3025,f521])).
fof(f521,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f9])).
fof(f9,axiom,(
  p(s(bool,'T_0'))),
  file('thm.integer.INT_MOD_COMMON_FACTOR','thm.bool.TRUTH')).
fof(f3304,plain,(
  $false | $spl132),
  inference(resolution,[],[f3025,f522])).
fof(f522,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.integer.INT_MOD_COMMON_FACTOR',aHL_TRUTH)).
fof(f3305,plain,(
  $false | $spl132),
  inference(resolution,[],[f3025,f524])).
fof(f3314,plain,(
  $false | $spl132),
  inference(resolution,[],[f3025,f523])).
fof(f3319,plain,(
  $false | ($spl6 | $spl132)),
  inference(resolution,[],[f3025,f552])).
fof(f3320,plain,(
  $false | ($spl16 | $spl132)),
  inference(resolution,[],[f3025,f572])).
fof(f3322,plain,(
  $false | ($spl4 | $spl132)),
  inference(resolution,[],[f3025,f549])).
fof(f3195,plain,(
  ( ! [X33,X31,X34,X32] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X31),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X34),s('type.integer.int',sK13))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X32),s(bool,X33)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | p(s(bool,X33)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X32) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X31))))) )),
  inference(forward_demodulation,[],[f3194,f353])).
fof(f3194,plain,(
  ( ! [X33,X31,X34,X32] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X34),s('type.integer.int',sK13))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X32),s(bool,X33)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | p(s(bool,X33)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X32) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X31)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X31))))) )),
  inference(forward_demodulation,[],[f2065,f354])).
fof(f2065,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X32),s(bool,X33)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | p(s(bool,X33)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X32) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X34),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X31)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X31))))) )),
  inference(trivial_inequality_removal,[],[f2049])).
fof(f2049,plain,(
  ( ! [X33,X31,X34,X32] : (s('type.integer.int',X31) != s('type.integer.int',X31) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X32),s(bool,X33)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | p(s(bool,X33)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X32) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X34),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X31)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X31))))) )),
  inference(superposition,[],[f536,f352])).
fof(f3192,plain,(
  ( ! [X2,X0,X1] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',sK13))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X0),s(bool,X1)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | p(s(bool,X1)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X0)) )),
  inference(forward_demodulation,[],[f2064,f354])).
fof(f2064,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X0),s(bool,X1)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | p(s(bool,X1)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X0) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(equality_resolution,[],[f536])).
fof(f3191,plain,(
  ( ! [X33,X31,X34,X32] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X31),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X34),s('type.integer.int',sK13))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X32),s(bool,X33)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | ~p(s(bool,X33)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X32) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X31))))) )),
  inference(forward_demodulation,[],[f3190,f353])).
fof(f3190,plain,(
  ( ! [X33,X31,X34,X32] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X34),s('type.integer.int',sK13))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X32),s(bool,X33)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | ~p(s(bool,X33)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X32) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X31)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X31))))) )),
  inference(forward_demodulation,[],[f3189,f354])).
fof(f3189,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X32),s(bool,X33)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | ~p(s(bool,X33)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X32) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X34),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X31)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X31))))) )),
  inference(subsumption_resolution,[],[f2328,f523])).
fof(f2328,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X32),s(bool,X33)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,X33)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X32) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X34),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X31)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X31))))) )),
  inference(trivial_inequality_removal,[],[f2312])).
fof(f2312,plain,(
  ( ! [X33,X31,X34,X32] : (s('type.integer.int',X31) != s('type.integer.int',X31) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X32),s(bool,X33)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,X33)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X32) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X34),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X31)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X31))))) )),
  inference(superposition,[],[f537,f352])).
fof(f3188,plain,(
  ( ! [X2,X0,X1] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',sK13))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X0),s(bool,X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | ~p(s(bool,X1)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X0)) )),
  inference(forward_demodulation,[],[f3187,f354])).
fof(f3187,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X0),s(bool,X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | ~p(s(bool,X1)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X0) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(subsumption_resolution,[],[f2327,f523])).
fof(f2327,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X0),s(bool,X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,X1)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X0) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(equality_resolution,[],[f537])).
fof(f3180,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl147),
  inference(cnf_transformation,[],[f3180_D])).
fof(f3180_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl147),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl147])])).
fof(f3186,plain,(
  ( ! [X22] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X22))) ) | $spl130),
  inference(subsumption_resolution,[],[f2836,f347])).
fof(f2836,plain,(
  ( ! [X22] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,X22) | ~p(s(bool,X22))) ) | $spl130),
  inference(superposition,[],[f2792,f347])).
fof(f3185,plain,(
  ( ! [X23] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X23))) ) | $spl130),
  inference(subsumption_resolution,[],[f2837,f347])).
fof(f2837,plain,(
  ( ! [X23] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,X23) | ~p(s(bool,X23))) ) | $spl130),
  inference(superposition,[],[f2792,f349])).
fof(f349,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f216])).
fof(f216,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X8] : (s(bool,'T_0') = s(bool,X8) <=> p(s(bool,X8)))),
  file('thm.integer.INT_MOD_COMMON_FACTOR','thm.bool.EQ_CLAUSES|split|0')).
fof(f3183,plain,(
  ( ! [X26] : (s(bool,'T_0') != s(bool,X26)) ) | $spl150),
  inference(cnf_transformation,[],[f3183_D])).
fof(f3183_D,plain,(
  ( ! [X26] : (s(bool,'T_0') != s(bool,X26)) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f3184,plain,(
  ( ! [X24] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,X24)) ) | $spl130),
  inference(subsumption_resolution,[],[f2838,f2792])).
fof(f2838,plain,(
  ( ! [X24] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,X24) | s(bool,'F_0') = s(bool,X24)) ) | $spl130),
  inference(superposition,[],[f2792,f350])).
fof(f350,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f86])).
fof(f86,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X8] : (s(bool,'T_0') = s(bool,X8) | s(bool,'F_0') = s(bool,X8))),
  file('thm.integer.INT_MOD_COMMON_FACTOR','thm.bool.BOOL_CASES_AX')).
fof(f3182,plain,(
  ( ! [X26] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,X26)) ) | $spl130),
  inference(subsumption_resolution,[],[f2840,f2792])).
fof(f2840,plain,(
  ( ! [X26] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,X26) | s(bool,'F_0') = s(bool,X26)) ) | $spl130),
  inference(superposition,[],[f2792,f351])).
fof(f351,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.integer.INT_MOD_COMMON_FACTOR',aHL_BOOL_CASES)).
fof(f3181,plain,(
  ( ! [X33] : (s(bool,'T_0') = s(bool,X33)) ) | $spl148),
  inference(cnf_transformation,[],[f3181_D])).
fof(f3181_D,plain,(
  ( ! [X33] : (s(bool,'T_0') = s(bool,X33)) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f2859,plain,(
  ( ! [X32] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X32)) ) | $spl130),
  inference(trivial_inequality_removal,[],[f2842])).
fof(f2842,plain,(
  ( ! [X32] : (s(bool,X32) != s(bool,X32) | s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X32)) ) | $spl130),
  inference(superposition,[],[f2792,f351])).
fof(f2858,plain,(
  ( ! [X33] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X33)) ) | $spl130),
  inference(trivial_inequality_removal,[],[f2843])).
fof(f2843,plain,(
  ( ! [X33] : (s(bool,X33) != s(bool,X33) | s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X33)) ) | $spl130),
  inference(superposition,[],[f2792,f350])).
fof(f2857,plain,(
  ( ! [X34] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X34))) ) | $spl130),
  inference(trivial_inequality_removal,[],[f2844])).
fof(f2844,plain,(
  ( ! [X34] : (s(bool,X34) != s(bool,X34) | s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X34))) ) | $spl130),
  inference(superposition,[],[f2792,f345])).
fof(f345,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f214])).
fof(f214,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f82])).
fof(f82,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X8] : (s(bool,'F_0') = s(bool,X8) <=> ~p(s(bool,X8)))),
  file('thm.integer.INT_MOD_COMMON_FACTOR','thm.bool.EQ_CLAUSES|split|2')).
fof(f2856,plain,(
  ( ! [X35] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X35))) ) | $spl130),
  inference(trivial_inequality_removal,[],[f2845])).
fof(f2845,plain,(
  ( ! [X35] : (s(bool,X35) != s(bool,X35) | s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X35))) ) | $spl130),
  inference(superposition,[],[f2792,f343])).
fof(f3172,plain,(
  ( ! [X9] : (s(bool,'F_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))) | ~p(s(bool,X9))) ) | ($spl104 | $spl121)),
  inference(resolution,[],[f2963,f2505])).
fof(f2963,plain,(
  ( ! [X2,X3] : (~sP3(X2,X3,'F_0') | s(bool,'F_0') = s(bool,X3) | ~p(s(bool,X2))) ) | $spl104),
  inference(subsumption_resolution,[],[f2957,f321])).
fof(f2957,plain,(
  ( ! [X2,X3] : (~sP3(X2,X3,'F_0') | ~sP1(X3,'F_0',X2) | s(bool,'F_0') = s(bool,X3) | ~p(s(bool,X2))) ) | $spl104),
  inference(resolution,[],[f2923,f326])).
fof(f2923,plain,(
  ( ! [X4,X5] : (sP4(X4,'F_0',X5) | ~sP3(X4,X5,'F_0') | ~sP1(X5,'F_0',X4)) ) | $spl104),
  inference(subsumption_resolution,[],[f2921,f514])).
fof(f2921,plain,(
  ( ! [X4,X5] : (sP4(X4,'F_0',X5) | p(s(bool,'F_0')) | ~sP3(X4,X5,'F_0') | ~sP1(X5,'F_0',X4)) ) | $spl104),
  inference(resolution,[],[f2918,f310])).
fof(f2918,plain,(
  ( ! [X4,X5] : (sP2('F_0',X4,X5)) ) | $spl104),
  inference(subsumption_resolution,[],[f2912,f317])).
fof(f2912,plain,(
  ( ! [X4,X5] : (sP2('F_0',X4,X5) | p(s(bool,X5))) ) | $spl104),
  inference(resolution,[],[f2909,f315])).
fof(f2909,plain,(
  ( ! [X6,X7] : (~sP3(X7,X6,'F_0') | sP2('F_0',X6,X7)) ) | $spl104),
  inference(subsumption_resolution,[],[f2905,f318])).
fof(f2905,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | ~sP3(X7,X6,'F_0') | sP2('F_0',X6,X7)) ) | $spl104),
  inference(resolution,[],[f2899,f306])).
fof(f2899,plain,(
  ( ! [X2,X3] : (sP4(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP3(X3,X2,'F_0')) ) | $spl104),
  inference(subsumption_resolution,[],[f2898,f347])).
fof(f2898,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP4(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP3(X3,X2,'F_0')) ) | $spl104),
  inference(subsumption_resolution,[],[f2895,f2873])).
fof(f2873,plain,(
  ( ! [X4,X5] : (sP1(X4,'F_0',X5) | s(bool,'T_0') != s(bool,X4)) ) | $spl104),
  inference(subsumption_resolution,[],[f2868,f321])).
fof(f2868,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | p(s(bool,X5)) | sP1(X4,'F_0',X5)) ) | $spl104),
  inference(resolution,[],[f2756,f305])).
fof(f2756,plain,(
  ( ! [X6,X7] : (sP4(X7,'F_0',X6) | s(bool,'T_0') != s(bool,X6) | p(s(bool,X7))) ) | $spl104),
  inference(superposition,[],[f2747,f324])).
fof(f2747,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl104),
  inference(subsumption_resolution,[],[f2736,f1903])).
fof(f2736,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | $spl104),
  inference(superposition,[],[f1927,f347])).
fof(f1927,plain,(
  s(bool,'F_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))) | $spl104),
  inference(resolution,[],[f1903,f342])).
fof(f342,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f213])).
fof(f2895,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP4(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP3(X3,X2,'F_0') | ~sP1(X2,'F_0',X3)) ) | $spl104),
  inference(resolution,[],[f2874,f309])).
fof(f2874,plain,(
  ( ! [X6,X7] : (sP2('F_0',X6,X7) | s(bool,'T_0') != s(bool,X6)) ) | $spl104),
  inference(subsumption_resolution,[],[f2869,f317])).
fof(f2869,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X6) | p(s(bool,X7)) | sP2('F_0',X6,X7)) ) | $spl104),
  inference(resolution,[],[f2756,f306])).
fof(f3171,plain,(
  ( ! [X8] : (s(bool,'F_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,X8))) ) | ($spl104 | $spl115)),
  inference(resolution,[],[f2963,f2194])).
fof(f3039,plain,(
  ( ! [X9] : (sP4(X9,'F_0','const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) ) | $spl140),
  inference(cnf_transformation,[],[f3039_D])).
fof(f3039_D,plain,(
  ( ! [X9] : (sP4(X9,'F_0','const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f3035,plain,(
  ( ! [X9] : (sP4(X9,'F_0','const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) ) | ($spl104 | $spl121)),
  inference(resolution,[],[f2893,f2505])).
fof(f2893,plain,(
  ( ! [X4,X5] : (~sP3(X5,X4,'F_0') | sP4(X5,'F_0',X4) | s(bool,'T_0') != s(bool,X4)) ) | $spl104),
  inference(subsumption_resolution,[],[f2892,f514])).
fof(f2892,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | sP4(X5,'F_0',X4) | p(s(bool,'F_0')) | ~sP3(X5,X4,'F_0')) ) | $spl104),
  inference(subsumption_resolution,[],[f2889,f319])).
fof(f2889,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | sP4(X5,'F_0',X4) | p(s(bool,'F_0')) | ~sP3(X5,X4,'F_0') | ~sP2('F_0',X4,X5)) ) | $spl104),
  inference(resolution,[],[f2873,f310])).
fof(f3038,plain,(
  ( ! [X8] : (sP4(X8,'F_0','const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl138),
  inference(cnf_transformation,[],[f3038_D])).
fof(f3038_D,plain,(
  ( ! [X8] : (sP4(X8,'F_0','const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f3034,plain,(
  ( ! [X8] : (sP4(X8,'F_0','const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl104 | $spl115)),
  inference(resolution,[],[f2893,f2194])).
fof(f3022,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))) | ~p(s(bool,X9))) ) | ($spl104 | $spl121)),
  inference(resolution,[],[f3016,f2505])).
fof(f3016,plain,(
  ( ! [X14,X15] : (~sP3(X14,X15,'F_0') | s(bool,'T_0') != s(bool,X15) | ~p(s(bool,X14))) ) | $spl104),
  inference(subsumption_resolution,[],[f3015,f321])).
fof(f3015,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | ~sP3(X14,X15,'F_0') | ~sP1(X15,'F_0',X14)) ) | $spl104),
  inference(subsumption_resolution,[],[f3014,f460])).
fof(f3014,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | p(s(bool,'F_0')) | ~sP3(X14,X15,'F_0') | ~sP1(X15,'F_0',X14)) ) | $spl104),
  inference(subsumption_resolution,[],[f3007,f319])).
fof(f3007,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | p(s(bool,'F_0')) | ~sP3(X14,X15,'F_0') | ~sP2('F_0',X15,X14) | ~sP1(X15,'F_0',X14)) ) | $spl104),
  inference(resolution,[],[f2758,f310])).
fof(f2758,plain,(
  ( ! [X10,X11] : (~sP4(X11,'F_0',X10) | ~p(s(bool,X11)) | s(bool,'T_0') != s(bool,X10)) ) | $spl104),
  inference(superposition,[],[f2747,f326])).
fof(f3021,plain,(
  ( ! [X8] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,X8))) ) | ($spl104 | $spl115)),
  inference(resolution,[],[f3016,f2194])).
fof(f2792,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | s(bool,'T_0') != s(bool,X2)) ) | $spl130),
  inference(cnf_transformation,[],[f2792_D])).
fof(f2792_D,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | s(bool,'T_0') != s(bool,X2)) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f2791,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) ) | $spl104),
  inference(duplicate_literal_removal,[],[f2786])).
fof(f2786,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) ) | $spl104),
  inference(resolution,[],[f2755,f327])).
fof(f2755,plain,(
  ( ! [X4,X5] : (sP4(X5,X4,'F_0') | s(bool,'T_0') != s(bool,X4) | p(s(bool,X5))) ) | $spl104),
  inference(superposition,[],[f2747,f324])).
fof(f2444,plain,(
  ( ! [X39,X38] : (sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13)),'const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),X38) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X38),s(bool,X39)))) | p(s(bool,X39))) ) | $spl126),
  inference(cnf_transformation,[],[f2444_D])).
fof(f2444_D,plain,(
  ( ! [X39,X38] : (sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13)),'const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),X38) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X38),s(bool,X39)))) | p(s(bool,X39))) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f2445,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X35)))) | s('type.integer.int',X36) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X37),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13))))) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X36),s('type.integer.int',X35)))) ) | $spl128),
  inference(cnf_transformation,[],[f2445_D])).
fof(f2445_D,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X35)))) | s('type.integer.int',X36) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X37),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13))))) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X36),s('type.integer.int',X35)))) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f2443,plain,(
  ( ! [X39,X37,X35,X38,X36] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X36),s('type.integer.int',X35))) | s('type.integer.int',X36) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X37),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13))))) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13)),'const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),X38) | p(s(bool,X39)) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X38),s(bool,X39)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X35))))) ) | $spl12),
  inference(forward_demodulation,[],[f2416,f353])).
fof(f2416,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s('type.integer.int',X35) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X36),s('type.integer.int',X35))) | s('type.integer.int',X36) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X37),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13))))) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13)),'const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),X38) | p(s(bool,X39)) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X38),s(bool,X39)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X35)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X35))))) ) | $spl12),
  inference(superposition,[],[f557,f352])).
fof(f2354,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | $spl121),
  inference(cnf_transformation,[],[f2354_D])).
fof(f2354_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) <=> ~$spl121),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl121])])).
fof(f2357,plain,(
  ( ! [X14,X12,X10,X13,X11] : (p(s(bool,X10)) | s('type.integer.int',X11) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X14),s('type.integer.int',X12))),s('type.integer.int',sK13))) | ~p(s(bool,X13)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',X12)))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X10),s(bool,X13)))) | s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X11),s('type.integer.int',X12)))) ) | $spl124),
  inference(cnf_transformation,[],[f2357_D])).
fof(f2357_D,plain,(
  ( ! [X14,X12,X10,X13,X11] : (p(s(bool,X10)) | s('type.integer.int',X11) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X14),s('type.integer.int',X12))),s('type.integer.int',sK13))) | ~p(s(bool,X13)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',X12)))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X10),s(bool,X13)))) | s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X11),s('type.integer.int',X12)))) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f2353,plain,(
  ( ! [X14,X12,X10,X13,X11] : (p(s(bool,X10)) | s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X11),s('type.integer.int',X12))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X10),s(bool,X13)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',X12)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | ~p(s(bool,X13)) | s('type.integer.int',X11) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X14),s('type.integer.int',X12))),s('type.integer.int',sK13)))) ) | $spl115),
  inference(resolution,[],[f2185,f304])).
fof(f2185,plain,(
  ( ! [X0,X1] : (~sP0(sK13,X1,X0) | p(s(bool,X0))) ) | $spl115),
  inference(resolution,[],[f2024,f300])).
fof(f300,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X2)) | ~sP0(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f196])).
fof(f2355,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | $spl120),
  inference(cnf_transformation,[],[f2355_D])).
fof(f2355_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f2356,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,X0)) | s('type.integer.int',X1) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X4),s('type.integer.int',X2))),s('type.integer.int',sK13))) | p(s(bool,X3)) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X0),s(bool,X3)))) | s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X1),s('type.integer.int',X2)))) ) | $spl122),
  inference(cnf_transformation,[],[f2356_D])).
fof(f2356_D,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,X0)) | s('type.integer.int',X1) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X4),s('type.integer.int',X2))),s('type.integer.int',sK13))) | p(s(bool,X3)) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X0),s(bool,X3)))) | s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X1),s('type.integer.int',X2)))) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f2351,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,X0)) | s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X1),s('type.integer.int',X2))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X0),s(bool,X3)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | p(s(bool,X3)) | s('type.integer.int',X1) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X4),s('type.integer.int',X2))),s('type.integer.int',sK13)))) ) | $spl115),
  inference(resolution,[],[f2185,f302])).
fof(f2350,plain,(
  ( ! [X2,X0,X1] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',sK13))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X0),s(bool,X1)))) | ~p(s(bool,X1)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X0)) ) | $spl104),
  inference(forward_demodulation,[],[f2349,f354])).
fof(f2349,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X0),s(bool,X1)))) | ~p(s(bool,X1)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X0) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl104),
  inference(subsumption_resolution,[],[f2348,f523])).
fof(f2348,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X0),s(bool,X1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,X1)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X0) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl104),
  inference(subsumption_resolution,[],[f2327,f1903])).
fof(f2335,plain,(
  ( ! [X33,X31,X34,X32] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X31),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X34),s('type.integer.int',sK13))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X32),s(bool,X33)))) | ~p(s(bool,X33)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X32) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X31))))) ) | $spl104),
  inference(forward_demodulation,[],[f2334,f353])).
fof(f2334,plain,(
  ( ! [X33,X31,X34,X32] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X34),s('type.integer.int',sK13))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X32),s(bool,X33)))) | ~p(s(bool,X33)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X32) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X31)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X31))))) ) | $spl104),
  inference(forward_demodulation,[],[f2333,f354])).
fof(f2333,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X32),s(bool,X33)))) | ~p(s(bool,X33)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X32) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X34),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X31)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X31))))) ) | $spl104),
  inference(subsumption_resolution,[],[f2332,f523])).
fof(f2332,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X32),s(bool,X33)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,X33)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X32) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X34),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X31)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X31))))) ) | $spl104),
  inference(subsumption_resolution,[],[f2328,f1903])).
fof(f2022,plain,(
  ( ! [X4] : (~sP0(sK13,'const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),X4) | ~p(s(bool,X4))) ) | $spl112),
  inference(cnf_transformation,[],[f2022_D])).
fof(f2022_D,plain,(
  ( ! [X4] : (~sP0(sK13,'const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),X4) | ~p(s(bool,X4))) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f2024,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl115),
  inference(cnf_transformation,[],[f2024_D])).
fof(f2024_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl115),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl115])])).
fof(f1907,plain,(
  ( ! [X4] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,X4)) | ~sP0(sK13,'const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),X4)) ) | $spl104),
  inference(resolution,[],[f1903,f301])).
fof(f2019,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int',sK13))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int',sK13))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | ~p(s(bool,X2)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X1)) ) | $spl104),
  inference(forward_demodulation,[],[f2018,f354])).
fof(f2018,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int',sK13))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | ~p(s(bool,X2)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X1) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl104),
  inference(subsumption_resolution,[],[f1906,f523])).
fof(f1906,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int',sK13))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,X2)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),sK13,X1) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl104),
  inference(resolution,[],[f1903,f304])).
fof(f1892,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))),
  inference(equality_resolution,[],[f533])).
fof(f533,plain,(
  ( ! [X0] : (s('type.integer.int',sK13) != s('type.integer.int',X0) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))))) )),
  inference(forward_demodulation,[],[f531,f353])).
fof(f531,plain,(
  ( ! [X0] : (s('type.integer.int',sK13) != s('type.integer.int',X0) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))))) )),
  inference(superposition,[],[f297,f352])).
fof(f297,plain,(
  s('type.integer.int',sK13) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),
  inference(cnf_transformation,[],[f193])).
fof(f1666,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))))) | $spl87),
  inference(cnf_transformation,[],[f1666_D])).
fof(f1666_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))))) <=> ~$spl87),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl87])])).
fof(f1840,plain,(
  ( ! [X10,X8] : (p(s(bool,X8)) | ~p(s(bool,X10)) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X8),s(bool,X10))))) ) | $spl102),
  inference(cnf_transformation,[],[f1840_D])).
fof(f1840_D,plain,(
  ( ! [X10,X8] : (p(s(bool,X8)) | ~p(s(bool,X10)) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X8),s(bool,X10))))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f1839,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X8)) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X8),s(bool,X10)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))))) | ~p(s(bool,X10)) | s('type.integer.int',X9) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))))) ) | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f1837,f549])).
fof(f1837,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X8)) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X8),s(bool,X10)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))))) | ~p(s(bool,X10)) | s('type.integer.int',X9) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))))) ) | $spl1),
  inference(resolution,[],[f1706,f304])).
fof(f1706,plain,(
  ( ! [X4] : (~sP0('const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),'const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),X4) | p(s(bool,X4))) ) | $spl1),
  inference(resolution,[],[f544,f299])).
fof(f1665,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))))) | $spl86),
  inference(cnf_transformation,[],[f1665_D])).
fof(f1665_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))))) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f1835,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X0)) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X0),s(bool,X2)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))))) | p(s(bool,X2)) | s('type.integer.int',X1) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))))) ) | $spl1),
  inference(resolution,[],[f1706,f302])).
fof(f1823,plain,(
  ( ! [X2,X1] : (~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),'const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),X1) | p(s(bool,X2))) ) | $spl98),
  inference(cnf_transformation,[],[f1823_D])).
fof(f1823_D,plain,(
  ( ! [X2,X1] : (~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),'const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),X1) | p(s(bool,X2))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f1822,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | p(s(bool,X2)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),'const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),X1)) ) | $spl1),
  inference(forward_demodulation,[],[f1705,f354])).
fof(f1705,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | p(s(bool,X2)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),'const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),X1) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl1),
  inference(resolution,[],[f544,f303])).
fof(f544,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))))) | $spl1),
  inference(cnf_transformation,[],[f544_D])).
fof(f544_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f1685,plain,(
  ( ! [X10,X8] : (~p(s(bool,X8)) | ~p(s(bool,X10)) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X8),s(bool,X10))))) ) | $spl96),
  inference(cnf_transformation,[],[f1685_D])).
fof(f1685_D,plain,(
  ( ! [X10,X8] : (~p(s(bool,X8)) | ~p(s(bool,X10)) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X8),s(bool,X10))))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f1684,plain,(
  ( ! [X10,X8,X11,X9] : (s('type.integer.int',X9) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))) | ~p(s(bool,X8)) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X8),s(bool,X10)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))))) | ~p(s(bool,X10))) ) | $spl4),
  inference(forward_demodulation,[],[f1683,f354])).
fof(f1683,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,X8)) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X8),s(bool,X10)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))))) | ~p(s(bool,X10)) | s('type.integer.int',X9) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl4),
  inference(subsumption_resolution,[],[f1679,f523])).
fof(f1679,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,X8)) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X8),s(bool,X10)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,X10)) | s('type.integer.int',X9) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl4),
  inference(resolution,[],[f1669,f304])).
fof(f1681,plain,(
  ( ! [X6,X4] : (~p(s(bool,X4)) | p(s(bool,X6)) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X4),s(bool,X6))))) ) | $spl92),
  inference(cnf_transformation,[],[f1681_D])).
fof(f1681_D,plain,(
  ( ! [X6,X4] : (~p(s(bool,X4)) | p(s(bool,X6)) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X4),s(bool,X6))))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f1680,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.integer.int',X5) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))) | ~p(s(bool,X4)) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X5),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X4),s(bool,X6)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))))) | p(s(bool,X6))) ) | $spl4),
  inference(forward_demodulation,[],[f1678,f354])).
fof(f1678,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,X4)) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X5),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X4),s(bool,X6)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))))) | p(s(bool,X6)) | s('type.integer.int',X5) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl4),
  inference(resolution,[],[f1669,f303])).
fof(f1667,plain,(
  ( ! [X2,X1] : (~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | sP0('const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),'const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),X1) | ~p(s(bool,X2))) ) | $spl88),
  inference(cnf_transformation,[],[f1667_D])).
fof(f1667_D,plain,(
  ( ! [X2,X1] : (~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | sP0('const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),'const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),X1) | ~p(s(bool,X2))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f1552,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))))) | ~p(s(bool,X2)) | sP0('const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)),'const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),X1) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))))) ) | $spl4),
  inference(resolution,[],[f549,f304])).
fof(f1546,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))))) | $spl11),
  inference(forward_demodulation,[],[f1514,f353])).
fof(f1514,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))))) | $spl11),
  inference(superposition,[],[f555,f352])).
fof(f1412,plain,(
  ( ! [X26,X24,X23,X25,X22] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X22),s('type.integer.int',X23))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X22) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X26),s('type.integer.int',X23))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X23,X24) | ~p(s(bool,X25)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))),s('type.integer.int',X23)))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X23),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X24),s(bool,X25))))) ) | $spl84),
  inference(cnf_transformation,[],[f1412_D])).
fof(f1412_D,plain,(
  ( ! [X26,X24,X23,X25,X22] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X22),s('type.integer.int',X23))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X22) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X26),s('type.integer.int',X23))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X23,X24) | ~p(s(bool,X25)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))),s('type.integer.int',X23)))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X23),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X24),s(bool,X25))))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f1385,plain,(
  ( ! [X26,X24,X23,X25,X22] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X22),s('type.integer.int',X23))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X23),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X24),s(bool,X25)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))),s('type.integer.int',X23)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) | ~p(s(bool,X25)) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X23,X24) | s('type.integer.int',X22) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X26),s('type.integer.int',X23))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) ) | $spl19),
  inference(superposition,[],[f573,f304])).
fof(f573,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl19),
  inference(cnf_transformation,[],[f573_D])).
fof(f573_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f1411,plain,(
  ( ! [X14,X12,X15,X13,X16] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X12),s('type.integer.int',X13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X16),s('type.integer.int',X13))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X13,X14) | p(s(bool,X15)) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X14),s(bool,X15))))) ) | $spl82),
  inference(cnf_transformation,[],[f1411_D])).
fof(f1411_D,plain,(
  ( ! [X14,X12,X15,X13,X16] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X12),s('type.integer.int',X13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X16),s('type.integer.int',X13))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X13,X14) | p(s(bool,X15)) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X14),s(bool,X15))))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f1383,plain,(
  ( ! [X14,X12,X15,X13,X16] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X12),s('type.integer.int',X13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X14),s(bool,X15)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) | p(s(bool,X15)) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X13,X14) | s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X16),s('type.integer.int',X13))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) ) | $spl19),
  inference(superposition,[],[f573,f302])).
fof(f1406,plain,(
  ( ! [X2,X1] : (~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),'const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),X1) | p(s(bool,X2))) ) | $spl80),
  inference(cnf_transformation,[],[f1406_D])).
fof(f1406_D,plain,(
  ( ! [X2,X1] : (~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),'const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),X1) | p(s(bool,X2))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f1284,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | p(s(bool,X2)) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),'const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),X1) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) ) | $spl19),
  inference(resolution,[],[f573,f303])).
fof(f1252,plain,(
  ( ! [X26,X24,X23,X25,X22] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X22),s('type.integer.int',X23))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X23),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X24),s(bool,X25)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))),s('type.integer.int',X23)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) | ~p(s(bool,X25)) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X23,X24) | s('type.integer.int',X22) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X26),s('type.integer.int',X23))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) ) | $spl18),
  inference(superposition,[],[f574,f304])).
fof(f1250,plain,(
  ( ! [X14,X12,X15,X13,X16] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X12),s('type.integer.int',X13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X14),s(bool,X15)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) | p(s(bool,X15)) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X13,X14) | s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X16),s('type.integer.int',X13))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) ) | $spl18),
  inference(superposition,[],[f574,f302])).
fof(f1154,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) | ~p(s(bool,X2)) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),'const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),X1) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) ) | $spl18),
  inference(resolution,[],[f574,f304])).
fof(f708,plain,(
  ( ! [X26,X24,X23,X25,X22] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X22),s('type.integer.int',X23)))))) | s('type.integer.int',X22) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X26),s('type.integer.int',X23))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X23,X24) | ~p(s(bool,X25)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))),s('type.integer.int',X23)))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X23),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X24),s(bool,X25))))) ) | $spl36),
  inference(cnf_transformation,[],[f708_D])).
fof(f708_D,plain,(
  ( ! [X26,X24,X23,X25,X22] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X22),s('type.integer.int',X23)))))) | s('type.integer.int',X22) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X26),s('type.integer.int',X23))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X23,X24) | ~p(s(bool,X25)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))),s('type.integer.int',X23)))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X23),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X24),s(bool,X25))))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f1128,plain,(
  ( ! [X26,X24,X23,X25,X22] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X22),s('type.integer.int',X23)))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X23),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X24),s(bool,X25)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))),s('type.integer.int',X23)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) | ~p(s(bool,X25)) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X23,X24) | s('type.integer.int',X22) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X26),s('type.integer.int',X23))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) ) | $spl16),
  inference(superposition,[],[f572,f304])).
fof(f707,plain,(
  ( ! [X14,X12,X15,X13,X16] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X12),s('type.integer.int',X13)))))) | s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X16),s('type.integer.int',X13))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X13,X14) | p(s(bool,X15)) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X14),s(bool,X15))))) ) | $spl34),
  inference(cnf_transformation,[],[f707_D])).
fof(f707_D,plain,(
  ( ! [X14,X12,X15,X13,X16] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X12),s('type.integer.int',X13)))))) | s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X16),s('type.integer.int',X13))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X13,X14) | p(s(bool,X15)) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X14),s(bool,X15))))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f1126,plain,(
  ( ! [X14,X12,X15,X13,X16] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X12),s('type.integer.int',X13)))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X14),s(bool,X15)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) | p(s(bool,X15)) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X13,X14) | s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X16),s('type.integer.int',X13))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) ) | $spl16),
  inference(superposition,[],[f572,f302])).
fof(f702,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | ~sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),'const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),X4)) ) | $spl30),
  inference(cnf_transformation,[],[f702_D])).
fof(f702_D,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | ~sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),'const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),X4)) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f704,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl33),
  inference(cnf_transformation,[],[f704_D])).
fof(f704_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f1029,plain,(
  ( ! [X4] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,X4)) | ~sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),'const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),X4)) ) | $spl16),
  inference(resolution,[],[f572,f301])).
fof(f700,plain,(
  ( ! [X2,X1] : (~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),'const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X1) | ~p(s(bool,X2))) ) | $spl26),
  inference(cnf_transformation,[],[f700_D])).
fof(f700_D,plain,(
  ( ! [X2,X1] : (~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),'const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X1) | ~p(s(bool,X2))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f701,plain,(
  ( ! [X0,X3] : (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) ) | $spl28),
  inference(cnf_transformation,[],[f701_D])).
fof(f701_D,plain,(
  ( ! [X0,X3] : (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f1148,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | ~p(s(bool,X2)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),'const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X1)) ) | $spl16),
  inference(forward_demodulation,[],[f1147,f354])).
fof(f1147,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | ~p(s(bool,X2)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),'const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X1) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl16),
  inference(subsumption_resolution,[],[f1028,f523])).
fof(f1028,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,X2)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),'const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X1) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl16),
  inference(resolution,[],[f572,f304])).
fof(f555,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13)))))) | $spl11),
  inference(cnf_transformation,[],[f555_D])).
fof(f555_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13)))))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f1023,plain,(
  ( ! [X26,X24,X23,X25,X22] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X22),s('type.integer.int',X23)))))) | s('type.integer.int',X22) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X26),s('type.integer.int',X23))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13))))) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13)),X23,X24) | ~p(s(bool,X25)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13))),s('type.integer.int',X23)))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X23),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X24),s(bool,X25))))) ) | $spl70),
  inference(cnf_transformation,[],[f1023_D])).
fof(f1023_D,plain,(
  ( ! [X26,X24,X23,X25,X22] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X22),s('type.integer.int',X23)))))) | s('type.integer.int',X22) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X26),s('type.integer.int',X23))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13))))) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13)),X23,X24) | ~p(s(bool,X25)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13))),s('type.integer.int',X23)))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X23),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X24),s(bool,X25))))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f993,plain,(
  ( ! [X26,X24,X23,X25,X22] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X22),s('type.integer.int',X23)))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X23),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X24),s(bool,X25)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13))),s('type.integer.int',X23)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13)))))) | ~p(s(bool,X25)) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13)),X23,X24) | s('type.integer.int',X22) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X26),s('type.integer.int',X23))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13)))))) ) | $spl6),
  inference(superposition,[],[f552,f304])).
fof(f556,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13)))))) | $spl10),
  inference(cnf_transformation,[],[f556_D])).
fof(f556_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13)))))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f1022,plain,(
  ( ! [X14,X12,X15,X13,X16] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X12),s('type.integer.int',X13)))))) | s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X16),s('type.integer.int',X13))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13))))) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13)),X13,X14) | p(s(bool,X15)) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X14),s(bool,X15))))) ) | $spl68),
  inference(cnf_transformation,[],[f1022_D])).
fof(f1022_D,plain,(
  ( ! [X14,X12,X15,X13,X16] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X12),s('type.integer.int',X13)))))) | s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X16),s('type.integer.int',X13))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13))))) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13)),X13,X14) | p(s(bool,X15)) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X14),s(bool,X15))))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f991,plain,(
  ( ! [X14,X12,X15,X13,X16] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X12),s('type.integer.int',X13)))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X14),s(bool,X15)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13)))))) | p(s(bool,X15)) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13)),X13,X14) | s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X16),s('type.integer.int',X13))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13)))))) ) | $spl6),
  inference(superposition,[],[f552,f302])).
fof(f572,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) | $spl16),
  inference(cnf_transformation,[],[f572_D])).
fof(f572_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f1021,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))))) | $spl6),
  inference(forward_demodulation,[],[f986,f353])).
fof(f986,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))))) | $spl6),
  inference(superposition,[],[f552,f352])).
fof(f1016,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | ~sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13)),'const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),X4)) ) | $spl64),
  inference(cnf_transformation,[],[f1016_D])).
fof(f1016_D,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | ~sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13)),'const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),X4)) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f1018,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl67),
  inference(cnf_transformation,[],[f1018_D])).
fof(f1018_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl67),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl67])])).
fof(f893,plain,(
  ( ! [X4] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,X4)) | ~sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13)),'const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),X4)) ) | $spl6),
  inference(resolution,[],[f552,f301])).
fof(f1014,plain,(
  ( ! [X2,X1] : (~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),'const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13)),X1) | ~p(s(bool,X2))) ) | $spl60),
  inference(cnf_transformation,[],[f1014_D])).
fof(f1014_D,plain,(
  ( ! [X2,X1] : (~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),'const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13)),X1) | ~p(s(bool,X2))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f1015,plain,(
  ( ! [X0,X3] : (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13)))))) ) | $spl62),
  inference(cnf_transformation,[],[f1015_D])).
fof(f1015_D,plain,(
  ( ! [X0,X3] : (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13)))))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f1013,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | ~p(s(bool,X2)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),'const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13)),X1)) ) | $spl6),
  inference(forward_demodulation,[],[f1012,f354])).
fof(f1012,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | ~p(s(bool,X2)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),'const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13)),X1) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl6),
  inference(subsumption_resolution,[],[f892,f523])).
fof(f892,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,X2)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),'const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13)),X1) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl6),
  inference(resolution,[],[f552,f304])).
fof(f891,plain,(
  ( ! [X30,X28,X33,X31,X29,X32] : (sP12('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X28),s('type.integer.int',X29)))),X30) | s('type.integer.int',X28) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X33),s('type.integer.int',X29))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X29,X31) | ~p(s(bool,X32)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))),s('type.integer.int',X29)))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X29),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X31),s(bool,X32))))) ) | $spl58),
  inference(cnf_transformation,[],[f891_D])).
fof(f891_D,plain,(
  ( ! [X30,X28,X33,X31,X29,X32] : (sP12('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X28),s('type.integer.int',X29)))),X30) | s('type.integer.int',X28) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X33),s('type.integer.int',X29))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X29,X31) | ~p(s(bool,X32)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))),s('type.integer.int',X29)))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X29),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X31),s(bool,X32))))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f884,plain,(
  ( ! [X30,X28,X33,X31,X29,X32] : (sP12('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X28),s('type.integer.int',X29)))),X30) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X29),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X31),s(bool,X32)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))),s('type.integer.int',X29)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) | ~p(s(bool,X32)) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X29,X31) | s('type.integer.int',X28) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X33),s('type.integer.int',X29))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) ) | $spl16),
  inference(superposition,[],[f658,f304])).
fof(f658,plain,(
  ( ! [X147] : (sP12('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))),X147)) ) | $spl16),
  inference(resolution,[],[f572,f444])).
fof(f444,plain,(
  ( ! [X0,X1] : (sP12(X0,X1) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f255])).
fof(f255,plain,(
  ! [X0,X1] : ((~sP12(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP12(X0,X1)))),
  inference(flattening,[],[f254])).
fof(f254,plain,(
  ! [X0,X1] : ((~sP12(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP12(X0,X1)))),
  inference(nnf_transformation,[],[f191])).
fof(f191,plain,(
  ! [X0,X1] : (sP12(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP12])])).
fof(f890,plain,(
  ( ! [X21,X19,X17,X20,X18,X16] : (sP12('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X16),s('type.integer.int',X17)))),X18) | s('type.integer.int',X16) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X21),s('type.integer.int',X17))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X17,X19) | p(s(bool,X20)) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X17),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X19),s(bool,X20))))) ) | $spl56),
  inference(cnf_transformation,[],[f890_D])).
fof(f890_D,plain,(
  ( ! [X21,X19,X17,X20,X18,X16] : (sP12('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X16),s('type.integer.int',X17)))),X18) | s('type.integer.int',X16) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X21),s('type.integer.int',X17))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X17,X19) | p(s(bool,X20)) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X17),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X19),s(bool,X20))))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f882,plain,(
  ( ! [X21,X19,X17,X20,X18,X16] : (sP12('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X16),s('type.integer.int',X17)))),X18) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X17),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X19),s(bool,X20)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) | p(s(bool,X20)) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X17,X19) | s('type.integer.int',X16) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X21),s('type.integer.int',X17))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) ) | $spl16),
  inference(superposition,[],[f658,f302])).
fof(f873,plain,(
  ( ! [X30,X28,X33,X31,X29,X32] : (sP12(X30,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X28),s('type.integer.int',X29))))) | s('type.integer.int',X28) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X33),s('type.integer.int',X29))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X29,X31) | ~p(s(bool,X32)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))),s('type.integer.int',X29)))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X29),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X31),s(bool,X32))))) ) | $spl52),
  inference(cnf_transformation,[],[f873_D])).
fof(f873_D,plain,(
  ( ! [X30,X28,X33,X31,X29,X32] : (sP12(X30,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X28),s('type.integer.int',X29))))) | s('type.integer.int',X28) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X33),s('type.integer.int',X29))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X29,X31) | ~p(s(bool,X32)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))),s('type.integer.int',X29)))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X29),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X31),s(bool,X32))))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f865,plain,(
  ( ! [X30,X28,X33,X31,X29,X32] : (sP12(X30,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X28),s('type.integer.int',X29))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X29),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X31),s(bool,X32)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))),s('type.integer.int',X29)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) | ~p(s(bool,X32)) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X29,X31) | s('type.integer.int',X28) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X33),s('type.integer.int',X29))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) ) | $spl16),
  inference(superposition,[],[f657,f304])).
fof(f657,plain,(
  ( ! [X146] : (sP12(X146,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) ) | $spl16),
  inference(resolution,[],[f572,f443])).
fof(f443,plain,(
  ( ! [X0,X1] : (sP12(X0,X1) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f255])).
fof(f872,plain,(
  ( ! [X21,X19,X17,X20,X18,X16] : (sP12(X18,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X16),s('type.integer.int',X17))))) | s('type.integer.int',X16) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X21),s('type.integer.int',X17))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X17,X19) | p(s(bool,X20)) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X17),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X19),s(bool,X20))))) ) | $spl50),
  inference(cnf_transformation,[],[f872_D])).
fof(f872_D,plain,(
  ( ! [X21,X19,X17,X20,X18,X16] : (sP12(X18,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X16),s('type.integer.int',X17))))) | s('type.integer.int',X16) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X21),s('type.integer.int',X17))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X17,X19) | p(s(bool,X20)) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X17),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X19),s(bool,X20))))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f863,plain,(
  ( ! [X21,X19,X17,X20,X18,X16] : (sP12(X18,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X16),s('type.integer.int',X17))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X17),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X19),s(bool,X20)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) | p(s(bool,X20)) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X17,X19) | s('type.integer.int',X16) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X21),s('type.integer.int',X17))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) ) | $spl16),
  inference(superposition,[],[f657,f302])).
fof(f850,plain,(
  ( ! [X26,X24,X23,X25,X22] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X22),s('type.integer.int',X23)))))) | s('type.integer.int',X22) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X26),s('type.integer.int',X23))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X23,X24) | ~p(s(bool,X25)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))),s('type.integer.int',X23)))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X23),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X24),s(bool,X25))))) ) | $spl44),
  inference(cnf_transformation,[],[f850_D])).
fof(f850_D,plain,(
  ( ! [X26,X24,X23,X25,X22] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X22),s('type.integer.int',X23)))))) | s('type.integer.int',X22) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X26),s('type.integer.int',X23))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X23,X24) | ~p(s(bool,X25)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))),s('type.integer.int',X23)))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X23),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X24),s(bool,X25))))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f817,plain,(
  ( ! [X26,X24,X23,X25,X22] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X22),s('type.integer.int',X23)))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X23),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X24),s(bool,X25)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))),s('type.integer.int',X23)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) | ~p(s(bool,X25)) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X23,X24) | s('type.integer.int',X22) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X26),s('type.integer.int',X23))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) ) | $spl17),
  inference(superposition,[],[f571,f304])).
fof(f571,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) | $spl17),
  inference(cnf_transformation,[],[f571_D])).
fof(f571_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f849,plain,(
  ( ! [X14,X12,X15,X13,X16] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X12),s('type.integer.int',X13)))))) | s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X16),s('type.integer.int',X13))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X13,X14) | p(s(bool,X15)) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X14),s(bool,X15))))) ) | $spl42),
  inference(cnf_transformation,[],[f849_D])).
fof(f849_D,plain,(
  ( ! [X14,X12,X15,X13,X16] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X12),s('type.integer.int',X13)))))) | s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X16),s('type.integer.int',X13))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X13,X14) | p(s(bool,X15)) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X14),s(bool,X15))))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f815,plain,(
  ( ! [X14,X12,X15,X13,X16] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X12),s('type.integer.int',X13)))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X14),s(bool,X15)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) | p(s(bool,X15)) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X13,X14) | s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X16),s('type.integer.int',X13))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) ) | $spl17),
  inference(superposition,[],[f571,f302])).
fof(f848,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl41),
  inference(cnf_transformation,[],[f848_D])).
fof(f848_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl41),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl41])])).
fof(f574,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl18),
  inference(cnf_transformation,[],[f574_D])).
fof(f574_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f846,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl17),
  inference(forward_demodulation,[],[f845,f353])).
fof(f845,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))))) | $spl17),
  inference(subsumption_resolution,[],[f814,f571])).
fof(f814,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) | $spl17),
  inference(superposition,[],[f571,f352])).
fof(f839,plain,(
  ( ! [X2,X1] : (~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),'const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X1) | p(s(bool,X2))) ) | $spl38),
  inference(cnf_transformation,[],[f839_D])).
fof(f839_D,plain,(
  ( ! [X2,X1] : (~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),'const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X1) | p(s(bool,X2))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f838,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | p(s(bool,X2)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),'const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X1)) ) | $spl17),
  inference(forward_demodulation,[],[f714,f354])).
fof(f714,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | p(s(bool,X2)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),'const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X1) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl17),
  inference(resolution,[],[f571,f303])).
fof(f679,plain,(
  ( ! [X26,X24,X23,X25,X22] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X22),s('type.integer.int',X23)))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X23),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X24),s(bool,X25)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))),s('type.integer.int',X23)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) | ~p(s(bool,X25)) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X23,X24) | s('type.integer.int',X22) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X26),s('type.integer.int',X23))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) ) | $spl16),
  inference(superposition,[],[f572,f304])).
fof(f677,plain,(
  ( ! [X14,X12,X15,X13,X16] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X12),s('type.integer.int',X13)))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X14),s(bool,X15)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) | p(s(bool,X15)) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X13,X14) | s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X16),s('type.integer.int',X13))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) ) | $spl16),
  inference(superposition,[],[f572,f302])).
fof(f580,plain,(
  ( ! [X4] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,X4)) | ~sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),'const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),X4)) ) | $spl16),
  inference(resolution,[],[f572,f301])).
fof(f699,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | ~p(s(bool,X2)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),'const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X1)) ) | $spl16),
  inference(forward_demodulation,[],[f698,f354])).
fof(f698,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | ~p(s(bool,X2)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),'const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X1) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl16),
  inference(subsumption_resolution,[],[f579,f523])).
fof(f579,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,X2)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,X2)) | sP0('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')),'const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X1) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl16),
  inference(resolution,[],[f572,f304])).
fof(f578,plain,(
  ( ! [X26,X24,X23,X25,X22] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X22),s('type.integer.int',X23))) | s('type.integer.int',X22) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X26),s('type.integer.int',X23))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X23,X24) | ~p(s(bool,X25)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))),s('type.integer.int',X23)))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X23),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X24),s(bool,X25))))) ) | $spl24),
  inference(cnf_transformation,[],[f578_D])).
fof(f578_D,plain,(
  ( ! [X26,X24,X23,X25,X22] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X22),s('type.integer.int',X23))) | s('type.integer.int',X22) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X26),s('type.integer.int',X23))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X23,X24) | ~p(s(bool,X25)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))),s('type.integer.int',X23)))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X23),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X24),s(bool,X25))))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f567,plain,(
  ( ! [X26,X24,X23,X25,X22] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X22),s('type.integer.int',X23))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X23),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X24),s(bool,X25)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))),s('type.integer.int',X23)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) | ~p(s(bool,X25)) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X23,X24) | s('type.integer.int',X22) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X26),s('type.integer.int',X23))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) ) | $spl3),
  inference(superposition,[],[f547,f304])).
fof(f577,plain,(
  ( ! [X14,X12,X15,X13,X16] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X12),s('type.integer.int',X13))) | s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X16),s('type.integer.int',X13))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X13,X14) | p(s(bool,X15)) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X14),s(bool,X15))))) ) | $spl22),
  inference(cnf_transformation,[],[f577_D])).
fof(f577_D,plain,(
  ( ! [X14,X12,X15,X13,X16] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X12),s('type.integer.int',X13))) | s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X16),s('type.integer.int',X13))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X13,X14) | p(s(bool,X15)) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X14),s(bool,X15))))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f565,plain,(
  ( ! [X14,X12,X15,X13,X16] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X12),s('type.integer.int',X13))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X14),s(bool,X15)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) | p(s(bool,X15)) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)),X13,X14) | s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X16),s('type.integer.int',X13))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) ) | $spl3),
  inference(superposition,[],[f547,f302])).
fof(f570,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) | $spl3),
  inference(forward_demodulation,[],[f568,f353])).
fof(f568,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) | $spl3),
  inference(trivial_inequality_removal,[],[f564])).
fof(f564,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))))) | $spl3),
  inference(superposition,[],[f547,f352])).
fof(f558,plain,(
  ( ! [X26,X24,X23,X25,X22] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X22),s('type.integer.int',X23))) | s('type.integer.int',X22) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X26),s('type.integer.int',X23))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13))))) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13)),X23,X24) | ~p(s(bool,X25)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13))),s('type.integer.int',X23)))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X23),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X24),s(bool,X25))))) ) | $spl14),
  inference(cnf_transformation,[],[f558_D])).
fof(f558_D,plain,(
  ( ! [X26,X24,X23,X25,X22] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X22),s('type.integer.int',X23))) | s('type.integer.int',X22) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X26),s('type.integer.int',X23))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13))))) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13)),X23,X24) | ~p(s(bool,X25)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13))),s('type.integer.int',X23)))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X23),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X24),s(bool,X25))))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f541,plain,(
  ( ! [X26,X24,X23,X25,X22] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X22),s('type.integer.int',X23))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X23),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X24),s(bool,X25)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13))),s('type.integer.int',X23)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13)))))) | ~p(s(bool,X25)) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13)),X23,X24) | s('type.integer.int',X22) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X26),s('type.integer.int',X23))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13)))))) )),
  inference(superposition,[],[f298,f304])).
fof(f557,plain,(
  ( ! [X14,X12,X15,X13,X16] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X12),s('type.integer.int',X13))) | s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X16),s('type.integer.int',X13))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13))))) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13)),X13,X14) | p(s(bool,X15)) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X14),s(bool,X15))))) ) | $spl12),
  inference(cnf_transformation,[],[f557_D])).
fof(f557_D,plain,(
  ( ! [X14,X12,X15,X13,X16] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X12),s('type.integer.int',X13))) | s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X16),s('type.integer.int',X13))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13))))) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13)),X13,X14) | p(s(bool,X15)) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X14),s(bool,X15))))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f539,plain,(
  ( ! [X14,X12,X15,X13,X16] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X12),s('type.integer.int',X13))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X14),s(bool,X15)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13)))))) | p(s(bool,X15)) | sP0('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13)),X13,X14) | s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X16),s('type.integer.int',X13))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13)))))) )),
  inference(superposition,[],[f298,f302])).
fof(f552,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13)))))) | $spl6),
  inference(cnf_transformation,[],[f552_D])).
fof(f552_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13)))))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f554,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl8),
  inference(cnf_transformation,[],[f554_D])).
fof(f554_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f550,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13))))))),
  inference(forward_demodulation,[],[f542,f353])).
fof(f542,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13)))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13))))))),
  inference(trivial_inequality_removal,[],[f538])).
fof(f538,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13)))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int',sK13))))))),
  inference(superposition,[],[f298,f352])).
fof(f547,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))) | $spl3),
  inference(cnf_transformation,[],[f547_D])).
fof(f547_D,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f543,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))))),
  inference(forward_demodulation,[],[f534,f353])).
fof(f534,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13)))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))))))),
  inference(superposition,[],[f298,f352])).
% SZS output end Proof for thm.integer
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 395
% Passive clauses: 3642
% Generated clauses: 9279
% Final active clauses: 238
% Final passive clauses: 1603
% Input formulas: 62
% Initial clauses: 227
% 
% Selected by SInE selection: 59
% SInE iterations: 1
% 
% Duplicate literals: 167
% Trivial inequalities: 75
% Fw subsumption resolutions: 1278
% Fw demodulations: 781
% Bw demodulations: 15
% 
% Simple tautologies: 352
% Equational tautologies: 1
% Forward subsumptions: 2785
% 
% Binary resolution: 3647
% Forward superposition: 1477
% Backward superposition: 1512
% Equality resolution: 12
% 
% Split clauses: 181
% Split components: 410
% SAT solver clauses: 218
% SAT solver unit clauses: 16
% SAT solver binary clauses: 19
% 
% Sat splits: 181
% Sat splitting refutations: 33
% 
% Memory used [KB]: 4221
% Time elapsed: 0.191 s
% ------------------------------
% ------------------------------
% Success in time 11.927 s
