Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.semi_ring
% (60778)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 140
% Passive clauses: 2495
% Generated clauses: 17018
% Final active clauses: 113
% Final passive clauses: 2244
% Input formulas: 36
% Initial clauses: 40
% 
% Duplicate literals: 534
% Trivial inequalities: 130
% Fw subsumption resolutions: 851
% Fw demodulations: 2000
% Bw demodulations: 22
% 
% Simple tautologies: 125
% Equational tautologies: 23
% Forward subsumptions: 10167
% Fw demodulations to eq. taut.: 4
% 
% Binary resolution: 119
% Forward superposition: 6500
% Backward superposition: 6423
% Self superposition: 240
% Equality factoring: 62
% Equality resolution: 7
% 
% Split clauses: 280
% Split components: 560
% SAT solver clauses: 398
% SAT solver unit clauses: 3
% SAT solver binary clauses: 57
% 
% Sat splits: 280
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 398
% TWLsolver calls for satisfiability: 86
% 
% Memory used [KB]: 21492
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 144 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.semi_ring
% (60799)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 137
% Passive clauses: 1289
% Generated clauses: 14423
% Final active clauses: 122
% Final passive clauses: 952
% Input formulas: 36
% Initial clauses: 40
% 
% Duplicate literals: 443
% Trivial inequalities: 69
% Fw subsumption resolutions: 482
% Fw demodulations: 1228
% Bw demodulations: 212
% Condensations: 613
% 
% Simple tautologies: 132
% Equational tautologies: 27
% Forward subsumptions: 8220
% Fw demodulations to eq. taut.: 4
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 163
% Forward superposition: 5502
% Backward superposition: 4283
% Self superposition: 195
% Equality factoring: 40
% Equality resolution: 8
% 
% Split clauses: 228
% Split components: 463
% SAT solver clauses: 1947
% SAT solver unit clauses: 66
% SAT solver binary clauses: 34
% 
% Sat splits: 228
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 1947
% TWLsolver calls for satisfiability: 202
% 
% Memory used [KB]: 26225
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 135 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.semi_ring
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.semi_ring
% SZS output start Proof for thm.semi_ring
fof(f35478,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f35461])).
fof(f35461,plain,(
  s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK1),'type.semi_ring.semi_ring'(sK1)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(sK1,sK1),'const.combin.o_2'(s(fun(sK1,sK1),sK4),s(fun(sK1,sK1),sK3))))),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),sK2))) != s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK1),'type.semi_ring.semi_ring'(sK1)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(sK1,sK1),'const.combin.o_2'(s(fun(sK1,sK1),sK4),s(fun(sK1,sK1),sK3))))),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),sK2)))),
  inference(unit_resulting_resolution,[],[f3438,f138])).
fof(f138,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) != s(X0,X2)) )),
  inference(cnf_transformation,[],[f113])).
fof(f113,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2))))))),
  inference(nnf_transformation,[],[f70])).
fof(f70,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X32,X33,X34] : (p(s(bool,happ(s(fun(X32,bool),happ(s(fun(X32,fun(X32,bool)),'=_0'),s(X32,X33))),s(X32,X34)))) <=> s(X32,X33) = s(X32,X34))),
  file('thm.semi_ring.semi_ring_fupdfupds_comp|split|3',arity_0_HL_EQ)).
fof(f3438,plain,(
  ~p(s(bool,happ(s(fun(fun(sK0,'type.semi_ring.semi_ring'(sK1)),bool),happ(s(fun(fun(sK0,'type.semi_ring.semi_ring'(sK1)),fun(fun(sK0,'type.semi_ring.semi_ring'(sK1)),bool)),'=_0'),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK1),'type.semi_ring.semi_ring'(sK1)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(sK1,sK1),'const.combin.o_2'(s(fun(sK1,sK1),sK4),s(fun(sK1,sK1),sK3))))),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),sK2))))),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK1),'type.semi_ring.semi_ring'(sK1)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(sK1,sK1),'const.combin.o_2'(s(fun(sK1,sK1),sK4),s(fun(sK1,sK1),sK3))))),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),sK2))))))),
  inference(backward_demodulation,[],[f3404,f1312])).
fof(f1312,plain,(
  ~p(s(bool,happ(s(fun(fun(sK0,'type.semi_ring.semi_ring'(sK1)),bool),happ(s(fun(fun(sK0,'type.semi_ring.semi_ring'(sK1)),fun(fun(sK0,'type.semi_ring.semi_ring'(sK1)),bool)),'=_0'),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK1),'type.semi_ring.semi_ring'(sK1)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(sK1,sK1),'const.combin.o_2'(s(fun(sK1,sK1),sK4),s(fun(sK1,sK1),sK3))))),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),sK2))))),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK1),'type.semi_ring.semi_ring'(sK1)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(sK1,sK1),sK4))),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK1),'type.semi_ring.semi_ring'(sK1)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(sK1,sK1),sK3))),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),sK2))))))))),
  inference(unit_resulting_resolution,[],[f114,f137])).
fof(f137,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) )),
  inference(cnf_transformation,[],[f113])).
fof(f114,plain,(
  s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK1),'type.semi_ring.semi_ring'(sK1)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(sK1,sK1),'const.combin.o_2'(s(fun(sK1,sK1),sK4),s(fun(sK1,sK1),sK3))))),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),sK2))) != s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK1),'type.semi_ring.semi_ring'(sK1)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(sK1,sK1),sK4))),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK1),'type.semi_ring.semi_ring'(sK1)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(sK1,sK1),sK3))),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),sK2)))))),
  inference(cnf_transformation,[],[f103])).
fof(f103,plain,(
  s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK1),'type.semi_ring.semi_ring'(sK1)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(sK1,sK1),'const.combin.o_2'(s(fun(sK1,sK1),sK4),s(fun(sK1,sK1),sK3))))),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),sK2))) != s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK1),'type.semi_ring.semi_ring'(sK1)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(sK1,sK1),sK4))),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK1),'type.semi_ring.semi_ring'(sK1)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(sK1,sK1),sK3))),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),sK2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4])],[f101])).
fof(f101,plain,(
  ? [X0,X1,X2,X3,X4] : s(fun(X0,'type.semi_ring.semi_ring'(X1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X1),'type.semi_ring.semi_ring'(X1)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(X1,X1),'const.combin.o_2'(s(fun(X1,X1),X4),s(fun(X1,X1),X3))))),s(fun(X0,'type.semi_ring.semi_ring'(X1)),X2))) != s(fun(X0,'type.semi_ring.semi_ring'(X1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X1),'type.semi_ring.semi_ring'(X1)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(X1,X1),X4))),s(fun(X0,'type.semi_ring.semi_ring'(X1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X1),'type.semi_ring.semi_ring'(X1)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(X1,X1),X3))),s(fun(X0,'type.semi_ring.semi_ring'(X1)),X2)))))),
  inference(ennf_transformation,[],[f39])).
fof(f39,plain,(
  ~! [X0,X1,X2,X3,X4] : s(fun(X0,'type.semi_ring.semi_ring'(X1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X1),'type.semi_ring.semi_ring'(X1)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(X1,X1),'const.combin.o_2'(s(fun(X1,X1),X4),s(fun(X1,X1),X3))))),s(fun(X0,'type.semi_ring.semi_ring'(X1)),X2))) = s(fun(X0,'type.semi_ring.semi_ring'(X1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X1),'type.semi_ring.semi_ring'(X1)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(X1,X1),X4))),s(fun(X0,'type.semi_ring.semi_ring'(X1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X1),'type.semi_ring.semi_ring'(X1)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(X1,X1),X3))),s(fun(X0,'type.semi_ring.semi_ring'(X1)),X2)))))),
  inference(flattening,[],[f38])).
fof(f38,plain,(
  ~! [X0,X1] : ! [X2,X3,X4] : s(fun(X0,'type.semi_ring.semi_ring'(X1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X1),'type.semi_ring.semi_ring'(X1)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(X1,X1),'const.combin.o_2'(s(fun(X1,X1),X4),s(fun(X1,X1),X3))))),s(fun(X0,'type.semi_ring.semi_ring'(X1)),X2))) = s(fun(X0,'type.semi_ring.semi_ring'(X1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X1),'type.semi_ring.semi_ring'(X1)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(X1,X1),X4))),s(fun(X0,'type.semi_ring.semi_ring'(X1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X1),'type.semi_ring.semi_ring'(X1)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(X1,X1),X3))),s(fun(X0,'type.semi_ring.semi_ring'(X1)),X2)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10] : s(fun(X6,'type.semi_ring.semi_ring'(X7)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X7),'type.semi_ring.semi_ring'(X7)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(X7,X7),'const.combin.o_2'(s(fun(X7,X7),X10),s(fun(X7,X7),X9))))),s(fun(X6,'type.semi_ring.semi_ring'(X7)),X8))) = s(fun(X6,'type.semi_ring.semi_ring'(X7)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X7),'type.semi_ring.semi_ring'(X7)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(X7,X7),X10))),s(fun(X6,'type.semi_ring.semi_ring'(X7)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X7),'type.semi_ring.semi_ring'(X7)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(X7,X7),X9))),s(fun(X6,'type.semi_ring.semi_ring'(X7)),X8)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10] : s(fun(X6,'type.semi_ring.semi_ring'(X7)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X7),'type.semi_ring.semi_ring'(X7)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(X7,X7),'const.combin.o_2'(s(fun(X7,X7),X10),s(fun(X7,X7),X9))))),s(fun(X6,'type.semi_ring.semi_ring'(X7)),X8))) = s(fun(X6,'type.semi_ring.semi_ring'(X7)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X7),'type.semi_ring.semi_ring'(X7)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(X7,X7),X10))),s(fun(X6,'type.semi_ring.semi_ring'(X7)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X7),'type.semi_ring.semi_ring'(X7)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(X7,X7),X9))),s(fun(X6,'type.semi_ring.semi_ring'(X7)),X8)))))),
  file('thm.semi_ring.semi_ring_fupdfupds_comp|split|3','thm.semi_ring.semi_ring_fupdfupds_comp|split|3')).
fof(f3404,plain,(
  ( ! [X47,X45,X48,X46,X44] : (s(fun(X47,'type.semi_ring.semi_ring'(X44)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X44),'type.semi_ring.semi_ring'(X44)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(X44,X44),'const.combin.o_2'(s(fun(X44,X44),X45),s(fun(X44,X44),X46))))),s(fun(X47,'type.semi_ring.semi_ring'(X44)),X48))) = s(fun(X47,'type.semi_ring.semi_ring'(X44)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X44),'type.semi_ring.semi_ring'(X44)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(X44,X44),X45))),s(fun(X47,'type.semi_ring.semi_ring'(X44)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X44),'type.semi_ring.semi_ring'(X44)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(X44,X44),X46))),s(fun(X47,'type.semi_ring.semi_ring'(X44)),X48)))))) )),
  inference(superposition,[],[f142,f121])).
fof(f121,plain,(
  ( ! [X2,X0,X1] : (s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(X0,X0),'const.combin.o_2'(s(fun(X0,X0),X2),s(fun(X0,X0),X1))))) = s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(X0,X0),X2))),s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(X0,X0),X1)))))) )),
  inference(cnf_transformation,[],[f48])).
fof(f48,plain,(
  ! [X0,X1,X2] : s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(X0,X0),'const.combin.o_2'(s(fun(X0,X0),X2),s(fun(X0,X0),X1))))) = s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(X0,X0),X2))),s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(X0,X0),X1)))))),
  inference(flattening,[],[f47])).
fof(f47,plain,(
  ! [X0] : ! [X1,X2] : s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(X0,X0),'const.combin.o_2'(s(fun(X0,X0),X2),s(fun(X0,X0),X1))))) = s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(X0,X0),X2))),s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(X0,X0),X1)))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X7] : ! [X9,X10] : s(fun('type.semi_ring.semi_ring'(X7),'type.semi_ring.semi_ring'(X7)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(X7,X7),'const.combin.o_2'(s(fun(X7,X7),X10),s(fun(X7,X7),X9))))) = s(fun('type.semi_ring.semi_ring'(X7),'type.semi_ring.semi_ring'(X7)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X7),'type.semi_ring.semi_ring'(X7)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(X7,X7),X10))),s(fun('type.semi_ring.semi_ring'(X7),'type.semi_ring.semi_ring'(X7)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(X7,X7),X9)))))),
  file('thm.semi_ring.semi_ring_fupdfupds_comp|split|3','thm.semi_ring.semi_ring_fupdfupds_comp|split|2')).
fof(f142,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))) )),
  inference(cnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))),
  inference(flattening,[],[f77])).
fof(f77,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6] : s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6,X7,X13,X14] : ! [X10,X9,X8] : s(fun(X13,X6),'const.combin.o_2'(s(fun(X7,X6),X10),s(fun(X13,X7),'const.combin.o_2'(s(fun(X14,X7),X9),s(fun(X13,X14),X8))))) = s(fun(X13,X6),'const.combin.o_2'(s(fun(X14,X6),'const.combin.o_2'(s(fun(X7,X6),X10),s(fun(X14,X7),X9))),s(fun(X13,X14),X8)))),
  file('thm.semi_ring.semi_ring_fupdfupds_comp|split|3','thm.combin.o_ASSOC')).
% SZS output end Proof for thm.semi_ring
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 254
% Passive clauses: 4743
% Generated clauses: 39250
% Final active clauses: 99
% Final passive clauses: 1493
% Input formulas: 36
% Initial clauses: 40
% 
% Duplicate literals: 937
% Trivial inequalities: 79
% Fw subsumption resolutions: 1581
% Fw demodulations: 8137
% Bw demodulations: 2742
% Condensations: 1254
% Global subsumptions: 18
% 
% Simple tautologies: 196
% Equational tautologies: 141
% Forward subsumptions: 21737
% Fw demodulations to eq. taut.: 11
% Bw demodulations to eq. taut.: 43
% 
% Unit resulting resolution: 1596
% Forward superposition: 11254
% Backward superposition: 7061
% Self superposition: 235
% Equality factoring: 64
% Equality resolution: 3
% 
% Split clauses: 120
% Split components: 247
% SAT solver clauses: 5293
% SAT solver unit clauses: 575
% SAT solver binary clauses: 2837
% 
% Sat splits: 120
% Sat splitting refutations: 290
% 
% Memory used [KB]: 39786
% Time elapsed: 2.389 s
% ------------------------------
% ------------------------------
% Success in time 3.79 s
