Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.ringNorm
% (75224)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 136
% Passive clauses: 2542
% Generated clauses: 25833
% Final active clauses: 131
% Final passive clauses: 2273
% Input formulas: 27
% Initial clauses: 28
% 
% Duplicate literals: 1528
% Trivial inequalities: 57
% Fw subsumption resolutions: 2156
% Fw demodulations: 4227
% Bw demodulations: 3
% 
% Simple tautologies: 205
% Equational tautologies: 19
% Forward subsumptions: 13600
% 
% Binary resolution: 13
% Forward superposition: 9837
% Backward superposition: 7535
% Self superposition: 200
% Equality factoring: 36
% Equality resolution: 5
% 
% Split clauses: 774
% Split components: 1553
% SAT solver clauses: 1260
% SAT solver unit clauses: 1
% SAT solver binary clauses: 20
% 
% Sat splits: 774
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1260
% TWLsolver calls for satisfiability: 86
% 
% Memory used [KB]: 22771
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.ringNorm
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.ringNorm
% SZS output start Proof for thm.ringNorm
fof(f9178,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f9150,f83,f112,f8710,f8857,f93,f122,f3015,f3016,f88,f119,f104,f133,f105,f134])).
fof(f134,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.ringNorm.r_ics_aux_4'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s(X0,X3),s('type.canonical.canonical_sum'(X0),X4))) = s(X0,happ(s(fun('type.canonical.canonical_sum'(X0),X0),happ(s(fun(X0,fun('type.canonical.canonical_sum'(X0),X0)),happ(s(fun('type.quote.varmap'(X0),fun(X0,fun('type.canonical.canonical_sum'(X0),X0))),'const.ringNorm.r_ics_aux_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s(X0,X3))),s('type.canonical.canonical_sum'(X0),X4)))) ) | $spl40),
  inference(cnf_transformation,[],[f134_D])).
fof(f134_D,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.ringNorm.r_ics_aux_4'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s(X0,X3),s('type.canonical.canonical_sum'(X0),X4))) = s(X0,happ(s(fun('type.canonical.canonical_sum'(X0),X0),happ(s(fun(X0,fun('type.canonical.canonical_sum'(X0),X0)),happ(s(fun('type.quote.varmap'(X0),fun(X0,fun('type.canonical.canonical_sum'(X0),X0))),'const.ringNorm.r_ics_aux_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s(X0,X3))),s('type.canonical.canonical_sum'(X0),X4)))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f105,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.ringNorm.r_ics_aux_4'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s(X0,X3),s('type.canonical.canonical_sum'(X0),X4))) = s(X0,happ(s(fun('type.canonical.canonical_sum'(X0),X0),happ(s(fun(X0,fun('type.canonical.canonical_sum'(X0),X0)),happ(s(fun('type.quote.varmap'(X0),fun(X0,fun('type.canonical.canonical_sum'(X0),X0))),'const.ringNorm.r_ics_aux_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s(X0,X3))),s('type.canonical.canonical_sum'(X0),X4)))) )),
  inference(cnf_transformation,[],[f67])).
fof(f67,plain,(
  ! [X0,X1,X2,X3,X4] : s(X0,'const.ringNorm.r_ics_aux_4'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s(X0,X3),s('type.canonical.canonical_sum'(X0),X4))) = s(X0,happ(s(fun('type.canonical.canonical_sum'(X0),X0),happ(s(fun(X0,fun('type.canonical.canonical_sum'(X0),X0)),happ(s(fun('type.quote.varmap'(X0),fun(X0,fun('type.canonical.canonical_sum'(X0),X0))),'const.ringNorm.r_ics_aux_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s(X0,X3))),s('type.canonical.canonical_sum'(X0),X4)))),
  inference(flattening,[],[f66])).
fof(f66,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s(X0,'const.ringNorm.r_ics_aux_4'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s(X0,X3),s('type.canonical.canonical_sum'(X0),X4))) = s(X0,happ(s(fun('type.canonical.canonical_sum'(X0),X0),happ(s(fun(X0,fun('type.canonical.canonical_sum'(X0),X0)),happ(s(fun('type.quote.varmap'(X0),fun(X0,fun('type.canonical.canonical_sum'(X0),X0))),'const.ringNorm.r_ics_aux_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s(X0,X3))),s('type.canonical.canonical_sum'(X0),X4)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X6] : ! [X19,X20,X21,X22] : s(X6,'const.ringNorm.r_ics_aux_4'(s('type.ring.ring'(X6),X19),s('type.quote.varmap'(X6),X20),s(X6,X21),s('type.canonical.canonical_sum'(X6),X22))) = s(X6,happ(s(fun('type.canonical.canonical_sum'(X6),X6),happ(s(fun(X6,fun('type.canonical.canonical_sum'(X6),X6)),happ(s(fun('type.quote.varmap'(X6),fun(X6,fun('type.canonical.canonical_sum'(X6),X6))),'const.ringNorm.r_ics_aux_1'(s('type.ring.ring'(X6),X19))),s('type.quote.varmap'(X6),X20))),s(X6,X21))),s('type.canonical.canonical_sum'(X6),X22)))),
  file('thm.ringNorm.ics_aux_def|split|0','ARITY_1_4_const.ringNorm.r_ics_aux')).
fof(f133,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.canonical.ics_aux_4'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s(X0,X3),s('type.canonical.canonical_sum'(X0),X4))) = s(X0,happ(s(fun('type.canonical.canonical_sum'(X0),X0),happ(s(fun(X0,fun('type.canonical.canonical_sum'(X0),X0)),happ(s(fun('type.quote.varmap'(X0),fun(X0,fun('type.canonical.canonical_sum'(X0),X0))),'const.canonical.ics_aux_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s(X0,X3))),s('type.canonical.canonical_sum'(X0),X4)))) ) | $spl38),
  inference(cnf_transformation,[],[f133_D])).
fof(f133_D,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.canonical.ics_aux_4'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s(X0,X3),s('type.canonical.canonical_sum'(X0),X4))) = s(X0,happ(s(fun('type.canonical.canonical_sum'(X0),X0),happ(s(fun(X0,fun('type.canonical.canonical_sum'(X0),X0)),happ(s(fun('type.quote.varmap'(X0),fun(X0,fun('type.canonical.canonical_sum'(X0),X0))),'const.canonical.ics_aux_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s(X0,X3))),s('type.canonical.canonical_sum'(X0),X4)))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f104,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.canonical.ics_aux_4'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s(X0,X3),s('type.canonical.canonical_sum'(X0),X4))) = s(X0,happ(s(fun('type.canonical.canonical_sum'(X0),X0),happ(s(fun(X0,fun('type.canonical.canonical_sum'(X0),X0)),happ(s(fun('type.quote.varmap'(X0),fun(X0,fun('type.canonical.canonical_sum'(X0),X0))),'const.canonical.ics_aux_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s(X0,X3))),s('type.canonical.canonical_sum'(X0),X4)))) )),
  inference(cnf_transformation,[],[f65])).
fof(f65,plain,(
  ! [X0,X1,X2,X3,X4] : s(X0,'const.canonical.ics_aux_4'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s(X0,X3),s('type.canonical.canonical_sum'(X0),X4))) = s(X0,happ(s(fun('type.canonical.canonical_sum'(X0),X0),happ(s(fun(X0,fun('type.canonical.canonical_sum'(X0),X0)),happ(s(fun('type.quote.varmap'(X0),fun(X0,fun('type.canonical.canonical_sum'(X0),X0))),'const.canonical.ics_aux_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s(X0,X3))),s('type.canonical.canonical_sum'(X0),X4)))),
  inference(flattening,[],[f64])).
fof(f64,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s(X0,'const.canonical.ics_aux_4'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s(X0,X3),s('type.canonical.canonical_sum'(X0),X4))) = s(X0,happ(s(fun('type.canonical.canonical_sum'(X0),X0),happ(s(fun(X0,fun('type.canonical.canonical_sum'(X0),X0)),happ(s(fun('type.quote.varmap'(X0),fun(X0,fun('type.canonical.canonical_sum'(X0),X0))),'const.canonical.ics_aux_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s(X0,X3))),s('type.canonical.canonical_sum'(X0),X4)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X6] : ! [X19,X20,X21,X22] : s(X6,'const.canonical.ics_aux_4'(s('type.semi_ring.semi_ring'(X6),X19),s('type.quote.varmap'(X6),X20),s(X6,X21),s('type.canonical.canonical_sum'(X6),X22))) = s(X6,happ(s(fun('type.canonical.canonical_sum'(X6),X6),happ(s(fun(X6,fun('type.canonical.canonical_sum'(X6),X6)),happ(s(fun('type.quote.varmap'(X6),fun(X6,fun('type.canonical.canonical_sum'(X6),X6))),'const.canonical.ics_aux_1'(s('type.semi_ring.semi_ring'(X6),X19))),s('type.quote.varmap'(X6),X20))),s(X6,X21))),s('type.canonical.canonical_sum'(X6),X22)))),
  file('thm.ringNorm.ics_aux_def|split|0','ARITY_1_4_const.canonical.ics_aux')).
fof(f119,plain,(
  ( ! [X0,X1] : (s(fun('type.quote.varmap'(X0),fun(X0,fun('type.canonical.canonical_sum'(X0),X0))),'const.ringNorm.r_ics_aux_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.quote.varmap'(X0),fun(X0,fun('type.canonical.canonical_sum'(X0),X0))),'const.canonical.ics_aux_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))) ) | $spl10),
  inference(cnf_transformation,[],[f119_D])).
fof(f119_D,plain,(
  ( ! [X0,X1] : (s(fun('type.quote.varmap'(X0),fun(X0,fun('type.canonical.canonical_sum'(X0),X0))),'const.ringNorm.r_ics_aux_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.quote.varmap'(X0),fun(X0,fun('type.canonical.canonical_sum'(X0),X0))),'const.canonical.ics_aux_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f88,plain,(
  ( ! [X0,X1] : (s(fun('type.quote.varmap'(X0),fun(X0,fun('type.canonical.canonical_sum'(X0),X0))),'const.ringNorm.r_ics_aux_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.quote.varmap'(X0),fun(X0,fun('type.canonical.canonical_sum'(X0),X0))),'const.canonical.ics_aux_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0,X1] : s(fun('type.quote.varmap'(X0),fun(X0,fun('type.canonical.canonical_sum'(X0),X0))),'const.ringNorm.r_ics_aux_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.quote.varmap'(X0),fun(X0,fun('type.canonical.canonical_sum'(X0),X0))),'const.canonical.ics_aux_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f34])).
fof(f34,plain,(
  ! [X0] : ! [X1] : s(fun('type.quote.varmap'(X0),fun(X0,fun('type.canonical.canonical_sum'(X0),X0))),'const.ringNorm.r_ics_aux_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.quote.varmap'(X0),fun(X0,fun('type.canonical.canonical_sum'(X0),X0))),'const.canonical.ics_aux_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6] : ! [X7] : s(fun('type.quote.varmap'(X6),fun(X6,fun('type.canonical.canonical_sum'(X6),X6))),'const.ringNorm.r_ics_aux_1'(s('type.ring.ring'(X6),X7))) = s(fun('type.quote.varmap'(X6),fun(X6,fun('type.canonical.canonical_sum'(X6),X6))),'const.canonical.ics_aux_1'(s('type.semi_ring.semi_ring'(X6),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X6),X7)))))),
  file('thm.ringNorm.ics_aux_def|split|0','thm.ringNorm.r_ics_aux_def')).
fof(f3016,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.ringNorm.r_ics_aux_4'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s(X0,X3),s('type.canonical.canonical_sum'(X0),X4))) = s(X0,'const.canonical.ics_aux_4'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2),s(X0,X3),s('type.canonical.canonical_sum'(X0),X4)))) ) | $spl456),
  inference(cnf_transformation,[],[f3016_D])).
fof(f3016_D,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.ringNorm.r_ics_aux_4'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s(X0,X3),s('type.canonical.canonical_sum'(X0),X4))) = s(X0,'const.canonical.ics_aux_4'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2),s(X0,X3),s('type.canonical.canonical_sum'(X0),X4)))) ) <=> ~$spl456),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl456])])).
fof(f3015,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.ringNorm.r_ics_aux_4'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s(X0,X3),s('type.canonical.canonical_sum'(X0),X4))) = s(X0,'const.canonical.ics_aux_4'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2),s(X0,X3),s('type.canonical.canonical_sum'(X0),X4)))) ) | ($spl10 | $spl38 | $spl40)),
  inference(forward_demodulation,[],[f2875,f134])).
fof(f2875,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.canonical.ics_aux_4'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2),s(X0,X3),s('type.canonical.canonical_sum'(X0),X4))) = s(X0,happ(s(fun('type.canonical.canonical_sum'(X0),X0),happ(s(fun(X0,fun('type.canonical.canonical_sum'(X0),X0)),happ(s(fun('type.quote.varmap'(X0),fun(X0,fun('type.canonical.canonical_sum'(X0),X0))),'const.ringNorm.r_ics_aux_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s(X0,X3))),s('type.canonical.canonical_sum'(X0),X4)))) ) | ($spl10 | $spl38)),
  inference(superposition,[],[f133,f119])).
fof(f122,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X3) = s(X0,'const.canonical.ics_aux_4'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s(X0,X3),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))) ) | $spl16),
  inference(cnf_transformation,[],[f122_D])).
fof(f122_D,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X3) = s(X0,'const.canonical.ics_aux_4'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s(X0,X3),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f93,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X3) = s(X0,'const.canonical.ics_aux_4'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s(X0,X3),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))) )),
  inference(cnf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0,X1,X2,X3] : s(X0,X3) = s(X0,'const.canonical.ics_aux_4'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s(X0,X3),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))),
  inference(flattening,[],[f42])).
fof(f42,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,X3) = s(X0,'const.canonical.ics_aux_4'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s(X0,X3),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : ! [X12,X8,X9] : s(X6,X9) = s(X6,'const.canonical.ics_aux_4'(s('type.semi_ring.semi_ring'(X6),X12),s('type.quote.varmap'(X6),X8),s(X6,X9),s('type.canonical.canonical_sum'(X6),'const.canonical.Nil_monom_0')))),
  file('thm.ringNorm.ics_aux_def|split|0','thm.canonical.ics_aux_def|split|0')).
fof(f8857,plain,(
  ( ! [X6,X8,X7,X5] : (s(X5,X8) = s(X5,'const.ringNorm.r_ics_aux_4'(s('type.ring.ring'(X5),X6),s('type.quote.varmap'(X5),X7),s(X5,X8),s('type.canonical.canonical_sum'(X5),'const.canonical.Nil_monom_0')))) ) | $spl1318),
  inference(cnf_transformation,[],[f8857_D])).
fof(f8857_D,plain,(
  ( ! [X6,X8,X7,X5] : (s(X5,X8) = s(X5,'const.ringNorm.r_ics_aux_4'(s('type.ring.ring'(X5),X6),s('type.quote.varmap'(X5),X7),s(X5,X8),s('type.canonical.canonical_sum'(X5),'const.canonical.Nil_monom_0')))) ) <=> ~$spl1318),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1318])])).
fof(f8710,plain,(
  ( ! [X6,X8,X7,X5] : (s(X5,X8) = s(X5,'const.ringNorm.r_ics_aux_4'(s('type.ring.ring'(X5),X6),s('type.quote.varmap'(X5),X7),s(X5,X8),s('type.canonical.canonical_sum'(X5),'const.canonical.Nil_monom_0')))) ) | ($spl16 | $spl456)),
  inference(superposition,[],[f3016,f122])).
fof(f112,plain,(
  s(sK0,sK3) != s(sK0,'const.ringNorm.r_ics_aux_4'(s('type.ring.ring'(sK0),sK1),s('type.quote.varmap'(sK0),sK2),s(sK0,sK3),s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0'))) | $spl1),
  inference(cnf_transformation,[],[f112_D])).
fof(f112_D,plain,(
  s(sK0,sK3) != s(sK0,'const.ringNorm.r_ics_aux_4'(s('type.ring.ring'(sK0),sK1),s('type.quote.varmap'(sK0),sK2),s(sK0,sK3),s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0'))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f83,plain,(
  s(sK0,sK3) != s(sK0,'const.ringNorm.r_ics_aux_4'(s('type.ring.ring'(sK0),sK1),s('type.quote.varmap'(sK0),sK2),s(sK0,sK3),s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0')))),
  inference(cnf_transformation,[],[f80])).
fof(f80,plain,(
  s(sK0,sK3) != s(sK0,'const.ringNorm.r_ics_aux_4'(s('type.ring.ring'(sK0),sK1),s('type.quote.varmap'(sK0),sK2),s(sK0,sK3),s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0')))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f78])).
fof(f78,plain,(
  ? [X0,X1,X2,X3] : s(X0,X3) != s(X0,'const.ringNorm.r_ics_aux_4'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s(X0,X3),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))),
  inference(ennf_transformation,[],[f30])).
fof(f30,plain,(
  ~! [X0,X1,X2,X3] : s(X0,X3) = s(X0,'const.ringNorm.r_ics_aux_4'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s(X0,X3),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))),
  inference(flattening,[],[f29])).
fof(f29,plain,(
  ~! [X0] : ! [X1,X2,X3] : s(X0,X3) = s(X0,'const.ringNorm.r_ics_aux_4'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s(X0,X3),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : s(X6,X9) = s(X6,'const.ringNorm.r_ics_aux_4'(s('type.ring.ring'(X6),X7),s('type.quote.varmap'(X6),X8),s(X6,X9),s('type.canonical.canonical_sum'(X6),'const.canonical.Nil_monom_0')))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : s(X6,X9) = s(X6,'const.ringNorm.r_ics_aux_4'(s('type.ring.ring'(X6),X7),s('type.quote.varmap'(X6),X8),s(X6,X9),s('type.canonical.canonical_sum'(X6),'const.canonical.Nil_monom_0')))),
  file('thm.ringNorm.ics_aux_def|split|0','thm.ringNorm.ics_aux_def|split|0')).
fof(f9150,plain,(
  $false | ($spl1 | $spl1318)),
  inference(trivial_inequality_removal,[],[f9149])).
fof(f9149,plain,(
  s(sK0,sK3) != s(sK0,sK3) | ($spl1 | $spl1318)),
  inference(superposition,[],[f112,f8857])).
% SZS output end Proof for thm.ringNorm
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 84
% Passive clauses: 672
% Generated clauses: 9011
% Final active clauses: 79
% Final passive clauses: 578
% Input formulas: 27
% Initial clauses: 28
% 
% Duplicate literals: 406
% Trivial inequalities: 22
% Fw subsumption resolutions: 264
% Fw demodulations: 797
% Bw demodulations: 11
% Condensations: 530
% 
% Simple tautologies: 91
% Equational tautologies: 16
% Forward subsumptions: 5238
% 
% Binary resolution: 8
% Forward superposition: 3617
% Backward superposition: 2488
% Self superposition: 128
% Equality factoring: 30
% Equality resolution: 2
% 
% Split clauses: 14
% Split components: 33
% SAT solver clauses: 963
% SAT solver unit clauses: 50
% SAT solver binary clauses: 14
% SAT solver learnt clauses: 1
% 
% Sat splits: 14
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 962
% TWLsolver calls for satisfiability: 105
% 
% Memory used [KB]: 13944
% Time elapsed: 0.388 s
% ------------------------------
% ------------------------------
% Success in time 1.024 s
