Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.integer
% (116082)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 734
% Passive clauses: 8419
% Generated clauses: 19718
% Final active clauses: 677
% Final passive clauses: 7018
% Input formulas: 52
% Initial clauses: 207
% 
% Duplicate literals: 439
% Trivial inequalities: 118
% Fw subsumption resolutions: 1862
% Fw demodulations: 556
% Bw demodulations: 49
% 
% Simple tautologies: 256
% Equational tautologies: 62
% Forward subsumptions: 7419
% 
% Binary resolution: 9976
% Factoring: 1284
% Forward superposition: 3039
% Backward superposition: 1818
% Self superposition: 35
% Equality factoring: 28
% Equality resolution: 54
% 
% Split clauses: 570
% Split components: 1147
% SAT solver clauses: 1018
% SAT solver unit clauses: 24
% SAT solver binary clauses: 179
% SAT solver learnt clauses: 1
% SAT solver learnt literals: 1
% 
% Sat splits: 570
% Sat splitting refutations: 6
% 
% TWLsolver clauses: 1015
% TWLsolver calls for satisfiability: 402
% 
% Memory used [KB]: 8443
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.integer
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.integer
% SZS output start Proof for thm.integer
fof(f10729,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f10691,f268,f487,f9067,f6201,f7021,f7041,f267,f488,f486,f2275,f2289,f392,f535,f349,f518,f6198,f6227,f2215,f2224,f393,f534,f6723,f489])).
fof(f489,plain,(
  s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK13),s('type.integer.int',sK14))) | $spl3),
  inference(cnf_transformation,[],[f489_D])).
fof(f489_D,plain,(
  s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK13),s('type.integer.int',sK14))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f6723,plain,(
  s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK13),s('type.integer.int',sK14))) | ($spl22 | $spl1801)),
  inference(superposition,[],[f5835,f518])).
fof(f5835,plain,(
  s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) | $spl1801),
  inference(cnf_transformation,[],[f5835_D])).
fof(f5835_D,plain,(
  s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) <=> ~$spl1801),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1801])])).
fof(f534,plain,(
  ( ! [X2,X0,X1] : (s('type.integer.int',X2) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X2)))) ) | $spl44),
  inference(cnf_transformation,[],[f534_D])).
fof(f534_D,plain,(
  ( ! [X2,X0,X1] : (s('type.integer.int',X2) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X2)))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f393,plain,(
  ( ! [X2,X0,X1] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X2))) | s('type.integer.int',X2) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) )),
  inference(cnf_transformation,[],[f229])).
fof(f229,plain,(
  ! [X0,X1,X2] : ((s('type.integer.int',X0) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X2))) | s('type.integer.int',X2) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & (s('type.integer.int',X2) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X2)))))),
  inference(nnf_transformation,[],[f115])).
fof(f115,plain,(
  ! [X0,X1,X2] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X2))) <=> s('type.integer.int',X2) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1))))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X6,X7,X8] : (s('type.integer.int',X6) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X7))),s('type.integer.int',X8))) <=> s('type.integer.int',X8) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X6),s('type.integer.int',X7))))),
  file('thm.integer.integer_325_prove_ltb','thm.integer.integer_323_prove_ltb')).
fof(f2224,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))) ) | $spl830),
  inference(cnf_transformation,[],[f2224_D])).
fof(f2224_D,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))) ) <=> ~$spl830),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl830])])).
fof(f2215,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))) ) | $spl44),
  inference(equality_resolution,[],[f534])).
fof(f6227,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12))) | $spl1982),
  inference(cnf_transformation,[],[f6227_D])).
fof(f6227_D,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12))) <=> ~$spl1982),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1982])])).
fof(f6198,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12))) | ($spl830 | $spl1800)),
  inference(superposition,[],[f2224,f5836])).
fof(f5836,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) | $spl1800),
  inference(cnf_transformation,[],[f5836_D])).
fof(f5836_D,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) <=> ~$spl1800),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1800])])).
fof(f518,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) ) | $spl22),
  inference(cnf_transformation,[],[f518_D])).
fof(f518_D,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f349,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) )),
  inference(cnf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0,X1] : s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',X0)))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X7,X6] : s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X6),s('type.integer.int',X7))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X7),s('type.integer.int',X6)))),
  file('thm.integer.integer_325_prove_ltb','thm.integer.INT_ADD_SYM')).
fof(f535,plain,(
  ( ! [X2,X0,X1] : (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X2))) | s('type.integer.int',X2) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) ) | $spl46),
  inference(cnf_transformation,[],[f535_D])).
fof(f535_D,plain,(
  ( ! [X2,X0,X1] : (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X2))) | s('type.integer.int',X2) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f392,plain,(
  ( ! [X2,X0,X1] : (s('type.integer.int',X2) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1))) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X2)))) )),
  inference(cnf_transformation,[],[f229])).
fof(f2289,plain,(
  ( ! [X2,X3,X1] : (s('type.integer.int',X3) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))))) | s('type.integer.int',X2) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X3),s('type.integer.int',X1)))) ) | $spl858),
  inference(cnf_transformation,[],[f2289_D])).
fof(f2289_D,plain,(
  ( ! [X2,X3,X1] : (s('type.integer.int',X3) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))))) | s('type.integer.int',X2) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X3),s('type.integer.int',X1)))) ) <=> ~$spl858),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl858])])).
fof(f2275,plain,(
  ( ! [X2,X3,X1] : (s('type.integer.int',X3) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))))) | s('type.integer.int',X2) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X3),s('type.integer.int',X1)))) ) | ($spl22 | $spl46)),
  inference(superposition,[],[f535,f518])).
fof(f486,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))))) | $spl0),
  inference(cnf_transformation,[],[f486_D])).
fof(f486_D,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f488,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK13),s('type.integer.int',sK14))) | $spl2),
  inference(cnf_transformation,[],[f488_D])).
fof(f488_D,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK13),s('type.integer.int',sK14))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f267,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK13),s('type.integer.int',sK14))) | s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))),
  inference(cnf_transformation,[],[f169])).
fof(f169,plain,(
  (s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))))) | s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK13),s('type.integer.int',sK14)))) & (s('type.integer.int',sK14) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))))) | s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK13),s('type.integer.int',sK14))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13,sK14])],[f168])).
fof(f168,plain,(
  ? [X0,X1,X2] : ((s('type.integer.int',X2) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',X2)))) & (s('type.integer.int',X2) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))))) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',X2)))))),
  inference(nnf_transformation,[],[f125])).
fof(f125,plain,(
  ? [X0,X1,X2] : (s('type.integer.int',X2) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))))) <~> s('type.integer.int',X0) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',X2))))),
  inference(ennf_transformation,[],[f54])).
fof(f54,plain,(
  ~! [X0,X1,X2] : (s('type.integer.int',X2) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))))) <=> s('type.integer.int',X0) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',X2))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : (s('type.integer.int',X8) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X6),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X7))))) <=> s('type.integer.int',X6) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X7),s('type.integer.int',X8))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : (s('type.integer.int',X8) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X6),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X7))))) <=> s('type.integer.int',X6) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X7),s('type.integer.int',X8))))),
  file('thm.integer.integer_325_prove_ltb','thm.integer.integer_325_prove_ltb')).
fof(f7041,plain,(
  ( ! [X9] : (s('type.integer.int',sK14) != s('type.integer.int',X9) | s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X9),s('type.integer.int',sK13)))) ) | $spl2310),
  inference(cnf_transformation,[],[f7041_D])).
fof(f7041_D,plain,(
  ( ! [X9] : (s('type.integer.int',sK14) != s('type.integer.int',X9) | s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X9),s('type.integer.int',sK13)))) ) <=> ~$spl2310),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2310])])).
fof(f7021,plain,(
  ( ! [X9] : (s('type.integer.int',sK14) != s('type.integer.int',X9) | s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X9),s('type.integer.int',sK13)))) ) | ($spl0 | $spl858)),
  inference(superposition,[],[f2289,f486])).
fof(f6201,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK13),s('type.integer.int',sK14))) | ($spl22 | $spl1800)),
  inference(superposition,[],[f518,f5836])).
fof(f9067,plain,(
  $false | ($spl1801 | $spl2310)),
  inference(subsumption_resolution,[],[f9057,f5835])).
fof(f9057,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK14),s('type.integer.int',sK13))) | $spl2310),
  inference(equality_resolution,[],[f7041])).
fof(f487,plain,(
  s('type.integer.int',sK14) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))))) | $spl1),
  inference(cnf_transformation,[],[f487_D])).
fof(f487_D,plain,(
  s('type.integer.int',sK14) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f268,plain,(
  s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK13),s('type.integer.int',sK14))) | s('type.integer.int',sK14) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))),
  inference(cnf_transformation,[],[f169])).
fof(f10691,plain,(
  $false | ($spl1 | $spl22 | $spl1982)),
  inference(subsumption_resolution,[],[f10625,f487])).
fof(f10625,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))))) | ($spl22 | $spl1982)),
  inference(superposition,[],[f6227,f518])).
% SZS output end Proof for thm.integer
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 394
% Passive clauses: 1909
% Generated clauses: 11483
% Final active clauses: 274
% Final passive clauses: 1114
% Input formulas: 52
% Initial clauses: 207
% 
% Duplicate literals: 141
% Trivial inequalities: 193
% Fw subsumption resolutions: 411
% Fw demodulations: 2063
% Bw demodulations: 453
% Condensations: 130
% Global subsumptions: 22
% 
% Simple tautologies: 229
% Equational tautologies: 18
% Forward subsumptions: 3835
% Fw demodulations to eq. taut.: 1
% 
% Binary resolution: 1559
% Factoring: 64
% Forward superposition: 1374
% Backward superposition: 1971
% Self superposition: 34
% Equality factoring: 26
% Equality resolution: 48
% 
% Split clauses: 41
% Split components: 87
% SAT solver clauses: 2419
% SAT solver unit clauses: 212
% SAT solver binary clauses: 117
% SAT solver learnt clauses: 53
% SAT solver learnt literals: 31
% 
% Sat splits: 41
% Sat splitting refutations: 100
% 
% TWLsolver clauses: 2318
% TWLsolver calls for satisfiability: 1759
% 
% Memory used [KB]: 9594
% Time elapsed: 0.425 s
% ------------------------------
% ------------------------------
% Success in time 1.044 s
