Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.integral
% (101722)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 87
% Passive clauses: 1521
% Generated clauses: 12700
% Final active clauses: 85
% Final passive clauses: 1327
% Input formulas: 18
% Initial clauses: 27
% 
% Duplicate literals: 790
% Trivial inequalities: 258
% Fw subsumption resolutions: 2206
% Fw demodulations: 659
% Bw demodulations: 1
% 
% Simple tautologies: 329
% Equational tautologies: 55
% Forward subsumptions: 5804
% 
% Binary resolution: 51
% Forward superposition: 4035
% Backward superposition: 4366
% Self superposition: 229
% Equality factoring: 154
% Equality resolution: 37
% 
% Split clauses: 602
% Split components: 1264
% SAT solver clauses: 1166
% SAT solver unit clauses: 67
% SAT solver binary clauses: 358
% 
% Sat splits: 602
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1144
% TWLsolver calls for satisfiability: 86
% 
% Memory used [KB]: 7036
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.integral
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.integral
% SZS output start Proof for thm.integral
fof(f4326,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f4325,f61,f89,f80,f102,f1450,f1460,f82,f103,f86,f105,f1831,f1834,f1415,f1435,f1376,f1377,f166,f85,f106,f174,f62,f92,f64,f94,f775,f797,f70,f99,f584,f585,f115,f134,f65,f95,f156,f158,f71,f98,f108,f109,f60,f90,f79,f101,f1242,f1257,f84,f107,f117,f135])).
fof(f135,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl38),
  inference(cnf_transformation,[],[f135_D])).
fof(f135_D,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f117,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | ~p(s(bool,X1)) | s(bool,'T_0') = s(bool,X0)) ) | ($spl8 | $spl16)),
  inference(superposition,[],[f99,f95])).
fof(f107,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',sK4(X3,X2,X1,X0))))) | s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',X2),s('type.num.num',X3))),s(fun('type.num.num','type.realax.real'),X0))) = s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',X2),s('type.num.num',X3))),s(fun('type.num.num','type.realax.real'),X1)))) ) | $spl32),
  inference(cnf_transformation,[],[f107_D])).
fof(f107_D,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',sK4(X3,X2,X1,X0))))) | s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',X2),s('type.num.num',X3))),s(fun('type.num.num','type.realax.real'),X0))) = s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',X2),s('type.num.num',X3))),s(fun('type.num.num','type.realax.real'),X1)))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f84,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',X2),s('type.num.num',X3))),s(fun('type.num.num','type.realax.real'),X0))) = s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',X2),s('type.num.num',X3))),s(fun('type.num.num','type.realax.real'),X1))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',sK4(X3,X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0,X1,X2,X3] : ((p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',sK4(X3,X2,X1,X0))))) & p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK4(X3,X2,X1,X0)),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X2)))))) & s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',sK4(X3,X2,X1,X0)))) != s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X1),s('type.num.num',sK4(X3,X2,X1,X0))))) | s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',X2),s('type.num.num',X3))),s(fun('type.num.num','type.realax.real'),X0))) = s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',X2),s('type.num.num',X3))),s(fun('type.num.num','type.realax.real'),X1))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f48])).
fof(f48,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X4)))) & p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X2)))))) & s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X4))) != s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X1),s('type.num.num',X4)))) | s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',X2),s('type.num.num',X3))),s(fun('type.num.num','type.realax.real'),X0))) = s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',X2),s('type.num.num',X3))),s(fun('type.num.num','type.realax.real'),X1))))),
  inference(flattening,[],[f47])).
fof(f47,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : ((p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X4)))) & p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X2))))))) & s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X4))) != s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X1),s('type.num.num',X4)))) | s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',X2),s('type.num.num',X3))),s(fun('type.num.num','type.realax.real'),X0))) = s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',X2),s('type.num.num',X3))),s(fun('type.num.num','type.realax.real'),X1))))),
  inference(ennf_transformation,[],[f41])).
fof(f41,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : ((p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X4)))) & p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X2))))))) => s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X4))) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X1),s('type.num.num',X4)))) => s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',X2),s('type.num.num',X3))),s(fun('type.num.num','type.realax.real'),X0))) = s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',X2),s('type.num.num',X3))),s(fun('type.num.num','type.realax.real'),X1))))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X8,X18,X7,X6] : (! [X9] : ((p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',X9)))) & p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X7))))))) => s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X8),s('type.num.num',X9))) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X18),s('type.num.num',X9)))) => s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',X7),s('type.num.num',X6))),s(fun('type.num.num','type.realax.real'),X8))) = s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',X7),s('type.num.num',X6))),s(fun('type.num.num','type.realax.real'),X18))))),
  file('thm.integral.SUM_EQ_|48|','thm.real.SUM_EQ')).
fof(f1257,plain,(
  ( ! [X10,X8,X7,X11,X9] : (s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',X7),s('type.num.num',X8))),s(fun('type.num.num','type.realax.real'),X9))) = s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',X7),s('type.num.num',X8))),s(fun('type.num.num','type.realax.real'),X10))) | s(bool,X11) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',sK4(X8,X7,X10,X9)))) | s(bool,'T_0') = s(bool,X11)) ) | $spl238),
  inference(cnf_transformation,[],[f1257_D])).
fof(f1257_D,plain,(
  ( ! [X10,X8,X7,X11,X9] : (s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',X7),s('type.num.num',X8))),s(fun('type.num.num','type.realax.real'),X9))) = s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',X7),s('type.num.num',X8))),s(fun('type.num.num','type.realax.real'),X10))) | s(bool,X11) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',sK4(X8,X7,X10,X9)))) | s(bool,'T_0') = s(bool,X11)) ) <=> ~$spl238),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl238])])).
fof(f1242,plain,(
  ( ! [X10,X8,X7,X11,X9] : (s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',X7),s('type.num.num',X8))),s(fun('type.num.num','type.realax.real'),X9))) = s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',X7),s('type.num.num',X8))),s(fun('type.num.num','type.realax.real'),X10))) | s(bool,X11) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',sK4(X8,X7,X10,X9)))) | s(bool,'T_0') = s(bool,X11)) ) | ($spl32 | $spl38)),
  inference(resolution,[],[f107,f135])).
fof(f101,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))) ) | $spl20),
  inference(cnf_transformation,[],[f101_D])).
fof(f101_D,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f79,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f33])).
fof(f33,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X7,X6] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X7))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num',X6)))),
  file('thm.integral.SUM_EQ_|48|','thm.arithmetic.ADD_SYM')).
fof(f90,plain,(
  ( ! [X3] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK2),s('type.num.num',X3))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK1),s('type.num.num',sK0)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK1),s('type.num.num',X3))))) ) | $spl2),
  inference(cnf_transformation,[],[f90_D])).
fof(f90_D,plain,(
  ( ! [X3] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK2),s('type.num.num',X3))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK1),s('type.num.num',sK0)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK1),s('type.num.num',X3))))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f60,plain,(
  ( ! [X3] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK2),s('type.num.num',X3))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK1),s('type.num.num',sK0)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK1),s('type.num.num',X3))))) )),
  inference(cnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X3] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK1),s('type.num.num',X3)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK1),s('type.num.num',sK0)))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK2),s('type.num.num',X3)))) & s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',sK1),s('type.num.num',sK0))),s(fun('type.num.num','type.realax.real'),sK2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f43])).
fof(f43,plain,(
  ? [X0,X1,X2] : (! [X3] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X3)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num',X3)))) & s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',X1),s('type.num.num',X0))),s(fun('type.num.num','type.realax.real'),X2))))),
  inference(flattening,[],[f42])).
fof(f42,plain,(
  ? [X0,X1,X2] : (! [X3] : ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X3)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num',X3)))) & s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',X1),s('type.num.num',X0))),s(fun('type.num.num','type.realax.real'),X2))))),
  inference(ennf_transformation,[],[f20])).
fof(f20,plain,(
  ~! [X0,X1,X2] : (! [X3] : ((p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X3)))) & p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0))))))) => s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num',X3)))) => s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',X1),s('type.num.num',X0))),s(fun('type.num.num','type.realax.real'),X2))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : (! [X9] : ((p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',X9)))) & p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num',X6))))))) => s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X8),s('type.num.num',X9)))) => s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',X7),s('type.num.num',X6))),s(fun('type.num.num','type.realax.real'),X8))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : (! [X9] : ((p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',X9)))) & p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num',X6))))))) => s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X8),s('type.num.num',X9)))) => s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',X7),s('type.num.num',X6))),s(fun('type.num.num','type.realax.real'),X8))))),
  file('thm.integral.SUM_EQ_|48|','thm.integral.SUM_EQ_|48|')).
fof(f109,plain,(
  ( ! [X3] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',sK1)))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK2),s('type.num.num',X3))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK1),s('type.num.num',X3))))) ) | $spl34),
  inference(cnf_transformation,[],[f109_D])).
fof(f109_D,plain,(
  ( ! [X3] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',sK1)))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK2),s('type.num.num',X3))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK1),s('type.num.num',X3))))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f108,plain,(
  ( ! [X3] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',sK1)))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK2),s('type.num.num',X3))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK1),s('type.num.num',X3))))) ) | ($spl2 | $spl20)),
  inference(forward_demodulation,[],[f90,f101])).
fof(f98,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) | $spl14),
  inference(cnf_transformation,[],[f98_D])).
fof(f98_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f71,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f25])).
fof(f25,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f24])).
fof(f24,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X14] : (s(bool,'F_0') = s(bool,X14) <=> ~p(s(bool,X14)))),
  file('thm.integral.SUM_EQ_|48|','thm.bool.EQ_CLAUSES|split|3')).
fof(f158,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK1),s('type.num.num',X0)))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK2),s('type.num.num',X0))) | s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',sK1)))))) ) | $spl50),
  inference(cnf_transformation,[],[f158_D])).
fof(f158_D,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK1),s('type.num.num',X0)))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK2),s('type.num.num',X0))) | s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',sK1)))))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f156,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK2),s('type.num.num',X0))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK1),s('type.num.num',X0)))) | s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',sK1)))))) ) | ($spl14 | $spl34)),
  inference(resolution,[],[f109,f98])).
fof(f95,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl8),
  inference(cnf_transformation,[],[f95_D])).
fof(f95_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f65,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.integral.SUM_EQ_|48|',aHL_BOOL_CASES)).
fof(f134,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl36),
  inference(cnf_transformation,[],[f134_D])).
fof(f134_D,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f115,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl8),
  inference(superposition,[],[f95,f95])).
fof(f585,plain,(
  ( ! [X23,X22] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X22),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',sK1))))) | s(bool,X23) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK1),s('type.num.num',X22))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK2),s('type.num.num',X22))) | s(bool,'T_0') = s(bool,X23)) ) | $spl116),
  inference(cnf_transformation,[],[f585_D])).
fof(f585_D,plain,(
  ( ! [X23,X22] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X22),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',sK1))))) | s(bool,X23) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK1),s('type.num.num',X22))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK2),s('type.num.num',X22))) | s(bool,'T_0') = s(bool,X23)) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f584,plain,(
  ( ! [X23,X22] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK2),s('type.num.num',X22))) | s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X22),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',sK1))))) | s(bool,X23) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK1),s('type.num.num',X22))) | s(bool,'T_0') = s(bool,X23)) ) | ($spl4 | $spl36 | $spl50)),
  inference(subsumption_resolution,[],[f421,f92])).
fof(f421,plain,(
  ( ! [X23,X22] : (~p(s(bool,'T_0')) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK2),s('type.num.num',X22))) | s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X22),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',sK1))))) | s(bool,X23) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK1),s('type.num.num',X22))) | s(bool,'T_0') = s(bool,X23)) ) | ($spl36 | $spl50)),
  inference(superposition,[],[f158,f134])).
fof(f99,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl16),
  inference(cnf_transformation,[],[f99_D])).
fof(f99_D,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f70,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f52])).
fof(f797,plain,(
  ( ! [X15,X16] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X15),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',sK1)))))) | s(bool,X16) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK1),s('type.num.num',X15))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK2),s('type.num.num',X15))) | s(bool,'T_0') = s(bool,X16)) ) | $spl152),
  inference(cnf_transformation,[],[f797_D])).
fof(f797_D,plain,(
  ( ! [X15,X16] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X15),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',sK1)))))) | s(bool,X16) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK1),s('type.num.num',X15))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK2),s('type.num.num',X15))) | s(bool,'T_0') = s(bool,X16)) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f775,plain,(
  ( ! [X15,X16] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X15),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',sK1)))))) | s(bool,X16) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK1),s('type.num.num',X15))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK2),s('type.num.num',X15))) | s(bool,'T_0') = s(bool,X16)) ) | ($spl16 | $spl116)),
  inference(trivial_inequality_removal,[],[f707])).
fof(f707,plain,(
  ( ! [X15,X16] : (s(bool,'F_0') != s(bool,'F_0') | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X15),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',sK1)))))) | s(bool,X16) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK1),s('type.num.num',X15))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK2),s('type.num.num',X15))) | s(bool,'T_0') = s(bool,X16)) ) | ($spl16 | $spl116)),
  inference(superposition,[],[f99,f585])).
fof(f94,plain,(
  ~p(s(bool,'F_0')) | $spl7),
  inference(cnf_transformation,[],[f94_D])).
fof(f94_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f64,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f21])).
fof(f21,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.integral.SUM_EQ_|48|',aHL_FALSITY)).
fof(f92,plain,(
  p(s(bool,'T_0')) | $spl4),
  inference(cnf_transformation,[],[f92_D])).
fof(f92_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f62,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.integral.SUM_EQ_|48|',aHL_TRUTH)).
fof(f174,plain,(
  $false | ($spl4 | $spl7 | $spl52)),
  inference(subsumption_resolution,[],[f170,f92])).
fof(f170,plain,(
  ~p(s(bool,'T_0')) | ($spl7 | $spl52)),
  inference(backward_demodulation,[],[f166,f94])).
fof(f106,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK4(X3,X2,X1,X0)),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X2)))))) | s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',X2),s('type.num.num',X3))),s(fun('type.num.num','type.realax.real'),X0))) = s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',X2),s('type.num.num',X3))),s(fun('type.num.num','type.realax.real'),X1)))) ) | $spl30),
  inference(cnf_transformation,[],[f106_D])).
fof(f106_D,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK4(X3,X2,X1,X0)),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X2)))))) | s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',X2),s('type.num.num',X3))),s(fun('type.num.num','type.realax.real'),X0))) = s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',X2),s('type.num.num',X3))),s(fun('type.num.num','type.realax.real'),X1)))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f85,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',X2),s('type.num.num',X3))),s(fun('type.num.num','type.realax.real'),X0))) = s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',X2),s('type.num.num',X3))),s(fun('type.num.num','type.realax.real'),X1))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK4(X3,X2,X1,X0)),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X2))))))) )),
  inference(cnf_transformation,[],[f59])).
fof(f166,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,X2)) ) | $spl52),
  inference(cnf_transformation,[],[f166_D])).
fof(f166_D,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f1377,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',sK1),s('type.num.num',sK0))),s(fun('type.num.num','type.realax.real'),X0))) = s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',sK1),s('type.num.num',sK0))),s(fun('type.num.num','type.realax.real'),X1))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK2),s('type.num.num',sK4(sK0,sK1,X1,X0))))) ) | $spl252),
  inference(cnf_transformation,[],[f1377_D])).
fof(f1377_D,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',sK1),s('type.num.num',sK0))),s(fun('type.num.num','type.realax.real'),X0))) = s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',sK1),s('type.num.num',sK0))),s(fun('type.num.num','type.realax.real'),X1))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK2),s('type.num.num',sK4(sK0,sK1,X1,X0))))) ) <=> ~$spl252),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl252])])).
fof(f1376,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',sK1),s('type.num.num',sK0))),s(fun('type.num.num','type.realax.real'),X0))) = s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',sK1),s('type.num.num',sK0))),s(fun('type.num.num','type.realax.real'),X1))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK2),s('type.num.num',sK4(sK0,sK1,X1,X0)))) | s(bool,'T_0') = s(bool,X2)) ) | ($spl30 | $spl152 | $spl238)),
  inference(subsumption_resolution,[],[f1353,f1257])).
fof(f1353,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',sK1),s('type.num.num',sK0))),s(fun('type.num.num','type.realax.real'),X0))) = s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',sK1),s('type.num.num',sK0))),s(fun('type.num.num','type.realax.real'),X1))) | s(bool,X2) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK1),s('type.num.num',sK4(sK0,sK1,X1,X0)))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK2),s('type.num.num',sK4(sK0,sK1,X1,X0)))) | s(bool,'T_0') = s(bool,X2)) ) | ($spl30 | $spl152)),
  inference(resolution,[],[f106,f797])).
fof(f1435,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',sK1),s('type.num.num',sK0))),s(fun('type.num.num','type.realax.real'),X0))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK2),s('type.num.num',sK4(sK0,sK1,sK2,X0))))) ) | $spl286),
  inference(cnf_transformation,[],[f1435_D])).
fof(f1435_D,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',sK1),s('type.num.num',sK0))),s(fun('type.num.num','type.realax.real'),X0))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK2),s('type.num.num',sK4(sK0,sK1,sK2,X0))))) ) <=> ~$spl286),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl286])])).
fof(f1415,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',sK1),s('type.num.num',sK0))),s(fun('type.num.num','type.realax.real'),X0))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK2),s('type.num.num',sK4(sK0,sK1,sK2,X0))))) ) | ($spl1 | $spl252)),
  inference(superposition,[],[f89,f1377])).
fof(f1834,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK2),s('type.num.num',sK4(sK0,sK1,sK2,'_dst_x0_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl454),
  inference(cnf_transformation,[],[f1834_D])).
fof(f1834_D,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK2),s('type.num.num',sK4(sK0,sK1,sK2,'_dst_x0_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) <=> ~$spl454),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl454])])).
fof(f1831,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK2),s('type.num.num',sK4(sK0,sK1,sK2,'_dst_x0_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | ($spl22 | $spl286)),
  inference(trivial_inequality_removal,[],[f1828])).
fof(f1828,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK2),s('type.num.num',sK4(sK0,sK1,sK2,'_dst_x0_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | ($spl22 | $spl286)),
  inference(superposition,[],[f1435,f102])).
fof(f105,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',sK4(X3,X2,X1,X0)))) != s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X1),s('type.num.num',sK4(X3,X2,X1,X0)))) | s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',X2),s('type.num.num',X3))),s(fun('type.num.num','type.realax.real'),X0))) = s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',X2),s('type.num.num',X3))),s(fun('type.num.num','type.realax.real'),X1)))) ) | $spl28),
  inference(cnf_transformation,[],[f105_D])).
fof(f105_D,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',sK4(X3,X2,X1,X0)))) != s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X1),s('type.num.num',sK4(X3,X2,X1,X0)))) | s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',X2),s('type.num.num',X3))),s(fun('type.num.num','type.realax.real'),X0))) = s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',X2),s('type.num.num',X3))),s(fun('type.num.num','type.realax.real'),X1)))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f86,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',X2),s('type.num.num',X3))),s(fun('type.num.num','type.realax.real'),X0))) = s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',X2),s('type.num.num',X3))),s(fun('type.num.num','type.realax.real'),X1))) | s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',sK4(X3,X2,X1,X0)))) != s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X1),s('type.num.num',sK4(X3,X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f59])).
fof(f103,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))) ) | $spl24),
  inference(cnf_transformation,[],[f103_D])).
fof(f103_D,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f82,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))),
  inference(flattening,[],[f37])).
fof(f37,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X19,X20] : ! [X21,X22] : s(X20,X21) = s(X20,happ(s(fun(X19,X20),'_dst_x0_1'(s(X20,X21))),s(X19,X22)))),
  file('thm.integral.SUM_EQ_|48|','_dst_x0_ax')).
fof(f1460,plain,(
  ( ! [X4,X2,X3,X1] : (s('type.realax.real',X1) != s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X4),s('type.num.num',sK4(X2,X3,X4,'_dst_x0_1'(s('type.realax.real',X1)))))) | s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',X3),s('type.num.num',X2))),s(fun('type.num.num','type.realax.real'),X4))) = s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',X3),s('type.num.num',X2))),s(fun('type.num.num','type.realax.real'),'_dst_x0_1'(s('type.realax.real',X1)))))) ) | $spl304),
  inference(cnf_transformation,[],[f1460_D])).
fof(f1460_D,plain,(
  ( ! [X4,X2,X3,X1] : (s('type.realax.real',X1) != s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X4),s('type.num.num',sK4(X2,X3,X4,'_dst_x0_1'(s('type.realax.real',X1)))))) | s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',X3),s('type.num.num',X2))),s(fun('type.num.num','type.realax.real'),X4))) = s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',X3),s('type.num.num',X2))),s(fun('type.num.num','type.realax.real'),'_dst_x0_1'(s('type.realax.real',X1)))))) ) <=> ~$spl304),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl304])])).
fof(f1450,plain,(
  ( ! [X4,X2,X3,X1] : (s('type.realax.real',X1) != s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X4),s('type.num.num',sK4(X2,X3,X4,'_dst_x0_1'(s('type.realax.real',X1)))))) | s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',X3),s('type.num.num',X2))),s(fun('type.num.num','type.realax.real'),X4))) = s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',X3),s('type.num.num',X2))),s(fun('type.num.num','type.realax.real'),'_dst_x0_1'(s('type.realax.real',X1)))))) ) | ($spl24 | $spl28)),
  inference(superposition,[],[f105,f103])).
fof(f102,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',X0),s('type.num.num',X1))),s(fun('type.num.num','type.realax.real'),'_dst_x0_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl22),
  inference(cnf_transformation,[],[f102_D])).
fof(f102_D,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',X0),s('type.num.num',X1))),s(fun('type.num.num','type.realax.real'),'_dst_x0_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f80,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',X0),s('type.num.num',X1))),s(fun('type.num.num','type.realax.real'),'_dst_x0_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(cnf_transformation,[],[f34])).
fof(f34,plain,(
  ! [X0,X1] : s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',X0),s('type.num.num',X1))),s(fun('type.num.num','type.realax.real'),'_dst_x0_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X7,X6] : s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',X7),s('type.num.num',X6))),s(fun('type.num.num','type.realax.real'),'_dst_x0_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))),
  file('thm.integral.SUM_EQ_|48|','thm.real.SUM_|48|')).
fof(f89,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',sK1),s('type.num.num',sK0))),s(fun('type.num.num','type.realax.real'),sK2))) | $spl1),
  inference(cnf_transformation,[],[f89_D])).
fof(f89_D,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',sK1),s('type.num.num',sK0))),s(fun('type.num.num','type.realax.real'),sK2))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f61,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',sK1),s('type.num.num',sK0))),s(fun('type.num.num','type.realax.real'),sK2)))),
  inference(cnf_transformation,[],[f49])).
fof(f4325,plain,(
  $false | ($spl1 | $spl22 | $spl304 | $spl454)),
  inference(subsumption_resolution,[],[f4324,f89])).
fof(f4324,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',sK1),s('type.num.num',sK0))),s(fun('type.num.num','type.realax.real'),sK2))) | ($spl22 | $spl304 | $spl454)),
  inference(forward_demodulation,[],[f4310,f102])).
fof(f4310,plain,(
  s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',sK1),s('type.num.num',sK0))),s(fun('type.num.num','type.realax.real'),sK2))) = s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',sK1),s('type.num.num',sK0))),s(fun('type.num.num','type.realax.real'),'_dst_x0_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl304 | $spl454)),
  inference(trivial_inequality_removal,[],[f4306])).
fof(f4306,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',sK1),s('type.num.num',sK0))),s(fun('type.num.num','type.realax.real'),sK2))) = s('type.realax.real','const.real.sum_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.pair.,_2'(s('type.num.num',sK1),s('type.num.num',sK0))),s(fun('type.num.num','type.realax.real'),'_dst_x0_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl304 | $spl454)),
  inference(superposition,[],[f1460,f1834])).
% SZS output end Proof for thm.integral
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 101
% Passive clauses: 531
% Generated clauses: 4195
% Final active clauses: 90
% Final passive clauses: 408
% Input formulas: 18
% Initial clauses: 27
% 
% Duplicate literals: 288
% Trivial inequalities: 123
% Fw subsumption resolutions: 270
% Fw demodulations: 265
% Bw demodulations: 27
% Condensations: 259
% 
% Simple tautologies: 103
% Equational tautologies: 22
% Forward subsumptions: 1552
% 
% Binary resolution: 45
% Forward superposition: 1321
% Backward superposition: 905
% Self superposition: 67
% Equality factoring: 34
% Equality resolution: 39
% 
% Split clauses: 78
% Split components: 161
% SAT solver clauses: 814
% SAT solver unit clauses: 31
% SAT solver binary clauses: 47
% SAT solver learnt clauses: 1
% 
% Sat splits: 78
% Sat splitting refutations: 7
% 
% TWLsolver clauses: 813
% TWLsolver calls for satisfiability: 187
% 
% Memory used [KB]: 5245
% Time elapsed: 0.339 s
% ------------------------------
% ------------------------------
% Success in time 0.961 s
