Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.quantHeuristics
% (48604)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 640
% Passive clauses: 10102
% Generated clauses: 18435
% Final active clauses: 635
% Final passive clauses: 9108
% Input formulas: 59
% Initial clauses: 230
% 
% Duplicate literals: 312
% Trivial inequalities: 42
% Fw subsumption resolutions: 1392
% Fw demodulations: 213
% Bw demodulations: 5
% 
% Simple tautologies: 244
% Equational tautologies: 56
% Forward subsumptions: 5861
% 
% Binary resolution: 11626
% Factoring: 1269
% Forward superposition: 1239
% Backward superposition: 1959
% Self superposition: 29
% Equality factoring: 26
% Equality resolution: 18
% 
% Split clauses: 305
% Split components: 612
% SAT solver clauses: 576
% SAT solver unit clauses: 22
% SAT solver binary clauses: 162
% 
% Sat splits: 305
% 
% TWLsolver clauses: 576
% TWLsolver calls for satisfiability: 346
% 
% Memory used [KB]: 9083
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.quantHeuristics
% (48686)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 413
% Passive clauses: 2555
% Generated clauses: 13366
% Final active clauses: 405
% Final passive clauses: 2126
% Input formulas: 59
% Initial clauses: 230
% 
% Duplicate literals: 373
% Trivial inequalities: 48
% Fw subsumption resolutions: 896
% Fw demodulations: 767
% Bw demodulations: 23
% Condensations: 530
% Global subsumptions: 22
% 
% Simple tautologies: 307
% Equational tautologies: 18
% Forward subsumptions: 4846
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 2984
% Factoring: 130
% Forward superposition: 2468
% Backward superposition: 2248
% Self superposition: 35
% Equality factoring: 30
% Equality resolution: 14
% 
% Split clauses: 23
% Split components: 46
% SAT solver clauses: 3277
% SAT solver unit clauses: 234
% SAT solver binary clauses: 91
% SAT solver learnt clauses: 44
% SAT solver learnt literals: 10
% 
% Sat splits: 23
% 
% TWLsolver clauses: 3189
% TWLsolver calls for satisfiability: 1934
% 
% Memory used [KB]: 14328
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.quantHeuristics
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.quantHeuristics
% SZS output start Proof for thm.quantHeuristics
fof(f90088,plain,(
  $false),
  inference(subsumption_resolution,[],[f90087,f569])).
fof(f569,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f323,f320])).
fof(f320,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f9])).
fof(f9,axiom,(
  p(s(bool,'T_0'))),
  file('thm.quantHeuristics.GUESS_RULES_DISJ|split|7','thm.bool.TRUTH')).
fof(f323,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f200])).
fof(f200,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f63])).
fof(f63,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f30])).
fof(f30,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.quantHeuristics.GUESS_RULES_DISJ|split|7','thm.bool.NOT_CLAUSES|split|1')).
fof(f90087,plain,(
  p(s(bool,'F_0'))),
  inference(forward_demodulation,[],[f90086,f8797])).
fof(f8797,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,happ(s(fun(sK13,bool),sK15),s(sK13,happ(s(fun(sK12,sK13),sK14),s(sK12,X0)))))) )),
  inference(forward_demodulation,[],[f8679,f2449])).
fof(f2449,plain,(
  s(bool,'F_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK12,sK13),sK14),s(fun(sK13,bool),sK16)))),
  inference(unit_resulting_resolution,[],[f569,f1811,f399])).
fof(f399,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f160])).
fof(f160,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f159])).
fof(f159,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X12,X13] : ((p(s(bool,X12)) => p(s(bool,X13))) => ((p(s(bool,X13)) => p(s(bool,X12))) => s(bool,X12) = s(bool,X13)))),
  file('thm.quantHeuristics.GUESS_RULES_DISJ|split|7','thm.bool.IMP_ANTISYM_AX')).
fof(f1811,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK12,sK13),sK14),s(fun(sK13,bool),sK16))))),
  inference(unit_resulting_resolution,[],[f1808,f442])).
fof(f442,plain,(
  ( ! [X2,X0,X5,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) )),
  inference(cnf_transformation,[],[f262])).
fof(f262,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) & (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK18(X3,X2,X1,X0))))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK18])],[f261])).
fof(f261,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) & (? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f260])).
fof(f260,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))))))) & (? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(nnf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))))),
  inference(flattening,[],[f122])).
fof(f122,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X7,X6] : ! [X21,X22] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X6,X7),X21),s(fun(X7,bool),X22)))) <=> ! [X20] : p(s(bool,happ(s(fun(X7,bool),X22),s(X7,happ(s(fun(X6,X7),X21),s(X6,X20)))))))),
  file('thm.quantHeuristics.GUESS_RULES_DISJ|split|7','thm.quantHeuristics.GUESS_REWRITES|split|2')).
fof(f1808,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK13,bool),sK16),s(sK13,happ(s(fun(sK12,sK13),sK14),s(sK12,X0))))))) )),
  inference(subsumption_resolution,[],[f1805,f320])).
fof(f1805,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(sK13,bool),sK16),s(sK13,happ(s(fun(sK12,sK13),sK14),s(sK12,X0))))))) )),
  inference(superposition,[],[f456,f617])).
fof(f617,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK12,sK13),sK14),s(fun(sK13,bool),sK16)))),
  inference(unit_resulting_resolution,[],[f317,f335])).
fof(f335,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f202])).
fof(f202,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X14] : (s(bool,'T_0') = s(bool,X14) <=> p(s(bool,X14)))),
  file('thm.quantHeuristics.GUESS_RULES_DISJ|split|7','thm.bool.EQ_CLAUSES|split|0')).
fof(f317,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK12,sK13),sK14),s(fun(sK13,bool),sK16))))),
  inference(cnf_transformation,[],[f199])).
fof(f199,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK12,sK13),sK14),s(fun(sK13,bool),sK16)))) & p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK12,sK13),sK14),s(fun(sK13,bool),sK15)))) & ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK12,sK13),sK14),s(fun(sK13,bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(sK13,bool),sK16),s(fun(sK13,bool),sK15))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13,sK14,sK15,sK16])],[f147])).
fof(f147,plain,(
  ? [X0,X1,X2,X3,X4] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X4)))) & p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X3)))) & ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(X1,bool),X4),s(fun(X1,bool),X3)))))))),
  inference(flattening,[],[f146])).
fof(f146,plain,(
  ? [X0,X1,X2,X3,X4] : ((p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X4)))) & p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X3))))) & ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(X1,bool),X4),s(fun(X1,bool),X3)))))))),
  inference(ennf_transformation,[],[f62])).
fof(f62,plain,(
  ~! [X0,X1,X2,X3,X4] : ((p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X4)))) & p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X3))))) => p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(X1,bool),X4),s(fun(X1,bool),X3)))))))),
  inference(flattening,[],[f61])).
fof(f61,plain,(
  ~! [X0,X1] : ! [X2,X3,X4] : ((p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X4)))) & p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X3))))) => p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(X1,bool),X4),s(fun(X1,bool),X3)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10] : ((p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X6,X7),X8),s(fun(X7,bool),X10)))) & p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X6,X7),X8),s(fun(X7,bool),X9))))) => p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X6,X7),X8),s(fun(X7,bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(X7,bool),X10),s(fun(X7,bool),X9)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10] : ((p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X6,X7),X8),s(fun(X7,bool),X10)))) & p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X6,X7),X8),s(fun(X7,bool),X9))))) => p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X6,X7),X8),s(fun(X7,bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(X7,bool),X10),s(fun(X7,bool),X9)))))))),
  file('thm.quantHeuristics.GUESS_RULES_DISJ|split|7','thm.quantHeuristics.GUESS_RULES_DISJ|split|7')).
fof(f456,plain,(
  ( ! [X2,X0,X5,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) )),
  inference(cnf_transformation,[],[f277])).
fof(f277,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) & (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK27(X3,X2,X1,X0))))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK27])],[f276])).
fof(f276,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) & (? [X4] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f275])).
fof(f275,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))))))) & (? [X4] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(nnf_transformation,[],[f133])).
fof(f133,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))))),
  inference(flattening,[],[f132])).
fof(f132,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X7,X6] : ! [X21,X22] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X6,X7),X21),s(fun(X7,bool),X22)))) <=> ! [X20] : ~p(s(bool,happ(s(fun(X7,bool),X22),s(X7,happ(s(fun(X6,X7),X21),s(X6,X20)))))))),
  file('thm.quantHeuristics.GUESS_RULES_DISJ|split|7','thm.quantHeuristics.GUESS_REWRITES|split|3')).
fof(f8679,plain,(
  ( ! [X0] : (s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK12,sK13),sK14),s(fun(sK13,bool),sK16))) = s(bool,happ(s(fun(sK13,bool),sK15),s(sK13,happ(s(fun(sK12,sK13),sK14),s(sK12,X0)))))) )),
  inference(unit_resulting_resolution,[],[f2629,f1809,f882])).
fof(f882,plain,(
  ( ! [X4,X5] : (p(s(bool,X5)) | s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f877,f569])).
fof(f877,plain,(
  ( ! [X4,X5] : (p(s(bool,'F_0')) | p(s(bool,X5)) | s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f399,f327])).
fof(f327,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.quantHeuristics.GUESS_RULES_DISJ|split|7',aHL_BOOL_CASES)).
fof(f1809,plain,(
  ( ! [X1] : (~p(s(bool,happ(s(fun(sK13,bool),sK15),s(sK13,happ(s(fun(sK12,sK13),sK14),s(sK12,X1))))))) )),
  inference(subsumption_resolution,[],[f1806,f320])).
fof(f1806,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(sK13,bool),sK15),s(sK13,happ(s(fun(sK12,sK13),sK14),s(sK12,X1))))))) )),
  inference(superposition,[],[f456,f618])).
fof(f618,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK12,sK13),sK14),s(fun(sK13,bool),sK15)))),
  inference(unit_resulting_resolution,[],[f318,f335])).
fof(f318,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK12,sK13),sK14),s(fun(sK13,bool),sK15))))),
  inference(cnf_transformation,[],[f199])).
fof(f2629,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK12,sK13),sK14),s(fun(sK13,bool),sK16)))),
  inference(unit_resulting_resolution,[],[f320,f1811,f558])).
fof(f558,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,X0) != s(bool,X1) | p(s(bool,X1))) )),
  inference(condensation,[],[f557])).
fof(f557,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2)) )),
  inference(condensation,[],[f545])).
fof(f545,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f316])).
fof(f316,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f315])).
fof(f315,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f180])).
fof(f180,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f179])).
fof(f179,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X11,X17,X15,X18] : ((s(bool,X11) = s(bool,X17) & (p(s(bool,X17)) => s(bool,X15) = s(bool,X18))) => ((p(s(bool,X11)) => p(s(bool,X15))) <=> (p(s(bool,X17)) => p(s(bool,X18)))))),
  file('thm.quantHeuristics.GUESS_RULES_DISJ|split|7','thm.bool.IMP_CONG')).
fof(f90086,plain,(
  p(s(bool,happ(s(fun(sK13,bool),sK15),s(sK13,happ(s(fun(sK12,sK13),sK14),s(sK12,sK27('_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(sK13,bool),sK16),s(fun(sK13,bool),sK15)),sK14,sK12,sK13)))))))),
  inference(subsumption_resolution,[],[f90085,f569])).
fof(f90085,plain,(
  p(s(bool,'F_0')) | p(s(bool,happ(s(fun(sK13,bool),sK15),s(sK13,happ(s(fun(sK12,sK13),sK14),s(sK12,sK27('_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(sK13,bool),sK16),s(fun(sK13,bool),sK15)),sK14,sK12,sK13)))))))),
  inference(forward_demodulation,[],[f90084,f1816])).
fof(f1816,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,happ(s(fun(sK13,bool),sK16),s(sK13,happ(s(fun(sK12,sK13),sK14),s(sK12,X0)))))) )),
  inference(unit_resulting_resolution,[],[f569,f1808,f399])).
fof(f90084,plain,(
  p(s(bool,happ(s(fun(sK13,bool),sK16),s(sK13,happ(s(fun(sK12,sK13),sK14),s(sK12,sK27('_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(sK13,bool),sK16),s(fun(sK13,bool),sK15)),sK14,sK12,sK13))))))) | p(s(bool,happ(s(fun(sK13,bool),sK15),s(sK13,happ(s(fun(sK12,sK13),sK14),s(sK12,sK27('_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(sK13,bool),sK16),s(fun(sK13,bool),sK15)),sK14,sK12,sK13)))))))),
  inference(subsumption_resolution,[],[f89760,f320])).
fof(f89760,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK13,bool),sK16),s(sK13,happ(s(fun(sK12,sK13),sK14),s(sK12,sK27('_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(sK13,bool),sK16),s(fun(sK13,bool),sK15)),sK14,sK12,sK13))))))) | p(s(bool,happ(s(fun(sK13,bool),sK15),s(sK13,happ(s(fun(sK12,sK13),sK14),s(sK12,sK27('_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(sK13,bool),sK16),s(fun(sK13,bool),sK15)),sK14,sK12,sK13)))))))),
  inference(superposition,[],[f2794,f14271])).
fof(f14271,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK13,bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(sK13,bool),sK16),s(fun(sK13,bool),sK15))),s(sK13,happ(s(fun(sK12,sK13),sK14),s(sK12,sK27('_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(sK13,bool),sK16),s(fun(sK13,bool),sK15)),sK14,sK12,sK13))))))),
  inference(subsumption_resolution,[],[f14263,f569])).
fof(f14263,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,happ(s(fun(sK13,bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(sK13,bool),sK16),s(fun(sK13,bool),sK15))),s(sK13,happ(s(fun(sK12,sK13),sK14),s(sK12,sK27('_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(sK13,bool),sK16),s(fun(sK13,bool),sK15)),sK14,sK12,sK13))))))),
  inference(superposition,[],[f2433,f976])).
fof(f976,plain,(
  s(bool,'F_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK12,sK13),sK14),s(fun(sK13,bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(sK13,bool),sK16),s(fun(sK13,bool),sK15)))))),
  inference(unit_resulting_resolution,[],[f569,f319,f399])).
fof(f319,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK12,sK13),sK14),s(fun(sK13,bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(sK13,bool),sK16),s(fun(sK13,bool),sK15))))))),
  inference(cnf_transformation,[],[f199])).
fof(f2433,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X2,X0),X3),s(fun(X0,bool),X1)))) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,sK27(X1,X3,X2,X0))))))) )),
  inference(subsumption_resolution,[],[f2432,f569])).
fof(f2432,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'F_0')) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X2,X0),X3),s(fun(X0,bool),X1)))) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,sK27(X1,X3,X2,X0))))))) )),
  inference(superposition,[],[f457,f327])).
fof(f457,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK27(X3,X2,X1,X0))))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f277])).
fof(f2794,plain,(
  ( ! [X111,X112,X110,X113] : (~p(s(bool,happ(s(fun(X110,bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(X110,bool),X111),s(fun(X110,bool),X113))),s(X110,X112)))) | p(s(bool,happ(s(fun(X110,bool),X111),s(X110,X112)))) | p(s(bool,happ(s(fun(X110,bool),X113),s(X110,X112))))) )),
  inference(superposition,[],[f405,f546])).
fof(f546,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun(X3,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X2,X0)),X4),s(fun(X3,X1),X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,happ(s(fun(X3,X1),X5),s(X3,X7))))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))) )),
  inference(cnf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun(X3,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X2,X0)),X4),s(fun(X3,X1),X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,happ(s(fun(X3,X1),X5),s(X3,X7))))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))),
  inference(flattening,[],[f144])).
fof(f144,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6,X7] : s(X0,happ(s(fun(X3,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X2,X0)),X4),s(fun(X3,X1),X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,happ(s(fun(X3,X1),X5),s(X3,X7))))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))),
  inference(rectify,[],[f60])).
fof(f60,axiom,(
  ! [X28,X29,X30,X31] : ! [X32,X33,X34,X35] : s(X28,happ(s(fun(X31,X28),'_dst_x0x1x1_3'(s(fun(X29,fun(X30,X28)),X32),s(fun(X31,X29),X33),s(fun(X31,X30),X34))),s(X31,X35))) = s(X28,happ(s(fun(X30,X28),happ(s(fun(X29,fun(X30,X28)),X32),s(X29,happ(s(fun(X31,X29),X33),s(X31,X35))))),s(X30,happ(s(fun(X31,X30),X34),s(X31,X35)))))),
  file('thm.quantHeuristics.GUESS_RULES_DISJ|split|7','_dst_x0x1x1_ax')).
fof(f405,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1)))) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f244])).
fof(f244,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1)))) | p(s(bool,X0)) | p(s(bool,X1))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1))))))),
  inference(flattening,[],[f243])).
fof(f243,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1)))) | (p(s(bool,X0)) | p(s(bool,X1)))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1))))))),
  inference(nnf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1)))) <=> (p(s(bool,X0)) | p(s(bool,X1))))),
  inference(rectify,[],[f59])).
fof(f59,axiom,(
  ! [X26,X27] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X26))),s(bool,X27)))) <=> (p(s(bool,X26)) | p(s(bool,X27))))),
  file('thm.quantHeuristics.GUESS_RULES_DISJ|split|7','|bslash||slash|_HL_0')).
% SZS output end Proof for thm.quantHeuristics
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 1775
% Passive clauses: 13774
% Generated clauses: 89012
% Final active clauses: 1651
% Final passive clauses: 4436
% Input formulas: 59
% Initial clauses: 230
% Discarded non-redundant clauses: 6812
% 
% Duplicate literals: 1022
% Trivial inequalities: 86
% Fw subsumption resolutions: 2362
% Fw demodulations: 15132
% Bw demodulations: 294
% Condensations: 1492
% Global subsumptions: 57
% 
% Simple tautologies: 162
% Equational tautologies: 1333
% Forward subsumptions: 52248
% Fw demodulations to eq. taut.: 16
% Bw demodulations to eq. taut.: 15
% 
% Unit resulting resolution: 33793
% Factoring: 4
% Forward superposition: 18553
% Backward superposition: 15524
% Self superposition: 170
% Equality factoring: 44
% Equality resolution: 3
% 
% Split clauses: 129
% Split components: 263
% SAT solver clauses: 14922
% SAT solver unit clauses: 5085
% SAT solver binary clauses: 3188
% 
% Sat splits: 129
% Sat splitting refutations: 831
% 
% Memory used [KB]: 103367
% Time elapsed: 8.338 s
% ------------------------------
% ------------------------------
% Success in time 9.664 s
