Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.util_prob
% (80442)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 85
% Passive clauses: 3906
% Generated clauses: 16014
% Final active clauses: 76
% Final passive clauses: 3706
% Input formulas: 11
% Initial clauses: 16
% 
% Duplicate literals: 879
% Trivial inequalities: 163
% Fw subsumption resolutions: 1859
% Fw demodulations: 950
% Bw demodulations: 7
% 
% Simple tautologies: 299
% Equational tautologies: 93
% Forward subsumptions: 5596
% 
% Binary resolution: 85
% Forward superposition: 7208
% Backward superposition: 4091
% Self superposition: 330
% Equality factoring: 241
% Equality resolution: 15
% 
% Split clauses: 876
% Split components: 1766
% SAT solver clauses: 1864
% SAT solver unit clauses: 296
% SAT solver binary clauses: 530
% 
% Sat splits: 876
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1854
% TWLsolver calls for satisfiability: 92
% 
% Memory used [KB]: 9850
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.util_prob
% (80461)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 47
% Passive clauses: 407
% Generated clauses: 6674
% Final active clauses: 43
% Final passive clauses: 351
% Input formulas: 11
% Initial clauses: 16
% 
% Duplicate literals: 983
% Trivial inequalities: 45
% Fw subsumption resolutions: 339
% Fw demodulations: 92
% Bw demodulations: 7
% Condensations: 1468
% 
% Simple tautologies: 126
% Equational tautologies: 63
% Forward subsumptions: 2328
% 
% Binary resolution: 33
% Forward superposition: 1783
% Backward superposition: 1212
% Self superposition: 195
% Equality factoring: 178
% Equality resolution: 6
% 
% Split clauses: 79
% Split components: 159
% SAT solver clauses: 891
% SAT solver unit clauses: 19
% SAT solver binary clauses: 82
% 
% Sat splits: 79
% Sat splitting refutations: 14
% 
% TWLsolver clauses: 891
% TWLsolver calls for satisfiability: 94
% 
% Memory used [KB]: 5628
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.util_prob
% (80478)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 289
% Passive clauses: 5303
% Generated clauses: 94685
% Final active clauses: 98
% Final passive clauses: 620
% Input formulas: 11
% Initial clauses: 16
% Discarded non-redundant clauses: 1370
% 
% Duplicate literals: 6820
% Trivial inequalities: 610
% Fw subsumption resolutions: 7578
% Fw demodulations: 6274
% Bw demodulations: 3247
% Condensations: 14166
% Global subsumptions: 41
% 
% Simple tautologies: 932
% Equational tautologies: 2082
% Forward subsumptions: 49938
% Bw demodulations to eq. taut.: 95
% 
% Unit resulting resolution: 4938
% Forward superposition: 27239
% Backward superposition: 17988
% Self superposition: 616
% Equality factoring: 404
% Equality resolution: 54
% 
% Split clauses: 215
% Split components: 431
% SAT solver clauses: 5863
% SAT solver unit clauses: 119
% SAT solver binary clauses: 552
% 
% Sat splits: 215
% Sat splitting refutations: 87
% 
% Memory used [KB]: 37483
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.util_prob
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.util_prob
% SZS output start Proof for thm.util_prob
fof(f52276,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f107,f110,f108,f5909,f5913,f5911,f5991,f5992,f5993,f5994,f6024,f10787,f10786,f10808,f10785,f10810,f10804,f10802,f10784,f10783,f10818,f10782,f10820,f10969,f10973,f10968,f10806,f10967,f10966,f10976,f10965,f10977,f10964,f10978,f10963,f10979,f10962,f10980,f10961,f10981,f10960,f10983,f10959,f10985,f10958,f10987,f10956,f10953,f10951,f10948,f10945,f10991,f10989,f10942,f10997,f10941,f10995,f10993,f10940,f11001,f10939,f11007,f11005,f10938,f10937,f11015,f11013,f10936,f10935,f11021,f10934,f10814,f10812,f10933,f10932,f11027,f10931,f11029,f10930,f10971,f11306,f11308,f11305,f10816,f11304,f11303,f11309,f11302,f11310,f11301,f11311,f11300,f11312,f11299,f11313,f11298,f11314,f11297,f11316,f11296,f11318,f11295,f11320,f11294,f11322,f11293,f11324,f11292,f11328,f11291,f11332,f11330,f11290,f11025,f11023,f11289,f11334,f11288,f11338,f11336,f11287,f11035,f11033,f11285,f11011,f11009,f11340,f11342,f11282,f11281,f11019,f11017,f10975,f11278,f11276,f11274,f11271,f11346,f11344,f11267,f11350,f11348,f10999,f11031,f18335,f18336,f18425,f18434,f18432,f18426,f19545,f19546,f19554,f19555,f19556,f19557,f20228,f20229,f20237,f20238,f20239,f20240,f20940,f20941,f20954,f20955,f20956,f20957,f21646,f21647,f21660,f21661,f21662,f21663,f19882,f22192,f22193,f22196,f22197,f22198,f22199,f22850,f22851,f22854,f22855,f22856,f22857,f23555,f23556,f23559,f23560,f23561,f23562,f24215,f24216,f24219,f24220,f24221,f24222,f22611,f25202,f25203,f25212,f25213,f25214,f25215,f25882,f25883,f25892,f25893,f25894,f25895,f26600,f26601,f26614,f26615,f26616,f26617,f27282,f27283,f27296,f27297,f27298,f27299,f25538,f27832,f27833,f27836,f27837,f27838,f27839,f28486,f28487,f28490,f28491,f28492,f28493,f29172,f29173,f29176,f29177,f29178,f29179,f29837,f29838,f29841,f29842,f29843,f29844,f28250,f30585,f30586,f30595,f30596,f30597,f30598,f31330,f31331,f31340,f31341,f31342,f31343,f32120,f32121,f32134,f32135,f32136,f32137,f32867,f32868,f32881,f32882,f32883,f32884,f30974,f33488,f33489,f33492,f33493,f33494,f33495,f34207,f34208,f34211,f34212,f34213,f34214,f34974,f34975,f34978,f34979,f34980,f34981,f35695,f35696,f35699,f35700,f35701,f35702,f33956,f36502,f19883,f36503,f19884,f36512,f19885,f36513,f19886,f36514,f19887,f36515,f19888,f19880,f37266,f20566,f37267,f20567,f37276,f20568,f37277,f20569,f37278,f20570,f37279,f20571,f20565,f38045,f21251,f38046,f21252,f38059,f21253,f38060,f21254,f38061,f21255,f38062,f21256,f21250,f38796,f21959,f38797,f21960,f38810,f21961,f38811,f21962,f38812,f21963,f38813,f21964,f11003,f36891,f21958,f39415,f22612,f39416,f22613,f39419,f22614,f39420,f22615,f39421,f22616,f39422,f22617,f22609,f40138,f23270,f40139,f23271,f40142,f23272,f40143,f23273,f40144,f23274,f40145,f23275,f23269,f40904,f23944,f40905,f23945,f40908,f23946,f40909,f23947,f40910,f23948,f40911,f23949,f23943,f41629,f24606,f41630,f24607,f41633,f24608,f41634,f24609,f41635,f24610,f41636,f24611,f11326,f39884,f24605,f42165,f42436,f42668,f42673,f42679,f42682,f42686,f42689,f42695,f42698,f42670,f42704,f42816,f42820,f42821,f42823,f42824,f42826,f42827,f42829,f42818,f42833,f42837,f42835,f43129,f43133,f43131,f43208,f6020,f43209,f6026,f43210,f6028,f43211,f43246,f6032,f50313,f50348,f50732,f50886,f18348,f18346,f51367,f51525,f18438,f18436,f52175,f42672,f52178,f42677,f52180,f42681,f6030,f52181,f42684,f42675,f52183,f42688,f52184,f42693,f52186,f42697,f6022,f52187,f42700,f52177,f42691,f52191,f52206,f42706,f42708,f52273,f50885,f18344,f18342,f52275])).
fof(f52275,plain,(
  $false | ($spl164 | $spl329)),
  inference(subsumption_resolution,[],[f42254,f42706])).
fof(f42254,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK4)))) | $spl164),
  inference(subsumption_resolution,[],[f42231,f55])).
fof(f55,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f42])).
fof(f42,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f33])).
fof(f33,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f18])).
fof(f18,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f17])).
fof(f17,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6] : ! [X11] : (s(X6,X11) = s(X6,X11) <=> p(s(bool,'T_0')))),
  file('thm.util_prob.PREIMAGE_INTER','thm.bool.REFL_CLAUSE')).
fof(f42231,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK4)))) | $spl164),
  inference(superposition,[],[f47,f18343])).
fof(f18343,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl164),
  inference(cnf_transformation,[],[f18343_D])).
fof(f18343_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f47,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))) )),
  inference(cnf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))))),
  inference(flattening,[],[f34])).
fof(f34,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))))),
  inference(nnf_transformation,[],[f22])).
fof(f22,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(flattening,[],[f21])).
fof(f21,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X6] : ! [X9,X10,X11] : (p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),X9),s(fun(X6,bool),X10)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X9)))) & p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X10))))))),
  file('thm.util_prob.PREIMAGE_INTER','thm.pred_set.IN_INTER')).
fof(f18342,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl162),
  inference(cnf_transformation,[],[f18342_D])).
fof(f18342_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f18344,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl165),
  inference(cnf_transformation,[],[f18344_D])).
fof(f18344_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) <=> ~$spl165),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl165])])).
fof(f50885,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))))))),
  inference(superposition,[],[f915,f44])).
fof(f44,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.util_prob.PREIMAGE_INTER',aHL_BOOL_CASES)).
fof(f915,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))))))),
  inference(forward_demodulation,[],[f914,f45])).
fof(f45,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X4))),s(fun(X1,bool),X3))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),'const.util_prob.PREIMAGE_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X3)))))) )),
  inference(cnf_transformation,[],[f20])).
fof(f20,plain,(
  ! [X0,X1,X2,X3,X4] : s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X4))),s(fun(X1,bool),X3))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),'const.util_prob.PREIMAGE_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X3)))))),
  inference(flattening,[],[f19])).
fof(f19,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X4))),s(fun(X1,bool),X3))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),'const.util_prob.PREIMAGE_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X3)))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6,X7] : ! [X8,X9,X11] : s(bool,'const.bool.IN_2'(s(X7,happ(s(fun(X6,X7),X8),s(X6,X11))),s(fun(X7,bool),X9))) = s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),'const.util_prob.PREIMAGE_2'(s(fun(X6,X7),X8),s(fun(X7,bool),X9)))))),
  file('thm.util_prob.PREIMAGE_INTER','thm.util_prob.IN_PREIMAGE')).
fof(f914,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))))))),
  inference(equality_resolution,[],[f74])).
fof(f74,plain,(
  ( ! [X0] : (s(fun(sK0,bool),X0) != s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),X0,sK0)),s(fun(sK0,bool),X0))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),X0,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))))))) )),
  inference(superposition,[],[f39,f41])).
fof(f41,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'const.bool.IN_2'(s(X0,sK5(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK5(X2,X1,X0)),s(fun(X0,bool),X2)))) )),
  inference(cnf_transformation,[],[f32])).
fof(f32,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (s(bool,'const.bool.IN_2'(s(X0,sK5(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK5(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5])],[f31])).
fof(f31,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(rectify,[],[f30])).
fof(f30,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(nnf_transformation,[],[f16])).
fof(f16,plain,(
  ! [X0,X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(flattening,[],[f15])).
fof(f15,plain,(
  ! [X0] : ! [X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6] : ! [X9,X10] : (s(fun(X6,bool),X9) = s(fun(X6,bool),X10) <=> ! [X11] : s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X9))) = s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X10))))),
  file('thm.util_prob.PREIMAGE_INTER','thm.pred_set.EXTENSION')).
fof(f39,plain,(
  s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))))),
  inference(cnf_transformation,[],[f29])).
fof(f29,plain,(
  s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4])],[f27])).
fof(f27,plain,(
  ? [X0,X1,X2,X3,X4] : s(fun(X0,bool),'const.util_prob.PREIMAGE_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'const.pred_set.INTER_2'(s(fun(X1,bool),X3),s(fun(X1,bool),X4))))) != s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),'const.util_prob.PREIMAGE_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X3))),s(fun(X0,bool),'const.util_prob.PREIMAGE_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X4)))))),
  inference(ennf_transformation,[],[f14])).
fof(f14,plain,(
  ~! [X0,X1,X2,X3,X4] : s(fun(X0,bool),'const.util_prob.PREIMAGE_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'const.pred_set.INTER_2'(s(fun(X1,bool),X3),s(fun(X1,bool),X4))))) = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),'const.util_prob.PREIMAGE_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X3))),s(fun(X0,bool),'const.util_prob.PREIMAGE_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X4)))))),
  inference(flattening,[],[f13])).
fof(f13,plain,(
  ~! [X0,X1] : ! [X2,X3,X4] : s(fun(X0,bool),'const.util_prob.PREIMAGE_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'const.pred_set.INTER_2'(s(fun(X1,bool),X3),s(fun(X1,bool),X4))))) = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),'const.util_prob.PREIMAGE_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X3))),s(fun(X0,bool),'const.util_prob.PREIMAGE_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X4)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10] : s(fun(X6,bool),'const.util_prob.PREIMAGE_2'(s(fun(X6,X7),X8),s(fun(X7,bool),'const.pred_set.INTER_2'(s(fun(X7,bool),X9),s(fun(X7,bool),X10))))) = s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),'const.util_prob.PREIMAGE_2'(s(fun(X6,X7),X8),s(fun(X7,bool),X9))),s(fun(X6,bool),'const.util_prob.PREIMAGE_2'(s(fun(X6,X7),X8),s(fun(X7,bool),X10)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10] : s(fun(X6,bool),'const.util_prob.PREIMAGE_2'(s(fun(X6,X7),X8),s(fun(X7,bool),'const.pred_set.INTER_2'(s(fun(X7,bool),X9),s(fun(X7,bool),X10))))) = s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),'const.util_prob.PREIMAGE_2'(s(fun(X6,X7),X8),s(fun(X7,bool),X9))),s(fun(X6,bool),'const.util_prob.PREIMAGE_2'(s(fun(X6,X7),X8),s(fun(X7,bool),X10)))))),
  file('thm.util_prob.PREIMAGE_INTER','thm.util_prob.PREIMAGE_INTER')).
fof(f52273,plain,(
  $false | ($spl166 | $spl329)),
  inference(subsumption_resolution,[],[f52193,f42706])).
fof(f52193,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK4)))) | $spl166),
  inference(forward_demodulation,[],[f52192,f45])).
fof(f52192,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))))) | $spl166),
  inference(subsumption_resolution,[],[f52150,f55])).
fof(f52150,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))))) | $spl166),
  inference(superposition,[],[f47,f18346])).
fof(f42708,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK3)))) | $spl331),
  inference(cnf_transformation,[],[f42708_D])).
fof(f42708_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK3)))) <=> ~$spl331),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl331])])).
fof(f42706,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK4)))) | $spl329),
  inference(cnf_transformation,[],[f42706_D])).
fof(f42706_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK4)))) <=> ~$spl329),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl329])])).
fof(f52206,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK4)))) | ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK3)))) | $spl168),
  inference(subsumption_resolution,[],[f51017,f54])).
fof(f54,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f26])).
fof(f26,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.util_prob.PREIMAGE_INTER',aHL_FALSITY)).
fof(f51017,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK4)))) | ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK3)))) | $spl168),
  inference(superposition,[],[f48,f18347])).
fof(f18347,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl168),
  inference(cnf_transformation,[],[f18347_D])).
fof(f18347_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f48,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f35])).
fof(f52191,plain,(
  $false | ($spl166 | $spl331)),
  inference(subsumption_resolution,[],[f52190,f42708])).
fof(f52190,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK3)))) | $spl166),
  inference(forward_demodulation,[],[f52189,f45])).
fof(f52189,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)))))) | $spl166),
  inference(subsumption_resolution,[],[f52149,f55])).
fof(f52149,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)))))) | $spl166),
  inference(superposition,[],[f46,f18346])).
fof(f46,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))) )),
  inference(cnf_transformation,[],[f35])).
fof(f42691,plain,(
  s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | $spl321),
  inference(cnf_transformation,[],[f42691_D])).
fof(f42691_D,plain,(
  s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) <=> ~$spl321),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl321])])).
fof(f52177,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | $spl345),
  inference(cnf_transformation,[],[f52177_D])).
fof(f52177_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) <=> ~$spl345),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl345])])).
fof(f42700,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),sK0)))),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl327),
  inference(cnf_transformation,[],[f42700_D])).
fof(f42700_D,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),sK0)))),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) <=> ~$spl327),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl327])])).
fof(f52187,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),sK0)))),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | $spl166),
  inference(forward_demodulation,[],[f52143,f45])).
fof(f52143,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl166),
  inference(superposition,[],[f669,f18346])).
fof(f669,plain,(
  ( ! [X2,X3] : (s(bool,'const.bool.IN_2'(s(sK0,sK5(X3,X2,sK0)),s(fun(sK0,bool),X2))) != s(bool,'const.bool.IN_2'(s(sK0,sK5(X3,X2,sK0)),s(fun(sK0,bool),X3))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),X2))) | s(bool,'const.bool.IN_2'(s(sK0,sK5(X2,'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),sK0)),s(fun(sK0,bool),X3))) != s(bool,'const.bool.IN_2'(s(sK0,sK5(X2,'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))))) )),
  inference(superposition,[],[f71,f41])).
fof(f71,plain,(
  ( ! [X1] : (s(bool,'const.bool.IN_2'(s(sK0,sK5(X1,'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),sK0)),s(fun(sK0,bool),X1))) != s(bool,'const.bool.IN_2'(s(sK0,sK5(X1,'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),X1)))) )),
  inference(superposition,[],[f39,f41])).
fof(f6022,plain,(
  s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl11),
  inference(cnf_transformation,[],[f6022_D])).
fof(f6022_D,plain,(
  s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f42697,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),sK0)))),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl325),
  inference(cnf_transformation,[],[f42697_D])).
fof(f42697_D,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),sK0)))),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) <=> ~$spl325),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl325])])).
fof(f52186,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),sK0)))),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl166),
  inference(forward_demodulation,[],[f52185,f45])).
fof(f52185,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl166),
  inference(forward_demodulation,[],[f52142,f45])).
fof(f52142,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | $spl166),
  inference(superposition,[],[f668,f18346])).
fof(f668,plain,(
  ( ! [X0,X1] : (s(bool,'const.bool.IN_2'(s(sK0,sK5(X0,X1,sK0)),s(fun(sK0,bool),X0))) != s(bool,'const.bool.IN_2'(s(sK0,sK5(X0,X1,sK0)),s(fun(sK0,bool),X1))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),X0))) | s(bool,'const.bool.IN_2'(s(sK0,sK5(X0,'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),sK0)),s(fun(sK0,bool),X1))) != s(bool,'const.bool.IN_2'(s(sK0,sK5(X0,'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))))) )),
  inference(superposition,[],[f71,f41])).
fof(f42693,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl323),
  inference(cnf_transformation,[],[f42693_D])).
fof(f42693_D,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) <=> ~$spl323),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl323])])).
fof(f52184,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | $spl166),
  inference(forward_demodulation,[],[f52141,f45])).
fof(f52141,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl166),
  inference(superposition,[],[f591,f18346])).
fof(f591,plain,(
  ( ! [X2,X3] : (s(bool,'const.bool.IN_2'(s(sK0,sK5(X3,X2,sK0)),s(fun(sK0,bool),X2))) != s(bool,'const.bool.IN_2'(s(sK0,sK5(X3,X2,sK0)),s(fun(sK0,bool),X3))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),X2))) | s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),X2,sK0)),s(fun(sK0,bool),X3))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),X2,sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))))) )),
  inference(superposition,[],[f70,f41])).
fof(f70,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),X0,sK0)),s(fun(sK0,bool),X0))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),X0,sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),X0)))) )),
  inference(superposition,[],[f39,f41])).
fof(f42688,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl319),
  inference(cnf_transformation,[],[f42688_D])).
fof(f42688_D,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) <=> ~$spl319),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl319])])).
fof(f52183,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl166),
  inference(forward_demodulation,[],[f52182,f45])).
fof(f52182,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl166),
  inference(forward_demodulation,[],[f52140,f45])).
fof(f52140,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | $spl166),
  inference(superposition,[],[f590,f18346])).
fof(f590,plain,(
  ( ! [X0,X1] : (s(bool,'const.bool.IN_2'(s(sK0,sK5(X0,X1,sK0)),s(fun(sK0,bool),X0))) != s(bool,'const.bool.IN_2'(s(sK0,sK5(X0,X1,sK0)),s(fun(sK0,bool),X1))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),X0))) | s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),X0,sK0)),s(fun(sK0,bool),X1))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),X0,sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))))) )),
  inference(superposition,[],[f70,f41])).
fof(f42675,plain,(
  s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | $spl311),
  inference(cnf_transformation,[],[f42675_D])).
fof(f42675_D,plain,(
  s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) <=> ~$spl311),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl311])])).
fof(f42684,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),sK0)))),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl317),
  inference(cnf_transformation,[],[f42684_D])).
fof(f42684_D,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),sK0)))),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) <=> ~$spl317),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl317])])).
fof(f52181,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),sK0)))),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | $spl166),
  inference(forward_demodulation,[],[f52139,f45])).
fof(f52139,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl166),
  inference(superposition,[],[f519,f18346])).
fof(f519,plain,(
  ( ! [X2,X3] : (s(bool,'const.bool.IN_2'(s(sK0,sK5(X3,X2,sK0)),s(fun(sK0,bool),X2))) != s(bool,'const.bool.IN_2'(s(sK0,sK5(X3,X2,sK0)),s(fun(sK0,bool),X3))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X2),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK0,sK5(X2,'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),sK0)),s(fun(sK0,bool),X3))) != s(bool,'const.bool.IN_2'(s(sK0,sK5(X2,'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)))))) )),
  inference(superposition,[],[f61,f41])).
fof(f61,plain,(
  ( ! [X1] : (s(bool,'const.bool.IN_2'(s(sK0,sK5(X1,'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),sK0)),s(fun(sK0,bool),X1))) != s(bool,'const.bool.IN_2'(s(sK0,sK5(X1,'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X1),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))))) )),
  inference(superposition,[],[f39,f41])).
fof(f6030,plain,(
  s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | $spl19),
  inference(cnf_transformation,[],[f6030_D])).
fof(f6030_D,plain,(
  s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f42681,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),sK0)))),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl315),
  inference(cnf_transformation,[],[f42681_D])).
fof(f42681_D,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),sK0)))),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) <=> ~$spl315),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl315])])).
fof(f52180,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),sK0)))),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | $spl166),
  inference(forward_demodulation,[],[f52179,f45])).
fof(f52179,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | $spl166),
  inference(forward_demodulation,[],[f52138,f45])).
fof(f52138,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | $spl166),
  inference(superposition,[],[f518,f18346])).
fof(f518,plain,(
  ( ! [X0,X1] : (s(bool,'const.bool.IN_2'(s(sK0,sK5(X0,X1,sK0)),s(fun(sK0,bool),X0))) != s(bool,'const.bool.IN_2'(s(sK0,sK5(X0,X1,sK0)),s(fun(sK0,bool),X1))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X0),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK0,sK5(X0,'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),sK0)),s(fun(sK0,bool),X1))) != s(bool,'const.bool.IN_2'(s(sK0,sK5(X0,'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)))))) )),
  inference(superposition,[],[f61,f41])).
fof(f42677,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl313),
  inference(cnf_transformation,[],[f42677_D])).
fof(f42677_D,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) <=> ~$spl313),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl313])])).
fof(f52178,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | $spl166),
  inference(forward_demodulation,[],[f52137,f45])).
fof(f52137,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl166),
  inference(superposition,[],[f446,f18346])).
fof(f446,plain,(
  ( ! [X2,X3] : (s(bool,'const.bool.IN_2'(s(sK0,sK5(X3,X2,sK0)),s(fun(sK0,bool),X2))) != s(bool,'const.bool.IN_2'(s(sK0,sK5(X3,X2,sK0)),s(fun(sK0,bool),X3))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X2),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),X2,sK0)),s(fun(sK0,bool),X3))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),X2,sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)))))) )),
  inference(superposition,[],[f60,f41])).
fof(f60,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),X0,sK0)),s(fun(sK0,bool),X0))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),X0,sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X0),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))))) )),
  inference(superposition,[],[f39,f41])).
fof(f42672,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl309),
  inference(cnf_transformation,[],[f42672_D])).
fof(f42672_D,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) <=> ~$spl309),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl309])])).
fof(f52175,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | $spl166),
  inference(forward_demodulation,[],[f52174,f45])).
fof(f52174,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | $spl166),
  inference(forward_demodulation,[],[f52136,f45])).
fof(f52136,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | $spl166),
  inference(superposition,[],[f445,f18346])).
fof(f445,plain,(
  ( ! [X0,X1] : (s(bool,'const.bool.IN_2'(s(sK0,sK5(X0,X1,sK0)),s(fun(sK0,bool),X0))) != s(bool,'const.bool.IN_2'(s(sK0,sK5(X0,X1,sK0)),s(fun(sK0,bool),X1))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X0),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),X0,sK0)),s(fun(sK0,bool),X1))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),X0,sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)))))) )),
  inference(superposition,[],[f60,f41])).
fof(f18436,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl174),
  inference(cnf_transformation,[],[f18436_D])).
fof(f18436_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f18438,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl177),
  inference(cnf_transformation,[],[f18438_D])).
fof(f18438_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) <=> ~$spl177),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl177])])).
fof(f51525,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))))))),
  inference(superposition,[],[f984,f44])).
fof(f984,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))))))),
  inference(forward_demodulation,[],[f983,f45])).
fof(f983,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))))))),
  inference(equality_resolution,[],[f75])).
fof(f75,plain,(
  ( ! [X1] : (s(fun(sK0,bool),X1) != s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK0,sK5(X1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),X1))) != s(bool,'const.bool.IN_2'(s(sK0,sK5(X1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))))))) )),
  inference(superposition,[],[f39,f41])).
fof(f51367,plain,(
  $false | ($spl170 | $spl173)),
  inference(subsumption_resolution,[],[f51362,f50349])).
fof(f50349,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl170),
  inference(forward_demodulation,[],[f984,f18432])).
fof(f51362,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl173),
  inference(trivial_inequality_removal,[],[f51360])).
fof(f51360,plain,(
  s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl173),
  inference(superposition,[],[f18434,f44])).
fof(f18346,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl166),
  inference(cnf_transformation,[],[f18346_D])).
fof(f18346_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f18348,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl169),
  inference(cnf_transformation,[],[f18348_D])).
fof(f18348_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) <=> ~$spl169),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl169])])).
fof(f50886,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))))))),
  inference(superposition,[],[f915,f44])).
fof(f50732,plain,(
  $false | ($spl162 | $spl165)),
  inference(subsumption_resolution,[],[f50727,f50314])).
fof(f50314,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl162),
  inference(forward_demodulation,[],[f915,f18342])).
fof(f50727,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl165),
  inference(trivial_inequality_removal,[],[f50725])).
fof(f50725,plain,(
  s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl165),
  inference(superposition,[],[f18344,f44])).
fof(f50348,plain,(
  $false | ($spl172 | $spl337)),
  inference(subsumption_resolution,[],[f42837,f42505])).
fof(f42505,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK3)))) | $spl172),
  inference(subsumption_resolution,[],[f42482,f55])).
fof(f42482,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK3)))) | $spl172),
  inference(superposition,[],[f46,f18433])).
fof(f18433,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl172),
  inference(cnf_transformation,[],[f18433_D])).
fof(f18433_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f50313,plain,(
  $false | ($spl164 | $spl331)),
  inference(subsumption_resolution,[],[f42708,f42253])).
fof(f42253,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK3)))) | $spl164),
  inference(subsumption_resolution,[],[f42230,f55])).
fof(f42230,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK3)))) | $spl164),
  inference(superposition,[],[f46,f18343])).
fof(f6032,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK1)),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl21),
  inference(cnf_transformation,[],[f6032_D])).
fof(f6032_D,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK1)),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f43246,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl343),
  inference(cnf_transformation,[],[f43246_D])).
fof(f43246_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) <=> ~$spl343),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl343])])).
fof(f43211,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK1)),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl338),
  inference(superposition,[],[f367,f43170])).
fof(f43170,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl338),
  inference(trivial_inequality_removal,[],[f43157])).
fof(f43157,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl338),
  inference(superposition,[],[f864,f43131])).
fof(f864,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,f(s(sK0,sK6(X1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) | s(fun(sK0,bool),X1) != s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK6(X1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0))))) )),
  inference(superposition,[],[f77,f44])).
fof(f77,plain,(
  ( ! [X3] : (s(bool,f(s(sK0,sK6(X3,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) != s(bool,g(s(sK0,sK6(X3,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) | s(fun(sK0,bool),X3) != s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))))) )),
  inference(superposition,[],[f39,f53])).
fof(f53,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun(X0,X1),X2) = s(fun(X0,X1),X3) | s(X1,f(s(X0,sK6(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK6(X3,X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK6(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK6(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK6])],[f28])).
fof(f28,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f25])).
fof(f25,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f24])).
fof(f24,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.util_prob.PREIMAGE_INTER',aHL_EXT)).
fof(f367,plain,(
  ( ! [X0,X1] : (s(bool,f(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X0)),X1,bool,sK0)))) != s(bool,g(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X0)),X1,bool,sK0)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X1),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X0,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X0,sK1)),s(fun(sK1,bool),X0)))) )),
  inference(superposition,[],[f62,f41])).
fof(f62,plain,(
  ( ! [X2] : (s(bool,f(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),X2,bool,sK0)))) != s(bool,g(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),X2,bool,sK0)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X2),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))))) )),
  inference(superposition,[],[f39,f53])).
fof(f6028,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,sK5('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,sK1)),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl17),
  inference(cnf_transformation,[],[f6028_D])).
fof(f6028_D,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,sK5('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,sK1)),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f43210,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,sK1)),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl338),
  inference(superposition,[],[f368,f43170])).
fof(f368,plain,(
  ( ! [X2,X3] : (s(bool,f(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X2)),X3,bool,sK0)))) != s(bool,g(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X2)),X3,bool,sK0)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X3),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X2,sK3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X2,sK3,sK1)),s(fun(sK1,bool),X2)))) )),
  inference(superposition,[],[f62,f41])).
fof(f6026,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK1)),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl15),
  inference(cnf_transformation,[],[f6026_D])).
fof(f6026_D,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK1)),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f43209,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK1)),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl338),
  inference(superposition,[],[f383,f43170])).
fof(f383,plain,(
  ( ! [X0,X1] : (s(bool,f(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X0)),X1,bool,sK0)))) != s(bool,g(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X0)),X1,bool,sK0)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),X1))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X0,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X0,sK1)),s(fun(sK1,bool),X0)))) )),
  inference(superposition,[],[f72,f41])).
fof(f72,plain,(
  ( ! [X2] : (s(bool,f(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),X2,bool,sK0)))) != s(bool,g(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),X2,bool,sK0)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),X2)))) )),
  inference(superposition,[],[f39,f53])).
fof(f6020,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,sK5('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,sK1)),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl9),
  inference(cnf_transformation,[],[f6020_D])).
fof(f6020_D,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,sK5('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,sK1)),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f43208,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,sK1)),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl338),
  inference(superposition,[],[f384,f43170])).
fof(f384,plain,(
  ( ! [X2,X3] : (s(bool,f(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X2)),X3,bool,sK0)))) != s(bool,g(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X2)),X3,bool,sK0)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),X3))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X2,sK4,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X2,sK4,sK1)),s(fun(sK1,bool),X2)))) )),
  inference(superposition,[],[f72,f41])).
fof(f43131,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl338),
  inference(cnf_transformation,[],[f43131_D])).
fof(f43131_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) <=> ~$spl338),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl338])])).
fof(f43133,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl340),
  inference(cnf_transformation,[],[f43133_D])).
fof(f43133_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) <=> ~$spl340),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl340])])).
fof(f43129,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0))))),
  inference(equality_resolution,[],[f3425])).
fof(f3425,plain,(
  ( ! [X0] : (s(fun(sK0,bool),X0) != s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') = s(bool,g(s(sK0,sK6(X0,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK6(X0,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0))))) )),
  inference(trivial_inequality_removal,[],[f3423])).
fof(f3423,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(fun(sK0,bool),X0) != s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') = s(bool,g(s(sK0,sK6(X0,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK6(X0,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0))))) )),
  inference(superposition,[],[f863,f44])).
fof(f863,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(sK0,sK6(X0,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) | s(fun(sK0,bool),X0) != s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') = s(bool,g(s(sK0,sK6(X0,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0))))) )),
  inference(superposition,[],[f77,f44])).
fof(f42835,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK4)))) | $spl335),
  inference(cnf_transformation,[],[f42835_D])).
fof(f42835_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK4)))) <=> ~$spl335),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl335])])).
fof(f42837,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK3)))) | $spl337),
  inference(cnf_transformation,[],[f42837_D])).
fof(f42837_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK3)))) <=> ~$spl337),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl337])])).
fof(f42833,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK3)))) | ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK4)))) | $spl170),
  inference(forward_demodulation,[],[f42832,f45])).
fof(f42832,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK4)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)))))) | $spl170),
  inference(forward_demodulation,[],[f42831,f45])).
fof(f42831,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)))))) | $spl170),
  inference(subsumption_resolution,[],[f42793,f54])).
fof(f42793,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)))))) | $spl170),
  inference(superposition,[],[f48,f18432])).
fof(f42818,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | $spl333),
  inference(cnf_transformation,[],[f42818_D])).
fof(f42818_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) <=> ~$spl333),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl333])])).
fof(f42829,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),sK0)))),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl170),
  inference(forward_demodulation,[],[f42828,f45])).
fof(f42828,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl170),
  inference(forward_demodulation,[],[f42785,f45])).
fof(f42785,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | $spl170),
  inference(superposition,[],[f669,f18432])).
fof(f42827,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),sK0)))),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | $spl170),
  inference(forward_demodulation,[],[f42784,f45])).
fof(f42784,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl170),
  inference(superposition,[],[f668,f18432])).
fof(f42826,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl170),
  inference(forward_demodulation,[],[f42825,f45])).
fof(f42825,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl170),
  inference(forward_demodulation,[],[f42783,f45])).
fof(f42783,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | $spl170),
  inference(superposition,[],[f591,f18432])).
fof(f42824,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | $spl170),
  inference(forward_demodulation,[],[f42782,f45])).
fof(f42782,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl170),
  inference(superposition,[],[f590,f18432])).
fof(f42823,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),sK0)))),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | $spl170),
  inference(forward_demodulation,[],[f42822,f45])).
fof(f42822,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | $spl170),
  inference(forward_demodulation,[],[f42781,f45])).
fof(f42781,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | $spl170),
  inference(superposition,[],[f519,f18432])).
fof(f42821,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),sK0)))),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | $spl170),
  inference(forward_demodulation,[],[f42780,f45])).
fof(f42780,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl170),
  inference(superposition,[],[f518,f18432])).
fof(f42820,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | $spl170),
  inference(forward_demodulation,[],[f42819,f45])).
fof(f42819,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | $spl170),
  inference(forward_demodulation,[],[f42779,f45])).
fof(f42779,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | $spl170),
  inference(superposition,[],[f446,f18432])).
fof(f42816,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | $spl170),
  inference(forward_demodulation,[],[f42778,f45])).
fof(f42778,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl170),
  inference(superposition,[],[f445,f18432])).
fof(f42704,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK3)))) | ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK4)))) | $spl162),
  inference(forward_demodulation,[],[f42703,f45])).
fof(f42703,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK4)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)))))) | $spl162),
  inference(forward_demodulation,[],[f42702,f45])).
fof(f42702,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)))))) | $spl162),
  inference(subsumption_resolution,[],[f42644,f54])).
fof(f42644,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)))))) | $spl162),
  inference(superposition,[],[f48,f18342])).
fof(f42670,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | $spl307),
  inference(cnf_transformation,[],[f42670_D])).
fof(f42670_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) <=> ~$spl307),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl307])])).
fof(f42698,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),sK0)))),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | $spl162),
  inference(forward_demodulation,[],[f42636,f45])).
fof(f42636,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl162),
  inference(superposition,[],[f669,f18342])).
fof(f42695,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),sK0)))),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl162),
  inference(forward_demodulation,[],[f42694,f45])).
fof(f42694,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl162),
  inference(forward_demodulation,[],[f42635,f45])).
fof(f42635,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | $spl162),
  inference(superposition,[],[f668,f18342])).
fof(f42689,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | $spl162),
  inference(forward_demodulation,[],[f42634,f45])).
fof(f42634,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl162),
  inference(superposition,[],[f591,f18342])).
fof(f42686,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl162),
  inference(forward_demodulation,[],[f42685,f45])).
fof(f42685,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl162),
  inference(forward_demodulation,[],[f42633,f45])).
fof(f42633,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | $spl162),
  inference(superposition,[],[f590,f18342])).
fof(f42682,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),sK0)))),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | $spl162),
  inference(forward_demodulation,[],[f42632,f45])).
fof(f42632,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl162),
  inference(superposition,[],[f519,f18342])).
fof(f42679,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),sK0)))),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | $spl162),
  inference(forward_demodulation,[],[f42678,f45])).
fof(f42678,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | $spl162),
  inference(forward_demodulation,[],[f42631,f45])).
fof(f42631,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | $spl162),
  inference(superposition,[],[f518,f18342])).
fof(f42673,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | $spl162),
  inference(forward_demodulation,[],[f42630,f45])).
fof(f42630,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl162),
  inference(superposition,[],[f446,f18342])).
fof(f42668,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | $spl162),
  inference(forward_demodulation,[],[f42667,f45])).
fof(f42667,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | $spl162),
  inference(forward_demodulation,[],[f42629,f45])).
fof(f42629,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | $spl162),
  inference(superposition,[],[f445,f18342])).
fof(f42436,plain,(
  $false | ($spl173 | $spl177)),
  inference(subsumption_resolution,[],[f42431,f18434])).
fof(f42431,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl177),
  inference(trivial_inequality_removal,[],[f42430])).
fof(f42430,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl177),
  inference(superposition,[],[f18438,f44])).
fof(f42165,plain,(
  $false | ($spl165 | $spl169)),
  inference(subsumption_resolution,[],[f42160,f18344])).
fof(f42160,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl169),
  inference(trivial_inequality_removal,[],[f42159])).
fof(f42159,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl169),
  inference(superposition,[],[f18348,f44])).
fof(f24605,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK4))) | $spl280),
  inference(cnf_transformation,[],[f24605_D])).
fof(f24605_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK4))) <=> ~$spl280),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl280])])).
fof(f39884,plain,(
  s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK4,bool,sK1)))) | $spl304),
  inference(cnf_transformation,[],[f39884_D])).
fof(f39884_D,plain,(
  s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK4,bool,sK1)))) <=> ~$spl304),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl304])])).
fof(f11326,plain,(
  s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | $spl139),
  inference(cnf_transformation,[],[f11326_D])).
fof(f11326_D,plain,(
  s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) <=> ~$spl139),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl139])])).
fof(f24611,plain,(
  ( ! [X112] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X112))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X112,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X112,sK4,bool,sK1))))) ) | $spl292),
  inference(cnf_transformation,[],[f24611_D])).
fof(f24611_D,plain,(
  ( ! [X112] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X112))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X112,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X112,sK4,bool,sK1))))) ) <=> ~$spl292),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl292])])).
fof(f41636,plain,(
  ( ! [X118] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X118))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X118,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X118,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl33),
  inference(superposition,[],[f4425,f18919])).
fof(f18919,plain,(
  ( ! [X8] : (s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X8)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X8,sK4,sK1)),s(fun(sK1,bool),X8))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X8,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl33),
  inference(trivial_inequality_removal,[],[f18918])).
fof(f18918,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X8)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X8,sK4,sK1)),s(fun(sK1,bool),sK4))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X8,sK4,sK1)),s(fun(sK1,bool),X8)))) ) | $spl33),
  inference(superposition,[],[f17714,f44])).
fof(f17714,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(X5,sK4,sK1)),s(fun(sK1,bool),X5))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X5)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X5,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl33),
  inference(superposition,[],[f11657,f44])).
fof(f11657,plain,(
  ( ! [X1] : (s(bool,'const.bool.IN_2'(s(sK1,sK5(X1,sK4,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X1,sK4,sK1)),s(fun(sK1,bool),X1))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X1)),sK4,bool,sK1))))) ) | $spl33),
  inference(superposition,[],[f11371,f41])).
fof(f11371,plain,(
  s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | $spl33),
  inference(trivial_inequality_removal,[],[f11370])).
fof(f11370,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | $spl33),
  inference(superposition,[],[f10811,f44])).
fof(f10811,plain,(
  s(bool,'F_0') != s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | $spl33),
  inference(cnf_transformation,[],[f10811_D])).
fof(f10811_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f4425,plain,(
  ( ! [X2,X3] : (s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X3,sK1)),s(fun(sK1,bool),X2))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X2,X3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X2,X3,bool,sK1))))) )),
  inference(trivial_inequality_removal,[],[f4424])).
fof(f4424,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X3,sK1)),s(fun(sK1,bool),X2))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X2,X3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X2,X3,bool,sK1))))) )),
  inference(superposition,[],[f1452,f44])).
fof(f1452,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,f(s(sK1,sK6(X2,X3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X3,sK1)),s(fun(sK1,bool),X2))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X2,X3,bool,sK1))))) )),
  inference(superposition,[],[f97,f44])).
fof(f97,plain,(
  ( ! [X6,X7] : (s(bool,f(s(sK1,sK6(X7,X6,bool,sK1)))) != s(bool,g(s(sK1,sK6(X7,X6,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X6),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X6))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X6,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X6,sK1)),s(fun(sK1,bool),X7)))) )),
  inference(superposition,[],[f56,f53])).
fof(f56,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X0,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X0,sK1)),s(fun(sK1,bool),X0))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X0),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X0))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))))) )),
  inference(superposition,[],[f39,f41])).
fof(f24610,plain,(
  ( ! [X111] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X111))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X111,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X111,sK4,bool,sK1))))) ) | $spl290),
  inference(cnf_transformation,[],[f24610_D])).
fof(f24610_D,plain,(
  ( ! [X111] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X111))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X111,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X111,sK4,bool,sK1))))) ) <=> ~$spl290),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl290])])).
fof(f41635,plain,(
  ( ! [X117] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X117))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X117,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X117,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl33),
  inference(superposition,[],[f4413,f18919])).
fof(f4413,plain,(
  ( ! [X0,X1] : (s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X1,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X1,sK1)),s(fun(sK1,bool),X0))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X1),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X1))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X0,X1,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X0,X1,bool,sK1))))) )),
  inference(trivial_inequality_removal,[],[f4411])).
fof(f4411,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'T_0') | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X1),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X1))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X1,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X1,sK1)),s(fun(sK1,bool),X0))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X0,X1,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X0,X1,bool,sK1))))) )),
  inference(superposition,[],[f1451,f44])).
fof(f1451,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,f(s(sK1,sK6(X0,X1,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X1),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X1))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X1,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X1,sK1)),s(fun(sK1,bool),X0))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X0,X1,bool,sK1))))) )),
  inference(superposition,[],[f97,f44])).
fof(f24609,plain,(
  ( ! [X110] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X110))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK4,X110,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK4,X110,bool,sK1))))) ) | $spl288),
  inference(cnf_transformation,[],[f24609_D])).
fof(f24609_D,plain,(
  ( ! [X110] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X110))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK4,X110,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK4,X110,bool,sK1))))) ) <=> ~$spl288),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl288])])).
fof(f41634,plain,(
  ( ! [X116] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X116))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK4,X116,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK4,X116,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl33),
  inference(superposition,[],[f4410,f18919])).
fof(f4410,plain,(
  ( ! [X2,X3] : (s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X2,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X2,sK1)),s(fun(sK1,bool),X3))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X2),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X2))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X2,X3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X2,X3,bool,sK1))))) )),
  inference(trivial_inequality_removal,[],[f4409])).
fof(f4409,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X2),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X2))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X2,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X2,sK1)),s(fun(sK1,bool),X3))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X2,X3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X2,X3,bool,sK1))))) )),
  inference(superposition,[],[f1432,f44])).
fof(f1432,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,f(s(sK1,sK6(X2,X3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X2),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X2))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X2,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X2,sK1)),s(fun(sK1,bool),X3))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X2,X3,bool,sK1))))) )),
  inference(superposition,[],[f96,f44])).
fof(f96,plain,(
  ( ! [X4,X5] : (s(bool,f(s(sK1,sK6(X4,X5,bool,sK1)))) != s(bool,g(s(sK1,sK6(X4,X5,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X4,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X4,sK1)),s(fun(sK1,bool),X5)))) )),
  inference(superposition,[],[f56,f53])).
fof(f24608,plain,(
  ( ! [X109] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X109))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK4,X109,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK4,X109,bool,sK1))))) ) | $spl286),
  inference(cnf_transformation,[],[f24608_D])).
fof(f24608_D,plain,(
  ( ! [X109] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X109))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK4,X109,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK4,X109,bool,sK1))))) ) <=> ~$spl286),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl286])])).
fof(f41633,plain,(
  ( ! [X115] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X115))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK4,X115,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK4,X115,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl33),
  inference(superposition,[],[f4407,f18919])).
fof(f4407,plain,(
  ( ! [X0,X1] : (s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X0,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X0,sK1)),s(fun(sK1,bool),X1))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X0),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X0))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X0,X1,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X0,X1,bool,sK1))))) )),
  inference(trivial_inequality_removal,[],[f4405])).
fof(f4405,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'T_0') | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X0),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X0))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X0,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X0,sK1)),s(fun(sK1,bool),X1))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X0,X1,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X0,X1,bool,sK1))))) )),
  inference(superposition,[],[f1431,f44])).
fof(f1431,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,f(s(sK1,sK6(X0,X1,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X0),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X0))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X0,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X0,sK1)),s(fun(sK1,bool),X1))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X0,X1,bool,sK1))))) )),
  inference(superposition,[],[f96,f44])).
fof(f24607,plain,(
  ( ! [X108] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X108))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X108,sK4,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X108,sK4,sK1)),s(fun(sK1,bool),X108)))) ) | $spl284),
  inference(cnf_transformation,[],[f24607_D])).
fof(f24607_D,plain,(
  ( ! [X108] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X108))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X108,sK4,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X108,sK4,sK1)),s(fun(sK1,bool),X108)))) ) <=> ~$spl284),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl284])])).
fof(f41630,plain,(
  ( ! [X114] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X114))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X114,sK4,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X114,sK4,sK1)),s(fun(sK1,bool),X114))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl33),
  inference(superposition,[],[f95,f18919])).
fof(f95,plain,(
  ( ! [X2,X3] : (s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X2,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X2,sK1)),s(fun(sK1,bool),X3))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X2),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X2))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X3,X2,sK1)),s(fun(sK1,bool),X2))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X3,X2,sK1)),s(fun(sK1,bool),X3)))) )),
  inference(superposition,[],[f56,f41])).
fof(f24606,plain,(
  ( ! [X107] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X107))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X107,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X107,sK1)),s(fun(sK1,bool),X107)))) ) | $spl282),
  inference(cnf_transformation,[],[f24606_D])).
fof(f24606_D,plain,(
  ( ! [X107] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X107))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X107,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X107,sK1)),s(fun(sK1,bool),X107)))) ) <=> ~$spl282),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl282])])).
fof(f41629,plain,(
  ( ! [X113] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X113))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X113,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X113,sK1)),s(fun(sK1,bool),X113))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl33),
  inference(superposition,[],[f94,f18919])).
fof(f94,plain,(
  ( ! [X0,X1] : (s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X0,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X0,sK1)),s(fun(sK1,bool),X1))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X0),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X0))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X0,X1,sK1)),s(fun(sK1,bool),X0))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X0,X1,sK1)),s(fun(sK1,bool),X1)))) )),
  inference(superposition,[],[f56,f41])).
fof(f23943,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK4))) | $spl266),
  inference(cnf_transformation,[],[f23943_D])).
fof(f23943_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK4))) <=> ~$spl266),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl266])])).
fof(f23949,plain,(
  ( ! [X111] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X111))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X111,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X111,sK4,bool,sK1))))) ) | $spl278),
  inference(cnf_transformation,[],[f23949_D])).
fof(f23949_D,plain,(
  ( ! [X111] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X111))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X111,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X111,sK4,bool,sK1))))) ) <=> ~$spl278),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl278])])).
fof(f40911,plain,(
  ( ! [X117] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X117))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X117,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X117,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl33),
  inference(superposition,[],[f4425,f18909])).
fof(f18909,plain,(
  ( ! [X7] : (s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X7)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X7,sK4,sK1)),s(fun(sK1,bool),X7))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X7,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl33),
  inference(trivial_inequality_removal,[],[f18907])).
fof(f18907,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X7)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X7,sK4,sK1)),s(fun(sK1,bool),sK4))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X7,sK4,sK1)),s(fun(sK1,bool),X7)))) ) | $spl33),
  inference(superposition,[],[f17713,f44])).
fof(f17713,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(X4,sK4,sK1)),s(fun(sK1,bool),X4))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X4)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X4,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl33),
  inference(superposition,[],[f11657,f44])).
fof(f23948,plain,(
  ( ! [X110] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X110))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X110,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X110,sK4,bool,sK1))))) ) | $spl276),
  inference(cnf_transformation,[],[f23948_D])).
fof(f23948_D,plain,(
  ( ! [X110] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X110))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X110,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X110,sK4,bool,sK1))))) ) <=> ~$spl276),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl276])])).
fof(f40910,plain,(
  ( ! [X116] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X116))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X116,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X116,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl33),
  inference(superposition,[],[f4413,f18909])).
fof(f23947,plain,(
  ( ! [X109] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X109))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK4,X109,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK4,X109,bool,sK1))))) ) | $spl274),
  inference(cnf_transformation,[],[f23947_D])).
fof(f23947_D,plain,(
  ( ! [X109] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X109))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK4,X109,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK4,X109,bool,sK1))))) ) <=> ~$spl274),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl274])])).
fof(f40909,plain,(
  ( ! [X115] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X115))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK4,X115,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK4,X115,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl33),
  inference(superposition,[],[f4410,f18909])).
fof(f23946,plain,(
  ( ! [X108] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X108))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK4,X108,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK4,X108,bool,sK1))))) ) | $spl272),
  inference(cnf_transformation,[],[f23946_D])).
fof(f23946_D,plain,(
  ( ! [X108] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X108))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK4,X108,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK4,X108,bool,sK1))))) ) <=> ~$spl272),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl272])])).
fof(f40908,plain,(
  ( ! [X114] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X114))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK4,X114,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK4,X114,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl33),
  inference(superposition,[],[f4407,f18909])).
fof(f23945,plain,(
  ( ! [X107] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X107))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X107,sK4,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X107,sK4,sK1)),s(fun(sK1,bool),X107)))) ) | $spl270),
  inference(cnf_transformation,[],[f23945_D])).
fof(f23945_D,plain,(
  ( ! [X107] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X107))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X107,sK4,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X107,sK4,sK1)),s(fun(sK1,bool),X107)))) ) <=> ~$spl270),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl270])])).
fof(f40905,plain,(
  ( ! [X113] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X113))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X113,sK4,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X113,sK4,sK1)),s(fun(sK1,bool),X113))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl33),
  inference(superposition,[],[f95,f18909])).
fof(f23944,plain,(
  ( ! [X106] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X106))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X106,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X106,sK1)),s(fun(sK1,bool),X106)))) ) | $spl268),
  inference(cnf_transformation,[],[f23944_D])).
fof(f23944_D,plain,(
  ( ! [X106] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X106))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X106,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X106,sK1)),s(fun(sK1,bool),X106)))) ) <=> ~$spl268),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl268])])).
fof(f40904,plain,(
  ( ! [X112] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X112))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X112,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X112,sK1)),s(fun(sK1,bool),X112))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl33),
  inference(superposition,[],[f94,f18909])).
fof(f23269,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK4))) | $spl252),
  inference(cnf_transformation,[],[f23269_D])).
fof(f23269_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK4))) <=> ~$spl252),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl252])])).
fof(f23275,plain,(
  ( ! [X96] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X96))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X96,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X96,sK4,bool,sK1))))) ) | $spl264),
  inference(cnf_transformation,[],[f23275_D])).
fof(f23275_D,plain,(
  ( ! [X96] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X96))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X96,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X96,sK4,bool,sK1))))) ) <=> ~$spl264),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl264])])).
fof(f40145,plain,(
  ( ! [X102] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X102))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X102,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X102,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl33),
  inference(superposition,[],[f4449,f18862])).
fof(f18862,plain,(
  ( ! [X8] : (s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X8)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X8,sK1)),s(fun(sK1,bool),X8))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X8,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl33),
  inference(trivial_inequality_removal,[],[f18861])).
fof(f18861,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X8)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X8,sK1)),s(fun(sK1,bool),sK4))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X8,sK1)),s(fun(sK1,bool),X8)))) ) | $spl33),
  inference(superposition,[],[f17696,f44])).
fof(f17696,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X5,sK1)),s(fun(sK1,bool),X5))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X5)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X5,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl33),
  inference(superposition,[],[f11656,f44])).
fof(f11656,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X0,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X0,sK1)),s(fun(sK1,bool),X0))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X0)),sK4,bool,sK1))))) ) | $spl33),
  inference(superposition,[],[f11371,f41])).
fof(f4449,plain,(
  ( ! [X2,X3] : (s(bool,'const.bool.IN_2'(s(sK1,sK5(X3,sK3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X3,sK3,sK1)),s(fun(sK1,bool),X2))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X2,X3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X2,X3,bool,sK1))))) )),
  inference(trivial_inequality_removal,[],[f4448])).
fof(f4448,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X3,sK3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X3,sK3,sK1)),s(fun(sK1,bool),X2))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X2,X3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X2,X3,bool,sK1))))) )),
  inference(superposition,[],[f1536,f44])).
fof(f1536,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,f(s(sK1,sK6(X2,X3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X3,sK3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X3,sK3,sK1)),s(fun(sK1,bool),X2))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X2,X3,bool,sK1))))) )),
  inference(superposition,[],[f134,f44])).
fof(f134,plain,(
  ( ! [X6,X7] : (s(bool,f(s(sK1,sK6(X7,X6,bool,sK1)))) != s(bool,g(s(sK1,sK6(X7,X6,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X6),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X6))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X6,sK3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X6,sK3,sK1)),s(fun(sK1,bool),X7)))) )),
  inference(superposition,[],[f57,f53])).
fof(f57,plain,(
  ( ! [X1] : (s(bool,'const.bool.IN_2'(s(sK1,sK5(X1,sK3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X1,sK3,sK1)),s(fun(sK1,bool),X1))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X1),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X1))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))))) )),
  inference(superposition,[],[f39,f41])).
fof(f23274,plain,(
  ( ! [X95] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X95))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X95,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X95,sK4,bool,sK1))))) ) | $spl262),
  inference(cnf_transformation,[],[f23274_D])).
fof(f23274_D,plain,(
  ( ! [X95] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X95))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X95,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X95,sK4,bool,sK1))))) ) <=> ~$spl262),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl262])])).
fof(f40144,plain,(
  ( ! [X101] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X101))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X101,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X101,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl33),
  inference(superposition,[],[f4446,f18862])).
fof(f4446,plain,(
  ( ! [X0,X1] : (s(bool,'const.bool.IN_2'(s(sK1,sK5(X1,sK3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X1,sK3,sK1)),s(fun(sK1,bool),X0))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X1),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X1))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X0,X1,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X0,X1,bool,sK1))))) )),
  inference(trivial_inequality_removal,[],[f4444])).
fof(f4444,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'T_0') | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X1),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X1))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X1,sK3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X1,sK3,sK1)),s(fun(sK1,bool),X0))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X0,X1,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X0,X1,bool,sK1))))) )),
  inference(superposition,[],[f1535,f44])).
fof(f1535,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,f(s(sK1,sK6(X0,X1,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X1),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X1))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X1,sK3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X1,sK3,sK1)),s(fun(sK1,bool),X0))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X0,X1,bool,sK1))))) )),
  inference(superposition,[],[f134,f44])).
fof(f23273,plain,(
  ( ! [X94] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X94))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK4,X94,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK4,X94,bool,sK1))))) ) | $spl260),
  inference(cnf_transformation,[],[f23273_D])).
fof(f23273_D,plain,(
  ( ! [X94] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X94))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK4,X94,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK4,X94,bool,sK1))))) ) <=> ~$spl260),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl260])])).
fof(f40143,plain,(
  ( ! [X100] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X100))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK4,X100,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK4,X100,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl33),
  inference(superposition,[],[f4443,f18862])).
fof(f4443,plain,(
  ( ! [X2,X3] : (s(bool,'const.bool.IN_2'(s(sK1,sK5(X2,sK3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X2,sK3,sK1)),s(fun(sK1,bool),X3))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X2),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X2))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X2,X3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X2,X3,bool,sK1))))) )),
  inference(trivial_inequality_removal,[],[f4442])).
fof(f4442,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X2),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X2))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X2,sK3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X2,sK3,sK1)),s(fun(sK1,bool),X3))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X2,X3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X2,X3,bool,sK1))))) )),
  inference(superposition,[],[f1530,f44])).
fof(f1530,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,f(s(sK1,sK6(X2,X3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X2),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X2))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X2,sK3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X2,sK3,sK1)),s(fun(sK1,bool),X3))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X2,X3,bool,sK1))))) )),
  inference(superposition,[],[f133,f44])).
fof(f133,plain,(
  ( ! [X4,X5] : (s(bool,f(s(sK1,sK6(X4,X5,bool,sK1)))) != s(bool,g(s(sK1,sK6(X4,X5,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X4,sK3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X4,sK3,sK1)),s(fun(sK1,bool),X5)))) )),
  inference(superposition,[],[f57,f53])).
fof(f23272,plain,(
  ( ! [X93] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X93))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK4,X93,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK4,X93,bool,sK1))))) ) | $spl258),
  inference(cnf_transformation,[],[f23272_D])).
fof(f23272_D,plain,(
  ( ! [X93] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X93))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK4,X93,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK4,X93,bool,sK1))))) ) <=> ~$spl258),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl258])])).
fof(f40142,plain,(
  ( ! [X99] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X99))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK4,X99,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK4,X99,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl33),
  inference(superposition,[],[f4440,f18862])).
fof(f4440,plain,(
  ( ! [X0,X1] : (s(bool,'const.bool.IN_2'(s(sK1,sK5(X0,sK3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X0,sK3,sK1)),s(fun(sK1,bool),X1))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X0),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X0))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X0,X1,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X0,X1,bool,sK1))))) )),
  inference(trivial_inequality_removal,[],[f4438])).
fof(f4438,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'T_0') | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X0),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X0))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X0,sK3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X0,sK3,sK1)),s(fun(sK1,bool),X1))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X0,X1,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X0,X1,bool,sK1))))) )),
  inference(superposition,[],[f1529,f44])).
fof(f1529,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,f(s(sK1,sK6(X0,X1,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X0),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X0))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X0,sK3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X0,sK3,sK1)),s(fun(sK1,bool),X1))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X0,X1,bool,sK1))))) )),
  inference(superposition,[],[f133,f44])).
fof(f23271,plain,(
  ( ! [X92] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X92))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X92,sK4,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X92,sK4,sK1)),s(fun(sK1,bool),X92)))) ) | $spl256),
  inference(cnf_transformation,[],[f23271_D])).
fof(f23271_D,plain,(
  ( ! [X92] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X92))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X92,sK4,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X92,sK4,sK1)),s(fun(sK1,bool),X92)))) ) <=> ~$spl256),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl256])])).
fof(f40139,plain,(
  ( ! [X98] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X98))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X98,sK4,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X98,sK4,sK1)),s(fun(sK1,bool),X98))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl33),
  inference(superposition,[],[f132,f18862])).
fof(f132,plain,(
  ( ! [X2,X3] : (s(bool,'const.bool.IN_2'(s(sK1,sK5(X2,sK3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X2,sK3,sK1)),s(fun(sK1,bool),X3))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X2),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X2))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X3,X2,sK1)),s(fun(sK1,bool),X2))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X3,X2,sK1)),s(fun(sK1,bool),X3)))) )),
  inference(superposition,[],[f57,f41])).
fof(f23270,plain,(
  ( ! [X91] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X91))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X91,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X91,sK1)),s(fun(sK1,bool),X91)))) ) | $spl254),
  inference(cnf_transformation,[],[f23270_D])).
fof(f23270_D,plain,(
  ( ! [X91] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X91))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X91,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X91,sK1)),s(fun(sK1,bool),X91)))) ) <=> ~$spl254),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl254])])).
fof(f40138,plain,(
  ( ! [X97] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X97))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X97,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X97,sK1)),s(fun(sK1,bool),X97))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl33),
  inference(superposition,[],[f131,f18862])).
fof(f131,plain,(
  ( ! [X0,X1] : (s(bool,'const.bool.IN_2'(s(sK1,sK5(X0,sK3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X0,sK3,sK1)),s(fun(sK1,bool),X1))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X0),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X0))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X0,X1,sK1)),s(fun(sK1,bool),X0))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X0,X1,sK1)),s(fun(sK1,bool),X1)))) )),
  inference(superposition,[],[f57,f41])).
fof(f22609,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK4))) | $spl236),
  inference(cnf_transformation,[],[f22609_D])).
fof(f22609_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK4))) <=> ~$spl236),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl236])])).
fof(f22617,plain,(
  ( ! [X95] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X95))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X95,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X95,sK4,bool,sK1))))) ) | $spl250),
  inference(cnf_transformation,[],[f22617_D])).
fof(f22617_D,plain,(
  ( ! [X95] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X95))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X95,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X95,sK4,bool,sK1))))) ) <=> ~$spl250),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl250])])).
fof(f39422,plain,(
  ( ! [X101] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X101))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X101,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X101,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl33),
  inference(superposition,[],[f4449,f18852])).
fof(f18852,plain,(
  ( ! [X7] : (s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X7)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X7,sK1)),s(fun(sK1,bool),X7))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X7,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl33),
  inference(trivial_inequality_removal,[],[f18850])).
fof(f18850,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X7)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X7,sK1)),s(fun(sK1,bool),sK4))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X7,sK1)),s(fun(sK1,bool),X7)))) ) | $spl33),
  inference(superposition,[],[f17695,f44])).
fof(f17695,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X4,sK1)),s(fun(sK1,bool),X4))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X4)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl33),
  inference(superposition,[],[f11656,f44])).
fof(f22616,plain,(
  ( ! [X94] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X94))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X94,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X94,sK4,bool,sK1))))) ) | $spl248),
  inference(cnf_transformation,[],[f22616_D])).
fof(f22616_D,plain,(
  ( ! [X94] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X94))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X94,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X94,sK4,bool,sK1))))) ) <=> ~$spl248),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl248])])).
fof(f39421,plain,(
  ( ! [X100] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X100))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X100,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X100,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl33),
  inference(superposition,[],[f4446,f18852])).
fof(f22615,plain,(
  ( ! [X93] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X93))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK4,X93,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK4,X93,bool,sK1))))) ) | $spl246),
  inference(cnf_transformation,[],[f22615_D])).
fof(f22615_D,plain,(
  ( ! [X93] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X93))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK4,X93,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK4,X93,bool,sK1))))) ) <=> ~$spl246),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl246])])).
fof(f39420,plain,(
  ( ! [X99] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X99))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK4,X99,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK4,X99,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl33),
  inference(superposition,[],[f4443,f18852])).
fof(f22614,plain,(
  ( ! [X92] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X92))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK4,X92,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK4,X92,bool,sK1))))) ) | $spl244),
  inference(cnf_transformation,[],[f22614_D])).
fof(f22614_D,plain,(
  ( ! [X92] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X92))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK4,X92,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK4,X92,bool,sK1))))) ) <=> ~$spl244),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl244])])).
fof(f39419,plain,(
  ( ! [X98] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X98))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK4,X98,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK4,X98,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl33),
  inference(superposition,[],[f4440,f18852])).
fof(f22613,plain,(
  ( ! [X91] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X91))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X91,sK4,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X91,sK4,sK1)),s(fun(sK1,bool),X91)))) ) | $spl242),
  inference(cnf_transformation,[],[f22613_D])).
fof(f22613_D,plain,(
  ( ! [X91] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X91))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X91,sK4,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X91,sK4,sK1)),s(fun(sK1,bool),X91)))) ) <=> ~$spl242),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl242])])).
fof(f39416,plain,(
  ( ! [X97] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X97))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X97,sK4,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X97,sK4,sK1)),s(fun(sK1,bool),X97))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl33),
  inference(superposition,[],[f132,f18852])).
fof(f22612,plain,(
  ( ! [X90] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X90))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X90,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X90,sK1)),s(fun(sK1,bool),X90)))) ) | $spl240),
  inference(cnf_transformation,[],[f22612_D])).
fof(f22612_D,plain,(
  ( ! [X90] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X90))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X90,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X90,sK1)),s(fun(sK1,bool),X90)))) ) <=> ~$spl240),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl240])])).
fof(f39415,plain,(
  ( ! [X96] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X96))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X96,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X96,sK1)),s(fun(sK1,bool),X96))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl33),
  inference(superposition,[],[f131,f18852])).
fof(f21958,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK3))) | $spl222),
  inference(cnf_transformation,[],[f21958_D])).
fof(f21958_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK3))) <=> ~$spl222),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl222])])).
fof(f36891,plain,(
  s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | $spl302),
  inference(cnf_transformation,[],[f36891_D])).
fof(f36891_D,plain,(
  s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) <=> ~$spl302),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl302])])).
fof(f11003,plain,(
  s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | $spl81),
  inference(cnf_transformation,[],[f11003_D])).
fof(f11003_D,plain,(
  s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) <=> ~$spl81),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl81])])).
fof(f21964,plain,(
  ( ! [X158] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X158))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X158,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X158,sK3,bool,sK1))))) ) | $spl234),
  inference(cnf_transformation,[],[f21964_D])).
fof(f21964_D,plain,(
  ( ! [X158] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X158))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X158,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X158,sK3,bool,sK1))))) ) <=> ~$spl234),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl234])])).
fof(f38813,plain,(
  ( ! [X164] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X164))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X164,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X164,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl33),
  inference(superposition,[],[f4488,f18806])).
fof(f18806,plain,(
  ( ! [X8] : (s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),X8),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X8,sK3,sK1)),s(fun(sK1,bool),X8))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X8,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl33),
  inference(trivial_inequality_removal,[],[f18805])).
fof(f18805,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),X8),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X8,sK3,sK1)),s(fun(sK1,bool),sK3))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X8,sK3,sK1)),s(fun(sK1,bool),X8)))) ) | $spl33),
  inference(superposition,[],[f17678,f44])).
fof(f17678,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(X5,sK3,sK1)),s(fun(sK1,bool),X5))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),X5),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X5,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl33),
  inference(superposition,[],[f11653,f44])).
fof(f11653,plain,(
  ( ! [X1] : (s(bool,'const.bool.IN_2'(s(sK1,sK5(X1,sK3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X1,sK3,sK1)),s(fun(sK1,bool),X1))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),X1),s(fun(sK1,bool),sK4)),sK4,bool,sK1))))) ) | $spl33),
  inference(superposition,[],[f11371,f41])).
fof(f4488,plain,(
  ( ! [X2,X3] : (s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X3,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X3,sK1)),s(fun(sK1,bool),X2))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X3))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X2,X3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X2,X3,bool,sK1))))) )),
  inference(trivial_inequality_removal,[],[f4487])).
fof(f4487,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X3))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X3,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X3,sK1)),s(fun(sK1,bool),X2))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X2,X3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X2,X3,bool,sK1))))) )),
  inference(superposition,[],[f1572,f44])).
fof(f1572,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,f(s(sK1,sK6(X2,X3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X3))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X3,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X3,sK1)),s(fun(sK1,bool),X2))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X2,X3,bool,sK1))))) )),
  inference(superposition,[],[f186,f44])).
fof(f186,plain,(
  ( ! [X6,X7] : (s(bool,f(s(sK1,sK6(X7,X6,bool,sK1)))) != s(bool,g(s(sK1,sK6(X7,X6,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X6))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X6))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X6,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X6,sK1)),s(fun(sK1,bool),X7)))) )),
  inference(superposition,[],[f66,f53])).
fof(f66,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X0,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X0,sK1)),s(fun(sK1,bool),X0))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X0))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X0)))))) )),
  inference(superposition,[],[f39,f41])).
fof(f21963,plain,(
  ( ! [X157] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X157))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X157,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X157,sK3,bool,sK1))))) ) | $spl232),
  inference(cnf_transformation,[],[f21963_D])).
fof(f21963_D,plain,(
  ( ! [X157] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X157))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X157,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X157,sK3,bool,sK1))))) ) <=> ~$spl232),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl232])])).
fof(f38812,plain,(
  ( ! [X163] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X163))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X163,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X163,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl33),
  inference(superposition,[],[f4485,f18806])).
fof(f4485,plain,(
  ( ! [X0,X1] : (s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X1,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X1,sK1)),s(fun(sK1,bool),X0))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X1))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X1))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X0,X1,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X0,X1,bool,sK1))))) )),
  inference(trivial_inequality_removal,[],[f4483])).
fof(f4483,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'T_0') | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X1))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X1))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X1,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X1,sK1)),s(fun(sK1,bool),X0))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X0,X1,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X0,X1,bool,sK1))))) )),
  inference(superposition,[],[f1571,f44])).
fof(f1571,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,f(s(sK1,sK6(X0,X1,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X1))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X1))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X1,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X1,sK1)),s(fun(sK1,bool),X0))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X0,X1,bool,sK1))))) )),
  inference(superposition,[],[f186,f44])).
fof(f21962,plain,(
  ( ! [X156] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X156))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK3,X156,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK3,X156,bool,sK1))))) ) | $spl230),
  inference(cnf_transformation,[],[f21962_D])).
fof(f21962_D,plain,(
  ( ! [X156] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X156))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK3,X156,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK3,X156,bool,sK1))))) ) <=> ~$spl230),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl230])])).
fof(f38811,plain,(
  ( ! [X162] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X162))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK3,X162,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK3,X162,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl33),
  inference(superposition,[],[f4482,f18806])).
fof(f4482,plain,(
  ( ! [X2,X3] : (s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X2,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X2,sK1)),s(fun(sK1,bool),X3))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X2))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X2))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X2,X3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X2,X3,bool,sK1))))) )),
  inference(trivial_inequality_removal,[],[f4481])).
fof(f4481,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X2))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X2))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X2,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X2,sK1)),s(fun(sK1,bool),X3))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X2,X3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X2,X3,bool,sK1))))) )),
  inference(superposition,[],[f1564,f44])).
fof(f1564,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,f(s(sK1,sK6(X2,X3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X2))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X2))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X2,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X2,sK1)),s(fun(sK1,bool),X3))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X2,X3,bool,sK1))))) )),
  inference(superposition,[],[f185,f44])).
fof(f185,plain,(
  ( ! [X4,X5] : (s(bool,f(s(sK1,sK6(X4,X5,bool,sK1)))) != s(bool,g(s(sK1,sK6(X4,X5,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X4,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X4,sK1)),s(fun(sK1,bool),X5)))) )),
  inference(superposition,[],[f66,f53])).
fof(f21961,plain,(
  ( ! [X155] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X155))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK3,X155,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK3,X155,bool,sK1))))) ) | $spl228),
  inference(cnf_transformation,[],[f21961_D])).
fof(f21961_D,plain,(
  ( ! [X155] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X155))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK3,X155,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK3,X155,bool,sK1))))) ) <=> ~$spl228),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl228])])).
fof(f38810,plain,(
  ( ! [X161] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X161))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK3,X161,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK3,X161,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl33),
  inference(superposition,[],[f4479,f18806])).
fof(f4479,plain,(
  ( ! [X0,X1] : (s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X0,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X0,sK1)),s(fun(sK1,bool),X1))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X0))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X0))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X0,X1,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X0,X1,bool,sK1))))) )),
  inference(trivial_inequality_removal,[],[f4477])).
fof(f4477,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'T_0') | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X0))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X0))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X0,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X0,sK1)),s(fun(sK1,bool),X1))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X0,X1,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X0,X1,bool,sK1))))) )),
  inference(superposition,[],[f1563,f44])).
fof(f1563,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,f(s(sK1,sK6(X0,X1,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X0))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X0))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X0,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X0,sK1)),s(fun(sK1,bool),X1))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X0,X1,bool,sK1))))) )),
  inference(superposition,[],[f185,f44])).
fof(f21960,plain,(
  ( ! [X144] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X144))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X144,sK3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X144,sK3,sK1)),s(fun(sK1,bool),X144)))) ) | $spl226),
  inference(cnf_transformation,[],[f21960_D])).
fof(f21960_D,plain,(
  ( ! [X144] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X144))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X144,sK3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X144,sK3,sK1)),s(fun(sK1,bool),X144)))) ) <=> ~$spl226),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl226])])).
fof(f38797,plain,(
  ( ! [X150] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X150))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X150,sK3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X150,sK3,sK1)),s(fun(sK1,bool),X150))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl33),
  inference(superposition,[],[f184,f18806])).
fof(f184,plain,(
  ( ! [X2,X3] : (s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X2,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X2,sK1)),s(fun(sK1,bool),X3))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X2))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X2))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X3,X2,sK1)),s(fun(sK1,bool),X2))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X3,X2,sK1)),s(fun(sK1,bool),X3)))) )),
  inference(superposition,[],[f66,f41])).
fof(f21959,plain,(
  ( ! [X143] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X143))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X143,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X143,sK1)),s(fun(sK1,bool),X143)))) ) | $spl224),
  inference(cnf_transformation,[],[f21959_D])).
fof(f21959_D,plain,(
  ( ! [X143] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X143))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X143,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X143,sK1)),s(fun(sK1,bool),X143)))) ) <=> ~$spl224),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl224])])).
fof(f38796,plain,(
  ( ! [X149] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X149))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X149,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X149,sK1)),s(fun(sK1,bool),X149))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl33),
  inference(superposition,[],[f183,f18806])).
fof(f183,plain,(
  ( ! [X0,X1] : (s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X0,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X0,sK1)),s(fun(sK1,bool),X1))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X0))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X0))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X0,X1,sK1)),s(fun(sK1,bool),X0))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X0,X1,sK1)),s(fun(sK1,bool),X1)))) )),
  inference(superposition,[],[f66,f41])).
fof(f21250,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK3))) | $spl208),
  inference(cnf_transformation,[],[f21250_D])).
fof(f21250_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK3))) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f21256,plain,(
  ( ! [X157] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X157))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X157,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X157,sK3,bool,sK1))))) ) | $spl220),
  inference(cnf_transformation,[],[f21256_D])).
fof(f21256_D,plain,(
  ( ! [X157] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X157))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X157,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X157,sK3,bool,sK1))))) ) <=> ~$spl220),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl220])])).
fof(f38062,plain,(
  ( ! [X163] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X163))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X163,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X163,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl33),
  inference(superposition,[],[f4488,f18796])).
fof(f18796,plain,(
  ( ! [X7] : (s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),X7),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X7,sK3,sK1)),s(fun(sK1,bool),X7))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X7,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl33),
  inference(trivial_inequality_removal,[],[f18794])).
fof(f18794,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),X7),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X7,sK3,sK1)),s(fun(sK1,bool),sK3))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X7,sK3,sK1)),s(fun(sK1,bool),X7)))) ) | $spl33),
  inference(superposition,[],[f17677,f44])).
fof(f17677,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(X4,sK3,sK1)),s(fun(sK1,bool),X4))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),X4),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X4,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl33),
  inference(superposition,[],[f11653,f44])).
fof(f21255,plain,(
  ( ! [X156] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X156))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X156,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X156,sK3,bool,sK1))))) ) | $spl218),
  inference(cnf_transformation,[],[f21255_D])).
fof(f21255_D,plain,(
  ( ! [X156] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X156))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X156,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X156,sK3,bool,sK1))))) ) <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f38061,plain,(
  ( ! [X162] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X162))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X162,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X162,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl33),
  inference(superposition,[],[f4485,f18796])).
fof(f21254,plain,(
  ( ! [X155] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X155))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK3,X155,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK3,X155,bool,sK1))))) ) | $spl216),
  inference(cnf_transformation,[],[f21254_D])).
fof(f21254_D,plain,(
  ( ! [X155] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X155))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK3,X155,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK3,X155,bool,sK1))))) ) <=> ~$spl216),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl216])])).
fof(f38060,plain,(
  ( ! [X161] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X161))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK3,X161,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK3,X161,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl33),
  inference(superposition,[],[f4482,f18796])).
fof(f21253,plain,(
  ( ! [X154] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X154))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK3,X154,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK3,X154,bool,sK1))))) ) | $spl214),
  inference(cnf_transformation,[],[f21253_D])).
fof(f21253_D,plain,(
  ( ! [X154] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X154))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK3,X154,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK3,X154,bool,sK1))))) ) <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f38059,plain,(
  ( ! [X160] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X160))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK3,X160,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK3,X160,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl33),
  inference(superposition,[],[f4479,f18796])).
fof(f21252,plain,(
  ( ! [X143] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X143))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X143,sK3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X143,sK3,sK1)),s(fun(sK1,bool),X143)))) ) | $spl212),
  inference(cnf_transformation,[],[f21252_D])).
fof(f21252_D,plain,(
  ( ! [X143] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X143))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X143,sK3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X143,sK3,sK1)),s(fun(sK1,bool),X143)))) ) <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f38046,plain,(
  ( ! [X149] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X149))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X149,sK3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X149,sK3,sK1)),s(fun(sK1,bool),X149))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl33),
  inference(superposition,[],[f184,f18796])).
fof(f21251,plain,(
  ( ! [X142] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X142))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X142,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X142,sK1)),s(fun(sK1,bool),X142)))) ) | $spl210),
  inference(cnf_transformation,[],[f21251_D])).
fof(f21251_D,plain,(
  ( ! [X142] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X142))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X142,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X142,sK1)),s(fun(sK1,bool),X142)))) ) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f38045,plain,(
  ( ! [X148] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X148))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X148,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X148,sK1)),s(fun(sK1,bool),X148))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl33),
  inference(superposition,[],[f183,f18796])).
fof(f20565,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK3))) | $spl194),
  inference(cnf_transformation,[],[f20565_D])).
fof(f20565_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK3))) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f20571,plain,(
  ( ! [X141] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X141))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X141,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X141,sK3,bool,sK1))))) ) | $spl206),
  inference(cnf_transformation,[],[f20571_D])).
fof(f20571_D,plain,(
  ( ! [X141] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X141))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X141,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X141,sK3,bool,sK1))))) ) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f37279,plain,(
  ( ! [X148] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X148))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X148,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X148,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl33),
  inference(superposition,[],[f4521,f18750])).
fof(f18750,plain,(
  ( ! [X8] : (s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),X8),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X8,sK1)),s(fun(sK1,bool),X8))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X8,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl33),
  inference(trivial_inequality_removal,[],[f18749])).
fof(f18749,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),X8),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X8,sK1)),s(fun(sK1,bool),sK3))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X8,sK1)),s(fun(sK1,bool),X8)))) ) | $spl33),
  inference(superposition,[],[f17648,f44])).
fof(f17648,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X5,sK1)),s(fun(sK1,bool),X5))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),X5),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X5,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl33),
  inference(superposition,[],[f11652,f44])).
fof(f11652,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X0,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X0,sK1)),s(fun(sK1,bool),X0))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),X0),s(fun(sK1,bool),sK4)),sK4,bool,sK1))))) ) | $spl33),
  inference(superposition,[],[f11371,f41])).
fof(f4521,plain,(
  ( ! [X2,X3] : (s(bool,'const.bool.IN_2'(s(sK1,sK5(X3,sK4,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X3,sK4,sK1)),s(fun(sK1,bool),X2))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X3))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X2,X3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X2,X3,bool,sK1))))) )),
  inference(trivial_inequality_removal,[],[f4520])).
fof(f4520,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X3))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X3,sK4,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X3,sK4,sK1)),s(fun(sK1,bool),X2))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X2,X3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X2,X3,bool,sK1))))) )),
  inference(superposition,[],[f1608,f44])).
fof(f1608,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,f(s(sK1,sK6(X2,X3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X3))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X3,sK4,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X3,sK4,sK1)),s(fun(sK1,bool),X2))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X2,X3,bool,sK1))))) )),
  inference(superposition,[],[f314,f44])).
fof(f314,plain,(
  ( ! [X6,X7] : (s(bool,f(s(sK1,sK6(X7,X6,bool,sK1)))) != s(bool,g(s(sK1,sK6(X7,X6,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X6))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X6))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X6,sK4,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X6,sK4,sK1)),s(fun(sK1,bool),X7)))) )),
  inference(superposition,[],[f67,f53])).
fof(f67,plain,(
  ( ! [X1] : (s(bool,'const.bool.IN_2'(s(sK1,sK5(X1,sK4,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X1,sK4,sK1)),s(fun(sK1,bool),X1))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X1))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X1)))))) )),
  inference(superposition,[],[f39,f41])).
fof(f20570,plain,(
  ( ! [X140] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X140))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X140,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X140,sK3,bool,sK1))))) ) | $spl204),
  inference(cnf_transformation,[],[f20570_D])).
fof(f20570_D,plain,(
  ( ! [X140] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X140))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X140,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X140,sK3,bool,sK1))))) ) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f37278,plain,(
  ( ! [X147] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X147))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X147,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X147,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl33),
  inference(superposition,[],[f4518,f18750])).
fof(f4518,plain,(
  ( ! [X0,X1] : (s(bool,'const.bool.IN_2'(s(sK1,sK5(X1,sK4,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X1,sK4,sK1)),s(fun(sK1,bool),X0))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X1))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X1))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X0,X1,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X0,X1,bool,sK1))))) )),
  inference(trivial_inequality_removal,[],[f4516])).
fof(f4516,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'T_0') | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X1))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X1))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X1,sK4,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X1,sK4,sK1)),s(fun(sK1,bool),X0))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X0,X1,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X0,X1,bool,sK1))))) )),
  inference(superposition,[],[f1607,f44])).
fof(f1607,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,f(s(sK1,sK6(X0,X1,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X1))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X1))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X1,sK4,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X1,sK4,sK1)),s(fun(sK1,bool),X0))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X0,X1,bool,sK1))))) )),
  inference(superposition,[],[f314,f44])).
fof(f20569,plain,(
  ( ! [X139] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X139))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK3,X139,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK3,X139,bool,sK1))))) ) | $spl202),
  inference(cnf_transformation,[],[f20569_D])).
fof(f20569_D,plain,(
  ( ! [X139] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X139))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK3,X139,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK3,X139,bool,sK1))))) ) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f37277,plain,(
  ( ! [X146] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X146))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK3,X146,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK3,X146,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl33),
  inference(superposition,[],[f4515,f18750])).
fof(f4515,plain,(
  ( ! [X2,X3] : (s(bool,'const.bool.IN_2'(s(sK1,sK5(X2,sK4,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X2,sK4,sK1)),s(fun(sK1,bool),X3))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X2))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X2))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X2,X3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X2,X3,bool,sK1))))) )),
  inference(trivial_inequality_removal,[],[f4514])).
fof(f4514,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X2))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X2))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X2,sK4,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X2,sK4,sK1)),s(fun(sK1,bool),X3))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X2,X3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X2,X3,bool,sK1))))) )),
  inference(superposition,[],[f1598,f44])).
fof(f1598,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,f(s(sK1,sK6(X2,X3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X2))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X2))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X2,sK4,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X2,sK4,sK1)),s(fun(sK1,bool),X3))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X2,X3,bool,sK1))))) )),
  inference(superposition,[],[f313,f44])).
fof(f313,plain,(
  ( ! [X4,X5] : (s(bool,f(s(sK1,sK6(X4,X5,bool,sK1)))) != s(bool,g(s(sK1,sK6(X4,X5,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X4,sK4,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X4,sK4,sK1)),s(fun(sK1,bool),X5)))) )),
  inference(superposition,[],[f67,f53])).
fof(f20568,plain,(
  ( ! [X138] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X138))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK3,X138,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK3,X138,bool,sK1))))) ) | $spl200),
  inference(cnf_transformation,[],[f20568_D])).
fof(f20568_D,plain,(
  ( ! [X138] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X138))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK3,X138,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK3,X138,bool,sK1))))) ) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f37276,plain,(
  ( ! [X145] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X145))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK3,X145,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK3,X145,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl33),
  inference(superposition,[],[f4512,f18750])).
fof(f4512,plain,(
  ( ! [X0,X1] : (s(bool,'const.bool.IN_2'(s(sK1,sK5(X0,sK4,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X0,sK4,sK1)),s(fun(sK1,bool),X1))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X0))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X0))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X0,X1,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X0,X1,bool,sK1))))) )),
  inference(trivial_inequality_removal,[],[f4510])).
fof(f4510,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'T_0') | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X0))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X0))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X0,sK4,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X0,sK4,sK1)),s(fun(sK1,bool),X1))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X0,X1,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X0,X1,bool,sK1))))) )),
  inference(superposition,[],[f1597,f44])).
fof(f1597,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,f(s(sK1,sK6(X0,X1,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X0))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X0))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X0,sK4,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X0,sK4,sK1)),s(fun(sK1,bool),X1))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X0,X1,bool,sK1))))) )),
  inference(superposition,[],[f313,f44])).
fof(f20567,plain,(
  ( ! [X132] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X132))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X132,sK3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X132,sK3,sK1)),s(fun(sK1,bool),X132)))) ) | $spl198),
  inference(cnf_transformation,[],[f20567_D])).
fof(f20567_D,plain,(
  ( ! [X132] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X132))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X132,sK3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X132,sK3,sK1)),s(fun(sK1,bool),X132)))) ) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f37267,plain,(
  ( ! [X138] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X138))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X138,sK3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X138,sK3,sK1)),s(fun(sK1,bool),X138))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl33),
  inference(superposition,[],[f312,f18750])).
fof(f312,plain,(
  ( ! [X2,X3] : (s(bool,'const.bool.IN_2'(s(sK1,sK5(X2,sK4,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X2,sK4,sK1)),s(fun(sK1,bool),X3))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X2))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X2))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X3,X2,sK1)),s(fun(sK1,bool),X2))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X3,X2,sK1)),s(fun(sK1,bool),X3)))) )),
  inference(superposition,[],[f67,f41])).
fof(f20566,plain,(
  ( ! [X131] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X131))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X131,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X131,sK1)),s(fun(sK1,bool),X131)))) ) | $spl196),
  inference(cnf_transformation,[],[f20566_D])).
fof(f20566_D,plain,(
  ( ! [X131] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X131))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X131,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X131,sK1)),s(fun(sK1,bool),X131)))) ) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f37266,plain,(
  ( ! [X137] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X137))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X137,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X137,sK1)),s(fun(sK1,bool),X137))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl33),
  inference(superposition,[],[f311,f18750])).
fof(f311,plain,(
  ( ! [X0,X1] : (s(bool,'const.bool.IN_2'(s(sK1,sK5(X0,sK4,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X0,sK4,sK1)),s(fun(sK1,bool),X1))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X0))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X0))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X0,X1,sK1)),s(fun(sK1,bool),X0))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X0,X1,sK1)),s(fun(sK1,bool),X1)))) )),
  inference(superposition,[],[f67,f41])).
fof(f19880,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK3))) | $spl178),
  inference(cnf_transformation,[],[f19880_D])).
fof(f19880_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK3))) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f19888,plain,(
  ( ! [X140] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X140))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X140,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X140,sK3,bool,sK1))))) ) | $spl192),
  inference(cnf_transformation,[],[f19888_D])).
fof(f19888_D,plain,(
  ( ! [X140] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X140))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X140,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X140,sK3,bool,sK1))))) ) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f36515,plain,(
  ( ! [X147] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X147))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X147,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X147,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl33),
  inference(superposition,[],[f4521,f18740])).
fof(f18740,plain,(
  ( ! [X7] : (s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),X7),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X7,sK1)),s(fun(sK1,bool),X7))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X7,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl33),
  inference(trivial_inequality_removal,[],[f18738])).
fof(f18738,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),X7),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X7,sK1)),s(fun(sK1,bool),sK3))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X7,sK1)),s(fun(sK1,bool),X7)))) ) | $spl33),
  inference(superposition,[],[f17647,f44])).
fof(f17647,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X4,sK1)),s(fun(sK1,bool),X4))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),X4),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X4,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl33),
  inference(superposition,[],[f11652,f44])).
fof(f19887,plain,(
  ( ! [X139] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X139))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X139,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X139,sK3,bool,sK1))))) ) | $spl190),
  inference(cnf_transformation,[],[f19887_D])).
fof(f19887_D,plain,(
  ( ! [X139] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X139))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X139,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X139,sK3,bool,sK1))))) ) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f36514,plain,(
  ( ! [X146] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X146))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X146,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X146,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl33),
  inference(superposition,[],[f4518,f18740])).
fof(f19886,plain,(
  ( ! [X138] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X138))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK3,X138,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK3,X138,bool,sK1))))) ) | $spl188),
  inference(cnf_transformation,[],[f19886_D])).
fof(f19886_D,plain,(
  ( ! [X138] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X138))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK3,X138,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK3,X138,bool,sK1))))) ) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f36513,plain,(
  ( ! [X145] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X145))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK3,X145,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK3,X145,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl33),
  inference(superposition,[],[f4515,f18740])).
fof(f19885,plain,(
  ( ! [X137] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X137))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK3,X137,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK3,X137,bool,sK1))))) ) | $spl186),
  inference(cnf_transformation,[],[f19885_D])).
fof(f19885_D,plain,(
  ( ! [X137] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X137))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK3,X137,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK3,X137,bool,sK1))))) ) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f36512,plain,(
  ( ! [X144] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X144))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK3,X144,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK3,X144,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl33),
  inference(superposition,[],[f4512,f18740])).
fof(f19884,plain,(
  ( ! [X131] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X131))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X131,sK3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X131,sK3,sK1)),s(fun(sK1,bool),X131)))) ) | $spl184),
  inference(cnf_transformation,[],[f19884_D])).
fof(f19884_D,plain,(
  ( ! [X131] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X131))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X131,sK3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X131,sK3,sK1)),s(fun(sK1,bool),X131)))) ) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f36503,plain,(
  ( ! [X137] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X137))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X137,sK3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X137,sK3,sK1)),s(fun(sK1,bool),X137))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl33),
  inference(superposition,[],[f312,f18740])).
fof(f19883,plain,(
  ( ! [X130] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X130))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X130,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X130,sK1)),s(fun(sK1,bool),X130)))) ) | $spl182),
  inference(cnf_transformation,[],[f19883_D])).
fof(f19883_D,plain,(
  ( ! [X130] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X130))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X130,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X130,sK1)),s(fun(sK1,bool),X130)))) ) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f36502,plain,(
  ( ! [X136] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X136))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X136,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X136,sK1)),s(fun(sK1,bool),X136))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl33),
  inference(superposition,[],[f311,f18740])).
fof(f33956,plain,(
  s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK3,bool,sK1)))) | $spl300),
  inference(cnf_transformation,[],[f33956_D])).
fof(f33956_D,plain,(
  s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK3,bool,sK1)))) <=> ~$spl300),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl300])])).
fof(f35702,plain,(
  ( ! [X116] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X116))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X116,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X116,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl23),
  inference(superposition,[],[f4425,f18712])).
fof(f18712,plain,(
  ( ! [X8] : (s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X8)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X8,sK4,sK1)),s(fun(sK1,bool),X8))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X8,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl23),
  inference(trivial_inequality_removal,[],[f18711])).
fof(f18711,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X8)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X8,sK4,sK1)),s(fun(sK1,bool),sK4))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X8,sK4,sK1)),s(fun(sK1,bool),X8)))) ) | $spl23),
  inference(superposition,[],[f17630,f44])).
fof(f17630,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(X5,sK4,sK1)),s(fun(sK1,bool),X5))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X5)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X5,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl23),
  inference(superposition,[],[f11515,f44])).
fof(f11515,plain,(
  ( ! [X1] : (s(bool,'const.bool.IN_2'(s(sK1,sK5(X1,sK4,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X1,sK4,sK1)),s(fun(sK1,bool),X1))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X1)),sK3,bool,sK1))))) ) | $spl23),
  inference(superposition,[],[f11056,f41])).
fof(f11056,plain,(
  s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | $spl23),
  inference(trivial_inequality_removal,[],[f11055])).
fof(f11055,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | $spl23),
  inference(superposition,[],[f10801,f44])).
fof(f10801,plain,(
  s(bool,'F_0') != s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | $spl23),
  inference(cnf_transformation,[],[f10801_D])).
fof(f10801_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f35701,plain,(
  ( ! [X115] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X115))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X115,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X115,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl23),
  inference(superposition,[],[f4413,f18712])).
fof(f35700,plain,(
  ( ! [X114] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X114))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK4,X114,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK4,X114,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl23),
  inference(superposition,[],[f4410,f18712])).
fof(f35699,plain,(
  ( ! [X113] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X113))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK4,X113,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK4,X113,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl23),
  inference(superposition,[],[f4407,f18712])).
fof(f35696,plain,(
  ( ! [X112] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X112))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X112,sK4,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X112,sK4,sK1)),s(fun(sK1,bool),X112))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl23),
  inference(superposition,[],[f95,f18712])).
fof(f35695,plain,(
  ( ! [X111] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X111))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X111,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X111,sK1)),s(fun(sK1,bool),X111))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl23),
  inference(superposition,[],[f94,f18712])).
fof(f34981,plain,(
  ( ! [X115] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X115))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X115,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X115,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl23),
  inference(superposition,[],[f4425,f18684])).
fof(f18684,plain,(
  ( ! [X7] : (s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X7)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X7,sK4,sK1)),s(fun(sK1,bool),X7))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X7,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl23),
  inference(trivial_inequality_removal,[],[f18682])).
fof(f18682,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X7)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X7,sK4,sK1)),s(fun(sK1,bool),sK4))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X7,sK4,sK1)),s(fun(sK1,bool),X7)))) ) | $spl23),
  inference(superposition,[],[f17629,f44])).
fof(f17629,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(X4,sK4,sK1)),s(fun(sK1,bool),X4))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X4)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X4,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl23),
  inference(superposition,[],[f11515,f44])).
fof(f34980,plain,(
  ( ! [X114] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X114))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X114,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X114,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl23),
  inference(superposition,[],[f4413,f18684])).
fof(f34979,plain,(
  ( ! [X113] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X113))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK4,X113,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK4,X113,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl23),
  inference(superposition,[],[f4410,f18684])).
fof(f34978,plain,(
  ( ! [X112] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X112))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK4,X112,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK4,X112,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl23),
  inference(superposition,[],[f4407,f18684])).
fof(f34975,plain,(
  ( ! [X111] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X111))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X111,sK4,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X111,sK4,sK1)),s(fun(sK1,bool),X111))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl23),
  inference(superposition,[],[f95,f18684])).
fof(f34974,plain,(
  ( ! [X110] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X110))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X110,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X110,sK1)),s(fun(sK1,bool),X110))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl23),
  inference(superposition,[],[f94,f18684])).
fof(f34214,plain,(
  ( ! [X100] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X100))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X100,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X100,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl23),
  inference(superposition,[],[f4449,f18656])).
fof(f18656,plain,(
  ( ! [X8] : (s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X8)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X8,sK1)),s(fun(sK1,bool),X8))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X8,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl23),
  inference(trivial_inequality_removal,[],[f18655])).
fof(f18655,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X8)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X8,sK1)),s(fun(sK1,bool),sK4))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X8,sK1)),s(fun(sK1,bool),X8)))) ) | $spl23),
  inference(superposition,[],[f17612,f44])).
fof(f17612,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X5,sK1)),s(fun(sK1,bool),X5))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X5)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X5,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl23),
  inference(superposition,[],[f11514,f44])).
fof(f11514,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X0,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X0,sK1)),s(fun(sK1,bool),X0))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X0)),sK3,bool,sK1))))) ) | $spl23),
  inference(superposition,[],[f11056,f41])).
fof(f34213,plain,(
  ( ! [X99] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X99))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X99,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X99,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl23),
  inference(superposition,[],[f4446,f18656])).
fof(f34212,plain,(
  ( ! [X98] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X98))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK4,X98,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK4,X98,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl23),
  inference(superposition,[],[f4443,f18656])).
fof(f34211,plain,(
  ( ! [X97] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X97))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK4,X97,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK4,X97,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl23),
  inference(superposition,[],[f4440,f18656])).
fof(f34208,plain,(
  ( ! [X96] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X96))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X96,sK4,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X96,sK4,sK1)),s(fun(sK1,bool),X96))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl23),
  inference(superposition,[],[f132,f18656])).
fof(f34207,plain,(
  ( ! [X95] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X95))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X95,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X95,sK1)),s(fun(sK1,bool),X95))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl23),
  inference(superposition,[],[f131,f18656])).
fof(f33495,plain,(
  ( ! [X99] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X99))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X99,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X99,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl23),
  inference(superposition,[],[f4449,f18646])).
fof(f18646,plain,(
  ( ! [X7] : (s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X7)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X7,sK1)),s(fun(sK1,bool),X7))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X7,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl23),
  inference(trivial_inequality_removal,[],[f18644])).
fof(f18644,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X7)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X7,sK1)),s(fun(sK1,bool),sK4))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X7,sK1)),s(fun(sK1,bool),X7)))) ) | $spl23),
  inference(superposition,[],[f17611,f44])).
fof(f17611,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X4,sK1)),s(fun(sK1,bool),X4))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X4)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl23),
  inference(superposition,[],[f11514,f44])).
fof(f33494,plain,(
  ( ! [X98] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X98))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X98,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X98,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl23),
  inference(superposition,[],[f4446,f18646])).
fof(f33493,plain,(
  ( ! [X97] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X97))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK4,X97,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK4,X97,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl23),
  inference(superposition,[],[f4443,f18646])).
fof(f33492,plain,(
  ( ! [X96] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X96))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK4,X96,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK4,X96,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl23),
  inference(superposition,[],[f4440,f18646])).
fof(f33489,plain,(
  ( ! [X95] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X95))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X95,sK4,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X95,sK4,sK1)),s(fun(sK1,bool),X95))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl23),
  inference(superposition,[],[f132,f18646])).
fof(f33488,plain,(
  ( ! [X94] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X94))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X94,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X94,sK1)),s(fun(sK1,bool),X94))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl23),
  inference(superposition,[],[f131,f18646])).
fof(f30974,plain,(
  s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | $spl298),
  inference(cnf_transformation,[],[f30974_D])).
fof(f30974_D,plain,(
  s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) <=> ~$spl298),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl298])])).
fof(f32884,plain,(
  ( ! [X162] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X162))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X162,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X162,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl23),
  inference(superposition,[],[f4488,f18600])).
fof(f18600,plain,(
  ( ! [X8] : (s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),X8),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X8,sK3,sK1)),s(fun(sK1,bool),X8))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X8,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl23),
  inference(trivial_inequality_removal,[],[f18599])).
fof(f18599,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),X8),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X8,sK3,sK1)),s(fun(sK1,bool),sK3))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X8,sK3,sK1)),s(fun(sK1,bool),X8)))) ) | $spl23),
  inference(superposition,[],[f17594,f44])).
fof(f17594,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(X5,sK3,sK1)),s(fun(sK1,bool),X5))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),X5),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X5,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl23),
  inference(superposition,[],[f11511,f44])).
fof(f11511,plain,(
  ( ! [X1] : (s(bool,'const.bool.IN_2'(s(sK1,sK5(X1,sK3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X1,sK3,sK1)),s(fun(sK1,bool),X1))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),X1),s(fun(sK1,bool),sK4)),sK3,bool,sK1))))) ) | $spl23),
  inference(superposition,[],[f11056,f41])).
fof(f32883,plain,(
  ( ! [X161] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X161))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X161,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X161,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl23),
  inference(superposition,[],[f4485,f18600])).
fof(f32882,plain,(
  ( ! [X160] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X160))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK3,X160,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK3,X160,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl23),
  inference(superposition,[],[f4482,f18600])).
fof(f32881,plain,(
  ( ! [X159] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X159))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK3,X159,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK3,X159,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl23),
  inference(superposition,[],[f4479,f18600])).
fof(f32868,plain,(
  ( ! [X148] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X148))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X148,sK3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X148,sK3,sK1)),s(fun(sK1,bool),X148))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl23),
  inference(superposition,[],[f184,f18600])).
fof(f32867,plain,(
  ( ! [X147] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X147))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X147,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X147,sK1)),s(fun(sK1,bool),X147))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl23),
  inference(superposition,[],[f183,f18600])).
fof(f32137,plain,(
  ( ! [X161] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X161))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X161,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X161,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl23),
  inference(superposition,[],[f4488,f18590])).
fof(f18590,plain,(
  ( ! [X7] : (s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),X7),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X7,sK3,sK1)),s(fun(sK1,bool),X7))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X7,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl23),
  inference(trivial_inequality_removal,[],[f18588])).
fof(f18588,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),X7),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X7,sK3,sK1)),s(fun(sK1,bool),sK3))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X7,sK3,sK1)),s(fun(sK1,bool),X7)))) ) | $spl23),
  inference(superposition,[],[f17593,f44])).
fof(f17593,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(X4,sK3,sK1)),s(fun(sK1,bool),X4))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),X4),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X4,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl23),
  inference(superposition,[],[f11511,f44])).
fof(f32136,plain,(
  ( ! [X160] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X160))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X160,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X160,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl23),
  inference(superposition,[],[f4485,f18590])).
fof(f32135,plain,(
  ( ! [X159] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X159))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK3,X159,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK3,X159,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl23),
  inference(superposition,[],[f4482,f18590])).
fof(f32134,plain,(
  ( ! [X158] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X158))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK3,X158,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK3,X158,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl23),
  inference(superposition,[],[f4479,f18590])).
fof(f32121,plain,(
  ( ! [X147] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X147))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X147,sK3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X147,sK3,sK1)),s(fun(sK1,bool),X147))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl23),
  inference(superposition,[],[f184,f18590])).
fof(f32120,plain,(
  ( ! [X146] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X146))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X146,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X146,sK1)),s(fun(sK1,bool),X146))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl23),
  inference(superposition,[],[f183,f18590])).
fof(f31343,plain,(
  ( ! [X146] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X146))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X146,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X146,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl23),
  inference(superposition,[],[f4521,f18543])).
fof(f18543,plain,(
  ( ! [X8] : (s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),X8),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X8,sK1)),s(fun(sK1,bool),X8))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X8,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl23),
  inference(trivial_inequality_removal,[],[f18542])).
fof(f18542,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),X8),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X8,sK1)),s(fun(sK1,bool),sK3))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X8,sK1)),s(fun(sK1,bool),X8)))) ) | $spl23),
  inference(superposition,[],[f17576,f44])).
fof(f17576,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X5,sK1)),s(fun(sK1,bool),X5))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),X5),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X5,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl23),
  inference(superposition,[],[f11510,f44])).
fof(f11510,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X0,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X0,sK1)),s(fun(sK1,bool),X0))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),X0),s(fun(sK1,bool),sK4)),sK3,bool,sK1))))) ) | $spl23),
  inference(superposition,[],[f11056,f41])).
fof(f31342,plain,(
  ( ! [X145] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X145))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X145,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X145,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl23),
  inference(superposition,[],[f4518,f18543])).
fof(f31341,plain,(
  ( ! [X144] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X144))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK3,X144,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK3,X144,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl23),
  inference(superposition,[],[f4515,f18543])).
fof(f31340,plain,(
  ( ! [X143] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X143))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK3,X143,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK3,X143,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl23),
  inference(superposition,[],[f4512,f18543])).
fof(f31331,plain,(
  ( ! [X136] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X136))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X136,sK3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X136,sK3,sK1)),s(fun(sK1,bool),X136))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl23),
  inference(superposition,[],[f312,f18543])).
fof(f31330,plain,(
  ( ! [X135] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X135))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X135,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X135,sK1)),s(fun(sK1,bool),X135))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl23),
  inference(superposition,[],[f311,f18543])).
fof(f30598,plain,(
  ( ! [X145] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X145))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X145,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X145,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl23),
  inference(superposition,[],[f4521,f18533])).
fof(f18533,plain,(
  ( ! [X7] : (s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),X7),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X7,sK1)),s(fun(sK1,bool),X7))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X7,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl23),
  inference(trivial_inequality_removal,[],[f18531])).
fof(f18531,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),X7),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X7,sK1)),s(fun(sK1,bool),sK3))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X7,sK1)),s(fun(sK1,bool),X7)))) ) | $spl23),
  inference(superposition,[],[f17575,f44])).
fof(f17575,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X4,sK1)),s(fun(sK1,bool),X4))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),X4),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X4,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl23),
  inference(superposition,[],[f11510,f44])).
fof(f30597,plain,(
  ( ! [X144] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X144))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X144,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X144,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl23),
  inference(superposition,[],[f4518,f18533])).
fof(f30596,plain,(
  ( ! [X143] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X143))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK3,X143,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK3,X143,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl23),
  inference(superposition,[],[f4515,f18533])).
fof(f30595,plain,(
  ( ! [X142] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X142))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK3,X142,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK3,X142,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl23),
  inference(superposition,[],[f4512,f18533])).
fof(f30586,plain,(
  ( ! [X135] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X135))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X135,sK3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X135,sK3,sK1)),s(fun(sK1,bool),X135))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl23),
  inference(superposition,[],[f312,f18533])).
fof(f30585,plain,(
  ( ! [X134] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X134))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X134,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X134,sK1)),s(fun(sK1,bool),X134))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl23),
  inference(superposition,[],[f311,f18533])).
fof(f28250,plain,(
  s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK4,bool,sK1)))) | $spl296),
  inference(cnf_transformation,[],[f28250_D])).
fof(f28250_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK4,bool,sK1)))) <=> ~$spl296),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl296])])).
fof(f29844,plain,(
  ( ! [X114] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X114))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X114,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X114,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl34),
  inference(superposition,[],[f4425,f18486])).
fof(f18486,plain,(
  ( ! [X8] : (s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X8)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X8,sK4,sK1)),s(fun(sK1,bool),X8))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X8,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl34),
  inference(trivial_inequality_removal,[],[f18485])).
fof(f18485,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X8)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X8,sK4,sK1)),s(fun(sK1,bool),sK4))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X8,sK4,sK1)),s(fun(sK1,bool),X8)))) ) | $spl34),
  inference(superposition,[],[f17556,f44])).
fof(f17556,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(X5,sK4,sK1)),s(fun(sK1,bool),X5))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X5)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X5,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl34),
  inference(superposition,[],[f11395,f44])).
fof(f11395,plain,(
  ( ! [X1] : (s(bool,'const.bool.IN_2'(s(sK1,sK5(X1,sK4,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X1,sK4,sK1)),s(fun(sK1,bool),X1))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X1)),sK4,bool,sK1))))) ) | $spl34),
  inference(superposition,[],[f10814,f41])).
fof(f29843,plain,(
  ( ! [X113] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X113))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X113,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X113,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl34),
  inference(superposition,[],[f4413,f18486])).
fof(f29842,plain,(
  ( ! [X112] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X112))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK4,X112,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK4,X112,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl34),
  inference(superposition,[],[f4410,f18486])).
fof(f29841,plain,(
  ( ! [X111] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X111))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK4,X111,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK4,X111,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl34),
  inference(superposition,[],[f4407,f18486])).
fof(f29838,plain,(
  ( ! [X110] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X110))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X110,sK4,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X110,sK4,sK1)),s(fun(sK1,bool),X110))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl34),
  inference(superposition,[],[f95,f18486])).
fof(f29837,plain,(
  ( ! [X109] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X109))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X109,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X109,sK1)),s(fun(sK1,bool),X109))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl34),
  inference(superposition,[],[f94,f18486])).
fof(f29179,plain,(
  ( ! [X113] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X113))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X113,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X113,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl34),
  inference(superposition,[],[f4425,f18476])).
fof(f18476,plain,(
  ( ! [X7] : (s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X7)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X7,sK4,sK1)),s(fun(sK1,bool),X7))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X7,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl34),
  inference(trivial_inequality_removal,[],[f18474])).
fof(f18474,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X7)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X7,sK4,sK1)),s(fun(sK1,bool),sK4))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X7,sK4,sK1)),s(fun(sK1,bool),X7)))) ) | $spl34),
  inference(superposition,[],[f17555,f44])).
fof(f17555,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(X4,sK4,sK1)),s(fun(sK1,bool),X4))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X4)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X4,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl34),
  inference(superposition,[],[f11395,f44])).
fof(f29178,plain,(
  ( ! [X112] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X112))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X112,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X112,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl34),
  inference(superposition,[],[f4413,f18476])).
fof(f29177,plain,(
  ( ! [X111] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X111))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK4,X111,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK4,X111,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl34),
  inference(superposition,[],[f4410,f18476])).
fof(f29176,plain,(
  ( ! [X110] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X110))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK4,X110,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK4,X110,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl34),
  inference(superposition,[],[f4407,f18476])).
fof(f29173,plain,(
  ( ! [X109] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X109))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X109,sK4,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X109,sK4,sK1)),s(fun(sK1,bool),X109))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl34),
  inference(superposition,[],[f95,f18476])).
fof(f29172,plain,(
  ( ! [X108] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X108))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X108,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X108,sK1)),s(fun(sK1,bool),X108))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl34),
  inference(superposition,[],[f94,f18476])).
fof(f28493,plain,(
  ( ! [X98] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X98))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X98,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X98,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl34),
  inference(superposition,[],[f4449,f18448])).
fof(f18448,plain,(
  ( ! [X8] : (s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X8)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X8,sK1)),s(fun(sK1,bool),X8))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X8,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl34),
  inference(trivial_inequality_removal,[],[f18447])).
fof(f18447,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X8)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X8,sK1)),s(fun(sK1,bool),sK4))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X8,sK1)),s(fun(sK1,bool),X8)))) ) | $spl34),
  inference(superposition,[],[f17538,f44])).
fof(f17538,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X5,sK1)),s(fun(sK1,bool),X5))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X5)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X5,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl34),
  inference(superposition,[],[f11394,f44])).
fof(f11394,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X0,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X0,sK1)),s(fun(sK1,bool),X0))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X0)),sK4,bool,sK1))))) ) | $spl34),
  inference(superposition,[],[f10814,f41])).
fof(f28492,plain,(
  ( ! [X97] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X97))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X97,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X97,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl34),
  inference(superposition,[],[f4446,f18448])).
fof(f28491,plain,(
  ( ! [X96] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X96))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK4,X96,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK4,X96,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl34),
  inference(superposition,[],[f4443,f18448])).
fof(f28490,plain,(
  ( ! [X95] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X95))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK4,X95,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK4,X95,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl34),
  inference(superposition,[],[f4440,f18448])).
fof(f28487,plain,(
  ( ! [X94] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X94))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X94,sK4,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X94,sK4,sK1)),s(fun(sK1,bool),X94))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl34),
  inference(superposition,[],[f132,f18448])).
fof(f28486,plain,(
  ( ! [X93] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X93))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X93,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X93,sK1)),s(fun(sK1,bool),X93))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl34),
  inference(superposition,[],[f131,f18448])).
fof(f27839,plain,(
  ( ! [X97] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X97))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X97,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X97,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl34),
  inference(superposition,[],[f4449,f18396])).
fof(f18396,plain,(
  ( ! [X7] : (s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X7)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X7,sK1)),s(fun(sK1,bool),X7))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X7,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl34),
  inference(trivial_inequality_removal,[],[f18394])).
fof(f18394,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X7)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X7,sK1)),s(fun(sK1,bool),sK4))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X7,sK1)),s(fun(sK1,bool),X7)))) ) | $spl34),
  inference(superposition,[],[f17537,f44])).
fof(f17537,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X4,sK1)),s(fun(sK1,bool),X4))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X4)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl34),
  inference(superposition,[],[f11394,f44])).
fof(f27838,plain,(
  ( ! [X96] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X96))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X96,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X96,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl34),
  inference(superposition,[],[f4446,f18396])).
fof(f27837,plain,(
  ( ! [X95] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X95))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK4,X95,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK4,X95,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl34),
  inference(superposition,[],[f4443,f18396])).
fof(f27836,plain,(
  ( ! [X94] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X94))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK4,X94,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK4,X94,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl34),
  inference(superposition,[],[f4440,f18396])).
fof(f27833,plain,(
  ( ! [X93] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X93))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X93,sK4,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X93,sK4,sK1)),s(fun(sK1,bool),X93))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl34),
  inference(superposition,[],[f132,f18396])).
fof(f27832,plain,(
  ( ! [X92] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X92))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X92,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X92,sK1)),s(fun(sK1,bool),X92))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl34),
  inference(superposition,[],[f131,f18396])).
fof(f25538,plain,(
  s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | $spl294),
  inference(cnf_transformation,[],[f25538_D])).
fof(f25538_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) <=> ~$spl294),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl294])])).
fof(f27299,plain,(
  ( ! [X160] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X160))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X160,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X160,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl34),
  inference(superposition,[],[f4488,f18368])).
fof(f18368,plain,(
  ( ! [X8] : (s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),X8),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X8,sK3,sK1)),s(fun(sK1,bool),X8))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X8,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl34),
  inference(trivial_inequality_removal,[],[f18367])).
fof(f18367,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),X8),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X8,sK3,sK1)),s(fun(sK1,bool),sK3))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X8,sK3,sK1)),s(fun(sK1,bool),X8)))) ) | $spl34),
  inference(superposition,[],[f17520,f44])).
fof(f17520,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(X5,sK3,sK1)),s(fun(sK1,bool),X5))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),X5),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X5,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl34),
  inference(superposition,[],[f11391,f44])).
fof(f11391,plain,(
  ( ! [X1] : (s(bool,'const.bool.IN_2'(s(sK1,sK5(X1,sK3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X1,sK3,sK1)),s(fun(sK1,bool),X1))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),X1),s(fun(sK1,bool),sK4)),sK4,bool,sK1))))) ) | $spl34),
  inference(superposition,[],[f10814,f41])).
fof(f27298,plain,(
  ( ! [X159] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X159))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X159,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X159,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl34),
  inference(superposition,[],[f4485,f18368])).
fof(f27297,plain,(
  ( ! [X158] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X158))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK3,X158,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK3,X158,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl34),
  inference(superposition,[],[f4482,f18368])).
fof(f27296,plain,(
  ( ! [X157] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X157))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK3,X157,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK3,X157,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl34),
  inference(superposition,[],[f4479,f18368])).
fof(f27283,plain,(
  ( ! [X146] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X146))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X146,sK3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X146,sK3,sK1)),s(fun(sK1,bool),X146))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl34),
  inference(superposition,[],[f184,f18368])).
fof(f27282,plain,(
  ( ! [X145] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X145))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X145,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X145,sK1)),s(fun(sK1,bool),X145))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl34),
  inference(superposition,[],[f183,f18368])).
fof(f26617,plain,(
  ( ! [X159] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X159))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X159,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X159,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl34),
  inference(superposition,[],[f4488,f18358])).
fof(f18358,plain,(
  ( ! [X7] : (s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),X7),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X7,sK3,sK1)),s(fun(sK1,bool),X7))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X7,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl34),
  inference(trivial_inequality_removal,[],[f18356])).
fof(f18356,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),X7),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X7,sK3,sK1)),s(fun(sK1,bool),sK3))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X7,sK3,sK1)),s(fun(sK1,bool),X7)))) ) | $spl34),
  inference(superposition,[],[f17519,f44])).
fof(f17519,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(X4,sK3,sK1)),s(fun(sK1,bool),X4))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),X4),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X4,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl34),
  inference(superposition,[],[f11391,f44])).
fof(f26616,plain,(
  ( ! [X158] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X158))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X158,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X158,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl34),
  inference(superposition,[],[f4485,f18358])).
fof(f26615,plain,(
  ( ! [X157] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X157))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK3,X157,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK3,X157,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl34),
  inference(superposition,[],[f4482,f18358])).
fof(f26614,plain,(
  ( ! [X156] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X156))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK3,X156,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK3,X156,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl34),
  inference(superposition,[],[f4479,f18358])).
fof(f26601,plain,(
  ( ! [X145] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X145))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X145,sK3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X145,sK3,sK1)),s(fun(sK1,bool),X145))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl34),
  inference(superposition,[],[f184,f18358])).
fof(f26600,plain,(
  ( ! [X144] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X144))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X144,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X144,sK1)),s(fun(sK1,bool),X144))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl34),
  inference(superposition,[],[f183,f18358])).
fof(f25895,plain,(
  ( ! [X144] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X144))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X144,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X144,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl34),
  inference(superposition,[],[f4521,f18288])).
fof(f18288,plain,(
  ( ! [X8] : (s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),X8),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X8,sK1)),s(fun(sK1,bool),X8))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X8,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl34),
  inference(trivial_inequality_removal,[],[f18287])).
fof(f18287,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),X8),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X8,sK1)),s(fun(sK1,bool),sK3))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X8,sK1)),s(fun(sK1,bool),X8)))) ) | $spl34),
  inference(superposition,[],[f17502,f44])).
fof(f17502,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X5,sK1)),s(fun(sK1,bool),X5))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),X5),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X5,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl34),
  inference(superposition,[],[f11390,f44])).
fof(f11390,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X0,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X0,sK1)),s(fun(sK1,bool),X0))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),X0),s(fun(sK1,bool),sK4)),sK4,bool,sK1))))) ) | $spl34),
  inference(superposition,[],[f10814,f41])).
fof(f25894,plain,(
  ( ! [X143] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X143))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X143,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X143,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl34),
  inference(superposition,[],[f4518,f18288])).
fof(f25893,plain,(
  ( ! [X142] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X142))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK3,X142,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK3,X142,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl34),
  inference(superposition,[],[f4515,f18288])).
fof(f25892,plain,(
  ( ! [X141] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X141))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK3,X141,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK3,X141,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl34),
  inference(superposition,[],[f4512,f18288])).
fof(f25883,plain,(
  ( ! [X134] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X134))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X134,sK3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X134,sK3,sK1)),s(fun(sK1,bool),X134))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl34),
  inference(superposition,[],[f312,f18288])).
fof(f25882,plain,(
  ( ! [X133] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X133))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X133,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X133,sK1)),s(fun(sK1,bool),X133))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl34),
  inference(superposition,[],[f311,f18288])).
fof(f25215,plain,(
  ( ! [X143] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X143))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X143,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X143,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl34),
  inference(superposition,[],[f4521,f18278])).
fof(f18278,plain,(
  ( ! [X7] : (s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),X7),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X7,sK1)),s(fun(sK1,bool),X7))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X7,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl34),
  inference(trivial_inequality_removal,[],[f18276])).
fof(f18276,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),X7),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X7,sK1)),s(fun(sK1,bool),sK3))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X7,sK1)),s(fun(sK1,bool),X7)))) ) | $spl34),
  inference(superposition,[],[f17501,f44])).
fof(f17501,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X4,sK1)),s(fun(sK1,bool),X4))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),X4),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X4,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl34),
  inference(superposition,[],[f11390,f44])).
fof(f25214,plain,(
  ( ! [X142] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X142))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X142,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X142,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl34),
  inference(superposition,[],[f4518,f18278])).
fof(f25213,plain,(
  ( ! [X141] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X141))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK3,X141,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK3,X141,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl34),
  inference(superposition,[],[f4515,f18278])).
fof(f25212,plain,(
  ( ! [X140] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X140))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK3,X140,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK3,X140,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl34),
  inference(superposition,[],[f4512,f18278])).
fof(f25203,plain,(
  ( ! [X133] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X133))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X133,sK3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X133,sK3,sK1)),s(fun(sK1,bool),X133))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl34),
  inference(superposition,[],[f312,f18278])).
fof(f25202,plain,(
  ( ! [X132] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X132))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X132,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X132,sK1)),s(fun(sK1,bool),X132))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl34),
  inference(superposition,[],[f311,f18278])).
fof(f22611,plain,(
  s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK3,bool,sK1)))) | $spl238),
  inference(cnf_transformation,[],[f22611_D])).
fof(f22611_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK3,bool,sK1)))) <=> ~$spl238),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl238])])).
fof(f24222,plain,(
  ( ! [X112] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X112))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X112,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X112,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl24),
  inference(superposition,[],[f4425,f18070])).
fof(f18070,plain,(
  ( ! [X8] : (s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X8)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X8,sK4,sK1)),s(fun(sK1,bool),X8))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X8,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl24),
  inference(trivial_inequality_removal,[],[f18069])).
fof(f18069,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X8)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X8,sK4,sK1)),s(fun(sK1,bool),sK4))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X8,sK4,sK1)),s(fun(sK1,bool),X8)))) ) | $spl24),
  inference(superposition,[],[f17408,f44])).
fof(f17408,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(X5,sK4,sK1)),s(fun(sK1,bool),X5))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X5)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X5,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl24),
  inference(superposition,[],[f11062,f44])).
fof(f11062,plain,(
  ( ! [X1] : (s(bool,'const.bool.IN_2'(s(sK1,sK5(X1,sK4,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X1,sK4,sK1)),s(fun(sK1,bool),X1))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X1)),sK3,bool,sK1))))) ) | $spl24),
  inference(superposition,[],[f10804,f41])).
fof(f24221,plain,(
  ( ! [X111] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X111))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X111,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X111,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl24),
  inference(superposition,[],[f4413,f18070])).
fof(f24220,plain,(
  ( ! [X110] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X110))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK4,X110,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK4,X110,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl24),
  inference(superposition,[],[f4410,f18070])).
fof(f24219,plain,(
  ( ! [X109] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X109))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK4,X109,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK4,X109,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl24),
  inference(superposition,[],[f4407,f18070])).
fof(f24216,plain,(
  ( ! [X108] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X108))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X108,sK4,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X108,sK4,sK1)),s(fun(sK1,bool),X108))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl24),
  inference(superposition,[],[f95,f18070])).
fof(f24215,plain,(
  ( ! [X107] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X107))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X107,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X107,sK1)),s(fun(sK1,bool),X107))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl24),
  inference(superposition,[],[f94,f18070])).
fof(f23562,plain,(
  ( ! [X111] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X111))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X111,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X111,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl24),
  inference(superposition,[],[f4425,f18060])).
fof(f18060,plain,(
  ( ! [X7] : (s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X7)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X7,sK4,sK1)),s(fun(sK1,bool),X7))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X7,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl24),
  inference(trivial_inequality_removal,[],[f18058])).
fof(f18058,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X7)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X7,sK4,sK1)),s(fun(sK1,bool),sK4))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X7,sK4,sK1)),s(fun(sK1,bool),X7)))) ) | $spl24),
  inference(superposition,[],[f17407,f44])).
fof(f17407,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(X4,sK4,sK1)),s(fun(sK1,bool),X4))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X4)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X4,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl24),
  inference(superposition,[],[f11062,f44])).
fof(f23561,plain,(
  ( ! [X110] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X110))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X110,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X110,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl24),
  inference(superposition,[],[f4413,f18060])).
fof(f23560,plain,(
  ( ! [X109] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X109))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK4,X109,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK4,X109,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl24),
  inference(superposition,[],[f4410,f18060])).
fof(f23559,plain,(
  ( ! [X108] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X108))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK4,X108,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK4,X108,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl24),
  inference(superposition,[],[f4407,f18060])).
fof(f23556,plain,(
  ( ! [X107] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X107))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X107,sK4,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X107,sK4,sK1)),s(fun(sK1,bool),X107))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl24),
  inference(superposition,[],[f95,f18060])).
fof(f23555,plain,(
  ( ! [X106] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X106))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X106,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X106,sK1)),s(fun(sK1,bool),X106))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl24),
  inference(superposition,[],[f94,f18060])).
fof(f22857,plain,(
  ( ! [X96] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X96))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X96,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X96,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl24),
  inference(superposition,[],[f4449,f18032])).
fof(f18032,plain,(
  ( ! [X8] : (s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X8)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X8,sK1)),s(fun(sK1,bool),X8))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X8,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl24),
  inference(trivial_inequality_removal,[],[f18031])).
fof(f18031,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X8)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X8,sK1)),s(fun(sK1,bool),sK4))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X8,sK1)),s(fun(sK1,bool),X8)))) ) | $spl24),
  inference(superposition,[],[f17390,f44])).
fof(f17390,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X5,sK1)),s(fun(sK1,bool),X5))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X5)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X5,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl24),
  inference(superposition,[],[f11061,f44])).
fof(f11061,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X0,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X0,sK1)),s(fun(sK1,bool),X0))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X0)),sK3,bool,sK1))))) ) | $spl24),
  inference(superposition,[],[f10804,f41])).
fof(f22856,plain,(
  ( ! [X95] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X95))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X95,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X95,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl24),
  inference(superposition,[],[f4446,f18032])).
fof(f22855,plain,(
  ( ! [X94] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X94))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK4,X94,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK4,X94,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl24),
  inference(superposition,[],[f4443,f18032])).
fof(f22854,plain,(
  ( ! [X93] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X93))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK4,X93,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK4,X93,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl24),
  inference(superposition,[],[f4440,f18032])).
fof(f22851,plain,(
  ( ! [X92] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X92))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X92,sK4,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X92,sK4,sK1)),s(fun(sK1,bool),X92))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl24),
  inference(superposition,[],[f132,f18032])).
fof(f22850,plain,(
  ( ! [X91] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X91))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X91,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X91,sK1)),s(fun(sK1,bool),X91))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl24),
  inference(superposition,[],[f131,f18032])).
fof(f22199,plain,(
  ( ! [X95] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X95))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X95,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X95,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl24),
  inference(superposition,[],[f4449,f18022])).
fof(f18022,plain,(
  ( ! [X7] : (s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X7)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X7,sK1)),s(fun(sK1,bool),X7))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X7,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl24),
  inference(trivial_inequality_removal,[],[f18020])).
fof(f18020,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X7)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X7,sK1)),s(fun(sK1,bool),sK4))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X7,sK1)),s(fun(sK1,bool),X7)))) ) | $spl24),
  inference(superposition,[],[f17389,f44])).
fof(f17389,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X4,sK1)),s(fun(sK1,bool),X4))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X4)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X4,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl24),
  inference(superposition,[],[f11061,f44])).
fof(f22198,plain,(
  ( ! [X94] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X94))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X94,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X94,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl24),
  inference(superposition,[],[f4446,f18022])).
fof(f22197,plain,(
  ( ! [X93] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X93))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK4,X93,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK4,X93,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl24),
  inference(superposition,[],[f4443,f18022])).
fof(f22196,plain,(
  ( ! [X92] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X92))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK4,X92,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK4,X92,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl24),
  inference(superposition,[],[f4440,f18022])).
fof(f22193,plain,(
  ( ! [X91] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X91))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X91,sK4,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X91,sK4,sK1)),s(fun(sK1,bool),X91))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl24),
  inference(superposition,[],[f132,f18022])).
fof(f22192,plain,(
  ( ! [X90] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X90))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X90,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X90,sK1)),s(fun(sK1,bool),X90))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK4)))) ) | $spl24),
  inference(superposition,[],[f131,f18022])).
fof(f19882,plain,(
  s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | $spl180),
  inference(cnf_transformation,[],[f19882_D])).
fof(f19882_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f21663,plain,(
  ( ! [X158] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X158))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X158,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X158,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl24),
  inference(superposition,[],[f4488,f17994])).
fof(f17994,plain,(
  ( ! [X8] : (s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),X8),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X8,sK3,sK1)),s(fun(sK1,bool),X8))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X8,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl24),
  inference(trivial_inequality_removal,[],[f17993])).
fof(f17993,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),X8),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X8,sK3,sK1)),s(fun(sK1,bool),sK3))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X8,sK3,sK1)),s(fun(sK1,bool),X8)))) ) | $spl24),
  inference(superposition,[],[f17372,f44])).
fof(f17372,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(X5,sK3,sK1)),s(fun(sK1,bool),X5))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),X5),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X5,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl24),
  inference(superposition,[],[f11058,f44])).
fof(f11058,plain,(
  ( ! [X1] : (s(bool,'const.bool.IN_2'(s(sK1,sK5(X1,sK3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X1,sK3,sK1)),s(fun(sK1,bool),X1))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),X1),s(fun(sK1,bool),sK4)),sK3,bool,sK1))))) ) | $spl24),
  inference(superposition,[],[f10804,f41])).
fof(f21662,plain,(
  ( ! [X157] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X157))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X157,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X157,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl24),
  inference(superposition,[],[f4485,f17994])).
fof(f21661,plain,(
  ( ! [X156] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X156))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK3,X156,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK3,X156,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl24),
  inference(superposition,[],[f4482,f17994])).
fof(f21660,plain,(
  ( ! [X155] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X155))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK3,X155,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK3,X155,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl24),
  inference(superposition,[],[f4479,f17994])).
fof(f21647,plain,(
  ( ! [X144] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X144))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X144,sK3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X144,sK3,sK1)),s(fun(sK1,bool),X144))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl24),
  inference(superposition,[],[f184,f17994])).
fof(f21646,plain,(
  ( ! [X143] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X143))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X143,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X143,sK1)),s(fun(sK1,bool),X143))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl24),
  inference(superposition,[],[f183,f17994])).
fof(f20957,plain,(
  ( ! [X157] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X157))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X157,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X157,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl24),
  inference(superposition,[],[f4488,f17984])).
fof(f17984,plain,(
  ( ! [X7] : (s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),X7),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X7,sK3,sK1)),s(fun(sK1,bool),X7))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X7,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl24),
  inference(trivial_inequality_removal,[],[f17982])).
fof(f17982,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),X7),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X7,sK3,sK1)),s(fun(sK1,bool),sK3))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X7,sK3,sK1)),s(fun(sK1,bool),X7)))) ) | $spl24),
  inference(superposition,[],[f17371,f44])).
fof(f17371,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(X4,sK3,sK1)),s(fun(sK1,bool),X4))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),X4),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(X4,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl24),
  inference(superposition,[],[f11058,f44])).
fof(f20956,plain,(
  ( ! [X156] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X156))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X156,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X156,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl24),
  inference(superposition,[],[f4485,f17984])).
fof(f20955,plain,(
  ( ! [X155] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X155))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK3,X155,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK3,X155,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl24),
  inference(superposition,[],[f4482,f17984])).
fof(f20954,plain,(
  ( ! [X154] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X154))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK3,X154,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK3,X154,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl24),
  inference(superposition,[],[f4479,f17984])).
fof(f20941,plain,(
  ( ! [X143] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X143))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X143,sK3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X143,sK3,sK1)),s(fun(sK1,bool),X143))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl24),
  inference(superposition,[],[f184,f17984])).
fof(f20940,plain,(
  ( ! [X142] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),X142))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X142,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X142,sK1)),s(fun(sK1,bool),X142))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl24),
  inference(superposition,[],[f183,f17984])).
fof(f20240,plain,(
  ( ! [X141] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X141))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X141,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X141,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl24),
  inference(superposition,[],[f4521,f17956])).
fof(f17956,plain,(
  ( ! [X8] : (s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),X8),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X8,sK1)),s(fun(sK1,bool),X8))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X8,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl24),
  inference(trivial_inequality_removal,[],[f17955])).
fof(f17955,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),X8),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X8,sK1)),s(fun(sK1,bool),sK3))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X8,sK1)),s(fun(sK1,bool),X8)))) ) | $spl24),
  inference(superposition,[],[f17354,f44])).
fof(f17354,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X5,sK1)),s(fun(sK1,bool),X5))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),X5),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X5,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl24),
  inference(superposition,[],[f11057,f44])).
fof(f11057,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X0,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X0,sK1)),s(fun(sK1,bool),X0))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),X0),s(fun(sK1,bool),sK4)),sK3,bool,sK1))))) ) | $spl24),
  inference(superposition,[],[f10804,f41])).
fof(f20239,plain,(
  ( ! [X140] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X140))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X140,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X140,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl24),
  inference(superposition,[],[f4518,f17956])).
fof(f20238,plain,(
  ( ! [X139] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X139))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK3,X139,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK3,X139,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl24),
  inference(superposition,[],[f4515,f17956])).
fof(f20237,plain,(
  ( ! [X138] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X138))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK3,X138,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK3,X138,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl24),
  inference(superposition,[],[f4512,f17956])).
fof(f20229,plain,(
  ( ! [X132] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X132))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X132,sK3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X132,sK3,sK1)),s(fun(sK1,bool),X132))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl24),
  inference(superposition,[],[f312,f17956])).
fof(f20228,plain,(
  ( ! [X131] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X131))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X131,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X131,sK1)),s(fun(sK1,bool),X131))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl24),
  inference(superposition,[],[f311,f17956])).
fof(f19557,plain,(
  ( ! [X140] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X140))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X140,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X140,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl24),
  inference(superposition,[],[f4521,f17923])).
fof(f17923,plain,(
  ( ! [X7] : (s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),X7),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X7,sK1)),s(fun(sK1,bool),X7))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X7,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl24),
  inference(trivial_inequality_removal,[],[f17921])).
fof(f17921,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),X7),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X7,sK1)),s(fun(sK1,bool),sK3))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X7,sK1)),s(fun(sK1,bool),X7)))) ) | $spl24),
  inference(superposition,[],[f17353,f44])).
fof(f17353,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X4,sK1)),s(fun(sK1,bool),X4))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),X4),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X4,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl24),
  inference(superposition,[],[f11057,f44])).
fof(f19556,plain,(
  ( ! [X139] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X139))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X139,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X139,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl24),
  inference(superposition,[],[f4518,f17923])).
fof(f19555,plain,(
  ( ! [X138] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X138))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK3,X138,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK3,X138,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl24),
  inference(superposition,[],[f4515,f17923])).
fof(f19554,plain,(
  ( ! [X137] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X137))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK3,X137,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK3,X137,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl24),
  inference(superposition,[],[f4512,f17923])).
fof(f19546,plain,(
  ( ! [X131] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X131))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X131,sK3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X131,sK3,sK1)),s(fun(sK1,bool),X131))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl24),
  inference(superposition,[],[f312,f17923])).
fof(f19545,plain,(
  ( ! [X130] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),X130))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X130,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X130,sK1)),s(fun(sK1,bool),X130))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK3)))) ) | $spl24),
  inference(superposition,[],[f311,f17923])).
fof(f18426,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))))))),
  inference(superposition,[],[f984,f44])).
fof(f18432,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl170),
  inference(cnf_transformation,[],[f18432_D])).
fof(f18432_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f18434,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl173),
  inference(cnf_transformation,[],[f18434_D])).
fof(f18434_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) <=> ~$spl173),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl173])])).
fof(f18425,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))))))),
  inference(superposition,[],[f984,f44])).
fof(f18336,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))))))),
  inference(superposition,[],[f915,f44])).
fof(f18335,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))))))),
  inference(superposition,[],[f915,f44])).
fof(f11031,plain,(
  s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | $spl108),
  inference(cnf_transformation,[],[f11031_D])).
fof(f11031_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f10999,plain,(
  s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | $spl77),
  inference(cnf_transformation,[],[f10999_D])).
fof(f10999_D,plain,(
  s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) <=> ~$spl77),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl77])])).
fof(f11348,plain,(
  s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK4,sK4,bool,sK1)))) | $spl158),
  inference(cnf_transformation,[],[f11348_D])).
fof(f11348_D,plain,(
  s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK4,sK4,bool,sK1)))) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f11350,plain,(
  s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK4,sK4,bool,sK1)))) | $spl160),
  inference(cnf_transformation,[],[f11350_D])).
fof(f11350_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK4,sK4,bool,sK1)))) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f11267,plain,(
  s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK4,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK4,sK4,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | $spl32),
  inference(trivial_inequality_removal,[],[f11256])).
fof(f11256,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK4,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK4,sK4,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | $spl32),
  inference(superposition,[],[f1612,f10812])).
fof(f1612,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,f(s(sK1,sK6(X2,X3,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X3,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X3,sK4,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X2))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X2))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X2,X3,bool,sK1))))) )),
  inference(superposition,[],[f356,f44])).
fof(f356,plain,(
  ( ! [X6,X7] : (s(bool,f(s(sK1,sK6(X7,X6,bool,sK1)))) != s(bool,g(s(sK1,sK6(X7,X6,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X6,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X6,sK4,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X7))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X7)))))) )),
  inference(superposition,[],[f152,f53])).
fof(f152,plain,(
  ( ! [X1] : (s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X1))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X1))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X1,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X1,sK4,bool,sK1))))) )),
  inference(trivial_inequality_removal,[],[f151])).
fof(f151,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X1))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X1))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X1,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X1,sK4,bool,sK1))))) )),
  inference(superposition,[],[f85,f44])).
fof(f85,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,f(s(sK1,sK6(X1,sK4,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X1))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X1))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X1,sK4,bool,sK1))))) )),
  inference(superposition,[],[f69,f44])).
fof(f69,plain,(
  ( ! [X3] : (s(bool,f(s(sK1,sK6(X3,sK4,bool,sK1)))) != s(bool,g(s(sK1,sK6(X3,sK4,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X3)))))) )),
  inference(superposition,[],[f39,f53])).
fof(f11344,plain,(
  s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK4,sK4,bool,sK1)))) | $spl154),
  inference(cnf_transformation,[],[f11344_D])).
fof(f11344_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK4,sK4,bool,sK1)))) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f11346,plain,(
  s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK4,sK4,bool,sK1)))) | $spl156),
  inference(cnf_transformation,[],[f11346_D])).
fof(f11346_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK4,sK4,bool,sK1)))) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f11271,plain,(
  s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK4,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK4,sK4,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | $spl32),
  inference(trivial_inequality_removal,[],[f11250])).
fof(f11250,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK4,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK4,sK4,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | $spl32),
  inference(superposition,[],[f1602,f10812])).
fof(f1602,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,f(s(sK1,sK6(X2,X3,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X3,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X3,sK4,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X2))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X2))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X2,X3,bool,sK1))))) )),
  inference(superposition,[],[f334,f44])).
fof(f334,plain,(
  ( ! [X6,X7] : (s(bool,f(s(sK1,sK6(X7,X6,bool,sK1)))) != s(bool,g(s(sK1,sK6(X7,X6,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X6,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X6,sK4,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X7))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X7)))))) )),
  inference(superposition,[],[f149,f53])).
fof(f149,plain,(
  ( ! [X0] : (s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X0))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X0))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X0,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X0,sK4,bool,sK1))))) )),
  inference(trivial_inequality_removal,[],[f147])).
fof(f147,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X0))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X0))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X0,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X0,sK4,bool,sK1))))) )),
  inference(superposition,[],[f84,f44])).
fof(f84,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(sK1,sK6(X0,sK4,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X0))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X0))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X0,sK4,bool,sK1))))) )),
  inference(superposition,[],[f69,f44])).
fof(f11274,plain,(
  s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | $spl32),
  inference(trivial_inequality_removal,[],[f11246])).
fof(f11246,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,sK1)),s(fun(sK1,bool),sK4))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | $spl32),
  inference(superposition,[],[f1598,f10812])).
fof(f11276,plain,(
  s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK4,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK4,sK4,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | $spl32),
  inference(trivial_inequality_removal,[],[f11244])).
fof(f11244,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK4,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK4,sK4,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | $spl32),
  inference(superposition,[],[f1592,f10812])).
fof(f1592,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,f(s(sK1,sK6(X2,X3,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK4,X3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK4,X3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X2))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X2))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X2,X3,bool,sK1))))) )),
  inference(superposition,[],[f294,f44])).
fof(f294,plain,(
  ( ! [X6,X7] : (s(bool,f(s(sK1,sK6(X7,X6,bool,sK1)))) != s(bool,g(s(sK1,sK6(X7,X6,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK4,X6,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK4,X6,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X7))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X7)))))) )),
  inference(superposition,[],[f146,f53])).
fof(f146,plain,(
  ( ! [X1] : (s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X1))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X1))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK4,X1,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK4,X1,bool,sK1))))) )),
  inference(trivial_inequality_removal,[],[f145])).
fof(f145,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X1))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X1))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK4,X1,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK4,X1,bool,sK1))))) )),
  inference(superposition,[],[f83,f44])).
fof(f83,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,f(s(sK1,sK6(sK4,X1,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X1))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X1))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK4,X1,bool,sK1))))) )),
  inference(superposition,[],[f68,f44])).
fof(f68,plain,(
  ( ! [X2] : (s(bool,f(s(sK1,sK6(sK4,X2,bool,sK1)))) != s(bool,g(s(sK1,sK6(sK4,X2,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X2))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X2)))))) )),
  inference(superposition,[],[f39,f53])).
fof(f11278,plain,(
  s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK4,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK4,sK4,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | $spl32),
  inference(trivial_inequality_removal,[],[f11240])).
fof(f11240,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK4,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK4,sK4,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | $spl32),
  inference(superposition,[],[f1582,f10812])).
fof(f1582,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,f(s(sK1,sK6(X2,X3,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK4,X3,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK4,X3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X2))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X2))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X2,X3,bool,sK1))))) )),
  inference(superposition,[],[f272,f44])).
fof(f272,plain,(
  ( ! [X6,X7] : (s(bool,f(s(sK1,sK6(X7,X6,bool,sK1)))) != s(bool,g(s(sK1,sK6(X7,X6,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK4,X6,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK4,X6,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X7))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X7)))))) )),
  inference(superposition,[],[f143,f53])).
fof(f143,plain,(
  ( ! [X0] : (s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X0))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X0))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK4,X0,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK4,X0,bool,sK1))))) )),
  inference(trivial_inequality_removal,[],[f141])).
fof(f141,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X0))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X0))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK4,X0,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK4,X0,bool,sK1))))) )),
  inference(superposition,[],[f82,f44])).
fof(f82,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(sK1,sK6(sK4,X0,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X0))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X0))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK4,X0,bool,sK1))))) )),
  inference(superposition,[],[f68,f44])).
fof(f10975,plain,(
  s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | $spl47),
  inference(cnf_transformation,[],[f10975_D])).
fof(f10975_D,plain,(
  s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) <=> ~$spl47),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl47])])).
fof(f11017,plain,(
  s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK4,sK3,bool,sK1)))) | $spl94),
  inference(cnf_transformation,[],[f11017_D])).
fof(f11017_D,plain,(
  s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK4,sK3,bool,sK1)))) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f11019,plain,(
  s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK4,sK3,bool,sK1)))) | $spl96),
  inference(cnf_transformation,[],[f11019_D])).
fof(f11019_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK4,sK3,bool,sK1)))) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f11281,plain,(
  s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK4,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK4,sK3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | $spl32),
  inference(trivial_inequality_removal,[],[f11234])).
fof(f11234,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK4,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK4,sK3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | $spl32),
  inference(superposition,[],[f1566,f10812])).
fof(f1566,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,f(s(sK1,sK6(X2,X3,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X3,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X3,sK3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X2),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X2))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X2,X3,bool,sK1))))) )),
  inference(superposition,[],[f240,f44])).
fof(f240,plain,(
  ( ! [X6,X7] : (s(bool,f(s(sK1,sK6(X7,X6,bool,sK1)))) != s(bool,g(s(sK1,sK6(X7,X6,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X6,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X6,sK3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X7),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X7))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))))) )),
  inference(superposition,[],[f122,f53])).
fof(f122,plain,(
  ( ! [X1] : (s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X1),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X1))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X1,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X1,sK3,bool,sK1))))) )),
  inference(trivial_inequality_removal,[],[f121])).
fof(f121,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X1),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X1))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X1,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X1,sK3,bool,sK1))))) )),
  inference(superposition,[],[f81,f44])).
fof(f81,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,f(s(sK1,sK6(X1,sK3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X1),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X1))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X1,sK3,bool,sK1))))) )),
  inference(superposition,[],[f59,f44])).
fof(f59,plain,(
  ( ! [X3] : (s(bool,f(s(sK1,sK6(X3,sK3,bool,sK1)))) != s(bool,g(s(sK1,sK6(X3,sK3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))))) )),
  inference(superposition,[],[f39,f53])).
fof(f11282,plain,(
  s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | $spl32),
  inference(trivial_inequality_removal,[],[f11232])).
fof(f11232,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK1)),s(fun(sK1,bool),sK4))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | $spl32),
  inference(superposition,[],[f1564,f10812])).
fof(f11342,plain,(
  s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | $spl152),
  inference(cnf_transformation,[],[f11342_D])).
fof(f11342_D,plain,(
  s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f11340,plain,(
  s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | ($spl3 | $spl24 | $spl32)),
  inference(subsumption_resolution,[],[f11339,f110])).
fof(f11339,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | ($spl24 | $spl32)),
  inference(forward_demodulation,[],[f11284,f10804])).
fof(f11284,plain,(
  s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | $spl32),
  inference(trivial_inequality_removal,[],[f11230])).
fof(f11230,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | $spl32),
  inference(superposition,[],[f1562,f10812])).
fof(f1562,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,f(s(sK1,sK6(X2,X3,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X2,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X2,sK3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X2,X3,bool,sK1))))) )),
  inference(superposition,[],[f239,f44])).
fof(f239,plain,(
  ( ! [X4,X5] : (s(bool,f(s(sK1,sK6(X4,X5,bool,sK1)))) != s(bool,g(s(sK1,sK6(X4,X5,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X4,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X4,sK3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X5),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X5))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))))) )),
  inference(superposition,[],[f122,f53])).
fof(f11009,plain,(
  s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK4,sK3,bool,sK1)))) | $spl86),
  inference(cnf_transformation,[],[f11009_D])).
fof(f11009_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK4,sK3,bool,sK1)))) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f11011,plain,(
  s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK4,sK3,bool,sK1)))) | $spl88),
  inference(cnf_transformation,[],[f11011_D])).
fof(f11011_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK4,sK3,bool,sK1)))) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f11285,plain,(
  s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK4,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK4,sK3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | $spl32),
  inference(trivial_inequality_removal,[],[f11228])).
fof(f11228,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK4,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK4,sK3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | $spl32),
  inference(superposition,[],[f1554,f10812])).
fof(f1554,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,f(s(sK1,sK6(X2,X3,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X3,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X3,sK3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X2),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X2))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X2,X3,bool,sK1))))) )),
  inference(superposition,[],[f218,f44])).
fof(f218,plain,(
  ( ! [X6,X7] : (s(bool,f(s(sK1,sK6(X7,X6,bool,sK1)))) != s(bool,g(s(sK1,sK6(X7,X6,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X6,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X6,sK3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X7),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X7))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))))) )),
  inference(superposition,[],[f119,f53])).
fof(f119,plain,(
  ( ! [X0] : (s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X0),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X0))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X0,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X0,sK3,bool,sK1))))) )),
  inference(trivial_inequality_removal,[],[f117])).
fof(f117,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X0),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X0))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X0,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X0,sK3,bool,sK1))))) )),
  inference(superposition,[],[f80,f44])).
fof(f80,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(sK1,sK6(X0,sK3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X0),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X0))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X0,sK3,bool,sK1))))) )),
  inference(superposition,[],[f59,f44])).
fof(f11033,plain,(
  s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK3,sK4,bool,sK1)))) | $spl110),
  inference(cnf_transformation,[],[f11033_D])).
fof(f11033_D,plain,(
  s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK3,sK4,bool,sK1)))) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f11035,plain,(
  s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK3,sK4,bool,sK1)))) | $spl112),
  inference(cnf_transformation,[],[f11035_D])).
fof(f11035_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK3,sK4,bool,sK1)))) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f11287,plain,(
  s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK3,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK3,sK4,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | $spl32),
  inference(trivial_inequality_removal,[],[f11224])).
fof(f11224,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK3,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK3,sK4,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | $spl32),
  inference(superposition,[],[f1544,f10812])).
fof(f1544,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,f(s(sK1,sK6(X2,X3,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK3,X3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK3,X3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X2),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X2))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X2,X3,bool,sK1))))) )),
  inference(superposition,[],[f196,f44])).
fof(f196,plain,(
  ( ! [X6,X7] : (s(bool,f(s(sK1,sK6(X7,X6,bool,sK1)))) != s(bool,g(s(sK1,sK6(X7,X6,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK3,X6,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK3,X6,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X7),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X7))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))))) )),
  inference(superposition,[],[f116,f53])).
fof(f116,plain,(
  ( ! [X1] : (s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X1),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X1))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK3,X1,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK3,X1,bool,sK1))))) )),
  inference(trivial_inequality_removal,[],[f115])).
fof(f115,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X1),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X1))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK3,X1,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK3,X1,bool,sK1))))) )),
  inference(superposition,[],[f79,f44])).
fof(f79,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,f(s(sK1,sK6(sK3,X1,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X1),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X1))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK3,X1,bool,sK1))))) )),
  inference(superposition,[],[f58,f44])).
fof(f58,plain,(
  ( ! [X2] : (s(bool,f(s(sK1,sK6(sK3,X2,bool,sK1)))) != s(bool,g(s(sK1,sK6(sK3,X2,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X2),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X2))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))))) )),
  inference(superposition,[],[f39,f53])).
fof(f11336,plain,(
  s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK3,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),bool,sK1)))) | $spl148),
  inference(cnf_transformation,[],[f11336_D])).
fof(f11336_D,plain,(
  s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK3,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),bool,sK1)))) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f11338,plain,(
  s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK3,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),bool,sK1)))) | $spl150),
  inference(cnf_transformation,[],[f11338_D])).
fof(f11338_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK3,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),bool,sK1)))) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f11288,plain,(
  s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK3,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK3,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | $spl32),
  inference(trivial_inequality_removal,[],[f11222])).
fof(f11222,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK3,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK3,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | $spl32),
  inference(superposition,[],[f1542,f10812])).
fof(f1542,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,f(s(sK1,sK6(X2,X3,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK3,X2,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK3,X2,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X2,X3,bool,sK1))))) )),
  inference(superposition,[],[f195,f44])).
fof(f195,plain,(
  ( ! [X4,X5] : (s(bool,f(s(sK1,sK6(X4,X5,bool,sK1)))) != s(bool,g(s(sK1,sK6(X4,X5,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK3,X4,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK3,X4,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X5),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X5))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))))) )),
  inference(superposition,[],[f116,f53])).
fof(f11334,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl147),
  inference(cnf_transformation,[],[f11334_D])).
fof(f11334_D,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) <=> ~$spl147),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl147])])).
fof(f11289,plain,(
  s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | $spl32),
  inference(trivial_inequality_removal,[],[f11220])).
fof(f11220,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | $spl32),
  inference(superposition,[],[f1536,f10812])).
fof(f11023,plain,(
  s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK3,sK4,bool,sK1)))) | $spl100),
  inference(cnf_transformation,[],[f11023_D])).
fof(f11023_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK3,sK4,bool,sK1)))) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f11025,plain,(
  s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK3,sK4,bool,sK1)))) | $spl102),
  inference(cnf_transformation,[],[f11025_D])).
fof(f11025_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK3,sK4,bool,sK1)))) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f11290,plain,(
  s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK3,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK3,sK4,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | $spl32),
  inference(trivial_inequality_removal,[],[f11218])).
fof(f11218,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK3,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK3,sK4,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | $spl32),
  inference(superposition,[],[f1534,f10812])).
fof(f1534,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,f(s(sK1,sK6(X2,X3,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK3,X3,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK3,X3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X2),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X2))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X2,X3,bool,sK1))))) )),
  inference(superposition,[],[f156,f44])).
fof(f156,plain,(
  ( ! [X6,X7] : (s(bool,f(s(sK1,sK6(X7,X6,bool,sK1)))) != s(bool,g(s(sK1,sK6(X7,X6,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK3,X6,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK3,X6,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X7),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X7))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))))) )),
  inference(superposition,[],[f106,f53])).
fof(f106,plain,(
  ( ! [X0] : (s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X0),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X0))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK3,X0,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK3,X0,bool,sK1))))) )),
  inference(trivial_inequality_removal,[],[f104])).
fof(f104,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X0),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X0))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK3,X0,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK3,X0,bool,sK1))))) )),
  inference(superposition,[],[f78,f44])).
fof(f78,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(sK1,sK6(sK3,X0,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X0),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X0))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK3,X0,bool,sK1))))) )),
  inference(superposition,[],[f58,f44])).
fof(f11330,plain,(
  s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK3,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),bool,sK1)))) | $spl142),
  inference(cnf_transformation,[],[f11330_D])).
fof(f11330_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK3,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),bool,sK1)))) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f11332,plain,(
  s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK3,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),bool,sK1)))) | $spl144),
  inference(cnf_transformation,[],[f11332_D])).
fof(f11332_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK3,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),bool,sK1)))) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f11291,plain,(
  s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK3,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK3,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | $spl32),
  inference(trivial_inequality_removal,[],[f11216])).
fof(f11216,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK3,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK3,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | $spl32),
  inference(superposition,[],[f1532,f10812])).
fof(f1532,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,f(s(sK1,sK6(X2,X3,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK3,X2,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK3,X2,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X2,X3,bool,sK1))))) )),
  inference(superposition,[],[f155,f44])).
fof(f155,plain,(
  ( ! [X4,X5] : (s(bool,f(s(sK1,sK6(X4,X5,bool,sK1)))) != s(bool,g(s(sK1,sK6(X4,X5,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK3,X4,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK3,X4,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X5),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X5))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))))) )),
  inference(superposition,[],[f106,f53])).
fof(f11328,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,sK5('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,sK1)),s(fun(sK1,bool),sK4))) | $spl141),
  inference(cnf_transformation,[],[f11328_D])).
fof(f11328_D,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,sK5('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,sK1)),s(fun(sK1,bool),sK4))) <=> ~$spl141),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl141])])).
fof(f11292,plain,(
  s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,sK1)),s(fun(sK1,bool),sK4))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | $spl32),
  inference(trivial_inequality_removal,[],[f11214])).
fof(f11214,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,sK1)),s(fun(sK1,bool),sK4))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | $spl32),
  inference(superposition,[],[f1530,f10812])).
fof(f11324,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl137),
  inference(cnf_transformation,[],[f11324_D])).
fof(f11324_D,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) <=> ~$spl137),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl137])])).
fof(f11293,plain,(
  s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | $spl32),
  inference(trivial_inequality_removal,[],[f11212])).
fof(f11212,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | $spl32),
  inference(superposition,[],[f1452,f10812])).
fof(f11322,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK1)),s(fun(sK1,bool),sK4))) | $spl135),
  inference(cnf_transformation,[],[f11322_D])).
fof(f11322_D,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK1)),s(fun(sK1,bool),sK4))) <=> ~$spl135),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl135])])).
fof(f11294,plain,(
  s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK1)),s(fun(sK1,bool),sK4))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | $spl32),
  inference(trivial_inequality_removal,[],[f11210])).
fof(f11210,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK1)),s(fun(sK1,bool),sK4))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | $spl32),
  inference(superposition,[],[f1432,f10812])).
fof(f11320,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))))),bool,sK0)))) | $spl133),
  inference(cnf_transformation,[],[f11320_D])).
fof(f11320_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))))),bool,sK0)))) <=> ~$spl133),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl133])])).
fof(f11295,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))))),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | ($spl4 | $spl32)),
  inference(trivial_inequality_removal,[],[f11208])).
fof(f11208,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))))),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | ($spl4 | $spl32)),
  inference(superposition,[],[f6125,f10812])).
fof(f6125,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,f(s(sK1,sK6(X1,sK4,bool,sK1)))) | s(bool,'T_0') != s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X1)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X1)))),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X1,sK4,bool,sK1))))) ) | $spl4),
  inference(superposition,[],[f6069,f44])).
fof(f6069,plain,(
  ( ! [X3] : (s(bool,f(s(sK1,sK6(X3,sK4,bool,sK1)))) != s(bool,g(s(sK1,sK6(X3,sK4,bool,sK1)))) | s(bool,'T_0') != s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X3)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X3)))),bool,sK0))))) ) | $spl4),
  inference(superposition,[],[f6014,f53])).
fof(f6014,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl4),
  inference(trivial_inequality_removal,[],[f5990])).
fof(f5990,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl4),
  inference(superposition,[],[f76,f5950])).
fof(f5950,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl4),
  inference(trivial_inequality_removal,[],[f5937])).
fof(f5937,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl4),
  inference(superposition,[],[f778,f5911])).
fof(f778,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),X1,bool,sK0)))) | s(fun(sK0,bool),X1) != s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),X1,bool,sK0))))) )),
  inference(superposition,[],[f76,f44])).
fof(f76,plain,(
  ( ! [X2] : (s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),X2,bool,sK0)))) != s(bool,g(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),X2,bool,sK0)))) | s(fun(sK0,bool),X2) != s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))))) )),
  inference(superposition,[],[f39,f53])).
fof(f11318,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))))),bool,sK0)))) | $spl130),
  inference(cnf_transformation,[],[f11318_D])).
fof(f11318_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))))),bool,sK0)))) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f11296,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))))),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | ($spl4 | $spl32)),
  inference(trivial_inequality_removal,[],[f11206])).
fof(f11206,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))))),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | ($spl4 | $spl32)),
  inference(superposition,[],[f6109,f10812])).
fof(f6109,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,f(s(sK1,sK6(X1,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X1)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X1)))),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X1,sK4,bool,sK1))))) ) | $spl4),
  inference(superposition,[],[f5984,f44])).
fof(f5984,plain,(
  ( ! [X3] : (s(bool,f(s(sK1,sK6(X3,sK4,bool,sK1)))) != s(bool,g(s(sK1,sK6(X3,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X3)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X3)))),bool,sK0))))) ) | $spl4),
  inference(superposition,[],[f5950,f53])).
fof(f11316,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))))),bool,sK0)))) | $spl128),
  inference(cnf_transformation,[],[f11316_D])).
fof(f11316_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))))),bool,sK0)))) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f11297,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))))),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | ($spl4 | $spl32)),
  inference(trivial_inequality_removal,[],[f11204])).
fof(f11204,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))))),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | ($spl4 | $spl32)),
  inference(superposition,[],[f6091,f10812])).
fof(f6091,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,f(s(sK1,sK6(X1,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X1)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X1)))),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X1,sK4,bool,sK1))))) ) | $spl4),
  inference(superposition,[],[f5931,f44])).
fof(f5931,plain,(
  ( ! [X3] : (s(bool,f(s(sK1,sK6(X3,sK4,bool,sK1)))) != s(bool,g(s(sK1,sK6(X3,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X3)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X3)))),bool,sK0))))) ) | $spl4),
  inference(superposition,[],[f5911,f53])).
fof(f11314,plain,(
  ( ! [X11] : (s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X11))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X11))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X11,sK4,sK1)),s(fun(sK1,bool),X11))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X11,sK4,sK1)),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))))) ) | $spl126),
  inference(cnf_transformation,[],[f11314_D])).
fof(f11314_D,plain,(
  ( ! [X11] : (s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X11))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X11))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X11,sK4,sK1)),s(fun(sK1,bool),X11))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X11,sK4,sK1)),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))))) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f11298,plain,(
  ( ! [X11] : (s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X11))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X11))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X11,sK4,sK1)),s(fun(sK1,bool),X11))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X11,sK4,sK1)),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1))))) ) | $spl32),
  inference(trivial_inequality_removal,[],[f11202])).
fof(f11202,plain,(
  ( ! [X11] : (s(bool,'F_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X11))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X11))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X11,sK4,sK1)),s(fun(sK1,bool),X11))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X11,sK4,sK1)),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1))))) ) | $spl32),
  inference(superposition,[],[f1588,f10812])).
fof(f1588,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,f(s(sK1,sK6(X2,sK4,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X3))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X3,sK4,sK1)),s(fun(sK1,bool),X2))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X3,sK4,sK1)),s(fun(sK1,bool),X3))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X2,sK4,bool,sK1))))) )),
  inference(superposition,[],[f306,f44])).
fof(f306,plain,(
  ( ! [X6,X7] : (s(bool,f(s(sK1,sK6(X6,sK4,bool,sK1)))) != s(bool,g(s(sK1,sK6(X6,sK4,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X7))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X7))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X7,sK4,sK1)),s(fun(sK1,bool),X6))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X7,sK4,sK1)),s(fun(sK1,bool),X7)))) )),
  inference(superposition,[],[f67,f53])).
fof(f11313,plain,(
  ( ! [X9] : (s(bool,'T_0') = s(bool,g(s(sK1,sK6(X9,sK3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X9),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X9))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X9,sK3,bool,sK1))))) ) | $spl124),
  inference(cnf_transformation,[],[f11313_D])).
fof(f11313_D,plain,(
  ( ! [X9] : (s(bool,'T_0') = s(bool,g(s(sK1,sK6(X9,sK3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X9),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X9))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X9,sK3,bool,sK1))))) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f11299,plain,(
  ( ! [X9] : (s(bool,'T_0') = s(bool,g(s(sK1,sK6(X9,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X9,sK3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X9),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X9))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1))))) ) | $spl32),
  inference(trivial_inequality_removal,[],[f11200])).
fof(f11200,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X9,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X9,sK3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X9),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X9))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1))))) ) | $spl32),
  inference(superposition,[],[f1570,f10812])).
fof(f1570,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,f(s(sK1,sK6(X2,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X3,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X3,sK3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X3),s(fun(sK1,bool),X2))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X2))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X2,sK4,bool,sK1))))) )),
  inference(superposition,[],[f250,f44])).
fof(f250,plain,(
  ( ! [X6,X7] : (s(bool,f(s(sK1,sK6(X6,sK4,bool,sK1)))) != s(bool,g(s(sK1,sK6(X6,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X7,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X7,sK3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X7),s(fun(sK1,bool),X6))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X7))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X6)))))) )),
  inference(superposition,[],[f122,f53])).
fof(f11312,plain,(
  ( ! [X7] : (s(bool,'F_0') = s(bool,g(s(sK1,sK6(X7,sK3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X7),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X7))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X7,sK3,bool,sK1))))) ) | $spl122),
  inference(cnf_transformation,[],[f11312_D])).
fof(f11312_D,plain,(
  ( ! [X7] : (s(bool,'F_0') = s(bool,g(s(sK1,sK6(X7,sK3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X7),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X7))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X7,sK3,bool,sK1))))) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f11300,plain,(
  ( ! [X7] : (s(bool,'F_0') = s(bool,g(s(sK1,sK6(X7,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X7,sK3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X7),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X7))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1))))) ) | $spl32),
  inference(trivial_inequality_removal,[],[f11198])).
fof(f11198,plain,(
  ( ! [X7] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X7,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X7,sK3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X7),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X7))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1))))) ) | $spl32),
  inference(superposition,[],[f1560,f10812])).
fof(f1560,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,f(s(sK1,sK6(X2,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X3,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X3,sK3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X3),s(fun(sK1,bool),X2))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X2))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X2,sK4,bool,sK1))))) )),
  inference(superposition,[],[f228,f44])).
fof(f228,plain,(
  ( ! [X6,X7] : (s(bool,f(s(sK1,sK6(X6,sK4,bool,sK1)))) != s(bool,g(s(sK1,sK6(X6,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X7,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X7,sK3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X7),s(fun(sK1,bool),X6))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X7))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X6)))))) )),
  inference(superposition,[],[f119,f53])).
fof(f11311,plain,(
  ( ! [X5] : (s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X5))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X5))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X5,sK1)),s(fun(sK1,bool),X5))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X5,sK1)),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))))) ) | $spl120),
  inference(cnf_transformation,[],[f11311_D])).
fof(f11311_D,plain,(
  ( ! [X5] : (s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X5))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X5))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X5,sK1)),s(fun(sK1,bool),X5))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X5,sK1)),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))))) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f11301,plain,(
  ( ! [X5] : (s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X5))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X5))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X5,sK1)),s(fun(sK1,bool),X5))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X5,sK1)),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1))))) ) | $spl32),
  inference(trivial_inequality_removal,[],[f11196])).
fof(f11196,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X5))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X5))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X5,sK1)),s(fun(sK1,bool),X5))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X5,sK1)),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1))))) ) | $spl32),
  inference(superposition,[],[f1556,f10812])).
fof(f1556,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,f(s(sK1,sK6(X2,sK4,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X3))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X3,sK1)),s(fun(sK1,bool),X2))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X3,sK1)),s(fun(sK1,bool),X3))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X2,sK4,bool,sK1))))) )),
  inference(superposition,[],[f178,f44])).
fof(f178,plain,(
  ( ! [X6,X7] : (s(bool,f(s(sK1,sK6(X6,sK4,bool,sK1)))) != s(bool,g(s(sK1,sK6(X6,sK4,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X7))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X7))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X7,sK1)),s(fun(sK1,bool),X6))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X7,sK1)),s(fun(sK1,bool),X7)))) )),
  inference(superposition,[],[f66,f53])).
fof(f11310,plain,(
  ( ! [X3] : (s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK3,X3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X3),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK3,X3,bool,sK1))))) ) | $spl118),
  inference(cnf_transformation,[],[f11310_D])).
fof(f11310_D,plain,(
  ( ! [X3] : (s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK3,X3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X3),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK3,X3,bool,sK1))))) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f11302,plain,(
  ( ! [X3] : (s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK3,X3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK3,X3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X3),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1))))) ) | $spl32),
  inference(trivial_inequality_removal,[],[f11194])).
fof(f11194,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK3,X3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK3,X3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X3),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1))))) ) | $spl32),
  inference(superposition,[],[f1550,f10812])).
fof(f1550,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,f(s(sK1,sK6(X2,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK3,X3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK3,X3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X3),s(fun(sK1,bool),X2))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X2))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X2,sK4,bool,sK1))))) )),
  inference(superposition,[],[f206,f44])).
fof(f206,plain,(
  ( ! [X6,X7] : (s(bool,f(s(sK1,sK6(X6,sK4,bool,sK1)))) != s(bool,g(s(sK1,sK6(X6,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK3,X7,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK3,X7,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X7),s(fun(sK1,bool),X6))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X7))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X6)))))) )),
  inference(superposition,[],[f116,f53])).
fof(f11309,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK3,X1,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X1),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X1))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK3,X1,bool,sK1))))) ) | $spl116),
  inference(cnf_transformation,[],[f11309_D])).
fof(f11309_D,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK3,X1,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X1),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X1))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK3,X1,bool,sK1))))) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f11303,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK3,X1,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK3,X1,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X1),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X1))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1))))) ) | $spl32),
  inference(trivial_inequality_removal,[],[f11192])).
fof(f11192,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK3,X1,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK3,X1,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X1),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X1))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1))))) ) | $spl32),
  inference(superposition,[],[f1540,f10812])).
fof(f1540,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,f(s(sK1,sK6(X2,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK3,X3,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK3,X3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X3),s(fun(sK1,bool),X2))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X2))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X2,sK4,bool,sK1))))) )),
  inference(superposition,[],[f166,f44])).
fof(f166,plain,(
  ( ! [X6,X7] : (s(bool,f(s(sK1,sK6(X6,sK4,bool,sK1)))) != s(bool,g(s(sK1,sK6(X6,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK3,X7,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK3,X7,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X7),s(fun(sK1,bool),X6))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X7))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X6)))))) )),
  inference(superposition,[],[f106,f53])).
fof(f11304,plain,(
  s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | $spl32),
  inference(trivial_inequality_removal,[],[f11190])).
fof(f11190,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | $spl32),
  inference(superposition,[],[f85,f10812])).
fof(f10816,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),bool,sK0)))) | $spl36),
  inference(cnf_transformation,[],[f10816_D])).
fof(f10816_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),bool,sK0)))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f11305,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | ($spl4 | $spl32)),
  inference(trivial_inequality_removal,[],[f11188])).
fof(f11188,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | ($spl4 | $spl32)),
  inference(superposition,[],[f7655,f10812])).
fof(f7655,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),X1,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X1)),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),X1,bool,sK1))))) ) | $spl4),
  inference(superposition,[],[f5987,f44])).
fof(f5987,plain,(
  ( ! [X2] : (s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),X2,bool,sK1)))) != s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),X2,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X2)),bool,sK0))))) ) | $spl4),
  inference(superposition,[],[f5950,f53])).
fof(f11308,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),bool,sK0)))) | $spl114),
  inference(cnf_transformation,[],[f11308_D])).
fof(f11308_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),bool,sK0)))) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f11306,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | ($spl4 | $spl32)),
  inference(trivial_inequality_removal,[],[f11186])).
fof(f11186,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | ($spl4 | $spl32)),
  inference(superposition,[],[f7587,f10812])).
fof(f7587,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),X1,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X1)),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),X1,bool,sK1))))) ) | $spl4),
  inference(superposition,[],[f5934,f44])).
fof(f5934,plain,(
  ( ! [X2] : (s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),X2,bool,sK1)))) != s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),X2,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X2)),bool,sK0))))) ) | $spl4),
  inference(superposition,[],[f5911,f53])).
fof(f10971,plain,(
  s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | $spl42),
  inference(cnf_transformation,[],[f10971_D])).
fof(f10971_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f10930,plain,(
  s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK3,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK3,sK4,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | $spl22),
  inference(trivial_inequality_removal,[],[f10919])).
fof(f10919,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK3,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK3,sK4,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | $spl22),
  inference(superposition,[],[f1612,f10802])).
fof(f11029,plain,(
  s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | $spl106),
  inference(cnf_transformation,[],[f11029_D])).
fof(f11029_D,plain,(
  s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f10931,plain,(
  s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | $spl22),
  inference(trivial_inequality_removal,[],[f10917])).
fof(f10917,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | $spl22),
  inference(superposition,[],[f1610,f10802])).
fof(f1610,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,f(s(sK1,sK6(X2,X3,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X2,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X2,sK4,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X3))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X2,X3,bool,sK1))))) )),
  inference(superposition,[],[f355,f44])).
fof(f355,plain,(
  ( ! [X4,X5] : (s(bool,f(s(sK1,sK6(X4,X5,bool,sK1)))) != s(bool,g(s(sK1,sK6(X4,X5,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X4,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X4,sK4,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X5))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X5)))))) )),
  inference(superposition,[],[f152,f53])).
fof(f11027,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl105),
  inference(cnf_transformation,[],[f11027_D])).
fof(f11027_D,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) <=> ~$spl105),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl105])])).
fof(f10932,plain,(
  s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | $spl22),
  inference(trivial_inequality_removal,[],[f10915])).
fof(f10915,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,sK4,sK1)),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | $spl22),
  inference(superposition,[],[f1608,f10802])).
fof(f10933,plain,(
  s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK3,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK3,sK4,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | $spl22),
  inference(trivial_inequality_removal,[],[f10913])).
fof(f10913,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK3,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK3,sK4,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | $spl22),
  inference(superposition,[],[f1602,f10802])).
fof(f10812,plain,(
  s(bool,'F_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | $spl32),
  inference(cnf_transformation,[],[f10812_D])).
fof(f10812_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f10814,plain,(
  s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | $spl34),
  inference(cnf_transformation,[],[f10814_D])).
fof(f10814_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f10934,plain,(
  s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | $spl22),
  inference(trivial_inequality_removal,[],[f10911])).
fof(f10911,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | $spl22),
  inference(superposition,[],[f1600,f10802])).
fof(f1600,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,f(s(sK1,sK6(X2,X3,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X2,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X2,sK4,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X3))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X2,X3,bool,sK1))))) )),
  inference(superposition,[],[f333,f44])).
fof(f333,plain,(
  ( ! [X4,X5] : (s(bool,f(s(sK1,sK6(X4,X5,bool,sK1)))) != s(bool,g(s(sK1,sK6(X4,X5,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X4,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X4,sK4,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X5))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X5)))))) )),
  inference(superposition,[],[f149,f53])).
fof(f11021,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,sK5('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,sK1)),s(fun(sK1,bool),sK4))) | $spl99),
  inference(cnf_transformation,[],[f11021_D])).
fof(f11021_D,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,sK5('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,sK1)),s(fun(sK1,bool),sK4))) <=> ~$spl99),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl99])])).
fof(f10935,plain,(
  s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,sK1)),s(fun(sK1,bool),sK4))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | $spl22),
  inference(trivial_inequality_removal,[],[f10909])).
fof(f10909,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,sK1)),s(fun(sK1,bool),sK4))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | $spl22),
  inference(superposition,[],[f1598,f10802])).
fof(f10936,plain,(
  s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK4,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK4,sK3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | $spl22),
  inference(trivial_inequality_removal,[],[f10907])).
fof(f10907,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK4,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK4,sK3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | $spl22),
  inference(superposition,[],[f1592,f10802])).
fof(f11013,plain,(
  s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK4,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),bool,sK1)))) | $spl90),
  inference(cnf_transformation,[],[f11013_D])).
fof(f11013_D,plain,(
  s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK4,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),bool,sK1)))) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f11015,plain,(
  s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK4,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),bool,sK1)))) | $spl92),
  inference(cnf_transformation,[],[f11015_D])).
fof(f11015_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK4,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),bool,sK1)))) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f10937,plain,(
  s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK4,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK4,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | $spl22),
  inference(trivial_inequality_removal,[],[f10905])).
fof(f10905,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK4,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK4,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | $spl22),
  inference(superposition,[],[f1590,f10802])).
fof(f1590,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,f(s(sK1,sK6(X2,X3,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK4,X2,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK4,X2,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X3))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X2,X3,bool,sK1))))) )),
  inference(superposition,[],[f293,f44])).
fof(f293,plain,(
  ( ! [X4,X5] : (s(bool,f(s(sK1,sK6(X4,X5,bool,sK1)))) != s(bool,g(s(sK1,sK6(X4,X5,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK4,X4,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK4,X4,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X5))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X5)))))) )),
  inference(superposition,[],[f146,f53])).
fof(f10938,plain,(
  s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK4,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK4,sK3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | $spl22),
  inference(trivial_inequality_removal,[],[f10903])).
fof(f10903,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK4,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK4,sK3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | $spl22),
  inference(superposition,[],[f1582,f10802])).
fof(f11005,plain,(
  s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK4,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),bool,sK1)))) | $spl82),
  inference(cnf_transformation,[],[f11005_D])).
fof(f11005_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK4,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),bool,sK1)))) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f11007,plain,(
  s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK4,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),bool,sK1)))) | $spl84),
  inference(cnf_transformation,[],[f11007_D])).
fof(f11007_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK4,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),bool,sK1)))) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f10939,plain,(
  s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK4,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK4,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | $spl22),
  inference(trivial_inequality_removal,[],[f10901])).
fof(f10901,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK4,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK4,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | $spl22),
  inference(superposition,[],[f1578,f10802])).
fof(f1578,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,f(s(sK1,sK6(X2,X3,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK4,X2,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK4,X2,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X3))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X2,X3,bool,sK1))))) )),
  inference(superposition,[],[f271,f44])).
fof(f271,plain,(
  ( ! [X4,X5] : (s(bool,f(s(sK1,sK6(X4,X5,bool,sK1)))) != s(bool,g(s(sK1,sK6(X4,X5,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK4,X4,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK4,X4,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),X5))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X5)))))) )),
  inference(superposition,[],[f143,f53])).
fof(f11001,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl79),
  inference(cnf_transformation,[],[f11001_D])).
fof(f11001_D,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) <=> ~$spl79),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl79])])).
fof(f10940,plain,(
  s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | $spl22),
  inference(trivial_inequality_removal,[],[f10899])).
fof(f10899,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,sK3,sK1)),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | $spl22),
  inference(superposition,[],[f1572,f10802])).
fof(f10993,plain,(
  s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK3,sK3,bool,sK1)))) | $spl70),
  inference(cnf_transformation,[],[f10993_D])).
fof(f10993_D,plain,(
  s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK3,sK3,bool,sK1)))) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f10995,plain,(
  s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK3,sK3,bool,sK1)))) | $spl72),
  inference(cnf_transformation,[],[f10995_D])).
fof(f10995_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK3,sK3,bool,sK1)))) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f10941,plain,(
  s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK3,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK3,sK3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | $spl22),
  inference(trivial_inequality_removal,[],[f10897])).
fof(f10897,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK3,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK3,sK3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | $spl22),
  inference(superposition,[],[f1566,f10802])).
fof(f10997,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK1)),s(fun(sK1,bool),sK4))) | $spl75),
  inference(cnf_transformation,[],[f10997_D])).
fof(f10997_D,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK1)),s(fun(sK1,bool),sK4))) <=> ~$spl75),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl75])])).
fof(f10942,plain,(
  s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK1)),s(fun(sK1,bool),sK4))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | $spl22),
  inference(trivial_inequality_removal,[],[f10895])).
fof(f10895,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK1)),s(fun(sK1,bool),sK4))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | $spl22),
  inference(superposition,[],[f1564,f10802])).
fof(f10989,plain,(
  s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK3,sK3,bool,sK1)))) | $spl66),
  inference(cnf_transformation,[],[f10989_D])).
fof(f10989_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK3,sK3,bool,sK1)))) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f10991,plain,(
  s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK3,sK3,bool,sK1)))) | $spl68),
  inference(cnf_transformation,[],[f10991_D])).
fof(f10991_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK3,sK3,bool,sK1)))) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f10945,plain,(
  s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK3,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK3,sK3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | $spl22),
  inference(trivial_inequality_removal,[],[f10891])).
fof(f10891,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK3,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK3,sK3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | $spl22),
  inference(superposition,[],[f1554,f10802])).
fof(f10948,plain,(
  s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK3,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK3,sK3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | $spl22),
  inference(trivial_inequality_removal,[],[f10887])).
fof(f10887,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK3,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK3,sK3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | $spl22),
  inference(superposition,[],[f1544,f10802])).
fof(f10951,plain,(
  s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK3,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK3,sK3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | $spl22),
  inference(trivial_inequality_removal,[],[f10881])).
fof(f10881,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK3,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK3,sK3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | $spl22),
  inference(superposition,[],[f1534,f10802])).
fof(f10953,plain,(
  s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | $spl22),
  inference(trivial_inequality_removal,[],[f10877])).
fof(f10877,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,sK1)),s(fun(sK1,bool),sK3))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | $spl22),
  inference(superposition,[],[f1530,f10802])).
fof(f10956,plain,(
  s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | $spl22),
  inference(trivial_inequality_removal,[],[f10873])).
fof(f10873,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK1)),s(fun(sK1,bool),sK3))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | $spl22),
  inference(superposition,[],[f1432,f10802])).
fof(f10987,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl65),
  inference(cnf_transformation,[],[f10987_D])).
fof(f10987_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))),s(fun(sK1,bool),sK4)))),bool,sK0)))) <=> ~$spl65),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl65])])).
fof(f10958,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))),s(fun(sK1,bool),sK4)))),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | ($spl4 | $spl22)),
  inference(trivial_inequality_removal,[],[f10871])).
fof(f10871,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))),s(fun(sK1,bool),sK4)))),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | ($spl4 | $spl22)),
  inference(superposition,[],[f6113,f10802])).
fof(f6113,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,f(s(sK1,sK6(X1,sK3,bool,sK1)))) | s(bool,'T_0') != s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X1))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X1),s(fun(sK1,bool),sK4)))),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X1,sK3,bool,sK1))))) ) | $spl4),
  inference(superposition,[],[f6065,f44])).
fof(f6065,plain,(
  ( ! [X3] : (s(bool,f(s(sK1,sK6(X3,sK3,bool,sK1)))) != s(bool,g(s(sK1,sK6(X3,sK3,bool,sK1)))) | s(bool,'T_0') != s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X3),s(fun(sK1,bool),sK4)))),bool,sK0))))) ) | $spl4),
  inference(superposition,[],[f6014,f53])).
fof(f10985,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl62),
  inference(cnf_transformation,[],[f10985_D])).
fof(f10985_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))),s(fun(sK1,bool),sK4)))),bool,sK0)))) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f10959,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))),s(fun(sK1,bool),sK4)))),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | ($spl4 | $spl22)),
  inference(trivial_inequality_removal,[],[f10869])).
fof(f10869,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))),s(fun(sK1,bool),sK4)))),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | ($spl4 | $spl22)),
  inference(superposition,[],[f6097,f10802])).
fof(f6097,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,f(s(sK1,sK6(X1,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X1))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X1),s(fun(sK1,bool),sK4)))),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X1,sK3,bool,sK1))))) ) | $spl4),
  inference(superposition,[],[f5980,f44])).
fof(f5980,plain,(
  ( ! [X3] : (s(bool,f(s(sK1,sK6(X3,sK3,bool,sK1)))) != s(bool,g(s(sK1,sK6(X3,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X3),s(fun(sK1,bool),sK4)))),bool,sK0))))) ) | $spl4),
  inference(superposition,[],[f5950,f53])).
fof(f10983,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl60),
  inference(cnf_transformation,[],[f10983_D])).
fof(f10983_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))),s(fun(sK1,bool),sK4)))),bool,sK0)))) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f10960,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))),s(fun(sK1,bool),sK4)))),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | ($spl4 | $spl22)),
  inference(trivial_inequality_removal,[],[f10867])).
fof(f10867,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))),s(fun(sK1,bool),sK4)))),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | ($spl4 | $spl22)),
  inference(superposition,[],[f6087,f10802])).
fof(f6087,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,f(s(sK1,sK6(X1,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X1))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X1),s(fun(sK1,bool),sK4)))),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X1,sK3,bool,sK1))))) ) | $spl4),
  inference(superposition,[],[f5927,f44])).
fof(f5927,plain,(
  ( ! [X3] : (s(bool,f(s(sK1,sK6(X3,sK3,bool,sK1)))) != s(bool,g(s(sK1,sK6(X3,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X3),s(fun(sK1,bool),sK4)))),bool,sK0))))) ) | $spl4),
  inference(superposition,[],[f5911,f53])).
fof(f10981,plain,(
  ( ! [X11] : (s(bool,'T_0') = s(bool,g(s(sK1,sK6(X11,sK4,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))),s(fun(sK1,bool),X11))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X11))))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X11,sK4,bool,sK1))))) ) | $spl58),
  inference(cnf_transformation,[],[f10981_D])).
fof(f10981_D,plain,(
  ( ! [X11] : (s(bool,'T_0') = s(bool,g(s(sK1,sK6(X11,sK4,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))),s(fun(sK1,bool),X11))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X11))))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X11,sK4,bool,sK1))))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f10961,plain,(
  ( ! [X11] : (s(bool,'T_0') = s(bool,g(s(sK1,sK6(X11,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X11,sK4,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))),s(fun(sK1,bool),X11))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X11))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1))))) ) | $spl22),
  inference(trivial_inequality_removal,[],[f10865])).
fof(f10865,plain,(
  ( ! [X11] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X11,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X11,sK4,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))),s(fun(sK1,bool),X11))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X11))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1))))) ) | $spl22),
  inference(superposition,[],[f1606,f10802])).
fof(f1606,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,f(s(sK1,sK6(X2,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X3,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X3,sK4,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X2),s(fun(sK1,bool),X3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X2))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X3))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X2,sK3,bool,sK1))))) )),
  inference(superposition,[],[f346,f44])).
fof(f346,plain,(
  ( ! [X6,X7] : (s(bool,f(s(sK1,sK6(X6,sK3,bool,sK1)))) != s(bool,g(s(sK1,sK6(X6,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X7,sK4,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(X7,sK4,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X6),s(fun(sK1,bool),X7))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X6))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X7)))))) )),
  inference(superposition,[],[f152,f53])).
fof(f10980,plain,(
  ( ! [X9] : (s(bool,'F_0') = s(bool,g(s(sK1,sK6(X9,sK4,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))),s(fun(sK1,bool),X9))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X9))))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X9,sK4,bool,sK1))))) ) | $spl56),
  inference(cnf_transformation,[],[f10980_D])).
fof(f10980_D,plain,(
  ( ! [X9] : (s(bool,'F_0') = s(bool,g(s(sK1,sK6(X9,sK4,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))),s(fun(sK1,bool),X9))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X9))))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X9,sK4,bool,sK1))))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f10962,plain,(
  ( ! [X9] : (s(bool,'F_0') = s(bool,g(s(sK1,sK6(X9,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X9,sK4,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))),s(fun(sK1,bool),X9))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X9))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1))))) ) | $spl22),
  inference(trivial_inequality_removal,[],[f10863])).
fof(f10863,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X9,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X9,sK4,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))),s(fun(sK1,bool),X9))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X9))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1))))) ) | $spl22),
  inference(superposition,[],[f1596,f10802])).
fof(f1596,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,f(s(sK1,sK6(X2,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X3,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X3,sK4,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X2),s(fun(sK1,bool),X3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X2))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X3))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X2,sK3,bool,sK1))))) )),
  inference(superposition,[],[f324,f44])).
fof(f324,plain,(
  ( ! [X6,X7] : (s(bool,f(s(sK1,sK6(X6,sK3,bool,sK1)))) != s(bool,g(s(sK1,sK6(X6,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(X7,sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(X7,sK4,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X6),s(fun(sK1,bool),X7))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X6))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X7)))))) )),
  inference(superposition,[],[f149,f53])).
fof(f10979,plain,(
  ( ! [X7] : (s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK4,X7,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))),s(fun(sK1,bool),X7))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X7))))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK4,X7,bool,sK1))))) ) | $spl54),
  inference(cnf_transformation,[],[f10979_D])).
fof(f10979_D,plain,(
  ( ! [X7] : (s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK4,X7,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))),s(fun(sK1,bool),X7))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X7))))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK4,X7,bool,sK1))))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f10963,plain,(
  ( ! [X7] : (s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK4,X7,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK4,X7,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))),s(fun(sK1,bool),X7))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X7))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1))))) ) | $spl22),
  inference(trivial_inequality_removal,[],[f10861])).
fof(f10861,plain,(
  ( ! [X7] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK4,X7,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK4,X7,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))),s(fun(sK1,bool),X7))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X7))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1))))) ) | $spl22),
  inference(superposition,[],[f1586,f10802])).
fof(f1586,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,f(s(sK1,sK6(X2,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK4,X3,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK4,X3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X2),s(fun(sK1,bool),X3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X2))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X3))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X2,sK3,bool,sK1))))) )),
  inference(superposition,[],[f284,f44])).
fof(f284,plain,(
  ( ! [X6,X7] : (s(bool,f(s(sK1,sK6(X6,sK3,bool,sK1)))) != s(bool,g(s(sK1,sK6(X6,sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(sK4,X7,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK4,X7,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X6),s(fun(sK1,bool),X7))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X6))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X7)))))) )),
  inference(superposition,[],[f146,f53])).
fof(f10978,plain,(
  ( ! [X5] : (s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK4,X5,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))),s(fun(sK1,bool),X5))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X5))))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK4,X5,bool,sK1))))) ) | $spl52),
  inference(cnf_transformation,[],[f10978_D])).
fof(f10978_D,plain,(
  ( ! [X5] : (s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK4,X5,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))),s(fun(sK1,bool),X5))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X5))))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK4,X5,bool,sK1))))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f10964,plain,(
  ( ! [X5] : (s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK4,X5,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK4,X5,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))),s(fun(sK1,bool),X5))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X5))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1))))) ) | $spl22),
  inference(trivial_inequality_removal,[],[f10859])).
fof(f10859,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK4,X5,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK4,X5,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))),s(fun(sK1,bool),X5))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X5))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1))))) ) | $spl22),
  inference(superposition,[],[f1576,f10802])).
fof(f1576,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,f(s(sK1,sK6(X2,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK4,X3,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK4,X3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X2),s(fun(sK1,bool),X3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X2))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X3))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X2,sK3,bool,sK1))))) )),
  inference(superposition,[],[f262,f44])).
fof(f262,plain,(
  ( ! [X6,X7] : (s(bool,f(s(sK1,sK6(X6,sK3,bool,sK1)))) != s(bool,g(s(sK1,sK6(X6,sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK4,X7,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6(sK4,X7,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X6),s(fun(sK1,bool),X7))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X6))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X7)))))) )),
  inference(superposition,[],[f143,f53])).
fof(f10977,plain,(
  ( ! [X3] : (s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X3,sK3,sK1)),s(fun(sK1,bool),X3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X3,sK3,sK1)),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))))) ) | $spl50),
  inference(cnf_transformation,[],[f10977_D])).
fof(f10977_D,plain,(
  ( ! [X3] : (s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X3,sK3,sK1)),s(fun(sK1,bool),X3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X3,sK3,sK1)),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f10965,plain,(
  ( ! [X3] : (s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X3,sK3,sK1)),s(fun(sK1,bool),X3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X3,sK3,sK1)),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1))))) ) | $spl22),
  inference(trivial_inequality_removal,[],[f10857])).
fof(f10857,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X3,sK3,sK1)),s(fun(sK1,bool),X3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X3,sK3,sK1)),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1))))) ) | $spl22),
  inference(superposition,[],[f1501,f10802])).
fof(f1501,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,f(s(sK1,sK6(X2,sK3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X3,sK3,sK1)),s(fun(sK1,bool),X2))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X3,sK3,sK1)),s(fun(sK1,bool),X3))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X2,sK3,bool,sK1))))) )),
  inference(superposition,[],[f126,f44])).
fof(f126,plain,(
  ( ! [X6,X7] : (s(bool,f(s(sK1,sK6(X6,sK3,bool,sK1)))) != s(bool,g(s(sK1,sK6(X6,sK3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X7),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X7))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X7,sK3,sK1)),s(fun(sK1,bool),X6))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X7,sK3,sK1)),s(fun(sK1,bool),X7)))) )),
  inference(superposition,[],[f57,f53])).
fof(f10976,plain,(
  ( ! [X1] : (s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X1),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X1))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X1,sK1)),s(fun(sK1,bool),X1))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X1,sK1)),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))))) ) | $spl48),
  inference(cnf_transformation,[],[f10976_D])).
fof(f10976_D,plain,(
  ( ! [X1] : (s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X1),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X1))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X1,sK1)),s(fun(sK1,bool),X1))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X1,sK1)),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f10966,plain,(
  ( ! [X1] : (s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X1),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X1))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X1,sK1)),s(fun(sK1,bool),X1))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X1,sK1)),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1))))) ) | $spl22),
  inference(trivial_inequality_removal,[],[f10855])).
fof(f10855,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X1),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X1))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X1,sK1)),s(fun(sK1,bool),X1))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X1,sK1)),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1))))) ) | $spl22),
  inference(superposition,[],[f1403,f10802])).
fof(f1403,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,f(s(sK1,sK6(X2,sK3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X3,sK1)),s(fun(sK1,bool),X2))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X3,sK1)),s(fun(sK1,bool),X3))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6(X2,sK3,bool,sK1))))) )),
  inference(superposition,[],[f89,f44])).
fof(f89,plain,(
  ( ! [X6,X7] : (s(bool,f(s(sK1,sK6(X6,sK3,bool,sK1)))) != s(bool,g(s(sK1,sK6(X6,sK3,bool,sK1)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X7),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X7))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X7,sK1)),s(fun(sK1,bool),X6))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X7,sK1)),s(fun(sK1,bool),X7)))) )),
  inference(superposition,[],[f56,f53])).
fof(f10967,plain,(
  s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | $spl22),
  inference(trivial_inequality_removal,[],[f10853])).
fof(f10853,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | $spl22),
  inference(superposition,[],[f81,f10802])).
fof(f10806,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),bool,sK0)))) | $spl26),
  inference(cnf_transformation,[],[f10806_D])).
fof(f10806_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),bool,sK0)))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f10968,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | ($spl4 | $spl22)),
  inference(trivial_inequality_removal,[],[f10851])).
fof(f10851,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | ($spl4 | $spl22)),
  inference(superposition,[],[f7655,f10802])).
fof(f10973,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),bool,sK0)))) | $spl44),
  inference(cnf_transformation,[],[f10973_D])).
fof(f10973_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),bool,sK0)))) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f10969,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | ($spl4 | $spl22)),
  inference(trivial_inequality_removal,[],[f10849])).
fof(f10849,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | ($spl4 | $spl22)),
  inference(superposition,[],[f7587,f10802])).
fof(f10820,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl40),
  inference(cnf_transformation,[],[f10820_D])).
fof(f10820_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f10782,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),bool,sK0)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | $spl4),
  inference(trivial_inequality_removal,[],[f10771])).
fof(f10771,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),bool,sK0)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | $spl4),
  inference(superposition,[],[f7637,f10024])).
fof(f10024,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X0)),bool,sK0)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),X0,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),X0,bool,sK1))))) ) | $spl4),
  inference(trivial_inequality_removal,[],[f10022])).
fof(f10022,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X0)),bool,sK0)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),X0,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),X0,bool,sK1))))) ) | $spl4),
  inference(superposition,[],[f7586,f44])).
fof(f7586,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),X0,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X0)),bool,sK0)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),X0,bool,sK1))))) ) | $spl4),
  inference(superposition,[],[f5934,f44])).
fof(f7637,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,f(s(sK0,sK6(X1,'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),X1)),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK6(X1,'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),bool,sK0))))) ) | $spl4),
  inference(superposition,[],[f5974,f44])).
fof(f5974,plain,(
  ( ! [X3] : (s(bool,f(s(sK0,sK6(X3,'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),bool,sK0)))) != s(bool,g(s(sK0,sK6(X3,'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),X3)),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0))))) ) | $spl4),
  inference(superposition,[],[f5950,f53])).
fof(f10818,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl38),
  inference(cnf_transformation,[],[f10818_D])).
fof(f10818_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f10783,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),bool,sK0)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | $spl4),
  inference(trivial_inequality_removal,[],[f10769])).
fof(f10769,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),bool,sK0)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | $spl4),
  inference(superposition,[],[f7577,f10024])).
fof(f7577,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,f(s(sK0,sK6(X1,'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),X1)),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK6(X1,'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),bool,sK0))))) ) | $spl4),
  inference(superposition,[],[f5921,f44])).
fof(f5921,plain,(
  ( ! [X3] : (s(bool,f(s(sK0,sK6(X3,'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),bool,sK0)))) != s(bool,g(s(sK0,sK6(X3,'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),X3)),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0))))) ) | $spl4),
  inference(superposition,[],[f5911,f53])).
fof(f10784,plain,(
  s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),bool,sK0)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | $spl4),
  inference(trivial_inequality_removal,[],[f10767])).
fof(f10767,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),bool,sK0)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,bool,sK1)))) | $spl4),
  inference(superposition,[],[f396,f10024])).
fof(f396,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,f(s(sK0,sK6(X1,'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),bool,sK0)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),X1))) | s(bool,'T_0') = s(bool,g(s(sK0,sK6(X1,'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),bool,sK0))))) )),
  inference(superposition,[],[f73,f44])).
fof(f73,plain,(
  ( ! [X3] : (s(bool,f(s(sK0,sK6(X3,'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),bool,sK0)))) != s(bool,g(s(sK0,sK6(X3,'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)),bool,sK0)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),X3)))) )),
  inference(superposition,[],[f39,f53])).
fof(f10802,plain,(
  s(bool,'F_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | $spl22),
  inference(cnf_transformation,[],[f10802_D])).
fof(f10802_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f10804,plain,(
  s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | $spl24),
  inference(cnf_transformation,[],[f10804_D])).
fof(f10804_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f10810,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl30),
  inference(cnf_transformation,[],[f10810_D])).
fof(f10810_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f10785,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),bool,sK0)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | $spl4),
  inference(trivial_inequality_removal,[],[f10765])).
fof(f10765,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),bool,sK0)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | $spl4),
  inference(superposition,[],[f7621,f10024])).
fof(f7621,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,f(s(sK0,sK6(X1,'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),X1),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK6(X1,'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),bool,sK0))))) ) | $spl4),
  inference(superposition,[],[f5970,f44])).
fof(f5970,plain,(
  ( ! [X3] : (s(bool,f(s(sK0,sK6(X3,'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),bool,sK0)))) != s(bool,g(s(sK0,sK6(X3,'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),X3),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0))))) ) | $spl4),
  inference(superposition,[],[f5950,f53])).
fof(f10808,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl28),
  inference(cnf_transformation,[],[f10808_D])).
fof(f10808_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f10786,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),bool,sK0)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | $spl4),
  inference(trivial_inequality_removal,[],[f10763])).
fof(f10763,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),bool,sK0)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | $spl4),
  inference(superposition,[],[f7561,f10024])).
fof(f7561,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,f(s(sK0,sK6(X1,'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),X1),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK6(X1,'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),bool,sK0))))) ) | $spl4),
  inference(superposition,[],[f5917,f44])).
fof(f5917,plain,(
  ( ! [X3] : (s(bool,f(s(sK0,sK6(X3,'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),bool,sK0)))) != s(bool,g(s(sK0,sK6(X3,'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),X3),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0))))) ) | $spl4),
  inference(superposition,[],[f5911,f53])).
fof(f10787,plain,(
  s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),bool,sK0)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | $spl4),
  inference(trivial_inequality_removal,[],[f10761])).
fof(f10761,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),bool,sK0)))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | s(bool,'F_0') = s(bool,f(s(sK1,sK6('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,bool,sK1)))) | $spl4),
  inference(superposition,[],[f380,f10024])).
fof(f380,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,f(s(sK0,sK6(X1,'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),bool,sK0)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X1),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK6(X1,'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),bool,sK0))))) )),
  inference(superposition,[],[f63,f44])).
fof(f63,plain,(
  ( ! [X3] : (s(bool,f(s(sK0,sK6(X3,'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),bool,sK0)))) != s(bool,g(s(sK0,sK6(X3,'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)),bool,sK0)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X3),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))))) )),
  inference(superposition,[],[f39,f53])).
fof(f6024,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl13),
  inference(cnf_transformation,[],[f6024_D])).
fof(f6024_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f5994,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK1)),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl4),
  inference(superposition,[],[f375,f5950])).
fof(f375,plain,(
  ( ! [X0,X1] : (s(bool,f(s(sK0,sK6(X1,'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X0)),bool,sK0)))) != s(bool,g(s(sK0,sK6(X1,'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X0)),bool,sK0)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X1),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X0,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK3,X0,sK1)),s(fun(sK1,bool),X0)))) )),
  inference(superposition,[],[f63,f41])).
fof(f5993,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK3,sK1)),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl4),
  inference(superposition,[],[f376,f5950])).
fof(f376,plain,(
  ( ! [X2,X3] : (s(bool,f(s(sK0,sK6(X3,'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X2)),bool,sK0)))) != s(bool,g(s(sK0,sK6(X3,'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X2)),bool,sK0)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X3),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X2,sK3,sK1)),s(fun(sK1,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X2,sK3,sK1)),s(fun(sK1,bool),X2)))) )),
  inference(superposition,[],[f63,f41])).
fof(f5992,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK1)),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl4),
  inference(superposition,[],[f391,f5950])).
fof(f391,plain,(
  ( ! [X0,X1] : (s(bool,f(s(sK0,sK6(X1,'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X0)),bool,sK0)))) != s(bool,g(s(sK0,sK6(X1,'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X0)),bool,sK0)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),X1))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X0,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(sK4,X0,sK1)),s(fun(sK1,bool),X0)))) )),
  inference(superposition,[],[f73,f41])).
fof(f5991,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | s(bool,'const.bool.IN_2'(s(sK1,sK5('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5('const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)),sK4,sK1)),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl4),
  inference(superposition,[],[f392,f5950])).
fof(f392,plain,(
  ( ! [X2,X3] : (s(bool,f(s(sK0,sK6(X3,'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X2)),bool,sK0)))) != s(bool,g(s(sK0,sK6(X3,'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X2)),bool,sK0)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),X3))) | s(bool,'const.bool.IN_2'(s(sK1,sK5(X2,sK4,sK1)),s(fun(sK1,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK1,sK5(X2,sK4,sK1)),s(fun(sK1,bool),X2)))) )),
  inference(superposition,[],[f73,f41])).
fof(f5911,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl4),
  inference(cnf_transformation,[],[f5911_D])).
fof(f5911_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f5913,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl6),
  inference(cnf_transformation,[],[f5913_D])).
fof(f5913_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f5909,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0))))),
  inference(equality_resolution,[],[f3163])).
fof(f3163,plain,(
  ( ! [X0] : (s(fun(sK0,bool),X0) != s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') = s(bool,g(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),X0,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),X0,bool,sK0))))) )),
  inference(trivial_inequality_removal,[],[f3161])).
fof(f3161,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(fun(sK0,bool),X0) != s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') = s(bool,g(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),X0,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),X0,bool,sK0))))) )),
  inference(superposition,[],[f777,f44])).
fof(f777,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),X0,bool,sK0)))) | s(fun(sK0,bool),X0) != s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') = s(bool,g(s(sK0,sK6('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),X0,bool,sK0))))) )),
  inference(superposition,[],[f76,f44])).
fof(f108,plain,(
  ( ! [X1] : (s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X1),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X1))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK3,X1,bool,sK1))))) ) | $spl0),
  inference(cnf_transformation,[],[f108_D])).
fof(f108_D,plain,(
  ( ! [X1] : (s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X1),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X1))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK3,X1,bool,sK1))))) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f110,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl3),
  inference(cnf_transformation,[],[f110_D])).
fof(f110_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f107,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X1),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X1))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK3,X1,bool,sK1))))) )),
  inference(subsumption_resolution,[],[f105,f78])).
fof(f105,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),X1),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),X1))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | s(bool,'F_0') = s(bool,g(s(sK1,sK6(sK3,X1,bool,sK1)))) | s(bool,'T_0') = s(bool,f(s(sK1,sK6(sK3,X1,bool,sK1))))) )),
  inference(superposition,[],[f78,f44])).
% SZS output end Proof for thm.util_prob
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 1673
% Passive clauses: 17964
% Generated clauses: 51986
% Final active clauses: 1622
% Final passive clauses: 14981
% Input formulas: 11
% Initial clauses: 16
% 
% Selected by SInE selection: 11
% SInE iterations: 1
% 
% Duplicate literals: 3889
% Trivial inequalities: 7309
% Fw subsumption resolutions: 919
% Fw demodulations: 769
% Bw demodulations: 9
% 
% Simple tautologies: 2551
% Equational tautologies: 1
% Forward subsumptions: 19379
% 
% Binary resolution: 160
% Forward superposition: 18544
% Backward superposition: 20447
% Equality factoring: 32
% Equality resolution: 437
% 
% Split clauses: 302
% Split components: 1085
% SAT solver clauses: 312
% SAT solver binary clauses: 19
% 
% Sat splits: 302
% Sat splitting refutations: 9
% 
% Memory used [KB]: 26097
% Time elapsed: 2.413 s
% ------------------------------
% ------------------------------
% Success in time 14.121 s
