Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.prim_rec
% (95068)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 183
% Passive clauses: 4696
% Generated clauses: 26144
% Final active clauses: 161
% Final passive clauses: 4372
% Input formulas: 69
% Initial clauses: 116
% 
% Duplicate literals: 899
% Trivial inequalities: 368
% Fw subsumption resolutions: 2117
% Fw demodulations: 4472
% Bw demodulations: 21
% 
% Simple tautologies: 417
% Equational tautologies: 54
% Forward subsumptions: 12413
% 
% Binary resolution: 426
% Factoring: 1
% Forward superposition: 8214
% Backward superposition: 9281
% Self superposition: 170
% Equality factoring: 96
% Equality resolution: 32
% 
% Split clauses: 348
% Split components: 707
% SAT solver clauses: 839
% SAT solver unit clauses: 6
% SAT solver binary clauses: 133
% 
% Sat splits: 348
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 839
% TWLsolver calls for satisfiability: 175
% 
% Memory used [KB]: 16119
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.prim_rec
% (95076)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 129
% Passive clauses: 1490
% Generated clauses: 16635
% Final active clauses: 103
% Final passive clauses: 1101
% Input formulas: 69
% Initial clauses: 116
% 
% Duplicate literals: 602
% Trivial inequalities: 88
% Fw subsumption resolutions: 844
% Fw demodulations: 2067
% Bw demodulations: 282
% Condensations: 1269
% Global subsumptions: 9
% 
% Simple tautologies: 175
% Equational tautologies: 43
% Forward subsumptions: 7260
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 403
% Factoring: 3
% Forward superposition: 5321
% Backward superposition: 4131
% Self superposition: 135
% Equality factoring: 68
% Equality resolution: 15
% 
% Split clauses: 70
% Split components: 140
% SAT solver clauses: 2234
% SAT solver unit clauses: 158
% SAT solver binary clauses: 54
% SAT solver learnt clauses: 18
% 
% Sat splits: 70
% Sat splitting refutations: 7
% 
% TWLsolver clauses: 2198
% TWLsolver calls for satisfiability: 571
% 
% Memory used [KB]: 16375
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.prim_rec
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.prim_rec
% SZS output start Proof for thm.prim_rec
fof(f53546,plain,(
  $false),
  inference(subsumption_resolution,[],[f53545,f252])).
fof(f252,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.prim_rec.PRIM_REC_THM|split|1',aHL_TRUTH)).
fof(f53545,plain,(
  ~p(s(bool,'T_0'))),
  inference(forward_demodulation,[],[f53544,f3560])).
fof(f3560,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X1)))) )),
  inference(unit_resulting_resolution,[],[f3484,f264])).
fof(f264,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f198])).
fof(f198,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f48])).
fof(f48,axiom,(
  ! [X33] : (s(bool,'T_0') = s(bool,X33) <=> p(s(bool,X33)))),
  file('thm.prim_rec.PRIM_REC_THM|split|1','thm.bool.EQ_CLAUSES|split|1')).
fof(f3484,plain,(
  ( ! [X2,X0] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X2))),s(X0,X2))))) )),
  inference(forward_demodulation,[],[f3428,f310])).
fof(f310,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f106])).
fof(f106,plain,(
  ! [X0,X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f105])).
fof(f105,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f61])).
fof(f61,axiom,(
  ! [X6] : ! [X27,X28] : s(X6,X28) = s(X6,'const.bool.COND_3'(s(bool,'F_0'),s(X6,X27),s(X6,X28)))),
  file('thm.prim_rec.PRIM_REC_THM|split|1','thm.bool.COND_CLAUSES|split|1')).
fof(f3428,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2))))),s(X0,X2))))) )),
  inference(unit_resulting_resolution,[],[f310,f357])).
fof(f357,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) != s(X0,X2)) )),
  inference(cnf_transformation,[],[f250])).
fof(f250,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2))))))),
  inference(nnf_transformation,[],[f158])).
fof(f158,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X36,X37,X38] : (p(s(bool,happ(s(fun(X36,bool),happ(s(fun(X36,fun(X36,bool)),'=_0'),s(X36,X37))),s(X36,X38)))) <=> s(X36,X37) = s(X36,X38))),
  file('thm.prim_rec.PRIM_REC_THM|split|1',arity_0_HL_EQ)).
fof(f53544,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'=_0'),s(sK0,'const.prim_rec.PRIM_REC_3'(s(sK0,sK1),s(fun(sK0,fun('type.num.num',sK0)),sK2),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK3))))))),s(sK0,'const.prim_rec.PRIM_REC_3'(s(sK0,sK1),s(fun(sK0,fun('type.num.num',sK0)),sK2),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK3))))))))),
  inference(backward_demodulation,[],[f53543,f3213])).
fof(f3213,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'=_0'),s(sK0,happ(s(fun('type.num.num',sK0),happ(s(fun(sK0,fun('type.num.num',sK0)),sK2),s(sK0,'const.prim_rec.PRIM_REC_3'(s(sK0,sK1),s(fun(sK0,fun('type.num.num',sK0)),sK2),s('type.num.num',sK3))))),s('type.num.num',sK3))))),s(sK0,'const.prim_rec.PRIM_REC_3'(s(sK0,sK1),s(fun(sK0,fun('type.num.num',sK0)),sK2),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK3))))))))),
  inference(unit_resulting_resolution,[],[f251,f356])).
fof(f356,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) )),
  inference(cnf_transformation,[],[f250])).
fof(f251,plain,(
  s(sK0,'const.prim_rec.PRIM_REC_3'(s(sK0,sK1),s(fun(sK0,fun('type.num.num',sK0)),sK2),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK3))))) != s(sK0,happ(s(fun('type.num.num',sK0),happ(s(fun(sK0,fun('type.num.num',sK0)),sK2),s(sK0,'const.prim_rec.PRIM_REC_3'(s(sK0,sK1),s(fun(sK0,fun('type.num.num',sK0)),sK2),s('type.num.num',sK3))))),s('type.num.num',sK3)))),
  inference(cnf_transformation,[],[f195])).
fof(f195,plain,(
  s(sK0,'const.prim_rec.PRIM_REC_3'(s(sK0,sK1),s(fun(sK0,fun('type.num.num',sK0)),sK2),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK3))))) != s(sK0,happ(s(fun('type.num.num',sK0),happ(s(fun(sK0,fun('type.num.num',sK0)),sK2),s(sK0,'const.prim_rec.PRIM_REC_3'(s(sK0,sK1),s(fun(sK0,fun('type.num.num',sK0)),sK2),s('type.num.num',sK3))))),s('type.num.num',sK3)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f177])).
fof(f177,plain,(
  ? [X0,X1,X2,X3] : s(X0,'const.prim_rec.PRIM_REC_3'(s(X0,X1),s(fun(X0,fun('type.num.num',X0)),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))) != s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(X0,fun('type.num.num',X0)),X2),s(X0,'const.prim_rec.PRIM_REC_3'(s(X0,X1),s(fun(X0,fun('type.num.num',X0)),X2),s('type.num.num',X3))))),s('type.num.num',X3)))),
  inference(ennf_transformation,[],[f72])).
fof(f72,plain,(
  ~! [X0,X1,X2,X3] : s(X0,'const.prim_rec.PRIM_REC_3'(s(X0,X1),s(fun(X0,fun('type.num.num',X0)),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(X0,fun('type.num.num',X0)),X2),s(X0,'const.prim_rec.PRIM_REC_3'(s(X0,X1),s(fun(X0,fun('type.num.num',X0)),X2),s('type.num.num',X3))))),s('type.num.num',X3)))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ~! [X0] : ! [X1,X2,X3] : s(X0,'const.prim_rec.PRIM_REC_3'(s(X0,X1),s(fun(X0,fun('type.num.num',X0)),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(X0,fun('type.num.num',X0)),X2),s(X0,'const.prim_rec.PRIM_REC_3'(s(X0,X1),s(fun(X0,fun('type.num.num',X0)),X2),s('type.num.num',X3))))),s('type.num.num',X3)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : s(X6,'const.prim_rec.PRIM_REC_3'(s(X6,X7),s(fun(X6,fun('type.num.num',X6)),X8),s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))))) = s(X6,happ(s(fun('type.num.num',X6),happ(s(fun(X6,fun('type.num.num',X6)),X8),s(X6,'const.prim_rec.PRIM_REC_3'(s(X6,X7),s(fun(X6,fun('type.num.num',X6)),X8),s('type.num.num',X9))))),s('type.num.num',X9)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : s(X6,'const.prim_rec.PRIM_REC_3'(s(X6,X7),s(fun(X6,fun('type.num.num',X6)),X8),s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))))) = s(X6,happ(s(fun('type.num.num',X6),happ(s(fun(X6,fun('type.num.num',X6)),X8),s(X6,'const.prim_rec.PRIM_REC_3'(s(X6,X7),s(fun(X6,fun('type.num.num',X6)),X8),s('type.num.num',X9))))),s('type.num.num',X9)))),
  file('thm.prim_rec.PRIM_REC_THM|split|1','thm.prim_rec.PRIM_REC_THM|split|1')).
fof(f53543,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.prim_rec.PRIM_REC_3'(s(X0,X1),s(fun(X0,fun('type.num.num',X0)),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(X0,fun('type.num.num',X0)),X2),s(X0,'const.prim_rec.PRIM_REC_3'(s(X0,X1),s(fun(X0,fun('type.num.num',X0)),X2),s('type.num.num',X3))))),s('type.num.num',X3)))) )),
  inference(forward_demodulation,[],[f53305,f29312])).
fof(f29312,plain,(
  ( ! [X6,X4,X5,X3] : (s(X4,'const.prim_rec.PRIM_REC_3'(s(X4,X5),s(fun(X4,fun('type.num.num',X4)),X6),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))) = s(X4,'const.prim_rec.PRIM_REC_FUN_4'(s(X4,X5),s(fun(X4,fun('type.num.num',X4)),X6),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))),s('type.num.num',X3)))) )),
  inference(superposition,[],[f319,f256])).
fof(f256,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X9] : s('type.num.num',X9) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X9)))))),
  file('thm.prim_rec.PRIM_REC_THM|split|1','thm.prim_rec.PRE|split|1')).
fof(f319,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.prim_rec.PRIM_REC_3'(s(X0,X1),s(fun(X0,fun('type.num.num',X0)),X2),s('type.num.num',X3))) = s(X0,'const.prim_rec.PRIM_REC_FUN_4'(s(X0,X1),s(fun(X0,fun('type.num.num',X0)),X2),s('type.num.num',X3),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',X3)))))) )),
  inference(cnf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.prim_rec.PRIM_REC_3'(s(X0,X1),s(fun(X0,fun('type.num.num',X0)),X2),s('type.num.num',X3))) = s(X0,'const.prim_rec.PRIM_REC_FUN_4'(s(X0,X1),s(fun(X0,fun('type.num.num',X0)),X2),s('type.num.num',X3),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',X3)))))),
  inference(flattening,[],[f117])).
fof(f117,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.prim_rec.PRIM_REC_3'(s(X0,X1),s(fun(X0,fun('type.num.num',X0)),X2),s('type.num.num',X3))) = s(X0,'const.prim_rec.PRIM_REC_FUN_4'(s(X0,X1),s(fun(X0,fun('type.num.num',X0)),X2),s('type.num.num',X3),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',X3)))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6] : ! [X7,X8,X9] : s(X6,'const.prim_rec.PRIM_REC_3'(s(X6,X7),s(fun(X6,fun('type.num.num',X6)),X8),s('type.num.num',X9))) = s(X6,'const.prim_rec.PRIM_REC_FUN_4'(s(X6,X7),s(fun(X6,fun('type.num.num',X6)),X8),s('type.num.num',X9),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',X9)))))),
  file('thm.prim_rec.PRIM_REC_THM|split|1','thm.prim_rec.PRIM_REC')).
fof(f53305,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.prim_rec.PRIM_REC_FUN_4'(s(X0,X1),s(fun(X0,fun('type.num.num',X0)),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(X0,fun('type.num.num',X0)),X2),s(X0,'const.prim_rec.PRIM_REC_3'(s(X0,X1),s(fun(X0,fun('type.num.num',X0)),X2),s('type.num.num',X3))))),s('type.num.num',X3)))) )),
  inference(superposition,[],[f325,f319])).
fof(f325,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.prim_rec.PRIM_REC_FUN_4'(s(X0,X1),s(fun(X0,fun('type.num.num',X0)),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))),s('type.num.num',X4))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(X0,fun('type.num.num',X0)),X2),s(X0,'const.prim_rec.PRIM_REC_FUN_4'(s(X0,X1),s(fun(X0,fun('type.num.num',X0)),X2),s('type.num.num',X3),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',X4))))))),s('type.num.num',X4)))) )),
  inference(cnf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0,X1,X2,X3,X4] : s(X0,'const.prim_rec.PRIM_REC_FUN_4'(s(X0,X1),s(fun(X0,fun('type.num.num',X0)),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))),s('type.num.num',X4))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(X0,fun('type.num.num',X0)),X2),s(X0,'const.prim_rec.PRIM_REC_FUN_4'(s(X0,X1),s(fun(X0,fun('type.num.num',X0)),X2),s('type.num.num',X3),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',X4))))))),s('type.num.num',X4)))),
  inference(flattening,[],[f127])).
fof(f127,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s(X0,'const.prim_rec.PRIM_REC_FUN_4'(s(X0,X1),s(fun(X0,fun('type.num.num',X0)),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))),s('type.num.num',X4))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(X0,fun('type.num.num',X0)),X2),s(X0,'const.prim_rec.PRIM_REC_FUN_4'(s(X0,X1),s(fun(X0,fun('type.num.num',X0)),X2),s('type.num.num',X3),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',X4))))))),s('type.num.num',X4)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6] : ! [X7,X8,X9,X14] : s(X6,'const.prim_rec.PRIM_REC_FUN_4'(s(X6,X7),s(fun(X6,fun('type.num.num',X6)),X8),s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))),s('type.num.num',X14))) = s(X6,happ(s(fun('type.num.num',X6),happ(s(fun(X6,fun('type.num.num',X6)),X8),s(X6,'const.prim_rec.PRIM_REC_FUN_4'(s(X6,X7),s(fun(X6,fun('type.num.num',X6)),X8),s('type.num.num',X9),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',X14))))))),s('type.num.num',X14)))),
  file('thm.prim_rec.PRIM_REC_THM|split|1','thm.prim_rec.PRIM_REC_EQN|split|1')).
% SZS output end Proof for thm.prim_rec
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 362
% Passive clauses: 5495
% Generated clauses: 56744
% Final active clauses: 129
% Final passive clauses: 1804
% Input formulas: 69
% Initial clauses: 116
% 
% Duplicate literals: 1566
% Trivial inequalities: 373
% Fw subsumption resolutions: 2303
% Fw demodulations: 10431
% Bw demodulations: 3042
% Condensations: 2515
% Global subsumptions: 161
% 
% Simple tautologies: 429
% Equational tautologies: 441
% Forward subsumptions: 32249
% Fw demodulations to eq. taut.: 86
% Bw demodulations to eq. taut.: 37
% 
% Unit resulting resolution: 3984
% Factoring: 1
% Forward superposition: 15406
% Backward superposition: 12469
% Self superposition: 265
% Equality factoring: 150
% Equality resolution: 38
% 
% Split clauses: 70
% Split components: 148
% SAT solver clauses: 6196
% SAT solver unit clauses: 777
% SAT solver binary clauses: 2184
% 
% Sat splits: 70
% Sat splitting refutations: 398
% 
% Memory used [KB]: 33261
% Time elapsed: 2.454 s
% ------------------------------
% ------------------------------
% Success in time 3.781 s
