Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f8904,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f8846,f7984,f378,f660,f377,f662,f658,f376,f659,f661,f657,f8743])).
fof(f8743,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f8742,f658])).
fof(f8742,plain,(
  s(fun(sK14,bool),'const.pred_set.EMPTY_0') = s(fun(sK14,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK14,bool),bool),sK15))) | $spl2),
  inference(forward_demodulation,[],[f8738,f659])).
fof(f8738,plain,(
  s(fun(sK14,bool),'const.pred_set.EMPTY_0') = s(fun(sK14,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK14,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK14,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK14,bool),bool),'const.pred_set.EMPTY_0'))))) | $spl2),
  inference(trivial_inequality_removal,[],[f8689])).
fof(f8689,plain,(
  s(fun(fun(sK14,bool),bool),sK15) != s(fun(fun(sK14,bool),bool),sK15) | s(fun(sK14,bool),'const.pred_set.EMPTY_0') = s(fun(sK14,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK14,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK14,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK14,bool),bool),'const.pred_set.EMPTY_0'))))) | $spl2),
  inference(superposition,[],[f454,f659])).
fof(f454,plain,(
  ( ! [X0,X1] : (s(fun(fun(X0,bool),bool),X1) != s(fun(fun(X0,bool),bool),'const.pred_set.INSERT_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.pred_set.EMPTY_0'))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X1)))) )),
  inference(cnf_transformation,[],[f274])).
fof(f274,plain,(
  ! [X0,X1] : ((s(fun(X0,bool),'const.pred_set.EMPTY_0') != s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X1))) | s(fun(fun(X0,bool),bool),'const.pred_set.EMPTY_0') = s(fun(fun(X0,bool),bool),X1) | s(fun(fun(X0,bool),bool),X1) = s(fun(fun(X0,bool),bool),'const.pred_set.INSERT_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.pred_set.EMPTY_0')))) & ((s(fun(fun(X0,bool),bool),'const.pred_set.EMPTY_0') != s(fun(fun(X0,bool),bool),X1) & s(fun(fun(X0,bool),bool),X1) != s(fun(fun(X0,bool),bool),'const.pred_set.INSERT_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.pred_set.EMPTY_0')))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X1)))))),
  inference(flattening,[],[f273])).
fof(f273,plain,(
  ! [X0,X1] : ((s(fun(X0,bool),'const.pred_set.EMPTY_0') != s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X1))) | (s(fun(fun(X0,bool),bool),'const.pred_set.EMPTY_0') = s(fun(fun(X0,bool),bool),X1) | s(fun(fun(X0,bool),bool),X1) = s(fun(fun(X0,bool),bool),'const.pred_set.INSERT_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.pred_set.EMPTY_0'))))) & ((s(fun(fun(X0,bool),bool),'const.pred_set.EMPTY_0') != s(fun(fun(X0,bool),bool),X1) & s(fun(fun(X0,bool),bool),X1) != s(fun(fun(X0,bool),bool),'const.pred_set.INSERT_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.pred_set.EMPTY_0')))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X1)))))),
  inference(nnf_transformation,[],[f113])).
fof(f113,plain,(
  ! [X0,X1] : (s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X1))) <=> (s(fun(fun(X0,bool),bool),'const.pred_set.EMPTY_0') = s(fun(fun(X0,bool),bool),X1) | s(fun(fun(X0,bool),bool),X1) = s(fun(fun(X0,bool),bool),'const.pred_set.INSERT_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.pred_set.EMPTY_0')))))),
  inference(flattening,[],[f112])).
fof(f112,plain,(
  ! [X0] : ! [X1] : (s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X1))) <=> (s(fun(fun(X0,bool),bool),'const.pred_set.EMPTY_0') = s(fun(fun(X0,bool),bool),X1) | s(fun(fun(X0,bool),bool),X1) = s(fun(fun(X0,bool),bool),'const.pred_set.INSERT_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.pred_set.EMPTY_0')))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6] : ! [X7] : (s(fun(X6,bool),'const.pred_set.EMPTY_0') = s(fun(X6,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X6,bool),bool),X7))) <=> (s(fun(fun(X6,bool),bool),'const.pred_set.EMPTY_0') = s(fun(fun(X6,bool),bool),X7) | s(fun(fun(X6,bool),bool),X7) = s(fun(fun(X6,bool),bool),'const.pred_set.INSERT_2'(s(fun(X6,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X6,bool),bool),'const.pred_set.EMPTY_0')))))),
  file('thm.pred_set.BIGUNION_EQ_EMPTY|split|1','thm.pred_set.BIGUNION_EQ_EMPTY|split|0')).
fof(f657,plain,(
  s(fun(sK14,bool),'const.pred_set.EMPTY_0') = s(fun(sK14,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK14,bool),bool),sK15))) | $spl0),
  inference(cnf_transformation,[],[f657_D])).
fof(f657_D,plain,(
  s(fun(sK14,bool),'const.pred_set.EMPTY_0') = s(fun(sK14,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK14,bool),bool),sK15))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f661,plain,(
  s(fun(fun(sK14,bool),bool),'const.pred_set.EMPTY_0') = s(fun(fun(sK14,bool),bool),sK15) | $spl4),
  inference(cnf_transformation,[],[f661_D])).
fof(f661_D,plain,(
  s(fun(fun(sK14,bool),bool),'const.pred_set.EMPTY_0') = s(fun(fun(sK14,bool),bool),sK15) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f659,plain,(
  s(fun(fun(sK14,bool),bool),sK15) = s(fun(fun(sK14,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK14,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK14,bool),bool),'const.pred_set.EMPTY_0'))) | $spl2),
  inference(cnf_transformation,[],[f659_D])).
fof(f659_D,plain,(
  s(fun(fun(sK14,bool),bool),sK15) = s(fun(fun(sK14,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK14,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK14,bool),bool),'const.pred_set.EMPTY_0'))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f376,plain,(
  s(fun(fun(sK14,bool),bool),sK15) = s(fun(fun(sK14,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK14,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK14,bool),bool),'const.pred_set.EMPTY_0'))) | s(fun(fun(sK14,bool),bool),'const.pred_set.EMPTY_0') = s(fun(fun(sK14,bool),bool),sK15) | s(fun(sK14,bool),'const.pred_set.EMPTY_0') = s(fun(sK14,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK14,bool),bool),sK15)))),
  inference(cnf_transformation,[],[f235])).
fof(f235,plain,(
  (s(fun(sK14,bool),'const.pred_set.EMPTY_0') = s(fun(sK14,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK14,bool),bool),sK15))) | s(fun(fun(sK14,bool),bool),'const.pred_set.EMPTY_0') = s(fun(fun(sK14,bool),bool),sK15) | s(fun(fun(sK14,bool),bool),sK15) = s(fun(fun(sK14,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK14,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK14,bool),bool),'const.pred_set.EMPTY_0')))) & (s(fun(sK14,bool),'const.pred_set.EMPTY_0') != s(fun(sK14,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK14,bool),bool),sK15))) | (s(fun(fun(sK14,bool),bool),'const.pred_set.EMPTY_0') != s(fun(fun(sK14,bool),bool),sK15) & s(fun(fun(sK14,bool),bool),sK15) != s(fun(fun(sK14,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK14,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK14,bool),bool),'const.pred_set.EMPTY_0')))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK14,sK15])],[f234])).
fof(f234,plain,(
  ? [X0,X1] : ((s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X1))) | s(fun(fun(X0,bool),bool),'const.pred_set.EMPTY_0') = s(fun(fun(X0,bool),bool),X1) | s(fun(fun(X0,bool),bool),X1) = s(fun(fun(X0,bool),bool),'const.pred_set.INSERT_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.pred_set.EMPTY_0')))) & (s(fun(X0,bool),'const.pred_set.EMPTY_0') != s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X1))) | (s(fun(fun(X0,bool),bool),'const.pred_set.EMPTY_0') != s(fun(fun(X0,bool),bool),X1) & s(fun(fun(X0,bool),bool),X1) != s(fun(fun(X0,bool),bool),'const.pred_set.INSERT_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.pred_set.EMPTY_0'))))))),
  inference(flattening,[],[f233])).
fof(f233,plain,(
  ? [X0,X1] : ((s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X1))) | (s(fun(fun(X0,bool),bool),'const.pred_set.EMPTY_0') = s(fun(fun(X0,bool),bool),X1) | s(fun(fun(X0,bool),bool),X1) = s(fun(fun(X0,bool),bool),'const.pred_set.INSERT_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.pred_set.EMPTY_0'))))) & (s(fun(X0,bool),'const.pred_set.EMPTY_0') != s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X1))) | (s(fun(fun(X0,bool),bool),'const.pred_set.EMPTY_0') != s(fun(fun(X0,bool),bool),X1) & s(fun(fun(X0,bool),bool),X1) != s(fun(fun(X0,bool),bool),'const.pred_set.INSERT_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.pred_set.EMPTY_0'))))))),
  inference(nnf_transformation,[],[f174])).
fof(f174,plain,(
  ? [X0,X1] : (s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X1))) <~> (s(fun(fun(X0,bool),bool),'const.pred_set.EMPTY_0') = s(fun(fun(X0,bool),bool),X1) | s(fun(fun(X0,bool),bool),X1) = s(fun(fun(X0,bool),bool),'const.pred_set.INSERT_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.pred_set.EMPTY_0')))))),
  inference(ennf_transformation,[],[f72])).
fof(f72,plain,(
  ~! [X0,X1] : (s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X1))) <=> (s(fun(fun(X0,bool),bool),'const.pred_set.EMPTY_0') = s(fun(fun(X0,bool),bool),X1) | s(fun(fun(X0,bool),bool),X1) = s(fun(fun(X0,bool),bool),'const.pred_set.INSERT_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.pred_set.EMPTY_0')))))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ~! [X0] : ! [X1] : (s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X1))) <=> (s(fun(fun(X0,bool),bool),'const.pred_set.EMPTY_0') = s(fun(fun(X0,bool),bool),X1) | s(fun(fun(X0,bool),bool),X1) = s(fun(fun(X0,bool),bool),'const.pred_set.INSERT_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.pred_set.EMPTY_0')))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7] : (s(fun(X6,bool),'const.pred_set.EMPTY_0') = s(fun(X6,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X6,bool),bool),X7))) <=> (s(fun(fun(X6,bool),bool),'const.pred_set.EMPTY_0') = s(fun(fun(X6,bool),bool),X7) | s(fun(fun(X6,bool),bool),X7) = s(fun(fun(X6,bool),bool),'const.pred_set.INSERT_2'(s(fun(X6,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X6,bool),bool),'const.pred_set.EMPTY_0')))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7] : (s(fun(X6,bool),'const.pred_set.EMPTY_0') = s(fun(X6,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X6,bool),bool),X7))) <=> (s(fun(fun(X6,bool),bool),'const.pred_set.EMPTY_0') = s(fun(fun(X6,bool),bool),X7) | s(fun(fun(X6,bool),bool),X7) = s(fun(fun(X6,bool),bool),'const.pred_set.INSERT_2'(s(fun(X6,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X6,bool),bool),'const.pred_set.EMPTY_0')))))),
  file('thm.pred_set.BIGUNION_EQ_EMPTY|split|1','thm.pred_set.BIGUNION_EQ_EMPTY|split|1')).
fof(f658,plain,(
  s(fun(sK14,bool),'const.pred_set.EMPTY_0') != s(fun(sK14,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK14,bool),bool),sK15))) | $spl1),
  inference(cnf_transformation,[],[f658_D])).
fof(f658_D,plain,(
  s(fun(sK14,bool),'const.pred_set.EMPTY_0') != s(fun(sK14,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK14,bool),bool),sK15))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f662,plain,(
  s(fun(fun(sK14,bool),bool),'const.pred_set.EMPTY_0') != s(fun(fun(sK14,bool),bool),sK15) | $spl5),
  inference(cnf_transformation,[],[f662_D])).
fof(f662_D,plain,(
  s(fun(fun(sK14,bool),bool),'const.pred_set.EMPTY_0') != s(fun(fun(sK14,bool),bool),sK15) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f377,plain,(
  s(fun(fun(sK14,bool),bool),'const.pred_set.EMPTY_0') != s(fun(fun(sK14,bool),bool),sK15) | s(fun(sK14,bool),'const.pred_set.EMPTY_0') != s(fun(sK14,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK14,bool),bool),sK15)))),
  inference(cnf_transformation,[],[f235])).
fof(f660,plain,(
  s(fun(fun(sK14,bool),bool),sK15) != s(fun(fun(sK14,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK14,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK14,bool),bool),'const.pred_set.EMPTY_0'))) | $spl3),
  inference(cnf_transformation,[],[f660_D])).
fof(f660_D,plain,(
  s(fun(fun(sK14,bool),bool),sK15) != s(fun(fun(sK14,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK14,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK14,bool),bool),'const.pred_set.EMPTY_0'))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f378,plain,(
  s(fun(fun(sK14,bool),bool),sK15) != s(fun(fun(sK14,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK14,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK14,bool),bool),'const.pred_set.EMPTY_0'))) | s(fun(sK14,bool),'const.pred_set.EMPTY_0') != s(fun(sK14,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK14,bool),bool),sK15)))),
  inference(cnf_transformation,[],[f235])).
fof(f7984,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f7983,f660])).
fof(f7983,plain,(
  s(fun(fun(sK14,bool),bool),sK15) = s(fun(fun(sK14,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK14,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK14,bool),bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl5)),
  inference(subsumption_resolution,[],[f7982,f662])).
fof(f7982,plain,(
  s(fun(fun(sK14,bool),bool),'const.pred_set.EMPTY_0') = s(fun(fun(sK14,bool),bool),sK15) | s(fun(fun(sK14,bool),bool),sK15) = s(fun(fun(sK14,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK14,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK14,bool),bool),'const.pred_set.EMPTY_0'))) | $spl0),
  inference(trivial_inequality_removal,[],[f7980])).
fof(f7980,plain,(
  s(fun(sK14,bool),'const.pred_set.EMPTY_0') != s(fun(sK14,bool),'const.pred_set.EMPTY_0') | s(fun(fun(sK14,bool),bool),'const.pred_set.EMPTY_0') = s(fun(fun(sK14,bool),bool),sK15) | s(fun(fun(sK14,bool),bool),sK15) = s(fun(fun(sK14,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK14,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK14,bool),bool),'const.pred_set.EMPTY_0'))) | $spl0),
  inference(superposition,[],[f452,f657])).
fof(f452,plain,(
  ( ! [X0,X1] : (s(fun(X0,bool),'const.pred_set.EMPTY_0') != s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X1))) | s(fun(fun(X0,bool),bool),'const.pred_set.EMPTY_0') = s(fun(fun(X0,bool),bool),X1) | s(fun(fun(X0,bool),bool),X1) = s(fun(fun(X0,bool),bool),'const.pred_set.INSERT_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.pred_set.EMPTY_0')))) )),
  inference(cnf_transformation,[],[f274])).
fof(f8846,plain,(
  $false | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f8794,f386])).
fof(f386,plain,(
  ( ! [X0] : (s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.EMPTY_0')))) )),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0] : s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.EMPTY_0')))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : s(fun(X6,bool),'const.pred_set.EMPTY_0') = s(fun(X6,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X6,bool),bool),'const.pred_set.EMPTY_0')))),
  file('thm.pred_set.BIGUNION_EQ_EMPTY|split|1','thm.pred_set.BIGUNION_EMPTY')).
fof(f8794,plain,(
  s(fun(sK14,bool),'const.pred_set.EMPTY_0') != s(fun(sK14,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK14,bool),bool),'const.pred_set.EMPTY_0'))) | ($spl1 | $spl4)),
  inference(backward_demodulation,[],[f661,f658])).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 510
% Passive clauses: 3621
% Generated clauses: 8477
% Final active clauses: 498
% Final passive clauses: 2817
% Input formulas: 69
% Initial clauses: 272
% 
% Duplicate literals: 187
% Trivial inequalities: 75
% Fw subsumption resolutions: 494
% Fw demodulations: 262
% Bw demodulations: 4
% 
% Simple tautologies: 209
% Equational tautologies: 44
% Forward subsumptions: 3534
% 
% Binary resolution: 4916
% Factoring: 918
% Forward superposition: 631
% Backward superposition: 634
% Self superposition: 4
% Equality factoring: 2
% Equality resolution: 21
% 
% Split clauses: 162
% Split components: 327
% SAT solver clauses: 315
% SAT solver unit clauses: 9
% SAT solver binary clauses: 60
% SAT solver learnt clauses: 2
% SAT solver learnt literals: 1
% 
% Sat splits: 162
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 300
% TWLsolver calls for satisfiability: 219
% 
% Memory used [KB]: 4733
% Time elapsed: 0.283 s
% ------------------------------
% ------------------------------
% Success in time 0.279 s
