Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.relation
% (54190)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 690
% Passive clauses: 5295
% Generated clauses: 20654
% Final active clauses: 678
% Final passive clauses: 4269
% Input formulas: 33
% Initial clauses: 159
% 
% Duplicate literals: 899
% Trivial inequalities: 98
% Fw subsumption resolutions: 2019
% Fw demodulations: 779
% Bw demodulations: 12
% 
% Simple tautologies: 230
% Equational tautologies: 75
% Forward subsumptions: 7559
% 
% Binary resolution: 6163
% Factoring: 1295
% Forward superposition: 3013
% Backward superposition: 5806
% Self superposition: 112
% Equality factoring: 108
% Equality resolution: 19
% 
% Split clauses: 1108
% Split components: 2366
% SAT solver clauses: 1437
% SAT solver unit clauses: 14
% SAT solver binary clauses: 139
% 
% Sat splits: 1108
% 
% TWLsolver clauses: 1434
% TWLsolver calls for satisfiability: 393
% 
% Memory used [KB]: 10234
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.relation
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.relation
% SZS output start Proof for thm.relation
fof(f8995,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f8921,f8806,f197,f369,f2271,f2273,f372,f374,f243,f388,f499,f503,f199,f366,f218,f381,f8805,f8807,f200,f371,f241,f390,f489,f491,f198,f368,f214,f379,f8691,f8726,f2268,f2270,f242,f389])).
fof(f389,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.relation.SC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))))) ) | $spl24),
  inference(cnf_transformation,[],[f389_D])).
fof(f389_D,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.relation.SC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f242,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.relation.SC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.SC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2))))) & ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2))))) | p(s(bool,'const.relation.SC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))))),
  inference(flattening,[],[f144])).
fof(f144,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.SC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2)))))) & ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2))))) | p(s(bool,'const.relation.SC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))))),
  inference(nnf_transformation,[],[f63])).
fof(f63,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.relation.SC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) <=> (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2))))))),
  inference(flattening,[],[f62])).
fof(f62,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.relation.SC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) <=> (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2))))))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X6] : ! [X8,X9,X10] : (p(s(bool,'const.relation.SC_3'(s(fun(X6,fun(X6,bool)),X8),s(X6,X9),s(X6,X10)))) <=> (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,X9))),s(X6,X10)))) | p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,X10))),s(X6,X9))))))),
  file('thm.relation.SC_lifts_monotonicities','thm.relation.SC_DEF')).
fof(f2270,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)))))) | $spl933),
  inference(cnf_transformation,[],[f2270_D])).
fof(f2270_D,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)))))) <=> ~$spl933),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl933])])).
fof(f2268,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)))))) | ($spl9 | $spl24 | $spl200)),
  inference(subsumption_resolution,[],[f2221,f374])).
fof(f2221,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)))))) | ($spl24 | $spl200)),
  inference(superposition,[],[f389,f503])).
fof(f8726,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16))),s(sK13,sK17)))) | $spl3055),
  inference(cnf_transformation,[],[f8726_D])).
fof(f8726_D,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16))),s(sK13,sK17)))) <=> ~$spl3055),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3055])])).
fof(f8691,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16))),s(sK13,sK17)))) | ($spl4 | $spl933)),
  inference(resolution,[],[f2270,f369])).
fof(f379,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl12),
  inference(cnf_transformation,[],[f379_D])).
fof(f379_D,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f214,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f44])).
fof(f44,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X11] : (s(bool,'T_0') = s(bool,X11) <=> p(s(bool,X11)))),
  file('thm.relation.SC_lifts_monotonicities','thm.bool.EQ_CLAUSES|split|1')).
fof(f368,plain,(
  p(s(bool,'const.relation.SC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)))) | $spl2),
  inference(cnf_transformation,[],[f368_D])).
fof(f368_D,plain,(
  p(s(bool,'const.relation.SC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f198,plain,(
  p(s(bool,'const.relation.SC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17))))),
  inference(cnf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X5,X6] : (~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X5))),s(sK13,X6)))) | p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,X5))))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,X6))))))) & (p(s(bool,'const.relation.SC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)))) & ~p(s(bool,'const.relation.SC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13,sK14,sK15,sK16,sK17])],[f124])).
fof(f124,plain,(
  ? [X0,X1,X2] : (! [X5,X6] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X5))),s(X0,X6)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X0,X0),X1),s(X0,X5))))),s(X0,happ(s(fun(X0,X0),X1),s(X0,X6))))))) & ? [X3,X4] : (p(s(bool,'const.relation.SC_3'(s(fun(X0,fun(X0,bool)),X2),s(X0,X3),s(X0,X4)))) & ~p(s(bool,'const.relation.SC_3'(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X0,X0),X1),s(X0,X3))),s(X0,happ(s(fun(X0,X0),X1),s(X0,X4))))))))),
  inference(rectify,[],[f84])).
fof(f84,plain,(
  ? [X0,X1,X2] : (! [X3,X4] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X0,X0),X1),s(X0,X3))))),s(X0,happ(s(fun(X0,X0),X1),s(X0,X4))))))) & ? [X5,X6] : (p(s(bool,'const.relation.SC_3'(s(fun(X0,fun(X0,bool)),X2),s(X0,X5),s(X0,X6)))) & ~p(s(bool,'const.relation.SC_3'(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X0,X0),X1),s(X0,X5))),s(X0,happ(s(fun(X0,X0),X1),s(X0,X6))))))))),
  inference(ennf_transformation,[],[f36])).
fof(f36,plain,(
  ~! [X0,X1,X2] : (! [X3,X4] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X0,X0),X1),s(X0,X3))))),s(X0,happ(s(fun(X0,X0),X1),s(X0,X4))))))) => ! [X5,X6] : (p(s(bool,'const.relation.SC_3'(s(fun(X0,fun(X0,bool)),X2),s(X0,X5),s(X0,X6)))) => p(s(bool,'const.relation.SC_3'(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X0,X0),X1),s(X0,X5))),s(X0,happ(s(fun(X0,X0),X1),s(X0,X6))))))))),
  inference(flattening,[],[f35])).
fof(f35,plain,(
  ~! [X0] : ! [X1,X2] : (! [X3,X4] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X0,X0),X1),s(X0,X3))))),s(X0,happ(s(fun(X0,X0),X1),s(X0,X4))))))) => ! [X5,X6] : (p(s(bool,'const.relation.SC_3'(s(fun(X0,fun(X0,bool)),X2),s(X0,X5),s(X0,X6)))) => p(s(bool,'const.relation.SC_3'(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X0,X0),X1),s(X0,X5))),s(X0,happ(s(fun(X0,X0),X1),s(X0,X6))))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : (! [X9,X10] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,X9))),s(X6,X10)))) => p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,happ(s(fun(X6,X6),X7),s(X6,X9))))),s(X6,happ(s(fun(X6,X6),X7),s(X6,X10))))))) => ! [X9,X10] : (p(s(bool,'const.relation.SC_3'(s(fun(X6,fun(X6,bool)),X8),s(X6,X9),s(X6,X10)))) => p(s(bool,'const.relation.SC_3'(s(fun(X6,fun(X6,bool)),X8),s(X6,happ(s(fun(X6,X6),X7),s(X6,X9))),s(X6,happ(s(fun(X6,X6),X7),s(X6,X10))))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : (! [X9,X10] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,X9))),s(X6,X10)))) => p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,happ(s(fun(X6,X6),X7),s(X6,X9))))),s(X6,happ(s(fun(X6,X6),X7),s(X6,X10))))))) => ! [X9,X10] : (p(s(bool,'const.relation.SC_3'(s(fun(X6,fun(X6,bool)),X8),s(X6,X9),s(X6,X10)))) => p(s(bool,'const.relation.SC_3'(s(fun(X6,fun(X6,bool)),X8),s(X6,happ(s(fun(X6,X6),X7),s(X6,X9))),s(X6,happ(s(fun(X6,X6),X7),s(X6,X10))))))))),
  file('thm.relation.SC_lifts_monotonicities','thm.relation.SC_lifts_monotonicities')).
fof(f491,plain,(
  s(bool,'T_0') = s(bool,'const.relation.SC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17))) | $spl198),
  inference(cnf_transformation,[],[f491_D])).
fof(f491_D,plain,(
  s(bool,'T_0') = s(bool,'const.relation.SC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17))) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f489,plain,(
  s(bool,'T_0') = s(bool,'const.relation.SC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17))) | ($spl2 | $spl12)),
  inference(resolution,[],[f379,f368])).
fof(f390,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2)))) | ~p(s(bool,'const.relation.SC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))) ) | $spl26),
  inference(cnf_transformation,[],[f390_D])).
fof(f390_D,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2)))) | ~p(s(bool,'const.relation.SC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f241,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | ~p(s(bool,'const.relation.SC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f145])).
fof(f371,plain,(
  p(s(bool,'T_0')) | $spl6),
  inference(cnf_transformation,[],[f371_D])).
fof(f371_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f200,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.relation.SC_lifts_monotonicities',aHL_TRUTH)).
fof(f8807,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,sK16)))) | $spl3116),
  inference(cnf_transformation,[],[f8807_D])).
fof(f8807_D,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,sK16)))) <=> ~$spl3116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3116])])).
fof(f8805,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,sK16)))) | ($spl6 | $spl26 | $spl198 | $spl3055)),
  inference(subsumption_resolution,[],[f8804,f371])).
fof(f8804,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,sK16)))) | ($spl26 | $spl198 | $spl3055)),
  inference(forward_demodulation,[],[f8768,f491])).
fof(f8768,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,sK16)))) | ~p(s(bool,'const.relation.SC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)))) | ($spl26 | $spl3055)),
  inference(resolution,[],[f8726,f390])).
fof(f381,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) | $spl16),
  inference(cnf_transformation,[],[f381_D])).
fof(f381_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f218,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f46])).
fof(f46,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X11] : (s(bool,'F_0') = s(bool,X11) <=> ~p(s(bool,X11)))),
  file('thm.relation.SC_lifts_monotonicities','thm.bool.EQ_CLAUSES|split|3')).
fof(f366,plain,(
  ~p(s(bool,'const.relation.SC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)))))) | $spl1),
  inference(cnf_transformation,[],[f366_D])).
fof(f366_D,plain,(
  ~p(s(bool,'const.relation.SC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f199,plain,(
  ~p(s(bool,'const.relation.SC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17))))))),
  inference(cnf_transformation,[],[f125])).
fof(f503,plain,(
  s(bool,'F_0') = s(bool,'const.relation.SC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17))))) | $spl200),
  inference(cnf_transformation,[],[f503_D])).
fof(f503_D,plain,(
  s(bool,'F_0') = s(bool,'const.relation.SC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17))))) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f499,plain,(
  s(bool,'F_0') = s(bool,'const.relation.SC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17))))) | ($spl1 | $spl16)),
  inference(resolution,[],[f381,f366])).
fof(f388,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.relation.SC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2))))) ) | $spl22),
  inference(cnf_transformation,[],[f388_D])).
fof(f388_D,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.relation.SC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2))))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f243,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.relation.SC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2))))) )),
  inference(cnf_transformation,[],[f145])).
fof(f374,plain,(
  ~p(s(bool,'F_0')) | $spl9),
  inference(cnf_transformation,[],[f374_D])).
fof(f374_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f372,plain,(
  ~p(s(bool,'F_0'))),
  inference(global_subsumption,[],[f203,f200])).
fof(f203,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f126])).
fof(f126,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f37])).
fof(f37,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f15])).
fof(f15,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.relation.SC_lifts_monotonicities','thm.bool.NOT_CLAUSES|split|2')).
fof(f2273,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17))))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16)))))) | $spl935),
  inference(cnf_transformation,[],[f2273_D])).
fof(f2273_D,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17))))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16)))))) <=> ~$spl935),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl935])])).
fof(f2271,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17))))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16)))))) | ($spl9 | $spl22 | $spl200)),
  inference(subsumption_resolution,[],[f2222,f374])).
fof(f2222,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17))))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16)))))) | ($spl22 | $spl200)),
  inference(superposition,[],[f388,f503])).
fof(f369,plain,(
  ( ! [X6,X5] : (p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,X5))))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,X6)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X5))),s(sK13,X6))))) ) | $spl4),
  inference(cnf_transformation,[],[f369_D])).
fof(f369_D,plain,(
  ( ! [X6,X5] : (p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,X5))))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,X6)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X5))),s(sK13,X6))))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f197,plain,(
  ( ! [X6,X5] : (p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,X5))))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,X6)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X5))),s(sK13,X6))))) )),
  inference(cnf_transformation,[],[f125])).
fof(f8806,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,sK16)))) | $spl3117),
  inference(cnf_transformation,[],[f8806_D])).
fof(f8806_D,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,sK16)))) <=> ~$spl3117),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3117])])).
fof(f8921,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK17))),s(sK13,sK16)))) | ($spl4 | $spl935)),
  inference(resolution,[],[f2273,f369])).
% SZS output end Proof for thm.relation
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 377
% Passive clauses: 1780
% Generated clauses: 8969
% Final active clauses: 353
% Final passive clauses: 1148
% Input formulas: 33
% Initial clauses: 159
% 
% Duplicate literals: 188
% Trivial inequalities: 19
% Fw subsumption resolutions: 912
% Fw demodulations: 483
% Bw demodulations: 266
% Condensations: 219
% Global subsumptions: 13
% 
% Simple tautologies: 209
% Equational tautologies: 16
% Forward subsumptions: 2961
% Bw demodulations to eq. taut.: 4
% 
% Binary resolution: 2436
% Factoring: 301
% Forward superposition: 1248
% Backward superposition: 783
% Self superposition: 28
% Equality factoring: 26
% Equality resolution: 6
% 
% Split clauses: 71
% Split components: 147
% SAT solver clauses: 2243
% SAT solver unit clauses: 180
% SAT solver binary clauses: 72
% SAT solver learnt clauses: 27
% SAT solver learnt literals: 10
% 
% Sat splits: 71
% Sat splitting refutations: 9
% 
% TWLsolver clauses: 2166
% TWLsolver calls for satisfiability: 1726
% 
% Memory used [KB]: 8955
% Time elapsed: 0.525 s
% ------------------------------
% ------------------------------
% Success in time 1.146 s
