Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.quotient
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.quotient
% SZS output start Proof for thm.quotient
fof(f7662,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f7580,f76,f116,f78,f118,f156,f322,f361,f359])).
fof(f359,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK2,bool),happ(s(fun(sK2,fun(sK2,bool)),sK10),s(sK2,sK4))),s(sK2,sK3))) | $spl10),
  inference(cnf_transformation,[],[f359_D])).
fof(f359_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK2,bool),happ(s(fun(sK2,fun(sK2,bool)),sK10),s(sK2,sK4))),s(sK2,sK3))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f361,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK2,bool),happ(s(fun(sK2,fun(sK2,bool)),sK10),s(sK2,sK4))),s(sK2,sK3))) | $spl5),
  inference(subsumption_resolution,[],[f356,f118])).
fof(f356,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,happ(s(fun(sK2,bool),happ(s(fun(sK2,fun(sK2,bool)),sK10),s(sK2,sK4))),s(sK2,sK3)))),
  inference(superposition,[],[f74,f79])).
fof(f79,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.quotient.quotient_47_prove_ltb',aHL_BOOL_CASES)).
fof(f74,plain,(
  p(s(bool,happ(s(fun(sK2,bool),happ(s(fun(sK2,fun(sK2,bool)),sK10),s(sK2,sK4))),s(sK2,sK3))))),
  inference(cnf_transformation,[],[f54])).
fof(f54,plain,(
  p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK2,fun(sK2,bool)),sK10),s(fun(sK2,sK0),sK8),s(fun(sK0,sK2),sK5)))) & p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK9),s(sK1,happ(s(fun(sK2,sK1),sK7),s(sK2,sK4))))),s(sK1,happ(s(fun(sK2,sK1),sK6),s(sK2,sK3)))))) & p(s(bool,happ(s(fun(sK2,bool),happ(s(fun(sK2,fun(sK2,bool)),sK10),s(sK2,sK4))),s(sK2,sK3)))) & ~p(s(bool,'const.bool.IN_2'(s(sK2,sK4),s(fun(sK2,bool),'const.quotient.respects_1'(s(fun(sK2,fun(sK2,bool)),sK10))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK9,sK10])],[f51])).
fof(f51,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X2,fun(X2,bool)),X10),s(fun(X2,X0),X8),s(fun(X0,X2),X5)))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X9),s(X1,happ(s(fun(X2,X1),X7),s(X2,X4))))),s(X1,happ(s(fun(X2,X1),X6),s(X2,X3)))))) & p(s(bool,happ(s(fun(X2,bool),happ(s(fun(X2,fun(X2,bool)),X10),s(X2,X4))),s(X2,X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X2,X4),s(fun(X2,bool),'const.quotient.respects_1'(s(fun(X2,fun(X2,bool)),X10)))))))),
  inference(flattening,[],[f50])).
fof(f50,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X2,fun(X2,bool)),X10),s(fun(X2,X0),X8),s(fun(X0,X2),X5)))) & (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X9),s(X1,happ(s(fun(X2,X1),X7),s(X2,X4))))),s(X1,happ(s(fun(X2,X1),X6),s(X2,X3)))))) & (p(s(bool,happ(s(fun(X2,bool),happ(s(fun(X2,fun(X2,bool)),X10),s(X2,X4))),s(X2,X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X2,X4),s(fun(X2,bool),'const.quotient.respects_1'(s(fun(X2,fun(X2,bool)),X10)))))))))),
  inference(ennf_transformation,[],[f24])).
fof(f24,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X2,fun(X2,bool)),X10),s(fun(X2,X0),X8),s(fun(X0,X2),X5)))) => (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X9),s(X1,happ(s(fun(X2,X1),X7),s(X2,X4))))),s(X1,happ(s(fun(X2,X1),X6),s(X2,X3)))))) => (p(s(bool,happ(s(fun(X2,bool),happ(s(fun(X2,fun(X2,bool)),X10),s(X2,X4))),s(X2,X3)))) => p(s(bool,'const.bool.IN_2'(s(X2,X4),s(fun(X2,bool),'const.quotient.respects_1'(s(fun(X2,fun(X2,bool)),X10)))))))))),
  inference(flattening,[],[f23])).
fof(f23,plain,(
  ~! [X0,X1,X2] : ! [X3,X4,X5,X6,X7,X8,X9,X10] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X2,fun(X2,bool)),X10),s(fun(X2,X0),X8),s(fun(X0,X2),X5)))) => (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X9),s(X1,happ(s(fun(X2,X1),X7),s(X2,X4))))),s(X1,happ(s(fun(X2,X1),X6),s(X2,X3)))))) => (p(s(bool,happ(s(fun(X2,bool),happ(s(fun(X2,fun(X2,bool)),X10),s(X2,X4))),s(X2,X3)))) => p(s(bool,'const.bool.IN_2'(s(X2,X4),s(fun(X2,bool),'const.quotient.respects_1'(s(fun(X2,fun(X2,bool)),X10)))))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : ! [X9,X10,X11,X12,X13,X14,X15,X16] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X8,fun(X8,bool)),X16),s(fun(X8,X6),X14),s(fun(X6,X8),X11)))) => (p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X7,fun(X7,bool)),X15),s(X7,happ(s(fun(X8,X7),X13),s(X8,X10))))),s(X7,happ(s(fun(X8,X7),X12),s(X8,X9)))))) => (p(s(bool,happ(s(fun(X8,bool),happ(s(fun(X8,fun(X8,bool)),X16),s(X8,X10))),s(X8,X9)))) => p(s(bool,'const.bool.IN_2'(s(X8,X10),s(fun(X8,bool),'const.quotient.respects_1'(s(fun(X8,fun(X8,bool)),X16)))))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : ! [X9,X10,X11,X12,X13,X14,X15,X16] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X8,fun(X8,bool)),X16),s(fun(X8,X6),X14),s(fun(X6,X8),X11)))) => (p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X7,fun(X7,bool)),X15),s(X7,happ(s(fun(X8,X7),X13),s(X8,X10))))),s(X7,happ(s(fun(X8,X7),X12),s(X8,X9)))))) => (p(s(bool,happ(s(fun(X8,bool),happ(s(fun(X8,fun(X8,bool)),X16),s(X8,X10))),s(X8,X9)))) => p(s(bool,'const.bool.IN_2'(s(X8,X10),s(fun(X8,bool),'const.quotient.respects_1'(s(fun(X8,fun(X8,bool)),X16)))))))))),
  file('thm.quotient.quotient_47_prove_ltb','thm.quotient.quotient_47_prove_ltb')).
fof(f322,plain,(
  s(bool,'T_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK2,fun(sK2,bool)),sK10),s(fun(sK2,sK0),sK8),s(fun(sK0,sK2),sK5))) | $spl8),
  inference(cnf_transformation,[],[f322_D])).
fof(f322_D,plain,(
  s(bool,'T_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK2,fun(sK2,bool)),sK10),s(fun(sK2,sK0),sK8),s(fun(sK0,sK2),sK5))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f156,plain,(
  s(bool,'T_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK2,fun(sK2,bool)),sK10),s(fun(sK2,sK0),sK8),s(fun(sK0,sK2),sK5)))),
  inference(resolution,[],[f72,f81])).
fof(f81,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f26])).
fof(f26,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X17] : (s(bool,'T_0') = s(bool,X17) <=> p(s(bool,X17)))),
  file('thm.quotient.quotient_47_prove_ltb','thm.bool.EQ_CLAUSES|split|1')).
fof(f72,plain,(
  p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK2,fun(sK2,bool)),sK10),s(fun(sK2,sK0),sK8),s(fun(sK0,sK2),sK5))))),
  inference(cnf_transformation,[],[f54])).
fof(f118,plain,(
  ~p(s(bool,'F_0')) | $spl5),
  inference(cnf_transformation,[],[f118_D])).
fof(f118_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f78,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f25])).
fof(f25,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.quotient.quotient_47_prove_ltb',aHL_FALSITY)).
fof(f116,plain,(
  p(s(bool,'T_0')) | $spl2),
  inference(cnf_transformation,[],[f116_D])).
fof(f116_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f76,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.quotient.quotient_47_prove_ltb',aHL_TRUTH)).
fof(f7580,plain,(
  $false | ($spl2 | $spl5 | $spl8 | $spl10)),
  inference(subsumption_resolution,[],[f7493,f120])).
fof(f120,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl5),
  inference(resolution,[],[f80,f118])).
fof(f80,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f55])).
fof(f7493,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl2 | $spl5 | $spl8 | $spl10)),
  inference(superposition,[],[f7303,f322])).
fof(f7303,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK2,fun(sK2,bool)),sK10),s(fun(sK2,X0),X1),s(fun(X0,sK2),X2)))) ) | ($spl2 | $spl5 | $spl10)),
  inference(subsumption_resolution,[],[f7302,f118])).
fof(f7302,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK2,fun(sK2,bool)),sK10),s(fun(sK2,X0),X1),s(fun(X0,sK2),X2)))) ) | ($spl2 | $spl10)),
  inference(forward_demodulation,[],[f7301,f1096])).
fof(f1096,plain,(
  s(bool,'F_0') = s(bool,'const.combin.W_2'(s(fun(sK2,fun(sK2,bool)),sK10),s(sK2,sK4)))),
  inference(superposition,[],[f809,f450])).
fof(f450,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK2,bool),'const.quotient.respects_1'(s(fun(sK2,fun(sK2,bool)),sK10))),s(sK2,sK4)))),
  inference(superposition,[],[f103,f152])).
fof(f152,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,sK4),s(fun(sK2,bool),'const.quotient.respects_1'(s(fun(sK2,fun(sK2,bool)),sK10)))))),
  inference(resolution,[],[f75,f85])).
fof(f85,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f29])).
fof(f29,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f28])).
fof(f28,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X17] : (s(bool,'F_0') = s(bool,X17) <=> ~p(s(bool,X17)))),
  file('thm.quotient.quotient_47_prove_ltb','thm.bool.EQ_CLAUSES|split|2')).
fof(f75,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK2,sK4),s(fun(sK2,bool),'const.quotient.respects_1'(s(fun(sK2,fun(sK2,bool)),sK10))))))),
  inference(cnf_transformation,[],[f54])).
fof(f103,plain,(
  ( ! [X2,X0,X1] : (s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) = s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) )),
  inference(cnf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0,X1,X2] : s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) = s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))),
  inference(flattening,[],[f37])).
fof(f37,plain,(
  ! [X0] : ! [X1,X2] : s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) = s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X8] : ! [X19,X10] : s(bool,happ(s(fun(X8,bool),X19),s(X8,X10))) = s(bool,'const.bool.IN_2'(s(X8,X10),s(fun(X8,bool),X19)))),
  file('thm.quotient.quotient_47_prove_ltb','thm.pred_set.SPECIFICATION')).
fof(f809,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.combin.W_2'(s(fun(X1,fun(X1,X0)),X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),'const.quotient.respects_1'(s(fun(X1,fun(X1,X0)),X2))),s(X1,X3)))) )),
  inference(forward_demodulation,[],[f808,f105])).
fof(f105,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,X1),'const.quotient.respects_1'(s(fun(X0,fun(X0,X1)),X2))) = s(fun(X0,X1),happ(s(fun(fun(X0,fun(X0,X1)),fun(X0,X1)),'const.quotient.respects_0'),s(fun(X0,fun(X0,X1)),X2)))) )),
  inference(cnf_transformation,[],[f41])).
fof(f41,plain,(
  ! [X0,X1,X2] : s(fun(X0,X1),'const.quotient.respects_1'(s(fun(X0,fun(X0,X1)),X2))) = s(fun(X0,X1),happ(s(fun(fun(X0,fun(X0,X1)),fun(X0,X1)),'const.quotient.respects_0'),s(fun(X0,fun(X0,X1)),X2)))),
  inference(flattening,[],[f40])).
fof(f40,plain,(
  ! [X0,X1] : ! [X2] : s(fun(X0,X1),'const.quotient.respects_1'(s(fun(X0,fun(X0,X1)),X2))) = s(fun(X0,X1),happ(s(fun(fun(X0,fun(X0,X1)),fun(X0,X1)),'const.quotient.respects_0'),s(fun(X0,fun(X0,X1)),X2)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X8,X7] : ! [X25] : s(fun(X8,X7),'const.quotient.respects_1'(s(fun(X8,fun(X8,X7)),X25))) = s(fun(X8,X7),happ(s(fun(fun(X8,fun(X8,X7)),fun(X8,X7)),'const.quotient.respects_0'),s(fun(X8,fun(X8,X7)),X25)))),
  file('thm.quotient.quotient_47_prove_ltb','ARITY_0_1_const.quotient.respects')).
fof(f808,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.combin.W_2'(s(fun(X1,fun(X1,X0)),X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(fun(X1,fun(X1,X0)),fun(X1,X0)),'const.quotient.respects_0'),s(fun(X1,fun(X1,X0)),X2))),s(X1,X3)))) )),
  inference(forward_demodulation,[],[f107,f104])).
fof(f104,plain,(
  ( ! [X0,X1] : (s(fun(fun(X0,fun(X0,X1)),fun(X0,X1)),'const.quotient.respects_0') = s(fun(fun(X0,fun(X0,X1)),fun(X0,X1)),'const.combin.W_0')) )),
  inference(cnf_transformation,[],[f39])).
fof(f39,plain,(
  ! [X0,X1] : s(fun(fun(X0,fun(X0,X1)),fun(X0,X1)),'const.quotient.respects_0') = s(fun(fun(X0,fun(X0,X1)),fun(X0,X1)),'const.combin.W_0')),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X8,X7] : s(fun(fun(X8,fun(X8,X7)),fun(X8,X7)),'const.quotient.respects_0') = s(fun(fun(X8,fun(X8,X7)),fun(X8,X7)),'const.combin.W_0')),
  file('thm.quotient.quotient_47_prove_ltb','thm.quotient.respects_def')).
fof(f107,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.combin.W_2'(s(fun(X1,fun(X1,X0)),X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(fun(X1,fun(X1,X0)),fun(X1,X0)),'const.combin.W_0'),s(fun(X1,fun(X1,X0)),X2))),s(X1,X3)))) )),
  inference(cnf_transformation,[],[f45])).
fof(f45,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.combin.W_2'(s(fun(X1,fun(X1,X0)),X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(fun(X1,fun(X1,X0)),fun(X1,X0)),'const.combin.W_0'),s(fun(X1,fun(X1,X0)),X2))),s(X1,X3)))),
  inference(flattening,[],[f44])).
fof(f44,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X0,'const.combin.W_2'(s(fun(X1,fun(X1,X0)),X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(fun(X1,fun(X1,X0)),fun(X1,X0)),'const.combin.W_0'),s(fun(X1,fun(X1,X0)),X2))),s(X1,X3)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X7,X8] : ! [X25,X26] : s(X7,'const.combin.W_2'(s(fun(X8,fun(X8,X7)),X25),s(X8,X26))) = s(X7,happ(s(fun(X8,X7),happ(s(fun(fun(X8,fun(X8,X7)),fun(X8,X7)),'const.combin.W_0'),s(fun(X8,fun(X8,X7)),X25))),s(X8,X26)))),
  file('thm.quotient.quotient_47_prove_ltb','ARITY_0_2_const.combin.W')).
fof(f7301,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.combin.W_2'(s(fun(sK2,fun(sK2,bool)),sK10),s(sK2,sK4)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK2,fun(sK2,bool)),sK10),s(fun(sK2,X0),X1),s(fun(X0,sK2),X2)))) ) | ($spl2 | $spl10)),
  inference(subsumption_resolution,[],[f7246,f116])).
fof(f7246,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.combin.W_2'(s(fun(sK2,fun(sK2,bool)),sK10),s(sK2,sK4)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK2,fun(sK2,bool)),sK10),s(fun(sK2,X0),X1),s(fun(X0,sK2),X2)))) ) | $spl10),
  inference(superposition,[],[f1373,f359])).
fof(f1373,plain,(
  ( ! [X14,X12,X10,X8,X13,X11,X9] : (~p(s(bool,happ(s(fun(X8,bool),happ(s(fun(X8,fun(X8,bool)),X9),s(X8,X10))),s(X8,X11)))) | p(s(bool,'const.combin.W_2'(s(fun(X8,fun(X8,bool)),X9),s(X8,X10)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X8,fun(X8,bool)),X9),s(fun(X8,X12),X13),s(fun(X12,X8),X14)))) )),
  inference(resolution,[],[f1198,f85])).
fof(f1198,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,X0),X3),s(fun(X0,X1),X4)))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X6)))) | p(s(bool,'const.combin.W_2'(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))))) )),
  inference(forward_demodulation,[],[f109,f106])).
fof(f106,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.combin.W_2'(s(fun(X1,fun(X1,X0)),X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,X3))),s(X1,X3)))) )),
  inference(cnf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.combin.W_2'(s(fun(X1,fun(X1,X0)),X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,X3))),s(X1,X3)))),
  inference(flattening,[],[f42])).
fof(f42,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X0,'const.combin.W_2'(s(fun(X1,fun(X1,X0)),X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,X3))),s(X1,X3)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X7,X8] : ! [X18,X10] : s(X7,'const.combin.W_2'(s(fun(X8,fun(X8,X7)),X18),s(X8,X10))) = s(X7,happ(s(fun(X8,X7),happ(s(fun(X8,fun(X8,X7)),X18),s(X8,X10))),s(X8,X10)))),
  file('thm.quotient.quotient_47_prove_ltb','thm.combin.W_THM')).
fof(f109,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X5)))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X6)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,X0),X3),s(fun(X0,X1),X4))))) )),
  inference(cnf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,X0),X3),s(fun(X0,X1),X4)))) | ! [X5,X6] : ((~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X6)))) | (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X5)))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X6))),s(X1,X6)))) & s(X0,happ(s(fun(X1,X0),X3),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X6))))) & (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X5)))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X6))),s(X1,X6)))) | s(X0,happ(s(fun(X1,X0),X3),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X6))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X6)))))))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,X0),X3),s(fun(X0,X1),X4)))) | ! [X5,X6] : ((~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X6)))) | (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X5)))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X6))),s(X1,X6)))) & s(X0,happ(s(fun(X1,X0),X3),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X6))))) & ((~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X5)))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X6))),s(X1,X6)))) | s(X0,happ(s(fun(X1,X0),X3),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X6)))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X6)))))))),
  inference(nnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,X0),X3),s(fun(X0,X1),X4)))) | ! [X5,X6] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X6)))) <=> (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X5)))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X6))),s(X1,X6)))) & s(X0,happ(s(fun(X1,X0),X3),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X6))))))),
  inference(ennf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,X0),X3),s(fun(X0,X1),X4)))) => ! [X5,X6] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X6)))) <=> (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X5)))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X6))),s(X1,X6)))) & s(X0,happ(s(fun(X1,X0),X3),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X6))))))),
  inference(flattening,[],[f48])).
fof(f48,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,X0),X3),s(fun(X0,X1),X4)))) => ! [X5,X6] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X6)))) <=> (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X5)))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X6))),s(X1,X6)))) & s(X0,happ(s(fun(X1,X0),X3),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X6))))))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X7,X8] : ! [X20,X21,X22] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X8,fun(X8,bool)),X20),s(fun(X8,X7),X21),s(fun(X7,X8),X22)))) => ! [X23,X24] : (p(s(bool,happ(s(fun(X8,bool),happ(s(fun(X8,fun(X8,bool)),X20),s(X8,X23))),s(X8,X24)))) <=> (p(s(bool,happ(s(fun(X8,bool),happ(s(fun(X8,fun(X8,bool)),X20),s(X8,X23))),s(X8,X23)))) & p(s(bool,happ(s(fun(X8,bool),happ(s(fun(X8,fun(X8,bool)),X20),s(X8,X24))),s(X8,X24)))) & s(X7,happ(s(fun(X8,X7),X21),s(X8,X23))) = s(X7,happ(s(fun(X8,X7),X21),s(X8,X24))))))),
  file('thm.quotient.quotient_47_prove_ltb','thm.quotient.QUOTIENT_REL')).
% SZS output end Proof for thm.quotient
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 99
% Passive clauses: 972
% Generated clauses: 7485
% Final active clauses: 93
% Final passive clauses: 752
% Input formulas: 21
% Initial clauses: 41
% 
% Duplicate literals: 390
% Trivial inequalities: 90
% Fw subsumption resolutions: 794
% Fw demodulations: 1583
% Bw demodulations: 6
% 
% Simple tautologies: 194
% Equational tautologies: 19
% Forward subsumptions: 3253
% 
% Binary resolution: 147
% Forward superposition: 2632
% Backward superposition: 1652
% Self superposition: 58
% Equality factoring: 34
% Equality resolution: 15
% 
% Split clauses: 111
% Split components: 228
% SAT solver clauses: 313
% SAT solver unit clauses: 12
% SAT solver binary clauses: 64
% SAT solver learnt clauses: 1
% 
% Sat splits: 111
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 306
% TWLsolver calls for satisfiability: 106
% 
% Memory used [KB]: 4733
% Time elapsed: 0.196 s
% ------------------------------
% ------------------------------
% Success in time 0.236 s
