Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.integer
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.integer
% SZS output start Proof for thm.integer
fof(f8495,plain,(
  $false),
  inference(subsumption_resolution,[],[f8494,f249])).
fof(f249,plain,(
  ( ! [X2] : (s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X2)))) )),
  inference(cnf_transformation,[],[f155])).
fof(f155,plain,(
  ! [X2] : s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X2))) & ! [X1] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X1) | s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12])],[f112])).
fof(f112,plain,(
  ? [X0] : (! [X2] : s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X2))) & ! [X1] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X1) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))))),
  inference(ennf_transformation,[],[f46])).
fof(f46,plain,(
  ~! [X0] : (? [X2] : s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X2))) | ? [X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X1) & s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))))),
  inference(flattening,[],[f45])).
fof(f45,plain,(
  ~! [X0] : (? [X2] : s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X2))) | ? [X1] : (~s('type.num.num','const.num.|48|_0') = s('type.num.num',X1) & s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : (? [X7] : s('type.integer.int',X6) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))) | ? [X7] : (~s('type.num.num','const.num.|48|_0') = s('type.num.num',X7) & s('type.integer.int',X6) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : (? [X7] : s('type.integer.int',X6) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))) | ? [X7] : (~s('type.num.num','const.num.|48|_0') = s('type.num.num',X7) & s('type.integer.int',X6) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7)))))))),
  file('thm.integer.integer_25|48|_prove_ltb','thm.integer.integer_25|48|_prove_ltb')).
fof(f8494,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),
  inference(subsumption_resolution,[],[f8493,f249])).
fof(f8493,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13(sK12))))),
  inference(equality_resolution,[],[f4548])).
fof(f4548,plain,(
  ( ! [X0] : (s('type.integer.int',sK12) != s('type.integer.int',X0) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13(X0))))) )),
  inference(subsumption_resolution,[],[f4543,f261])).
fof(f261,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',sK14(X0)) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13(X0))))) )),
  inference(cnf_transformation,[],[f158])).
fof(f158,plain,(
  ! [X0] : ((s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13(X0)))) & s('type.num.num','const.num.|48|_0') != s('type.num.num',sK13(X0))) | (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK14(X0)))))) & s('type.num.num','const.num.|48|_0') != s('type.num.num',sK14(X0))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13,sK14])],[f51])).
fof(f51,plain,(
  ! [X0] : (? [X2] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X2))) & s('type.num.num','const.num.|48|_0') != s('type.num.num',X2)) | ? [X1] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))))) & s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(flattening,[],[f50])).
fof(f50,plain,(
  ! [X0] : (? [X2] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X2))) & ~s('type.num.num','const.num.|48|_0') = s('type.num.num',X2)) | ? [X1] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))))) & ~s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X19] : (? [X7] : (s('type.integer.int',X19) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))) & ~s('type.num.num','const.num.|48|_0') = s('type.num.num',X7)) | ? [X7] : (s('type.integer.int',X19) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))) & ~s('type.num.num','const.num.|48|_0') = s('type.num.num',X7)) | s('type.integer.int',X19) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  file('thm.integer.integer_25|48|_prove_ltb','thm.integer.INT_NUM_CASES')).
fof(f4543,plain,(
  ( ! [X0] : (s('type.integer.int',sK12) != s('type.integer.int',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK14(X0)) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13(X0))))) )),
  inference(superposition,[],[f250,f259])).
fof(f259,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK14(X0)))))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13(X0))))) )),
  inference(cnf_transformation,[],[f158])).
fof(f250,plain,(
  ( ! [X1] : (s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f155])).
% SZS output end Proof for thm.integer
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 448
% Passive clauses: 3342
% Generated clauses: 8226
% Final active clauses: 448
% Final passive clauses: 2668
% Input formulas: 43
% Initial clauses: 193
% 
% Duplicate literals: 220
% Trivial inequalities: 37
% Fw subsumption resolutions: 723
% 
% Simple tautologies: 194
% Equational tautologies: 54
% Forward subsumptions: 3553
% 
% Binary resolution: 4782
% Factoring: 906
% Forward superposition: 361
% Backward superposition: 870
% Self superposition: 22
% Equality factoring: 26
% Equality resolution: 18
% 
% Split clauses: 162
% Split components: 325
% SAT solver clauses: 337
% SAT solver unit clauses: 10
% SAT solver binary clauses: 65
% 
% Sat splits: 162
% 
% TWLsolver clauses: 337
% TWLsolver calls for satisfiability: 222
% 
% Memory used [KB]: 3965
% Time elapsed: 0.221 s
% ------------------------------
% ------------------------------
% Success in time 0.224 s
