Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.extreal
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.extreal
% SZS output start Proof for thm.extreal
fof(f13855,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f13794,f204,f349,f351,f208,f352,f12967,f353,f12771,f3841,f12772,f13691,f12769,f12770,f12768,f13756])).
fof(f13756,plain,(
  $false | $spl154),
  inference(equality_resolution,[],[f12770])).
fof(f12768,plain,(
  ( ! [X104] : (~p(s(bool,happ(s(fun('type.realax.real',bool),X104),s('type.realax.real',sK11))))) ) | $spl152),
  inference(cnf_transformation,[],[f12768_D])).
fof(f12768_D,plain,(
  ( ! [X104] : (~p(s(bool,happ(s(fun('type.realax.real',bool),X104),s('type.realax.real',sK11))))) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f12770,plain,(
  ( ! [X105] : (s(bool,'F_0') != s(bool,X105)) ) | $spl154),
  inference(cnf_transformation,[],[f12770_D])).
fof(f12770_D,plain,(
  ( ! [X105] : (s(bool,'F_0') != s(bool,X105)) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f12769,plain,(
  ( ! [X107,X105] : (s(bool,'F_0') != s(bool,X105) | ~p(s(bool,happ(s(fun('type.realax.real',bool),X107),s('type.realax.real',sK11))))) )),
  inference(forward_demodulation,[],[f12644,f250])).
fof(f250,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X1) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(X0,X1),s(X0,X2),s(fun('type.realax.real',X0),X3)))) )),
  inference(cnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0,X1,X2,X3] : s(X0,X1) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(X0,X1),s(X0,X2),s(fun('type.realax.real',X0),X3)))),
  inference(flattening,[],[f76])).
fof(f76,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,X1) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(X0,X1),s(X0,X2),s(fun('type.realax.real',X0),X3)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X10] : ! [X14,X15,X16] : s(X10,X14) = s(X10,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(X10,X14),s(X10,X15),s(fun('type.realax.real',X10),X16)))),
  file('thm.extreal.extreal_not_infty|split|0','thm.extreal.extreal_case_def|split|0')).
fof(f12644,plain,(
  ( ! [X107,X105,X106] : (~p(s(bool,happ(s(fun('type.realax.real',bool),X107),s('type.realax.real',sK11)))) | s(bool,'F_0') != s(bool,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(bool,X105),s(bool,X106),s(fun('type.realax.real',bool),X107)))) )),
  inference(superposition,[],[f223,f1901])).
fof(f1901,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun('type.realax.real',X0),X1),s('type.realax.real',sK11))) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(X0,X2),s(X0,X3),s(fun('type.realax.real',X0),X1)))) )),
  inference(superposition,[],[f252,f202])).
fof(f202,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK11)))),
  inference(cnf_transformation,[],[f135])).
fof(f135,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK11)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK11])],[f100])).
fof(f100,plain,(
  ? [X0] : s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))),
  inference(ennf_transformation,[],[f40])).
fof(f40,plain,(
  ~! [X0] : s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ~! [X0] : ~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X6)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X6)))),
  file('thm.extreal.extreal_not_infty|split|0','thm.extreal.extreal_not_infty|split|0')).
fof(f252,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,happ(s(fun('type.realax.real',X0),X4),s('type.realax.real',X1))) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1))),s(X0,X2),s(X0,X3),s(fun('type.realax.real',X0),X4)))) )),
  inference(cnf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0,X1,X2,X3,X4] : s(X0,happ(s(fun('type.realax.real',X0),X4),s('type.realax.real',X1))) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1))),s(X0,X2),s(X0,X3),s(fun('type.realax.real',X0),X4)))),
  inference(flattening,[],[f80])).
fof(f80,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s(X0,happ(s(fun('type.realax.real',X0),X4),s('type.realax.real',X1))) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1))),s(X0,X2),s(X0,X3),s(fun('type.realax.real',X0),X4)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X10] : ! [X9,X14,X15,X16] : s(X10,happ(s(fun('type.realax.real',X10),X16),s('type.realax.real',X9))) = s(X10,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X9))),s(X10,X14),s(X10,X15),s(fun('type.realax.real',X10),X16)))),
  file('thm.extreal.extreal_not_infty|split|0','thm.extreal.extreal_case_def|split|2')).
fof(f223,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f141])).
fof(f141,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f53])).
fof(f53,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X21] : (s(bool,'F_0') = s(bool,X21) <=> ~p(s(bool,X21)))),
  file('thm.extreal.extreal_not_infty|split|0','thm.bool.EQ_CLAUSES|split|2')).
fof(f13691,plain,(
  $false | ($spl152 | $spl156)),
  inference(subsumption_resolution,[],[f12772,f12768])).
fof(f12772,plain,(
  ( ! [X110] : (p(s(bool,happ(s(fun('type.realax.real',bool),X110),s('type.realax.real',sK11))))) ) | $spl156),
  inference(cnf_transformation,[],[f12772_D])).
fof(f12772_D,plain,(
  ( ! [X110] : (p(s(bool,happ(s(fun('type.realax.real',bool),X110),s('type.realax.real',sK11))))) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f3841,plain,(
  ( ! [X54] : (s(bool,'F_0') = s(bool,X54)) ) | $spl80),
  inference(cnf_transformation,[],[f3841_D])).
fof(f3841_D,plain,(
  ( ! [X54] : (s(bool,'F_0') = s(bool,X54)) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f12771,plain,(
  ( ! [X110,X108] : (s(bool,'F_0') = s(bool,X108) | p(s(bool,happ(s(fun('type.realax.real',bool),X110),s('type.realax.real',sK11))))) )),
  inference(forward_demodulation,[],[f12645,f250])).
fof(f12645,plain,(
  ( ! [X109,X110,X108] : (p(s(bool,happ(s(fun('type.realax.real',bool),X110),s('type.realax.real',sK11)))) | s(bool,'F_0') = s(bool,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(bool,X108),s(bool,X109),s(fun('type.realax.real',bool),X110)))) )),
  inference(superposition,[],[f224,f1901])).
fof(f224,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f141])).
fof(f353,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl4),
  inference(cnf_transformation,[],[f353_D])).
fof(f353_D,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f12967,plain,(
  ( ! [X2] : (~p(s(bool,X2))) ) | ($spl3 | $spl80)),
  inference(superposition,[],[f352,f3841])).
fof(f352,plain,(
  ~p(s(bool,'F_0')) | $spl3),
  inference(cnf_transformation,[],[f352_D])).
fof(f352_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f208,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f43])).
fof(f43,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.extreal.extreal_not_infty|split|0',aHL_FALSITY)).
fof(f351,plain,(
  p(s(bool,'F_0')) | $spl2),
  inference(cnf_transformation,[],[f351_D])).
fof(f351_D,plain,(
  p(s(bool,'F_0')) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f349,plain,(
  p(s(bool,'T_0')) | $spl0),
  inference(cnf_transformation,[],[f349_D])).
fof(f349_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f204,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f136])).
fof(f136,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f41])).
fof(f41,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f34])).
fof(f34,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.extreal.extreal_not_infty|split|0','thm.bool.NOT_CLAUSES|split|2')).
fof(f13794,plain,(
  $false | ($spl0 | $spl4)),
  inference(resolution,[],[f353,f349])).
% SZS output end Proof for thm.extreal
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 274
% Passive clauses: 1687
% Generated clauses: 13587
% Final active clauses: 214
% Final passive clauses: 1010
% Input formulas: 37
% Initial clauses: 143
% 
% Duplicate literals: 471
% Trivial inequalities: 91
% Fw subsumption resolutions: 1387
% Fw demodulations: 2291
% Bw demodulations: 36
% 
% Simple tautologies: 249
% Equational tautologies: 27
% Forward subsumptions: 7129
% 
% Binary resolution: 693
% Factoring: 4
% Forward superposition: 3561
% Backward superposition: 4592
% Self superposition: 85
% Equality factoring: 44
% Equality resolution: 12
% 
% Split clauses: 224
% Split components: 448
% SAT solver clauses: 375
% SAT solver unit clauses: 6
% SAT solver binary clauses: 77
% SAT solver learnt clauses: 1
% 
% Sat splits: 224
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 363
% TWLsolver calls for satisfiability: 93
% 
% Memory used [KB]: 7164
% Time elapsed: 0.272 s
% ------------------------------
% ------------------------------
% Success in time 0.28 s
