Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.measure
% (76843)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 143
% Passive clauses: 2212
% Generated clauses: 19374
% Final active clauses: 132
% Final passive clauses: 1932
% Input formulas: 24
% Initial clauses: 53
% 
% Duplicate literals: 974
% Trivial inequalities: 426
% Fw subsumption resolutions: 3389
% Fw demodulations: 2068
% Bw demodulations: 7
% 
% Simple tautologies: 524
% Equational tautologies: 59
% Forward subsumptions: 7695
% 
% Binary resolution: 280
% Forward superposition: 5559
% Backward superposition: 6288
% Self superposition: 210
% Equality factoring: 169
% Equality resolution: 39
% 
% Split clauses: 1442
% Split components: 2999
% SAT solver clauses: 2164
% SAT solver unit clauses: 35
% SAT solver binary clauses: 267
% 
% Sat splits: 1442
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 2163
% TWLsolver calls for satisfiability: 197
% 
% Memory used [KB]: 9338
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.measure
% (76861)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 104
% Passive clauses: 514
% Generated clauses: 13744
% Final active clauses: 96
% Final passive clauses: 380
% Input formulas: 24
% Initial clauses: 53
% 
% Duplicate literals: 1046
% Trivial inequalities: 330
% Fw subsumption resolutions: 1369
% Fw demodulations: 983
% Bw demodulations: 28
% Condensations: 1540
% Global subsumptions: 4
% 
% Simple tautologies: 365
% Equational tautologies: 46
% Forward subsumptions: 5590
% Bw demodulations to eq. taut.: 8
% 
% Binary resolution: 57
% Forward superposition: 4133
% Backward superposition: 4449
% Self superposition: 208
% Equality factoring: 158
% Equality resolution: 34
% 
% Split clauses: 172
% Split components: 355
% SAT solver clauses: 998
% SAT solver unit clauses: 59
% SAT solver binary clauses: 37
% SAT solver learnt clauses: 8
% 
% Sat splits: 172
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 982
% TWLsolver calls for satisfiability: 261
% 
% Memory used [KB]: 8955
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.measure
% (76877)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 552
% Passive clauses: 5895
% Generated clauses: 146074
% Final active clauses: 145
% Final passive clauses: 1123
% Input formulas: 24
% Initial clauses: 53
% Discarded non-redundant clauses: 42
% 
% Duplicate literals: 7945
% Trivial inequalities: 1518
% Fw subsumption resolutions: 14758
% Fw demodulations: 16252
% Bw demodulations: 3655
% Condensations: 15311
% Global subsumptions: 804
% 
% Simple tautologies: 1820
% Equational tautologies: 2132
% Forward subsumptions: 77316
% Fw demodulations to eq. taut.: 79
% Bw demodulations to eq. taut.: 207
% 
% Unit resulting resolution: 9938
% Forward superposition: 38010
% Backward superposition: 24197
% Self superposition: 706
% Equality factoring: 371
% Equality resolution: 133
% 
% Split clauses: 408
% Split components: 835
% SAT solver clauses: 9540
% SAT solver unit clauses: 896
% SAT solver binary clauses: 1206
% 
% Sat splits: 408
% Sat splitting refutations: 2023
% 
% Memory used [KB]: 49636
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.measure
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.measure
% SZS output start Proof for thm.measure
fof(f3673,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f143,f144,f145,f146,f147,f174,f173,f148,f150,f159,f171,f204,f201,f223,f292,f293,f294,f295,f297,f325,f320,f319,f326,f328,f330,f331,f333,f335,f344,f375,f386,f388,f389,f390,f393,f401,f425,f586,f589,f590,f604,f606,f607,f608,f611,f619,f643,f645,f624,f650,f642,f641,f646,f888,f887,f886,f889,f895,f896,f1126,f1128,f1129,f1131,f1133,f1142,f1104,f1103,f1183,f587,f1072,f1184,f1070,f1185,f1186,f1188,f1189,f588,f1187,f272,f245,f265,f270,f271,f263,f262,f266,f1235,f1236,f1267,f1268,f1298,f1299,f1335,f1336,f1337,f1356,f1358,f1359,f1361,f1363,f1372,f1175,f1445,f1447,f1448,f1450,f1452,f1461,f1494,f1880,f1882,f1883,f1884,f1887,f1895,f1924,f1991,f1993,f1994,f1996,f1998,f2007,f2042,f2234,f2233,f2232,f2242,f2243,f2246,f2248,f2249,f2251,f2253,f2262,f2297,f2409,f2411,f2412,f2414,f2416,f2425,f2480,f2482,f2483,f2485,f2487,f2496,f2660,f2666,f2668,f2672,f2677,f2681,f2450,f2686,f2833,f2839,f2841,f2845,f2850,f2854,f2521,f2859,f2235,f2687,f2872,f2874,f2875,f2877,f2879,f2888,f2903,f2912,f2914,f2915,f2917,f2919,f2928,f2938,f2945,f2947,f2948,f2950,f2952,f2961,f2986,f2998,f3000,f3001,f3003,f3005,f3014,f3024,f3031,f3033,f3034,f3036,f3038,f3047,f3062,f3071,f3073,f3074,f3076,f3078,f3087,f3112,f3170,f3172,f3173,f3174,f3177,f3185,f3226,f3236,f3238,f3239,f3241,f3243,f3252,f3268,f3317,f3319,f3320,f3322,f3324,f3333,f3349,f3357,f3359,f3360,f3362,f3364,f3373,f3411,f3474,f3476,f3477,f3479,f3481,f3490,f3518,f3570,f3572,f3573,f172,f3575,f176,f3577,f175,f3586,f3614,f177,f3664])).
fof(f3664,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f3663,f132])).
fof(f132,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.measure.UNIV_SIGMA_ALGEBRA',aHL_TRUTH)).
fof(f3663,plain,(
  ~p(s(bool,'T_0')) | $spl30),
  inference(forward_demodulation,[],[f3662,f896])).
fof(f3662,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1('const.pred_set.UNIV_0','const.pred_set.POW_1'(s(fun(sK0,bool),'const.pred_set.UNIV_0')),fun(sK0,bool))),s(fun(sK0,bool),'const.pred_set.UNIV_0'))))),
  inference(forward_demodulation,[],[f3629,f141])).
fof(f141,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.pred_set.POW_1'(s(fun(X0,bool),X1)))))) )),
  inference(cnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0,X1,X2] : s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.pred_set.POW_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f53])).
fof(f53,plain,(
  ! [X0] : ! [X1,X2] : s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.pred_set.POW_1'(s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X6] : ! [X17,X18] : s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X18),s(fun(X6,bool),X17))) = s(bool,'const.bool.IN_2'(s(fun(X6,bool),X18),s(fun(fun(X6,bool),bool),'const.pred_set.POW_1'(s(fun(X6,bool),X17)))))),
  file('thm.measure.UNIV_SIGMA_ALGEBRA','thm.pred_set.IN_POW')).
fof(f3629,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK1('const.pred_set.UNIV_0','const.pred_set.POW_1'(s(fun(sK0,bool),'const.pred_set.UNIV_0')),fun(sK0,bool))),s(fun(fun(sK0,bool),bool),'const.pred_set.POW_1'(s(fun(sK0,bool),'const.pred_set.UNIV_0'))))))),
  inference(resolution,[],[f1986,f133])).
fof(f133,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),'const.pair.,_2'(s(fun(X0,bool),X1),s(fun(fun(X0,bool),bool),'const.pred_set.POW_1'(s(fun(X0,bool),X1))))))))) )),
  inference(cnf_transformation,[],[f44])).
fof(f44,plain,(
  ! [X0,X1] : p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),'const.pair.,_2'(s(fun(X0,bool),X1),s(fun(fun(X0,bool),bool),'const.pred_set.POW_1'(s(fun(X0,bool),X1))))))))),
  inference(flattening,[],[f43])).
fof(f43,plain,(
  ! [X0] : ! [X1] : p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),'const.pair.,_2'(s(fun(X0,bool),X1),s(fun(fun(X0,bool),bool),'const.pred_set.POW_1'(s(fun(X0,bool),X1))))))))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X6] : ! [X15] : p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),'const.pair.,_2'(s(fun(X6,bool),X15),s(fun(fun(X6,bool),bool),'const.pred_set.POW_1'(s(fun(X6,bool),X15))))))))),
  file('thm.measure.UNIV_SIGMA_ALGEBRA','thm.measure.POW_SIGMA_ALGEBRA')).
fof(f1986,plain,(
  ( ! [X8] : (~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(fun(sK0,bool),bool),X8)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK1('const.pred_set.UNIV_0',X8,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),X8))))) )),
  inference(trivial_inequality_removal,[],[f1983])).
fof(f1983,plain,(
  ( ! [X8] : (s(bool,'T_0') != s(bool,'T_0') | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(fun(sK0,bool),bool),X8)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK1('const.pred_set.UNIV_0',X8,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),X8))))) )),
  inference(superposition,[],[f1932,f127])).
fof(f127,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f40])).
fof(f40,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X7] : (s(bool,'T_0') = s(bool,X7) <=> p(s(bool,X7)))),
  file('thm.measure.UNIV_SIGMA_ALGEBRA','thm.bool.EQ_CLAUSES|split|0')).
fof(f1932,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK1('const.pred_set.UNIV_0',X0,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),X0))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(fun(sK0,bool),bool),X0))))))) )),
  inference(forward_demodulation,[],[f163,f684])).
fof(f684,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X8,X9),s(fun(X8,bool),'const.pred_set.UNIV_0')))) )),
  inference(trivial_inequality_removal,[],[f677])).
fof(f677,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X8,X9),s(fun(X8,bool),'const.pred_set.UNIV_0')))) )),
  inference(superposition,[],[f622,f136])).
fof(f136,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.measure.UNIV_SIGMA_ALGEBRA',aHL_BOOL_CASES)).
fof(f622,plain,(
  ( ! [X81,X82] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(X81,X82),s(fun(X81,bool),'const.pred_set.UNIV_0')))) )),
  inference(resolution,[],[f356,f130])).
fof(f130,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f42])).
fof(f42,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.measure.UNIV_SIGMA_ALGEBRA',aHL_FALSITY)).
fof(f356,plain,(
  ( ! [X83,X81,X82] : (p(s(bool,X81)) | s(bool,X81) != s(bool,'const.bool.IN_2'(s(X82,X83),s(fun(X82,bool),'const.pred_set.UNIV_0')))) )),
  inference(resolution,[],[f160,f138])).
fof(f138,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0'))))) )),
  inference(cnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0,X1] : p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0'))))),
  inference(flattening,[],[f49])).
fof(f49,plain,(
  ! [X0] : ! [X1] : p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0'))))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X6] : ! [X8] : p(s(bool,'const.bool.IN_2'(s(X6,X8),s(fun(X6,bool),'const.pred_set.UNIV_0'))))),
  file('thm.measure.UNIV_SIGMA_ALGEBRA','thm.pred_set.IN_UNIV')).
fof(f160,plain,(
  ( ! [X26,X27] : (~p(s(bool,X26)) | p(s(bool,X27)) | s(bool,X26) != s(bool,X27)) )),
  inference(resolution,[],[f90,f142])).
fof(f142,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f95])).
fof(f95,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f67])).
fof(f67,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f66])).
fof(f66,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f56])).
fof(f56,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f27])).
fof(f27,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X8,X13,X9,X14] : ((s(bool,X8) = s(bool,X13) & (p(s(bool,X13)) => s(bool,X9) = s(bool,X14))) => ((p(s(bool,X8)) => p(s(bool,X9))) <=> (p(s(bool,X13)) => p(s(bool,X14)))))),
  file('thm.measure.UNIV_SIGMA_ALGEBRA','thm.bool.IMP_CONG')).
fof(f90,plain,(
  ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.UNIV_0'))))))),
  inference(cnf_transformation,[],[f65])).
fof(f65,plain,(
  ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.UNIV_0'))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f55])).
fof(f55,plain,(
  ? [X0] : ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),'const.pair.,_2'(s(fun(X0,bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.pred_set.UNIV_0'))))))),
  inference(ennf_transformation,[],[f26])).
fof(f26,plain,(
  ~! [X0] : p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),'const.pair.,_2'(s(fun(X0,bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.pred_set.UNIV_0'))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),'const.pair.,_2'(s(fun(X6,bool),'const.pred_set.UNIV_0'),s(fun(fun(X6,bool),bool),'const.pred_set.UNIV_0'))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),'const.pair.,_2'(s(fun(X6,bool),'const.pred_set.UNIV_0'),s(fun(fun(X6,bool),bool),'const.pred_set.UNIV_0'))))))),
  file('thm.measure.UNIV_SIGMA_ALGEBRA','thm.measure.UNIV_SIGMA_ALGEBRA')).
fof(f163,plain,(
  ( ! [X0] : (~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(fun(sK0,bool),bool),X0)))))) | s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK1('const.pred_set.UNIV_0',X0,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),'const.pred_set.UNIV_0'))) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK1('const.pred_set.UNIV_0',X0,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),X0)))) )),
  inference(superposition,[],[f90,f140])).
fof(f140,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'const.bool.IN_2'(s(X0,sK1(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK1(X2,X1,X0)),s(fun(X0,bool),X2)))) )),
  inference(cnf_transformation,[],[f89])).
fof(f89,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (s(bool,'const.bool.IN_2'(s(X0,sK1(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK1(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1])],[f88])).
fof(f88,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(rectify,[],[f87])).
fof(f87,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(nnf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0,X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(flattening,[],[f51])).
fof(f51,plain,(
  ! [X0] : ! [X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X6] : ! [X16,X7] : (s(fun(X6,bool),X7) = s(fun(X6,bool),X16) <=> ! [X8] : s(bool,'const.bool.IN_2'(s(X6,X8),s(fun(X6,bool),X7))) = s(bool,'const.bool.IN_2'(s(X6,X8),s(fun(X6,bool),X16))))),
  file('thm.measure.UNIV_SIGMA_ALGEBRA','thm.pred_set.EXTENSION')).
fof(f177,plain,(
  ( ! [X15] : (p(s(bool,X15))) ) | $spl12),
  inference(cnf_transformation,[],[f177_D])).
fof(f177_D,plain,(
  ( ! [X15] : (p(s(bool,X15))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f3614,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2))) | s(bool,X4) != s(bool,'const.bool.IN_2'(s(X0,sK1(X2,'const.pred_set.UNIV_0',X0)),s(fun(X0,bool),X2))) | ~p(s(bool,X4))) ) | $spl82),
  inference(cnf_transformation,[],[f3614_D])).
fof(f3614_D,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2))) | s(bool,X4) != s(bool,'const.bool.IN_2'(s(X0,sK1(X2,'const.pred_set.UNIV_0',X0)),s(fun(X0,bool),X2))) | ~p(s(bool,X4))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f3586,plain,(
  ( ! [X76,X74,X72,X75,X73] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X72,X73),s(fun(X72,bool),X74))) | p(s(bool,X75)) | ~p(s(bool,X76)) | s(bool,X76) != s(bool,'const.bool.IN_2'(s(X72,sK1(X74,'const.pred_set.UNIV_0',X72)),s(fun(X72,bool),X74)))) )),
  inference(resolution,[],[f3549,f142])).
fof(f3549,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'const.bool.IN_2'(s(X28,sK1(X29,'const.pred_set.UNIV_0',X28)),s(fun(X28,bool),X29)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X28,X30),s(fun(X28,bool),X29)))) )),
  inference(trivial_inequality_removal,[],[f3546])).
fof(f3546,plain,(
  ( ! [X30,X28,X29] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X28,X30),s(fun(X28,bool),X29))) | ~p(s(bool,'const.bool.IN_2'(s(X28,sK1(X29,'const.pred_set.UNIV_0',X28)),s(fun(X28,bool),X29))))) )),
  inference(superposition,[],[f792,f127])).
fof(f792,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X3,sK1(X4,'const.pred_set.UNIV_0',X3)),s(fun(X3,bool),X4))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X3,X5),s(fun(X3,bool),X4)))) )),
  inference(forward_demodulation,[],[f708,f684])).
fof(f708,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X3,X5),s(fun(X3,bool),X4))) | s(bool,'const.bool.IN_2'(s(X3,sK1(X4,'const.pred_set.UNIV_0',X3)),s(fun(X3,bool),'const.pred_set.UNIV_0'))) != s(bool,'const.bool.IN_2'(s(X3,sK1(X4,'const.pred_set.UNIV_0',X3)),s(fun(X3,bool),X4)))) )),
  inference(superposition,[],[f684,f140])).
fof(f175,plain,(
  ( ! [X6,X8] : (~p(s(bool,X6)) | s(bool,X6) != s(bool,X8)) ) | $spl8),
  inference(cnf_transformation,[],[f175_D])).
fof(f175_D,plain,(
  ( ! [X6,X8] : (~p(s(bool,X6)) | s(bool,X6) != s(bool,X8)) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f3577,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X42,X43),s(fun(X42,bool),X44))) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X46) != s(bool,'const.bool.IN_2'(s(X42,sK1(X44,'const.pred_set.UNIV_0',X42)),s(fun(X42,bool),X44))) | s(bool,X45) != s(bool,X47)) )),
  inference(resolution,[],[f3549,f98])).
fof(f98,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f67])).
fof(f176,plain,(
  ( ! [X17] : (~p(s(bool,X17))) ) | $spl10),
  inference(cnf_transformation,[],[f176_D])).
fof(f176_D,plain,(
  ( ! [X17] : (~p(s(bool,X17))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f3575,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X30,X31),s(fun(X30,bool),X32))) | p(s(bool,X33)) | ~p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X34) != s(bool,'const.bool.IN_2'(s(X30,sK1(X32,'const.pred_set.UNIV_0',X30)),s(fun(X30,bool),X32)))) )),
  inference(resolution,[],[f3549,f97])).
fof(f97,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f67])).
fof(f172,plain,(
  ( ! [X2,X0] : (p(s(bool,X0)) | s(bool,X0) != s(bool,X2)) ) | $spl2),
  inference(cnf_transformation,[],[f172_D])).
fof(f172_D,plain,(
  ( ! [X2,X0] : (p(s(bool,X0)) | s(bool,X0) != s(bool,X2)) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f3573,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X18,X19),s(fun(X18,bool),X20))) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,'const.bool.IN_2'(s(X18,sK1(X20,'const.pred_set.UNIV_0',X18)),s(fun(X18,bool),X20)))) )),
  inference(resolution,[],[f3549,f96])).
fof(f96,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f67])).
fof(f3572,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X12,X13),s(fun(X12,bool),X14))) | ~p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X16) != s(bool,'const.bool.IN_2'(s(X12,sK1(X14,'const.pred_set.UNIV_0',X12)),s(fun(X12,bool),X14))) | s(bool,X15) != s(bool,X17)) )),
  inference(resolution,[],[f3549,f94])).
fof(f94,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f67])).
fof(f3570,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2))) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.bool.IN_2'(s(X0,sK1(X2,'const.pred_set.UNIV_0',X0)),s(fun(X0,bool),X2)))) )),
  inference(resolution,[],[f3549,f92])).
fof(f92,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f67])).
fof(f3518,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2))) | s(bool,X4) != s(bool,'const.bool.IN_2'(s(X0,sK1('const.pred_set.UNIV_0',X2,X0)),s(fun(X0,bool),X2))) | ~p(s(bool,X4))) ) | $spl80),
  inference(cnf_transformation,[],[f3518_D])).
fof(f3518_D,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2))) | s(bool,X4) != s(bool,'const.bool.IN_2'(s(X0,sK1('const.pred_set.UNIV_0',X2,X0)),s(fun(X0,bool),X2))) | ~p(s(bool,X4))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f3490,plain,(
  ( ! [X76,X74,X72,X75,X73] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X72,X73),s(fun(X72,bool),X74))) | p(s(bool,X75)) | ~p(s(bool,X76)) | s(bool,X76) != s(bool,'const.bool.IN_2'(s(X72,sK1('const.pred_set.UNIV_0',X74,X72)),s(fun(X72,bool),X74)))) )),
  inference(resolution,[],[f3469,f142])).
fof(f3469,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'const.bool.IN_2'(s(X28,sK1('const.pred_set.UNIV_0',X29,X28)),s(fun(X28,bool),X29)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X28,X30),s(fun(X28,bool),X29)))) )),
  inference(trivial_inequality_removal,[],[f3466])).
fof(f3466,plain,(
  ( ! [X30,X28,X29] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X28,X30),s(fun(X28,bool),X29))) | ~p(s(bool,'const.bool.IN_2'(s(X28,sK1('const.pred_set.UNIV_0',X29,X28)),s(fun(X28,bool),X29))))) )),
  inference(superposition,[],[f791,f127])).
fof(f791,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,sK1('const.pred_set.UNIV_0',X1,X0)),s(fun(X0,bool),X1))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) )),
  inference(forward_demodulation,[],[f707,f684])).
fof(f707,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1))) | s(bool,'const.bool.IN_2'(s(X0,sK1('const.pred_set.UNIV_0',X1,X0)),s(fun(X0,bool),'const.pred_set.UNIV_0'))) != s(bool,'const.bool.IN_2'(s(X0,sK1('const.pred_set.UNIV_0',X1,X0)),s(fun(X0,bool),X1)))) )),
  inference(superposition,[],[f684,f140])).
fof(f3481,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X42,X43),s(fun(X42,bool),X44))) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X46) != s(bool,'const.bool.IN_2'(s(X42,sK1('const.pred_set.UNIV_0',X44,X42)),s(fun(X42,bool),X44))) | s(bool,X45) != s(bool,X47)) )),
  inference(resolution,[],[f3469,f98])).
fof(f3479,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X30,X31),s(fun(X30,bool),X32))) | p(s(bool,X33)) | ~p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X34) != s(bool,'const.bool.IN_2'(s(X30,sK1('const.pred_set.UNIV_0',X32,X30)),s(fun(X30,bool),X32)))) )),
  inference(resolution,[],[f3469,f97])).
fof(f3477,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X18,X19),s(fun(X18,bool),X20))) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,'const.bool.IN_2'(s(X18,sK1('const.pred_set.UNIV_0',X20,X18)),s(fun(X18,bool),X20)))) )),
  inference(resolution,[],[f3469,f96])).
fof(f3476,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X12,X13),s(fun(X12,bool),X14))) | ~p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X16) != s(bool,'const.bool.IN_2'(s(X12,sK1('const.pred_set.UNIV_0',X14,X12)),s(fun(X12,bool),X14))) | s(bool,X15) != s(bool,X17)) )),
  inference(resolution,[],[f3469,f94])).
fof(f3474,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2))) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.bool.IN_2'(s(X0,sK1('const.pred_set.UNIV_0',X2,X0)),s(fun(X0,bool),X2)))) )),
  inference(resolution,[],[f3469,f92])).
fof(f3411,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) != s(bool,X5) | ~p(s(bool,X5)) | s(bool,'T_0') = s(bool,X3) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl78),
  inference(cnf_transformation,[],[f3411_D])).
fof(f3411_D,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) != s(bool,X5) | ~p(s(bool,X5)) | s(bool,'T_0') = s(bool,X3) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f3373,plain,(
  ( ! [X92,X90,X88,X93,X91,X89] : (s(bool,X88) != s(bool,X89) | s(bool,X88) != s(bool,X90) | s(bool,X90) != s(bool,X91) | s(bool,'T_0') = s(bool,X91) | p(s(bool,X92)) | ~p(s(bool,X93)) | s(bool,X89) != s(bool,X93)) ) | $spl46),
  inference(resolution,[],[f2042,f142])).
fof(f3364,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (s(bool,X49) != s(bool,X50) | s(bool,X49) != s(bool,X51) | s(bool,X51) != s(bool,X52) | s(bool,'T_0') = s(bool,X52) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X50) != s(bool,X54) | s(bool,X53) != s(bool,X55)) ) | $spl46),
  inference(resolution,[],[f2042,f98])).
fof(f3362,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X40] : (s(bool,X35) != s(bool,X36) | s(bool,X35) != s(bool,X37) | s(bool,X37) != s(bool,X38) | s(bool,'T_0') = s(bool,X38) | p(s(bool,X39)) | ~p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X36) != s(bool,X40)) ) | $spl46),
  inference(resolution,[],[f2042,f97])).
fof(f3360,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s(bool,X21) != s(bool,X22) | s(bool,X21) != s(bool,X23) | s(bool,X23) != s(bool,X24) | s(bool,'T_0') = s(bool,X24) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X22) != s(bool,X26)) ) | $spl46),
  inference(resolution,[],[f2042,f96])).
fof(f3359,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s(bool,X14) != s(bool,X15) | s(bool,X14) != s(bool,X16) | s(bool,X16) != s(bool,X17) | s(bool,'T_0') = s(bool,X17) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X15) != s(bool,X19) | s(bool,X18) != s(bool,X20)) ) | $spl46),
  inference(resolution,[],[f2042,f94])).
fof(f3357,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2) | s(bool,X2) != s(bool,X3) | s(bool,'T_0') = s(bool,X3) | p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X1) != s(bool,X5)) ) | $spl46),
  inference(resolution,[],[f2042,f92])).
fof(f3349,plain,(
  ( ! [X2,X0] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1(X0,'const.pred_set.UNIV_0',sK0)),s(fun(sK0,bool),X0))) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),X0),s(fun(fun(sK0,bool),bool),'const.pred_set.UNIV_0')))))) ) | $spl76),
  inference(cnf_transformation,[],[f3349_D])).
fof(f3349_D,plain,(
  ( ! [X2,X0] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1(X0,'const.pred_set.UNIV_0',sK0)),s(fun(sK0,bool),X0))) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),X0),s(fun(fun(sK0,bool),bool),'const.pred_set.UNIV_0')))))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f3333,plain,(
  ( ! [X41,X42,X40] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1(X40,'const.pred_set.UNIV_0',sK0)),s(fun(sK0,bool),X40))) | p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X42) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),X40),s(fun(fun(sK0,bool),bool),'const.pred_set.UNIV_0')))))) )),
  inference(resolution,[],[f1877,f142])).
fof(f1877,plain,(
  ( ! [X4] : (~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),X4),s(fun(fun(sK0,bool),bool),'const.pred_set.UNIV_0')))))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1(X4,'const.pred_set.UNIV_0',sK0)),s(fun(sK0,bool),X4)))) )),
  inference(trivial_inequality_removal,[],[f1869])).
fof(f1869,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,'T_0') | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),X4),s(fun(fun(sK0,bool),bool),'const.pred_set.UNIV_0')))))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1(X4,'const.pred_set.UNIV_0',sK0)),s(fun(sK0,bool),X4)))) )),
  inference(superposition,[],[f1633,f136])).
fof(f1633,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK1(X1,'const.pred_set.UNIV_0',sK0)),s(fun(sK0,bool),X1))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),X1),s(fun(fun(sK0,bool),bool),'const.pred_set.UNIV_0'))))))) )),
  inference(forward_demodulation,[],[f162,f684])).
fof(f162,plain,(
  ( ! [X1] : (~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),X1),s(fun(fun(sK0,bool),bool),'const.pred_set.UNIV_0')))))) | s(bool,'const.bool.IN_2'(s(sK0,sK1(X1,'const.pred_set.UNIV_0',sK0)),s(fun(sK0,bool),'const.pred_set.UNIV_0'))) != s(bool,'const.bool.IN_2'(s(sK0,sK1(X1,'const.pred_set.UNIV_0',sK0)),s(fun(sK0,bool),X1)))) )),
  inference(superposition,[],[f90,f140])).
fof(f3324,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1(X28,'const.pred_set.UNIV_0',sK0)),s(fun(sK0,bool),X28))) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X30) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),X28),s(fun(fun(sK0,bool),bool),'const.pred_set.UNIV_0'))))) | s(bool,X29) != s(bool,X31)) )),
  inference(resolution,[],[f1877,f98])).
fof(f3322,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1(X20,'const.pred_set.UNIV_0',sK0)),s(fun(sK0,bool),X20))) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),X20),s(fun(fun(sK0,bool),bool),'const.pred_set.UNIV_0')))))) )),
  inference(resolution,[],[f1877,f97])).
fof(f3320,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1(X12,'const.pred_set.UNIV_0',sK0)),s(fun(sK0,bool),X12))) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),X12),s(fun(fun(sK0,bool),bool),'const.pred_set.UNIV_0')))))) )),
  inference(resolution,[],[f1877,f96])).
fof(f3319,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1(X8,'const.pred_set.UNIV_0',sK0)),s(fun(sK0,bool),X8))) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),X8),s(fun(fun(sK0,bool),bool),'const.pred_set.UNIV_0'))))) | s(bool,X9) != s(bool,X11)) )),
  inference(resolution,[],[f1877,f94])).
fof(f3317,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1(X0,'const.pred_set.UNIV_0',sK0)),s(fun(sK0,bool),X0))) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),X0),s(fun(fun(sK0,bool),bool),'const.pred_set.UNIV_0')))))) )),
  inference(resolution,[],[f1877,f92])).
fof(f3268,plain,(
  ( ! [X2,X0] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1('const.pred_set.UNIV_0',X0,sK0)),s(fun(sK0,bool),X0))) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),X0),s(fun(fun(sK0,bool),bool),'const.pred_set.UNIV_0')))))) ) | $spl74),
  inference(cnf_transformation,[],[f3268_D])).
fof(f3268_D,plain,(
  ( ! [X2,X0] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1('const.pred_set.UNIV_0',X0,sK0)),s(fun(sK0,bool),X0))) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),X0),s(fun(fun(sK0,bool),bool),'const.pred_set.UNIV_0')))))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f3252,plain,(
  ( ! [X41,X42,X40] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1('const.pred_set.UNIV_0',X40,sK0)),s(fun(sK0,bool),X40))) | p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X42) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),X40),s(fun(fun(sK0,bool),bool),'const.pred_set.UNIV_0')))))) )),
  inference(resolution,[],[f1568,f142])).
fof(f1568,plain,(
  ( ! [X4] : (~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),X4),s(fun(fun(sK0,bool),bool),'const.pred_set.UNIV_0')))))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1('const.pred_set.UNIV_0',X4,sK0)),s(fun(sK0,bool),X4)))) )),
  inference(trivial_inequality_removal,[],[f1560])).
fof(f1560,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,'T_0') | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),X4),s(fun(fun(sK0,bool),bool),'const.pred_set.UNIV_0')))))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1('const.pred_set.UNIV_0',X4,sK0)),s(fun(sK0,bool),X4)))) )),
  inference(superposition,[],[f1444,f136])).
fof(f1444,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK1('const.pred_set.UNIV_0',X0,sK0)),s(fun(sK0,bool),X0))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),X0),s(fun(fun(sK0,bool),bool),'const.pred_set.UNIV_0'))))))) )),
  inference(forward_demodulation,[],[f161,f684])).
fof(f161,plain,(
  ( ! [X0] : (~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),X0),s(fun(fun(sK0,bool),bool),'const.pred_set.UNIV_0')))))) | s(bool,'const.bool.IN_2'(s(sK0,sK1('const.pred_set.UNIV_0',X0,sK0)),s(fun(sK0,bool),'const.pred_set.UNIV_0'))) != s(bool,'const.bool.IN_2'(s(sK0,sK1('const.pred_set.UNIV_0',X0,sK0)),s(fun(sK0,bool),X0)))) )),
  inference(superposition,[],[f90,f140])).
fof(f3243,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1('const.pred_set.UNIV_0',X28,sK0)),s(fun(sK0,bool),X28))) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X30) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),X28),s(fun(fun(sK0,bool),bool),'const.pred_set.UNIV_0'))))) | s(bool,X29) != s(bool,X31)) )),
  inference(resolution,[],[f1568,f98])).
fof(f3241,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1('const.pred_set.UNIV_0',X20,sK0)),s(fun(sK0,bool),X20))) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),X20),s(fun(fun(sK0,bool),bool),'const.pred_set.UNIV_0')))))) )),
  inference(resolution,[],[f1568,f97])).
fof(f3239,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1('const.pred_set.UNIV_0',X12,sK0)),s(fun(sK0,bool),X12))) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),X12),s(fun(fun(sK0,bool),bool),'const.pred_set.UNIV_0')))))) )),
  inference(resolution,[],[f1568,f96])).
fof(f3238,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1('const.pred_set.UNIV_0',X8,sK0)),s(fun(sK0,bool),X8))) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),X8),s(fun(fun(sK0,bool),bool),'const.pred_set.UNIV_0'))))) | s(bool,X9) != s(bool,X11)) )),
  inference(resolution,[],[f1568,f94])).
fof(f3236,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1('const.pred_set.UNIV_0',X0,sK0)),s(fun(sK0,bool),X0))) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),X0),s(fun(fun(sK0,bool),bool),'const.pred_set.UNIV_0')))))) )),
  inference(resolution,[],[f1568,f92])).
fof(f3226,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X2) != s(bool,X5) | p(s(bool,X5)) | s(bool,'F_0') = s(bool,X3) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl72),
  inference(cnf_transformation,[],[f3226_D])).
fof(f3226_D,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X2) != s(bool,X5) | p(s(bool,X5)) | s(bool,'F_0') = s(bool,X3) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f3185,plain,(
  ( ! [X88,X87,X85,X89,X86,X84] : (s(bool,X84) != s(bool,X85) | s(bool,X84) != s(bool,X86) | s(bool,X85) != s(bool,X87) | s(bool,'F_0') = s(bool,X87) | p(s(bool,X88)) | p(s(bool,X89)) | s(bool,X86) != s(bool,X89)) ) | $spl18),
  inference(resolution,[],[f1463,f142])).
fof(f1463,plain,(
  ( ! [X85,X83,X84,X82] : (p(s(bool,X84)) | s(bool,X82) != s(bool,X83) | s(bool,X82) != s(bool,X84) | s(bool,X83) != s(bool,X85) | s(bool,'F_0') = s(bool,X85)) ) | $spl18),
  inference(resolution,[],[f375,f340])).
fof(f340,plain,(
  ( ! [X48,X49] : (p(s(bool,X48)) | s(bool,X48) != s(bool,X49) | s(bool,'F_0') = s(bool,X49)) )),
  inference(resolution,[],[f160,f123])).
fof(f123,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f37])).
fof(f37,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f36])).
fof(f36,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X7] : (s(bool,'F_0') = s(bool,X7) <=> ~p(s(bool,X7)))),
  file('thm.measure.UNIV_SIGMA_ALGEBRA','thm.bool.EQ_CLAUSES|split|2')).
fof(f3177,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (s(bool,X49) != s(bool,X50) | s(bool,X49) != s(bool,X51) | s(bool,X50) != s(bool,X52) | s(bool,'F_0') = s(bool,X52) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X51) != s(bool,X53) | s(bool,X54) != s(bool,X55)) ) | $spl18),
  inference(resolution,[],[f1463,f98])).
fof(f3174,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (s(bool,X28) != s(bool,X29) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31) | s(bool,'F_0') = s(bool,X31) | p(s(bool,X32)) | ~p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X30) != s(bool,X34)) ) | $spl18),
  inference(resolution,[],[f1463,f97])).
fof(f3173,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s(bool,X21) != s(bool,X22) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,X24) | s(bool,'F_0') = s(bool,X24) | p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X23) != s(bool,X26)) ) | $spl18),
  inference(resolution,[],[f1463,f96])).
fof(f3172,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s(bool,X14) != s(bool,X15) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,X17) | s(bool,'F_0') = s(bool,X17) | p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X16) != s(bool,X18) | s(bool,X19) != s(bool,X20)) ) | $spl18),
  inference(resolution,[],[f1463,f94])).
fof(f3170,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,X3) | s(bool,'F_0') = s(bool,X3) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X2) != s(bool,X5)) ) | $spl18),
  inference(resolution,[],[f1463,f92])).
fof(f3112,plain,(
  ( ! [X2,X0] : (s(fun(bool,bool),'const.pred_set.UNIV_0') = s(fun(bool,bool),X0) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X0))) | p(s(bool,sK1('const.pred_set.UNIV_0',X0,bool)))) ) | $spl70),
  inference(cnf_transformation,[],[f3112_D])).
fof(f3112_D,plain,(
  ( ! [X2,X0] : (s(fun(bool,bool),'const.pred_set.UNIV_0') = s(fun(bool,bool),X0) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X0))) | p(s(bool,sK1('const.pred_set.UNIV_0',X0,bool)))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f3087,plain,(
  ( ! [X41,X42,X40] : (s(fun(bool,bool),'const.pred_set.UNIV_0') = s(fun(bool,bool),X40) | p(s(bool,sK1('const.pred_set.UNIV_0',X40,bool))) | p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X42) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X40)))) )),
  inference(resolution,[],[f2504,f142])).
fof(f2504,plain,(
  ( ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X4)))) | s(fun(bool,bool),'const.pred_set.UNIV_0') = s(fun(bool,bool),X4) | p(s(bool,sK1('const.pred_set.UNIV_0',X4,bool)))) )),
  inference(superposition,[],[f2477,f125])).
fof(f125,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f39])).
fof(f39,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f38])).
fof(f38,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X7] : (s(bool,'F_0') = s(bool,X7) <=> ~p(s(bool,X7)))),
  file('thm.measure.UNIV_SIGMA_ALGEBRA','thm.bool.EQ_CLAUSES|split|3')).
fof(f2477,plain,(
  ( ! [X15,X16] : (~p(s(bool,'const.bool.IN_2'(s(X15,sK1('const.pred_set.UNIV_0',X16,X15)),s(fun(X15,bool),X16)))) | s(fun(X15,bool),'const.pred_set.UNIV_0') = s(fun(X15,bool),X16)) )),
  inference(trivial_inequality_removal,[],[f2474])).
fof(f2474,plain,(
  ( ! [X15,X16] : (s(bool,'T_0') != s(bool,'T_0') | s(fun(X15,bool),'const.pred_set.UNIV_0') = s(fun(X15,bool),X16) | ~p(s(bool,'const.bool.IN_2'(s(X15,sK1('const.pred_set.UNIV_0',X16,X15)),s(fun(X15,bool),X16))))) )),
  inference(superposition,[],[f717,f127])).
fof(f717,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X8,sK1('const.pred_set.UNIV_0',X9,X8)),s(fun(X8,bool),X9))) | s(fun(X8,bool),'const.pred_set.UNIV_0') = s(fun(X8,bool),X9)) )),
  inference(superposition,[],[f140,f684])).
fof(f3078,plain,(
  ( ! [X30,X28,X31,X29] : (s(fun(bool,bool),'const.pred_set.UNIV_0') = s(fun(bool,bool),X28) | p(s(bool,sK1('const.pred_set.UNIV_0',X28,bool))) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X30) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X28))) | s(bool,X29) != s(bool,X31)) )),
  inference(resolution,[],[f2504,f98])).
fof(f3076,plain,(
  ( ! [X23,X21,X22,X20] : (s(fun(bool,bool),'const.pred_set.UNIV_0') = s(fun(bool,bool),X20) | p(s(bool,sK1('const.pred_set.UNIV_0',X20,bool))) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X20)))) )),
  inference(resolution,[],[f2504,f97])).
fof(f3074,plain,(
  ( ! [X14,X12,X15,X13] : (s(fun(bool,bool),'const.pred_set.UNIV_0') = s(fun(bool,bool),X12) | p(s(bool,sK1('const.pred_set.UNIV_0',X12,bool))) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X12)))) )),
  inference(resolution,[],[f2504,f96])).
fof(f3073,plain,(
  ( ! [X10,X8,X11,X9] : (s(fun(bool,bool),'const.pred_set.UNIV_0') = s(fun(bool,bool),X8) | p(s(bool,sK1('const.pred_set.UNIV_0',X8,bool))) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X8))) | s(bool,X9) != s(bool,X11)) )),
  inference(resolution,[],[f2504,f94])).
fof(f3071,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun(bool,bool),'const.pred_set.UNIV_0') = s(fun(bool,bool),X0) | p(s(bool,sK1('const.pred_set.UNIV_0',X0,bool))) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))) )),
  inference(resolution,[],[f2504,f92])).
fof(f3062,plain,(
  ( ! [X2,X0] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK1(X0,'const.pred_set.UNIV_0',sK0)),s(fun(sK0,bool),X0)))) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),X0),s(fun(fun(sK0,bool),bool),'const.pred_set.UNIV_0')))))) ) | $spl68),
  inference(cnf_transformation,[],[f3062_D])).
fof(f3062_D,plain,(
  ( ! [X2,X0] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK1(X0,'const.pred_set.UNIV_0',sK0)),s(fun(sK0,bool),X0)))) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),X0),s(fun(fun(sK0,bool),bool),'const.pred_set.UNIV_0')))))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f3047,plain,(
  ( ! [X41,X42,X40] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK1(X40,'const.pred_set.UNIV_0',sK0)),s(fun(sK0,bool),X40)))) | p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X42) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),X40),s(fun(fun(sK0,bool),bool),'const.pred_set.UNIV_0')))))) )),
  inference(resolution,[],[f1875,f142])).
fof(f1875,plain,(
  ( ! [X7] : (~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),X7),s(fun(fun(sK0,bool),bool),'const.pred_set.UNIV_0')))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK1(X7,'const.pred_set.UNIV_0',sK0)),s(fun(sK0,bool),X7))))) )),
  inference(trivial_inequality_removal,[],[f1872])).
fof(f1872,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,'T_0') | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),X7),s(fun(fun(sK0,bool),bool),'const.pred_set.UNIV_0')))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK1(X7,'const.pred_set.UNIV_0',sK0)),s(fun(sK0,bool),X7))))) )),
  inference(superposition,[],[f1633,f127])).
fof(f3038,plain,(
  ( ! [X30,X28,X31,X29] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK1(X28,'const.pred_set.UNIV_0',sK0)),s(fun(sK0,bool),X28)))) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X30) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),X28),s(fun(fun(sK0,bool),bool),'const.pred_set.UNIV_0'))))) | s(bool,X29) != s(bool,X31)) )),
  inference(resolution,[],[f1875,f98])).
fof(f3036,plain,(
  ( ! [X23,X21,X22,X20] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK1(X20,'const.pred_set.UNIV_0',sK0)),s(fun(sK0,bool),X20)))) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),X20),s(fun(fun(sK0,bool),bool),'const.pred_set.UNIV_0')))))) )),
  inference(resolution,[],[f1875,f97])).
fof(f3034,plain,(
  ( ! [X14,X12,X15,X13] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK1(X12,'const.pred_set.UNIV_0',sK0)),s(fun(sK0,bool),X12)))) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),X12),s(fun(fun(sK0,bool),bool),'const.pred_set.UNIV_0')))))) )),
  inference(resolution,[],[f1875,f96])).
fof(f3033,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK1(X8,'const.pred_set.UNIV_0',sK0)),s(fun(sK0,bool),X8)))) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),X8),s(fun(fun(sK0,bool),bool),'const.pred_set.UNIV_0'))))) | s(bool,X9) != s(bool,X11)) )),
  inference(resolution,[],[f1875,f94])).
fof(f3031,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK1(X0,'const.pred_set.UNIV_0',sK0)),s(fun(sK0,bool),X0)))) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),X0),s(fun(fun(sK0,bool),bool),'const.pred_set.UNIV_0')))))) )),
  inference(resolution,[],[f1875,f92])).
fof(f3024,plain,(
  ( ! [X2,X0] : (s(fun(bool,bool),'const.pred_set.UNIV_0') = s(fun(bool,bool),X0) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,sK1('const.pred_set.UNIV_0',X0,bool)) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X0))))) ) | $spl66),
  inference(cnf_transformation,[],[f3024_D])).
fof(f3024_D,plain,(
  ( ! [X2,X0] : (s(fun(bool,bool),'const.pred_set.UNIV_0') = s(fun(bool,bool),X0) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,sK1('const.pred_set.UNIV_0',X0,bool)) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X0))))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f3014,plain,(
  ( ! [X41,X42,X40] : (s(fun(bool,bool),'const.pred_set.UNIV_0') = s(fun(bool,bool),X40) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X40)))) | p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X42) != s(bool,sK1('const.pred_set.UNIV_0',X40,bool))) )),
  inference(resolution,[],[f2502,f142])).
fof(f2502,plain,(
  ( ! [X2] : (~p(s(bool,sK1('const.pred_set.UNIV_0',X2,bool))) | s(fun(bool,bool),'const.pred_set.UNIV_0') = s(fun(bool,bool),X2) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X2))))) )),
  inference(superposition,[],[f2477,f129])).
fof(f129,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f41])).
fof(f41,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X7] : (s(bool,'T_0') = s(bool,X7) <=> p(s(bool,X7)))),
  file('thm.measure.UNIV_SIGMA_ALGEBRA','thm.bool.EQ_CLAUSES|split|1')).
fof(f3005,plain,(
  ( ! [X30,X28,X31,X29] : (s(fun(bool,bool),'const.pred_set.UNIV_0') = s(fun(bool,bool),X28) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X28)))) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X30) != s(bool,sK1('const.pred_set.UNIV_0',X28,bool)) | s(bool,X29) != s(bool,X31)) )),
  inference(resolution,[],[f2502,f98])).
fof(f3003,plain,(
  ( ! [X23,X21,X22,X20] : (s(fun(bool,bool),'const.pred_set.UNIV_0') = s(fun(bool,bool),X20) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X20)))) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,sK1('const.pred_set.UNIV_0',X20,bool))) )),
  inference(resolution,[],[f2502,f97])).
fof(f3001,plain,(
  ( ! [X14,X12,X15,X13] : (s(fun(bool,bool),'const.pred_set.UNIV_0') = s(fun(bool,bool),X12) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X12)))) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,sK1('const.pred_set.UNIV_0',X12,bool))) )),
  inference(resolution,[],[f2502,f96])).
fof(f3000,plain,(
  ( ! [X10,X8,X11,X9] : (s(fun(bool,bool),'const.pred_set.UNIV_0') = s(fun(bool,bool),X8) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X8)))) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,sK1('const.pred_set.UNIV_0',X8,bool)) | s(bool,X9) != s(bool,X11)) )),
  inference(resolution,[],[f2502,f94])).
fof(f2998,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun(bool,bool),'const.pred_set.UNIV_0') = s(fun(bool,bool),X0) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,sK1('const.pred_set.UNIV_0',X0,bool))) )),
  inference(resolution,[],[f2502,f92])).
fof(f2986,plain,(
  ( ! [X2,X0] : (s(fun(bool,bool),'const.pred_set.UNIV_0') = s(fun(bool,bool),X0) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X0))) | p(s(bool,sK1(X0,'const.pred_set.UNIV_0',bool)))) ) | $spl64),
  inference(cnf_transformation,[],[f2986_D])).
fof(f2986_D,plain,(
  ( ! [X2,X0] : (s(fun(bool,bool),'const.pred_set.UNIV_0') = s(fun(bool,bool),X0) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X0))) | p(s(bool,sK1(X0,'const.pred_set.UNIV_0',bool)))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f2961,plain,(
  ( ! [X41,X42,X40] : (s(fun(bool,bool),'const.pred_set.UNIV_0') = s(fun(bool,bool),X40) | p(s(bool,sK1(X40,'const.pred_set.UNIV_0',bool))) | p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X42) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X40)))) )),
  inference(resolution,[],[f2433,f142])).
fof(f2433,plain,(
  ( ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X4)))) | s(fun(bool,bool),'const.pred_set.UNIV_0') = s(fun(bool,bool),X4) | p(s(bool,sK1(X4,'const.pred_set.UNIV_0',bool)))) )),
  inference(superposition,[],[f2401,f125])).
fof(f2401,plain,(
  ( ! [X15,X16] : (~p(s(bool,'const.bool.IN_2'(s(X15,sK1(X16,'const.pred_set.UNIV_0',X15)),s(fun(X15,bool),X16)))) | s(fun(X15,bool),'const.pred_set.UNIV_0') = s(fun(X15,bool),X16)) )),
  inference(trivial_inequality_removal,[],[f2398])).
fof(f2398,plain,(
  ( ! [X15,X16] : (s(bool,'T_0') != s(bool,'T_0') | s(fun(X15,bool),'const.pred_set.UNIV_0') = s(fun(X15,bool),X16) | ~p(s(bool,'const.bool.IN_2'(s(X15,sK1(X16,'const.pred_set.UNIV_0',X15)),s(fun(X15,bool),X16))))) )),
  inference(superposition,[],[f716,f127])).
fof(f716,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X6,sK1(X7,'const.pred_set.UNIV_0',X6)),s(fun(X6,bool),X7))) | s(fun(X6,bool),'const.pred_set.UNIV_0') = s(fun(X6,bool),X7)) )),
  inference(superposition,[],[f140,f684])).
fof(f2952,plain,(
  ( ! [X30,X28,X31,X29] : (s(fun(bool,bool),'const.pred_set.UNIV_0') = s(fun(bool,bool),X28) | p(s(bool,sK1(X28,'const.pred_set.UNIV_0',bool))) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X30) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X28))) | s(bool,X29) != s(bool,X31)) )),
  inference(resolution,[],[f2433,f98])).
fof(f2950,plain,(
  ( ! [X23,X21,X22,X20] : (s(fun(bool,bool),'const.pred_set.UNIV_0') = s(fun(bool,bool),X20) | p(s(bool,sK1(X20,'const.pred_set.UNIV_0',bool))) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X20)))) )),
  inference(resolution,[],[f2433,f97])).
fof(f2948,plain,(
  ( ! [X14,X12,X15,X13] : (s(fun(bool,bool),'const.pred_set.UNIV_0') = s(fun(bool,bool),X12) | p(s(bool,sK1(X12,'const.pred_set.UNIV_0',bool))) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X12)))) )),
  inference(resolution,[],[f2433,f96])).
fof(f2947,plain,(
  ( ! [X10,X8,X11,X9] : (s(fun(bool,bool),'const.pred_set.UNIV_0') = s(fun(bool,bool),X8) | p(s(bool,sK1(X8,'const.pred_set.UNIV_0',bool))) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X8))) | s(bool,X9) != s(bool,X11)) )),
  inference(resolution,[],[f2433,f94])).
fof(f2945,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun(bool,bool),'const.pred_set.UNIV_0') = s(fun(bool,bool),X0) | p(s(bool,sK1(X0,'const.pred_set.UNIV_0',bool))) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))) )),
  inference(resolution,[],[f2433,f92])).
fof(f2938,plain,(
  ( ! [X2,X0] : (s(fun(bool,bool),'const.pred_set.UNIV_0') = s(fun(bool,bool),X0) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,sK1(X0,'const.pred_set.UNIV_0',bool)) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X0))))) ) | $spl62),
  inference(cnf_transformation,[],[f2938_D])).
fof(f2938_D,plain,(
  ( ! [X2,X0] : (s(fun(bool,bool),'const.pred_set.UNIV_0') = s(fun(bool,bool),X0) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,sK1(X0,'const.pred_set.UNIV_0',bool)) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X0))))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f2928,plain,(
  ( ! [X41,X42,X40] : (s(fun(bool,bool),'const.pred_set.UNIV_0') = s(fun(bool,bool),X40) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X40)))) | p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X42) != s(bool,sK1(X40,'const.pred_set.UNIV_0',bool))) )),
  inference(resolution,[],[f2431,f142])).
fof(f2431,plain,(
  ( ! [X2] : (~p(s(bool,sK1(X2,'const.pred_set.UNIV_0',bool))) | s(fun(bool,bool),'const.pred_set.UNIV_0') = s(fun(bool,bool),X2) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X2))))) )),
  inference(superposition,[],[f2401,f129])).
fof(f2919,plain,(
  ( ! [X30,X28,X31,X29] : (s(fun(bool,bool),'const.pred_set.UNIV_0') = s(fun(bool,bool),X28) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X28)))) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X30) != s(bool,sK1(X28,'const.pred_set.UNIV_0',bool)) | s(bool,X29) != s(bool,X31)) )),
  inference(resolution,[],[f2431,f98])).
fof(f2917,plain,(
  ( ! [X23,X21,X22,X20] : (s(fun(bool,bool),'const.pred_set.UNIV_0') = s(fun(bool,bool),X20) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X20)))) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,sK1(X20,'const.pred_set.UNIV_0',bool))) )),
  inference(resolution,[],[f2431,f97])).
fof(f2915,plain,(
  ( ! [X14,X12,X15,X13] : (s(fun(bool,bool),'const.pred_set.UNIV_0') = s(fun(bool,bool),X12) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X12)))) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,sK1(X12,'const.pred_set.UNIV_0',bool))) )),
  inference(resolution,[],[f2431,f96])).
fof(f2914,plain,(
  ( ! [X10,X8,X11,X9] : (s(fun(bool,bool),'const.pred_set.UNIV_0') = s(fun(bool,bool),X8) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X8)))) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,sK1(X8,'const.pred_set.UNIV_0',bool)) | s(bool,X9) != s(bool,X11)) )),
  inference(resolution,[],[f2431,f94])).
fof(f2912,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun(bool,bool),'const.pred_set.UNIV_0') = s(fun(bool,bool),X0) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,sK1(X0,'const.pred_set.UNIV_0',bool))) )),
  inference(resolution,[],[f2431,f92])).
fof(f2903,plain,(
  ( ! [X2,X0] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK1('const.pred_set.UNIV_0',X0,sK0)),s(fun(sK0,bool),X0)))) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),X0),s(fun(fun(sK0,bool),bool),'const.pred_set.UNIV_0')))))) ) | $spl60),
  inference(cnf_transformation,[],[f2903_D])).
fof(f2903_D,plain,(
  ( ! [X2,X0] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK1('const.pred_set.UNIV_0',X0,sK0)),s(fun(sK0,bool),X0)))) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),X0),s(fun(fun(sK0,bool),bool),'const.pred_set.UNIV_0')))))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f2888,plain,(
  ( ! [X41,X42,X40] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK1('const.pred_set.UNIV_0',X40,sK0)),s(fun(sK0,bool),X40)))) | p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X42) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),X40),s(fun(fun(sK0,bool),bool),'const.pred_set.UNIV_0')))))) )),
  inference(resolution,[],[f1566,f142])).
fof(f1566,plain,(
  ( ! [X7] : (~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),X7),s(fun(fun(sK0,bool),bool),'const.pred_set.UNIV_0')))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK1('const.pred_set.UNIV_0',X7,sK0)),s(fun(sK0,bool),X7))))) )),
  inference(trivial_inequality_removal,[],[f1563])).
fof(f1563,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,'T_0') | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),X7),s(fun(fun(sK0,bool),bool),'const.pred_set.UNIV_0')))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK1('const.pred_set.UNIV_0',X7,sK0)),s(fun(sK0,bool),X7))))) )),
  inference(superposition,[],[f1444,f127])).
fof(f2879,plain,(
  ( ! [X30,X28,X31,X29] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK1('const.pred_set.UNIV_0',X28,sK0)),s(fun(sK0,bool),X28)))) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X30) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),X28),s(fun(fun(sK0,bool),bool),'const.pred_set.UNIV_0'))))) | s(bool,X29) != s(bool,X31)) )),
  inference(resolution,[],[f1566,f98])).
fof(f2877,plain,(
  ( ! [X23,X21,X22,X20] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK1('const.pred_set.UNIV_0',X20,sK0)),s(fun(sK0,bool),X20)))) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),X20),s(fun(fun(sK0,bool),bool),'const.pred_set.UNIV_0')))))) )),
  inference(resolution,[],[f1566,f97])).
fof(f2875,plain,(
  ( ! [X14,X12,X15,X13] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK1('const.pred_set.UNIV_0',X12,sK0)),s(fun(sK0,bool),X12)))) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),X12),s(fun(fun(sK0,bool),bool),'const.pred_set.UNIV_0')))))) )),
  inference(resolution,[],[f1566,f96])).
fof(f2874,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK1('const.pred_set.UNIV_0',X8,sK0)),s(fun(sK0,bool),X8)))) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),X8),s(fun(fun(sK0,bool),bool),'const.pred_set.UNIV_0'))))) | s(bool,X9) != s(bool,X11)) )),
  inference(resolution,[],[f1566,f94])).
fof(f2872,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK1('const.pred_set.UNIV_0',X0,sK0)),s(fun(sK0,bool),X0)))) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),X0),s(fun(fun(sK0,bool),bool),'const.pred_set.UNIV_0')))))) )),
  inference(resolution,[],[f1566,f92])).
fof(f2687,plain,(
  ( ! [X271,X272] : (s(fun(X271,bool),'const.pred_set.UNIV_0') = s(fun(X271,bool),X272)) ) | $spl58),
  inference(cnf_transformation,[],[f2687_D])).
fof(f2687_D,plain,(
  ( ! [X271,X272] : (s(fun(X271,bool),'const.pred_set.UNIV_0') = s(fun(X271,bool),X272)) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f2235,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X4,bool),fun(fun(X4,bool),bool)),'const.pair.,_2'(s(fun(X4,bool),X5),s(fun(fun(X4,bool),bool),'const.pred_set.POW_1'(s(fun(X4,bool),X5)))))))) ) | $spl48),
  inference(cnf_transformation,[],[f2235_D])).
fof(f2235_D,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X4,bool),fun(fun(X4,bool),bool)),'const.pair.,_2'(s(fun(X4,bool),X5),s(fun(fun(X4,bool),bool),'const.pred_set.POW_1'(s(fun(X4,bool),X5)))))))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f2859,plain,(
  ( ! [X273,X275,X272,X274] : (s(bool,'F_0') != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X274,bool),fun(fun(X274,bool),bool)),'const.pair.,_2'(s(fun(X274,bool),X275),s(fun(fun(X274,bool),bool),'const.pred_set.POW_1'(s(fun(X274,bool),X275))))))) | s(fun(X272,bool),'const.pred_set.UNIV_0') = s(fun(X272,bool),X273)) )),
  inference(subsumption_resolution,[],[f2798,f2477])).
fof(f2798,plain,(
  ( ! [X273,X275,X272,X274] : (s(bool,'F_0') != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X274,bool),fun(fun(X274,bool),bool)),'const.pair.,_2'(s(fun(X274,bool),X275),s(fun(fun(X274,bool),bool),'const.pred_set.POW_1'(s(fun(X274,bool),X275))))))) | p(s(bool,'const.bool.IN_2'(s(X272,sK1('const.pred_set.UNIV_0',X273,X272)),s(fun(X272,bool),X273)))) | s(fun(X272,bool),'const.pred_set.UNIV_0') = s(fun(X272,bool),X273)) )),
  inference(superposition,[],[f355,f2479])).
fof(f2479,plain,(
  ( ! [X10,X9] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X9,sK1('const.pred_set.UNIV_0',X10,X9)),s(fun(X9,bool),X10))) | s(fun(X9,bool),'const.pred_set.UNIV_0') = s(fun(X9,bool),X10)) )),
  inference(trivial_inequality_removal,[],[f2471])).
fof(f2471,plain,(
  ( ! [X10,X9] : (s(bool,'T_0') != s(bool,'T_0') | s(fun(X9,bool),'const.pred_set.UNIV_0') = s(fun(X9,bool),X10) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X9,sK1('const.pred_set.UNIV_0',X10,X9)),s(fun(X9,bool),X10)))) )),
  inference(superposition,[],[f717,f136])).
fof(f355,plain,(
  ( ! [X80,X78,X79] : (s(bool,X78) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X79,bool),fun(fun(X79,bool),bool)),'const.pair.,_2'(s(fun(X79,bool),X80),s(fun(fun(X79,bool),bool),'const.pred_set.POW_1'(s(fun(X79,bool),X80))))))) | p(s(bool,X78))) )),
  inference(resolution,[],[f160,f133])).
fof(f2521,plain,(
  ( ! [X0,X3,X1] : (s(fun(X0,bool),'const.pred_set.UNIV_0') = s(fun(X0,bool),X1) | s(bool,X3) != s(bool,'const.bool.IN_2'(s(X0,sK1('const.pred_set.UNIV_0',X1,X0)),s(fun(X0,bool),X1))) | ~p(s(bool,X3))) ) | $spl56),
  inference(cnf_transformation,[],[f2521_D])).
fof(f2521_D,plain,(
  ( ! [X0,X3,X1] : (s(fun(X0,bool),'const.pred_set.UNIV_0') = s(fun(X0,bool),X1) | s(bool,X3) != s(bool,'const.bool.IN_2'(s(X0,sK1('const.pred_set.UNIV_0',X1,X0)),s(fun(X0,bool),X1))) | ~p(s(bool,X3))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f2854,plain,(
  ( ! [X253,X251,X254,X252] : (p(s(bool,X253)) | ~p(s(bool,X254)) | s(bool,X254) != s(bool,'const.bool.IN_2'(s(X251,sK1('const.pred_set.UNIV_0',X252,X251)),s(fun(X251,bool),X252))) | s(fun(X251,bool),'const.pred_set.UNIV_0') = s(fun(X251,bool),X252)) )),
  inference(subsumption_resolution,[],[f2792,f108])).
fof(f108,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f72])).
fof(f72,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f30])).
fof(f30,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X7] : ((p(s(bool,X7)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X7)))),
  file('thm.measure.UNIV_SIGMA_ALGEBRA','thm.bool.IMP_CLAUSES|split|4')).
fof(f2792,plain,(
  ( ! [X253,X251,X254,X252] : (p(s(bool,'F_0')) | p(s(bool,X253)) | ~p(s(bool,X254)) | s(bool,X254) != s(bool,'const.bool.IN_2'(s(X251,sK1('const.pred_set.UNIV_0',X252,X251)),s(fun(X251,bool),X252))) | s(fun(X251,bool),'const.pred_set.UNIV_0') = s(fun(X251,bool),X252)) )),
  inference(superposition,[],[f142,f2479])).
fof(f2850,plain,(
  ( ! [X182,X180,X178,X181,X179] : (~p(s(bool,X180)) | ~p(s(bool,X181)) | s(bool,X181) != s(bool,'const.bool.IN_2'(s(X178,sK1('const.pred_set.UNIV_0',X179,X178)),s(fun(X178,bool),X179))) | s(bool,X180) != s(bool,X182) | s(fun(X178,bool),'const.pred_set.UNIV_0') = s(fun(X178,bool),X179)) )),
  inference(subsumption_resolution,[],[f2765,f108])).
fof(f2765,plain,(
  ( ! [X182,X180,X178,X181,X179] : (p(s(bool,'F_0')) | ~p(s(bool,X180)) | ~p(s(bool,X181)) | s(bool,X181) != s(bool,'const.bool.IN_2'(s(X178,sK1('const.pred_set.UNIV_0',X179,X178)),s(fun(X178,bool),X179))) | s(bool,X180) != s(bool,X182) | s(fun(X178,bool),'const.pred_set.UNIV_0') = s(fun(X178,bool),X179)) )),
  inference(superposition,[],[f98,f2479])).
fof(f2845,plain,(
  ( ! [X142,X140,X138,X141,X139] : (p(s(bool,X140)) | ~p(s(bool,X141)) | ~p(s(bool,X142)) | s(bool,X141) != s(bool,'const.bool.IN_2'(s(X138,sK1('const.pred_set.UNIV_0',X139,X138)),s(fun(X138,bool),X139))) | s(fun(X138,bool),'const.pred_set.UNIV_0') = s(fun(X138,bool),X139)) )),
  inference(subsumption_resolution,[],[f2757,f108])).
fof(f2757,plain,(
  ( ! [X142,X140,X138,X141,X139] : (p(s(bool,'F_0')) | p(s(bool,X140)) | ~p(s(bool,X141)) | ~p(s(bool,X142)) | s(bool,X141) != s(bool,'const.bool.IN_2'(s(X138,sK1('const.pred_set.UNIV_0',X139,X138)),s(fun(X138,bool),X139))) | s(fun(X138,bool),'const.pred_set.UNIV_0') = s(fun(X138,bool),X139)) )),
  inference(superposition,[],[f97,f2479])).
fof(f2841,plain,(
  ( ! [X116,X114,X117,X115,X113] : (p(s(bool,X115)) | ~p(s(bool,X116)) | s(bool,X115) != s(bool,X117) | s(bool,X116) != s(bool,'const.bool.IN_2'(s(X113,sK1('const.pred_set.UNIV_0',X114,X113)),s(fun(X113,bool),X114))) | s(fun(X113,bool),'const.pred_set.UNIV_0') = s(fun(X113,bool),X114)) )),
  inference(subsumption_resolution,[],[f2752,f108])).
fof(f2752,plain,(
  ( ! [X116,X114,X117,X115,X113] : (p(s(bool,'F_0')) | p(s(bool,X115)) | ~p(s(bool,X116)) | s(bool,X115) != s(bool,X117) | s(bool,X116) != s(bool,'const.bool.IN_2'(s(X113,sK1('const.pred_set.UNIV_0',X114,X113)),s(fun(X113,bool),X114))) | s(fun(X113,bool),'const.pred_set.UNIV_0') = s(fun(X113,bool),X114)) )),
  inference(superposition,[],[f96,f2479])).
fof(f2839,plain,(
  ( ! [X80,X78,X81,X79,X82] : (~p(s(bool,X80)) | ~p(s(bool,X81)) | s(bool,X81) != s(bool,'const.bool.IN_2'(s(X78,sK1('const.pred_set.UNIV_0',X79,X78)),s(fun(X78,bool),X79))) | s(bool,X80) != s(bool,X82) | s(fun(X78,bool),'const.pred_set.UNIV_0') = s(fun(X78,bool),X79)) )),
  inference(subsumption_resolution,[],[f2745,f108])).
fof(f2745,plain,(
  ( ! [X80,X78,X81,X79,X82] : (p(s(bool,'F_0')) | ~p(s(bool,X80)) | ~p(s(bool,X81)) | s(bool,X81) != s(bool,'const.bool.IN_2'(s(X78,sK1('const.pred_set.UNIV_0',X79,X78)),s(fun(X78,bool),X79))) | s(bool,X80) != s(bool,X82) | s(fun(X78,bool),'const.pred_set.UNIV_0') = s(fun(X78,bool),X79)) )),
  inference(superposition,[],[f94,f2479])).
fof(f2833,plain,(
  ( ! [X47,X45,X43,X46,X44] : (p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X45) != s(bool,X47) | s(bool,X46) != s(bool,'const.bool.IN_2'(s(X43,sK1('const.pred_set.UNIV_0',X44,X43)),s(fun(X43,bool),X44))) | s(fun(X43,bool),'const.pred_set.UNIV_0') = s(fun(X43,bool),X44)) )),
  inference(subsumption_resolution,[],[f2738,f108])).
fof(f2738,plain,(
  ( ! [X47,X45,X43,X46,X44] : (p(s(bool,'F_0')) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X45) != s(bool,X47) | s(bool,X46) != s(bool,'const.bool.IN_2'(s(X43,sK1('const.pred_set.UNIV_0',X44,X43)),s(fun(X43,bool),X44))) | s(fun(X43,bool),'const.pred_set.UNIV_0') = s(fun(X43,bool),X44)) )),
  inference(superposition,[],[f92,f2479])).
fof(f2686,plain,(
  ( ! [X271,X273,X272,X274] : (s(bool,'F_0') != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X273,bool),fun(fun(X273,bool),bool)),'const.pair.,_2'(s(fun(X273,bool),X274),s(fun(fun(X273,bool),bool),'const.pred_set.POW_1'(s(fun(X273,bool),X274))))))) | s(fun(X271,bool),'const.pred_set.UNIV_0') = s(fun(X271,bool),X272)) )),
  inference(subsumption_resolution,[],[f2625,f2401])).
fof(f2625,plain,(
  ( ! [X271,X273,X272,X274] : (s(bool,'F_0') != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X273,bool),fun(fun(X273,bool),bool)),'const.pair.,_2'(s(fun(X273,bool),X274),s(fun(fun(X273,bool),bool),'const.pred_set.POW_1'(s(fun(X273,bool),X274))))))) | p(s(bool,'const.bool.IN_2'(s(X271,sK1(X272,'const.pred_set.UNIV_0',X271)),s(fun(X271,bool),X272)))) | s(fun(X271,bool),'const.pred_set.UNIV_0') = s(fun(X271,bool),X272)) )),
  inference(superposition,[],[f355,f2403])).
fof(f2403,plain,(
  ( ! [X10,X9] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X9,sK1(X10,'const.pred_set.UNIV_0',X9)),s(fun(X9,bool),X10))) | s(fun(X9,bool),'const.pred_set.UNIV_0') = s(fun(X9,bool),X10)) )),
  inference(trivial_inequality_removal,[],[f2395])).
fof(f2395,plain,(
  ( ! [X10,X9] : (s(bool,'T_0') != s(bool,'T_0') | s(fun(X9,bool),'const.pred_set.UNIV_0') = s(fun(X9,bool),X10) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X9,sK1(X10,'const.pred_set.UNIV_0',X9)),s(fun(X9,bool),X10)))) )),
  inference(superposition,[],[f716,f136])).
fof(f2450,plain,(
  ( ! [X0,X3,X1] : (s(fun(X0,bool),'const.pred_set.UNIV_0') = s(fun(X0,bool),X1) | s(bool,X3) != s(bool,'const.bool.IN_2'(s(X0,sK1(X1,'const.pred_set.UNIV_0',X0)),s(fun(X0,bool),X1))) | ~p(s(bool,X3))) ) | $spl54),
  inference(cnf_transformation,[],[f2450_D])).
fof(f2450_D,plain,(
  ( ! [X0,X3,X1] : (s(fun(X0,bool),'const.pred_set.UNIV_0') = s(fun(X0,bool),X1) | s(bool,X3) != s(bool,'const.bool.IN_2'(s(X0,sK1(X1,'const.pred_set.UNIV_0',X0)),s(fun(X0,bool),X1))) | ~p(s(bool,X3))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f2681,plain,(
  ( ! [X253,X251,X252,X250] : (p(s(bool,X252)) | ~p(s(bool,X253)) | s(bool,X253) != s(bool,'const.bool.IN_2'(s(X250,sK1(X251,'const.pred_set.UNIV_0',X250)),s(fun(X250,bool),X251))) | s(fun(X250,bool),'const.pred_set.UNIV_0') = s(fun(X250,bool),X251)) )),
  inference(subsumption_resolution,[],[f2619,f108])).
fof(f2619,plain,(
  ( ! [X253,X251,X252,X250] : (p(s(bool,'F_0')) | p(s(bool,X252)) | ~p(s(bool,X253)) | s(bool,X253) != s(bool,'const.bool.IN_2'(s(X250,sK1(X251,'const.pred_set.UNIV_0',X250)),s(fun(X250,bool),X251))) | s(fun(X250,bool),'const.pred_set.UNIV_0') = s(fun(X250,bool),X251)) )),
  inference(superposition,[],[f142,f2403])).
fof(f2677,plain,(
  ( ! [X177,X180,X178,X181,X179] : (~p(s(bool,X179)) | ~p(s(bool,X180)) | s(bool,X180) != s(bool,'const.bool.IN_2'(s(X177,sK1(X178,'const.pred_set.UNIV_0',X177)),s(fun(X177,bool),X178))) | s(bool,X179) != s(bool,X181) | s(fun(X177,bool),'const.pred_set.UNIV_0') = s(fun(X177,bool),X178)) )),
  inference(subsumption_resolution,[],[f2592,f108])).
fof(f2592,plain,(
  ( ! [X177,X180,X178,X181,X179] : (p(s(bool,'F_0')) | ~p(s(bool,X179)) | ~p(s(bool,X180)) | s(bool,X180) != s(bool,'const.bool.IN_2'(s(X177,sK1(X178,'const.pred_set.UNIV_0',X177)),s(fun(X177,bool),X178))) | s(bool,X179) != s(bool,X181) | s(fun(X177,bool),'const.pred_set.UNIV_0') = s(fun(X177,bool),X178)) )),
  inference(superposition,[],[f98,f2403])).
fof(f2672,plain,(
  ( ! [X140,X138,X141,X139,X137] : (p(s(bool,X139)) | ~p(s(bool,X140)) | ~p(s(bool,X141)) | s(bool,X140) != s(bool,'const.bool.IN_2'(s(X137,sK1(X138,'const.pred_set.UNIV_0',X137)),s(fun(X137,bool),X138))) | s(fun(X137,bool),'const.pred_set.UNIV_0') = s(fun(X137,bool),X138)) )),
  inference(subsumption_resolution,[],[f2584,f108])).
fof(f2584,plain,(
  ( ! [X140,X138,X141,X139,X137] : (p(s(bool,'F_0')) | p(s(bool,X139)) | ~p(s(bool,X140)) | ~p(s(bool,X141)) | s(bool,X140) != s(bool,'const.bool.IN_2'(s(X137,sK1(X138,'const.pred_set.UNIV_0',X137)),s(fun(X137,bool),X138))) | s(fun(X137,bool),'const.pred_set.UNIV_0') = s(fun(X137,bool),X138)) )),
  inference(superposition,[],[f97,f2403])).
fof(f2668,plain,(
  ( ! [X116,X114,X112,X115,X113] : (p(s(bool,X114)) | ~p(s(bool,X115)) | s(bool,X114) != s(bool,X116) | s(bool,X115) != s(bool,'const.bool.IN_2'(s(X112,sK1(X113,'const.pred_set.UNIV_0',X112)),s(fun(X112,bool),X113))) | s(fun(X112,bool),'const.pred_set.UNIV_0') = s(fun(X112,bool),X113)) )),
  inference(subsumption_resolution,[],[f2579,f108])).
fof(f2579,plain,(
  ( ! [X116,X114,X112,X115,X113] : (p(s(bool,'F_0')) | p(s(bool,X114)) | ~p(s(bool,X115)) | s(bool,X114) != s(bool,X116) | s(bool,X115) != s(bool,'const.bool.IN_2'(s(X112,sK1(X113,'const.pred_set.UNIV_0',X112)),s(fun(X112,bool),X113))) | s(fun(X112,bool),'const.pred_set.UNIV_0') = s(fun(X112,bool),X113)) )),
  inference(superposition,[],[f96,f2403])).
fof(f2666,plain,(
  ( ! [X80,X78,X81,X79,X77] : (~p(s(bool,X79)) | ~p(s(bool,X80)) | s(bool,X80) != s(bool,'const.bool.IN_2'(s(X77,sK1(X78,'const.pred_set.UNIV_0',X77)),s(fun(X77,bool),X78))) | s(bool,X79) != s(bool,X81) | s(fun(X77,bool),'const.pred_set.UNIV_0') = s(fun(X77,bool),X78)) )),
  inference(subsumption_resolution,[],[f2572,f108])).
fof(f2572,plain,(
  ( ! [X80,X78,X81,X79,X77] : (p(s(bool,'F_0')) | ~p(s(bool,X79)) | ~p(s(bool,X80)) | s(bool,X80) != s(bool,'const.bool.IN_2'(s(X77,sK1(X78,'const.pred_set.UNIV_0',X77)),s(fun(X77,bool),X78))) | s(bool,X79) != s(bool,X81) | s(fun(X77,bool),'const.pred_set.UNIV_0') = s(fun(X77,bool),X78)) )),
  inference(superposition,[],[f94,f2403])).
fof(f2660,plain,(
  ( ! [X45,X43,X46,X44,X42] : (p(s(bool,X44)) | ~p(s(bool,X45)) | s(bool,X44) != s(bool,X46) | s(bool,X45) != s(bool,'const.bool.IN_2'(s(X42,sK1(X43,'const.pred_set.UNIV_0',X42)),s(fun(X42,bool),X43))) | s(fun(X42,bool),'const.pred_set.UNIV_0') = s(fun(X42,bool),X43)) )),
  inference(subsumption_resolution,[],[f2565,f108])).
fof(f2565,plain,(
  ( ! [X45,X43,X46,X44,X42] : (p(s(bool,'F_0')) | p(s(bool,X44)) | ~p(s(bool,X45)) | s(bool,X44) != s(bool,X46) | s(bool,X45) != s(bool,'const.bool.IN_2'(s(X42,sK1(X43,'const.pred_set.UNIV_0',X42)),s(fun(X42,bool),X43))) | s(fun(X42,bool),'const.pred_set.UNIV_0') = s(fun(X42,bool),X43)) )),
  inference(superposition,[],[f92,f2403])).
fof(f2496,plain,(
  ( ! [X59,X57,X58,X56] : (s(fun(X56,bool),'const.pred_set.UNIV_0') = s(fun(X56,bool),X57) | p(s(bool,X58)) | ~p(s(bool,X59)) | s(bool,X59) != s(bool,'const.bool.IN_2'(s(X56,sK1('const.pred_set.UNIV_0',X57,X56)),s(fun(X56,bool),X57)))) )),
  inference(resolution,[],[f2477,f142])).
fof(f2487,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s(fun(X35,bool),'const.pred_set.UNIV_0') = s(fun(X35,bool),X36) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X38) != s(bool,'const.bool.IN_2'(s(X35,sK1('const.pred_set.UNIV_0',X36,X35)),s(fun(X35,bool),X36))) | s(bool,X37) != s(bool,X39)) )),
  inference(resolution,[],[f2477,f98])).
fof(f2485,plain,(
  ( ! [X28,X26,X29,X27,X25] : (s(fun(X25,bool),'const.pred_set.UNIV_0') = s(fun(X25,bool),X26) | p(s(bool,X27)) | ~p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X28) != s(bool,'const.bool.IN_2'(s(X25,sK1('const.pred_set.UNIV_0',X26,X25)),s(fun(X25,bool),X26)))) )),
  inference(resolution,[],[f2477,f97])).
fof(f2483,plain,(
  ( ! [X19,X17,X15,X18,X16] : (s(fun(X15,bool),'const.pred_set.UNIV_0') = s(fun(X15,bool),X16) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,'const.bool.IN_2'(s(X15,sK1('const.pred_set.UNIV_0',X16,X15)),s(fun(X15,bool),X16)))) )),
  inference(resolution,[],[f2477,f96])).
fof(f2482,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s(fun(X10,bool),'const.pred_set.UNIV_0') = s(fun(X10,bool),X11) | ~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X13) != s(bool,'const.bool.IN_2'(s(X10,sK1('const.pred_set.UNIV_0',X11,X10)),s(fun(X10,bool),X11))) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f2477,f94])).
fof(f2480,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun(X0,bool),'const.pred_set.UNIV_0') = s(fun(X0,bool),X1) | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.bool.IN_2'(s(X0,sK1('const.pred_set.UNIV_0',X1,X0)),s(fun(X0,bool),X1)))) )),
  inference(resolution,[],[f2477,f92])).
fof(f2425,plain,(
  ( ! [X59,X57,X58,X56] : (s(fun(X56,bool),'const.pred_set.UNIV_0') = s(fun(X56,bool),X57) | p(s(bool,X58)) | ~p(s(bool,X59)) | s(bool,X59) != s(bool,'const.bool.IN_2'(s(X56,sK1(X57,'const.pred_set.UNIV_0',X56)),s(fun(X56,bool),X57)))) )),
  inference(resolution,[],[f2401,f142])).
fof(f2416,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s(fun(X35,bool),'const.pred_set.UNIV_0') = s(fun(X35,bool),X36) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X38) != s(bool,'const.bool.IN_2'(s(X35,sK1(X36,'const.pred_set.UNIV_0',X35)),s(fun(X35,bool),X36))) | s(bool,X37) != s(bool,X39)) )),
  inference(resolution,[],[f2401,f98])).
fof(f2414,plain,(
  ( ! [X28,X26,X29,X27,X25] : (s(fun(X25,bool),'const.pred_set.UNIV_0') = s(fun(X25,bool),X26) | p(s(bool,X27)) | ~p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X28) != s(bool,'const.bool.IN_2'(s(X25,sK1(X26,'const.pred_set.UNIV_0',X25)),s(fun(X25,bool),X26)))) )),
  inference(resolution,[],[f2401,f97])).
fof(f2412,plain,(
  ( ! [X19,X17,X15,X18,X16] : (s(fun(X15,bool),'const.pred_set.UNIV_0') = s(fun(X15,bool),X16) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,'const.bool.IN_2'(s(X15,sK1(X16,'const.pred_set.UNIV_0',X15)),s(fun(X15,bool),X16)))) )),
  inference(resolution,[],[f2401,f96])).
fof(f2411,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s(fun(X10,bool),'const.pred_set.UNIV_0') = s(fun(X10,bool),X11) | ~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X13) != s(bool,'const.bool.IN_2'(s(X10,sK1(X11,'const.pred_set.UNIV_0',X10)),s(fun(X10,bool),X11))) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f2401,f94])).
fof(f2409,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun(X0,bool),'const.pred_set.UNIV_0') = s(fun(X0,bool),X1) | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.bool.IN_2'(s(X0,sK1(X1,'const.pred_set.UNIV_0',X0)),s(fun(X0,bool),X1)))) )),
  inference(resolution,[],[f2401,f92])).
fof(f2297,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) != s(bool,X5) | ~p(s(bool,X5)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl52),
  inference(cnf_transformation,[],[f2297_D])).
fof(f2297_D,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) != s(bool,X5) | ~p(s(bool,X5)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f2262,plain,(
  ( ! [X92,X90,X88,X93,X91,X89] : (s(bool,X88) != s(bool,X89) | s(bool,X88) != s(bool,X90) | s(bool,X90) != s(bool,X91) | p(s(bool,X91)) | p(s(bool,X92)) | ~p(s(bool,X93)) | s(bool,X89) != s(bool,X93)) ) | $spl42),
  inference(resolution,[],[f1494,f142])).
fof(f2253,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (s(bool,X49) != s(bool,X50) | s(bool,X49) != s(bool,X51) | s(bool,X51) != s(bool,X52) | p(s(bool,X52)) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X50) != s(bool,X54) | s(bool,X53) != s(bool,X55)) ) | $spl42),
  inference(resolution,[],[f1494,f98])).
fof(f2251,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X40] : (s(bool,X35) != s(bool,X36) | s(bool,X35) != s(bool,X37) | s(bool,X37) != s(bool,X38) | p(s(bool,X38)) | p(s(bool,X39)) | ~p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X36) != s(bool,X40)) ) | $spl42),
  inference(resolution,[],[f1494,f97])).
fof(f2249,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s(bool,X21) != s(bool,X22) | s(bool,X21) != s(bool,X23) | s(bool,X23) != s(bool,X24) | p(s(bool,X24)) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X22) != s(bool,X26)) ) | $spl42),
  inference(resolution,[],[f1494,f96])).
fof(f2248,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s(bool,X14) != s(bool,X15) | s(bool,X14) != s(bool,X16) | s(bool,X16) != s(bool,X17) | p(s(bool,X17)) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X15) != s(bool,X19) | s(bool,X18) != s(bool,X20)) ) | $spl42),
  inference(resolution,[],[f1494,f94])).
fof(f2246,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2) | s(bool,X2) != s(bool,X3) | p(s(bool,X3)) | p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X1) != s(bool,X5)) ) | $spl42),
  inference(resolution,[],[f1494,f92])).
fof(f2243,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') = s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X3,bool),fun(fun(X3,bool),bool)),'const.pair.,_2'(s(fun(X3,bool),X4),s(fun(fun(X3,bool),bool),'const.pred_set.POW_1'(s(fun(X3,bool),X4)))))))) ) | $spl50),
  inference(cnf_transformation,[],[f2243_D])).
fof(f2243_D,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') = s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X3,bool),fun(fun(X3,bool),bool)),'const.pair.,_2'(s(fun(X3,bool),X4),s(fun(fun(X3,bool),bool),'const.pred_set.POW_1'(s(fun(X3,bool),X4)))))))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f2242,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X5)) | s(bool,'T_0') = s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X3,bool),fun(fun(X3,bool),bool)),'const.pair.,_2'(s(fun(X3,bool),X4),s(fun(fun(X3,bool),bool),'const.pred_set.POW_1'(s(fun(X3,bool),X4)))))))) )),
  inference(subsumption_resolution,[],[f2226,f123])).
fof(f2226,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X5) | p(s(bool,X5)) | s(bool,'T_0') = s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X3,bool),fun(fun(X3,bool),bool)),'const.pair.,_2'(s(fun(X3,bool),X4),s(fun(fun(X3,bool),bool),'const.pred_set.POW_1'(s(fun(X3,bool),X4)))))))) )),
  inference(superposition,[],[f355,f136])).
fof(f2232,plain,(
  ( ! [X17,X15,X16] : (s(bool,'F_0') != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X16,bool),fun(fun(X16,bool),bool)),'const.pair.,_2'(s(fun(X16,bool),X17),s(fun(fun(X16,bool),bool),'const.pred_set.POW_1'(s(fun(X16,bool),X17))))))) | p(s(bool,X15))) )),
  inference(duplicate_literal_removal,[],[f2214])).
fof(f2214,plain,(
  ( ! [X17,X15,X16] : (s(bool,'F_0') != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X16,bool),fun(fun(X16,bool),bool)),'const.pair.,_2'(s(fun(X16,bool),X17),s(fun(fun(X16,bool),bool),'const.pred_set.POW_1'(s(fun(X16,bool),X17))))))) | p(s(bool,X15)) | p(s(bool,X15))) )),
  inference(superposition,[],[f355,f123])).
fof(f2233,plain,(
  ( ! [X14,X12,X13] : (s(bool,'F_0') != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X13,bool),fun(fun(X13,bool),bool)),'const.pair.,_2'(s(fun(X13,bool),X14),s(fun(fun(X13,bool),bool),'const.pred_set.POW_1'(s(fun(X13,bool),X14))))))) | p(s(bool,X12))) )),
  inference(duplicate_literal_removal,[],[f2213])).
fof(f2213,plain,(
  ( ! [X14,X12,X13] : (s(bool,'F_0') != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X13,bool),fun(fun(X13,bool),bool)),'const.pair.,_2'(s(fun(X13,bool),X14),s(fun(fun(X13,bool),bool),'const.pred_set.POW_1'(s(fun(X13,bool),X14))))))) | p(s(bool,X12)) | p(s(bool,X12))) )),
  inference(superposition,[],[f355,f125])).
fof(f2234,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X4,bool),fun(fun(X4,bool),bool)),'const.pair.,_2'(s(fun(X4,bool),X5),s(fun(fun(X4,bool),bool),'const.pred_set.POW_1'(s(fun(X4,bool),X5))))))) | p(s(bool,X3))) )),
  inference(subsumption_resolution,[],[f2210,f126])).
fof(f126,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f84])).
fof(f2210,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X4,bool),fun(fun(X4,bool),bool)),'const.pair.,_2'(s(fun(X4,bool),X5),s(fun(fun(X4,bool),bool),'const.pred_set.POW_1'(s(fun(X4,bool),X5))))))) | p(s(bool,X3)) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f355,f136])).
fof(f2042,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,X4)) | s(bool,X1) != s(bool,X4) | s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl46),
  inference(cnf_transformation,[],[f2042_D])).
fof(f2042_D,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,X4)) | s(bool,X1) != s(bool,X4) | s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f2007,plain,(
  ( ! [X76,X74,X72,X75,X73] : (s(bool,X72) != s(bool,X73) | s(bool,'T_0') = s(bool,X74) | s(bool,X72) != s(bool,X74) | p(s(bool,X75)) | ~p(s(bool,X76)) | s(bool,X73) != s(bool,X76)) ) | $spl32),
  inference(resolution,[],[f1175,f142])).
fof(f1998,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(bool,X42) != s(bool,X43) | s(bool,'T_0') = s(bool,X44) | s(bool,X42) != s(bool,X44) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X43) != s(bool,X46) | s(bool,X45) != s(bool,X47)) ) | $spl32),
  inference(resolution,[],[f1175,f98])).
fof(f1996,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (s(bool,X30) != s(bool,X31) | s(bool,'T_0') = s(bool,X32) | s(bool,X30) != s(bool,X32) | p(s(bool,X33)) | ~p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X31) != s(bool,X34)) ) | $spl32),
  inference(resolution,[],[f1175,f97])).
fof(f1994,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,X18) != s(bool,X19) | s(bool,'T_0') = s(bool,X20) | s(bool,X18) != s(bool,X20) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X19) != s(bool,X22)) ) | $spl32),
  inference(resolution,[],[f1175,f96])).
fof(f1993,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s(bool,X12) != s(bool,X13) | s(bool,'T_0') = s(bool,X14) | s(bool,X12) != s(bool,X14) | ~p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X13) != s(bool,X16) | s(bool,X15) != s(bool,X17)) ) | $spl32),
  inference(resolution,[],[f1175,f94])).
fof(f1991,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) != s(bool,X2) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X1) != s(bool,X4)) ) | $spl32),
  inference(resolution,[],[f1175,f92])).
fof(f1924,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,X2) != s(bool,X4) | p(s(bool,X4)) | s(bool,X1) != s(bool,X2) | s(bool,X0) != s(bool,X1)) ) | $spl44),
  inference(cnf_transformation,[],[f1924_D])).
fof(f1924_D,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,X2) != s(bool,X4) | p(s(bool,X4)) | s(bool,X1) != s(bool,X2) | s(bool,X0) != s(bool,X1)) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f1895,plain,(
  ( ! [X70,X72,X71,X69,X73] : (s(bool,'F_0') = s(bool,X69) | s(bool,X69) != s(bool,X70) | s(bool,X70) != s(bool,X71) | p(s(bool,X72)) | p(s(bool,X73)) | s(bool,X71) != s(bool,X73)) )),
  inference(resolution,[],[f402,f142])).
fof(f402,plain,(
  ( ! [X59,X60,X58] : (p(s(bool,X60)) | s(bool,'F_0') = s(bool,X59) | s(bool,X58) != s(bool,X59) | s(bool,X58) != s(bool,X60)) )),
  inference(resolution,[],[f340,f160])).
fof(f1887,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(bool,'F_0') = s(bool,X42) | s(bool,X42) != s(bool,X43) | s(bool,X43) != s(bool,X44) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X44) != s(bool,X45) | s(bool,X46) != s(bool,X47)) )),
  inference(resolution,[],[f402,f98])).
fof(f1884,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (s(bool,'F_0') = s(bool,X24) | s(bool,X24) != s(bool,X25) | s(bool,X25) != s(bool,X26) | p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X26) != s(bool,X29)) )),
  inference(resolution,[],[f402,f97])).
fof(f1883,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,'F_0') = s(bool,X18) | s(bool,X18) != s(bool,X19) | s(bool,X19) != s(bool,X20) | p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X20) != s(bool,X22)) )),
  inference(resolution,[],[f402,f96])).
fof(f1882,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s(bool,'F_0') = s(bool,X12) | s(bool,X12) != s(bool,X13) | s(bool,X13) != s(bool,X14) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X14) != s(bool,X15) | s(bool,X16) != s(bool,X17)) )),
  inference(resolution,[],[f402,f94])).
fof(f1880,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,X0) != s(bool,X1) | s(bool,X1) != s(bool,X2) | p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X2) != s(bool,X4)) )),
  inference(resolution,[],[f402,f92])).
fof(f1494,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,X4)) | s(bool,X1) != s(bool,X4) | s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) | $spl42),
  inference(cnf_transformation,[],[f1494_D])).
fof(f1494_D,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,X4)) | s(bool,X1) != s(bool,X4) | s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f1461,plain,(
  ( ! [X76,X74,X72,X75,X73] : (s(bool,X72) != s(bool,X73) | p(s(bool,X74)) | s(bool,X72) != s(bool,X74) | p(s(bool,X75)) | ~p(s(bool,X76)) | s(bool,X73) != s(bool,X76)) ) | $spl18),
  inference(resolution,[],[f375,f142])).
fof(f1452,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(bool,X42) != s(bool,X43) | p(s(bool,X44)) | s(bool,X42) != s(bool,X44) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X43) != s(bool,X46) | s(bool,X45) != s(bool,X47)) ) | $spl18),
  inference(resolution,[],[f375,f98])).
fof(f1450,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (s(bool,X30) != s(bool,X31) | p(s(bool,X32)) | s(bool,X30) != s(bool,X32) | p(s(bool,X33)) | ~p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X31) != s(bool,X34)) ) | $spl18),
  inference(resolution,[],[f375,f97])).
fof(f1448,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,X18) != s(bool,X19) | p(s(bool,X20)) | s(bool,X18) != s(bool,X20) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X19) != s(bool,X22)) ) | $spl18),
  inference(resolution,[],[f375,f96])).
fof(f1447,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s(bool,X12) != s(bool,X13) | p(s(bool,X14)) | s(bool,X12) != s(bool,X14) | ~p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X13) != s(bool,X16) | s(bool,X15) != s(bool,X17)) ) | $spl18),
  inference(resolution,[],[f375,f94])).
fof(f1445,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | p(s(bool,X2)) | s(bool,X0) != s(bool,X2) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X1) != s(bool,X4)) ) | $spl18),
  inference(resolution,[],[f375,f92])).
fof(f1175,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,X3)) | s(bool,X1) != s(bool,X3) | s(bool,'T_0') = s(bool,X0) | s(bool,X0) != s(bool,X1)) ) | $spl32),
  inference(cnf_transformation,[],[f1175_D])).
fof(f1175_D,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,X3)) | s(bool,X1) != s(bool,X3) | s(bool,'T_0') = s(bool,X0) | s(bool,X0) != s(bool,X1)) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f1372,plain,(
  ( ! [X59,X57,X58,X56] : (s(bool,X56) != s(bool,X57) | s(bool,'T_0') = s(bool,X57) | p(s(bool,X58)) | ~p(s(bool,X59)) | s(bool,X56) != s(bool,X59)) ) | $spl34),
  inference(resolution,[],[f1183,f142])).
fof(f1363,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s(bool,X35) != s(bool,X36) | s(bool,'T_0') = s(bool,X36) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X35) != s(bool,X38) | s(bool,X37) != s(bool,X39)) ) | $spl34),
  inference(resolution,[],[f1183,f98])).
fof(f1361,plain,(
  ( ! [X28,X26,X29,X27,X25] : (s(bool,X25) != s(bool,X26) | s(bool,'T_0') = s(bool,X26) | p(s(bool,X27)) | ~p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X25) != s(bool,X28)) ) | $spl34),
  inference(resolution,[],[f1183,f97])).
fof(f1359,plain,(
  ( ! [X19,X17,X15,X18,X16] : (s(bool,X15) != s(bool,X16) | s(bool,'T_0') = s(bool,X16) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X15) != s(bool,X18)) ) | $spl34),
  inference(resolution,[],[f1183,f96])).
fof(f1358,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s(bool,X10) != s(bool,X11) | s(bool,'T_0') = s(bool,X11) | ~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X10) != s(bool,X13) | s(bool,X12) != s(bool,X14)) ) | $spl34),
  inference(resolution,[],[f1183,f94])).
fof(f1356,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X1) | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X0) != s(bool,X3)) ) | $spl34),
  inference(resolution,[],[f1183,f92])).
fof(f1337,plain,(
  $false | $spl16),
  inference(trivial_inequality_removal,[],[f1332])).
fof(f1332,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl16),
  inference(superposition,[],[f271,f684])).
fof(f1336,plain,(
  $false | ($spl16 | $spl30)),
  inference(trivial_inequality_removal,[],[f1334])).
fof(f1334,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl16 | $spl30)),
  inference(superposition,[],[f271,f896])).
fof(f1335,plain,(
  $false | $spl16),
  inference(equality_resolution,[],[f271])).
fof(f1299,plain,(
  $false | $spl4),
  inference(trivial_inequality_removal,[],[f1281])).
fof(f1281,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl4),
  inference(superposition,[],[f173,f155])).
fof(f155,plain,(
  s(bool,'F_0') = s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.UNIV_0')))))),
  inference(resolution,[],[f90,f123])).
fof(f1298,plain,(
  $false | $spl4),
  inference(equality_resolution,[],[f173])).
fof(f1268,plain,(
  $false | $spl14),
  inference(trivial_inequality_removal,[],[f1261])).
fof(f1261,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl14),
  inference(superposition,[],[f266,f155])).
fof(f1267,plain,(
  $false | $spl14),
  inference(equality_resolution,[],[f266])).
fof(f1236,plain,(
  $false | $spl2),
  inference(trivial_inequality_removal,[],[f1229])).
fof(f1229,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl2),
  inference(superposition,[],[f243,f155])).
fof(f243,plain,(
  ( ! [X60] : (s(bool,'F_0') != s(bool,X60)) ) | $spl2),
  inference(resolution,[],[f172,f130])).
fof(f1235,plain,(
  $false | $spl2),
  inference(equality_resolution,[],[f243])).
fof(f266,plain,(
  ( ! [X61] : (s(bool,'F_0') != s(bool,X61)) ) | $spl14),
  inference(cnf_transformation,[],[f266_D])).
fof(f266_D,plain,(
  ( ! [X61] : (s(bool,'F_0') != s(bool,X61)) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f262,plain,(
  ( ! [X23,X22] : (p(s(bool,X22)) | s(bool,'F_0') != s(bool,X23)) ) | $spl2),
  inference(duplicate_literal_removal,[],[f258])).
fof(f258,plain,(
  ( ! [X23,X22] : (p(s(bool,X22)) | s(bool,'F_0') != s(bool,X23) | p(s(bool,X22))) ) | $spl2),
  inference(superposition,[],[f172,f123])).
fof(f263,plain,(
  ( ! [X21,X20] : (p(s(bool,X20)) | s(bool,'F_0') != s(bool,X21)) ) | $spl2),
  inference(duplicate_literal_removal,[],[f257])).
fof(f257,plain,(
  ( ! [X21,X20] : (p(s(bool,X20)) | s(bool,'F_0') != s(bool,X21) | p(s(bool,X20))) ) | $spl2),
  inference(superposition,[],[f172,f125])).
fof(f271,plain,(
  ( ! [X13] : (s(bool,'T_0') != s(bool,X13)) ) | $spl16),
  inference(cnf_transformation,[],[f271_D])).
fof(f271_D,plain,(
  ( ! [X13] : (s(bool,'T_0') != s(bool,X13)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f270,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | s(bool,'T_0') != s(bool,X13)) ) | $spl2),
  inference(subsumption_resolution,[],[f253,f172])).
fof(f253,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | s(bool,'T_0') != s(bool,X13) | s(bool,'F_0') = s(bool,X12)) ) | $spl2),
  inference(superposition,[],[f172,f136])).
fof(f265,plain,(
  ( ! [X61,X62] : (s(bool,'F_0') != s(bool,X61) | p(s(bool,X62))) ) | $spl2),
  inference(subsumption_resolution,[],[f244,f110])).
fof(f110,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f74])).
fof(f74,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (~p(s(bool,'T_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f61])).
fof(f61,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(ennf_transformation,[],[f32])).
fof(f32,plain,(
  ! [X0] : ((p(s(bool,'T_0')) => p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X7] : ((p(s(bool,'T_0')) => p(s(bool,X7))) <=> p(s(bool,X7)))),
  file('thm.measure.UNIV_SIGMA_ALGEBRA','thm.bool.IMP_CLAUSES|split|0')).
fof(f244,plain,(
  ( ! [X61,X62] : (s(bool,'F_0') != s(bool,X61) | p(s(bool,X62)) | ~p(s(bool,'T_0'))) ) | $spl2),
  inference(resolution,[],[f172,f121])).
fof(f121,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f80])).
fof(f80,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f64])).
fof(f64,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X7] : ((p(s(bool,'F_0')) => p(s(bool,X7))) <=> p(s(bool,'T_0')))),
  file('thm.measure.UNIV_SIGMA_ALGEBRA','thm.bool.IMP_CLAUSES|split|2')).
fof(f245,plain,(
  ( ! [X64,X63] : (s(bool,'F_0') != s(bool,X63) | ~p(s(bool,X64))) ) | $spl2),
  inference(resolution,[],[f172,f108])).
fof(f272,plain,(
  ( ! [X19,X18] : (p(s(bool,X18)) | s(bool,'F_0') != s(bool,X19)) ) | $spl2),
  inference(subsumption_resolution,[],[f256,f172])).
fof(f256,plain,(
  ( ! [X19,X18] : (p(s(bool,X18)) | s(bool,'F_0') != s(bool,X19) | s(bool,'T_0') = s(bool,X18)) ) | $spl2),
  inference(superposition,[],[f172,f136])).
fof(f1187,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') = s(bool,X7) | s(bool,X6) != s(bool,X7)) ) | $spl40),
  inference(cnf_transformation,[],[f1187_D])).
fof(f1187_D,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') = s(bool,X7) | s(bool,X6) != s(bool,X7)) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f588,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl23),
  inference(cnf_transformation,[],[f588_D])).
fof(f588_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f1189,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X1) | s(bool,X0) != s(bool,X1)) )),
  inference(subsumption_resolution,[],[f1058,f397])).
fof(f397,plain,(
  ( ! [X47,X46] : (s(bool,'F_0') != s(bool,X46) | s(bool,'F_0') = s(bool,X47) | s(bool,X46) != s(bool,X47)) )),
  inference(resolution,[],[f340,f122])).
fof(f122,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f82])).
fof(f1058,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X1) | s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f397,f136])).
fof(f1188,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X5) | s(bool,X4) != s(bool,X5)) )),
  inference(subsumption_resolution,[],[f1060,f340])).
fof(f1060,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X5) | s(bool,X4) != s(bool,X5) | ~p(s(bool,X4))) )),
  inference(superposition,[],[f397,f129])).
fof(f1186,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X7) | s(bool,X6) != s(bool,X7)) )),
  inference(subsumption_resolution,[],[f1061,f340])).
fof(f1061,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X7) | s(bool,X6) != s(bool,X7) | ~p(s(bool,X6))) )),
  inference(superposition,[],[f397,f127])).
fof(f1185,plain,(
  ( ! [X33,X34,X32] : (s(bool,'F_0') = s(bool,X34) | s(bool,X34) != s(bool,'const.bool.IN_2'(s(X32,X33),s(fun(X32,bool),'const.pred_set.UNIV_0')))) ) | $spl38),
  inference(cnf_transformation,[],[f1185_D])).
fof(f1185_D,plain,(
  ( ! [X33,X34,X32] : (s(bool,'F_0') = s(bool,X34) | s(bool,X34) != s(bool,'const.bool.IN_2'(s(X32,X33),s(fun(X32,bool),'const.pred_set.UNIV_0')))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f1070,plain,(
  ( ! [X33,X34,X32] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X34) | s(bool,X34) != s(bool,'const.bool.IN_2'(s(X32,X33),s(fun(X32,bool),'const.pred_set.UNIV_0')))) )),
  inference(superposition,[],[f397,f684])).
fof(f1184,plain,(
  ( ! [X39,X41,X40] : (s(bool,'F_0') = s(bool,X41) | s(bool,X41) != s(bool,'const.pred_set.SUBSET_2'(s(fun(X39,bool),X40),s(fun(X39,bool),'const.pred_set.UNIV_0')))) ) | $spl36),
  inference(cnf_transformation,[],[f1184_D])).
fof(f1184_D,plain,(
  ( ! [X39,X41,X40] : (s(bool,'F_0') = s(bool,X41) | s(bool,X41) != s(bool,'const.pred_set.SUBSET_2'(s(fun(X39,bool),X40),s(fun(X39,bool),'const.pred_set.UNIV_0')))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f1072,plain,(
  ( ! [X39,X41,X40] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X41) | s(bool,X41) != s(bool,'const.pred_set.SUBSET_2'(s(fun(X39,bool),X40),s(fun(X39,bool),'const.pred_set.UNIV_0')))) ) | $spl30),
  inference(superposition,[],[f397,f896])).
fof(f587,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl22),
  inference(cnf_transformation,[],[f587_D])).
fof(f587_D,plain,(
  s(bool,'T_0') = s(bool,'F_0') <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f1183,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | s(bool,X16) != s(bool,X17) | s(bool,'T_0') = s(bool,X17)) ) | $spl34),
  inference(cnf_transformation,[],[f1183_D])).
fof(f1183_D,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | s(bool,X16) != s(bool,X17) | s(bool,'T_0') = s(bool,X17)) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f1103,plain,(
  ( ! [X14,X15] : (s(bool,X14) != s(bool,X15) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X15) | ~p(s(bool,X14))) )),
  inference(superposition,[],[f399,f129])).
fof(f399,plain,(
  ( ! [X50,X51] : (s(bool,X50) != s(bool,X51) | s(bool,'F_0') = s(bool,X51) | s(bool,'T_0') = s(bool,X50)) )),
  inference(resolution,[],[f340,f127])).
fof(f1104,plain,(
  ( ! [X17,X16] : (s(bool,X16) != s(bool,X17) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X17) | ~p(s(bool,X16))) )),
  inference(superposition,[],[f399,f127])).
fof(f1142,plain,(
  ( ! [X59,X57,X58,X56] : (s(bool,'T_0') = s(bool,X56) | s(bool,X56) != s(bool,X57) | p(s(bool,X58)) | ~p(s(bool,X59)) | s(bool,X57) != s(bool,X59)) ) | $spl23),
  inference(resolution,[],[f1124,f142])).
fof(f1124,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') = s(bool,X15) | s(bool,X14) != s(bool,X15)) ) | $spl23),
  inference(subsumption_resolution,[],[f1103,f588])).
fof(f1133,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s(bool,'T_0') = s(bool,X35) | s(bool,X35) != s(bool,X36) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X38) | s(bool,X37) != s(bool,X39)) ) | $spl23),
  inference(resolution,[],[f1124,f98])).
fof(f1131,plain,(
  ( ! [X28,X26,X29,X27,X25] : (s(bool,'T_0') = s(bool,X25) | s(bool,X25) != s(bool,X26) | p(s(bool,X27)) | ~p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X26) != s(bool,X28)) ) | $spl23),
  inference(resolution,[],[f1124,f97])).
fof(f1129,plain,(
  ( ! [X19,X17,X15,X18,X16] : (s(bool,'T_0') = s(bool,X15) | s(bool,X15) != s(bool,X16) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | $spl23),
  inference(resolution,[],[f1124,f96])).
fof(f1128,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s(bool,'T_0') = s(bool,X10) | s(bool,X10) != s(bool,X11) | ~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,X14)) ) | $spl23),
  inference(resolution,[],[f1124,f94])).
fof(f1126,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X0) != s(bool,X1) | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) ) | $spl23),
  inference(resolution,[],[f1124,f92])).
fof(f896,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(X3,bool),X4),s(fun(X3,bool),'const.pred_set.UNIV_0')))) ) | $spl30),
  inference(cnf_transformation,[],[f896_D])).
fof(f896_D,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(X3,bool),X4),s(fun(X3,bool),'const.pred_set.UNIV_0')))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f895,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X5)) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(X3,bool),X4),s(fun(X3,bool),'const.pred_set.UNIV_0')))) )),
  inference(subsumption_resolution,[],[f880,f123])).
fof(f880,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X5) | p(s(bool,X5)) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(X3,bool),X4),s(fun(X3,bool),'const.pred_set.UNIV_0')))) )),
  inference(superposition,[],[f357,f136])).
fof(f357,plain,(
  ( ! [X85,X86,X84] : (s(bool,X84) != s(bool,'const.pred_set.SUBSET_2'(s(fun(X85,bool),X86),s(fun(X85,bool),'const.pred_set.UNIV_0'))) | p(s(bool,X84))) )),
  inference(resolution,[],[f160,f137])).
fof(f137,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0'))))) )),
  inference(cnf_transformation,[],[f48])).
fof(f48,plain,(
  ! [X0,X1] : p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0'))))),
  inference(flattening,[],[f47])).
fof(f47,plain,(
  ! [X0] : ! [X1] : p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0'))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X6] : ! [X16] : p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X16),s(fun(X6,bool),'const.pred_set.UNIV_0'))))),
  file('thm.measure.UNIV_SIGMA_ALGEBRA','thm.pred_set.SUBSET_UNIV')).
fof(f889,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(X4,bool),X5),s(fun(X4,bool),'const.pred_set.UNIV_0')))) ) | $spl28),
  inference(cnf_transformation,[],[f889_D])).
fof(f889_D,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(X4,bool),X5),s(fun(X4,bool),'const.pred_set.UNIV_0')))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f886,plain,(
  ( ! [X17,X15,X16] : (s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(X16,bool),X17),s(fun(X16,bool),'const.pred_set.UNIV_0'))) | p(s(bool,X15))) )),
  inference(duplicate_literal_removal,[],[f869])).
fof(f869,plain,(
  ( ! [X17,X15,X16] : (s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(X16,bool),X17),s(fun(X16,bool),'const.pred_set.UNIV_0'))) | p(s(bool,X15)) | p(s(bool,X15))) )),
  inference(superposition,[],[f357,f123])).
fof(f887,plain,(
  ( ! [X14,X12,X13] : (s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(X13,bool),X14),s(fun(X13,bool),'const.pred_set.UNIV_0'))) | p(s(bool,X12))) )),
  inference(duplicate_literal_removal,[],[f868])).
fof(f868,plain,(
  ( ! [X14,X12,X13] : (s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(X13,bool),X14),s(fun(X13,bool),'const.pred_set.UNIV_0'))) | p(s(bool,X12)) | p(s(bool,X12))) )),
  inference(superposition,[],[f357,f125])).
fof(f888,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(X4,bool),X5),s(fun(X4,bool),'const.pred_set.UNIV_0'))) | p(s(bool,X3))) )),
  inference(subsumption_resolution,[],[f865,f126])).
fof(f865,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(X4,bool),X5),s(fun(X4,bool),'const.pred_set.UNIV_0'))) | p(s(bool,X3)) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f357,f136])).
fof(f646,plain,(
  ( ! [X83,X84] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(X83,X84),s(fun(X83,bool),'const.pred_set.UNIV_0')))) ) | $spl26),
  inference(cnf_transformation,[],[f646_D])).
fof(f646_D,plain,(
  ( ! [X83,X84] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(X83,X84),s(fun(X83,bool),'const.pred_set.UNIV_0')))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f641,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X33)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(X34,X35),s(fun(X34,bool),'const.pred_set.UNIV_0')))) )),
  inference(duplicate_literal_removal,[],[f636])).
fof(f636,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X33)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(X34,X35),s(fun(X34,bool),'const.pred_set.UNIV_0'))) | p(s(bool,X33))) )),
  inference(superposition,[],[f356,f123])).
fof(f642,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(X31,X32),s(fun(X31,bool),'const.pred_set.UNIV_0')))) )),
  inference(duplicate_literal_removal,[],[f635])).
fof(f635,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(X31,X32),s(fun(X31,bool),'const.pred_set.UNIV_0'))) | p(s(bool,X30))) )),
  inference(superposition,[],[f356,f125])).
fof(f650,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(X28,X29),s(fun(X28,bool),'const.pred_set.UNIV_0')))) )),
  inference(subsumption_resolution,[],[f634,f126])).
fof(f634,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(X28,X29),s(fun(X28,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f356,f136])).
fof(f624,plain,(
  ( ! [X88,X87,X86] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(X86,X87),s(fun(X86,bool),'const.pred_set.UNIV_0'))) | ~p(s(bool,X88))) )),
  inference(resolution,[],[f356,f108])).
fof(f645,plain,(
  ( ! [X85,X83,X84] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(X83,X84),s(fun(X83,bool),'const.pred_set.UNIV_0'))) | p(s(bool,X85))) )),
  inference(subsumption_resolution,[],[f623,f110])).
fof(f623,plain,(
  ( ! [X85,X83,X84] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(X83,X84),s(fun(X83,bool),'const.pred_set.UNIV_0'))) | p(s(bool,X85)) | ~p(s(bool,'T_0'))) )),
  inference(resolution,[],[f356,f121])).
fof(f643,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,X0) != s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.UNIV_0'))) | s(bool,X0) != s(bool,X4) | p(s(bool,X4))) ) | $spl24),
  inference(cnf_transformation,[],[f643_D])).
fof(f643_D,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,X0) != s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.UNIV_0'))) | s(bool,X0) != s(bool,X4) | p(s(bool,X4))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f619,plain,(
  ( ! [X70,X72,X71,X69,X73] : (s(bool,X69) != s(bool,'const.bool.IN_2'(s(X70,X71),s(fun(X70,bool),'const.pred_set.UNIV_0'))) | p(s(bool,X72)) | p(s(bool,X73)) | s(bool,X69) != s(bool,X73)) )),
  inference(resolution,[],[f356,f142])).
fof(f611,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(bool,X42) != s(bool,'const.bool.IN_2'(s(X43,X44),s(fun(X43,bool),'const.pred_set.UNIV_0'))) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X42) != s(bool,X45) | s(bool,X46) != s(bool,X47)) )),
  inference(resolution,[],[f356,f98])).
fof(f608,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (s(bool,X24) != s(bool,'const.bool.IN_2'(s(X25,X26),s(fun(X25,bool),'const.pred_set.UNIV_0'))) | p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X24) != s(bool,X29)) )),
  inference(resolution,[],[f356,f97])).
fof(f607,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,X18) != s(bool,'const.bool.IN_2'(s(X19,X20),s(fun(X19,bool),'const.pred_set.UNIV_0'))) | p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X18) != s(bool,X22)) )),
  inference(resolution,[],[f356,f96])).
fof(f606,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s(bool,X12) != s(bool,'const.bool.IN_2'(s(X13,X14),s(fun(X13,bool),'const.pred_set.UNIV_0'))) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X12) != s(bool,X15) | s(bool,X16) != s(bool,X17)) )),
  inference(resolution,[],[f356,f94])).
fof(f604,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.UNIV_0'))) | p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X0) != s(bool,X4)) )),
  inference(resolution,[],[f356,f92])).
fof(f590,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X10) != s(bool,X11) | p(s(bool,X10))) )),
  inference(subsumption_resolution,[],[f568,f160])).
fof(f568,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X10) != s(bool,X11) | p(s(bool,X11)) | p(s(bool,X10))) )),
  inference(superposition,[],[f342,f123])).
fof(f342,plain,(
  ( ! [X52,X53] : (s(bool,'T_0') != s(bool,X53) | s(bool,X52) != s(bool,X53) | p(s(bool,X52))) )),
  inference(resolution,[],[f160,f126])).
fof(f589,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X8) != s(bool,X9) | p(s(bool,X8))) )),
  inference(subsumption_resolution,[],[f567,f160])).
fof(f567,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X8) != s(bool,X9) | p(s(bool,X9)) | p(s(bool,X8))) )),
  inference(superposition,[],[f342,f125])).
fof(f586,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X2) != s(bool,X3) | p(s(bool,X3))) )),
  inference(subsumption_resolution,[],[f564,f342])).
fof(f564,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X2) != s(bool,X3) | p(s(bool,X3)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f342,f136])).
fof(f425,plain,(
  ( ! [X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X3) | p(s(bool,X3)) | s(bool,'F_0') = s(bool,X1)) ) | $spl20),
  inference(cnf_transformation,[],[f425_D])).
fof(f425_D,plain,(
  ( ! [X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X3) | p(s(bool,X3)) | s(bool,'F_0') = s(bool,X1)) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f401,plain,(
  ( ! [X57,X54,X56,X55] : (s(bool,X54) != s(bool,X55) | s(bool,'F_0') = s(bool,X55) | p(s(bool,X56)) | p(s(bool,X57)) | s(bool,X54) != s(bool,X57)) )),
  inference(resolution,[],[f340,f142])).
fof(f393,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s(bool,X35) != s(bool,X36) | s(bool,'F_0') = s(bool,X36) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X35) != s(bool,X37) | s(bool,X38) != s(bool,X39)) )),
  inference(resolution,[],[f340,f98])).
fof(f390,plain,(
  ( ! [X24,X23,X21,X22,X20] : (s(bool,X20) != s(bool,X21) | s(bool,'F_0') = s(bool,X21) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X20) != s(bool,X24)) )),
  inference(resolution,[],[f340,f97])).
fof(f389,plain,(
  ( ! [X19,X17,X15,X18,X16] : (s(bool,X15) != s(bool,X16) | s(bool,'F_0') = s(bool,X16) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X15) != s(bool,X18)) )),
  inference(resolution,[],[f340,f96])).
fof(f388,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s(bool,X10) != s(bool,X11) | s(bool,'F_0') = s(bool,X11) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X10) != s(bool,X12) | s(bool,X13) != s(bool,X14)) )),
  inference(resolution,[],[f340,f94])).
fof(f386,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X1) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X0) != s(bool,X3)) )),
  inference(resolution,[],[f340,f92])).
fof(f375,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,X3)) | s(bool,X1) != s(bool,X3) | p(s(bool,X0)) | s(bool,X0) != s(bool,X1)) ) | $spl18),
  inference(cnf_transformation,[],[f375_D])).
fof(f375_D,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,X3)) | s(bool,X1) != s(bool,X3) | p(s(bool,X0)) | s(bool,X0) != s(bool,X1)) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f344,plain,(
  ( ! [X59,X57,X58,X56] : (p(s(bool,X56)) | s(bool,X56) != s(bool,X57) | p(s(bool,X58)) | ~p(s(bool,X59)) | s(bool,X57) != s(bool,X59)) )),
  inference(resolution,[],[f160,f142])).
fof(f335,plain,(
  ( ! [X39,X37,X35,X38,X36] : (p(s(bool,X35)) | s(bool,X35) != s(bool,X36) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X38) | s(bool,X37) != s(bool,X39)) )),
  inference(resolution,[],[f160,f98])).
fof(f333,plain,(
  ( ! [X28,X26,X29,X27,X25] : (p(s(bool,X25)) | s(bool,X25) != s(bool,X26) | p(s(bool,X27)) | ~p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X26) != s(bool,X28)) )),
  inference(resolution,[],[f160,f97])).
fof(f331,plain,(
  ( ! [X19,X17,X15,X18,X16] : (p(s(bool,X15)) | s(bool,X15) != s(bool,X16) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) )),
  inference(resolution,[],[f160,f96])).
fof(f330,plain,(
  ( ! [X14,X12,X10,X13,X11] : (p(s(bool,X10)) | s(bool,X10) != s(bool,X11) | ~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f160,f94])).
fof(f328,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,X0)) | s(bool,X0) != s(bool,X1) | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) )),
  inference(resolution,[],[f160,f92])).
fof(f326,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,X19)) ) | $spl8),
  inference(subsumption_resolution,[],[f313,f175])).
fof(f313,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,X19) | s(bool,'T_0') = s(bool,X18)) ) | $spl8),
  inference(superposition,[],[f175,f136])).
fof(f319,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | s(bool,'T_0') != s(bool,X17)) ) | $spl8),
  inference(duplicate_literal_removal,[],[f312])).
fof(f312,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | s(bool,'T_0') != s(bool,X17) | ~p(s(bool,X16))) ) | $spl8),
  inference(superposition,[],[f175,f127])).
fof(f320,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15)) ) | $spl8),
  inference(duplicate_literal_removal,[],[f311])).
fof(f311,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | ~p(s(bool,X14))) ) | $spl8),
  inference(superposition,[],[f175,f129])).
fof(f325,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | s(bool,'T_0') != s(bool,X13)) ) | $spl8),
  inference(subsumption_resolution,[],[f310,f175])).
fof(f310,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | s(bool,'T_0') != s(bool,X13) | s(bool,'F_0') = s(bool,X12)) ) | $spl8),
  inference(superposition,[],[f175,f136])).
fof(f297,plain,(
  ( ! [X74,X73] : (s(bool,'T_0') != s(bool,X73) | ~p(s(bool,X74))) ) | $spl8),
  inference(resolution,[],[f175,f120])).
fof(f120,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f81])).
fof(f295,plain,(
  ( ! [X70,X71] : (s(bool,'T_0') != s(bool,X70) | ~p(s(bool,X71))) ) | $spl8),
  inference(resolution,[],[f175,f117])).
fof(f117,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))))),
  inference(flattening,[],[f78])).
fof(f78,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f63])).
fof(f63,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f34])).
fof(f34,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X7] : ((p(s(bool,X7)) => p(s(bool,X7))) <=> p(s(bool,'T_0')))),
  file('thm.measure.UNIV_SIGMA_ALGEBRA','thm.bool.IMP_CLAUSES|split|3')).
fof(f294,plain,(
  ( ! [X68,X69] : (s(bool,'T_0') != s(bool,X68) | p(s(bool,X69))) ) | $spl8),
  inference(resolution,[],[f175,f116])).
fof(f116,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f79])).
fof(f293,plain,(
  ( ! [X66,X67] : (s(bool,'T_0') != s(bool,X66) | p(s(bool,X67))) ) | $spl8),
  inference(resolution,[],[f175,f113])).
fof(f113,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,'T_0'))))),
  inference(flattening,[],[f76])).
fof(f76,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f62])).
fof(f62,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f33])).
fof(f33,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X7] : ((p(s(bool,X7)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  file('thm.measure.UNIV_SIGMA_ALGEBRA','thm.bool.IMP_CLAUSES|split|1')).
fof(f292,plain,(
  ( ! [X64,X65] : (s(bool,'T_0') != s(bool,X64) | p(s(bool,X65))) ) | $spl8),
  inference(resolution,[],[f175,f110])).
fof(f223,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f203,f177])).
fof(f203,plain,(
  ( ! [X40] : (~p(s(bool,X40))) ) | $spl12),
  inference(resolution,[],[f177,f108])).
fof(f201,plain,(
  $false | $spl12),
  inference(resolution,[],[f177,f130])).
fof(f204,plain,(
  $false | $spl12),
  inference(resolution,[],[f177,f90])).
fof(f171,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.UNIV_0')))))) ) | $spl0),
  inference(cnf_transformation,[],[f171_D])).
fof(f171_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.UNIV_0')))))) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f159,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.UNIV_0')))))) )),
  inference(resolution,[],[f90,f142])).
fof(f150,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.UNIV_0'))))) | s(bool,X21) != s(bool,X23)) )),
  inference(resolution,[],[f90,f98])).
fof(f148,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.UNIV_0')))))) )),
  inference(resolution,[],[f90,f97])).
fof(f173,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.UNIV_0')))))) ) | $spl4),
  inference(cnf_transformation,[],[f173_D])).
fof(f173_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.UNIV_0')))))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f174,plain,(
  ( ! [X4,X3] : (~p(s(bool,X3)) | s(bool,X3) != s(bool,X4) | p(s(bool,X4))) ) | $spl6),
  inference(cnf_transformation,[],[f174_D])).
fof(f174_D,plain,(
  ( ! [X4,X3] : (~p(s(bool,X3)) | s(bool,X3) != s(bool,X4) | p(s(bool,X4))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f147,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.UNIV_0'))))) | s(bool,X12) != s(bool,X13)) )),
  inference(resolution,[],[f90,f96])).
fof(f146,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.UNIV_0')))))) )),
  inference(resolution,[],[f90,f96])).
fof(f145,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.UNIV_0'))))) | s(bool,X6) != s(bool,X8)) )),
  inference(resolution,[],[f90,f94])).
fof(f144,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.UNIV_0'))))) | s(bool,X3) != s(bool,X4)) )),
  inference(resolution,[],[f90,f92])).
fof(f143,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.UNIV_0')))))) )),
  inference(resolution,[],[f90,f92])).
% SZS output end Proof for thm.measure
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 106
% Passive clauses: 678
% Generated clauses: 3510
% Final active clauses: 90
% Final passive clauses: 469
% Input formulas: 24
% Initial clauses: 52
% 
% Selected by SInE selection: 23
% SInE iterations: 1
% 
% Duplicate literals: 61
% Trivial inequalities: 143
% Fw subsumption resolutions: 468
% Fw demodulations: 286
% Bw demodulations: 7
% 
% Simple tautologies: 146
% Forward subsumptions: 1576
% 
% Binary resolution: 663
% Forward superposition: 1297
% Backward superposition: 476
% Self superposition: 2
% Equality resolution: 25
% 
% Split clauses: 204
% Split components: 434
% SAT solver clauses: 218
% SAT solver unit clauses: 12
% SAT solver binary clauses: 122
% 
% Sat splits: 204
% Sat splitting refutations: 13
% 
% Memory used [KB]: 2430
% Time elapsed: 0.099 s
% ------------------------------
% ------------------------------
% Success in time 11.82 s
