Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.pred_set
% (55739)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1175
% Passive clauses: 8788
% Generated clauses: 41697
% Final active clauses: 897
% Final passive clauses: 7048
% Input formulas: 260
% Initial clauses: 576
% 
% Duplicate literals: 297
% Trivial inequalities: 77
% Fw subsumption resolutions: 2293
% Fw demodulations: 3443
% Bw demodulations: 255
% 
% Simple tautologies: 308
% Equational tautologies: 35
% Forward subsumptions: 26260
% Fw demodulations to eq. taut.: 14
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 8793
% Factoring: 1544
% Forward superposition: 11836
% Backward superposition: 12472
% Self superposition: 56
% Equality factoring: 2
% Equality resolution: 19
% 
% Split clauses: 364
% Split components: 730
% SAT solver clauses: 616
% SAT solver unit clauses: 21
% SAT solver binary clauses: 98
% 
% Sat splits: 364
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 611
% TWLsolver calls for satisfiability: 488
% 
% Memory used [KB]: 19445
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pred_set
% (55750)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 662
% Passive clauses: 3380
% Generated clauses: 27521
% Final active clauses: 552
% Final passive clauses: 2199
% Input formulas: 260
% Initial clauses: 576
% 
% Duplicate literals: 102
% Trivial inequalities: 66
% Fw subsumption resolutions: 1076
% Fw demodulations: 3800
% Bw demodulations: 586
% Condensations: 89
% Global subsumptions: 25
% 
% Simple tautologies: 328
% Equational tautologies: 9
% Forward subsumptions: 14729
% Fw demodulations to eq. taut.: 21
% Bw demodulations to eq. taut.: 14
% 
% Binary resolution: 2602
% Factoring: 162
% Forward superposition: 7742
% Backward superposition: 6802
% Self superposition: 82
% Equality factoring: 2
% Equality resolution: 10
% 
% Split clauses: 19
% Split components: 38
% SAT solver clauses: 4227
% SAT solver unit clauses: 665
% SAT solver binary clauses: 211
% SAT solver learnt clauses: 54
% SAT solver learnt literals: 19
% 
% Sat splits: 19
% Sat splitting refutations: 57
% 
% TWLsolver clauses: 4123
% TWLsolver calls for satisfiability: 3956
% 
% Memory used [KB]: 22387
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.pred_set
% (55770)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1439
% Passive clauses: 10128
% Generated clauses: 219525
% Final active clauses: 917
% Final passive clauses: 4263
% Input formulas: 260
% Initial clauses: 576
% Discarded non-redundant clauses: 847
% 
% Duplicate literals: 11
% Trivial inequalities: 65
% Fw subsumption resolutions: 2175
% Fw demodulations: 98800
% Bw demodulations: 3188
% Condensations: 207
% Global subsumptions: 786
% 
% Simple tautologies: 201
% Equational tautologies: 52
% Forward subsumptions: 59995
% Fw demodulations to eq. taut.: 45972
% Bw demodulations to eq. taut.: 14
% 
% Unit resulting resolution: 69349
% Factoring: 3
% Forward superposition: 27558
% Backward superposition: 13988
% Self superposition: 130
% Equality factoring: 2
% Equality resolution: 23
% 
% Split clauses: 29
% Split components: 58
% SAT solver clauses: 11859
% SAT solver unit clauses: 2819
% SAT solver binary clauses: 4989
% 
% Sat splits: 29
% Sat splitting refutations: 885
% 
% Memory used [KB]: 225028
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f53949,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f622,f916,f621,f915,f924,f925,f926,f927,f928,f931,f932,f933,f935,f936,f937,f940,f1038,f1101,f1102,f1059,f1108,f1109,f1062,f1152,f1153,f1154,f1157,f1161,f1162,f1163,f1165,f1148,f1175,f1177,f1178,f1180,f1182,f1280,f1479,f1457,f1579,f1580,f1581,f1582,f1604,f1605,f1606,f1608,f1609,f1610,f1613,f1711,f1753,f1755,f1756,f1758,f1760,f1858,f1911,f1912,f1913,f1915,f1916,f1917,f1920,f2018,f2058,f2060,f2061,f2063,f2065,f2163,f2217,f1158,f2218,f1155,f1159,f2219,f2221,f2222,f2223,f2226,f2324,f2368,f2387,f2406,f2426,f2636,f2614,f2736,f2737,f2738,f2739,f2754,f2772,f2791,f2811,f3001,f2999,f2986,f2984,f2983,f3105,f3106,f3115,f3171,f3181,f3191,f3201,f3211,f3221,f3231,f3242,f1156,f3271,f3290,f3309,f3329,f3349,f3370,f3382,f3404,f3421,f3439,f3451,f3474,f3486,f3509,f3521,f3544,f3565,f3588,f3609,f3621,f3646,f3667,f3679,f3703,f3723,f3745,f3767,f3786,f3798,f3825,f3843,f3855,f3882,f3903,f3930,f3948,f3960,f3987,f4012,f4034,f4046,f4070,f4092,f4115,f4137,f4160,f4178,f4190,f4215,f4236,f4259,f4280,f4299,f4320,f4328,f4343,f4351,f4362,f4370,f4385,f4393,f4404,f4406,f4407,f4409,f4411,f4509,f4528,f4572,f4593,f4595,f4596,f4598,f4600,f4698,f4717,f5123,f5207,f5206,f5497,f5205,f5204,f5499,f5203,f5202,f5500,f5201,f5200,f5501,f5199,f5198,f5503,f5197,f5196,f5504,f5195,f5194,f5505,f5193,f5192,f5506,f5191,f5190,f5507,f5189,f5188,f5508,f5187,f5186,f5509,f5185,f5184,f5510,f5183,f5182,f5511,f5181,f5180,f5512,f5179,f5178,f5513,f5177,f5176,f5514,f5175,f5174,f5515,f5173,f5172,f5516,f5167,f5166,f5517,f5165,f5164,f5518,f5163,f5162,f5519,f5161,f5160,f5520,f5159,f5158,f5521,f5157,f5156,f5522,f5472,f3160,f5528,f5530,f5531,f5533,f5535,f5633,f1315,f5679,f2377,f5697,f2396,f5715,f2416,f5732,f5733,f5734,f5735,f5741,f5743,f5744,f5745,f5748,f5846,f5867,f5905,f2435,f5923,f2763,f5940,f2782,f5958,f2802,f5977,f2822,f6032,f6024,f6041,f3126,f6104,f3554,f6125,f3575,f6146,f6158,f6181,f3598,f6204,f3656,f6225,f6237,f6260,f6282,f6303,f6323,f6335,f6358,f6380,f6403,f6425,f6448,f6467,f6479,f6504,f6514,f6526,f6537,f6549,f6560,f6569,f6581,f6602,f6610,f6619,f6627,f6636,f6648,f6669,f6681,f6703,f6711,f6720,f6728,f6741,f3834,f6758,f6770,f6798,f3891,f6819,f3914,f6846,f3939,f6863,f6875,f6902,f3998,f6928,f4023,f6949,f6961,f6983,f4080,f7004,f4102,f7028,f4124,f7049,f4147,f7073,f4169,f7090,f7102,f7125,f4225,f7147,f4244,f7170,f4267,f7188,f7190,f7191,f7193,f7195,f7293,f7312,f7356,f4288,f7376,f4307,f7397,f7399,f7400,f7402,f7404,f7502,f7564,f7565,f7567,f7568,f7570,f5523,f7572,f7573,f1090,f7591,f7592,f7593,f7595,f7596,f7597,f7600,f7698,f7742,f7751,f7761,f7770,f7780,f7789,f7799,f7809,f7820,f7822,f7823,f7825,f7827,f7925,f7944,f7978,f7987,f7997,f8006,f8016,f8025,f8036,f8045,f8056,f8065,f8076,f8462,f8464,f8111,f8485,f8112,f8487,f8494,f8496,f8499,f8502,f8500,f8503,f8505,f8504,f8506,f8508,f8507,f8509,f8511,f8510,f8512,f8514,f8513,f8515,f8517,f8516,f8522,f8524,f8523,f8525,f8527,f8526,f8528,f8531,f8529,f8532,f8534,f8533,f8535,f8537,f8536,f8538,f8540,f8539,f8541,f8543,f8542,f8544,f8546,f8545,f8547,f8549,f8548,f8550,f8552,f8551,f8553,f8555,f8554,f8556,f8558,f8557,f8559,f8561,f8560,f8562,f8564,f8563,f8501,f8565,f8567,f8566,f8568,f8570,f8569,f8571,f8574,f8573,f8575,f8578,f8577,f8789,f8790,f8812,f8821,f8813,f8890,f8898,f8907,f8915,f8923,f8924,f8925,f8926,f9085,f8932,f8934,f8935,f8936,f8939,f9037,f9058,f9092,f9139,f9141,f9129,f9495,f9851,f1747,f9852,f1110,f1746,f9853,f9855,f9856,f9857,f9860,f9958,f1748,f10002,f10011,f3180,f10021,f10030,f3200,f10040,f10049,f3220,f10059,f10069,f3241,f10077,f10078,f10079,f10080,f1086,f10239,f1088,f1094,f10086,f10088,f10089,f10090,f10093,f10191,f10212,f10246,f1106,f10255,f10265,f3281,f10274,f10284,f3300,f10293,f10304,f3320,f10313,f10324,f3340,f10333,f10344,f3360,f10391,f8463,f10393,f10381,f10402,f10410,f3412,f10419,f10427,f3429,f1111,f10473,f3391,f10485,f3395,f10581,f3460,f10593,f3464,f10659,f3495,f10671,f3499,f10695,f3531,f10707,f3535,f10754,f10756,f10757,f10758,f10761,f10859,f10897,f10868,f8530,f11020,f3631,f11032,f3635,f11048,f11058,f3713,f11077,f3733,f11098,f3755,f11119,f3777,f11137,f3689,f11149,f3693,f11173,f3809,f11185,f3813,f11210,f3866,f11222,f3870,f11248,f11261,f11274,f11288,f11301,f11314,f11326,f3971,f11338,f3975,f11364,f11378,f11390,f11404,f11416,f11430,f11442,f11457,f11469,f11484,f11496,f4054,f11508,f4058,f11530,f11545,f11557,f11569,f11581,f11593,f11605,f4198,f11617,f4202,f11644,f4582,f11666,f6270,f11687,f6291,f11708,f6313,f11728,f11730,f11731,f11733,f11735,f11833,f4554,f11852,f1105,f11900,f6368,f11924,f6390,f11947,f6412,f11970,f6435,f11993,f6457,f1104,f12013,f12015,f12016,f12018,f12020,f12118,f7552,f12261,f12263,f12264,f12266,f12268,f12366,f12416,f13899,f13901,f13902,f13903,f13906,f14004,f14043,f14052,f14054,f14055,f14057,f14059,f14157,f14194,f14204,f14206,f14207,f14208,f14211,f14309,f14348,f14357,f14359,f14360,f14361,f14364,f14462,f14501,f15895,f15896,f15899,f15900,f15902,f15904,f15906,f15908,f15909,f16159,f16161,f16162,f16164,f16166,f16264,f16448,f16450,f16451,f16453,f16455,f16553,f16650,f17505,f17538,f17707,f17783,f17814,f18083,f18113,f18147,f18177,f18297,f18328,f18470,f18543,f18574,f18825,f18855,f18889,f18919,f20408,f20412,f20415,f20417,f20421,f20478,f20503,f20504,f8498,f20509,f8815,f20510,f20514,f20517,f20519,f20522,f20525,f20623,f20409,f20637,f20636,f20634,f19606,f20807,f19553,f20808,f19500,f20809,f19449,f20810,f19204,f20811,f19162,f20812,f19120,f20813,f19080,f20814,f17289,f20815,f17266,f20816,f17209,f20817,f17186,f20818,f17163,f20819,f17134,f20820,f17105,f20821,f17030,f20822,f16999,f20823,f16964,f20824,f16933,f20825,f16895,f20826,f16817,f20827,f16786,f20828,f16759,f20829,f16728,f20830,f16068,f16067,f16064,f16063,f1092,f15844,f20831,f15777,f20832,f15756,f20833,f15732,f20834,f15708,f20835,f15684,f20836,f15615,f20837,f15592,f20838,f15568,f20839,f15545,f20840,f15522,f20841,f15167,f20842,f15143,f20843,f15078,f20844,f15054,f20845,f15031,f20846,f15008,f20847,f14985,f20848,f14918,f20849,f14895,f20850,f14619,f20851,f14589,f20852,f14559,f20853,f14531,f20854,f14503,f20855,f14197,f20856,f13882,f20857,f13836,f20858,f13790,f20859,f13744,f20860,f13699,f20861,f13654,f20862,f13609,f20863,f13563,f20864,f13518,f20865,f12956,f20866,f12919,f20867,f12882,f20868,f12807,f20869,f12771,f20870,f12735,f20871,f12699,f20872,f12662,f20873,f12608,f20874,f12419,f20875,f11884,f20876,f11595,f20877,f11571,f20878,f11547,f20879,f11486,f20880,f11459,f20881,f11432,f20882,f11406,f20883,f11380,f20884,f11316,f20885,f11290,f20886,f11263,f20887,f10899,f20888,f3261,f20889,f3142,f20890,f1091,f20897,f20899,f20900,f20902,f20904,f21002,f2198,f21044,f21038,f21213,f21232,f21234,f21235,f21236,f21239,f21337,f2053,f21389,f21391,f21392,f21393,f21396,f21494,f1095,f21692,f1168,f21683,f1167,f21786,f21788,f21789,f21790,f21793,f21891,f21917,f21964,f21966,f21967,f21969,f21971,f22069,f22107,f22119,f22151,f22153,f22154,f22155,f22158,f22256,f22297,f23097,f25079,f25081,f25082,f25084,f25086,f25184,f16361,f26281,f26290,f26344,f26350,f26360,f26366,f26374,f26382,f26391,f26532,f26565,f26761,f26842,f26875,f27132,f27164,f27200,f27232,f17531,f27510,f27512,f27513,f27515,f27517,f27615,f27713,f29321,f29323,f29324,f29326,f29328,f29426,f29526,f29639,f29652,f29822,f29841,f29848,f29861,f29868,f29879,f29901,f29903,f29904,f29905,f29908,f30006,f30074,f30083,f30094,f29643,f30304,f31363,f31365,f31364,f31366,f31368,f31367,f31369,f31371,f31370,f31372,f31374,f31373,f31375,f31377,f31376,f31378,f31380,f31379,f31385,f31387,f31386,f31388,f31390,f31389,f31391,f31393,f31392,f31394,f31396,f31395,f31397,f31399,f31398,f31400,f31402,f31401,f31403,f31405,f31404,f31406,f31408,f31407,f31409,f31411,f31410,f31412,f31414,f31413,f31415,f31417,f31416,f31418,f31420,f31419,f31421,f31423,f31422,f31424,f31426,f31425,f5502,f31427,f31429,f31428,f31430,f31432,f31431,f31433,f31436,f31435,f31437,f31440,f31439,f5495,f31731,f31733,f31734,f31736,f31738,f31836,f31873,f33243,f29900,f1169,f914,f29899,f913,f1170,f33589,f33591,f33592,f33594,f33596,f33694,f4743,f34860,f34862,f34863,f34864,f34867,f34965,f35007,f35016,f35018,f35019,f35021,f35023,f35121,f35158,f38231,f38227,f38414,f38378,f38376,f38374,f38368,f38367,f38366,f38365,f38356,f38355,f38487,f38488,f38498,f38497,f38506,f38505,f38518,f38513,f38529,f38527,f38539,f38535,f38550,f38547,f38560,f38556,f38571,f38568,f39644,f39740,f39939,f40085,f40130,f40596,f40640,f40690,f40734,f45191,f45194,f45193,f45436,f45439,f45438,f45665,f45668,f45667,f46868,f46870,f46871,f46872,f46875,f46973,f47104,f47614,f47618,f47643,f47647,f47794,f47796,f47857,f47861,f47886,f47890,f48179,f48183,f48207,f48211,f48241,f48245,f48269,f38233,f48273,f39678,f48419,f48421,f48422,f48424,f48426,f48524,f48587,f48712,f48714,f48715,f48717,f48719,f48817,f48881,f53115,f53117,f53118,f53120,f53122,f53220,f53318,f53349,f53351,f53352,f53354,f53356,f53454,f53535,f53584,f53586,f53587,f1096,f53588,f1098,f53591,f1097,f53689,f1099,f53786,f53942])).
fof(f53942,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f53789,f33584])).
fof(f33584,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK13,sK39(sK15,sK14,sK13)),s(fun(sK13,bool),sK15)))) | $spl3),
  inference(resolution,[],[f916,f851])).
fof(f851,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,sK39(X2,X1,X0)),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f606])).
fof(f606,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))))) & ((p(s(bool,'const.bool.IN_2'(s(X0,sK39(X2,X1,X0)),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,sK39(X2,X1,X0)),s(fun(X0,bool),X2))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK39])],[f605])).
fof(f605,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))))) & (? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f604])).
fof(f604,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & (? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(nnf_transformation,[],[f437])).
fof(f437,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(ennf_transformation,[],[f338])).
fof(f338,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) => p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(flattening,[],[f337])).
fof(f337,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) => p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X6] : ! [X7,X8] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))) <=> ! [X9] : (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X7)))) => p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X8))))))),
  file('thm.pred_set.SUBSET_DIFF_EMPTY','thm.pred_set.SUBSET_DEF')).
fof(f53789,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK13,sK39(sK15,sK14,sK13)),s(fun(sK13,bool),sK15)))) | ($spl0 | $spl3)),
  inference(resolution,[],[f34306,f33583])).
fof(f33583,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK13,sK39(sK15,sK14,sK13)),s(fun(sK13,bool),sK14)))) | $spl3),
  inference(resolution,[],[f916,f850])).
fof(f850,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK39(X2,X1,X0)),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f606])).
fof(f34306,plain,(
  ( ! [X2] : (~p(s(bool,'const.bool.IN_2'(s(sK13,X2),s(fun(sK13,bool),sK14)))) | p(s(bool,'const.bool.IN_2'(s(sK13,X2),s(fun(sK13,bool),sK15))))) ) | $spl0),
  inference(subsumption_resolution,[],[f33942,f892])).
fof(f892,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(cnf_transformation,[],[f402])).
fof(f402,plain,(
  ! [X0,X1] : ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),
  inference(flattening,[],[f401])).
fof(f401,plain,(
  ! [X0] : ! [X1] : ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X6] : ! [X9] : ~p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),'const.pred_set.EMPTY_0'))))),
  file('thm.pred_set.SUBSET_DIFF_EMPTY','thm.pred_set.NOT_IN_EMPTY')).
fof(f33942,plain,(
  ( ! [X2] : (p(s(bool,'const.bool.IN_2'(s(sK13,X2),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,'const.bool.IN_2'(s(sK13,X2),s(fun(sK13,bool),sK15)))) | ~p(s(bool,'const.bool.IN_2'(s(sK13,X2),s(fun(sK13,bool),sK14))))) ) | $spl0),
  inference(superposition,[],[f874,f913])).
fof(f874,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f614])).
fof(f614,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))))),
  inference(flattening,[],[f613])).
fof(f613,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))))),
  inference(nnf_transformation,[],[f366])).
fof(f366,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(flattening,[],[f365])).
fof(f365,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6] : ! [X7,X8,X9] : (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),'const.pred_set.DIFF_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X7)))) & ~p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X8))))))),
  file('thm.pred_set.SUBSET_DIFF_EMPTY','thm.pred_set.IN_DIFF')).
fof(f53786,plain,(
  ( ! [X38,X36] : (p(s(bool,X36)) | s(bool,X36) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X38)))))) ) | $spl676),
  inference(cnf_transformation,[],[f53786_D])).
fof(f53786_D,plain,(
  ( ! [X38,X36] : (p(s(bool,X36)) | s(bool,X36) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X38)))))) ) <=> ~$spl676),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl676])])).
fof(f1099,plain,(
  ( ! [X15] : (p(s(bool,X15))) ) | $spl22),
  inference(cnf_transformation,[],[f1099_D])).
fof(f1099_D,plain,(
  ( ! [X15] : (p(s(bool,X15))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f53689,plain,(
  ( ! [X350,X349,X351] : (p(s(bool,X349)) | p(s(bool,X350)) | s(bool,X350) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X351)))))) ) | $spl8),
  inference(resolution,[],[f7586,f905])).
fof(f905,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f627])).
fof(f627,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f481])).
fof(f481,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f480])).
fof(f480,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f417])).
fof(f417,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f416])).
fof(f416,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f264])).
fof(f264,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f71])).
fof(f71,axiom,(
  ! [X9,X26,X16,X27] : ((s(bool,X9) = s(bool,X26) & (p(s(bool,X26)) => s(bool,X16) = s(bool,X27))) => ((p(s(bool,X9)) => p(s(bool,X16))) <=> (p(s(bool,X26)) => p(s(bool,X27)))))),
  file('thm.pred_set.SUBSET_DIFF_EMPTY','thm.bool.IMP_CONG')).
fof(f7586,plain,(
  ( ! [X11] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X11))))))) ) | $spl8),
  inference(resolution,[],[f1089,f888])).
fof(f888,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f454])).
fof(f454,plain,(
  ! [X0,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) | ! [X2] : p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))))),
  inference(ennf_transformation,[],[f394])).
fof(f394,plain,(
  ! [X0,X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) => ! [X2] : p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))))),
  inference(flattening,[],[f393])).
fof(f393,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) => ! [X2] : p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X6] : ! [X7] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X7)))) => ! [X8] : p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),'const.pred_set.DIFF_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))))))),
  file('thm.pred_set.SUBSET_DIFF_EMPTY','thm.pred_set.FINITE_DIFF')).
fof(f1089,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) | $spl8),
  inference(cnf_transformation,[],[f1089_D])).
fof(f1089_D,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f1097,plain,(
  ( ! [X10,X11] : (s(bool,X10) != s(bool,X11) | ~p(s(bool,X10))) ) | $spl18),
  inference(cnf_transformation,[],[f1097_D])).
fof(f1097_D,plain,(
  ( ! [X10,X11] : (s(bool,X10) != s(bool,X11) | ~p(s(bool,X10))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f53591,plain,(
  ( ! [X64,X62,X65,X63] : (p(s(bool,X62)) | ~p(s(bool,X63)) | s(bool,X62) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X64))))) | s(bool,X63) != s(bool,X65)) ) | $spl8),
  inference(resolution,[],[f7586,f630])).
fof(f630,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f481])).
fof(f1098,plain,(
  ( ! [X16] : (~p(s(bool,X16))) ) | $spl20),
  inference(cnf_transformation,[],[f1098_D])).
fof(f1098_D,plain,(
  ( ! [X16] : (~p(s(bool,X16))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f53588,plain,(
  ( ! [X54,X52,X53,X51] : (p(s(bool,X51)) | ~p(s(bool,X52)) | p(s(bool,X53)) | s(bool,X53) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X54)))))) ) | $spl8),
  inference(resolution,[],[f7586,f629])).
fof(f629,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f481])).
fof(f1096,plain,(
  ( ! [X5,X3] : (s(bool,X3) != s(bool,X5) | p(s(bool,X3))) ) | $spl16),
  inference(cnf_transformation,[],[f1096_D])).
fof(f1096_D,plain,(
  ( ! [X5,X3] : (s(bool,X3) != s(bool,X5) | p(s(bool,X3))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f53587,plain,(
  ( ! [X47,X50,X48,X49] : (p(s(bool,X47)) | p(s(bool,X48)) | s(bool,X47) != s(bool,X49) | s(bool,X48) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X50)))))) ) | $spl8),
  inference(resolution,[],[f7586,f628])).
fof(f628,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f481])).
fof(f53586,plain,(
  ( ! [X45,X43,X46,X44] : (p(s(bool,X43)) | ~p(s(bool,X44)) | s(bool,X43) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X45))))) | s(bool,X44) != s(bool,X46)) ) | $spl8),
  inference(resolution,[],[f7586,f626])).
fof(f626,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f481])).
fof(f53584,plain,(
  ( ! [X37,X35,X38,X36] : (p(s(bool,X35)) | p(s(bool,X36)) | s(bool,X35) != s(bool,X37) | s(bool,X36) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X38)))))) ) | $spl8),
  inference(resolution,[],[f7586,f624])).
fof(f624,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f481])).
fof(f53535,plain,(
  ( ! [X18,X16] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X16)))) | ~p(s(bool,X18)) | s(bool,X18) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) | $spl674),
  inference(cnf_transformation,[],[f53535_D])).
fof(f53535_D,plain,(
  ( ! [X18,X16] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X16)))) | ~p(s(bool,X18)) | s(bool,X18) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) <=> ~$spl674),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl674])])).
fof(f53454,plain,(
  ( ! [X333,X335,X334] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X333)))) | p(s(bool,X334)) | ~p(s(bool,X335)) | s(bool,X335) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X333),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) | $spl0),
  inference(resolution,[],[f34478,f905])).
fof(f34478,plain,(
  ( ! [X249] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X249),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X249))))) ) | $spl0),
  inference(forward_demodulation,[],[f34477,f913])).
fof(f34477,plain,(
  ( ! [X249] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X249)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X249),s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))))))) ) | $spl0),
  inference(subsumption_resolution,[],[f34107,f902])).
fof(f902,plain,(
  ( ! [X0] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(cnf_transformation,[],[f412])).
fof(f412,plain,(
  ! [X0] : p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),
  inference(rectify,[],[f125])).
fof(f125,axiom,(
  ! [X6] : p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),'const.pred_set.EMPTY_0'))))),
  file('thm.pred_set.SUBSET_DIFF_EMPTY','thm.pred_set.FINITE_EMPTY')).
fof(f34107,plain,(
  ( ! [X249] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X249)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X249),s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))))))) ) | $spl0),
  inference(superposition,[],[f860,f913])).
fof(f860,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f441])).
fof(f441,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))) | p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))))),
  inference(flattening,[],[f440])).
fof(f440,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))))) | p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))))),
  inference(ennf_transformation,[],[f348])).
fof(f348,plain,(
  ! [X0,X1,X2] : ((p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))))) => p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))))),
  inference(flattening,[],[f347])).
fof(f347,plain,(
  ! [X0] : ! [X1,X2] : ((p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))))) => p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))))),
  inference(rectify,[],[f119])).
fof(f119,axiom,(
  ! [X6] : ! [X7,X8] : ((p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X7)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X7))))) => p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X8)))))),
  file('thm.pred_set.SUBSET_DIFF_EMPTY','thm.pred_set.SUBSET_FINITE_I')).
fof(f53356,plain,(
  ( ! [X47,X45,X46,X44] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X44)))) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X46) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X44),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s(bool,X45) != s(bool,X47)) ) | $spl0),
  inference(resolution,[],[f34478,f630])).
fof(f53354,plain,(
  ( ! [X39,X37,X38,X36] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X36)))) | p(s(bool,X37)) | ~p(s(bool,X38)) | ~p(s(bool,X39)) | s(bool,X38) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X36),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) | $spl0),
  inference(resolution,[],[f34478,f629])).
fof(f53352,plain,(
  ( ! [X30,X28,X31,X29] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X28)))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X28),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) | $spl0),
  inference(resolution,[],[f34478,f628])).
fof(f53351,plain,(
  ( ! [X26,X24,X27,X25] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X24)))) | ~p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X26) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X24),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s(bool,X25) != s(bool,X27)) ) | $spl0),
  inference(resolution,[],[f34478,f626])).
fof(f53349,plain,(
  ( ! [X19,X17,X18,X16] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X16)))) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) | $spl0),
  inference(resolution,[],[f34478,f624])).
fof(f53318,plain,(
  ( ! [X30,X29,X27] : (~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.INSERT_2'(s(sK13,X29),s(fun(sK13,bool),X30))),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) | $spl672),
  inference(cnf_transformation,[],[f53318_D])).
fof(f53318_D,plain,(
  ( ! [X30,X29,X27] : (~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.INSERT_2'(s(sK13,X29),s(fun(sK13,bool),X30))),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) <=> ~$spl672),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl672])])).
fof(f53220,plain,(
  ( ! [X447,X449,X446,X448] : (p(s(bool,X446)) | ~p(s(bool,X447)) | s(bool,X447) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.INSERT_2'(s(sK13,X448),s(fun(sK13,bool),X449))),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) | $spl0),
  inference(resolution,[],[f34387,f905])).
fof(f34387,plain,(
  ( ! [X161,X160] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.INSERT_2'(s(sK13,X160),s(fun(sK13,bool),X161))),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl0),
  inference(forward_demodulation,[],[f34386,f913])).
fof(f34386,plain,(
  ( ! [X161,X160] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.INSERT_2'(s(sK13,X160),s(fun(sK13,bool),X161))),s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))))))) ) | $spl0),
  inference(subsumption_resolution,[],[f34032,f892])).
fof(f34032,plain,(
  ( ! [X161,X160] : (p(s(bool,'const.bool.IN_2'(s(sK13,X160),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.INSERT_2'(s(sK13,X160),s(fun(sK13,bool),X161))),s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))))))) ) | $spl0),
  inference(superposition,[],[f842,f913])).
fof(f842,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f601])).
fof(f601,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X3)))) | (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))))) & (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X3))))))),
  inference(flattening,[],[f600])).
fof(f600,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X3)))) | (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))))) & ((~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X3))))))),
  inference(nnf_transformation,[],[f332])).
fof(f332,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X3)))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))))),
  inference(flattening,[],[f331])).
fof(f331,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X3)))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f254])).
fof(f254,axiom,(
  ! [X6] : ! [X9,X7,X8] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),'const.pred_set.INSERT_2'(s(X6,X9),s(fun(X6,bool),X7))),s(fun(X6,bool),X8)))) <=> (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X8)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8))))))),
  file('thm.pred_set.SUBSET_DIFF_EMPTY','thm.pred_set.INSERT_SUBSET')).
fof(f53122,plain,(
  ( ! [X61,X59,X62,X60,X63] : (~p(s(bool,X59)) | ~p(s(bool,X60)) | s(bool,X60) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.INSERT_2'(s(sK13,X61),s(fun(sK13,bool),X62))),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s(bool,X59) != s(bool,X63)) ) | $spl0),
  inference(resolution,[],[f34387,f630])).
fof(f53120,plain,(
  ( ! [X54,X52,X55,X53,X51] : (p(s(bool,X51)) | ~p(s(bool,X52)) | ~p(s(bool,X53)) | s(bool,X52) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.INSERT_2'(s(sK13,X54),s(fun(sK13,bool),X55))),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) | $spl0),
  inference(resolution,[],[f34387,f629])).
fof(f53118,plain,(
  ( ! [X45,X43,X41,X44,X42] : (p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X41) != s(bool,X43) | s(bool,X42) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.INSERT_2'(s(sK13,X44),s(fun(sK13,bool),X45))),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) | $spl0),
  inference(resolution,[],[f34387,f628])).
fof(f53117,plain,(
  ( ! [X39,X37,X38,X36,X40] : (~p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X37) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.INSERT_2'(s(sK13,X38),s(fun(sK13,bool),X39))),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s(bool,X36) != s(bool,X40)) ) | $spl0),
  inference(resolution,[],[f34387,f626])).
fof(f53115,plain,(
  ( ! [X30,X28,X26,X29,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,X28) | s(bool,X27) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.INSERT_2'(s(sK13,X29),s(fun(sK13,bool),X30))),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) | $spl0),
  inference(resolution,[],[f34387,f624])).
fof(f48881,plain,(
  ( ! [X14,X16] : (~p(s(bool,X14)) | s(bool,X14) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.INSERT_2'(s(sK13,X16),s(fun(sK13,bool),sK14))),s(fun(sK13,bool),sK15)))) ) | $spl670),
  inference(cnf_transformation,[],[f48881_D])).
fof(f48881_D,plain,(
  ( ! [X14,X16] : (~p(s(bool,X14)) | s(bool,X14) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.INSERT_2'(s(sK13,X16),s(fun(sK13,bool),sK14))),s(fun(sK13,bool),sK15)))) ) <=> ~$spl670),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl670])])).
fof(f48817,plain,(
  ( ! [X329,X331,X330] : (p(s(bool,X329)) | ~p(s(bool,X330)) | s(bool,X330) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.INSERT_2'(s(sK13,X331),s(fun(sK13,bool),sK14))),s(fun(sK13,bool),sK15)))) ) | $spl3),
  inference(resolution,[],[f33580,f905])).
fof(f33580,plain,(
  ( ! [X1] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.INSERT_2'(s(sK13,X1),s(fun(sK13,bool),sK14))),s(fun(sK13,bool),sK15))))) ) | $spl3),
  inference(resolution,[],[f916,f843])).
fof(f843,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f601])).
fof(f48719,plain,(
  ( ! [X43,X41,X42,X40] : (~p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X41) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.INSERT_2'(s(sK13,X42),s(fun(sK13,bool),sK14))),s(fun(sK13,bool),sK15))) | s(bool,X40) != s(bool,X43)) ) | $spl3),
  inference(resolution,[],[f33580,f630])).
fof(f48717,plain,(
  ( ! [X35,X33,X36,X34] : (p(s(bool,X33)) | ~p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X34) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.INSERT_2'(s(sK13,X36),s(fun(sK13,bool),sK14))),s(fun(sK13,bool),sK15)))) ) | $spl3),
  inference(resolution,[],[f33580,f629])).
fof(f48715,plain,(
  ( ! [X28,X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.INSERT_2'(s(sK13,X28),s(fun(sK13,bool),sK14))),s(fun(sK13,bool),sK15)))) ) | $spl3),
  inference(resolution,[],[f33580,f628])).
fof(f48714,plain,(
  ( ! [X24,X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.INSERT_2'(s(sK13,X23),s(fun(sK13,bool),sK14))),s(fun(sK13,bool),sK15))) | s(bool,X21) != s(bool,X24)) ) | $spl3),
  inference(resolution,[],[f33580,f626])).
fof(f48712,plain,(
  ( ! [X14,X15,X13,X16] : (p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.INSERT_2'(s(sK13,X16),s(fun(sK13,bool),sK14))),s(fun(sK13,bool),sK15)))) ) | $spl3),
  inference(resolution,[],[f33580,f624])).
fof(f48587,plain,(
  ( ! [X14,X16] : (~p(s(bool,X14)) | s(bool,X14) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.DELETE_2'(s(fun(sK13,bool),sK15),s(sK13,X16)))))) ) | $spl668),
  inference(cnf_transformation,[],[f48587_D])).
fof(f48587_D,plain,(
  ( ! [X14,X16] : (~p(s(bool,X14)) | s(bool,X14) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.DELETE_2'(s(fun(sK13,bool),sK15),s(sK13,X16)))))) ) <=> ~$spl668),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl668])])).
fof(f48524,plain,(
  ( ! [X329,X331,X330] : (p(s(bool,X329)) | ~p(s(bool,X330)) | s(bool,X330) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.DELETE_2'(s(fun(sK13,bool),sK15),s(sK13,X331)))))) ) | $spl3),
  inference(resolution,[],[f33579,f905])).
fof(f33579,plain,(
  ( ! [X0] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.DELETE_2'(s(fun(sK13,bool),sK15),s(sK13,X0))))))) ) | $spl3),
  inference(resolution,[],[f916,f840])).
fof(f840,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X3),s(X0,X1))))))) )),
  inference(cnf_transformation,[],[f599])).
fof(f599,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X3),s(X0,X1)))))) | (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))))) & (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X3),s(X0,X1))))))))),
  inference(flattening,[],[f598])).
fof(f598,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X3),s(X0,X1)))))) | (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))))) & ((p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X3),s(X0,X1))))))))),
  inference(nnf_transformation,[],[f330])).
fof(f330,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X3),s(X0,X1)))))) <=> (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))))),
  inference(flattening,[],[f329])).
fof(f329,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X3),s(X0,X1)))))) <=> (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f251])).
fof(f251,axiom,(
  ! [X6] : ! [X9,X7,X8] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),'const.pred_set.DELETE_2'(s(fun(X6,bool),X8),s(X6,X9)))))) <=> (~p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X7)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8))))))),
  file('thm.pred_set.SUBSET_DIFF_EMPTY','thm.pred_set.SUBSET_DELETE')).
fof(f48426,plain,(
  ( ! [X43,X41,X42,X40] : (~p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X41) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.DELETE_2'(s(fun(sK13,bool),sK15),s(sK13,X42))))) | s(bool,X40) != s(bool,X43)) ) | $spl3),
  inference(resolution,[],[f33579,f630])).
fof(f48424,plain,(
  ( ! [X35,X33,X36,X34] : (p(s(bool,X33)) | ~p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X34) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.DELETE_2'(s(fun(sK13,bool),sK15),s(sK13,X36)))))) ) | $spl3),
  inference(resolution,[],[f33579,f629])).
fof(f48422,plain,(
  ( ! [X28,X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.DELETE_2'(s(fun(sK13,bool),sK15),s(sK13,X28)))))) ) | $spl3),
  inference(resolution,[],[f33579,f628])).
fof(f48421,plain,(
  ( ! [X24,X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.DELETE_2'(s(fun(sK13,bool),sK15),s(sK13,X23))))) | s(bool,X21) != s(bool,X24)) ) | $spl3),
  inference(resolution,[],[f33579,f626])).
fof(f48419,plain,(
  ( ! [X14,X15,X13,X16] : (p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.DELETE_2'(s(fun(sK13,bool),sK15),s(sK13,X16)))))) ) | $spl3),
  inference(resolution,[],[f33579,f624])).
fof(f39678,plain,(
  ( ! [X19] : (s(bool,'const.pred_set.FINITE_1'(s(fun(X19,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl652),
  inference(cnf_transformation,[],[f39678_D])).
fof(f39678_D,plain,(
  ( ! [X19] : (s(bool,'const.pred_set.FINITE_1'(s(fun(X19,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl652),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl652])])).
fof(f48273,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,'const.pred_set.FINITE_1'(s(fun(X18,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl304)),
  inference(resolution,[],[f21846,f29381])).
fof(f29381,plain,(
  ( ! [X251,X249,X250,X248] : (sP7(X250,X248,X251) | s(bool,X248) = s(bool,'const.pred_set.FINITE_1'(s(fun(X249,bool),'const.pred_set.EMPTY_0')))) ) | $spl8),
  inference(resolution,[],[f28487,f695])).
fof(f695,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f513])).
fof(f513,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f512])).
fof(f512,plain,(
  ! [X1,X0,X2] : ((~sP7(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP7(X1,X0,X2)))),
  inference(flattening,[],[f511])).
fof(f511,plain,(
  ! [X1,X0,X2] : ((~sP7(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP7(X1,X0,X2)))),
  inference(nnf_transformation,[],[f468])).
fof(f468,plain,(
  ! [X1,X0,X2] : (sP7(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f28487,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,X18) = s(bool,'const.pred_set.FINITE_1'(s(fun(X19,bool),'const.pred_set.EMPTY_0')))) ) | $spl8),
  inference(subsumption_resolution,[],[f27734,f1089])).
fof(f27734,plain,(
  ( ! [X19,X18] : (s(bool,X18) = s(bool,'const.pred_set.FINITE_1'(s(fun(X19,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) | ~p(s(bool,X18))) ) | $spl8),
  inference(superposition,[],[f16588,f774])).
fof(f774,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f546])).
fof(f546,plain,(
  ! [X0,X1] : ((s(bool,X0) != s(bool,X1) | ((~p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | p(s(bool,X0))))) & ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(flattening,[],[f545])).
fof(f545,plain,(
  ! [X0,X1] : ((s(bool,X0) != s(bool,X1) | ((~p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | p(s(bool,X0))))) & (((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0)))) | s(bool,X0) = s(bool,X1)))),
  inference(nnf_transformation,[],[f424])).
fof(f424,plain,(
  ! [X0,X1] : (s(bool,X0) = s(bool,X1) <=> ((~p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | p(s(bool,X0)))))),
  inference(ennf_transformation,[],[f287])).
fof(f287,plain,(
  ! [X0,X1] : (s(bool,X0) = s(bool,X1) <=> ((p(s(bool,X0)) => p(s(bool,X1))) & (p(s(bool,X1)) => p(s(bool,X0)))))),
  inference(rectify,[],[f122])).
fof(f122,axiom,(
  ! [X10,X11] : (s(bool,X10) = s(bool,X11) <=> ((p(s(bool,X10)) => p(s(bool,X11))) & (p(s(bool,X11)) => p(s(bool,X10)))))),
  file('thm.pred_set.SUBSET_DIFF_EMPTY','thm.bool.EQ_IMP_THM')).
fof(f16588,plain,(
  ( ! [X419] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) = s(bool,'const.pred_set.FINITE_1'(s(fun(X419,bool),'const.pred_set.EMPTY_0')))) ) | $spl8),
  inference(resolution,[],[f7688,f902])).
fof(f7688,plain,(
  ( ! [X204] : (~p(s(bool,X204)) | s(bool,X204) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | $spl8),
  inference(resolution,[],[f1089,f774])).
fof(f21846,plain,(
  ( ! [X128,X129] : (~sP7(X128,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X129) | p(s(bool,X129)) | ~p(s(bool,X128))) ) | $spl304),
  inference(resolution,[],[f8498,f693])).
fof(f693,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f513])).
fof(f38233,plain,(
  s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | $spl650),
  inference(cnf_transformation,[],[f38233_D])).
fof(f38233_D,plain,(
  s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) <=> ~$spl650),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl650])])).
fof(f48269,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | ($spl10 | $spl304)),
  inference(resolution,[],[f21846,f25139])).
fof(f25139,plain,(
  ( ! [X189,X190,X188] : (sP7(X189,X188,X190) | s(bool,X188) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl10),
  inference(resolution,[],[f21327,f695])).
fof(f21327,plain,(
  ( ! [X204] : (~p(s(bool,X204)) | s(bool,X204) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl10),
  inference(resolution,[],[f1092,f774])).
fof(f48245,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,'const.pred_set.FINITE_1'(s(fun(X21,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl304)),
  inference(resolution,[],[f21845,f29382])).
fof(f29382,plain,(
  ( ! [X255,X253,X254,X252] : (sP7(X252,X254,X255) | s(bool,X252) = s(bool,'const.pred_set.FINITE_1'(s(fun(X253,bool),'const.pred_set.EMPTY_0')))) ) | $spl8),
  inference(resolution,[],[f28487,f696])).
fof(f696,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f513])).
fof(f21845,plain,(
  ( ! [X127,X126] : (~sP7('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X126,X127) | p(s(bool,X127)) | ~p(s(bool,X126))) ) | $spl304),
  inference(resolution,[],[f8498,f693])).
fof(f48241,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | ($spl10 | $spl304)),
  inference(resolution,[],[f21845,f25140])).
fof(f25140,plain,(
  ( ! [X191,X192,X193] : (sP7(X191,X192,X193) | s(bool,X191) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl10),
  inference(resolution,[],[f21327,f696])).
fof(f48211,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,'const.pred_set.FINITE_1'(s(fun(X18,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl304)),
  inference(resolution,[],[f21843,f29378])).
fof(f29378,plain,(
  ( ! [X239,X237,X238,X236] : (sP8(X238,X236,X239) | s(bool,X236) = s(bool,'const.pred_set.FINITE_1'(s(fun(X237,bool),'const.pred_set.EMPTY_0')))) ) | $spl8),
  inference(resolution,[],[f28487,f691])).
fof(f691,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f510])).
fof(f510,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f509])).
fof(f509,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f508])).
fof(f508,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f469])).
fof(f469,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f21843,plain,(
  ( ! [X123,X122] : (~sP8(X122,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X123) | p(s(bool,X123)) | ~p(s(bool,X122))) ) | $spl304),
  inference(resolution,[],[f8498,f689])).
fof(f689,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f510])).
fof(f48207,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | ($spl10 | $spl304)),
  inference(resolution,[],[f21843,f25136])).
fof(f25136,plain,(
  ( ! [X180,X181,X179] : (sP8(X180,X179,X181) | s(bool,X179) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl10),
  inference(resolution,[],[f21327,f691])).
fof(f48183,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,'const.pred_set.FINITE_1'(s(fun(X23,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl304)),
  inference(resolution,[],[f21842,f29379])).
fof(f29379,plain,(
  ( ! [X243,X241,X242,X240] : (sP8(X240,X242,X243) | s(bool,X240) = s(bool,'const.pred_set.FINITE_1'(s(fun(X241,bool),'const.pred_set.EMPTY_0')))) ) | $spl8),
  inference(resolution,[],[f28487,f692])).
fof(f692,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f510])).
fof(f21842,plain,(
  ( ! [X121,X120] : (~sP8('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X120,X121) | p(s(bool,X121)) | ~p(s(bool,X120))) ) | $spl304),
  inference(resolution,[],[f8498,f689])).
fof(f48179,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | ($spl10 | $spl304)),
  inference(resolution,[],[f21842,f25137])).
fof(f25137,plain,(
  ( ! [X182,X184,X183] : (sP8(X182,X183,X184) | s(bool,X182) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl10),
  inference(resolution,[],[f21327,f692])).
fof(f47890,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,'const.pred_set.FINITE_1'(s(fun(X18,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl304)),
  inference(resolution,[],[f21830,f29365])).
fof(f29365,plain,(
  ( ! [X187,X185,X186,X184] : (sP4(X186,X184,X187) | s(bool,X184) = s(bool,'const.pred_set.FINITE_1'(s(fun(X185,bool),'const.pred_set.EMPTY_0')))) ) | $spl8),
  inference(resolution,[],[f28487,f674])).
fof(f674,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f502])).
fof(f502,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(flattening,[],[f501])).
fof(f501,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(nnf_transformation,[],[f464])).
fof(f464,plain,(
  ! [X0,X1,X2] : (sP4(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f21830,plain,(
  ( ! [X97,X96] : (~sP4(X96,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X97) | p(s(bool,X97)) | ~p(s(bool,X96))) ) | $spl304),
  inference(resolution,[],[f8498,f672])).
fof(f672,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f502])).
fof(f47886,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | ($spl10 | $spl304)),
  inference(resolution,[],[f21830,f25123])).
fof(f25123,plain,(
  ( ! [X142,X140,X141] : (sP4(X141,X140,X142) | s(bool,X140) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl10),
  inference(resolution,[],[f21327,f674])).
fof(f47861,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,'const.pred_set.FINITE_1'(s(fun(X21,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl304)),
  inference(resolution,[],[f21829,f29366])).
fof(f29366,plain,(
  ( ! [X191,X189,X190,X188] : (sP4(X188,X190,X191) | s(bool,X188) = s(bool,'const.pred_set.FINITE_1'(s(fun(X189,bool),'const.pred_set.EMPTY_0')))) ) | $spl8),
  inference(resolution,[],[f28487,f675])).
fof(f675,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f502])).
fof(f21829,plain,(
  ( ! [X94,X95] : (~sP4('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X94,X95) | p(s(bool,X95)) | ~p(s(bool,X94))) ) | $spl304),
  inference(resolution,[],[f8498,f672])).
fof(f47857,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | ($spl10 | $spl304)),
  inference(resolution,[],[f21829,f25124])).
fof(f25124,plain,(
  ( ! [X144,X145,X143] : (sP4(X143,X144,X145) | s(bool,X143) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl10),
  inference(resolution,[],[f21327,f675])).
fof(f47796,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,'const.pred_set.FINITE_1'(s(fun(X14,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl304)),
  inference(resolution,[],[f21813,f29352])).
fof(f29352,plain,(
  ( ! [X134,X132,X135,X133] : (sP2(X132,X134,X135) | s(bool,X132) = s(bool,'const.pred_set.FINITE_1'(s(fun(X133,bool),'const.pred_set.EMPTY_0')))) ) | $spl8),
  inference(resolution,[],[f28487,f658])).
fof(f658,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f495])).
fof(f495,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(rectify,[],[f494])).
fof(f494,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(flattening,[],[f493])).
fof(f493,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(nnf_transformation,[],[f461])).
fof(f461,plain,(
  ! [X2,X0,X1] : (sP2(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f21813,plain,(
  ( ! [X62,X63] : (~sP2('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X62,X63) | p(s(bool,X63)) | p(s(bool,X62))) ) | $spl304),
  inference(resolution,[],[f8498,f655])).
fof(f655,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f495])).
fof(f47794,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | p(s(bool,X9)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | ($spl10 | $spl304)),
  inference(resolution,[],[f21813,f25110])).
fof(f25110,plain,(
  ( ! [X103,X101,X102] : (sP2(X101,X102,X103) | s(bool,X101) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl10),
  inference(resolution,[],[f21327,f658])).
fof(f47647,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,'const.pred_set.FINITE_1'(s(fun(X18,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl304)),
  inference(resolution,[],[f21800,f29335])).
fof(f29335,plain,(
  ( ! [X66,X64,X67,X65] : (sP0(X66,X67,X64) | s(bool,X64) = s(bool,'const.pred_set.FINITE_1'(s(fun(X65,bool),'const.pred_set.EMPTY_0')))) ) | $spl8),
  inference(resolution,[],[f28487,f639])).
fof(f639,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f487])).
fof(f487,plain,(
  ! [X0,X1,X2] : ((~sP0(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP0(X0,X1,X2)))),
  inference(rectify,[],[f486])).
fof(f486,plain,(
  ! [X2,X0,X1] : ((~sP0(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP0(X2,X0,X1)))),
  inference(flattening,[],[f485])).
fof(f485,plain,(
  ! [X2,X0,X1] : ((~sP0(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP0(X2,X0,X1)))),
  inference(nnf_transformation,[],[f458])).
fof(f458,plain,(
  ! [X2,X0,X1] : (sP0(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f21800,plain,(
  ( ! [X37,X36] : (~sP0(X36,X37,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X37)) | ~p(s(bool,X36))) ) | $spl304),
  inference(resolution,[],[f8498,f638])).
fof(f638,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP0(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f487])).
fof(f47643,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | ($spl10 | $spl304)),
  inference(resolution,[],[f21800,f25093])).
fof(f25093,plain,(
  ( ! [X52,X50,X51] : (sP0(X51,X52,X50) | s(bool,X50) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl10),
  inference(resolution,[],[f21327,f639])).
fof(f47618,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,X19)) | p(s(bool,X20)) | s(bool,'const.pred_set.FINITE_1'(s(fun(X21,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl304)),
  inference(resolution,[],[f21799,f29336])).
fof(f29336,plain,(
  ( ! [X70,X68,X71,X69] : (sP0(X68,X70,X71) | s(bool,X68) = s(bool,'const.pred_set.FINITE_1'(s(fun(X69,bool),'const.pred_set.EMPTY_0')))) ) | $spl8),
  inference(resolution,[],[f28487,f641])).
fof(f641,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f487])).
fof(f21799,plain,(
  ( ! [X35,X34] : (~sP0('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X34,X35) | ~p(s(bool,X35)) | p(s(bool,X34))) ) | $spl304),
  inference(resolution,[],[f8498,f638])).
fof(f47614,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | p(s(bool,X11)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | ($spl10 | $spl304)),
  inference(resolution,[],[f21799,f25094])).
fof(f25094,plain,(
  ( ! [X54,X55,X53] : (sP0(X53,X54,X55) | s(bool,X53) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl10),
  inference(resolution,[],[f21327,f641])).
fof(f47104,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) | $spl666),
  inference(cnf_transformation,[],[f47104_D])).
fof(f47104_D,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) <=> ~$spl666),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl666])])).
fof(f46973,plain,(
  ( ! [X315,X316,X317] : (s(bool,X315) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))) | p(s(bool,X316)) | p(s(bool,X317)) | s(bool,X315) != s(bool,X317)) ) | $spl3),
  inference(resolution,[],[f33682,f905])).
fof(f33682,plain,(
  ( ! [X186] : (p(s(bool,X186)) | s(bool,X186) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))) ) | $spl3),
  inference(resolution,[],[f916,f771])).
fof(f771,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f546])).
fof(f46875,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) ) | $spl3),
  inference(resolution,[],[f33682,f630])).
fof(f46872,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,X16) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) ) | $spl3),
  inference(resolution,[],[f33682,f629])).
fof(f46871,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl3),
  inference(resolution,[],[f33682,f628])).
fof(f46870,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) ) | $spl3),
  inference(resolution,[],[f33682,f626])).
fof(f46868,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl3),
  inference(resolution,[],[f33682,f624])).
fof(f45667,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.UNION_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))))) | $spl663),
  inference(cnf_transformation,[],[f45667_D])).
fof(f45667_D,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.UNION_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))))) <=> ~$spl663),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl663])])).
fof(f45668,plain,(
  ( ! [X5] : (s(fun(sK13,bool),sK15) != s(fun(sK13,bool),X5) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.UNION_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),'const.pred_set.UNION_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))))))) ) | $spl664),
  inference(cnf_transformation,[],[f45668_D])).
fof(f45668_D,plain,(
  ( ! [X5] : (s(fun(sK13,bool),sK15) != s(fun(sK13,bool),X5) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.UNION_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),'const.pred_set.UNION_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))))))) ) <=> ~$spl664),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl664])])).
fof(f45665,plain,(
  ( ! [X5] : (s(fun(sK13,bool),sK15) != s(fun(sK13,bool),X5) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),'const.pred_set.UNION_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.UNION_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.UNION_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))))))) ) | $spl3),
  inference(subsumption_resolution,[],[f45650,f860])).
fof(f45650,plain,(
  ( ! [X5] : (s(fun(sK13,bool),sK15) != s(fun(sK13,bool),X5) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),'const.pred_set.UNION_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.UNION_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.UNION_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X5))))) ) | $spl3),
  inference(superposition,[],[f33585,f866])).
fof(f866,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X2))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f444])).
fof(f444,plain,(
  ! [X0,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) | ! [X2] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X2))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(flattening,[],[f443])).
fof(f443,plain,(
  ! [X0,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) | ! [X2] : ((~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X2))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(ennf_transformation,[],[f356])).
fof(f356,plain,(
  ! [X0,X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) => ! [X2] : ((p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) & s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X2))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) => s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(flattening,[],[f355])).
fof(f355,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) => ! [X2] : ((p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) & s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X2))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) => s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(rectify,[],[f111])).
fof(f111,axiom,(
  ! [X6] : ! [X7] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X7)))) => ! [X8] : ((p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X8)))) & s('type.num.num','const.pred_set.CARD_1'(s(fun(X6,bool),X7))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(X6,bool),X8))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8))))) => s(fun(X6,bool),X7) = s(fun(X6,bool),X8)))),
  file('thm.pred_set.SUBSET_DIFF_EMPTY','thm.pred_set.SUBSET_EQ_CARD')).
fof(f33585,plain,(
  s(fun(sK13,bool),sK15) != s(fun(sK13,bool),'const.pred_set.UNION_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))) | $spl3),
  inference(resolution,[],[f916,f853])).
fof(f853,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | s(fun(X0,bool),X2) != s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) )),
  inference(cnf_transformation,[],[f607])).
fof(f607,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | s(fun(X0,bool),X2) = s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) & (s(fun(X0,bool),X2) != s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(nnf_transformation,[],[f340])).
fof(f340,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> s(fun(X0,bool),X2) = s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))),
  inference(flattening,[],[f339])).
fof(f339,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> s(fun(X0,bool),X2) = s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))),
  inference(rectify,[],[f110])).
fof(f110,axiom,(
  ! [X6] : ! [X7,X8] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))) <=> s(fun(X6,bool),X8) = s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8))))),
  file('thm.pred_set.SUBSET_DIFF_EMPTY','thm.pred_set.SUBSET_UNION_ABSORPTION')).
fof(f45438,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))))) | $spl659),
  inference(cnf_transformation,[],[f45438_D])).
fof(f45438_D,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))))) <=> ~$spl659),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl659])])).
fof(f45439,plain,(
  ( ! [X5] : (s(fun(sK13,bool),sK14) != s(fun(sK13,bool),X5) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))))))) ) | $spl660),
  inference(cnf_transformation,[],[f45439_D])).
fof(f45439_D,plain,(
  ( ! [X5] : (s(fun(sK13,bool),sK14) != s(fun(sK13,bool),X5) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))))))) ) <=> ~$spl660),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl660])])).
fof(f45436,plain,(
  ( ! [X5] : (s(fun(sK13,bool),sK14) != s(fun(sK13,bool),X5) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))))))) ) | $spl3),
  inference(subsumption_resolution,[],[f45420,f860])).
fof(f45420,plain,(
  ( ! [X5] : (s(fun(sK13,bool),sK14) != s(fun(sK13,bool),X5) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X5))))) ) | $spl3),
  inference(superposition,[],[f33582,f866])).
fof(f33582,plain,(
  s(fun(sK13,bool),sK14) != s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))) | $spl3),
  inference(resolution,[],[f916,f847])).
fof(f847,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))) | s(fun(X0,bool),X2) != s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))) )),
  inference(cnf_transformation,[],[f603])).
fof(f603,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X2) != s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))))) & (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))) | s(fun(X0,bool),X2) = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))))),
  inference(nnf_transformation,[],[f336])).
fof(f336,plain,(
  ! [X0,X1,X2] : (s(fun(X0,bool),X2) = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))) <=> p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f335])).
fof(f335,plain,(
  ! [X0] : ! [X1,X2] : (s(fun(X0,bool),X2) = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))) <=> p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f115])).
fof(f115,axiom,(
  ! [X6] : ! [X15,X46] : (s(fun(X6,bool),X46) = s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),X46),s(fun(X6,bool),X15))) <=> p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X46),s(fun(X6,bool),X15)))))),
  file('thm.pred_set.SUBSET_DIFF_EMPTY','thm.pred_set.INTER_SUBSET_EQN|split|0')).
fof(f45193,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)))))) | $spl655),
  inference(cnf_transformation,[],[f45193_D])).
fof(f45193_D,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)))))) <=> ~$spl655),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl655])])).
fof(f45194,plain,(
  ( ! [X5] : (s(fun(sK13,bool),sK14) != s(fun(sK13,bool),X5) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))))))) ) | $spl656),
  inference(cnf_transformation,[],[f45194_D])).
fof(f45194_D,plain,(
  ( ! [X5] : (s(fun(sK13,bool),sK14) != s(fun(sK13,bool),X5) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))))))) ) <=> ~$spl656),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl656])])).
fof(f45191,plain,(
  ( ! [X5] : (s(fun(sK13,bool),sK14) != s(fun(sK13,bool),X5) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))))))) ) | $spl3),
  inference(subsumption_resolution,[],[f45176,f860])).
fof(f45176,plain,(
  ( ! [X5] : (s(fun(sK13,bool),sK14) != s(fun(sK13,bool),X5) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X5))))) ) | $spl3),
  inference(superposition,[],[f33581,f866])).
fof(f33581,plain,(
  s(fun(sK13,bool),sK14) != s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))) | $spl3),
  inference(resolution,[],[f916,f845])).
fof(f845,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | s(fun(X0,bool),X1) != s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))) )),
  inference(cnf_transformation,[],[f602])).
fof(f602,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) & (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))))),
  inference(nnf_transformation,[],[f334])).
fof(f334,plain,(
  ! [X0,X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))) <=> p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))),
  inference(flattening,[],[f333])).
fof(f333,plain,(
  ! [X0] : ! [X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))) <=> p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))),
  inference(rectify,[],[f118])).
fof(f118,axiom,(
  ! [X6] : ! [X15,X46] : (s(fun(X6,bool),X15) = s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),X46),s(fun(X6,bool),X15))) <=> p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X15),s(fun(X6,bool),X46)))))),
  file('thm.pred_set.SUBSET_DIFF_EMPTY','thm.pred_set.INTER_SUBSET_EQN|split|1')).
fof(f40734,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,'const.pred_set.FINITE_1'(s(fun(X17,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl304)),
  inference(resolution,[],[f28677,f27570])).
fof(f27570,plain,(
  ( ! [X189,X190,X188] : (sP7(X189,X188,X190) | s(bool,X188) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl304),
  inference(resolution,[],[f21881,f695])).
fof(f21881,plain,(
  ( ! [X188] : (~p(s(bool,X188)) | s(bool,X188) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl304),
  inference(resolution,[],[f8498,f774])).
fof(f28677,plain,(
  ( ! [X494,X493,X495] : (~sP7(X494,'const.pred_set.FINITE_1'(s(fun(X493,bool),'const.pred_set.EMPTY_0')),X495) | p(s(bool,X495)) | ~p(s(bool,X494))) ) | $spl8),
  inference(subsumption_resolution,[],[f28286,f1089])).
fof(f28286,plain,(
  ( ! [X494,X493,X495] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) | ~p(s(bool,X494)) | p(s(bool,X495)) | ~sP7(X494,'const.pred_set.FINITE_1'(s(fun(X493,bool),'const.pred_set.EMPTY_0')),X495)) ) | $spl8),
  inference(superposition,[],[f693,f16588])).
fof(f40690,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,'const.pred_set.FINITE_1'(s(fun(X19,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl304)),
  inference(resolution,[],[f28676,f27571])).
fof(f27571,plain,(
  ( ! [X191,X192,X193] : (sP7(X191,X192,X193) | s(bool,X191) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl304),
  inference(resolution,[],[f21881,f696])).
fof(f28676,plain,(
  ( ! [X492,X490,X491] : (~sP7('const.pred_set.FINITE_1'(s(fun(X490,bool),'const.pred_set.EMPTY_0')),X491,X492) | p(s(bool,X492)) | ~p(s(bool,X491))) ) | $spl8),
  inference(subsumption_resolution,[],[f28285,f1089])).
fof(f28285,plain,(
  ( ! [X492,X490,X491] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) | ~p(s(bool,X491)) | p(s(bool,X492)) | ~sP7('const.pred_set.FINITE_1'(s(fun(X490,bool),'const.pred_set.EMPTY_0')),X491,X492)) ) | $spl8),
  inference(superposition,[],[f693,f16588])).
fof(f40640,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,'const.pred_set.FINITE_1'(s(fun(X17,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl304)),
  inference(resolution,[],[f28675,f27567])).
fof(f27567,plain,(
  ( ! [X180,X181,X179] : (sP8(X180,X179,X181) | s(bool,X179) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl304),
  inference(resolution,[],[f21881,f691])).
fof(f28675,plain,(
  ( ! [X476,X475,X477] : (~sP8(X476,'const.pred_set.FINITE_1'(s(fun(X475,bool),'const.pred_set.EMPTY_0')),X477) | p(s(bool,X477)) | ~p(s(bool,X476))) ) | $spl8),
  inference(subsumption_resolution,[],[f28280,f1089])).
fof(f28280,plain,(
  ( ! [X476,X475,X477] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) | ~p(s(bool,X476)) | p(s(bool,X477)) | ~sP8(X476,'const.pred_set.FINITE_1'(s(fun(X475,bool),'const.pred_set.EMPTY_0')),X477)) ) | $spl8),
  inference(superposition,[],[f689,f16588])).
fof(f40596,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,'const.pred_set.FINITE_1'(s(fun(X21,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl304)),
  inference(resolution,[],[f28674,f27568])).
fof(f27568,plain,(
  ( ! [X182,X184,X183] : (sP8(X182,X183,X184) | s(bool,X182) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl304),
  inference(resolution,[],[f21881,f692])).
fof(f28674,plain,(
  ( ! [X472,X474,X473] : (~sP8('const.pred_set.FINITE_1'(s(fun(X472,bool),'const.pred_set.EMPTY_0')),X473,X474) | p(s(bool,X474)) | ~p(s(bool,X473))) ) | $spl8),
  inference(subsumption_resolution,[],[f28279,f1089])).
fof(f28279,plain,(
  ( ! [X472,X474,X473] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) | ~p(s(bool,X473)) | p(s(bool,X474)) | ~sP8('const.pred_set.FINITE_1'(s(fun(X472,bool),'const.pred_set.EMPTY_0')),X473,X474)) ) | $spl8),
  inference(superposition,[],[f689,f16588])).
fof(f40130,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,'const.pred_set.FINITE_1'(s(fun(X17,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl304)),
  inference(resolution,[],[f28662,f27554])).
fof(f27554,plain,(
  ( ! [X142,X140,X141] : (sP4(X141,X140,X142) | s(bool,X140) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl304),
  inference(resolution,[],[f21881,f674])).
fof(f28662,plain,(
  ( ! [X397,X399,X398] : (~sP4(X398,'const.pred_set.FINITE_1'(s(fun(X397,bool),'const.pred_set.EMPTY_0')),X399) | p(s(bool,X399)) | ~p(s(bool,X398))) ) | $spl8),
  inference(subsumption_resolution,[],[f28254,f1089])).
fof(f28254,plain,(
  ( ! [X397,X399,X398] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) | ~p(s(bool,X398)) | p(s(bool,X399)) | ~sP4(X398,'const.pred_set.FINITE_1'(s(fun(X397,bool),'const.pred_set.EMPTY_0')),X399)) ) | $spl8),
  inference(superposition,[],[f672,f16588])).
fof(f40085,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,'const.pred_set.FINITE_1'(s(fun(X19,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl304)),
  inference(resolution,[],[f28661,f27555])).
fof(f27555,plain,(
  ( ! [X144,X145,X143] : (sP4(X143,X144,X145) | s(bool,X143) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl304),
  inference(resolution,[],[f21881,f675])).
fof(f28661,plain,(
  ( ! [X395,X396,X394] : (~sP4('const.pred_set.FINITE_1'(s(fun(X394,bool),'const.pred_set.EMPTY_0')),X395,X396) | p(s(bool,X396)) | ~p(s(bool,X395))) ) | $spl8),
  inference(subsumption_resolution,[],[f28253,f1089])).
fof(f28253,plain,(
  ( ! [X395,X396,X394] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) | ~p(s(bool,X395)) | p(s(bool,X396)) | ~sP4('const.pred_set.FINITE_1'(s(fun(X394,bool),'const.pred_set.EMPTY_0')),X395,X396)) ) | $spl8),
  inference(superposition,[],[f672,f16588])).
fof(f39939,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,'const.pred_set.FINITE_1'(s(fun(X15,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl304)),
  inference(resolution,[],[f28651,f27541])).
fof(f27541,plain,(
  ( ! [X103,X101,X102] : (sP2(X101,X102,X103) | s(bool,X101) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl304),
  inference(resolution,[],[f21881,f658])).
fof(f28651,plain,(
  ( ! [X308,X309,X307] : (~sP2('const.pred_set.FINITE_1'(s(fun(X307,bool),'const.pred_set.EMPTY_0')),X308,X309) | p(s(bool,X309)) | p(s(bool,X308))) ) | $spl8),
  inference(subsumption_resolution,[],[f28224,f1089])).
fof(f28224,plain,(
  ( ! [X308,X309,X307] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) | p(s(bool,X308)) | p(s(bool,X309)) | ~sP2('const.pred_set.FINITE_1'(s(fun(X307,bool),'const.pred_set.EMPTY_0')),X308,X309)) ) | $spl8),
  inference(superposition,[],[f655,f16588])).
fof(f39740,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,'const.pred_set.FINITE_1'(s(fun(X17,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl304)),
  inference(resolution,[],[f28640,f27524])).
fof(f27524,plain,(
  ( ! [X52,X50,X51] : (sP0(X51,X52,X50) | s(bool,X50) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl304),
  inference(resolution,[],[f21881,f639])).
fof(f28640,plain,(
  ( ! [X218,X219,X217] : (~sP0(X218,X219,'const.pred_set.FINITE_1'(s(fun(X217,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X219)) | ~p(s(bool,X218))) ) | $spl8),
  inference(subsumption_resolution,[],[f28194,f1089])).
fof(f28194,plain,(
  ( ! [X218,X219,X217] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) | ~p(s(bool,X218)) | p(s(bool,X219)) | ~sP0(X218,X219,'const.pred_set.FINITE_1'(s(fun(X217,bool),'const.pred_set.EMPTY_0')))) ) | $spl8),
  inference(superposition,[],[f638,f16588])).
fof(f39644,plain,(
  ( ! [X19,X17,X18] : (~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,'const.pred_set.FINITE_1'(s(fun(X19,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl304)),
  inference(resolution,[],[f28639,f27525])).
fof(f27525,plain,(
  ( ! [X54,X55,X53] : (sP0(X53,X54,X55) | s(bool,X53) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl304),
  inference(resolution,[],[f21881,f641])).
fof(f28639,plain,(
  ( ! [X216,X215,X214] : (~sP0('const.pred_set.FINITE_1'(s(fun(X214,bool),'const.pred_set.EMPTY_0')),X215,X216) | ~p(s(bool,X216)) | p(s(bool,X215))) ) | $spl8),
  inference(subsumption_resolution,[],[f28193,f1089])).
fof(f28193,plain,(
  ( ! [X216,X215,X214] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) | p(s(bool,X215)) | ~p(s(bool,X216)) | ~sP0('const.pred_set.FINITE_1'(s(fun(X214,bool),'const.pred_set.EMPTY_0')),X215,X216)) ) | $spl8),
  inference(superposition,[],[f638,f16588])).
fof(f38568,plain,(
  ( ! [X19,X18] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X18)) | ~p(s(bool,X19))) ) | ($spl10 | $spl304)),
  inference(resolution,[],[f27571,f21291])).
fof(f21291,plain,(
  ( ! [X142,X143] : (~sP7('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X142,X143) | p(s(bool,X143)) | ~p(s(bool,X142))) ) | $spl10),
  inference(resolution,[],[f1092,f693])).
fof(f38571,plain,(
  ( ! [X17,X16] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X16)) | ~p(s(bool,X17))) ) | ($spl8 | $spl10 | $spl304)),
  inference(forward_demodulation,[],[f38567,f21342])).
fof(f21342,plain,(
  s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) | ($spl8 | $spl10)),
  inference(resolution,[],[f1092,f7688])).
fof(f38567,plain,(
  ( ! [X17,X16] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X16)) | ~p(s(bool,X17))) ) | ($spl8 | $spl304)),
  inference(resolution,[],[f27571,f7652])).
fof(f7652,plain,(
  ( ! [X142,X143] : (~sP7('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)),X142,X143) | p(s(bool,X143)) | ~p(s(bool,X142))) ) | $spl8),
  inference(resolution,[],[f1089,f693])).
fof(f38556,plain,(
  ( ! [X14,X15] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X14)) | ~p(s(bool,X15))) ) | ($spl10 | $spl304)),
  inference(resolution,[],[f27570,f21292])).
fof(f21292,plain,(
  ( ! [X144,X145] : (~sP7(X144,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X145) | p(s(bool,X145)) | ~p(s(bool,X144))) ) | $spl10),
  inference(resolution,[],[f1092,f693])).
fof(f38560,plain,(
  ( ! [X12,X13] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X12)) | ~p(s(bool,X13))) ) | ($spl8 | $spl10 | $spl304)),
  inference(forward_demodulation,[],[f38555,f21342])).
fof(f38555,plain,(
  ( ! [X12,X13] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X12)) | ~p(s(bool,X13))) ) | ($spl8 | $spl304)),
  inference(resolution,[],[f27570,f7653])).
fof(f7653,plain,(
  ( ! [X144,X145] : (~sP7(X144,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)),X145) | p(s(bool,X145)) | ~p(s(bool,X144))) ) | $spl8),
  inference(resolution,[],[f1089,f693])).
fof(f38547,plain,(
  ( ! [X19,X18] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X18)) | ~p(s(bool,X19))) ) | ($spl10 | $spl304)),
  inference(resolution,[],[f27568,f21288])).
fof(f21288,plain,(
  ( ! [X136,X137] : (~sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X136,X137) | p(s(bool,X137)) | ~p(s(bool,X136))) ) | $spl10),
  inference(resolution,[],[f1092,f689])).
fof(f38550,plain,(
  ( ! [X17,X16] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X16)) | ~p(s(bool,X17))) ) | ($spl8 | $spl10 | $spl304)),
  inference(forward_demodulation,[],[f38546,f21342])).
fof(f38546,plain,(
  ( ! [X17,X16] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X16)) | ~p(s(bool,X17))) ) | ($spl8 | $spl304)),
  inference(resolution,[],[f27568,f7649])).
fof(f7649,plain,(
  ( ! [X136,X137] : (~sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)),X136,X137) | p(s(bool,X137)) | ~p(s(bool,X136))) ) | $spl8),
  inference(resolution,[],[f1089,f689])).
fof(f38535,plain,(
  ( ! [X14,X15] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X14)) | ~p(s(bool,X15))) ) | ($spl10 | $spl304)),
  inference(resolution,[],[f27567,f21289])).
fof(f21289,plain,(
  ( ! [X138,X139] : (~sP8(X138,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X139) | p(s(bool,X139)) | ~p(s(bool,X138))) ) | $spl10),
  inference(resolution,[],[f1092,f689])).
fof(f38539,plain,(
  ( ! [X12,X13] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X12)) | ~p(s(bool,X13))) ) | ($spl8 | $spl10 | $spl304)),
  inference(forward_demodulation,[],[f38534,f21342])).
fof(f38534,plain,(
  ( ! [X12,X13] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X12)) | ~p(s(bool,X13))) ) | ($spl8 | $spl304)),
  inference(resolution,[],[f27567,f7650])).
fof(f7650,plain,(
  ( ! [X138,X139] : (~sP8(X138,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)),X139) | p(s(bool,X139)) | ~p(s(bool,X138))) ) | $spl8),
  inference(resolution,[],[f1089,f689])).
fof(f38527,plain,(
  ( ! [X21,X22] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X21)) | ~p(s(bool,X22))) ) | ($spl10 | $spl304)),
  inference(resolution,[],[f27555,f21275])).
fof(f21275,plain,(
  ( ! [X111,X110] : (~sP4('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X110,X111) | p(s(bool,X111)) | ~p(s(bool,X110))) ) | $spl10),
  inference(resolution,[],[f1092,f672])).
fof(f38529,plain,(
  ( ! [X19,X20] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X19)) | ~p(s(bool,X20))) ) | ($spl8 | $spl10 | $spl304)),
  inference(forward_demodulation,[],[f38526,f21342])).
fof(f38526,plain,(
  ( ! [X19,X20] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X19)) | ~p(s(bool,X20))) ) | ($spl8 | $spl304)),
  inference(resolution,[],[f27555,f7636])).
fof(f7636,plain,(
  ( ! [X111,X110] : (~sP4('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)),X110,X111) | p(s(bool,X111)) | ~p(s(bool,X110))) ) | $spl8),
  inference(resolution,[],[f1089,f672])).
fof(f38513,plain,(
  ( ! [X17,X18] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X17)) | ~p(s(bool,X18))) ) | ($spl10 | $spl304)),
  inference(resolution,[],[f27554,f21276])).
fof(f21276,plain,(
  ( ! [X112,X113] : (~sP4(X112,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X113) | p(s(bool,X113)) | ~p(s(bool,X112))) ) | $spl10),
  inference(resolution,[],[f1092,f672])).
fof(f38518,plain,(
  ( ! [X15,X16] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X15)) | ~p(s(bool,X16))) ) | ($spl8 | $spl10 | $spl304)),
  inference(forward_demodulation,[],[f38512,f21342])).
fof(f38512,plain,(
  ( ! [X15,X16] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X15)) | ~p(s(bool,X16))) ) | ($spl8 | $spl304)),
  inference(resolution,[],[f27554,f7637])).
fof(f7637,plain,(
  ( ! [X112,X113] : (~sP4(X112,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)),X113) | p(s(bool,X113)) | ~p(s(bool,X112))) ) | $spl8),
  inference(resolution,[],[f1089,f672])).
fof(f38505,plain,(
  ( ! [X17,X18] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X17)) | p(s(bool,X18))) ) | ($spl10 | $spl304)),
  inference(resolution,[],[f27541,f21259])).
fof(f21259,plain,(
  ( ! [X78,X79] : (~sP2('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X78,X79) | p(s(bool,X79)) | p(s(bool,X78))) ) | $spl10),
  inference(resolution,[],[f1092,f655])).
fof(f38506,plain,(
  ( ! [X15,X16] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X15)) | p(s(bool,X16))) ) | ($spl8 | $spl10 | $spl304)),
  inference(forward_demodulation,[],[f38504,f21342])).
fof(f38504,plain,(
  ( ! [X15,X16] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X15)) | p(s(bool,X16))) ) | ($spl8 | $spl304)),
  inference(resolution,[],[f27541,f7620])).
fof(f7620,plain,(
  ( ! [X78,X79] : (~sP2('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)),X78,X79) | p(s(bool,X79)) | p(s(bool,X78))) ) | $spl8),
  inference(resolution,[],[f1089,f655])).
fof(f38497,plain,(
  ( ! [X21,X22] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X21)) | p(s(bool,X22))) ) | ($spl10 | $spl304)),
  inference(resolution,[],[f27525,f21245])).
fof(f21245,plain,(
  ( ! [X50,X51] : (~sP0('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X50,X51) | ~p(s(bool,X51)) | p(s(bool,X50))) ) | $spl10),
  inference(resolution,[],[f1092,f638])).
fof(f38498,plain,(
  ( ! [X19,X20] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X19)) | p(s(bool,X20))) ) | ($spl8 | $spl10 | $spl304)),
  inference(forward_demodulation,[],[f38496,f21342])).
fof(f38496,plain,(
  ( ! [X19,X20] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X19)) | p(s(bool,X20))) ) | ($spl8 | $spl304)),
  inference(resolution,[],[f27525,f7606])).
fof(f7606,plain,(
  ( ! [X50,X51] : (~sP0('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)),X50,X51) | ~p(s(bool,X51)) | p(s(bool,X50))) ) | $spl8),
  inference(resolution,[],[f1089,f638])).
fof(f38488,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f38363,f1098])).
fof(f38363,plain,(
  ( ! [X391,X390,X392] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X390,bool),X391),s(fun(X390,bool),X392))))) ) | $spl20),
  inference(resolution,[],[f1098,f850])).
fof(f38487,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f38346,f1098])).
fof(f38346,plain,(
  ( ! [X339,X340,X338] : (p(s(bool,'const.bool.IN_2'(s(X338,sK39(X339,X340,X338)),s(fun(X338,bool),X340))))) ) | $spl20),
  inference(resolution,[],[f1098,f850])).
fof(f38355,plain,(
  $false | $spl20),
  inference(resolution,[],[f1098,f863])).
fof(f863,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f352])).
fof(f352,plain,(
  ! [X0,X1,X2] : p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X1))))),
  inference(flattening,[],[f351])).
fof(f351,plain,(
  ! [X0] : ! [X1,X2] : p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X1))))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X6] : ! [X7,X8] : p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),'const.pred_set.DIFF_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8))),s(fun(X6,bool),X7))))),
  file('thm.pred_set.SUBSET_DIFF_EMPTY','thm.pred_set.DIFF_SUBSET')).
fof(f38356,plain,(
  $false | $spl20),
  inference(resolution,[],[f1098,f868])).
fof(f868,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f360])).
fof(f360,plain,(
  ! [X0,X1] : p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X0,bool),X1))))),
  inference(flattening,[],[f359])).
fof(f359,plain,(
  ! [X0] : ! [X1] : p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X0,bool),X1))))),
  inference(rectify,[],[f107])).
fof(f107,axiom,(
  ! [X6] : ! [X7] : p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),'const.pred_set.EMPTY_0'),s(fun(X6,bool),X7))))),
  file('thm.pred_set.SUBSET_DIFF_EMPTY','thm.pred_set.EMPTY_SUBSET')).
fof(f38365,plain,(
  $false | ($spl5 | $spl20)),
  inference(resolution,[],[f1098,f11722])).
fof(f11722,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),sK15)))) | $spl5),
  inference(resolution,[],[f1086,f850])).
fof(f38366,plain,(
  $false | ($spl20 | $spl27)),
  inference(resolution,[],[f1098,f12007])).
fof(f12007,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK15,sK13)),s(fun(sK13,bool),sK15)))) | $spl27),
  inference(resolution,[],[f1104,f850])).
fof(f38367,plain,(
  $false | ($spl20 | $spl29)),
  inference(resolution,[],[f1098,f12255])).
fof(f12255,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),sK14)))) | $spl29),
  inference(resolution,[],[f1105,f850])).
fof(f38368,plain,(
  $false | ($spl3 | $spl20)),
  inference(resolution,[],[f1098,f33583])).
fof(f38374,plain,(
  $false | $spl20),
  inference(resolution,[],[f1098,f902])).
fof(f38376,plain,(
  $false | ($spl10 | $spl20)),
  inference(resolution,[],[f1098,f1092])).
fof(f38378,plain,(
  $false | ($spl20 | $spl304)),
  inference(resolution,[],[f1098,f8498])).
fof(f38414,plain,(
  $false | ($spl20 | $spl326)),
  inference(resolution,[],[f1098,f8530])).
fof(f38227,plain,(
  ( ! [X17,X18] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X17)) | ~p(s(bool,X18))) ) | ($spl10 | $spl304)),
  inference(resolution,[],[f27524,f21246])).
fof(f21246,plain,(
  ( ! [X52,X53] : (~sP0(X52,X53,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | p(s(bool,X53)) | ~p(s(bool,X52))) ) | $spl10),
  inference(resolution,[],[f1092,f638])).
fof(f38231,plain,(
  ( ! [X15,X16] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X15)) | ~p(s(bool,X16))) ) | ($spl8 | $spl10 | $spl304)),
  inference(forward_demodulation,[],[f38226,f21342])).
fof(f38226,plain,(
  ( ! [X15,X16] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X15)) | ~p(s(bool,X16))) ) | ($spl8 | $spl304)),
  inference(resolution,[],[f27524,f7607])).
fof(f7607,plain,(
  ( ! [X52,X53] : (~sP0(X52,X53,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) | p(s(bool,X53)) | ~p(s(bool,X52))) ) | $spl8),
  inference(resolution,[],[f1089,f638])).
fof(f35158,plain,(
  ( ! [X6] : (~p(s(bool,X6)) | s(bool,X6) != s(bool,'const.bool.IN_2'(s(sK13,sK39(sK15,sK14,sK13)),s(fun(sK13,bool),sK15)))) ) | $spl648),
  inference(cnf_transformation,[],[f35158_D])).
fof(f35158_D,plain,(
  ( ! [X6] : (~p(s(bool,X6)) | s(bool,X6) != s(bool,'const.bool.IN_2'(s(sK13,sK39(sK15,sK14,sK13)),s(fun(sK13,bool),sK15)))) ) <=> ~$spl648),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl648])])).
fof(f35121,plain,(
  ( ! [X218,X217] : (p(s(bool,X217)) | ~p(s(bool,X218)) | s(bool,X218) != s(bool,'const.bool.IN_2'(s(sK13,sK39(sK15,sK14,sK13)),s(fun(sK13,bool),sK15)))) ) | $spl3),
  inference(resolution,[],[f33584,f905])).
fof(f35023,plain,(
  ( ! [X28,X26,X27] : (~p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.bool.IN_2'(s(sK13,sK39(sK15,sK14,sK13)),s(fun(sK13,bool),sK15))) | s(bool,X26) != s(bool,X28)) ) | $spl3),
  inference(resolution,[],[f33584,f630])).
fof(f35021,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.bool.IN_2'(s(sK13,sK39(sK15,sK14,sK13)),s(fun(sK13,bool),sK15)))) ) | $spl3),
  inference(resolution,[],[f33584,f629])).
fof(f35019,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.bool.IN_2'(s(sK13,sK39(sK15,sK14,sK13)),s(fun(sK13,bool),sK15)))) ) | $spl3),
  inference(resolution,[],[f33584,f628])).
fof(f35018,plain,(
  ( ! [X12,X13,X11] : (~p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X12) != s(bool,'const.bool.IN_2'(s(sK13,sK39(sK15,sK14,sK13)),s(fun(sK13,bool),sK15))) | s(bool,X11) != s(bool,X13)) ) | $spl3),
  inference(resolution,[],[f33584,f626])).
fof(f35016,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,'const.bool.IN_2'(s(sK13,sK39(sK15,sK14,sK13)),s(fun(sK13,bool),sK15)))) ) | $spl3),
  inference(resolution,[],[f33584,f624])).
fof(f35007,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,X6) != s(bool,'const.bool.IN_2'(s(sK13,sK39(sK15,sK14,sK13)),s(fun(sK13,bool),sK14)))) ) | $spl646),
  inference(cnf_transformation,[],[f35007_D])).
fof(f35007_D,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,X6) != s(bool,'const.bool.IN_2'(s(sK13,sK39(sK15,sK14,sK13)),s(fun(sK13,bool),sK14)))) ) <=> ~$spl646),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl646])])).
fof(f34965,plain,(
  ( ! [X216,X215] : (p(s(bool,X215)) | p(s(bool,X216)) | s(bool,X216) != s(bool,'const.bool.IN_2'(s(sK13,sK39(sK15,sK14,sK13)),s(fun(sK13,bool),sK14)))) ) | $spl3),
  inference(resolution,[],[f33583,f905])).
fof(f34867,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,'const.bool.IN_2'(s(sK13,sK39(sK15,sK14,sK13)),s(fun(sK13,bool),sK14))) | s(bool,X27) != s(bool,X28)) ) | $spl3),
  inference(resolution,[],[f33583,f630])).
fof(f34864,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X19) != s(bool,'const.bool.IN_2'(s(sK13,sK39(sK15,sK14,sK13)),s(fun(sK13,bool),sK14)))) ) | $spl3),
  inference(resolution,[],[f33583,f629])).
fof(f34863,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.bool.IN_2'(s(sK13,sK39(sK15,sK14,sK13)),s(fun(sK13,bool),sK14)))) ) | $spl3),
  inference(resolution,[],[f33583,f628])).
fof(f34862,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,'const.bool.IN_2'(s(sK13,sK39(sK15,sK14,sK13)),s(fun(sK13,bool),sK14))) | s(bool,X12) != s(bool,X13)) ) | $spl3),
  inference(resolution,[],[f33583,f626])).
fof(f34860,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,'const.bool.IN_2'(s(sK13,sK39(sK15,sK14,sK13)),s(fun(sK13,bool),sK14)))) ) | $spl3),
  inference(resolution,[],[f33583,f624])).
fof(f4743,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))) | ~p(s(bool,X3))) ) | $spl202),
  inference(cnf_transformation,[],[f4743_D])).
fof(f4743_D,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))) | ~p(s(bool,X3))) ) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f33694,plain,(
  ( ! [X215,X214] : (p(s(bool,X214)) | ~p(s(bool,X215)) | s(bool,X215) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))) ) | $spl3),
  inference(resolution,[],[f916,f905])).
fof(f33596,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))) | s(bool,X23) != s(bool,X25)) ) | $spl3),
  inference(resolution,[],[f916,f630])).
fof(f33594,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))) ) | $spl3),
  inference(resolution,[],[f916,f629])).
fof(f33592,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))) ) | $spl3),
  inference(resolution,[],[f916,f628])).
fof(f33591,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X9) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))) | s(bool,X8) != s(bool,X10)) ) | $spl3),
  inference(resolution,[],[f916,f626])).
fof(f33589,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))) ) | $spl3),
  inference(resolution,[],[f916,f624])).
fof(f1170,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))))) | $spl47),
  inference(cnf_transformation,[],[f1170_D])).
fof(f1170_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))))) <=> ~$spl47),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl47])])).
fof(f913,plain,(
  s(fun(sK13,bool),'const.pred_set.EMPTY_0') = s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))) | $spl0),
  inference(cnf_transformation,[],[f913_D])).
fof(f913_D,plain,(
  s(fun(sK13,bool),'const.pred_set.EMPTY_0') = s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f29899,plain,(
  s(fun(sK13,bool),'const.pred_set.EMPTY_0') = s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))))) | $spl42),
  inference(resolution,[],[f1166,f897])).
fof(f897,plain,(
  ( ! [X0,X1] : (s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f617])).
fof(f617,plain,(
  ! [X0,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) | ((s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1)) & (s(fun(X0,bool),'const.pred_set.EMPTY_0') != s(fun(X0,bool),X1) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))))))),
  inference(nnf_transformation,[],[f457])).
fof(f457,plain,(
  ! [X0,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) | (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))) <=> s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1)))),
  inference(ennf_transformation,[],[f406])).
fof(f406,plain,(
  ! [X0,X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) => (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))) <=> s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1)))),
  inference(flattening,[],[f405])).
fof(f405,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) => (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))) <=> s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1)))),
  inference(rectify,[],[f250])).
fof(f250,axiom,(
  ! [X6] : ! [X7] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X7)))) => (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(X6,bool),X7))) <=> s(fun(X6,bool),'const.pred_set.EMPTY_0') = s(fun(X6,bool),X7)))),
  file('thm.pred_set.SUBSET_DIFF_EMPTY','thm.pred_set.CARD_EQ_|48|')).
fof(f1166,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))))) | $spl42),
  inference(cnf_transformation,[],[f1166_D])).
fof(f1166_D,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))))) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f914,plain,(
  s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))) | $spl1),
  inference(cnf_transformation,[],[f914_D])).
fof(f914_D,plain,(
  s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f1169,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))))) | $spl46),
  inference(cnf_transformation,[],[f1169_D])).
fof(f1169_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))))) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f29900,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))))) | s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))) | $spl42),
  inference(resolution,[],[f1166,f898])).
fof(f898,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') != s(fun(X0,bool),X1) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f617])).
fof(f33243,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f33242,f914])).
fof(f33242,plain,(
  s(fun(sK13,bool),'const.pred_set.EMPTY_0') = s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))) | $spl2),
  inference(forward_demodulation,[],[f32880,f889])).
fof(f889,plain,(
  ( ! [X0,X1] : (s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)))) )),
  inference(cnf_transformation,[],[f396])).
fof(f396,plain,(
  ! [X0,X1] : s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)))),
  inference(flattening,[],[f395])).
fof(f395,plain,(
  ! [X0] : ! [X1] : s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X6] : ! [X7] : s(fun(X6,bool),'const.pred_set.EMPTY_0') = s(fun(X6,bool),'const.pred_set.DIFF_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X7)))),
  file('thm.pred_set.SUBSET_DIFF_EMPTY','thm.pred_set.DIFF_EQ_EMPTY')).
fof(f32880,plain,(
  s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))) = s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK15))) | $spl2),
  inference(superposition,[],[f882,f21378])).
fof(f21378,plain,(
  s(fun(sK13,bool),sK15) = s(fun(sK13,bool),'const.pred_set.UNION_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))) | $spl2),
  inference(resolution,[],[f915,f852])).
fof(f852,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X2) = s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f607])).
fof(f882,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) = s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X2)))) )),
  inference(cnf_transformation,[],[f382])).
fof(f382,plain,(
  ! [X0,X1,X2] : s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) = s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X2)))),
  inference(flattening,[],[f381])).
fof(f381,plain,(
  ! [X0] : ! [X1,X2] : s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) = s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X2)))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X6] : ! [X9,X16] : s(fun(X6,bool),'const.pred_set.DIFF_2'(s(fun(X6,bool),X9),s(fun(X6,bool),X16))) = s(fun(X6,bool),'const.pred_set.DIFF_2'(s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),X9),s(fun(X6,bool),X16))),s(fun(X6,bool),X16)))),
  file('thm.pred_set.SUBSET_DIFF_EMPTY','thm.pred_set.DIFF_SAME_UNION|split|1')).
fof(f31873,plain,(
  ( ! [X48,X46] : (~p(s(bool,X46)) | s(bool,X46) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X48)))) ) | $spl644),
  inference(cnf_transformation,[],[f31873_D])).
fof(f31873_D,plain,(
  ( ! [X48,X46] : (~p(s(bool,X46)) | s(bool,X46) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X48)))) ) <=> ~$spl644),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl644])])).
fof(f31836,plain,(
  ( ! [X362,X361,X363] : (p(s(bool,X361)) | ~p(s(bool,X362)) | s(bool,X362) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X363)))) ) | $spl214),
  inference(resolution,[],[f5502,f905])).
fof(f31738,plain,(
  ( ! [X74,X72,X75,X73] : (~p(s(bool,X72)) | ~p(s(bool,X73)) | s(bool,X73) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X74))) | s(bool,X72) != s(bool,X75)) ) | $spl214),
  inference(resolution,[],[f5502,f630])).
fof(f31736,plain,(
  ( ! [X68,X66,X67,X65] : (p(s(bool,X65)) | ~p(s(bool,X66)) | ~p(s(bool,X67)) | s(bool,X66) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X68)))) ) | $spl214),
  inference(resolution,[],[f5502,f629])).
fof(f31734,plain,(
  ( ! [X59,X57,X60,X58] : (p(s(bool,X57)) | ~p(s(bool,X58)) | s(bool,X57) != s(bool,X59) | s(bool,X58) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X60)))) ) | $spl214),
  inference(resolution,[],[f5502,f628])).
fof(f31733,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X54) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X55))) | s(bool,X53) != s(bool,X56)) ) | $spl214),
  inference(resolution,[],[f5502,f626])).
fof(f31731,plain,(
  ( ! [X47,X45,X48,X46] : (p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X45) != s(bool,X47) | s(bool,X46) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X48)))) ) | $spl214),
  inference(resolution,[],[f5502,f624])).
fof(f5495,plain,(
  ( ! [X114] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X114))))) ) | $spl204),
  inference(cnf_transformation,[],[f5495_D])).
fof(f5495_D,plain,(
  ( ! [X114] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X114))))) ) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f31439,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK38('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13))))) | $spl643),
  inference(cnf_transformation,[],[f31439_D])).
fof(f31439_D,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK38('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13))))) <=> ~$spl643),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl643])])).
fof(f31440,plain,(
  ( ! [X114] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK38('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X114))))) ) | $spl2),
  inference(forward_demodulation,[],[f31059,f21375])).
fof(f21375,plain,(
  s(fun(sK13,bool),sK14) = s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))) | $spl2),
  inference(resolution,[],[f915,f846])).
fof(f846,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f602])).
fof(f31059,plain,(
  ( ! [X114] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X114)))) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))),s(sK13,sK38('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13)))))) ) | $spl2),
  inference(superposition,[],[f825,f21375])).
fof(f825,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK38(X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f590])).
fof(f590,plain,(
  ! [X0,X1] : ((! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK37(X1,X0)))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK38(X1,X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK37,sK38])],[f589])).
fof(f589,plain,(
  ! [X0,X1] : ((! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(rectify,[],[f588])).
fof(f588,plain,(
  ! [X0,X1] : ((! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(nnf_transformation,[],[f435])).
fof(f435,plain,(
  ! [X0,X1] : (? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(ennf_transformation,[],[f318])).
fof(f318,plain,(
  ! [X0,X1] : (~! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(flattening,[],[f317])).
fof(f317,plain,(
  ! [X0] : ! [X1] : (~! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(rectify,[],[f116])).
fof(f116,axiom,(
  ! [X6] : ! [X18] : (~! [X9] : p(s(bool,happ(s(fun(X6,bool),X18),s(X6,X9)))) <=> ? [X9] : ~p(s(bool,happ(s(fun(X6,bool),X18),s(X6,X9)))))),
  file('thm.pred_set.SUBSET_DIFF_EMPTY','thm.bool.NOT_FORALL_THM')).
fof(f31437,plain,(
  ( ! [X113] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X113)))) | ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK38('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13)))))) ) | $spl2),
  inference(forward_demodulation,[],[f31058,f21375])).
fof(f31058,plain,(
  ( ! [X113] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK38('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13))))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))),s(sK13,X113))))) ) | $spl2),
  inference(superposition,[],[f825,f21375])).
fof(f31435,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK37('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13))))) | $spl641),
  inference(cnf_transformation,[],[f31435_D])).
fof(f31435_D,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK37('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13))))) <=> ~$spl641),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl641])])).
fof(f31436,plain,(
  ( ! [X112] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK37('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X112))))) ) | $spl2),
  inference(forward_demodulation,[],[f31057,f21375])).
fof(f31057,plain,(
  ( ! [X112] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X112)))) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))),s(sK13,sK37('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13)))))) ) | $spl2),
  inference(superposition,[],[f824,f21375])).
fof(f824,plain,(
  ( ! [X0,X5,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK37(X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))) )),
  inference(cnf_transformation,[],[f590])).
fof(f31433,plain,(
  ( ! [X111] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X111)))) | ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK37('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13)))))) ) | $spl2),
  inference(forward_demodulation,[],[f31056,f21375])).
fof(f31056,plain,(
  ( ! [X111] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK37('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13))))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))),s(sK13,X111))))) ) | $spl2),
  inference(superposition,[],[f824,f21375])).
fof(f31431,plain,(
  ( ! [X107] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK36('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X107,sK13))))) | ~p(s(bool,X107))) ) | $spl638),
  inference(cnf_transformation,[],[f31431_D])).
fof(f31431_D,plain,(
  ( ! [X107] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK36('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X107,sK13))))) | ~p(s(bool,X107))) ) <=> ~$spl638),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl638])])).
fof(f31432,plain,(
  ( ! [X109,X110] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK36('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X110,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X109)))) | ~p(s(bool,X110))) ) | $spl2),
  inference(forward_demodulation,[],[f31055,f21375])).
fof(f31055,plain,(
  ( ! [X109,X110] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X109)))) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))),s(sK13,sK36('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X110,sK13))))) | ~p(s(bool,X110))) ) | $spl2),
  inference(superposition,[],[f821,f21375])).
fof(f821,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK36(X2,X1,X0))))) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f585])).
fof(f585,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK35(X2,X1,X0))))) | ! [X5] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))))) & ((~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK36(X2,X1,X0)))))) | (p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK35,sK36])],[f584])).
fof(f584,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,X1)) | ? [X6] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) | ! [X5] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))))) & (? [X4] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(rectify,[],[f583])).
fof(f583,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,X1)) | ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | ! [X4] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (? [X4] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(flattening,[],[f582])).
fof(f582,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) | ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ! [X4] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (? [X4] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f315])).
fof(f315,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> ! [X4] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f314])).
fof(f314,plain,(
  ! [X0] : ! [X1,X2] : ((p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> ! [X4] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f240])).
fof(f240,axiom,(
  ! [X6] : ! [X18,X19] : ((p(s(bool,X18)) & ! [X9] : p(s(bool,happ(s(fun(X6,bool),X19),s(X6,X9))))) <=> ! [X9] : (p(s(bool,X18)) & p(s(bool,happ(s(fun(X6,bool),X19),s(X6,X9))))))),
  file('thm.pred_set.SUBSET_DIFF_EMPTY','thm.bool.RIGHT_AND_FORALL_THM')).
fof(f31430,plain,(
  ( ! [X107,X108] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X108)))) | ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK36('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X107,sK13))))) | ~p(s(bool,X107))) ) | $spl2),
  inference(forward_demodulation,[],[f31054,f21375])).
fof(f31054,plain,(
  ( ! [X107,X108] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK36('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X107,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))),s(sK13,X108)))) | ~p(s(bool,X107))) ) | $spl2),
  inference(superposition,[],[f821,f21375])).
fof(f31428,plain,(
  ( ! [X103] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK35('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X103,sK13))))) | ~p(s(bool,X103))) ) | $spl636),
  inference(cnf_transformation,[],[f31428_D])).
fof(f31428_D,plain,(
  ( ! [X103] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK35('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X103,sK13))))) | ~p(s(bool,X103))) ) <=> ~$spl636),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl636])])).
fof(f31429,plain,(
  ( ! [X105,X106] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK35('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X106,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X105)))) | ~p(s(bool,X106))) ) | $spl2),
  inference(forward_demodulation,[],[f31053,f21375])).
fof(f31053,plain,(
  ( ! [X105,X106] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X105)))) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))),s(sK13,sK35('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X106,sK13))))) | ~p(s(bool,X106))) ) | $spl2),
  inference(superposition,[],[f819,f21375])).
fof(f819,plain,(
  ( ! [X2,X0,X5,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK35(X2,X1,X0))))) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f585])).
fof(f31427,plain,(
  ( ! [X103,X104] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X104)))) | ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK35('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X103,sK13))))) | ~p(s(bool,X103))) ) | $spl2),
  inference(forward_demodulation,[],[f31052,f21375])).
fof(f31052,plain,(
  ( ! [X103,X104] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK35('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X103,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))),s(sK13,X104)))) | ~p(s(bool,X103))) ) | $spl2),
  inference(superposition,[],[f819,f21375])).
fof(f5502,plain,(
  ( ! [X101] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X101))))) ) | $spl214),
  inference(cnf_transformation,[],[f5502_D])).
fof(f5502_D,plain,(
  ( ! [X101] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X101))))) ) <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f31425,plain,(
  ( ! [X99] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK34('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X99,sK13))))) | p(s(bool,X99))) ) | $spl634),
  inference(cnf_transformation,[],[f31425_D])).
fof(f31425_D,plain,(
  ( ! [X99] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK34('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X99,sK13))))) | p(s(bool,X99))) ) <=> ~$spl634),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl634])])).
fof(f31426,plain,(
  ( ! [X101,X102] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK34('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X102,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X101)))) | p(s(bool,X102))) ) | $spl2),
  inference(forward_demodulation,[],[f31051,f21375])).
fof(f31051,plain,(
  ( ! [X101,X102] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X101)))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))),s(sK13,sK34('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X102,sK13))))) | p(s(bool,X102))) ) | $spl2),
  inference(superposition,[],[f817,f21375])).
fof(f817,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK34(X2,X1,X0))))) | p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) )),
  inference(cnf_transformation,[],[f581])).
fof(f581,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK33(X2,X1,X0))))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK34(X2,X1,X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK33,sK34])],[f580])).
fof(f580,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | ? [X5] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(rectify,[],[f579])).
fof(f579,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(flattening,[],[f578])).
fof(f578,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f313])).
fof(f313,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f312])).
fof(f312,plain,(
  ! [X0] : ! [X1,X2] : ((p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f156])).
fof(f156,axiom,(
  ! [X6] : ! [X18,X19] : ((p(s(bool,X18)) | ? [X9] : p(s(bool,happ(s(fun(X6,bool),X19),s(X6,X9))))) <=> ? [X9] : (p(s(bool,X18)) | p(s(bool,happ(s(fun(X6,bool),X19),s(X6,X9))))))),
  file('thm.pred_set.SUBSET_DIFF_EMPTY','thm.bool.RIGHT_OR_EXISTS_THM')).
fof(f31424,plain,(
  ( ! [X99,X100] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X100)))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK34('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X99,sK13))))) | p(s(bool,X99))) ) | $spl2),
  inference(forward_demodulation,[],[f31050,f21375])).
fof(f31050,plain,(
  ( ! [X99,X100] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK34('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X99,sK13))))) | p(s(bool,X99)) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))),s(sK13,X100))))) ) | $spl2),
  inference(superposition,[],[f817,f21375])).
fof(f31422,plain,(
  ( ! [X95] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK33('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X95,sK13))))) | p(s(bool,X95))) ) | $spl632),
  inference(cnf_transformation,[],[f31422_D])).
fof(f31422_D,plain,(
  ( ! [X95] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK33('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X95,sK13))))) | p(s(bool,X95))) ) <=> ~$spl632),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl632])])).
fof(f31423,plain,(
  ( ! [X97,X98] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK33('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X98,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X97)))) | p(s(bool,X98))) ) | $spl2),
  inference(forward_demodulation,[],[f31049,f21375])).
fof(f31049,plain,(
  ( ! [X97,X98] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X97)))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))),s(sK13,sK33('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X98,sK13))))) | p(s(bool,X98))) ) | $spl2),
  inference(superposition,[],[f815,f21375])).
fof(f815,plain,(
  ( ! [X6,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK33(X2,X1,X0))))) | p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) )),
  inference(cnf_transformation,[],[f581])).
fof(f31421,plain,(
  ( ! [X95,X96] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X96)))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK33('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X95,sK13))))) | p(s(bool,X95))) ) | $spl2),
  inference(forward_demodulation,[],[f31048,f21375])).
fof(f31048,plain,(
  ( ! [X95,X96] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK33('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X95,sK13))))) | p(s(bool,X95)) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))),s(sK13,X96))))) ) | $spl2),
  inference(superposition,[],[f815,f21375])).
fof(f31419,plain,(
  ( ! [X91] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK32(X91,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13))))) | p(s(bool,X91))) ) | $spl630),
  inference(cnf_transformation,[],[f31419_D])).
fof(f31419_D,plain,(
  ( ! [X91] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK32(X91,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13))))) | p(s(bool,X91))) ) <=> ~$spl630),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl630])])).
fof(f31420,plain,(
  ( ! [X94,X93] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK32(X94,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X93)))) | p(s(bool,X94))) ) | $spl2),
  inference(forward_demodulation,[],[f31047,f21375])).
fof(f31047,plain,(
  ( ! [X94,X93] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X93)))) | p(s(bool,X94)) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))),s(sK13,sK32(X94,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13)))))) ) | $spl2),
  inference(superposition,[],[f812,f21375])).
fof(f812,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,X2)) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK32(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) )),
  inference(cnf_transformation,[],[f577])).
fof(f577,plain,(
  ! [X0,X1,X2] : (((! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) & ~p(s(bool,X2))) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK31(X2,X1,X0))))) | p(s(bool,X2)))) & (! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ~p(s(bool,X2))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK32(X2,X1,X0))))) | p(s(bool,X2))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK31,sK32])],[f576])).
fof(f576,plain,(
  ! [X0,X1,X2] : (((! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) & ~p(s(bool,X2))) | ? [X5] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | p(s(bool,X2)))) & (! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ~p(s(bool,X2))) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2))))),
  inference(rectify,[],[f575])).
fof(f575,plain,(
  ! [X0,X1,X2] : (((! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,X2))) | ? [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,X2)))) & (! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ~p(s(bool,X2))) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2))))),
  inference(flattening,[],[f574])).
fof(f574,plain,(
  ! [X0,X1,X2] : (((! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,X2))) | ? [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,X2)))) & (! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ~p(s(bool,X2))) | (? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2)))))),
  inference(nnf_transformation,[],[f311])).
fof(f311,plain,(
  ! [X0,X1,X2] : ((? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2))) <=> ? [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,X2))))),
  inference(flattening,[],[f310])).
fof(f310,plain,(
  ! [X0] : ! [X1,X2] : ((? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2))) <=> ? [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,X2))))),
  inference(rectify,[],[f155])).
fof(f155,axiom,(
  ! [X6] : ! [X18,X19] : ((? [X9] : p(s(bool,happ(s(fun(X6,bool),X18),s(X6,X9)))) | p(s(bool,X19))) <=> ? [X9] : (p(s(bool,happ(s(fun(X6,bool),X18),s(X6,X9)))) | p(s(bool,X19))))),
  file('thm.pred_set.SUBSET_DIFF_EMPTY','thm.bool.LEFT_OR_EXISTS_THM')).
fof(f31418,plain,(
  ( ! [X92,X91] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X92)))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK32(X91,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13))))) | p(s(bool,X91))) ) | $spl2),
  inference(forward_demodulation,[],[f31046,f21375])).
fof(f31046,plain,(
  ( ! [X92,X91] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK32(X91,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13))))) | p(s(bool,X91)) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))),s(sK13,X92))))) ) | $spl2),
  inference(superposition,[],[f812,f21375])).
fof(f31416,plain,(
  ( ! [X87] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK31(X87,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13))))) | p(s(bool,X87))) ) | $spl628),
  inference(cnf_transformation,[],[f31416_D])).
fof(f31416_D,plain,(
  ( ! [X87] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK31(X87,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13))))) | p(s(bool,X87))) ) <=> ~$spl628),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl628])])).
fof(f31417,plain,(
  ( ! [X90,X89] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK31(X90,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X89)))) | p(s(bool,X90))) ) | $spl2),
  inference(forward_demodulation,[],[f31045,f21375])).
fof(f31045,plain,(
  ( ! [X90,X89] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X89)))) | p(s(bool,X90)) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))),s(sK13,sK31(X90,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13)))))) ) | $spl2),
  inference(superposition,[],[f810,f21375])).
fof(f810,plain,(
  ( ! [X6,X2,X0,X1] : (p(s(bool,X2)) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK31(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6))))) )),
  inference(cnf_transformation,[],[f577])).
fof(f31415,plain,(
  ( ! [X88,X87] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X88)))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK31(X87,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13))))) | p(s(bool,X87))) ) | $spl2),
  inference(forward_demodulation,[],[f31044,f21375])).
fof(f31044,plain,(
  ( ! [X88,X87] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK31(X87,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13))))) | p(s(bool,X87)) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))),s(sK13,X88))))) ) | $spl2),
  inference(superposition,[],[f810,f21375])).
fof(f31413,plain,(
  ( ! [X83] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK30('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X83,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X83),s(sK13,sK29('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X83,sK13)))))) ) | $spl626),
  inference(cnf_transformation,[],[f31413_D])).
fof(f31413_D,plain,(
  ( ! [X83] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK30('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X83,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X83),s(sK13,sK29('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X83,sK13)))))) ) <=> ~$spl626),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl626])])).
fof(f31414,plain,(
  ( ! [X85,X86] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK30('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X86,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X85)))) | ~p(s(bool,happ(s(fun(sK13,bool),X86),s(sK13,sK29('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X86,sK13)))))) ) | $spl2),
  inference(forward_demodulation,[],[f31043,f21375])).
fof(f31043,plain,(
  ( ! [X85,X86] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X85)))) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))),s(sK13,sK30('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X86,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X86),s(sK13,sK29('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X86,sK13)))))) ) | $spl2),
  inference(superposition,[],[f809,f21375])).
fof(f809,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK30(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK29(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f573])).
fof(f573,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK28(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK28(X2,X1,X0)))))) | (! [X7] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X7)))) & ! [X6] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK29(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK30(X2,X1,X0))))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK28,sK29,sK30])],[f572])).
fof(f572,plain,(
  ! [X0,X1,X2] : ((? [X8] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X8)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X8))))) | (! [X7] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X7)))) & ! [X6] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))))) & (? [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(rectify,[],[f571])).
fof(f571,plain,(
  ! [X0,X1,X2] : ((? [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (? [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(flattening,[],[f570])).
fof(f570,plain,(
  ! [X0,X1,X2] : ((? [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & ((? [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f309])).
fof(f309,plain,(
  ! [X0,X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f308])).
fof(f308,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f258])).
fof(f258,axiom,(
  ! [X6] : ! [X18,X19] : (! [X9] : (p(s(bool,happ(s(fun(X6,bool),X18),s(X6,X9)))) & p(s(bool,happ(s(fun(X6,bool),X19),s(X6,X9))))) <=> (! [X9] : p(s(bool,happ(s(fun(X6,bool),X18),s(X6,X9)))) & ! [X9] : p(s(bool,happ(s(fun(X6,bool),X19),s(X6,X9))))))),
  file('thm.pred_set.SUBSET_DIFF_EMPTY','thm.bool.FORALL_AND_THM')).
fof(f31412,plain,(
  ( ! [X83,X84] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X84)))) | ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK30('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X83,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X83),s(sK13,sK29('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X83,sK13)))))) ) | $spl2),
  inference(forward_demodulation,[],[f31042,f21375])).
fof(f31042,plain,(
  ( ! [X83,X84] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK30('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X83,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))),s(sK13,X84)))) | ~p(s(bool,happ(s(fun(sK13,bool),X83),s(sK13,sK29('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X83,sK13)))))) ) | $spl2),
  inference(superposition,[],[f809,f21375])).
fof(f31410,plain,(
  ( ! [X75] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK29(X75,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X75),s(sK13,sK30(X75,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13)))))) ) | $spl624),
  inference(cnf_transformation,[],[f31410_D])).
fof(f31410_D,plain,(
  ( ! [X75] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK29(X75,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X75),s(sK13,sK30(X75,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13)))))) ) <=> ~$spl624),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl624])])).
fof(f31411,plain,(
  ( ! [X80,X79] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK29(X80,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X79)))) | ~p(s(bool,happ(s(fun(sK13,bool),X80),s(sK13,sK30(X80,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13)))))) ) | $spl2),
  inference(forward_demodulation,[],[f31040,f21375])).
fof(f31040,plain,(
  ( ! [X80,X79] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X79)))) | ~p(s(bool,happ(s(fun(sK13,bool),X80),s(sK13,sK30(X80,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))),s(sK13,sK29(X80,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13)))))) ) | $spl2),
  inference(superposition,[],[f808,f21375])).
fof(f808,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK30(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK29(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f573])).
fof(f31409,plain,(
  ( ! [X76,X75] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X76)))) | ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK29(X75,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X75),s(sK13,sK30(X75,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13)))))) ) | $spl2),
  inference(forward_demodulation,[],[f31038,f21375])).
fof(f31038,plain,(
  ( ! [X76,X75] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK29(X75,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13))))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))),s(sK13,X76)))) | ~p(s(bool,happ(s(fun(sK13,bool),X75),s(sK13,sK30(X75,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13)))))) ) | $spl2),
  inference(superposition,[],[f808,f21375])).
fof(f31407,plain,(
  ( ! [X71] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK28('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X71,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X71),s(sK13,sK28('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X71,sK13)))))) ) | $spl622),
  inference(cnf_transformation,[],[f31407_D])).
fof(f31407_D,plain,(
  ( ! [X71] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK28('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X71,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X71),s(sK13,sK28('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X71,sK13)))))) ) <=> ~$spl622),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl622])])).
fof(f31408,plain,(
  ( ! [X74,X73] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK28('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X74,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X73)))) | ~p(s(bool,happ(s(fun(sK13,bool),X74),s(sK13,sK28('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X74,sK13)))))) ) | $spl2),
  inference(forward_demodulation,[],[f31037,f21375])).
fof(f31037,plain,(
  ( ! [X74,X73] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X73)))) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))),s(sK13,sK28('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X74,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X74),s(sK13,sK28('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X74,sK13)))))) ) | $spl2),
  inference(superposition,[],[f807,f21375])).
fof(f807,plain,(
  ( ! [X6,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK28(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK28(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f573])).
fof(f31406,plain,(
  ( ! [X72,X71] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X72)))) | ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK28('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X71,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X71),s(sK13,sK28('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X71,sK13)))))) ) | $spl2),
  inference(forward_demodulation,[],[f31036,f21375])).
fof(f31036,plain,(
  ( ! [X72,X71] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK28('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X71,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))),s(sK13,X72)))) | ~p(s(bool,happ(s(fun(sK13,bool),X71),s(sK13,sK28('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X71,sK13)))))) ) | $spl2),
  inference(superposition,[],[f807,f21375])).
fof(f31404,plain,(
  ( ! [X63] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK28(X63,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X63),s(sK13,sK28(X63,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13)))))) ) | $spl620),
  inference(cnf_transformation,[],[f31404_D])).
fof(f31404_D,plain,(
  ( ! [X63] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK28(X63,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X63),s(sK13,sK28(X63,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13)))))) ) <=> ~$spl620),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl620])])).
fof(f31405,plain,(
  ( ! [X68,X67] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK28(X68,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X67)))) | ~p(s(bool,happ(s(fun(sK13,bool),X68),s(sK13,sK28(X68,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13)))))) ) | $spl2),
  inference(forward_demodulation,[],[f31034,f21375])).
fof(f31034,plain,(
  ( ! [X68,X67] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X67)))) | ~p(s(bool,happ(s(fun(sK13,bool),X68),s(sK13,sK28(X68,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))),s(sK13,sK28(X68,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13)))))) ) | $spl2),
  inference(superposition,[],[f806,f21375])).
fof(f806,plain,(
  ( ! [X2,X0,X7,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X7)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK28(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK28(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f573])).
fof(f31403,plain,(
  ( ! [X64,X63] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X64)))) | ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK28(X63,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X63),s(sK13,sK28(X63,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13)))))) ) | $spl2),
  inference(forward_demodulation,[],[f31032,f21375])).
fof(f31032,plain,(
  ( ! [X64,X63] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK28(X63,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13))))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))),s(sK13,X64)))) | ~p(s(bool,happ(s(fun(sK13,bool),X63),s(sK13,sK28(X63,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13)))))) ) | $spl2),
  inference(superposition,[],[f806,f21375])).
fof(f31401,plain,(
  ( ! [X59] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK27('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X59,sK13))))) | p(s(bool,X59))) ) | $spl618),
  inference(cnf_transformation,[],[f31401_D])).
fof(f31401_D,plain,(
  ( ! [X59] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK27('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X59,sK13))))) | p(s(bool,X59))) ) <=> ~$spl618),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl618])])).
fof(f31402,plain,(
  ( ! [X61,X62] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK27('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X62,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X61)))) | p(s(bool,X62))) ) | $spl2),
  inference(forward_demodulation,[],[f31031,f21375])).
fof(f31031,plain,(
  ( ! [X61,X62] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X61)))) | p(s(bool,X62)) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))),s(sK13,sK27('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X62,sK13)))))) ) | $spl2),
  inference(superposition,[],[f805,f21375])).
fof(f805,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK27(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f569])).
fof(f569,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK26(X2,X1,X0)))))) | p(s(bool,X1)) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5))))) & ((~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK27(X2,X1,X0)))))) | ! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK26,sK27])],[f568])).
fof(f568,plain,(
  ! [X0,X1,X2] : ((? [X6] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | p(s(bool,X1)) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5))))) & ((~p(s(bool,X1)) & ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(rectify,[],[f567])).
fof(f567,plain,(
  ! [X0,X1,X2] : ((? [X3] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | p(s(bool,X1)) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) & ((~p(s(bool,X1)) & ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(flattening,[],[f566])).
fof(f566,plain,(
  ! [X0,X1,X2] : ((? [X3] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (p(s(bool,X1)) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & ((~p(s(bool,X1)) & ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f307])).
fof(f307,plain,(
  ! [X0,X1,X2] : (! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (p(s(bool,X1)) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f306])).
fof(f306,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (p(s(bool,X1)) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f131])).
fof(f131,axiom,(
  ! [X6] : ! [X18,X19] : (! [X9] : (p(s(bool,X18)) | p(s(bool,happ(s(fun(X6,bool),X19),s(X6,X9))))) <=> (p(s(bool,X18)) | ! [X9] : p(s(bool,happ(s(fun(X6,bool),X19),s(X6,X9))))))),
  file('thm.pred_set.SUBSET_DIFF_EMPTY','thm.bool.RIGHT_FORALL_OR_THM')).
fof(f31400,plain,(
  ( ! [X59,X60] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X60)))) | ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK27('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X59,sK13))))) | p(s(bool,X59))) ) | $spl2),
  inference(forward_demodulation,[],[f31030,f21375])).
fof(f31030,plain,(
  ( ! [X59,X60] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK27('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X59,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))),s(sK13,X60)))) | p(s(bool,X59))) ) | $spl2),
  inference(superposition,[],[f805,f21375])).
fof(f31398,plain,(
  ( ! [X55] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK26('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X55,sK13))))) | p(s(bool,X55))) ) | $spl616),
  inference(cnf_transformation,[],[f31398_D])).
fof(f31398_D,plain,(
  ( ! [X55] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK26('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X55,sK13))))) | p(s(bool,X55))) ) <=> ~$spl616),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl616])])).
fof(f31399,plain,(
  ( ! [X57,X58] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK26('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X58,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X57)))) | p(s(bool,X58))) ) | $spl2),
  inference(forward_demodulation,[],[f31029,f21375])).
fof(f31029,plain,(
  ( ! [X57,X58] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X57)))) | p(s(bool,X58)) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))),s(sK13,sK26('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X58,sK13)))))) ) | $spl2),
  inference(superposition,[],[f803,f21375])).
fof(f803,plain,(
  ( ! [X2,X0,X5,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) | p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK26(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f569])).
fof(f31397,plain,(
  ( ! [X56,X55] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X56)))) | ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK26('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X55,sK13))))) | p(s(bool,X55))) ) | $spl2),
  inference(forward_demodulation,[],[f31028,f21375])).
fof(f31028,plain,(
  ( ! [X56,X55] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK26('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X55,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))),s(sK13,X56)))) | p(s(bool,X55))) ) | $spl2),
  inference(superposition,[],[f803,f21375])).
fof(f31395,plain,(
  ( ! [X51] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK25('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X51,sK13))))) | p(s(bool,X51))) ) | $spl614),
  inference(cnf_transformation,[],[f31395_D])).
fof(f31395_D,plain,(
  ( ! [X51] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK25('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X51,sK13))))) | p(s(bool,X51))) ) <=> ~$spl614),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl614])])).
fof(f31396,plain,(
  ( ! [X54,X53] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK25('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X54,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X53)))) | p(s(bool,X54))) ) | $spl2),
  inference(forward_demodulation,[],[f31027,f21375])).
fof(f31027,plain,(
  ( ! [X54,X53] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X53)))) | p(s(bool,X54)) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))),s(sK13,sK25('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X54,sK13)))))) ) | $spl2),
  inference(superposition,[],[f800,f21375])).
fof(f800,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK25(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f565])).
fof(f565,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK24(X2,X1,X0))))) & ~p(s(bool,X1))) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) | p(s(bool,X1))) & ((~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK25(X2,X1,X0))))) & ~p(s(bool,X1))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK24,sK25])],[f564])).
fof(f564,plain,(
  ! [X0,X1,X2] : ((? [X6] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) & ~p(s(bool,X1))) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) | p(s(bool,X1))) & ((? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) & ~p(s(bool,X1))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1)))))),
  inference(rectify,[],[f563])).
fof(f563,plain,(
  ! [X0,X1,X2] : ((? [X3] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) & ~p(s(bool,X1))) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,X1))) & ((? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) & ~p(s(bool,X1))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1)))))),
  inference(flattening,[],[f562])).
fof(f562,plain,(
  ! [X0,X1,X2] : ((? [X3] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) & ~p(s(bool,X1))) | (! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,X1)))) & ((? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) & ~p(s(bool,X1))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1)))))),
  inference(nnf_transformation,[],[f305])).
fof(f305,plain,(
  ! [X0,X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1))) <=> (! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,X1))))),
  inference(flattening,[],[f304])).
fof(f304,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1))) <=> (! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,X1))))),
  inference(rectify,[],[f149])).
fof(f149,axiom,(
  ! [X6] : ! [X19,X18] : (! [X9] : (p(s(bool,happ(s(fun(X6,bool),X18),s(X6,X9)))) | p(s(bool,X19))) <=> (! [X9] : p(s(bool,happ(s(fun(X6,bool),X18),s(X6,X9)))) | p(s(bool,X19))))),
  file('thm.pred_set.SUBSET_DIFF_EMPTY','thm.bool.LEFT_FORALL_OR_THM')).
fof(f31394,plain,(
  ( ! [X52,X51] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X52)))) | ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK25('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X51,sK13))))) | p(s(bool,X51))) ) | $spl2),
  inference(forward_demodulation,[],[f31026,f21375])).
fof(f31026,plain,(
  ( ! [X52,X51] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK25('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X51,sK13))))) | p(s(bool,X51)) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))),s(sK13,X52))))) ) | $spl2),
  inference(superposition,[],[f800,f21375])).
fof(f31392,plain,(
  ( ! [X47] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK24('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X47,sK13))))) | p(s(bool,X47))) ) | $spl612),
  inference(cnf_transformation,[],[f31392_D])).
fof(f31392_D,plain,(
  ( ! [X47] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK24('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X47,sK13))))) | p(s(bool,X47))) ) <=> ~$spl612),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl612])])).
fof(f31393,plain,(
  ( ! [X50,X49] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK24('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X50,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X49)))) | p(s(bool,X50))) ) | $spl2),
  inference(forward_demodulation,[],[f31025,f21375])).
fof(f31025,plain,(
  ( ! [X50,X49] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X49)))) | p(s(bool,X50)) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))),s(sK13,sK24('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X50,sK13)))))) ) | $spl2),
  inference(superposition,[],[f798,f21375])).
fof(f798,plain,(
  ( ! [X2,X0,X5,X1] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK24(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f565])).
fof(f31391,plain,(
  ( ! [X47,X48] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X48)))) | ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK24('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X47,sK13))))) | p(s(bool,X47))) ) | $spl2),
  inference(forward_demodulation,[],[f31024,f21375])).
fof(f31024,plain,(
  ( ! [X47,X48] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK24('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X47,sK13))))) | p(s(bool,X47)) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))),s(sK13,X48))))) ) | $spl2),
  inference(superposition,[],[f798,f21375])).
fof(f31389,plain,(
  ( ! [X43] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X43,sK13))))) | ~p(s(bool,X43))) ) | $spl610),
  inference(cnf_transformation,[],[f31389_D])).
fof(f31389_D,plain,(
  ( ! [X43] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X43,sK13))))) | ~p(s(bool,X43))) ) <=> ~$spl610),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl610])])).
fof(f31390,plain,(
  ( ! [X45,X46] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X46,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X45)))) | ~p(s(bool,X46))) ) | $spl2),
  inference(forward_demodulation,[],[f31023,f21375])).
fof(f31023,plain,(
  ( ! [X45,X46] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X45)))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))),s(sK13,sK23('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X46,sK13))))) | ~p(s(bool,X46))) ) | $spl2),
  inference(superposition,[],[f797,f21375])).
fof(f797,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK23(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f561])).
fof(f561,plain,(
  ! [X0,X1,X2] : ((! [X6] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK22(X2,X1,X0))))))) & (~p(s(bool,X1)) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK23(X2,X1,X0))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK22,sK23])],[f560])).
fof(f560,plain,(
  ! [X0,X1,X2] : ((! [X6] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | (p(s(bool,X1)) & ? [X5] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))))) & (~p(s(bool,X1)) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X3] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(rectify,[],[f559])).
fof(f559,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (p(s(bool,X1)) & ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (~p(s(bool,X1)) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X3] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(flattening,[],[f558])).
fof(f558,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (p(s(bool,X1)) & ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & ((~p(s(bool,X1)) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ? [X3] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f303])).
fof(f303,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (p(s(bool,X1)) & ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f302])).
fof(f302,plain,(
  ! [X0] : ! [X1,X2] : (? [X3] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (p(s(bool,X1)) & ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f133])).
fof(f133,axiom,(
  ! [X6] : ! [X18,X19] : (? [X9] : (p(s(bool,X18)) & p(s(bool,happ(s(fun(X6,bool),X19),s(X6,X9))))) <=> (p(s(bool,X18)) & ? [X9] : p(s(bool,happ(s(fun(X6,bool),X19),s(X6,X9))))))),
  file('thm.pred_set.SUBSET_DIFF_EMPTY','thm.bool.RIGHT_EXISTS_AND_THM')).
fof(f31388,plain,(
  ( ! [X43,X44] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X44)))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X43,sK13))))) | ~p(s(bool,X43))) ) | $spl2),
  inference(forward_demodulation,[],[f31022,f21375])).
fof(f31022,plain,(
  ( ! [X43,X44] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X43,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))),s(sK13,X44)))) | ~p(s(bool,X43))) ) | $spl2),
  inference(superposition,[],[f797,f21375])).
fof(f31386,plain,(
  ( ! [X39] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK22('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X39,sK13))))) | ~p(s(bool,X39))) ) | $spl608),
  inference(cnf_transformation,[],[f31386_D])).
fof(f31386_D,plain,(
  ( ! [X39] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK22('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X39,sK13))))) | ~p(s(bool,X39))) ) <=> ~$spl608),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl608])])).
fof(f31387,plain,(
  ( ! [X41,X42] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK22('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X42,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X41)))) | ~p(s(bool,X42))) ) | $spl2),
  inference(forward_demodulation,[],[f31021,f21375])).
fof(f31021,plain,(
  ( ! [X41,X42] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X41)))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))),s(sK13,sK22('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X42,sK13))))) | ~p(s(bool,X42))) ) | $spl2),
  inference(superposition,[],[f795,f21375])).
fof(f795,plain,(
  ( ! [X6,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK22(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f561])).
fof(f31385,plain,(
  ( ! [X39,X40] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X40)))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK22('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X39,sK13))))) | ~p(s(bool,X39))) ) | $spl2),
  inference(forward_demodulation,[],[f31020,f21375])).
fof(f31020,plain,(
  ( ! [X39,X40] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK22('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X39,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))),s(sK13,X40)))) | ~p(s(bool,X39))) ) | $spl2),
  inference(superposition,[],[f795,f21375])).
fof(f31379,plain,(
  ( ! [X28] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK20(X28,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13))))) | ~p(s(bool,X28))) ) | $spl606),
  inference(cnf_transformation,[],[f31379_D])).
fof(f31379_D,plain,(
  ( ! [X28] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK20(X28,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13))))) | ~p(s(bool,X28))) ) <=> ~$spl606),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl606])])).
fof(f31380,plain,(
  ( ! [X30,X31] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK20(X31,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X30)))) | ~p(s(bool,X31))) ) | $spl2),
  inference(forward_demodulation,[],[f31014,f21375])).
fof(f31014,plain,(
  ( ! [X30,X31] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X30)))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))),s(sK13,sK20(X31,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13))))) | ~p(s(bool,X31))) ) | $spl2),
  inference(superposition,[],[f789,f21375])).
fof(f789,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK20(X2,X1,X0))))) | ~p(s(bool,X2)) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) )),
  inference(cnf_transformation,[],[f554])).
fof(f554,plain,(
  ! [X0,X1,X2] : ((! [X6] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) | ~p(s(bool,X2))) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK19(X2,X1,X0))))) & p(s(bool,X2)))) & (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,X2)) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK20(X2,X1,X0))))) & p(s(bool,X2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK19,sK20])],[f553])).
fof(f553,plain,(
  ! [X0,X1,X2] : ((! [X6] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) | ~p(s(bool,X2))) | (? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & p(s(bool,X2)))) & (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,X2)) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2)))))),
  inference(rectify,[],[f552])).
fof(f552,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,X2))) | (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2)))) & (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,X2)) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2)))))),
  inference(flattening,[],[f551])).
fof(f551,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,X2))) | (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2)))) & ((! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,X2))) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2)))))),
  inference(nnf_transformation,[],[f299])).
fof(f299,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2))) <=> (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2))))),
  inference(flattening,[],[f298])).
fof(f298,plain,(
  ! [X0] : ! [X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2))) <=> (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2))))),
  inference(rectify,[],[f153])).
fof(f153,axiom,(
  ! [X6] : ! [X18,X19] : (? [X9] : (p(s(bool,happ(s(fun(X6,bool),X18),s(X6,X9)))) & p(s(bool,X19))) <=> (? [X9] : p(s(bool,happ(s(fun(X6,bool),X18),s(X6,X9)))) & p(s(bool,X19))))),
  file('thm.pred_set.SUBSET_DIFF_EMPTY','thm.bool.LEFT_EXISTS_AND_THM')).
fof(f31378,plain,(
  ( ! [X28,X29] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X29)))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK20(X28,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13))))) | ~p(s(bool,X28))) ) | $spl2),
  inference(forward_demodulation,[],[f31013,f21375])).
fof(f31013,plain,(
  ( ! [X28,X29] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK20(X28,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13))))) | ~p(s(bool,X28)) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))),s(sK13,X29))))) ) | $spl2),
  inference(superposition,[],[f789,f21375])).
fof(f31376,plain,(
  ( ! [X24] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK19(X24,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13))))) | ~p(s(bool,X24))) ) | $spl604),
  inference(cnf_transformation,[],[f31376_D])).
fof(f31376_D,plain,(
  ( ! [X24] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK19(X24,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13))))) | ~p(s(bool,X24))) ) <=> ~$spl604),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl604])])).
fof(f31377,plain,(
  ( ! [X26,X27] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK19(X27,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X26)))) | ~p(s(bool,X27))) ) | $spl2),
  inference(forward_demodulation,[],[f31012,f21375])).
fof(f31012,plain,(
  ( ! [X26,X27] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X26)))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))),s(sK13,sK19(X27,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13))))) | ~p(s(bool,X27))) ) | $spl2),
  inference(superposition,[],[f787,f21375])).
fof(f787,plain,(
  ( ! [X6,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK19(X2,X1,X0))))) | ~p(s(bool,X2)) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6))))) )),
  inference(cnf_transformation,[],[f554])).
fof(f31375,plain,(
  ( ! [X24,X25] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X25)))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK19(X24,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13))))) | ~p(s(bool,X24))) ) | $spl2),
  inference(forward_demodulation,[],[f31011,f21375])).
fof(f31011,plain,(
  ( ! [X24,X25] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK19(X24,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13))))) | ~p(s(bool,X24)) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))),s(sK13,X25))))) ) | $spl2),
  inference(superposition,[],[f787,f21375])).
fof(f31373,plain,(
  ( ! [X20] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK18('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X20,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X20),s(sK13,sK18('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X20,sK13)))))) ) | $spl602),
  inference(cnf_transformation,[],[f31373_D])).
fof(f31373_D,plain,(
  ( ! [X20] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK18('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X20,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X20),s(sK13,sK18('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X20,sK13)))))) ) <=> ~$spl602),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl602])])).
fof(f31374,plain,(
  ( ! [X23,X22] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK18('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X23,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X22)))) | p(s(bool,happ(s(fun(sK13,bool),X23),s(sK13,sK18('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X23,sK13)))))) ) | $spl2),
  inference(forward_demodulation,[],[f31010,f21375])).
fof(f31010,plain,(
  ( ! [X23,X22] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X22)))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))),s(sK13,sK18('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X23,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X23),s(sK13,sK18('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X23,sK13)))))) ) | $spl2),
  inference(superposition,[],[f786,f21375])).
fof(f786,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK18(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK18(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) )),
  inference(cnf_transformation,[],[f550])).
fof(f550,plain,(
  ! [X0,X1,X2] : ((! [X8] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X8)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X8))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK16(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK17(X2,X1,X0)))))) & ((! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK18(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK18(X2,X1,X0))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK16,sK17,sK18])],[f549])).
fof(f549,plain,(
  ! [X0,X1,X2] : ((! [X8] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X8)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X8))))) | ? [X7] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X7)))) | ? [X6] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) & ((! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(rectify,[],[f548])).
fof(f548,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) & ((! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(flattening,[],[f547])).
fof(f547,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & ((! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f297])).
fof(f297,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f296])).
fof(f296,plain,(
  ! [X0] : ! [X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f154])).
fof(f154,axiom,(
  ! [X6] : ! [X18,X19] : (? [X9] : (p(s(bool,happ(s(fun(X6,bool),X18),s(X6,X9)))) | p(s(bool,happ(s(fun(X6,bool),X19),s(X6,X9))))) <=> (? [X9] : p(s(bool,happ(s(fun(X6,bool),X18),s(X6,X9)))) | ? [X9] : p(s(bool,happ(s(fun(X6,bool),X19),s(X6,X9))))))),
  file('thm.pred_set.SUBSET_DIFF_EMPTY','thm.bool.EXISTS_OR_THM')).
fof(f31372,plain,(
  ( ! [X21,X20] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X21)))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK18('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X20,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X20),s(sK13,sK18('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X20,sK13)))))) ) | $spl2),
  inference(forward_demodulation,[],[f31009,f21375])).
fof(f31009,plain,(
  ( ! [X21,X20] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK18('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X20,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X20),s(sK13,sK18('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X20,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))),s(sK13,X21))))) ) | $spl2),
  inference(superposition,[],[f786,f21375])).
fof(f31370,plain,(
  ( ! [X12] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK18(X12,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X12),s(sK13,sK18(X12,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13)))))) ) | $spl600),
  inference(cnf_transformation,[],[f31370_D])).
fof(f31370_D,plain,(
  ( ! [X12] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK18(X12,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X12),s(sK13,sK18(X12,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13)))))) ) <=> ~$spl600),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl600])])).
fof(f31371,plain,(
  ( ! [X17,X16] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK18(X17,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X16)))) | p(s(bool,happ(s(fun(sK13,bool),X17),s(sK13,sK18(X17,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13)))))) ) | $spl2),
  inference(forward_demodulation,[],[f31007,f21375])).
fof(f31007,plain,(
  ( ! [X17,X16] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X16)))) | p(s(bool,happ(s(fun(sK13,bool),X17),s(sK13,sK18(X17,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13))))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))),s(sK13,sK18(X17,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13)))))) ) | $spl2),
  inference(superposition,[],[f785,f21375])).
fof(f785,plain,(
  ( ! [X2,X0,X5,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK18(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK18(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))) )),
  inference(cnf_transformation,[],[f550])).
fof(f31369,plain,(
  ( ! [X12,X13] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X13)))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK18(X12,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X12),s(sK13,sK18(X12,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13)))))) ) | $spl2),
  inference(forward_demodulation,[],[f31005,f21375])).
fof(f31005,plain,(
  ( ! [X12,X13] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK18(X12,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X12),s(sK13,sK18(X12,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))),s(sK13,X13))))) ) | $spl2),
  inference(superposition,[],[f785,f21375])).
fof(f31367,plain,(
  ( ! [X8] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK17('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X8,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X8),s(sK13,sK16('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X8,sK13)))))) ) | $spl598),
  inference(cnf_transformation,[],[f31367_D])).
fof(f31367_D,plain,(
  ( ! [X8] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK17('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X8,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X8),s(sK13,sK16('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X8,sK13)))))) ) <=> ~$spl598),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl598])])).
fof(f31368,plain,(
  ( ! [X10,X11] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK17('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X11,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X10)))) | p(s(bool,happ(s(fun(sK13,bool),X11),s(sK13,sK16('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X11,sK13)))))) ) | $spl2),
  inference(forward_demodulation,[],[f31004,f21375])).
fof(f31004,plain,(
  ( ! [X10,X11] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X10)))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))),s(sK13,sK17('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X11,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X11),s(sK13,sK16('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X11,sK13)))))) ) | $spl2),
  inference(superposition,[],[f784,f21375])).
fof(f784,plain,(
  ( ! [X2,X0,X8,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK17(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK16(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X8))))) )),
  inference(cnf_transformation,[],[f550])).
fof(f31366,plain,(
  ( ! [X8,X9] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X9)))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK17('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X8,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X8),s(sK13,sK16('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X8,sK13)))))) ) | $spl2),
  inference(forward_demodulation,[],[f31003,f21375])).
fof(f31003,plain,(
  ( ! [X8,X9] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK17('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X8,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X8),s(sK13,sK16('const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X8,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))),s(sK13,X9))))) ) | $spl2),
  inference(superposition,[],[f784,f21375])).
fof(f31364,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK16(X0,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X0),s(sK13,sK17(X0,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13)))))) ) | $spl596),
  inference(cnf_transformation,[],[f31364_D])).
fof(f31364_D,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK16(X0,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X0),s(sK13,sK17(X0,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13)))))) ) <=> ~$spl596),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl596])])).
fof(f31365,plain,(
  ( ! [X4,X5] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK16(X5,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X4)))) | p(s(bool,happ(s(fun(sK13,bool),X5),s(sK13,sK17(X5,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13)))))) ) | $spl2),
  inference(forward_demodulation,[],[f31001,f21375])).
fof(f31001,plain,(
  ( ! [X4,X5] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X4)))) | p(s(bool,happ(s(fun(sK13,bool),X5),s(sK13,sK17(X5,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13))))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))),s(sK13,sK16(X5,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13)))))) ) | $spl2),
  inference(superposition,[],[f783,f21375])).
fof(f783,plain,(
  ( ! [X2,X0,X8,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK17(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK16(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X8))))) )),
  inference(cnf_transformation,[],[f550])).
fof(f31363,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X1)))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK16(X0,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X0),s(sK13,sK17(X0,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13)))))) ) | $spl2),
  inference(forward_demodulation,[],[f30999,f21375])).
fof(f30999,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK16(X0,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X0),s(sK13,sK17(X0,'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))),s(sK13,X1))))) ) | $spl2),
  inference(superposition,[],[f783,f21375])).
fof(f30304,plain,(
  $false | ($spl14 | $spl202)),
  inference(equality_resolution,[],[f30247])).
fof(f30247,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))) ) | ($spl14 | $spl202)),
  inference(subsumption_resolution,[],[f4743,f1095])).
fof(f29643,plain,(
  ( ! [X23] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.pred_set.FINITE_1'(s(fun(X23,bool),'const.pred_set.EMPTY_0')))) ) | $spl592),
  inference(cnf_transformation,[],[f29643_D])).
fof(f29643_D,plain,(
  ( ! [X23] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.pred_set.FINITE_1'(s(fun(X23,bool),'const.pred_set.EMPTY_0')))) ) <=> ~$spl592),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl592])])).
fof(f30094,plain,(
  ( ! [X26,X27,X25] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.pred_set.FINITE_1'(s(fun(X25,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X26)) | ~p(s(bool,X27))) ) | ($spl8 | $spl10)),
  inference(resolution,[],[f29382,f21291])).
fof(f30083,plain,(
  ( ! [X21,X19,X20] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.pred_set.FINITE_1'(s(fun(X19,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X20)) | ~p(s(bool,X21))) ) | ($spl8 | $spl10)),
  inference(resolution,[],[f29381,f21292])).
fof(f30074,plain,(
  ( ! [X17] : (p(s(bool,X17)) | s(bool,X17) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))))) ) | $spl594),
  inference(cnf_transformation,[],[f30074_D])).
fof(f30074_D,plain,(
  ( ! [X17] : (p(s(bool,X17)) | s(bool,X17) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))))) ) <=> ~$spl594),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl594])])).
fof(f30006,plain,(
  ( ! [X227,X226] : (p(s(bool,X226)) | p(s(bool,X227)) | s(bool,X227) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))))) ) | $spl42),
  inference(resolution,[],[f1166,f905])).
fof(f29908,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))))) | s(bool,X38) != s(bool,X39)) ) | $spl42),
  inference(resolution,[],[f1166,f630])).
fof(f29905,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X28)) | ~p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X30) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))))) ) | $spl42),
  inference(resolution,[],[f1166,f629])).
fof(f29904,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))))) ) | $spl42),
  inference(resolution,[],[f1166,f628])).
fof(f29903,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))))) | s(bool,X23) != s(bool,X24)) ) | $spl42),
  inference(resolution,[],[f1166,f626])).
fof(f29901,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X16) != s(bool,X18) | s(bool,X17) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))))) ) | $spl42),
  inference(resolution,[],[f1166,f624])).
fof(f29879,plain,(
  ( ! [X26,X27,X25] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.pred_set.FINITE_1'(s(fun(X25,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X26)) | ~p(s(bool,X27))) ) | ($spl8 | $spl10)),
  inference(resolution,[],[f29379,f21288])).
fof(f29868,plain,(
  ( ! [X21,X19,X20] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.pred_set.FINITE_1'(s(fun(X19,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X20)) | ~p(s(bool,X21))) ) | ($spl8 | $spl10)),
  inference(resolution,[],[f29378,f21289])).
fof(f29861,plain,(
  ( ! [X30,X31,X29] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.pred_set.FINITE_1'(s(fun(X29,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X30)) | ~p(s(bool,X31))) ) | ($spl8 | $spl10)),
  inference(resolution,[],[f29366,f21275])).
fof(f29848,plain,(
  ( ! [X24,X23,X25] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.pred_set.FINITE_1'(s(fun(X23,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X24)) | ~p(s(bool,X25))) ) | ($spl8 | $spl10)),
  inference(resolution,[],[f29365,f21276])).
fof(f29841,plain,(
  ( ! [X24,X23,X25] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.pred_set.FINITE_1'(s(fun(X23,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X24)) | p(s(bool,X25))) ) | ($spl8 | $spl10)),
  inference(resolution,[],[f29352,f21259])).
fof(f29822,plain,(
  $false | ($spl10 | $spl43)),
  inference(subsumption_resolution,[],[f29653,f1092])).
fof(f29653,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | $spl43),
  inference(resolution,[],[f1167,f888])).
fof(f29652,plain,(
  ( ! [X30,X31,X29] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.pred_set.FINITE_1'(s(fun(X29,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X30)) | p(s(bool,X31))) ) | ($spl8 | $spl10)),
  inference(resolution,[],[f29336,f21245])).
fof(f29639,plain,(
  ( ! [X24,X23,X25] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.pred_set.FINITE_1'(s(fun(X23,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X24)) | ~p(s(bool,X25))) ) | ($spl8 | $spl10)),
  inference(resolution,[],[f29335,f21246])).
fof(f29526,plain,(
  ( ! [X0,X3,X1] : (s(bool,X0) = s(bool,'const.pred_set.FINITE_1'(s(fun(X1,bool),'const.pred_set.EMPTY_0'))) | s(bool,X0) != s(bool,X3) | ~p(s(bool,X3))) ) | $spl590),
  inference(cnf_transformation,[],[f29526_D])).
fof(f29526_D,plain,(
  ( ! [X0,X3,X1] : (s(bool,X0) = s(bool,'const.pred_set.FINITE_1'(s(fun(X1,bool),'const.pred_set.EMPTY_0'))) | s(bool,X0) != s(bool,X3) | ~p(s(bool,X3))) ) <=> ~$spl590),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl590])])).
fof(f29426,plain,(
  ( ! [X422,X424,X423,X425] : (s(bool,X422) = s(bool,'const.pred_set.FINITE_1'(s(fun(X423,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X424)) | ~p(s(bool,X425)) | s(bool,X422) != s(bool,X425)) ) | $spl8),
  inference(resolution,[],[f28487,f905])).
fof(f29328,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s(bool,X35) = s(bool,'const.pred_set.FINITE_1'(s(fun(X36,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X35) != s(bool,X38) | s(bool,X37) != s(bool,X39)) ) | $spl8),
  inference(resolution,[],[f28487,f630])).
fof(f29326,plain,(
  ( ! [X28,X26,X29,X27,X25] : (s(bool,X25) = s(bool,'const.pred_set.FINITE_1'(s(fun(X26,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X27)) | ~p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X25) != s(bool,X28)) ) | $spl8),
  inference(resolution,[],[f28487,f629])).
fof(f29324,plain,(
  ( ! [X19,X17,X15,X18,X16] : (s(bool,X15) = s(bool,'const.pred_set.FINITE_1'(s(fun(X16,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X15) != s(bool,X18)) ) | $spl8),
  inference(resolution,[],[f28487,f628])).
fof(f29323,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s(bool,X10) = s(bool,'const.pred_set.FINITE_1'(s(fun(X11,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X10) != s(bool,X13) | s(bool,X12) != s(bool,X14)) ) | $spl8),
  inference(resolution,[],[f28487,f626])).
fof(f29321,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.pred_set.FINITE_1'(s(fun(X1,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X0) != s(bool,X3)) ) | $spl8),
  inference(resolution,[],[f28487,f624])).
fof(f27713,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) | $spl588),
  inference(cnf_transformation,[],[f27713_D])).
fof(f27713_D,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) <=> ~$spl588),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl588])])).
fof(f27615,plain,(
  ( ! [X318,X317,X319] : (s(bool,X317) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X318)) | ~p(s(bool,X319)) | s(bool,X317) != s(bool,X319)) ) | $spl304),
  inference(resolution,[],[f21881,f905])).
fof(f27517,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) ) | $spl304),
  inference(resolution,[],[f21881,f630])).
fof(f27515,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,X20) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X20) != s(bool,X22)) ) | $spl304),
  inference(resolution,[],[f21881,f629])).
fof(f27513,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl304),
  inference(resolution,[],[f21881,f628])).
fof(f27512,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) ) | $spl304),
  inference(resolution,[],[f21881,f626])).
fof(f27510,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl304),
  inference(resolution,[],[f21881,f624])).
fof(f17531,plain,(
  s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) | $spl406),
  inference(cnf_transformation,[],[f17531_D])).
fof(f17531_D,plain,(
  s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) <=> ~$spl406),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl406])])).
fof(f27232,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | ($spl8 | $spl10)),
  inference(resolution,[],[f21292,f16508])).
fof(f16508,plain,(
  ( ! [X189,X190,X188] : (sP7(X189,X188,X190) | s(bool,X188) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | $spl8),
  inference(resolution,[],[f7688,f695])).
fof(f27200,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | ($spl8 | $spl10)),
  inference(resolution,[],[f21291,f16509])).
fof(f16509,plain,(
  ( ! [X191,X192,X193] : (sP7(X191,X192,X193) | s(bool,X191) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | $spl8),
  inference(resolution,[],[f7688,f696])).
fof(f27164,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | ($spl8 | $spl10)),
  inference(resolution,[],[f21289,f16505])).
fof(f16505,plain,(
  ( ! [X180,X181,X179] : (sP8(X180,X179,X181) | s(bool,X179) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | $spl8),
  inference(resolution,[],[f7688,f691])).
fof(f27132,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | ($spl8 | $spl10)),
  inference(resolution,[],[f21288,f16506])).
fof(f16506,plain,(
  ( ! [X182,X184,X183] : (sP8(X182,X183,X184) | s(bool,X182) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | $spl8),
  inference(resolution,[],[f7688,f692])).
fof(f26875,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | ($spl8 | $spl10)),
  inference(resolution,[],[f21276,f16492])).
fof(f16492,plain,(
  ( ! [X142,X140,X141] : (sP4(X141,X140,X142) | s(bool,X140) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | $spl8),
  inference(resolution,[],[f7688,f674])).
fof(f26842,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | ($spl8 | $spl10)),
  inference(resolution,[],[f21275,f16493])).
fof(f16493,plain,(
  ( ! [X144,X145,X143] : (sP4(X143,X144,X145) | s(bool,X143) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | $spl8),
  inference(resolution,[],[f7688,f675])).
fof(f26761,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | p(s(bool,X3)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | ($spl8 | $spl10)),
  inference(resolution,[],[f21259,f16479])).
fof(f16479,plain,(
  ( ! [X103,X101,X102] : (sP2(X101,X102,X103) | s(bool,X101) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | $spl8),
  inference(resolution,[],[f7688,f658])).
fof(f26565,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | ($spl8 | $spl10)),
  inference(resolution,[],[f21246,f16462])).
fof(f16462,plain,(
  ( ! [X52,X50,X51] : (sP0(X51,X52,X50) | s(bool,X50) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | $spl8),
  inference(resolution,[],[f7688,f639])).
fof(f26532,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | p(s(bool,X11)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | ($spl8 | $spl10)),
  inference(resolution,[],[f21245,f16463])).
fof(f16463,plain,(
  ( ! [X54,X55,X53] : (sP0(X53,X54,X55) | s(bool,X53) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | $spl8),
  inference(resolution,[],[f7688,f641])).
fof(f26391,plain,(
  ( ! [X14,X15] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) | p(s(bool,X14)) | ~p(s(bool,X15))) ) | ($spl8 | $spl10)),
  inference(resolution,[],[f25140,f7652])).
fof(f26382,plain,(
  ( ! [X12,X13] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) | p(s(bool,X12)) | ~p(s(bool,X13))) ) | ($spl8 | $spl10)),
  inference(resolution,[],[f25139,f7653])).
fof(f26374,plain,(
  ( ! [X14,X15] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) | p(s(bool,X14)) | ~p(s(bool,X15))) ) | ($spl8 | $spl10)),
  inference(resolution,[],[f25137,f7649])).
fof(f26366,plain,(
  ( ! [X12,X13] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) | p(s(bool,X12)) | ~p(s(bool,X13))) ) | ($spl8 | $spl10)),
  inference(resolution,[],[f25136,f7650])).
fof(f26360,plain,(
  ( ! [X17,X18] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) | p(s(bool,X17)) | ~p(s(bool,X18))) ) | ($spl8 | $spl10)),
  inference(resolution,[],[f25124,f7636])).
fof(f26350,plain,(
  ( ! [X15,X16] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) | p(s(bool,X15)) | ~p(s(bool,X16))) ) | ($spl8 | $spl10)),
  inference(resolution,[],[f25123,f7637])).
fof(f26344,plain,(
  ( ! [X15,X16] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) | p(s(bool,X15)) | p(s(bool,X16))) ) | ($spl8 | $spl10)),
  inference(resolution,[],[f25110,f7620])).
fof(f26290,plain,(
  ( ! [X17,X18] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) | ~p(s(bool,X17)) | p(s(bool,X18))) ) | ($spl8 | $spl10)),
  inference(resolution,[],[f25094,f7606])).
fof(f26281,plain,(
  ( ! [X15,X16] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) | p(s(bool,X15)) | ~p(s(bool,X16))) ) | ($spl8 | $spl10)),
  inference(resolution,[],[f25093,f7607])).
fof(f16361,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) | $spl402),
  inference(cnf_transformation,[],[f16361_D])).
fof(f16361_D,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) <=> ~$spl402),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl402])])).
fof(f25184,plain,(
  ( ! [X318,X317,X319] : (s(bool,X317) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | p(s(bool,X318)) | ~p(s(bool,X319)) | s(bool,X317) != s(bool,X319)) ) | $spl10),
  inference(resolution,[],[f21327,f905])).
fof(f25086,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) ) | $spl10),
  inference(resolution,[],[f21327,f630])).
fof(f25084,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,X20) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X20) != s(bool,X22)) ) | $spl10),
  inference(resolution,[],[f21327,f629])).
fof(f25082,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl10),
  inference(resolution,[],[f21327,f628])).
fof(f25081,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) ) | $spl10),
  inference(resolution,[],[f21327,f626])).
fof(f25079,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl10),
  inference(resolution,[],[f21327,f624])).
fof(f23097,plain,(
  $false | ($spl54 | $spl56)),
  inference(equality_resolution,[],[f20891])).
fof(f20891,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | ($spl54 | $spl56)),
  inference(subsumption_resolution,[],[f2198,f2053])).
fof(f22297,plain,(
  ( ! [X47,X49] : (p(s(bool,X47)) | s(bool,X47) != s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X49)))) ) | $spl586),
  inference(cnf_transformation,[],[f22297_D])).
fof(f22297_D,plain,(
  ( ! [X47,X49] : (p(s(bool,X47)) | s(bool,X47) != s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X49)))) ) <=> ~$spl586),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl586])])).
fof(f22256,plain,(
  ( ! [X360,X362,X361] : (p(s(bool,X360)) | p(s(bool,X361)) | s(bool,X361) != s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X362)))) ) | $spl326),
  inference(resolution,[],[f8530,f905])).
fof(f22158,plain,(
  ( ! [X76,X74,X75,X73] : (p(s(bool,X73)) | ~p(s(bool,X74)) | s(bool,X73) != s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X75))) | s(bool,X74) != s(bool,X76)) ) | $spl326),
  inference(resolution,[],[f8530,f630])).
fof(f22155,plain,(
  ( ! [X64,X62,X65,X63] : (p(s(bool,X62)) | ~p(s(bool,X63)) | p(s(bool,X64)) | s(bool,X64) != s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X65)))) ) | $spl326),
  inference(resolution,[],[f8530,f629])).
fof(f22154,plain,(
  ( ! [X61,X59,X60,X58] : (p(s(bool,X58)) | p(s(bool,X59)) | s(bool,X58) != s(bool,X60) | s(bool,X59) != s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X61)))) ) | $spl326),
  inference(resolution,[],[f8530,f628])).
fof(f22153,plain,(
  ( ! [X57,X54,X56,X55] : (p(s(bool,X54)) | ~p(s(bool,X55)) | s(bool,X54) != s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X56))) | s(bool,X55) != s(bool,X57)) ) | $spl326),
  inference(resolution,[],[f8530,f626])).
fof(f22151,plain,(
  ( ! [X47,X48,X46,X49] : (p(s(bool,X46)) | p(s(bool,X47)) | s(bool,X46) != s(bool,X48) | s(bool,X47) != s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X49)))) ) | $spl326),
  inference(resolution,[],[f8530,f624])).
fof(f22119,plain,(
  $false | ($spl308 | $spl326)),
  inference(subsumption_resolution,[],[f8530,f8501])).
fof(f22107,plain,(
  ( ! [X48,X46] : (~p(s(bool,X46)) | s(bool,X46) != s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X48)))) ) | $spl584),
  inference(cnf_transformation,[],[f22107_D])).
fof(f22107_D,plain,(
  ( ! [X48,X46] : (~p(s(bool,X46)) | s(bool,X46) != s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X48)))) ) <=> ~$spl584),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl584])])).
fof(f22069,plain,(
  ( ! [X362,X361,X363] : (p(s(bool,X361)) | ~p(s(bool,X362)) | s(bool,X362) != s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X363)))) ) | $spl308),
  inference(resolution,[],[f8501,f905])).
fof(f21971,plain,(
  ( ! [X74,X72,X75,X73] : (~p(s(bool,X72)) | ~p(s(bool,X73)) | s(bool,X73) != s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X74))) | s(bool,X72) != s(bool,X75)) ) | $spl308),
  inference(resolution,[],[f8501,f630])).
fof(f21969,plain,(
  ( ! [X68,X66,X67,X65] : (p(s(bool,X65)) | ~p(s(bool,X66)) | ~p(s(bool,X67)) | s(bool,X66) != s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X68)))) ) | $spl308),
  inference(resolution,[],[f8501,f629])).
fof(f21967,plain,(
  ( ! [X59,X57,X60,X58] : (p(s(bool,X57)) | ~p(s(bool,X58)) | s(bool,X57) != s(bool,X59) | s(bool,X58) != s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X60)))) ) | $spl308),
  inference(resolution,[],[f8501,f628])).
fof(f21966,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X54) != s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X55))) | s(bool,X53) != s(bool,X56)) ) | $spl308),
  inference(resolution,[],[f8501,f626])).
fof(f21964,plain,(
  ( ! [X47,X45,X48,X46] : (p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X45) != s(bool,X47) | s(bool,X46) != s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X48)))) ) | $spl308),
  inference(resolution,[],[f8501,f624])).
fof(f21917,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl582),
  inference(cnf_transformation,[],[f21917_D])).
fof(f21917_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl582),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl582])])).
fof(f21891,plain,(
  ( ! [X210,X211] : (p(s(bool,X210)) | p(s(bool,X211)) | s(bool,X211) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl304),
  inference(resolution,[],[f8498,f905])).
fof(f21793,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s(bool,X22) != s(bool,X23)) ) | $spl304),
  inference(resolution,[],[f8498,f630])).
fof(f21790,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl304),
  inference(resolution,[],[f8498,f629])).
fof(f21789,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl304),
  inference(resolution,[],[f8498,f628])).
fof(f21788,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s(bool,X7) != s(bool,X8)) ) | $spl304),
  inference(resolution,[],[f8498,f626])).
fof(f21786,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl304),
  inference(resolution,[],[f8498,f624])).
fof(f1167,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))))) | $spl43),
  inference(cnf_transformation,[],[f1167_D])).
fof(f1167_D,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))))) <=> ~$spl43),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl43])])).
fof(f21683,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))))) | $spl1),
  inference(trivial_inequality_removal,[],[f21678])).
fof(f21678,plain,(
  s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),'const.pred_set.EMPTY_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))))) | $spl1),
  inference(superposition,[],[f914,f897])).
fof(f1168,plain,(
  ( ! [X5] : (s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),X5) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))))))) ) | $spl44),
  inference(cnf_transformation,[],[f1168_D])).
fof(f1168_D,plain,(
  ( ! [X5] : (s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),X5) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))))))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f21692,plain,(
  ( ! [X5] : (s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),X5) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))))))) ) | $spl1),
  inference(subsumption_resolution,[],[f21677,f860])).
fof(f21677,plain,(
  ( ! [X5] : (s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),X5) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X5))))) ) | $spl1),
  inference(superposition,[],[f914,f866])).
fof(f1095,plain,(
  ( ! [X4] : (s(bool,X4) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))) | p(s(bool,X4))) ) | $spl14),
  inference(cnf_transformation,[],[f1095_D])).
fof(f1095_D,plain,(
  ( ! [X4] : (s(bool,X4) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))) | p(s(bool,X4))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f21494,plain,(
  ( ! [X213,X214] : (p(s(bool,X213)) | p(s(bool,X214)) | s(bool,X214) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))) ) | $spl2),
  inference(resolution,[],[f915,f905])).
fof(f21396,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))) | s(bool,X25) != s(bool,X26)) ) | $spl2),
  inference(resolution,[],[f915,f630])).
fof(f21393,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X17) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))) ) | $spl2),
  inference(resolution,[],[f915,f629])).
fof(f21392,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))) ) | $spl2),
  inference(resolution,[],[f915,f628])).
fof(f21391,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))) | s(bool,X10) != s(bool,X11)) ) | $spl2),
  inference(resolution,[],[f915,f626])).
fof(f21389,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))) ) | $spl2),
  inference(resolution,[],[f915,f624])).
fof(f2053,plain,(
  ( ! [X17] : (s(bool,X17) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | p(s(bool,X17))) ) | $spl54),
  inference(cnf_transformation,[],[f2053_D])).
fof(f2053_D,plain,(
  ( ! [X17] : (s(bool,X17) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | p(s(bool,X17))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f21337,plain,(
  ( ! [X227,X226] : (p(s(bool,X226)) | p(s(bool,X227)) | s(bool,X227) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl10),
  inference(resolution,[],[f1092,f905])).
fof(f21239,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | s(bool,X38) != s(bool,X39)) ) | $spl10),
  inference(resolution,[],[f1092,f630])).
fof(f21236,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X28)) | ~p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X30) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl10),
  inference(resolution,[],[f1092,f629])).
fof(f21235,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl10),
  inference(resolution,[],[f1092,f628])).
fof(f21234,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | s(bool,X23) != s(bool,X24)) ) | $spl10),
  inference(resolution,[],[f1092,f626])).
fof(f21232,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X16) != s(bool,X18) | s(bool,X17) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl10),
  inference(resolution,[],[f1092,f624])).
fof(f21213,plain,(
  $false | ($spl0 | $spl8 | $spl11)),
  inference(subsumption_resolution,[],[f20892,f1091])).
fof(f20892,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f8453,f1089])).
fof(f8453,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) | p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | $spl0),
  inference(trivial_inequality_removal,[],[f8121])).
fof(f8121,plain,(
  s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),'const.pred_set.EMPTY_0') | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) | p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | $spl0),
  inference(superposition,[],[f879,f913])).
fof(f879,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),'const.pred_set.EMPTY_0') != s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) | p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f447])).
fof(f447,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') != s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))),
  inference(flattening,[],[f446])).
fof(f446,plain,(
  ! [X0,X1,X2] : ((p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2))))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') != s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))),
  inference(ennf_transformation,[],[f376])).
fof(f376,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) & p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2))))) => s(fun(X0,bool),'const.pred_set.EMPTY_0') != s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))),
  inference(flattening,[],[f375])).
fof(f375,plain,(
  ! [X0] : ! [X1,X2] : ((~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) & p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2))))) => ~s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X6] : ! [X7,X8] : ((~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X7)))) & p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X8))))) => ~s(fun(X6,bool),'const.pred_set.EMPTY_0') = s(fun(X6,bool),'const.pred_set.DIFF_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8))))),
  file('thm.pred_set.SUBSET_DIFF_EMPTY','thm.pred_set.INFINITE_DIFF_FINITE')).
fof(f21038,plain,(
  $false | ($spl8 | $spl360)),
  inference(resolution,[],[f8813,f1089])).
fof(f21044,plain,(
  $false | $spl360),
  inference(resolution,[],[f8813,f902])).
fof(f2198,plain,(
  ( ! [X5] : (~p(s(bool,X5)) | s(bool,X5) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl56),
  inference(cnf_transformation,[],[f2198_D])).
fof(f2198_D,plain,(
  ( ! [X5] : (~p(s(bool,X5)) | s(bool,X5) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f21002,plain,(
  ( ! [X216,X217] : (p(s(bool,X216)) | ~p(s(bool,X217)) | s(bool,X217) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl11),
  inference(resolution,[],[f1091,f905])).
fof(f20904,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X26) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | s(bool,X25) != s(bool,X27)) ) | $spl11),
  inference(resolution,[],[f1091,f630])).
fof(f20902,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X20) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl11),
  inference(resolution,[],[f1091,f629])).
fof(f20900,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl11),
  inference(resolution,[],[f1091,f628])).
fof(f20899,plain,(
  ( ! [X12,X10,X11] : (~p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X11) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | s(bool,X10) != s(bool,X12)) ) | $spl11),
  inference(resolution,[],[f1091,f626])).
fof(f20897,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl11),
  inference(resolution,[],[f1091,f624])).
fof(f1091,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | $spl11),
  inference(cnf_transformation,[],[f1091_D])).
fof(f1091_D,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f20890,plain,(
  ( ! [X5] : (s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),X5) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK14))))) ) | $spl580),
  inference(cnf_transformation,[],[f20890_D])).
fof(f20890_D,plain,(
  ( ! [X5] : (s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),X5) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK14))))) ) <=> ~$spl580),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl580])])).
fof(f3142,plain,(
  ( ! [X5] : (s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),X5) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl39),
  inference(subsumption_resolution,[],[f3133,f860])).
fof(f3133,plain,(
  ( ! [X5] : (s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),X5) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X5))))) ) | $spl39),
  inference(superposition,[],[f1159,f866])).
fof(f20889,plain,(
  ( ! [X5] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK14))))) ) | $spl578),
  inference(cnf_transformation,[],[f20889_D])).
fof(f20889_D,plain,(
  ( ! [X5] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK14))))) ) <=> ~$spl578),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl578])])).
fof(f3261,plain,(
  ( ! [X5] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl37),
  inference(subsumption_resolution,[],[f3249,f860])).
fof(f3249,plain,(
  ( ! [X5] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X5))))) ) | $spl37),
  inference(superposition,[],[f1156,f866])).
fof(f20888,plain,(
  ( ! [X12,X11] : (p(s(bool,happ(s(fun(sK13,bool),X11),s(sK13,X12)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X11))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X11),s(fun(sK13,bool),sK14))))) ) | $spl576),
  inference(cnf_transformation,[],[f20888_D])).
fof(f20888_D,plain,(
  ( ! [X12,X11] : (p(s(bool,happ(s(fun(sK13,bool),X11),s(sK13,X12)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X11))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X11),s(fun(sK13,bool),sK14))))) ) <=> ~$spl576),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl576])])).
fof(f10899,plain,(
  ( ! [X12,X11] : (p(s(bool,happ(s(fun(sK13,bool),X11),s(sK13,X12)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X11),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X11))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl204),
  inference(subsumption_resolution,[],[f10870,f860])).
fof(f10870,plain,(
  ( ! [X12,X11] : (p(s(bool,happ(s(fun(sK13,bool),X11),s(sK13,X12)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X11),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X11))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X11))))) ) | $spl204),
  inference(superposition,[],[f5495,f866])).
fof(f20887,plain,(
  ( ! [X26,X24,X23,X25] : (sP0(X24,happ(s(fun(sK13,bool),X23),s(sK13,X25)),X26) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14))))) ) | $spl574),
  inference(cnf_transformation,[],[f20887_D])).
fof(f20887_D,plain,(
  ( ! [X26,X24,X23,X25] : (sP0(X24,happ(s(fun(sK13,bool),X23),s(sK13,X25)),X26) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14))))) ) <=> ~$spl574),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl574])])).
fof(f11263,plain,(
  ( ! [X26,X24,X23,X25] : (sP0(X24,happ(s(fun(sK13,bool),X23),s(sK13,X25)),X26) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl204),
  inference(subsumption_resolution,[],[f11250,f860])).
fof(f11250,plain,(
  ( ! [X26,X24,X23,X25] : (sP0(X24,happ(s(fun(sK13,bool),X23),s(sK13,X25)),X26) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X23))))) ) | $spl204),
  inference(superposition,[],[f10769,f866])).
fof(f10769,plain,(
  ( ! [X99,X100,X98] : (sP0(X98,happ(s(fun(sK13,bool),sK14),s(sK13,X99)),X100)) ) | $spl204),
  inference(resolution,[],[f5495,f640])).
fof(f640,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f487])).
fof(f20886,plain,(
  ( ! [X26,X24,X23,X25] : (sP2(X24,X25,happ(s(fun(sK13,bool),X23),s(sK13,X26))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14))))) ) | $spl572),
  inference(cnf_transformation,[],[f20886_D])).
fof(f20886_D,plain,(
  ( ! [X26,X24,X23,X25] : (sP2(X24,X25,happ(s(fun(sK13,bool),X23),s(sK13,X26))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14))))) ) <=> ~$spl572),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl572])])).
fof(f11290,plain,(
  ( ! [X26,X24,X23,X25] : (sP2(X24,X25,happ(s(fun(sK13,bool),X23),s(sK13,X26))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl204),
  inference(subsumption_resolution,[],[f11276,f860])).
fof(f11276,plain,(
  ( ! [X26,X24,X23,X25] : (sP2(X24,X25,happ(s(fun(sK13,bool),X23),s(sK13,X26))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X23))))) ) | $spl204),
  inference(superposition,[],[f10782,f866])).
fof(f10782,plain,(
  ( ! [X138,X139,X137] : (sP2(X137,X138,happ(s(fun(sK13,bool),sK14),s(sK13,X139)))) ) | $spl204),
  inference(resolution,[],[f5495,f656])).
fof(f656,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f495])).
fof(f20885,plain,(
  ( ! [X26,X24,X23,X25] : (sP2(X24,happ(s(fun(sK13,bool),X23),s(sK13,X25)),X26) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14))))) ) | $spl570),
  inference(cnf_transformation,[],[f20885_D])).
fof(f20885_D,plain,(
  ( ! [X26,X24,X23,X25] : (sP2(X24,happ(s(fun(sK13,bool),X23),s(sK13,X25)),X26) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14))))) ) <=> ~$spl570),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl570])])).
fof(f11316,plain,(
  ( ! [X26,X24,X23,X25] : (sP2(X24,happ(s(fun(sK13,bool),X23),s(sK13,X25)),X26) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl204),
  inference(subsumption_resolution,[],[f11303,f860])).
fof(f11303,plain,(
  ( ! [X26,X24,X23,X25] : (sP2(X24,happ(s(fun(sK13,bool),X23),s(sK13,X25)),X26) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X23))))) ) | $spl204),
  inference(superposition,[],[f10783,f866])).
fof(f10783,plain,(
  ( ! [X142,X140,X141] : (sP2(X140,happ(s(fun(sK13,bool),sK14),s(sK13,X141)),X142)) ) | $spl204),
  inference(resolution,[],[f5495,f657])).
fof(f657,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f495])).
fof(f20884,plain,(
  ( ! [X26,X24,X23,X25] : (sP4(X24,X25,happ(s(fun(sK13,bool),X23),s(sK13,X26))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14))))) ) | $spl568),
  inference(cnf_transformation,[],[f20884_D])).
fof(f20884_D,plain,(
  ( ! [X26,X24,X23,X25] : (sP4(X24,X25,happ(s(fun(sK13,bool),X23),s(sK13,X26))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14))))) ) <=> ~$spl568),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl568])])).
fof(f11380,plain,(
  ( ! [X26,X24,X23,X25] : (sP4(X24,X25,happ(s(fun(sK13,bool),X23),s(sK13,X26))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl204),
  inference(subsumption_resolution,[],[f11366,f860])).
fof(f11366,plain,(
  ( ! [X26,X24,X23,X25] : (sP4(X24,X25,happ(s(fun(sK13,bool),X23),s(sK13,X26))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X23))))) ) | $spl204),
  inference(superposition,[],[f10799,f866])).
fof(f10799,plain,(
  ( ! [X189,X190,X188] : (sP4(X188,X189,happ(s(fun(sK13,bool),sK14),s(sK13,X190)))) ) | $spl204),
  inference(resolution,[],[f5495,f673])).
fof(f673,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f502])).
fof(f20883,plain,(
  ( ! [X26,X24,X23,X25] : (sP8(X24,X25,happ(s(fun(sK13,bool),X23),s(sK13,X26))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14))))) ) | $spl566),
  inference(cnf_transformation,[],[f20883_D])).
fof(f20883_D,plain,(
  ( ! [X26,X24,X23,X25] : (sP8(X24,X25,happ(s(fun(sK13,bool),X23),s(sK13,X26))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14))))) ) <=> ~$spl566),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl566])])).
fof(f11406,plain,(
  ( ! [X26,X24,X23,X25] : (sP8(X24,X25,happ(s(fun(sK13,bool),X23),s(sK13,X26))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl204),
  inference(subsumption_resolution,[],[f11392,f860])).
fof(f11392,plain,(
  ( ! [X26,X24,X23,X25] : (sP8(X24,X25,happ(s(fun(sK13,bool),X23),s(sK13,X26))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X23))))) ) | $spl204),
  inference(superposition,[],[f10812,f866])).
fof(f10812,plain,(
  ( ! [X229,X227,X228] : (sP8(X227,X228,happ(s(fun(sK13,bool),sK14),s(sK13,X229)))) ) | $spl204),
  inference(resolution,[],[f5495,f690])).
fof(f690,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f510])).
fof(f20882,plain,(
  ( ! [X26,X24,X23,X25] : (sP7(X24,X25,happ(s(fun(sK13,bool),X23),s(sK13,X26))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14))))) ) | $spl564),
  inference(cnf_transformation,[],[f20882_D])).
fof(f20882_D,plain,(
  ( ! [X26,X24,X23,X25] : (sP7(X24,X25,happ(s(fun(sK13,bool),X23),s(sK13,X26))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14))))) ) <=> ~$spl564),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl564])])).
fof(f11432,plain,(
  ( ! [X26,X24,X23,X25] : (sP7(X24,X25,happ(s(fun(sK13,bool),X23),s(sK13,X26))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl204),
  inference(subsumption_resolution,[],[f11418,f860])).
fof(f11418,plain,(
  ( ! [X26,X24,X23,X25] : (sP7(X24,X25,happ(s(fun(sK13,bool),X23),s(sK13,X26))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X23))))) ) | $spl204),
  inference(superposition,[],[f10815,f866])).
fof(f10815,plain,(
  ( ! [X237,X238,X236] : (sP7(X236,X237,happ(s(fun(sK13,bool),sK14),s(sK13,X238)))) ) | $spl204),
  inference(resolution,[],[f5495,f694])).
fof(f694,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f513])).
fof(f20881,plain,(
  ( ! [X26,X24,X23,X25] : (sP6(X24,X25,happ(s(fun(sK13,bool),X23),s(sK13,X26))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14))))) ) | $spl562),
  inference(cnf_transformation,[],[f20881_D])).
fof(f20881_D,plain,(
  ( ! [X26,X24,X23,X25] : (sP6(X24,X25,happ(s(fun(sK13,bool),X23),s(sK13,X26))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14))))) ) <=> ~$spl562),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl562])])).
fof(f11459,plain,(
  ( ! [X26,X24,X23,X25] : (sP6(X24,X25,happ(s(fun(sK13,bool),X23),s(sK13,X26))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl204),
  inference(subsumption_resolution,[],[f11444,f860])).
fof(f11444,plain,(
  ( ! [X26,X24,X23,X25] : (sP6(X24,X25,happ(s(fun(sK13,bool),X23),s(sK13,X26))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X23))))) ) | $spl204),
  inference(superposition,[],[f10816,f866])).
fof(f10816,plain,(
  ( ! [X241,X239,X240] : (sP6(X239,X240,happ(s(fun(sK13,bool),sK14),s(sK13,X241)))) ) | $spl204),
  inference(resolution,[],[f5495,f698])).
fof(f698,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f515])).
fof(f515,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(flattening,[],[f514])).
fof(f514,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(nnf_transformation,[],[f467])).
fof(f467,plain,(
  ! [X0,X1,X2] : (sP6(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f20880,plain,(
  ( ! [X26,X24,X23,X25] : (sP6(X24,happ(s(fun(sK13,bool),X23),s(sK13,X25)),X26) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14))))) ) | $spl560),
  inference(cnf_transformation,[],[f20880_D])).
fof(f20880_D,plain,(
  ( ! [X26,X24,X23,X25] : (sP6(X24,happ(s(fun(sK13,bool),X23),s(sK13,X25)),X26) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14))))) ) <=> ~$spl560),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl560])])).
fof(f11486,plain,(
  ( ! [X26,X24,X23,X25] : (sP6(X24,happ(s(fun(sK13,bool),X23),s(sK13,X25)),X26) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl204),
  inference(subsumption_resolution,[],[f11471,f860])).
fof(f11471,plain,(
  ( ! [X26,X24,X23,X25] : (sP6(X24,happ(s(fun(sK13,bool),X23),s(sK13,X25)),X26) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X23))))) ) | $spl204),
  inference(superposition,[],[f10817,f866])).
fof(f10817,plain,(
  ( ! [X243,X244,X242] : (sP6(X242,happ(s(fun(sK13,bool),sK14),s(sK13,X243)),X244)) ) | $spl204),
  inference(resolution,[],[f5495,f699])).
fof(f699,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f515])).
fof(f20879,plain,(
  ( ! [X26,X24,X23,X25] : (sP6(happ(s(fun(sK13,bool),X23),s(sK13,X24)),X25,X26) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14))))) ) | $spl558),
  inference(cnf_transformation,[],[f20879_D])).
fof(f20879_D,plain,(
  ( ! [X26,X24,X23,X25] : (sP6(happ(s(fun(sK13,bool),X23),s(sK13,X24)),X25,X26) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14))))) ) <=> ~$spl558),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl558])])).
fof(f11547,plain,(
  ( ! [X26,X24,X23,X25] : (sP6(happ(s(fun(sK13,bool),X23),s(sK13,X24)),X25,X26) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl204),
  inference(subsumption_resolution,[],[f11532,f860])).
fof(f11532,plain,(
  ( ! [X26,X24,X23,X25] : (sP6(happ(s(fun(sK13,bool),X23),s(sK13,X24)),X25,X26) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X23))))) ) | $spl204),
  inference(superposition,[],[f10818,f866])).
fof(f10818,plain,(
  ( ! [X246,X247,X245] : (sP6(happ(s(fun(sK13,bool),sK14),s(sK13,X245)),X246,X247)) ) | $spl204),
  inference(resolution,[],[f5495,f700])).
fof(f700,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f515])).
fof(f20878,plain,(
  ( ! [X26,X24,X23,X25] : (sP10(happ(s(fun(sK13,bool),X23),s(sK13,X24)),X25,X26) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14))))) ) | $spl556),
  inference(cnf_transformation,[],[f20878_D])).
fof(f20878_D,plain,(
  ( ! [X26,X24,X23,X25] : (sP10(happ(s(fun(sK13,bool),X23),s(sK13,X24)),X25,X26) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14))))) ) <=> ~$spl556),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl556])])).
fof(f11571,plain,(
  ( ! [X26,X24,X23,X25] : (sP10(happ(s(fun(sK13,bool),X23),s(sK13,X24)),X25,X26) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl204),
  inference(subsumption_resolution,[],[f11559,f860])).
fof(f11559,plain,(
  ( ! [X26,X24,X23,X25] : (sP10(happ(s(fun(sK13,bool),X23),s(sK13,X24)),X25,X26) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X23))))) ) | $spl204),
  inference(superposition,[],[f10829,f866])).
fof(f10829,plain,(
  ( ! [X280,X279,X278] : (sP10(happ(s(fun(sK13,bool),sK14),s(sK13,X278)),X279,X280)) ) | $spl204),
  inference(resolution,[],[f5495,f718])).
fof(f718,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X2,X1) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f523])).
fof(f523,plain,(
  ! [X0,X1,X2] : ((((~p(s(bool,X1)) | ~p(s(bool,X2))) & ~p(s(bool,X0))) | sP10(X0,X2,X1)) & (~sP10(X0,X2,X1) | (p(s(bool,X1)) & p(s(bool,X2))) | p(s(bool,X0))))),
  inference(flattening,[],[f522])).
fof(f522,plain,(
  ! [X0,X1,X2] : ((((~p(s(bool,X1)) | ~p(s(bool,X2))) & ~p(s(bool,X0))) | sP10(X0,X2,X1)) & (~sP10(X0,X2,X1) | ((p(s(bool,X1)) & p(s(bool,X2))) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f473])).
fof(f473,plain,(
  ! [X0,X1,X2] : (((p(s(bool,X1)) & p(s(bool,X2))) | p(s(bool,X0))) <=> sP10(X0,X2,X1))),
  inference(definition_folding,[],[f274,f472])).
fof(f472,plain,(
  ! [X0,X2,X1] : (sP10(X0,X2,X1) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | p(s(bool,X0)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f274,plain,(
  ! [X0,X1,X2] : (((p(s(bool,X1)) & p(s(bool,X2))) | p(s(bool,X0))) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | p(s(bool,X0)))))),
  inference(rectify,[],[f121])).
fof(f121,axiom,(
  ! [X6,X15,X49] : (((p(s(bool,X15)) & p(s(bool,X49))) | p(s(bool,X6))) <=> ((p(s(bool,X15)) | p(s(bool,X6))) & (p(s(bool,X49)) | p(s(bool,X6)))))),
  file('thm.pred_set.SUBSET_DIFF_EMPTY','thm.bool.RIGHT_OR_OVER_AND')).
fof(f20877,plain,(
  ( ! [X26,X24,X23,X25] : (sP11(X24,happ(s(fun(sK13,bool),X23),s(sK13,X25)),X26) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14))))) ) | $spl554),
  inference(cnf_transformation,[],[f20877_D])).
fof(f20877_D,plain,(
  ( ! [X26,X24,X23,X25] : (sP11(X24,happ(s(fun(sK13,bool),X23),s(sK13,X25)),X26) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14))))) ) <=> ~$spl554),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl554])])).
fof(f11595,plain,(
  ( ! [X26,X24,X23,X25] : (sP11(X24,happ(s(fun(sK13,bool),X23),s(sK13,X25)),X26) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl204),
  inference(subsumption_resolution,[],[f11583,f860])).
fof(f11583,plain,(
  ( ! [X26,X24,X23,X25] : (sP11(X24,happ(s(fun(sK13,bool),X23),s(sK13,X25)),X26) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X23))))) ) | $spl204),
  inference(superposition,[],[f10836,f866])).
fof(f10836,plain,(
  ( ! [X300,X301,X299] : (sP11(X299,happ(s(fun(sK13,bool),sK14),s(sK13,X300)),X301)) ) | $spl204),
  inference(resolution,[],[f5495,f727])).
fof(f727,plain,(
  ( ! [X2,X0,X1] : (sP11(X2,X0,X1) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f528])).
fof(f528,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X0)) & (~p(s(bool,X1)) | ~p(s(bool,X2)))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | p(s(bool,X0)) | (p(s(bool,X1)) & p(s(bool,X2)))))),
  inference(flattening,[],[f527])).
fof(f527,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X0)) & (~p(s(bool,X1)) | ~p(s(bool,X2)))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | (p(s(bool,X0)) | (p(s(bool,X1)) & p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f475])).
fof(f475,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X0)) | (p(s(bool,X1)) & p(s(bool,X2)))) <=> sP11(X2,X0,X1))),
  inference(definition_folding,[],[f275,f474])).
fof(f474,plain,(
  ! [X2,X0,X1] : (sP11(X2,X0,X1) <=> ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f275,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X0)) | (p(s(bool,X1)) & p(s(bool,X2)))) <=> ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(rectify,[],[f238])).
fof(f238,axiom,(
  ! [X6,X15,X49] : ((p(s(bool,X6)) | (p(s(bool,X15)) & p(s(bool,X49)))) <=> ((p(s(bool,X6)) | p(s(bool,X15))) & (p(s(bool,X6)) | p(s(bool,X49)))))),
  file('thm.pred_set.SUBSET_DIFF_EMPTY','thm.bool.LEFT_OR_OVER_AND')).
fof(f20876,plain,(
  ( ! [X5] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X5)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK14))))) ) | $spl552),
  inference(cnf_transformation,[],[f20876_D])).
fof(f20876_D,plain,(
  ( ! [X5] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X5)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK14))))) ) <=> ~$spl552),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl552])])).
fof(f11884,plain,(
  ( ! [X5] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X5)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl5),
  inference(subsumption_resolution,[],[f11854,f860])).
fof(f11854,plain,(
  ( ! [X5] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X5)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X5))))) ) | $spl5),
  inference(superposition,[],[f1086,f866])).
fof(f20875,plain,(
  ( ! [X5] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK14))))) ) | $spl550),
  inference(cnf_transformation,[],[f20875_D])).
fof(f20875_D,plain,(
  ( ! [X5] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK14))))) ) <=> ~$spl550),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl550])])).
fof(f12419,plain,(
  ( ! [X5] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl29),
  inference(subsumption_resolution,[],[f12375,f860])).
fof(f12375,plain,(
  ( ! [X5] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X5))))) ) | $spl29),
  inference(superposition,[],[f1105,f866])).
fof(f20874,plain,(
  ( ! [X19,X17,X18] : (sP0(X18,X19,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X17))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) | $spl548),
  inference(cnf_transformation,[],[f20874_D])).
fof(f20874_D,plain,(
  ( ! [X19,X17,X18] : (sP0(X18,X19,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X17))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) <=> ~$spl548),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl548])])).
fof(f12608,plain,(
  ( ! [X19,X17,X18] : (sP0(X18,X19,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl5),
  inference(subsumption_resolution,[],[f12597,f860])).
fof(f12597,plain,(
  ( ! [X19,X17,X18] : (sP0(X18,X19,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X17))))) ) | $spl5),
  inference(superposition,[],[f4418,f866])).
fof(f4418,plain,(
  ( ! [X39,X38] : (sP0(X38,X39,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)))) ) | $spl5),
  inference(resolution,[],[f1086,f639])).
fof(f20873,plain,(
  ( ! [X19,X17,X18] : (sP0('const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X17)),X18,X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) | $spl546),
  inference(cnf_transformation,[],[f20873_D])).
fof(f20873_D,plain,(
  ( ! [X19,X17,X18] : (sP0('const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X17)),X18,X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) <=> ~$spl546),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl546])])).
fof(f12662,plain,(
  ( ! [X19,X17,X18] : (sP0('const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X17)),X18,X19) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl5),
  inference(subsumption_resolution,[],[f12651,f860])).
fof(f12651,plain,(
  ( ! [X19,X17,X18] : (sP0('const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X17)),X18,X19) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X17))))) ) | $spl5),
  inference(superposition,[],[f4419,f866])).
fof(f4419,plain,(
  ( ! [X41,X40] : (sP0('const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X40,X41)) ) | $spl5),
  inference(resolution,[],[f1086,f641])).
fof(f20872,plain,(
  ( ! [X19,X17,X18] : (sP2('const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X17)),X18,X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) | $spl544),
  inference(cnf_transformation,[],[f20872_D])).
fof(f20872_D,plain,(
  ( ! [X19,X17,X18] : (sP2('const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X17)),X18,X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) <=> ~$spl544),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl544])])).
fof(f12699,plain,(
  ( ! [X19,X17,X18] : (sP2('const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X17)),X18,X19) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl5),
  inference(subsumption_resolution,[],[f12687,f860])).
fof(f12687,plain,(
  ( ! [X19,X17,X18] : (sP2('const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X17)),X18,X19) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X17))))) ) | $spl5),
  inference(superposition,[],[f4435,f866])).
fof(f4435,plain,(
  ( ! [X72,X73] : (sP2('const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X72,X73)) ) | $spl5),
  inference(resolution,[],[f1086,f658])).
fof(f20871,plain,(
  ( ! [X19,X17,X18] : (sP4(X18,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X17)),X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) | $spl542),
  inference(cnf_transformation,[],[f20871_D])).
fof(f20871_D,plain,(
  ( ! [X19,X17,X18] : (sP4(X18,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X17)),X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) <=> ~$spl542),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl542])])).
fof(f12735,plain,(
  ( ! [X19,X17,X18] : (sP4(X18,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X17)),X19) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl5),
  inference(subsumption_resolution,[],[f12724,f860])).
fof(f12724,plain,(
  ( ! [X19,X17,X18] : (sP4(X18,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X17)),X19) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X17))))) ) | $spl5),
  inference(superposition,[],[f4448,f866])).
fof(f4448,plain,(
  ( ! [X99,X98] : (sP4(X98,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X99)) ) | $spl5),
  inference(resolution,[],[f1086,f674])).
fof(f20870,plain,(
  ( ! [X19,X17,X18] : (sP4('const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X17)),X18,X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) | $spl540),
  inference(cnf_transformation,[],[f20870_D])).
fof(f20870_D,plain,(
  ( ! [X19,X17,X18] : (sP4('const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X17)),X18,X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) <=> ~$spl540),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl540])])).
fof(f12771,plain,(
  ( ! [X19,X17,X18] : (sP4('const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X17)),X18,X19) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl5),
  inference(subsumption_resolution,[],[f12760,f860])).
fof(f12760,plain,(
  ( ! [X19,X17,X18] : (sP4('const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X17)),X18,X19) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X17))))) ) | $spl5),
  inference(superposition,[],[f4449,f866])).
fof(f4449,plain,(
  ( ! [X101,X100] : (sP4('const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X100,X101)) ) | $spl5),
  inference(resolution,[],[f1086,f675])).
fof(f20869,plain,(
  ( ! [X19,X17,X18] : (sP8(X18,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X17)),X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) | $spl538),
  inference(cnf_transformation,[],[f20869_D])).
fof(f20869_D,plain,(
  ( ! [X19,X17,X18] : (sP8(X18,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X17)),X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) <=> ~$spl538),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl538])])).
fof(f12807,plain,(
  ( ! [X19,X17,X18] : (sP8(X18,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X17)),X19) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl5),
  inference(subsumption_resolution,[],[f12795,f860])).
fof(f12795,plain,(
  ( ! [X19,X17,X18] : (sP8(X18,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X17)),X19) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X17))))) ) | $spl5),
  inference(superposition,[],[f4461,f866])).
fof(f4461,plain,(
  ( ! [X125,X124] : (sP8(X124,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X125)) ) | $spl5),
  inference(resolution,[],[f1086,f691])).
fof(f20868,plain,(
  ( ! [X19,X17,X18] : (sP8('const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X17)),X18,X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) | $spl536),
  inference(cnf_transformation,[],[f20868_D])).
fof(f20868_D,plain,(
  ( ! [X19,X17,X18] : (sP8('const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X17)),X18,X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) <=> ~$spl536),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl536])])).
fof(f12882,plain,(
  ( ! [X19,X17,X18] : (sP8('const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X17)),X18,X19) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl5),
  inference(subsumption_resolution,[],[f12869,f860])).
fof(f12869,plain,(
  ( ! [X19,X17,X18] : (sP8('const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X17)),X18,X19) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X17))))) ) | $spl5),
  inference(superposition,[],[f4462,f866])).
fof(f4462,plain,(
  ( ! [X127,X126] : (sP8('const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X126,X127)) ) | $spl5),
  inference(resolution,[],[f1086,f692])).
fof(f20867,plain,(
  ( ! [X19,X17,X18] : (sP7(X18,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X17)),X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) | $spl534),
  inference(cnf_transformation,[],[f20867_D])).
fof(f20867_D,plain,(
  ( ! [X19,X17,X18] : (sP7(X18,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X17)),X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) <=> ~$spl534),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl534])])).
fof(f12919,plain,(
  ( ! [X19,X17,X18] : (sP7(X18,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X17)),X19) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl5),
  inference(subsumption_resolution,[],[f12906,f860])).
fof(f12906,plain,(
  ( ! [X19,X17,X18] : (sP7(X18,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X17)),X19) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X17))))) ) | $spl5),
  inference(superposition,[],[f4464,f866])).
fof(f4464,plain,(
  ( ! [X130,X131] : (sP7(X130,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X131)) ) | $spl5),
  inference(resolution,[],[f1086,f695])).
fof(f20866,plain,(
  ( ! [X19,X17,X18] : (sP7('const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X17)),X18,X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) | $spl532),
  inference(cnf_transformation,[],[f20866_D])).
fof(f20866_D,plain,(
  ( ! [X19,X17,X18] : (sP7('const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X17)),X18,X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) <=> ~$spl532),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl532])])).
fof(f12956,plain,(
  ( ! [X19,X17,X18] : (sP7('const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X17)),X18,X19) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl5),
  inference(subsumption_resolution,[],[f12943,f860])).
fof(f12943,plain,(
  ( ! [X19,X17,X18] : (sP7('const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X17)),X18,X19) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X17))))) ) | $spl5),
  inference(superposition,[],[f4465,f866])).
fof(f4465,plain,(
  ( ! [X132,X133] : (sP7('const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X132,X133)) ) | $spl5),
  inference(resolution,[],[f1086,f696])).
fof(f20865,plain,(
  ( ! [X19,X17,X18] : (sP0(X18,X19,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) | $spl530),
  inference(cnf_transformation,[],[f20865_D])).
fof(f20865_D,plain,(
  ( ! [X19,X17,X18] : (sP0(X18,X19,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) <=> ~$spl530),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl530])])).
fof(f13518,plain,(
  ( ! [X19,X17,X18] : (sP0(X18,X19,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl29),
  inference(subsumption_resolution,[],[f13496,f860])).
fof(f13496,plain,(
  ( ! [X19,X17,X18] : (sP0(X18,X19,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X17))))) ) | $spl29),
  inference(superposition,[],[f12275,f866])).
fof(f12275,plain,(
  ( ! [X39,X38] : (sP0(X38,X39,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) | $spl29),
  inference(resolution,[],[f1105,f639])).
fof(f20864,plain,(
  ( ! [X19,X17,X18] : (sP0('const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X18,X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) | $spl528),
  inference(cnf_transformation,[],[f20864_D])).
fof(f20864_D,plain,(
  ( ! [X19,X17,X18] : (sP0('const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X18,X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) <=> ~$spl528),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl528])])).
fof(f13563,plain,(
  ( ! [X19,X17,X18] : (sP0('const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X18,X19) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl29),
  inference(subsumption_resolution,[],[f13541,f860])).
fof(f13541,plain,(
  ( ! [X19,X17,X18] : (sP0('const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X18,X19) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X17))))) ) | $spl29),
  inference(superposition,[],[f12276,f866])).
fof(f12276,plain,(
  ( ! [X41,X40] : (sP0('const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X40,X41)) ) | $spl29),
  inference(resolution,[],[f1105,f641])).
fof(f20863,plain,(
  ( ! [X19,X17,X18] : (sP2('const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X18,X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) | $spl526),
  inference(cnf_transformation,[],[f20863_D])).
fof(f20863_D,plain,(
  ( ! [X19,X17,X18] : (sP2('const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X18,X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) <=> ~$spl526),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl526])])).
fof(f13609,plain,(
  ( ! [X19,X17,X18] : (sP2('const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X18,X19) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl29),
  inference(subsumption_resolution,[],[f13586,f860])).
fof(f13586,plain,(
  ( ! [X19,X17,X18] : (sP2('const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X18,X19) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X17))))) ) | $spl29),
  inference(superposition,[],[f12292,f866])).
fof(f12292,plain,(
  ( ! [X72,X73] : (sP2('const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X72,X73)) ) | $spl29),
  inference(resolution,[],[f1105,f658])).
fof(f20862,plain,(
  ( ! [X19,X17,X18] : (sP4(X18,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) | $spl524),
  inference(cnf_transformation,[],[f20862_D])).
fof(f20862_D,plain,(
  ( ! [X19,X17,X18] : (sP4(X18,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) <=> ~$spl524),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl524])])).
fof(f13654,plain,(
  ( ! [X19,X17,X18] : (sP4(X18,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X19) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl29),
  inference(subsumption_resolution,[],[f13632,f860])).
fof(f13632,plain,(
  ( ! [X19,X17,X18] : (sP4(X18,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X19) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X17))))) ) | $spl29),
  inference(superposition,[],[f12305,f866])).
fof(f12305,plain,(
  ( ! [X99,X98] : (sP4(X98,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X99)) ) | $spl29),
  inference(resolution,[],[f1105,f674])).
fof(f20861,plain,(
  ( ! [X19,X17,X18] : (sP4('const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X18,X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) | $spl522),
  inference(cnf_transformation,[],[f20861_D])).
fof(f20861_D,plain,(
  ( ! [X19,X17,X18] : (sP4('const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X18,X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) <=> ~$spl522),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl522])])).
fof(f13699,plain,(
  ( ! [X19,X17,X18] : (sP4('const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X18,X19) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl29),
  inference(subsumption_resolution,[],[f13677,f860])).
fof(f13677,plain,(
  ( ! [X19,X17,X18] : (sP4('const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X18,X19) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X17))))) ) | $spl29),
  inference(superposition,[],[f12306,f866])).
fof(f12306,plain,(
  ( ! [X101,X100] : (sP4('const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X100,X101)) ) | $spl29),
  inference(resolution,[],[f1105,f675])).
fof(f20860,plain,(
  ( ! [X19,X17,X18] : (sP8(X18,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) | $spl520),
  inference(cnf_transformation,[],[f20860_D])).
fof(f20860_D,plain,(
  ( ! [X19,X17,X18] : (sP8(X18,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) <=> ~$spl520),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl520])])).
fof(f13744,plain,(
  ( ! [X19,X17,X18] : (sP8(X18,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X19) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl29),
  inference(subsumption_resolution,[],[f13721,f860])).
fof(f13721,plain,(
  ( ! [X19,X17,X18] : (sP8(X18,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X19) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X17))))) ) | $spl29),
  inference(superposition,[],[f12318,f866])).
fof(f12318,plain,(
  ( ! [X125,X124] : (sP8(X124,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X125)) ) | $spl29),
  inference(resolution,[],[f1105,f691])).
fof(f20859,plain,(
  ( ! [X19,X17,X18] : (sP8('const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X18,X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) | $spl518),
  inference(cnf_transformation,[],[f20859_D])).
fof(f20859_D,plain,(
  ( ! [X19,X17,X18] : (sP8('const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X18,X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) <=> ~$spl518),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl518])])).
fof(f13790,plain,(
  ( ! [X19,X17,X18] : (sP8('const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X18,X19) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl29),
  inference(subsumption_resolution,[],[f13766,f860])).
fof(f13766,plain,(
  ( ! [X19,X17,X18] : (sP8('const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X18,X19) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X17))))) ) | $spl29),
  inference(superposition,[],[f12319,f866])).
fof(f12319,plain,(
  ( ! [X127,X126] : (sP8('const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X126,X127)) ) | $spl29),
  inference(resolution,[],[f1105,f692])).
fof(f20858,plain,(
  ( ! [X19,X17,X18] : (sP7(X18,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) | $spl516),
  inference(cnf_transformation,[],[f20858_D])).
fof(f20858_D,plain,(
  ( ! [X19,X17,X18] : (sP7(X18,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) <=> ~$spl516),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl516])])).
fof(f13836,plain,(
  ( ! [X19,X17,X18] : (sP7(X18,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X19) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl29),
  inference(subsumption_resolution,[],[f13812,f860])).
fof(f13812,plain,(
  ( ! [X19,X17,X18] : (sP7(X18,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X19) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X17))))) ) | $spl29),
  inference(superposition,[],[f12321,f866])).
fof(f12321,plain,(
  ( ! [X130,X131] : (sP7(X130,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X131)) ) | $spl29),
  inference(resolution,[],[f1105,f695])).
fof(f20857,plain,(
  ( ! [X19,X17,X18] : (sP7('const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X18,X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) | $spl514),
  inference(cnf_transformation,[],[f20857_D])).
fof(f20857_D,plain,(
  ( ! [X19,X17,X18] : (sP7('const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X18,X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) <=> ~$spl514),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl514])])).
fof(f13882,plain,(
  ( ! [X19,X17,X18] : (sP7('const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X18,X19) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl29),
  inference(subsumption_resolution,[],[f13858,f860])).
fof(f13858,plain,(
  ( ! [X19,X17,X18] : (sP7('const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X18,X19) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X17))))) ) | $spl29),
  inference(superposition,[],[f12322,f866])).
fof(f12322,plain,(
  ( ! [X132,X133] : (sP7('const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X132,X133)) ) | $spl29),
  inference(resolution,[],[f1105,f696])).
fof(f20856,plain,(
  ( ! [X5] : (~p(s(bool,'const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),X5)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK14))))) ) | $spl512),
  inference(cnf_transformation,[],[f20856_D])).
fof(f20856_D,plain,(
  ( ! [X5] : (~p(s(bool,'const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),X5)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK14))))) ) <=> ~$spl512),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl512])])).
fof(f14197,plain,(
  ( ! [X5] : (~p(s(bool,'const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),X5)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl5),
  inference(subsumption_resolution,[],[f14168,f860])).
fof(f14168,plain,(
  ( ! [X5] : (~p(s(bool,'const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),X5)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X5))))) ) | $spl5),
  inference(superposition,[],[f11723,f866])).
fof(f11723,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),sK14)))) | $spl5),
  inference(resolution,[],[f1086,f851])).
fof(f20855,plain,(
  ( ! [X5] : (p(s(bool,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),X5)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK14))))) ) | $spl510),
  inference(cnf_transformation,[],[f20855_D])).
fof(f20855_D,plain,(
  ( ! [X5] : (p(s(bool,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),X5)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK14))))) ) <=> ~$spl510),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl510])])).
fof(f14503,plain,(
  ( ! [X5] : (p(s(bool,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),X5)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl29),
  inference(subsumption_resolution,[],[f14475,f860])).
fof(f14475,plain,(
  ( ! [X5] : (p(s(bool,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),X5)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X5))))) ) | $spl29),
  inference(superposition,[],[f12255,f866])).
fof(f20854,plain,(
  ( ! [X19,X17,X18] : (~sP12(happ(s(fun(sK13,bool),X17),s(sK13,X18)),X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | ~p(s(bool,X19))) ) | $spl508),
  inference(cnf_transformation,[],[f20854_D])).
fof(f20854_D,plain,(
  ( ! [X19,X17,X18] : (~sP12(happ(s(fun(sK13,bool),X17),s(sK13,X18)),X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | ~p(s(bool,X19))) ) <=> ~$spl508),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl508])])).
fof(f14531,plain,(
  ( ! [X19,X17,X18] : (~sP12(happ(s(fun(sK13,bool),X17),s(sK13,X18)),X19) | ~p(s(bool,X19)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl204),
  inference(subsumption_resolution,[],[f14518,f860])).
fof(f14518,plain,(
  ( ! [X19,X17,X18] : (~sP12(happ(s(fun(sK13,bool),X17),s(sK13,X18)),X19) | ~p(s(bool,X19)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X17))))) ) | $spl204),
  inference(superposition,[],[f10839,f866])).
fof(f10839,plain,(
  ( ! [X308,X309] : (~sP12(happ(s(fun(sK13,bool),sK14),s(sK13,X309)),X308) | ~p(s(bool,X308))) ) | $spl204),
  inference(resolution,[],[f5495,f752])).
fof(f752,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | ~sP12(X0,X1)) )),
  inference(cnf_transformation,[],[f539])).
fof(f539,plain,(
  ! [X0,X1] : ((~sP12(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP12(X0,X1)))),
  inference(rectify,[],[f538])).
fof(f538,plain,(
  ! [X1,X0] : ((~sP12(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP12(X1,X0)))),
  inference(flattening,[],[f537])).
fof(f537,plain,(
  ! [X1,X0] : ((~sP12(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP12(X1,X0)))),
  inference(nnf_transformation,[],[f476])).
fof(f476,plain,(
  ! [X1,X0] : (sP12(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP12])])).
fof(f20853,plain,(
  ( ! [X19,X17,X18] : (~sP12(X18,happ(s(fun(sK13,bool),X17),s(sK13,X19))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | ~p(s(bool,X18))) ) | $spl506),
  inference(cnf_transformation,[],[f20853_D])).
fof(f20853_D,plain,(
  ( ! [X19,X17,X18] : (~sP12(X18,happ(s(fun(sK13,bool),X17),s(sK13,X19))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | ~p(s(bool,X18))) ) <=> ~$spl506),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl506])])).
fof(f14559,plain,(
  ( ! [X19,X17,X18] : (~sP12(X18,happ(s(fun(sK13,bool),X17),s(sK13,X19))) | ~p(s(bool,X18)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl204),
  inference(subsumption_resolution,[],[f14546,f860])).
fof(f14546,plain,(
  ( ! [X19,X17,X18] : (~sP12(X18,happ(s(fun(sK13,bool),X17),s(sK13,X19))) | ~p(s(bool,X18)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X17))))) ) | $spl204),
  inference(superposition,[],[f10840,f866])).
fof(f10840,plain,(
  ( ! [X310,X311] : (~sP12(X310,happ(s(fun(sK13,bool),sK14),s(sK13,X311))) | ~p(s(bool,X310))) ) | $spl204),
  inference(resolution,[],[f5495,f752])).
fof(f20852,plain,(
  ( ! [X19,X17,X18] : (sP12(happ(s(fun(sK13,bool),X17),s(sK13,X18)),X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | p(s(bool,X19))) ) | $spl504),
  inference(cnf_transformation,[],[f20852_D])).
fof(f20852_D,plain,(
  ( ! [X19,X17,X18] : (sP12(happ(s(fun(sK13,bool),X17),s(sK13,X18)),X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | p(s(bool,X19))) ) <=> ~$spl504),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl504])])).
fof(f14589,plain,(
  ( ! [X19,X17,X18] : (sP12(happ(s(fun(sK13,bool),X17),s(sK13,X18)),X19) | p(s(bool,X19)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl204),
  inference(subsumption_resolution,[],[f14576,f860])).
fof(f14576,plain,(
  ( ! [X19,X17,X18] : (sP12(happ(s(fun(sK13,bool),X17),s(sK13,X18)),X19) | p(s(bool,X19)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X17))))) ) | $spl204),
  inference(superposition,[],[f10841,f866])).
fof(f10841,plain,(
  ( ! [X313,X312] : (sP12(happ(s(fun(sK13,bool),sK14),s(sK13,X312)),X313) | p(s(bool,X313))) ) | $spl204),
  inference(resolution,[],[f5495,f753])).
fof(f753,plain,(
  ( ! [X0,X1] : (sP12(X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f539])).
fof(f20851,plain,(
  ( ! [X19,X17,X18] : (sP12(X18,happ(s(fun(sK13,bool),X17),s(sK13,X19))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | p(s(bool,X18))) ) | $spl502),
  inference(cnf_transformation,[],[f20851_D])).
fof(f20851_D,plain,(
  ( ! [X19,X17,X18] : (sP12(X18,happ(s(fun(sK13,bool),X17),s(sK13,X19))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | p(s(bool,X18))) ) <=> ~$spl502),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl502])])).
fof(f14619,plain,(
  ( ! [X19,X17,X18] : (sP12(X18,happ(s(fun(sK13,bool),X17),s(sK13,X19))) | p(s(bool,X18)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl204),
  inference(subsumption_resolution,[],[f14606,f860])).
fof(f14606,plain,(
  ( ! [X19,X17,X18] : (sP12(X18,happ(s(fun(sK13,bool),X17),s(sK13,X19))) | p(s(bool,X18)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X17))))) ) | $spl204),
  inference(superposition,[],[f10842,f866])).
fof(f10842,plain,(
  ( ! [X315,X314] : (sP12(X314,happ(s(fun(sK13,bool),sK14),s(sK13,X315))) | p(s(bool,X314))) ) | $spl204),
  inference(resolution,[],[f5495,f756])).
fof(f756,plain,(
  ( ! [X0,X1] : (sP12(X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f539])).
fof(f20850,plain,(
  ( ! [X19,X17,X18] : (sP0(X18,X19,'const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),X17))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) | $spl500),
  inference(cnf_transformation,[],[f20850_D])).
fof(f20850_D,plain,(
  ( ! [X19,X17,X18] : (sP0(X18,X19,'const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),X17))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) <=> ~$spl500),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl500])])).
fof(f14895,plain,(
  ( ! [X19,X17,X18] : (sP0(X18,X19,'const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl5),
  inference(subsumption_resolution,[],[f14887,f860])).
fof(f14887,plain,(
  ( ! [X19,X17,X18] : (sP0(X18,X19,'const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X17))))) ) | $spl5),
  inference(superposition,[],[f14066,f866])).
fof(f14066,plain,(
  ( ! [X41,X42] : (sP0(X41,X42,'const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),sK14)))) ) | $spl5),
  inference(resolution,[],[f11723,f639])).
fof(f20849,plain,(
  ( ! [X19,X17,X18] : (sP0('const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),X17)),X18,X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) | $spl498),
  inference(cnf_transformation,[],[f20849_D])).
fof(f20849_D,plain,(
  ( ! [X19,X17,X18] : (sP0('const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),X17)),X18,X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) <=> ~$spl498),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl498])])).
fof(f14918,plain,(
  ( ! [X19,X17,X18] : (sP0('const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),X17)),X18,X19) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl5),
  inference(subsumption_resolution,[],[f14910,f860])).
fof(f14910,plain,(
  ( ! [X19,X17,X18] : (sP0('const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),X17)),X18,X19) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X17))))) ) | $spl5),
  inference(superposition,[],[f14067,f866])).
fof(f14067,plain,(
  ( ! [X43,X44] : (sP0('const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),sK14)),X43,X44)) ) | $spl5),
  inference(resolution,[],[f11723,f641])).
fof(f20848,plain,(
  ( ! [X19,X17,X18] : (sP2('const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),X17)),X18,X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) | $spl496),
  inference(cnf_transformation,[],[f20848_D])).
fof(f20848_D,plain,(
  ( ! [X19,X17,X18] : (sP2('const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),X17)),X18,X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) <=> ~$spl496),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl496])])).
fof(f14985,plain,(
  ( ! [X19,X17,X18] : (sP2('const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),X17)),X18,X19) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl5),
  inference(subsumption_resolution,[],[f14976,f860])).
fof(f14976,plain,(
  ( ! [X19,X17,X18] : (sP2('const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),X17)),X18,X19) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X17))))) ) | $spl5),
  inference(superposition,[],[f14083,f866])).
fof(f14083,plain,(
  ( ! [X76,X75] : (sP2('const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),sK14)),X75,X76)) ) | $spl5),
  inference(resolution,[],[f11723,f658])).
fof(f20847,plain,(
  ( ! [X19,X17,X18] : (sP4(X18,'const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),X17)),X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) | $spl494),
  inference(cnf_transformation,[],[f20847_D])).
fof(f20847_D,plain,(
  ( ! [X19,X17,X18] : (sP4(X18,'const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),X17)),X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) <=> ~$spl494),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl494])])).
fof(f15008,plain,(
  ( ! [X19,X17,X18] : (sP4(X18,'const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),X17)),X19) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl5),
  inference(subsumption_resolution,[],[f15000,f860])).
fof(f15000,plain,(
  ( ! [X19,X17,X18] : (sP4(X18,'const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),X17)),X19) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X17))))) ) | $spl5),
  inference(superposition,[],[f14096,f866])).
fof(f14096,plain,(
  ( ! [X101,X102] : (sP4(X101,'const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),sK14)),X102)) ) | $spl5),
  inference(resolution,[],[f11723,f674])).
fof(f20846,plain,(
  ( ! [X19,X17,X18] : (sP4('const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),X17)),X18,X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) | $spl492),
  inference(cnf_transformation,[],[f20846_D])).
fof(f20846_D,plain,(
  ( ! [X19,X17,X18] : (sP4('const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),X17)),X18,X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) <=> ~$spl492),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl492])])).
fof(f15031,plain,(
  ( ! [X19,X17,X18] : (sP4('const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),X17)),X18,X19) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl5),
  inference(subsumption_resolution,[],[f15023,f860])).
fof(f15023,plain,(
  ( ! [X19,X17,X18] : (sP4('const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),X17)),X18,X19) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X17))))) ) | $spl5),
  inference(superposition,[],[f14097,f866])).
fof(f14097,plain,(
  ( ! [X103,X104] : (sP4('const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),sK14)),X103,X104)) ) | $spl5),
  inference(resolution,[],[f11723,f675])).
fof(f20845,plain,(
  ( ! [X19,X17,X18] : (sP8(X18,'const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),X17)),X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) | $spl490),
  inference(cnf_transformation,[],[f20845_D])).
fof(f20845_D,plain,(
  ( ! [X19,X17,X18] : (sP8(X18,'const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),X17)),X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) <=> ~$spl490),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl490])])).
fof(f15054,plain,(
  ( ! [X19,X17,X18] : (sP8(X18,'const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),X17)),X19) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl5),
  inference(subsumption_resolution,[],[f15045,f860])).
fof(f15045,plain,(
  ( ! [X19,X17,X18] : (sP8(X18,'const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),X17)),X19) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X17))))) ) | $spl5),
  inference(superposition,[],[f14109,f866])).
fof(f14109,plain,(
  ( ! [X127,X128] : (sP8(X127,'const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),sK14)),X128)) ) | $spl5),
  inference(resolution,[],[f11723,f691])).
fof(f20844,plain,(
  ( ! [X19,X17,X18] : (sP8('const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),X17)),X18,X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) | $spl488),
  inference(cnf_transformation,[],[f20844_D])).
fof(f20844_D,plain,(
  ( ! [X19,X17,X18] : (sP8('const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),X17)),X18,X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) <=> ~$spl488),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl488])])).
fof(f15078,plain,(
  ( ! [X19,X17,X18] : (sP8('const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),X17)),X18,X19) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl5),
  inference(subsumption_resolution,[],[f15068,f860])).
fof(f15068,plain,(
  ( ! [X19,X17,X18] : (sP8('const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),X17)),X18,X19) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X17))))) ) | $spl5),
  inference(superposition,[],[f14110,f866])).
fof(f14110,plain,(
  ( ! [X130,X129] : (sP8('const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),sK14)),X129,X130)) ) | $spl5),
  inference(resolution,[],[f11723,f692])).
fof(f20843,plain,(
  ( ! [X19,X17,X18] : (sP7(X18,'const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),X17)),X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) | $spl486),
  inference(cnf_transformation,[],[f20843_D])).
fof(f20843_D,plain,(
  ( ! [X19,X17,X18] : (sP7(X18,'const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),X17)),X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) <=> ~$spl486),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl486])])).
fof(f15143,plain,(
  ( ! [X19,X17,X18] : (sP7(X18,'const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),X17)),X19) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl5),
  inference(subsumption_resolution,[],[f15133,f860])).
fof(f15133,plain,(
  ( ! [X19,X17,X18] : (sP7(X18,'const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),X17)),X19) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X17))))) ) | $spl5),
  inference(superposition,[],[f14112,f866])).
fof(f14112,plain,(
  ( ! [X134,X133] : (sP7(X133,'const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),sK14)),X134)) ) | $spl5),
  inference(resolution,[],[f11723,f695])).
fof(f20842,plain,(
  ( ! [X19,X17,X18] : (sP7('const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),X17)),X18,X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) | $spl484),
  inference(cnf_transformation,[],[f20842_D])).
fof(f20842_D,plain,(
  ( ! [X19,X17,X18] : (sP7('const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),X17)),X18,X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) <=> ~$spl484),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl484])])).
fof(f15167,plain,(
  ( ! [X19,X17,X18] : (sP7('const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),X17)),X18,X19) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl5),
  inference(subsumption_resolution,[],[f15157,f860])).
fof(f15157,plain,(
  ( ! [X19,X17,X18] : (sP7('const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),X17)),X18,X19) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X17))))) ) | $spl5),
  inference(superposition,[],[f14113,f866])).
fof(f14113,plain,(
  ( ! [X136,X135] : (sP7('const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),sK14)),X135,X136)) ) | $spl5),
  inference(resolution,[],[f11723,f696])).
fof(f20841,plain,(
  ( ! [X19,X17,X18] : (sP0(X18,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),X17)),X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) | $spl482),
  inference(cnf_transformation,[],[f20841_D])).
fof(f20841_D,plain,(
  ( ! [X19,X17,X18] : (sP0(X18,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),X17)),X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) <=> ~$spl482),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl482])])).
fof(f15522,plain,(
  ( ! [X19,X17,X18] : (sP0(X18,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),X17)),X19) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl29),
  inference(subsumption_resolution,[],[f15514,f860])).
fof(f15514,plain,(
  ( ! [X19,X17,X18] : (sP0(X18,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),X17)),X19) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X17))))) ) | $spl29),
  inference(superposition,[],[f14372,f866])).
fof(f14372,plain,(
  ( ! [X43,X44] : (sP0(X43,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),sK14)),X44)) ) | $spl29),
  inference(resolution,[],[f12255,f640])).
fof(f20840,plain,(
  ( ! [X19,X17,X18] : (sP2(X18,X19,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),X17))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) | $spl480),
  inference(cnf_transformation,[],[f20840_D])).
fof(f20840_D,plain,(
  ( ! [X19,X17,X18] : (sP2(X18,X19,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),X17))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) <=> ~$spl480),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl480])])).
fof(f15545,plain,(
  ( ! [X19,X17,X18] : (sP2(X18,X19,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl29),
  inference(subsumption_resolution,[],[f15537,f860])).
fof(f15537,plain,(
  ( ! [X19,X17,X18] : (sP2(X18,X19,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X17))))) ) | $spl29),
  inference(superposition,[],[f14385,f866])).
fof(f14385,plain,(
  ( ! [X70,X69] : (sP2(X69,X70,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),sK14)))) ) | $spl29),
  inference(resolution,[],[f12255,f656])).
fof(f20839,plain,(
  ( ! [X19,X17,X18] : (sP2(X18,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),X17)),X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) | $spl478),
  inference(cnf_transformation,[],[f20839_D])).
fof(f20839_D,plain,(
  ( ! [X19,X17,X18] : (sP2(X18,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),X17)),X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) <=> ~$spl478),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl478])])).
fof(f15568,plain,(
  ( ! [X19,X17,X18] : (sP2(X18,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),X17)),X19) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl29),
  inference(subsumption_resolution,[],[f15560,f860])).
fof(f15560,plain,(
  ( ! [X19,X17,X18] : (sP2(X18,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),X17)),X19) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X17))))) ) | $spl29),
  inference(superposition,[],[f14386,f866])).
fof(f14386,plain,(
  ( ! [X72,X71] : (sP2(X71,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),sK14)),X72)) ) | $spl29),
  inference(resolution,[],[f12255,f657])).
fof(f20838,plain,(
  ( ! [X19,X17,X18] : (sP4(X18,X19,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),X17))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) | $spl476),
  inference(cnf_transformation,[],[f20838_D])).
fof(f20838_D,plain,(
  ( ! [X19,X17,X18] : (sP4(X18,X19,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),X17))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) <=> ~$spl476),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl476])])).
fof(f15592,plain,(
  ( ! [X19,X17,X18] : (sP4(X18,X19,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl29),
  inference(subsumption_resolution,[],[f15583,f860])).
fof(f15583,plain,(
  ( ! [X19,X17,X18] : (sP4(X18,X19,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X17))))) ) | $spl29),
  inference(superposition,[],[f14402,f866])).
fof(f14402,plain,(
  ( ! [X103,X104] : (sP4(X103,X104,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),sK14)))) ) | $spl29),
  inference(resolution,[],[f12255,f673])).
fof(f20837,plain,(
  ( ! [X19,X17,X18] : (sP8(X18,X19,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),X17))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) | $spl474),
  inference(cnf_transformation,[],[f20837_D])).
fof(f20837_D,plain,(
  ( ! [X19,X17,X18] : (sP8(X18,X19,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),X17))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) <=> ~$spl474),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl474])])).
fof(f15615,plain,(
  ( ! [X19,X17,X18] : (sP8(X18,X19,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl29),
  inference(subsumption_resolution,[],[f15606,f860])).
fof(f15606,plain,(
  ( ! [X19,X17,X18] : (sP8(X18,X19,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X17))))) ) | $spl29),
  inference(superposition,[],[f14415,f866])).
fof(f14415,plain,(
  ( ! [X130,X129] : (sP8(X129,X130,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),sK14)))) ) | $spl29),
  inference(resolution,[],[f12255,f690])).
fof(f20836,plain,(
  ( ! [X19,X17,X18] : (sP7(X18,X19,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),X17))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) | $spl472),
  inference(cnf_transformation,[],[f20836_D])).
fof(f20836_D,plain,(
  ( ! [X19,X17,X18] : (sP7(X18,X19,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),X17))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) <=> ~$spl472),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl472])])).
fof(f15684,plain,(
  ( ! [X19,X17,X18] : (sP7(X18,X19,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl29),
  inference(subsumption_resolution,[],[f15675,f860])).
fof(f15675,plain,(
  ( ! [X19,X17,X18] : (sP7(X18,X19,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X17))))) ) | $spl29),
  inference(superposition,[],[f14418,f866])).
fof(f14418,plain,(
  ( ! [X136,X135] : (sP7(X135,X136,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),sK14)))) ) | $spl29),
  inference(resolution,[],[f12255,f694])).
fof(f20835,plain,(
  ( ! [X19,X17,X18] : (sP6(X18,X19,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),X17))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) | $spl470),
  inference(cnf_transformation,[],[f20835_D])).
fof(f20835_D,plain,(
  ( ! [X19,X17,X18] : (sP6(X18,X19,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),X17))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) <=> ~$spl470),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl470])])).
fof(f15708,plain,(
  ( ! [X19,X17,X18] : (sP6(X18,X19,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl29),
  inference(subsumption_resolution,[],[f15698,f860])).
fof(f15698,plain,(
  ( ! [X19,X17,X18] : (sP6(X18,X19,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X17))))) ) | $spl29),
  inference(superposition,[],[f14419,f866])).
fof(f14419,plain,(
  ( ! [X138,X137] : (sP6(X137,X138,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),sK14)))) ) | $spl29),
  inference(resolution,[],[f12255,f698])).
fof(f20834,plain,(
  ( ! [X19,X17,X18] : (sP6(X18,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),X17)),X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) | $spl468),
  inference(cnf_transformation,[],[f20834_D])).
fof(f20834_D,plain,(
  ( ! [X19,X17,X18] : (sP6(X18,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),X17)),X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) <=> ~$spl468),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl468])])).
fof(f15732,plain,(
  ( ! [X19,X17,X18] : (sP6(X18,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),X17)),X19) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl29),
  inference(subsumption_resolution,[],[f15722,f860])).
fof(f15722,plain,(
  ( ! [X19,X17,X18] : (sP6(X18,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),X17)),X19) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X17))))) ) | $spl29),
  inference(superposition,[],[f14420,f866])).
fof(f14420,plain,(
  ( ! [X140,X139] : (sP6(X139,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),sK14)),X140)) ) | $spl29),
  inference(resolution,[],[f12255,f699])).
fof(f20833,plain,(
  ( ! [X19,X17,X18] : (sP6('const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),X17)),X18,X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) | $spl466),
  inference(cnf_transformation,[],[f20833_D])).
fof(f20833_D,plain,(
  ( ! [X19,X17,X18] : (sP6('const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),X17)),X18,X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) <=> ~$spl466),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl466])])).
fof(f15756,plain,(
  ( ! [X19,X17,X18] : (sP6('const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),X17)),X18,X19) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl29),
  inference(subsumption_resolution,[],[f15746,f860])).
fof(f15746,plain,(
  ( ! [X19,X17,X18] : (sP6('const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),X17)),X18,X19) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X17))))) ) | $spl29),
  inference(superposition,[],[f14421,f866])).
fof(f14421,plain,(
  ( ! [X142,X141] : (sP6('const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),sK14)),X141,X142)) ) | $spl29),
  inference(resolution,[],[f12255,f700])).
fof(f20832,plain,(
  ( ! [X19,X17,X18] : (sP10('const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),X17)),X18,X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) | $spl464),
  inference(cnf_transformation,[],[f20832_D])).
fof(f20832_D,plain,(
  ( ! [X19,X17,X18] : (sP10('const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),X17)),X18,X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) <=> ~$spl464),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl464])])).
fof(f15777,plain,(
  ( ! [X19,X17,X18] : (sP10('const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),X17)),X18,X19) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl29),
  inference(subsumption_resolution,[],[f15770,f860])).
fof(f15770,plain,(
  ( ! [X19,X17,X18] : (sP10('const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),X17)),X18,X19) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X17))))) ) | $spl29),
  inference(superposition,[],[f14432,f866])).
fof(f14432,plain,(
  ( ! [X163,X164] : (sP10('const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),sK14)),X163,X164)) ) | $spl29),
  inference(resolution,[],[f12255,f718])).
fof(f20831,plain,(
  ( ! [X19,X17,X18] : (sP11(X18,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),X17)),X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) | $spl462),
  inference(cnf_transformation,[],[f20831_D])).
fof(f20831_D,plain,(
  ( ! [X19,X17,X18] : (sP11(X18,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),X17)),X19) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14))))) ) <=> ~$spl462),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl462])])).
fof(f15844,plain,(
  ( ! [X19,X17,X18] : (sP11(X18,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),X17)),X19) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl29),
  inference(subsumption_resolution,[],[f15837,f860])).
fof(f15837,plain,(
  ( ! [X19,X17,X18] : (sP11(X18,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),X17)),X19) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X17),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X17))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X17))))) ) | $spl29),
  inference(superposition,[],[f14439,f866])).
fof(f14439,plain,(
  ( ! [X177,X178] : (sP11(X177,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),sK14)),X178)) ) | $spl29),
  inference(resolution,[],[f12255,f727])).
fof(f1092,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | $spl10),
  inference(cnf_transformation,[],[f1092_D])).
fof(f1092_D,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f16063,plain,(
  ( ! [X12,X13] : (p(s(bool,X13)) | p(s(bool,X12)) | p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl54),
  inference(subsumption_resolution,[],[f16044,f771])).
fof(f16044,plain,(
  ( ! [X12,X13] : (s(bool,X12) != s(bool,X13) | p(s(bool,X13)) | p(s(bool,X12)) | p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl54),
  inference(superposition,[],[f2053,f771])).
fof(f16064,plain,(
  ( ! [X14,X15] : (p(s(bool,X15)) | p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | p(s(bool,X14))) ) | $spl54),
  inference(subsumption_resolution,[],[f16045,f771])).
fof(f16045,plain,(
  ( ! [X14,X15] : (s(bool,X14) != s(bool,X15) | p(s(bool,X15)) | p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | p(s(bool,X14))) ) | $spl54),
  inference(superposition,[],[f2053,f771])).
fof(f16067,plain,(
  ( ! [X21,X20] : (p(s(bool,X21)) | p(s(bool,X20)) | p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl54),
  inference(subsumption_resolution,[],[f16048,f771])).
fof(f16048,plain,(
  ( ! [X21,X20] : (s(bool,X20) != s(bool,X21) | p(s(bool,X21)) | p(s(bool,X20)) | p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl54),
  inference(superposition,[],[f2053,f779])).
fof(f779,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f434])).
fof(f434,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f433])).
fof(f433,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f295])).
fof(f295,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X10,X11] : ((p(s(bool,X10)) => p(s(bool,X11))) => ((p(s(bool,X11)) => p(s(bool,X10))) => s(bool,X10) = s(bool,X11)))),
  file('thm.pred_set.SUBSET_DIFF_EMPTY','thm.bool.IMP_ANTISYM_AX')).
fof(f16068,plain,(
  ( ! [X23,X22] : (p(s(bool,X23)) | p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | p(s(bool,X22))) ) | $spl54),
  inference(subsumption_resolution,[],[f16049,f771])).
fof(f16049,plain,(
  ( ! [X23,X22] : (s(bool,X22) != s(bool,X23) | p(s(bool,X23)) | p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | p(s(bool,X22))) ) | $spl54),
  inference(superposition,[],[f2053,f779])).
fof(f20830,plain,(
  ( ! [X26,X24,X23,X25] : (~sP1(X24,happ(s(fun(sK13,bool),X23),s(sK13,X25)),X26) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | p(s(bool,X24))) ) | $spl460),
  inference(cnf_transformation,[],[f20830_D])).
fof(f20830_D,plain,(
  ( ! [X26,X24,X23,X25] : (~sP1(X24,happ(s(fun(sK13,bool),X23),s(sK13,X25)),X26) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | p(s(bool,X24))) ) <=> ~$spl460),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl460])])).
fof(f16728,plain,(
  ( ! [X26,X24,X23,X25] : (~sP1(X24,happ(s(fun(sK13,bool),X23),s(sK13,X25)),X26) | p(s(bool,X24)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl204),
  inference(subsumption_resolution,[],[f16715,f860])).
fof(f16715,plain,(
  ( ! [X26,X24,X23,X25] : (~sP1(X24,happ(s(fun(sK13,bool),X23),s(sK13,X25)),X26) | p(s(bool,X24)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X23))))) ) | $spl204),
  inference(superposition,[],[f10762,f866])).
fof(f10762,plain,(
  ( ! [X78,X79,X77] : (~sP1(X77,happ(s(fun(sK13,bool),sK14),s(sK13,X78)),X79) | p(s(bool,X77))) ) | $spl204),
  inference(resolution,[],[f5495,f632])).
fof(f632,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | ~sP1(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f484])).
fof(f484,plain,(
  ! [X0,X1,X2] : ((~sP1(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP0(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP0(X0,X1,X2) | sP1(X0,X1,X2)))),
  inference(rectify,[],[f483])).
fof(f483,plain,(
  ! [X2,X0,X1] : ((~sP1(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP0(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP0(X2,X0,X1) | sP1(X2,X0,X1)))),
  inference(flattening,[],[f482])).
fof(f482,plain,(
  ! [X2,X0,X1] : ((~sP1(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP0(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP0(X2,X0,X1)) | sP1(X2,X0,X1)))),
  inference(nnf_transformation,[],[f459])).
fof(f459,plain,(
  ! [X2,X0,X1] : (sP1(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP0(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f20829,plain,(
  ( ! [X26,X24,X23,X25] : (sP1(happ(s(fun(sK13,bool),X23),s(sK13,X24)),X25,X26) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | p(s(bool,X26))) ) | $spl458),
  inference(cnf_transformation,[],[f20829_D])).
fof(f20829_D,plain,(
  ( ! [X26,X24,X23,X25] : (sP1(happ(s(fun(sK13,bool),X23),s(sK13,X24)),X25,X26) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | p(s(bool,X26))) ) <=> ~$spl458),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl458])])).
fof(f16759,plain,(
  ( ! [X26,X24,X23,X25] : (sP1(happ(s(fun(sK13,bool),X23),s(sK13,X24)),X25,X26) | p(s(bool,X26)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl204),
  inference(subsumption_resolution,[],[f16746,f860])).
fof(f16746,plain,(
  ( ! [X26,X24,X23,X25] : (sP1(happ(s(fun(sK13,bool),X23),s(sK13,X24)),X25,X26) | p(s(bool,X26)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X23))))) ) | $spl204),
  inference(superposition,[],[f10771,f866])).
fof(f10771,plain,(
  ( ! [X105,X106,X104] : (sP1(happ(s(fun(sK13,bool),sK14),s(sK13,X104)),X105,X106) | p(s(bool,X106))) ) | $spl204),
  inference(resolution,[],[f5495,f643])).
fof(f643,plain,(
  ( ! [X2,X0,X1] : (sP1(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f489])).
fof(f489,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP1(X2,X0,X1)) & (~sP1(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f488])).
fof(f488,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP1(X2,X0,X1)) & (~sP1(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f460])).
fof(f460,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP1(X2,X0,X1))),
  inference(definition_folding,[],[f418,f459,f458])).
fof(f418,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f266])).
fof(f266,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f265])).
fof(f265,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f70])).
fof(f70,axiom,(
  ! [X25,X23,X24] : ((p(s(bool,X24)) <=> (p(s(bool,X23)) => p(s(bool,X25)))) <=> ((p(s(bool,X24)) | p(s(bool,X23))) & (p(s(bool,X24)) | ~p(s(bool,X25))) & (~p(s(bool,X23)) | p(s(bool,X25)) | ~p(s(bool,X24)))))),
  file('thm.pred_set.SUBSET_DIFF_EMPTY','thm.sat.sat_27_prove_ltb')).
fof(f20828,plain,(
  ( ! [X26,X24,X23,X25] : (sP1(X24,happ(s(fun(sK13,bool),X23),s(sK13,X25)),X26) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | ~p(s(bool,X24))) ) | $spl456),
  inference(cnf_transformation,[],[f20828_D])).
fof(f20828_D,plain,(
  ( ! [X26,X24,X23,X25] : (sP1(X24,happ(s(fun(sK13,bool),X23),s(sK13,X25)),X26) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | ~p(s(bool,X24))) ) <=> ~$spl456),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl456])])).
fof(f16786,plain,(
  ( ! [X26,X24,X23,X25] : (sP1(X24,happ(s(fun(sK13,bool),X23),s(sK13,X25)),X26) | ~p(s(bool,X24)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl204),
  inference(subsumption_resolution,[],[f16774,f860])).
fof(f16774,plain,(
  ( ! [X26,X24,X23,X25] : (sP1(X24,happ(s(fun(sK13,bool),X23),s(sK13,X25)),X26) | ~p(s(bool,X24)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X23))))) ) | $spl204),
  inference(superposition,[],[f10772,f866])).
fof(f10772,plain,(
  ( ! [X109,X107,X108] : (sP1(X107,happ(s(fun(sK13,bool),sK14),s(sK13,X108)),X109) | ~p(s(bool,X107))) ) | $spl204),
  inference(resolution,[],[f5495,f644])).
fof(f644,plain,(
  ( ! [X2,X0,X1] : (sP1(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f489])).
fof(f20827,plain,(
  ( ! [X26,X24,X23,X25] : (sP1(happ(s(fun(sK13,bool),X23),s(sK13,X24)),X25,X26) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | ~p(s(bool,X25))) ) | $spl454),
  inference(cnf_transformation,[],[f20827_D])).
fof(f20827_D,plain,(
  ( ! [X26,X24,X23,X25] : (sP1(happ(s(fun(sK13,bool),X23),s(sK13,X24)),X25,X26) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | ~p(s(bool,X25))) ) <=> ~$spl454),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl454])])).
fof(f16817,plain,(
  ( ! [X26,X24,X23,X25] : (sP1(happ(s(fun(sK13,bool),X23),s(sK13,X24)),X25,X26) | ~p(s(bool,X25)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl204),
  inference(subsumption_resolution,[],[f16804,f860])).
fof(f16804,plain,(
  ( ! [X26,X24,X23,X25] : (sP1(happ(s(fun(sK13,bool),X23),s(sK13,X24)),X25,X26) | ~p(s(bool,X25)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X23))))) ) | $spl204),
  inference(superposition,[],[f10773,f866])).
fof(f10773,plain,(
  ( ! [X111,X112,X110] : (sP1(happ(s(fun(sK13,bool),sK14),s(sK13,X110)),X111,X112) | ~p(s(bool,X111))) ) | $spl204),
  inference(resolution,[],[f5495,f644])).
fof(f20826,plain,(
  ( ! [X26,X24,X23,X25] : (~sP3(X24,X25,happ(s(fun(sK13,bool),X23),s(sK13,X26))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | p(s(bool,X24))) ) | $spl452),
  inference(cnf_transformation,[],[f20826_D])).
fof(f20826_D,plain,(
  ( ! [X26,X24,X23,X25] : (~sP3(X24,X25,happ(s(fun(sK13,bool),X23),s(sK13,X26))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | p(s(bool,X24))) ) <=> ~$spl452),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl452])])).
fof(f16895,plain,(
  ( ! [X26,X24,X23,X25] : (~sP3(X24,X25,happ(s(fun(sK13,bool),X23),s(sK13,X26))) | p(s(bool,X24)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl204),
  inference(subsumption_resolution,[],[f16882,f860])).
fof(f16882,plain,(
  ( ! [X26,X24,X23,X25] : (~sP3(X24,X25,happ(s(fun(sK13,bool),X23),s(sK13,X26))) | p(s(bool,X24)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X23))))) ) | $spl204),
  inference(superposition,[],[f10777,f866])).
fof(f10777,plain,(
  ( ! [X123,X124,X122] : (~sP3(X122,X123,happ(s(fun(sK13,bool),sK14),s(sK13,X124))) | p(s(bool,X122))) ) | $spl204),
  inference(resolution,[],[f5495,f648])).
fof(f648,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X0)) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f492])).
fof(f492,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP2(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP2(X0,X1,X2) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f491])).
fof(f491,plain,(
  ! [X2,X0,X1] : ((~sP3(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP2(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP2(X2,X0,X1) | sP3(X2,X0,X1)))),
  inference(flattening,[],[f490])).
fof(f490,plain,(
  ! [X2,X0,X1] : ((~sP3(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP2(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP2(X2,X0,X1)) | sP3(X2,X0,X1)))),
  inference(nnf_transformation,[],[f462])).
fof(f462,plain,(
  ! [X2,X0,X1] : (sP3(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP2(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f20825,plain,(
  ( ! [X26,X24,X23,X25] : (~sP3(X24,happ(s(fun(sK13,bool),X23),s(sK13,X25)),X26) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | p(s(bool,X24))) ) | $spl450),
  inference(cnf_transformation,[],[f20825_D])).
fof(f20825_D,plain,(
  ( ! [X26,X24,X23,X25] : (~sP3(X24,happ(s(fun(sK13,bool),X23),s(sK13,X25)),X26) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | p(s(bool,X24))) ) <=> ~$spl450),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl450])])).
fof(f16933,plain,(
  ( ! [X26,X24,X23,X25] : (~sP3(X24,happ(s(fun(sK13,bool),X23),s(sK13,X25)),X26) | p(s(bool,X24)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl204),
  inference(subsumption_resolution,[],[f16920,f860])).
fof(f16920,plain,(
  ( ! [X26,X24,X23,X25] : (~sP3(X24,happ(s(fun(sK13,bool),X23),s(sK13,X25)),X26) | p(s(bool,X24)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X23))))) ) | $spl204),
  inference(superposition,[],[f10778,f866])).
fof(f10778,plain,(
  ( ! [X127,X125,X126] : (~sP3(X125,happ(s(fun(sK13,bool),sK14),s(sK13,X126)),X127) | p(s(bool,X125))) ) | $spl204),
  inference(resolution,[],[f5495,f649])).
fof(f649,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f492])).
fof(f20824,plain,(
  ( ! [X26,X24,X23,X25] : (sP3(X24,X25,happ(s(fun(sK13,bool),X23),s(sK13,X26))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | ~p(s(bool,X24))) ) | $spl448),
  inference(cnf_transformation,[],[f20824_D])).
fof(f20824_D,plain,(
  ( ! [X26,X24,X23,X25] : (sP3(X24,X25,happ(s(fun(sK13,bool),X23),s(sK13,X26))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | ~p(s(bool,X24))) ) <=> ~$spl448),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl448])])).
fof(f16964,plain,(
  ( ! [X26,X24,X23,X25] : (sP3(X24,X25,happ(s(fun(sK13,bool),X23),s(sK13,X26))) | ~p(s(bool,X24)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl204),
  inference(subsumption_resolution,[],[f16952,f860])).
fof(f16952,plain,(
  ( ! [X26,X24,X23,X25] : (sP3(X24,X25,happ(s(fun(sK13,bool),X23),s(sK13,X26))) | ~p(s(bool,X24)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X23))))) ) | $spl204),
  inference(superposition,[],[f10784,f866])).
fof(f10784,plain,(
  ( ! [X144,X145,X143] : (sP3(X143,X144,happ(s(fun(sK13,bool),sK14),s(sK13,X145))) | ~p(s(bool,X143))) ) | $spl204),
  inference(resolution,[],[f5495,f660])).
fof(f660,plain,(
  ( ! [X2,X0,X1] : (sP3(X2,X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f497])).
fof(f497,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP3(X2,X0,X1)) & (~sP3(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f496])).
fof(f496,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP3(X2,X0,X1)) & (~sP3(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f463])).
fof(f463,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP3(X2,X0,X1))),
  inference(definition_folding,[],[f268,f462,f461])).
fof(f268,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f267])).
fof(f267,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f82])).
fof(f82,axiom,(
  ! [X25,X23,X24] : ((p(s(bool,X24)) <=> (p(s(bool,X23)) | p(s(bool,X25)))) <=> ((p(s(bool,X24)) | ~p(s(bool,X23))) & (p(s(bool,X24)) | ~p(s(bool,X25))) & (p(s(bool,X23)) | p(s(bool,X25)) | ~p(s(bool,X24)))))),
  file('thm.pred_set.SUBSET_DIFF_EMPTY','thm.sat.sat_25_prove_ltb')).
fof(f20823,plain,(
  ( ! [X26,X24,X23,X25] : (sP3(happ(s(fun(sK13,bool),X23),s(sK13,X24)),X25,X26) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | ~p(s(bool,X26))) ) | $spl446),
  inference(cnf_transformation,[],[f20823_D])).
fof(f20823_D,plain,(
  ( ! [X26,X24,X23,X25] : (sP3(happ(s(fun(sK13,bool),X23),s(sK13,X24)),X25,X26) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | ~p(s(bool,X26))) ) <=> ~$spl446),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl446])])).
fof(f16999,plain,(
  ( ! [X26,X24,X23,X25] : (sP3(happ(s(fun(sK13,bool),X23),s(sK13,X24)),X25,X26) | ~p(s(bool,X26)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl204),
  inference(subsumption_resolution,[],[f16986,f860])).
fof(f16986,plain,(
  ( ! [X26,X24,X23,X25] : (sP3(happ(s(fun(sK13,bool),X23),s(sK13,X24)),X25,X26) | ~p(s(bool,X26)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X23))))) ) | $spl204),
  inference(superposition,[],[f10785,f866])).
fof(f10785,plain,(
  ( ! [X146,X147,X148] : (sP3(happ(s(fun(sK13,bool),sK14),s(sK13,X146)),X147,X148) | ~p(s(bool,X148))) ) | $spl204),
  inference(resolution,[],[f5495,f660])).
fof(f20822,plain,(
  ( ! [X26,X24,X23,X25] : (sP3(X24,happ(s(fun(sK13,bool),X23),s(sK13,X25)),X26) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | ~p(s(bool,X24))) ) | $spl444),
  inference(cnf_transformation,[],[f20822_D])).
fof(f20822_D,plain,(
  ( ! [X26,X24,X23,X25] : (sP3(X24,happ(s(fun(sK13,bool),X23),s(sK13,X25)),X26) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | ~p(s(bool,X24))) ) <=> ~$spl444),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl444])])).
fof(f17030,plain,(
  ( ! [X26,X24,X23,X25] : (sP3(X24,happ(s(fun(sK13,bool),X23),s(sK13,X25)),X26) | ~p(s(bool,X24)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl204),
  inference(subsumption_resolution,[],[f17018,f860])).
fof(f17018,plain,(
  ( ! [X26,X24,X23,X25] : (sP3(X24,happ(s(fun(sK13,bool),X23),s(sK13,X25)),X26) | ~p(s(bool,X24)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X23))))) ) | $spl204),
  inference(superposition,[],[f10786,f866])).
fof(f10786,plain,(
  ( ! [X151,X149,X150] : (sP3(X149,happ(s(fun(sK13,bool),sK14),s(sK13,X150)),X151) | ~p(s(bool,X149))) ) | $spl204),
  inference(resolution,[],[f5495,f661])).
fof(f661,plain,(
  ( ! [X2,X0,X1] : (sP3(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f497])).
fof(f20821,plain,(
  ( ! [X26,X24,X23,X25] : (sP3(happ(s(fun(sK13,bool),X23),s(sK13,X24)),X25,X26) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | ~p(s(bool,X25))) ) | $spl442),
  inference(cnf_transformation,[],[f20821_D])).
fof(f20821_D,plain,(
  ( ! [X26,X24,X23,X25] : (sP3(happ(s(fun(sK13,bool),X23),s(sK13,X24)),X25,X26) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | ~p(s(bool,X25))) ) <=> ~$spl442),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl442])])).
fof(f17105,plain,(
  ( ! [X26,X24,X23,X25] : (sP3(happ(s(fun(sK13,bool),X23),s(sK13,X24)),X25,X26) | ~p(s(bool,X25)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl204),
  inference(subsumption_resolution,[],[f17092,f860])).
fof(f17092,plain,(
  ( ! [X26,X24,X23,X25] : (sP3(happ(s(fun(sK13,bool),X23),s(sK13,X24)),X25,X26) | ~p(s(bool,X25)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X23))))) ) | $spl204),
  inference(superposition,[],[f10787,f866])).
fof(f10787,plain,(
  ( ! [X154,X152,X153] : (sP3(happ(s(fun(sK13,bool),sK14),s(sK13,X152)),X153,X154) | ~p(s(bool,X153))) ) | $spl204),
  inference(resolution,[],[f5495,f661])).
fof(f20820,plain,(
  ( ! [X26,X24,X23,X25] : (~sP5(happ(s(fun(sK13,bool),X23),s(sK13,X24)),X25,X26) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | p(s(bool,X26))) ) | $spl440),
  inference(cnf_transformation,[],[f20820_D])).
fof(f20820_D,plain,(
  ( ! [X26,X24,X23,X25] : (~sP5(happ(s(fun(sK13,bool),X23),s(sK13,X24)),X25,X26) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | p(s(bool,X26))) ) <=> ~$spl440),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl440])])).
fof(f17134,plain,(
  ( ! [X26,X24,X23,X25] : (~sP5(happ(s(fun(sK13,bool),X23),s(sK13,X24)),X25,X26) | p(s(bool,X26)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl204),
  inference(subsumption_resolution,[],[f17120,f860])).
fof(f17120,plain,(
  ( ! [X26,X24,X23,X25] : (~sP5(happ(s(fun(sK13,bool),X23),s(sK13,X24)),X25,X26) | p(s(bool,X26)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X23))))) ) | $spl204),
  inference(superposition,[],[f10791,f866])).
fof(f10791,plain,(
  ( ! [X165,X166,X164] : (~sP5(happ(s(fun(sK13,bool),sK14),s(sK13,X165)),X166,X164) | p(s(bool,X164))) ) | $spl204),
  inference(resolution,[],[f5495,f666])).
fof(f666,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X2)) | ~sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f500])).
fof(f500,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | (sP4(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP4(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f499])).
fof(f499,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | (sP4(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP4(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP5(X2,X0,X1)))),
  inference(flattening,[],[f498])).
fof(f498,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | (sP4(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP4(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP5(X2,X0,X1)))),
  inference(nnf_transformation,[],[f465])).
fof(f465,plain,(
  ! [X2,X0,X1] : (sP5(X2,X0,X1) <=> (sP4(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f20819,plain,(
  ( ! [X26,X24,X23,X25] : (~sP5(happ(s(fun(sK13,bool),X23),s(sK13,X24)),X25,X26) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | p(s(bool,X25))) ) | $spl438),
  inference(cnf_transformation,[],[f20819_D])).
fof(f20819_D,plain,(
  ( ! [X26,X24,X23,X25] : (~sP5(happ(s(fun(sK13,bool),X23),s(sK13,X24)),X25,X26) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | p(s(bool,X25))) ) <=> ~$spl438),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl438])])).
fof(f17163,plain,(
  ( ! [X26,X24,X23,X25] : (~sP5(happ(s(fun(sK13,bool),X23),s(sK13,X24)),X25,X26) | p(s(bool,X25)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl204),
  inference(subsumption_resolution,[],[f17149,f860])).
fof(f17149,plain,(
  ( ! [X26,X24,X23,X25] : (~sP5(happ(s(fun(sK13,bool),X23),s(sK13,X24)),X25,X26) | p(s(bool,X25)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X23))))) ) | $spl204),
  inference(superposition,[],[f10792,f866])).
fof(f10792,plain,(
  ( ! [X167,X169,X168] : (~sP5(happ(s(fun(sK13,bool),sK14),s(sK13,X168)),X167,X169) | p(s(bool,X167))) ) | $spl204),
  inference(resolution,[],[f5495,f667])).
fof(f667,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f500])).
fof(f20818,plain,(
  ( ! [X26,X24,X23,X25] : (sP10(X24,happ(s(fun(sK13,bool),X23),s(sK13,X25)),X26) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | ~p(s(bool,X26))) ) | $spl436),
  inference(cnf_transformation,[],[f20818_D])).
fof(f20818_D,plain,(
  ( ! [X26,X24,X23,X25] : (sP10(X24,happ(s(fun(sK13,bool),X23),s(sK13,X25)),X26) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | ~p(s(bool,X26))) ) <=> ~$spl436),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl436])])).
fof(f17186,plain,(
  ( ! [X26,X24,X23,X25] : (sP10(X24,happ(s(fun(sK13,bool),X23),s(sK13,X25)),X26) | ~p(s(bool,X26)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl204),
  inference(subsumption_resolution,[],[f17175,f860])).
fof(f17175,plain,(
  ( ! [X26,X24,X23,X25] : (sP10(X24,happ(s(fun(sK13,bool),X23),s(sK13,X25)),X26) | ~p(s(bool,X26)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X23))))) ) | $spl204),
  inference(superposition,[],[f10821,f866])).
fof(f10821,plain,(
  ( ! [X255,X254,X256] : (sP10(X254,happ(s(fun(sK13,bool),sK14),s(sK13,X255)),X256) | ~p(s(bool,X256))) ) | $spl204),
  inference(resolution,[],[f5495,f713])).
fof(f713,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f521])).
fof(f521,plain,(
  ! [X0,X1,X2] : ((~sP10(X0,X1,X2) | ((p(s(bool,X2)) | p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0))))) & ((~p(s(bool,X2)) & ~p(s(bool,X0))) | (~p(s(bool,X1)) & ~p(s(bool,X0))) | sP10(X0,X1,X2)))),
  inference(rectify,[],[f520])).
fof(f520,plain,(
  ! [X0,X2,X1] : ((~sP10(X0,X2,X1) | ((p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | p(s(bool,X0))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (~p(s(bool,X2)) & ~p(s(bool,X0))) | sP10(X0,X2,X1)))),
  inference(flattening,[],[f519])).
fof(f519,plain,(
  ! [X0,X2,X1] : ((~sP10(X0,X2,X1) | ((p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | p(s(bool,X0))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (~p(s(bool,X2)) & ~p(s(bool,X0)))) | sP10(X0,X2,X1)))),
  inference(nnf_transformation,[],[f472])).
fof(f20817,plain,(
  ( ! [X26,X24,X23,X25] : (sP10(X24,X25,happ(s(fun(sK13,bool),X23),s(sK13,X26))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | ~p(s(bool,X25))) ) | $spl434),
  inference(cnf_transformation,[],[f20817_D])).
fof(f20817_D,plain,(
  ( ! [X26,X24,X23,X25] : (sP10(X24,X25,happ(s(fun(sK13,bool),X23),s(sK13,X26))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | ~p(s(bool,X25))) ) <=> ~$spl434),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl434])])).
fof(f17209,plain,(
  ( ! [X26,X24,X23,X25] : (sP10(X24,X25,happ(s(fun(sK13,bool),X23),s(sK13,X26))) | ~p(s(bool,X25)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl204),
  inference(subsumption_resolution,[],[f17198,f860])).
fof(f17198,plain,(
  ( ! [X26,X24,X23,X25] : (sP10(X24,X25,happ(s(fun(sK13,bool),X23),s(sK13,X26))) | ~p(s(bool,X25)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X23))))) ) | $spl204),
  inference(superposition,[],[f10822,f866])).
fof(f10822,plain,(
  ( ! [X257,X259,X258] : (sP10(X257,X258,happ(s(fun(sK13,bool),sK14),s(sK13,X259))) | ~p(s(bool,X258))) ) | $spl204),
  inference(resolution,[],[f5495,f713])).
fof(f20816,plain,(
  ( ! [X26,X24,X23,X25] : (sP11(happ(s(fun(sK13,bool),X23),s(sK13,X24)),X25,X26) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | ~p(s(bool,X26))) ) | $spl432),
  inference(cnf_transformation,[],[f20816_D])).
fof(f20816_D,plain,(
  ( ! [X26,X24,X23,X25] : (sP11(happ(s(fun(sK13,bool),X23),s(sK13,X24)),X25,X26) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | ~p(s(bool,X26))) ) <=> ~$spl432),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl432])])).
fof(f17266,plain,(
  ( ! [X26,X24,X23,X25] : (sP11(happ(s(fun(sK13,bool),X23),s(sK13,X24)),X25,X26) | ~p(s(bool,X26)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl204),
  inference(subsumption_resolution,[],[f17255,f860])).
fof(f17255,plain,(
  ( ! [X26,X24,X23,X25] : (sP11(happ(s(fun(sK13,bool),X23),s(sK13,X24)),X25,X26) | ~p(s(bool,X26)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X23))))) ) | $spl204),
  inference(superposition,[],[f10834,f866])).
fof(f10834,plain,(
  ( ! [X294,X293,X295] : (sP11(happ(s(fun(sK13,bool),sK14),s(sK13,X293)),X294,X295) | ~p(s(bool,X295))) ) | $spl204),
  inference(resolution,[],[f5495,f726])).
fof(f726,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f526])).
fof(f526,plain,(
  ! [X0,X1,X2] : ((~sP11(X0,X1,X2) | ((p(s(bool,X1)) | p(s(bool,X2))) & (p(s(bool,X1)) | p(s(bool,X0))))) & ((~p(s(bool,X1)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) & ~p(s(bool,X0))) | sP11(X0,X1,X2)))),
  inference(rectify,[],[f525])).
fof(f525,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X2))) | sP11(X2,X0,X1)))),
  inference(flattening,[],[f524])).
fof(f524,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2))))) & (((~p(s(bool,X0)) & ~p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X2)))) | sP11(X2,X0,X1)))),
  inference(nnf_transformation,[],[f474])).
fof(f20815,plain,(
  ( ! [X26,X24,X23,X25] : (sP11(X24,X25,happ(s(fun(sK13,bool),X23),s(sK13,X26))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | ~p(s(bool,X24))) ) | $spl430),
  inference(cnf_transformation,[],[f20815_D])).
fof(f20815_D,plain,(
  ( ! [X26,X24,X23,X25] : (sP11(X24,X25,happ(s(fun(sK13,bool),X23),s(sK13,X26))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | ~p(s(bool,X24))) ) <=> ~$spl430),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl430])])).
fof(f17289,plain,(
  ( ! [X26,X24,X23,X25] : (sP11(X24,X25,happ(s(fun(sK13,bool),X23),s(sK13,X26))) | ~p(s(bool,X24)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl204),
  inference(subsumption_resolution,[],[f17278,f860])).
fof(f17278,plain,(
  ( ! [X26,X24,X23,X25] : (sP11(X24,X25,happ(s(fun(sK13,bool),X23),s(sK13,X26))) | ~p(s(bool,X24)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X23),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X23))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X23))))) ) | $spl204),
  inference(superposition,[],[f10835,f866])).
fof(f10835,plain,(
  ( ! [X296,X298,X297] : (sP11(X296,X297,happ(s(fun(sK13,bool),sK14),s(sK13,X298))) | ~p(s(bool,X296))) ) | $spl204),
  inference(resolution,[],[f5495,f726])).
fof(f20814,plain,(
  ( ! [X12,X11] : (~sP12(X12,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X11))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X11))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X11),s(fun(sK13,bool),sK14)))) | p(s(bool,X12))) ) | $spl428),
  inference(cnf_transformation,[],[f20814_D])).
fof(f20814_D,plain,(
  ( ! [X12,X11] : (~sP12(X12,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X11))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X11))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X11),s(fun(sK13,bool),sK14)))) | p(s(bool,X12))) ) <=> ~$spl428),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl428])])).
fof(f19080,plain,(
  ( ! [X12,X11] : (~sP12(X12,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X11))) | p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X11),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X11))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl5),
  inference(subsumption_resolution,[],[f19068,f860])).
fof(f19068,plain,(
  ( ! [X12,X11] : (~sP12(X12,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X11))) | p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X11),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X11))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X11))))) ) | $spl5),
  inference(superposition,[],[f11811,f866])).
fof(f11811,plain,(
  ( ! [X176] : (~sP12(X176,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))) | p(s(bool,X176))) ) | $spl5),
  inference(resolution,[],[f1086,f751])).
fof(f751,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | ~sP12(X0,X1)) )),
  inference(cnf_transformation,[],[f539])).
fof(f20813,plain,(
  ( ! [X12,X11] : (~sP12('const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X11)),X12) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X11))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X11),s(fun(sK13,bool),sK14)))) | p(s(bool,X12))) ) | $spl426),
  inference(cnf_transformation,[],[f20813_D])).
fof(f20813_D,plain,(
  ( ! [X12,X11] : (~sP12('const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X11)),X12) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X11))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X11),s(fun(sK13,bool),sK14)))) | p(s(bool,X12))) ) <=> ~$spl426),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl426])])).
fof(f19120,plain,(
  ( ! [X12,X11] : (~sP12('const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X11)),X12) | p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X11),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X11))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl5),
  inference(subsumption_resolution,[],[f19108,f860])).
fof(f19108,plain,(
  ( ! [X12,X11] : (~sP12('const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X11)),X12) | p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X11),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X11))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X11))))) ) | $spl5),
  inference(superposition,[],[f11812,f866])).
fof(f11812,plain,(
  ( ! [X177] : (~sP12('const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X177) | p(s(bool,X177))) ) | $spl5),
  inference(resolution,[],[f1086,f751])).
fof(f20812,plain,(
  ( ! [X12,X11] : (sP12(X12,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X11))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X11))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X11),s(fun(sK13,bool),sK14)))) | ~p(s(bool,X12))) ) | $spl424),
  inference(cnf_transformation,[],[f20812_D])).
fof(f20812_D,plain,(
  ( ! [X12,X11] : (sP12(X12,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X11))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X11))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X11),s(fun(sK13,bool),sK14)))) | ~p(s(bool,X12))) ) <=> ~$spl424),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl424])])).
fof(f19162,plain,(
  ( ! [X12,X11] : (sP12(X12,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X11))) | ~p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X11),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X11))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl5),
  inference(subsumption_resolution,[],[f19150,f860])).
fof(f19150,plain,(
  ( ! [X12,X11] : (sP12(X12,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X11))) | ~p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X11),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X11))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X11))))) ) | $spl5),
  inference(superposition,[],[f11813,f866])).
fof(f11813,plain,(
  ( ! [X178] : (sP12(X178,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))) | ~p(s(bool,X178))) ) | $spl5),
  inference(resolution,[],[f1086,f753])).
fof(f20811,plain,(
  ( ! [X12,X11] : (sP12('const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X11)),X12) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X11))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X11),s(fun(sK13,bool),sK14)))) | ~p(s(bool,X12))) ) | $spl422),
  inference(cnf_transformation,[],[f20811_D])).
fof(f20811_D,plain,(
  ( ! [X12,X11] : (sP12('const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X11)),X12) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X11))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X11),s(fun(sK13,bool),sK14)))) | ~p(s(bool,X12))) ) <=> ~$spl422),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl422])])).
fof(f19204,plain,(
  ( ! [X12,X11] : (sP12('const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X11)),X12) | ~p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X11),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X11))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl5),
  inference(subsumption_resolution,[],[f19192,f860])).
fof(f19192,plain,(
  ( ! [X12,X11] : (sP12('const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X11)),X12) | ~p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X11),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X11))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X11))))) ) | $spl5),
  inference(superposition,[],[f11814,f866])).
fof(f11814,plain,(
  ( ! [X179] : (sP12('const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)),X179) | ~p(s(bool,X179))) ) | $spl5),
  inference(resolution,[],[f1086,f756])).
fof(f20810,plain,(
  ( ! [X12,X11] : (~sP12(X12,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X11),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X11))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X11),s(fun(sK13,bool),sK14)))) | p(s(bool,X12))) ) | $spl420),
  inference(cnf_transformation,[],[f20810_D])).
fof(f20810_D,plain,(
  ( ! [X12,X11] : (~sP12(X12,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X11),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X11))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X11),s(fun(sK13,bool),sK14)))) | p(s(bool,X12))) ) <=> ~$spl420),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl420])])).
fof(f19449,plain,(
  ( ! [X12,X11] : (~sP12(X12,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X11),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X11),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X11))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl29),
  inference(subsumption_resolution,[],[f19426,f860])).
fof(f19426,plain,(
  ( ! [X12,X11] : (~sP12(X12,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X11),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X11),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X11))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X11))))) ) | $spl29),
  inference(superposition,[],[f12344,f866])).
fof(f12344,plain,(
  ( ! [X176] : (~sP12(X176,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X176))) ) | $spl29),
  inference(resolution,[],[f1105,f751])).
fof(f20809,plain,(
  ( ! [X12,X11] : (~sP12('const.pred_set.SUBSET_2'(s(fun(sK13,bool),X11),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X11))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X11),s(fun(sK13,bool),sK14)))) | p(s(bool,X12))) ) | $spl418),
  inference(cnf_transformation,[],[f20809_D])).
fof(f20809_D,plain,(
  ( ! [X12,X11] : (~sP12('const.pred_set.SUBSET_2'(s(fun(sK13,bool),X11),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X11))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X11),s(fun(sK13,bool),sK14)))) | p(s(bool,X12))) ) <=> ~$spl418),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl418])])).
fof(f19500,plain,(
  ( ! [X12,X11] : (~sP12('const.pred_set.SUBSET_2'(s(fun(sK13,bool),X11),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12) | p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X11),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X11))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl29),
  inference(subsumption_resolution,[],[f19477,f860])).
fof(f19477,plain,(
  ( ! [X12,X11] : (~sP12('const.pred_set.SUBSET_2'(s(fun(sK13,bool),X11),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12) | p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X11),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X11))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X11))))) ) | $spl29),
  inference(superposition,[],[f12345,f866])).
fof(f12345,plain,(
  ( ! [X177] : (~sP12('const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X177) | p(s(bool,X177))) ) | $spl29),
  inference(resolution,[],[f1105,f751])).
fof(f20808,plain,(
  ( ! [X12,X11] : (sP12(X12,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X11),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X11))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X11),s(fun(sK13,bool),sK14)))) | ~p(s(bool,X12))) ) | $spl416),
  inference(cnf_transformation,[],[f20808_D])).
fof(f20808_D,plain,(
  ( ! [X12,X11] : (sP12(X12,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X11),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X11))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X11),s(fun(sK13,bool),sK14)))) | ~p(s(bool,X12))) ) <=> ~$spl416),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl416])])).
fof(f19553,plain,(
  ( ! [X12,X11] : (sP12(X12,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X11),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X11),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X11))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl29),
  inference(subsumption_resolution,[],[f19530,f860])).
fof(f19530,plain,(
  ( ! [X12,X11] : (sP12(X12,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X11),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X11),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X11))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X11))))) ) | $spl29),
  inference(superposition,[],[f12346,f866])).
fof(f12346,plain,(
  ( ! [X178] : (sP12(X178,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X178))) ) | $spl29),
  inference(resolution,[],[f1105,f753])).
fof(f20807,plain,(
  ( ! [X12,X11] : (sP12('const.pred_set.SUBSET_2'(s(fun(sK13,bool),X11),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X11))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X11),s(fun(sK13,bool),sK14)))) | ~p(s(bool,X12))) ) | $spl414),
  inference(cnf_transformation,[],[f20807_D])).
fof(f20807_D,plain,(
  ( ! [X12,X11] : (sP12('const.pred_set.SUBSET_2'(s(fun(sK13,bool),X11),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X11))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X11),s(fun(sK13,bool),sK14)))) | ~p(s(bool,X12))) ) <=> ~$spl414),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl414])])).
fof(f19606,plain,(
  ( ! [X12,X11] : (sP12('const.pred_set.SUBSET_2'(s(fun(sK13,bool),X11),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12) | ~p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X11),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X11))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl29),
  inference(subsumption_resolution,[],[f19583,f860])).
fof(f19583,plain,(
  ( ! [X12,X11] : (sP12('const.pred_set.SUBSET_2'(s(fun(sK13,bool),X11),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12) | ~p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X11),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X11))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X11))))) ) | $spl29),
  inference(superposition,[],[f12347,f866])).
fof(f12347,plain,(
  ( ! [X179] : (sP12('const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X179) | ~p(s(bool,X179))) ) | $spl29),
  inference(resolution,[],[f1105,f756])).
fof(f20634,plain,(
  $false | $spl412),
  inference(resolution,[],[f20510,f902])).
fof(f20636,plain,(
  $false | ($spl10 | $spl412)),
  inference(resolution,[],[f20510,f1092])).
fof(f20637,plain,(
  $false | ($spl8 | $spl412)),
  inference(resolution,[],[f20510,f1089])).
fof(f20409,plain,(
  ( ! [X50,X49] : (s(bool,X50) != s(bool,'const.pred_set.FINITE_1'(s(fun(X49,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X50))) ) | $spl408),
  inference(cnf_transformation,[],[f20409_D])).
fof(f20409_D,plain,(
  ( ! [X50,X49] : (s(bool,X50) != s(bool,'const.pred_set.FINITE_1'(s(fun(X49,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X50))) ) <=> ~$spl408),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl408])])).
fof(f20623,plain,(
  ( ! [X808,X809,X807] : (p(s(bool,X808)) | p(s(bool,X809)) | s(bool,X809) != s(bool,'const.pred_set.FINITE_1'(s(fun(X807,bool),'const.pred_set.EMPTY_0')))) ) | $spl10),
  inference(subsumption_resolution,[],[f20299,f1092])).
fof(f20299,plain,(
  ( ! [X808,X809,X807] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | p(s(bool,X808)) | p(s(bool,X809)) | s(bool,X809) != s(bool,'const.pred_set.FINITE_1'(s(fun(X807,bool),'const.pred_set.EMPTY_0')))) ) | $spl10),
  inference(superposition,[],[f905,f16299])).
fof(f16299,plain,(
  ( ! [X419] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.pred_set.FINITE_1'(s(fun(X419,bool),'const.pred_set.EMPTY_0')))) ) | $spl10),
  inference(resolution,[],[f2314,f902])).
fof(f2314,plain,(
  ( ! [X204] : (~p(s(bool,X204)) | s(bool,X204) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl10),
  inference(resolution,[],[f1092,f774])).
fof(f20525,plain,(
  ( ! [X177,X175,X178,X176] : (p(s(bool,X176)) | ~p(s(bool,X177)) | s(bool,X176) != s(bool,'const.pred_set.FINITE_1'(s(fun(X175,bool),'const.pred_set.EMPTY_0'))) | s(bool,X177) != s(bool,X178)) ) | $spl10),
  inference(subsumption_resolution,[],[f20080,f1092])).
fof(f20080,plain,(
  ( ! [X177,X175,X178,X176] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | p(s(bool,X176)) | ~p(s(bool,X177)) | s(bool,X176) != s(bool,'const.pred_set.FINITE_1'(s(fun(X175,bool),'const.pred_set.EMPTY_0'))) | s(bool,X177) != s(bool,X178)) ) | $spl10),
  inference(superposition,[],[f630,f16299])).
fof(f20522,plain,(
  ( ! [X146,X144,X145,X143] : (p(s(bool,X144)) | ~p(s(bool,X145)) | p(s(bool,X146)) | s(bool,X146) != s(bool,'const.pred_set.FINITE_1'(s(fun(X143,bool),'const.pred_set.EMPTY_0')))) ) | $spl10),
  inference(subsumption_resolution,[],[f20072,f1092])).
fof(f20072,plain,(
  ( ! [X146,X144,X145,X143] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | p(s(bool,X144)) | ~p(s(bool,X145)) | p(s(bool,X146)) | s(bool,X146) != s(bool,'const.pred_set.FINITE_1'(s(fun(X143,bool),'const.pred_set.EMPTY_0')))) ) | $spl10),
  inference(superposition,[],[f629,f16299])).
fof(f20519,plain,(
  ( ! [X125,X123,X126,X124] : (p(s(bool,X124)) | p(s(bool,X125)) | s(bool,X124) != s(bool,X126) | s(bool,X125) != s(bool,'const.pred_set.FINITE_1'(s(fun(X123,bool),'const.pred_set.EMPTY_0')))) ) | $spl10),
  inference(subsumption_resolution,[],[f20067,f1092])).
fof(f20067,plain,(
  ( ! [X125,X123,X126,X124] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | p(s(bool,X124)) | p(s(bool,X125)) | s(bool,X124) != s(bool,X126) | s(bool,X125) != s(bool,'const.pred_set.FINITE_1'(s(fun(X123,bool),'const.pred_set.EMPTY_0')))) ) | $spl10),
  inference(superposition,[],[f628,f16299])).
fof(f20517,plain,(
  ( ! [X97,X95,X98,X96] : (p(s(bool,X96)) | ~p(s(bool,X97)) | s(bool,X96) != s(bool,'const.pred_set.FINITE_1'(s(fun(X95,bool),'const.pred_set.EMPTY_0'))) | s(bool,X97) != s(bool,X98)) ) | $spl10),
  inference(subsumption_resolution,[],[f20060,f1092])).
fof(f20060,plain,(
  ( ! [X97,X95,X98,X96] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | p(s(bool,X96)) | ~p(s(bool,X97)) | s(bool,X96) != s(bool,'const.pred_set.FINITE_1'(s(fun(X95,bool),'const.pred_set.EMPTY_0'))) | s(bool,X97) != s(bool,X98)) ) | $spl10),
  inference(superposition,[],[f626,f16299])).
fof(f20514,plain,(
  ( ! [X70,X68,X69,X67] : (p(s(bool,X68)) | p(s(bool,X69)) | s(bool,X68) != s(bool,X70) | s(bool,X69) != s(bool,'const.pred_set.FINITE_1'(s(fun(X67,bool),'const.pred_set.EMPTY_0')))) ) | $spl10),
  inference(subsumption_resolution,[],[f20053,f1092])).
fof(f20053,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | p(s(bool,X68)) | p(s(bool,X69)) | s(bool,X68) != s(bool,X70) | s(bool,X69) != s(bool,'const.pred_set.FINITE_1'(s(fun(X67,bool),'const.pred_set.EMPTY_0')))) ) | $spl10),
  inference(superposition,[],[f624,f16299])).
fof(f20510,plain,(
  ( ! [X30,X29] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X29,bool),X30))))) ) | $spl412),
  inference(cnf_transformation,[],[f20510_D])).
fof(f20510_D,plain,(
  ( ! [X30,X29] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X29,bool),X30))))) ) <=> ~$spl412),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl412])])).
fof(f8815,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | $spl362),
  inference(cnf_transformation,[],[f8815_D])).
fof(f8815_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) <=> ~$spl362),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl362])])).
fof(f20509,plain,(
  ( ! [X30,X29] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X29,bool),X30))))) ) | $spl10),
  inference(forward_demodulation,[],[f20508,f901])).
fof(f901,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(cnf_transformation,[],[f411])).
fof(f411,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.EMPTY_0')))),
  inference(rectify,[],[f132])).
fof(f132,axiom,(
  ! [X6] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(X6,bool),'const.pred_set.EMPTY_0')))),
  file('thm.pred_set.SUBSET_DIFF_EMPTY','thm.pred_set.CARD_DEF|split|0')).
fof(f20508,plain,(
  ( ! [X30,X29] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(X29,bool),'const.pred_set.EMPTY_0'))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X29,bool),X30))))) ) | $spl10),
  inference(forward_demodulation,[],[f20507,f891])).
fof(f891,plain,(
  ( ! [X0,X1] : (s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X0,bool),X1)))) )),
  inference(cnf_transformation,[],[f400])).
fof(f400,plain,(
  ! [X0,X1] : s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X0,bool),X1)))),
  inference(flattening,[],[f399])).
fof(f399,plain,(
  ! [X0] : ! [X1] : s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X0,bool),X1)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X6] : ! [X7] : s(fun(X6,bool),'const.pred_set.EMPTY_0') = s(fun(X6,bool),'const.pred_set.DIFF_2'(s(fun(X6,bool),'const.pred_set.EMPTY_0'),s(fun(X6,bool),X7)))),
  file('thm.pred_set.SUBSET_DIFF_EMPTY','thm.pred_set.EMPTY_DIFF')).
fof(f20507,plain,(
  ( ! [X30,X29] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(X29,bool),'const.pred_set.DIFF_2'(s(fun(X29,bool),'const.pred_set.EMPTY_0'),s(fun(X29,bool),X30))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X29,bool),X30))))) ) | $spl10),
  inference(forward_demodulation,[],[f20506,f901])).
fof(f20506,plain,(
  ( ! [X30,X29] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(X29,bool),'const.pred_set.DIFF_2'(s(fun(X29,bool),'const.pred_set.EMPTY_0'),s(fun(X29,bool),X30))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X29,bool),'const.pred_set.EMPTY_0'))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X29,bool),'const.pred_set.EMPTY_0'))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X29,bool),X30))))) ) | $spl10),
  inference(forward_demodulation,[],[f20505,f900])).
fof(f900,plain,(
  ( ! [X0,X1] : (s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X0,bool),X1)))) )),
  inference(cnf_transformation,[],[f410])).
fof(f410,plain,(
  ! [X0,X1] : s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X0,bool),X1)))),
  inference(flattening,[],[f409])).
fof(f409,plain,(
  ! [X0] : ! [X1] : s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X0,bool),X1)))),
  inference(rectify,[],[f138])).
fof(f138,axiom,(
  ! [X6] : ! [X7] : s(fun(X6,bool),'const.pred_set.EMPTY_0') = s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),'const.pred_set.EMPTY_0'),s(fun(X6,bool),X7)))),
  file('thm.pred_set.SUBSET_DIFF_EMPTY','thm.pred_set.INTER_EMPTY|split|0')).
fof(f20505,plain,(
  ( ! [X30,X29] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(X29,bool),'const.pred_set.DIFF_2'(s(fun(X29,bool),'const.pred_set.EMPTY_0'),s(fun(X29,bool),X30))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X29,bool),'const.pred_set.EMPTY_0'))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X29,bool),'const.pred_set.INTER_2'(s(fun(X29,bool),'const.pred_set.EMPTY_0'),s(fun(X29,bool),X30))))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X29,bool),X30))))) ) | $spl10),
  inference(subsumption_resolution,[],[f20039,f1092])).
fof(f20039,plain,(
  ( ! [X30,X29] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(X29,bool),'const.pred_set.DIFF_2'(s(fun(X29,bool),'const.pred_set.EMPTY_0'),s(fun(X29,bool),X30))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X29,bool),'const.pred_set.EMPTY_0'))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X29,bool),'const.pred_set.INTER_2'(s(fun(X29,bool),'const.pred_set.EMPTY_0'),s(fun(X29,bool),X30))))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X29,bool),X30))))) ) | $spl10),
  inference(superposition,[],[f887,f16299])).
fof(f887,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f453])).
fof(f453,plain,(
  ! [X0,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) | ! [X2] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))))))))),
  inference(ennf_transformation,[],[f392])).
fof(f392,plain,(
  ! [X0,X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) => ! [X2] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) => s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))))))))),
  inference(flattening,[],[f391])).
fof(f391,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) => ! [X2] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) => s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X2))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))))))))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X6] : ! [X8] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X8)))) => ! [X7] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X7)))) => s('type.num.num','const.pred_set.CARD_1'(s(fun(X6,bool),'const.pred_set.DIFF_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X6,bool),X7))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))))))))),
  file('thm.pred_set.SUBSET_DIFF_EMPTY','thm.pred_set.CARD_DIFF')).
fof(f8498,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | $spl304),
  inference(cnf_transformation,[],[f8498_D])).
fof(f8498_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) <=> ~$spl304),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl304])])).
fof(f20504,plain,(
  ( ! [X26,X25] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X25,bool),X26))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X25,bool),X26))))) ) | $spl410),
  inference(cnf_transformation,[],[f20504_D])).
fof(f20504_D,plain,(
  ( ! [X26,X25] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X25,bool),X26))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X25,bool),X26))))) ) <=> ~$spl410),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl410])])).
fof(f20503,plain,(
  ( ! [X26,X25] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X25,bool),X26))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X25,bool),X26))))) ) | $spl10),
  inference(forward_demodulation,[],[f20502,f901])).
fof(f20502,plain,(
  ( ! [X26,X25] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X25,bool),X26))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X25,bool),'const.pred_set.EMPTY_0')))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X25,bool),X26))))) ) | $spl10),
  inference(forward_demodulation,[],[f20501,f901])).
fof(f20501,plain,(
  ( ! [X26,X25] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(X25,bool),'const.pred_set.EMPTY_0')))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X25,bool),X26))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X25,bool),'const.pred_set.EMPTY_0')))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X25,bool),X26))))) ) | $spl10),
  inference(forward_demodulation,[],[f20500,f891])).
fof(f20500,plain,(
  ( ! [X26,X25] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(X25,bool),'const.pred_set.DIFF_2'(s(fun(X25,bool),'const.pred_set.EMPTY_0'),s(fun(X25,bool),X26)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X25,bool),X26))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X25,bool),'const.pred_set.EMPTY_0')))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X25,bool),X26))))) ) | $spl10),
  inference(subsumption_resolution,[],[f20037,f1092])).
fof(f20037,plain,(
  ( ! [X26,X25] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(X25,bool),'const.pred_set.DIFF_2'(s(fun(X25,bool),'const.pred_set.EMPTY_0'),s(fun(X25,bool),X26)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X25,bool),X26))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X25,bool),'const.pred_set.EMPTY_0')))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X25,bool),X26))))) ) | $spl10),
  inference(superposition,[],[f886,f16299])).
fof(f886,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X2)))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f452])).
fof(f452,plain,(
  ! [X0,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) | ! [X2] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))))))))))),
  inference(flattening,[],[f451])).
fof(f451,plain,(
  ! [X0,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) | ! [X2] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) | (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))))))))))),
  inference(ennf_transformation,[],[f390])).
fof(f390,plain,(
  ! [X0,X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) => ! [X2] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) => (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X2)))))) => p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))))))))))),
  inference(flattening,[],[f389])).
fof(f389,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) => ! [X2] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) => (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X2)))))) => p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))))))))))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X6] : ! [X8] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X8)))) => ! [X7] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X7)))) => (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X6,bool),X8))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X6,bool),X7)))))) => p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(X6,bool),'const.pred_set.DIFF_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))))))))))),
  file('thm.pred_set.SUBSET_DIFF_EMPTY','thm.pred_set.LESS_CARD_DIFF')).
fof(f20478,plain,(
  ( ! [X798,X799,X797] : (s(bool,X798) != s(bool,'const.pred_set.FINITE_1'(s(fun(X797,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X799)) | p(s(bool,X798))) ) | $spl10),
  inference(subsumption_resolution,[],[f19945,f1092])).
fof(f19945,plain,(
  ( ! [X798,X799,X797] : (s(bool,X798) != s(bool,'const.pred_set.FINITE_1'(s(fun(X797,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X799)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | p(s(bool,X798))) ) | $spl10),
  inference(superposition,[],[f905,f16299])).
fof(f20421,plain,(
  ( ! [X167,X165,X166,X168] : (s(bool,X166) != s(bool,'const.pred_set.FINITE_1'(s(fun(X165,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X166)) | ~p(s(bool,X167)) | s(bool,X167) != s(bool,X168)) ) | $spl10),
  inference(subsumption_resolution,[],[f19727,f1092])).
fof(f19727,plain,(
  ( ! [X167,X165,X166,X168] : (s(bool,X166) != s(bool,'const.pred_set.FINITE_1'(s(fun(X165,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X166)) | ~p(s(bool,X167)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | s(bool,X167) != s(bool,X168)) ) | $spl10),
  inference(superposition,[],[f630,f16299])).
fof(f20417,plain,(
  ( ! [X132,X130,X131,X129] : (s(bool,X130) != s(bool,'const.pred_set.FINITE_1'(s(fun(X129,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X131)) | ~p(s(bool,X132)) | p(s(bool,X130))) ) | $spl10),
  inference(subsumption_resolution,[],[f19718,f1092])).
fof(f19718,plain,(
  ( ! [X132,X130,X131,X129] : (s(bool,X130) != s(bool,'const.pred_set.FINITE_1'(s(fun(X129,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X131)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,X132)) | p(s(bool,X130))) ) | $spl10),
  inference(superposition,[],[f629,f16299])).
fof(f20415,plain,(
  ( ! [X103,X101,X102,X104] : (s(bool,X102) != s(bool,'const.pred_set.FINITE_1'(s(fun(X101,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X103)) | p(s(bool,X102)) | s(bool,X103) != s(bool,X104)) ) | $spl10),
  inference(subsumption_resolution,[],[f19711,f1092])).
fof(f19711,plain,(
  ( ! [X103,X101,X102,X104] : (s(bool,X102) != s(bool,'const.pred_set.FINITE_1'(s(fun(X101,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X103)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | p(s(bool,X102)) | s(bool,X103) != s(bool,X104)) ) | $spl10),
  inference(superposition,[],[f628,f16299])).
fof(f20412,plain,(
  ( ! [X83,X81,X84,X82] : (s(bool,X82) != s(bool,'const.pred_set.FINITE_1'(s(fun(X81,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X82)) | ~p(s(bool,X83)) | s(bool,X83) != s(bool,X84)) ) | $spl10),
  inference(subsumption_resolution,[],[f19706,f1092])).
fof(f19706,plain,(
  ( ! [X83,X81,X84,X82] : (s(bool,X82) != s(bool,'const.pred_set.FINITE_1'(s(fun(X81,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X82)) | ~p(s(bool,X83)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | s(bool,X83) != s(bool,X84)) ) | $spl10),
  inference(superposition,[],[f626,f16299])).
fof(f20408,plain,(
  ( ! [X52,X50,X51,X49] : (s(bool,X50) != s(bool,'const.pred_set.FINITE_1'(s(fun(X49,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X51)) | p(s(bool,X50)) | s(bool,X51) != s(bool,X52)) ) | $spl10),
  inference(subsumption_resolution,[],[f19698,f1092])).
fof(f19698,plain,(
  ( ! [X52,X50,X51,X49] : (s(bool,X50) != s(bool,'const.pred_set.FINITE_1'(s(fun(X49,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X51)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | p(s(bool,X50)) | s(bool,X51) != s(bool,X52)) ) | $spl10),
  inference(superposition,[],[f624,f16299])).
fof(f18919,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | ($spl8 | $spl10)),
  inference(resolution,[],[f7653,f16219])).
fof(f16219,plain,(
  ( ! [X189,X190,X188] : (sP7(X189,X188,X190) | s(bool,X188) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl10),
  inference(resolution,[],[f2314,f695])).
fof(f18889,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | ($spl8 | $spl10)),
  inference(resolution,[],[f7652,f16220])).
fof(f16220,plain,(
  ( ! [X191,X192,X193] : (sP7(X191,X192,X193) | s(bool,X191) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl10),
  inference(resolution,[],[f2314,f696])).
fof(f18855,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | ($spl8 | $spl10)),
  inference(resolution,[],[f7650,f16216])).
fof(f16216,plain,(
  ( ! [X180,X181,X179] : (sP8(X180,X179,X181) | s(bool,X179) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl10),
  inference(resolution,[],[f2314,f691])).
fof(f18825,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | ($spl8 | $spl10)),
  inference(resolution,[],[f7649,f16217])).
fof(f16217,plain,(
  ( ! [X182,X184,X183] : (sP8(X182,X183,X184) | s(bool,X182) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl10),
  inference(resolution,[],[f2314,f692])).
fof(f18574,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | ($spl8 | $spl10)),
  inference(resolution,[],[f7637,f16203])).
fof(f16203,plain,(
  ( ! [X142,X140,X141] : (sP4(X141,X140,X142) | s(bool,X140) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl10),
  inference(resolution,[],[f2314,f674])).
fof(f18543,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | ($spl8 | $spl10)),
  inference(resolution,[],[f7636,f16204])).
fof(f16204,plain,(
  ( ! [X144,X145,X143] : (sP4(X143,X144,X145) | s(bool,X143) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl10),
  inference(resolution,[],[f2314,f675])).
fof(f18470,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | p(s(bool,X3)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | ($spl8 | $spl10)),
  inference(resolution,[],[f7620,f16190])).
fof(f16190,plain,(
  ( ! [X103,X101,X102] : (sP2(X101,X102,X103) | s(bool,X101) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl10),
  inference(resolution,[],[f2314,f658])).
fof(f18328,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | ($spl8 | $spl10)),
  inference(resolution,[],[f7607,f16173])).
fof(f16173,plain,(
  ( ! [X52,X50,X51] : (sP0(X51,X52,X50) | s(bool,X50) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl10),
  inference(resolution,[],[f2314,f639])).
fof(f18297,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | p(s(bool,X11)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | ($spl8 | $spl10)),
  inference(resolution,[],[f7606,f16174])).
fof(f16174,plain,(
  ( ! [X54,X55,X53] : (sP0(X53,X54,X55) | s(bool,X53) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl10),
  inference(resolution,[],[f2314,f641])).
fof(f18177,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | ($spl8 | $spl10)),
  inference(resolution,[],[f2279,f16508])).
fof(f2279,plain,(
  ( ! [X144,X145] : (~sP7(X144,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X145) | p(s(bool,X145)) | ~p(s(bool,X144))) ) | $spl10),
  inference(resolution,[],[f1092,f693])).
fof(f18147,plain,(
  ( ! [X14,X15] : (p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | ($spl8 | $spl10)),
  inference(resolution,[],[f2278,f16509])).
fof(f2278,plain,(
  ( ! [X142,X143] : (~sP7('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X142,X143) | p(s(bool,X143)) | ~p(s(bool,X142))) ) | $spl10),
  inference(resolution,[],[f1092,f693])).
fof(f18113,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | ($spl8 | $spl10)),
  inference(resolution,[],[f2276,f16505])).
fof(f2276,plain,(
  ( ! [X138,X139] : (~sP8(X138,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X139) | p(s(bool,X139)) | ~p(s(bool,X138))) ) | $spl10),
  inference(resolution,[],[f1092,f689])).
fof(f18083,plain,(
  ( ! [X17,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | ($spl8 | $spl10)),
  inference(resolution,[],[f2275,f16506])).
fof(f2275,plain,(
  ( ! [X136,X137] : (~sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X136,X137) | p(s(bool,X137)) | ~p(s(bool,X136))) ) | $spl10),
  inference(resolution,[],[f1092,f689])).
fof(f17814,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | ($spl8 | $spl10)),
  inference(resolution,[],[f2263,f16492])).
fof(f2263,plain,(
  ( ! [X112,X113] : (~sP4(X112,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X113) | p(s(bool,X113)) | ~p(s(bool,X112))) ) | $spl10),
  inference(resolution,[],[f1092,f672])).
fof(f17783,plain,(
  ( ! [X14,X15] : (p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | ($spl8 | $spl10)),
  inference(resolution,[],[f2262,f16493])).
fof(f2262,plain,(
  ( ! [X111,X110] : (~sP4('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X110,X111) | p(s(bool,X111)) | ~p(s(bool,X110))) ) | $spl10),
  inference(resolution,[],[f1092,f672])).
fof(f17707,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | ($spl8 | $spl10)),
  inference(resolution,[],[f2246,f16479])).
fof(f2246,plain,(
  ( ! [X78,X79] : (~sP2('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X78,X79) | p(s(bool,X79)) | p(s(bool,X78))) ) | $spl10),
  inference(resolution,[],[f1092,f655])).
fof(f17538,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | ($spl8 | $spl10)),
  inference(resolution,[],[f2233,f16462])).
fof(f2233,plain,(
  ( ! [X52,X53] : (~sP0(X52,X53,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | p(s(bool,X53)) | ~p(s(bool,X52))) ) | $spl10),
  inference(resolution,[],[f1092,f638])).
fof(f17505,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | ($spl8 | $spl10)),
  inference(resolution,[],[f2232,f16463])).
fof(f2232,plain,(
  ( ! [X50,X51] : (~sP0('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X50,X51) | ~p(s(bool,X51)) | p(s(bool,X50))) ) | $spl10),
  inference(resolution,[],[f1092,f638])).
fof(f16650,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) | $spl404),
  inference(cnf_transformation,[],[f16650_D])).
fof(f16650_D,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) <=> ~$spl404),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl404])])).
fof(f16553,plain,(
  ( ! [X318,X317,X319] : (s(bool,X317) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) | p(s(bool,X318)) | ~p(s(bool,X319)) | s(bool,X317) != s(bool,X319)) ) | $spl8),
  inference(resolution,[],[f7688,f905])).
fof(f16455,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) ) | $spl8),
  inference(resolution,[],[f7688,f630])).
fof(f16453,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,X20) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X20) != s(bool,X22)) ) | $spl8),
  inference(resolution,[],[f7688,f629])).
fof(f16451,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl8),
  inference(resolution,[],[f7688,f628])).
fof(f16450,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) ) | $spl8),
  inference(resolution,[],[f7688,f626])).
fof(f16448,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl8),
  inference(resolution,[],[f7688,f624])).
fof(f16264,plain,(
  ( ! [X318,X317,X319] : (s(bool,X317) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | p(s(bool,X318)) | ~p(s(bool,X319)) | s(bool,X317) != s(bool,X319)) ) | $spl10),
  inference(resolution,[],[f2314,f905])).
fof(f16166,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) ) | $spl10),
  inference(resolution,[],[f2314,f630])).
fof(f16164,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,X20) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X20) != s(bool,X22)) ) | $spl10),
  inference(resolution,[],[f2314,f629])).
fof(f16162,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl10),
  inference(resolution,[],[f2314,f628])).
fof(f16161,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) ) | $spl10),
  inference(resolution,[],[f2314,f626])).
fof(f16159,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl10),
  inference(resolution,[],[f2314,f624])).
fof(f15909,plain,(
  $false | ($spl8 | $spl48)),
  inference(subsumption_resolution,[],[f15889,f1089])).
fof(f15889,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) | $spl48),
  inference(equality_resolution,[],[f1315])).
fof(f15908,plain,(
  ( ! [X26,X27] : (~p(s(bool,X27)) | ~p(s(bool,X26))) ) | ($spl8 | $spl48)),
  inference(subsumption_resolution,[],[f15907,f1089])).
fof(f15907,plain,(
  ( ! [X26,X27] : (~p(s(bool,X27)) | ~p(s(bool,X26)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))))) ) | $spl48),
  inference(subsumption_resolution,[],[f15886,f774])).
fof(f15886,plain,(
  ( ! [X26,X27] : (s(bool,X26) != s(bool,X27) | ~p(s(bool,X27)) | ~p(s(bool,X26)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))))) ) | $spl48),
  inference(superposition,[],[f1315,f782])).
fof(f782,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f434])).
fof(f15906,plain,(
  ( ! [X24,X25] : (~p(s(bool,X25)) | ~p(s(bool,X24))) ) | ($spl8 | $spl48)),
  inference(subsumption_resolution,[],[f15905,f1089])).
fof(f15905,plain,(
  ( ! [X24,X25] : (~p(s(bool,X25)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) | ~p(s(bool,X24))) ) | $spl48),
  inference(subsumption_resolution,[],[f15885,f774])).
fof(f15885,plain,(
  ( ! [X24,X25] : (s(bool,X24) != s(bool,X25) | ~p(s(bool,X25)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) | ~p(s(bool,X24))) ) | $spl48),
  inference(superposition,[],[f1315,f782])).
fof(f15904,plain,(
  ( ! [X19,X18] : (~p(s(bool,X19)) | ~p(s(bool,X18))) ) | ($spl8 | $spl48)),
  inference(subsumption_resolution,[],[f15903,f1089])).
fof(f15903,plain,(
  ( ! [X19,X18] : (~p(s(bool,X19)) | ~p(s(bool,X18)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))))) ) | $spl48),
  inference(subsumption_resolution,[],[f15882,f774])).
fof(f15882,plain,(
  ( ! [X19,X18] : (s(bool,X18) != s(bool,X19) | ~p(s(bool,X19)) | ~p(s(bool,X18)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))))) ) | $spl48),
  inference(superposition,[],[f1315,f774])).
fof(f15902,plain,(
  ( ! [X17,X16] : (~p(s(bool,X17)) | ~p(s(bool,X16))) ) | ($spl8 | $spl48)),
  inference(subsumption_resolution,[],[f15901,f1089])).
fof(f15901,plain,(
  ( ! [X17,X16] : (~p(s(bool,X17)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) | ~p(s(bool,X16))) ) | $spl48),
  inference(subsumption_resolution,[],[f15881,f774])).
fof(f15881,plain,(
  ( ! [X17,X16] : (s(bool,X16) != s(bool,X17) | ~p(s(bool,X17)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) | ~p(s(bool,X16))) ) | $spl48),
  inference(superposition,[],[f1315,f774])).
fof(f15900,plain,(
  ( ! [X11] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X11),s(fun(sK13,bool),sK15)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X11)))) ) | $spl400),
  inference(cnf_transformation,[],[f15900_D])).
fof(f15900_D,plain,(
  ( ! [X11] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X11),s(fun(sK13,bool),sK15)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X11)))) ) <=> ~$spl400),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl400])])).
fof(f15899,plain,(
  ( ! [X12,X11] : (~p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X11),s(fun(sK13,bool),sK15)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X11)))) ) | ($spl8 | $spl48)),
  inference(subsumption_resolution,[],[f15898,f1089])).
fof(f15898,plain,(
  ( ! [X12,X11] : (~p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X11),s(fun(sK13,bool),sK15)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X11))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))))) ) | $spl48),
  inference(subsumption_resolution,[],[f15897,f860])).
fof(f15897,plain,(
  ( ! [X12,X11] : (~p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X11),s(fun(sK13,bool),sK15)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X11))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X11))))) ) | $spl48),
  inference(subsumption_resolution,[],[f15869,f774])).
fof(f15869,plain,(
  ( ! [X12,X11] : (s(bool,X12) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X11))) | ~p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X11),s(fun(sK13,bool),sK15)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X11))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X11))))) ) | $spl48),
  inference(superposition,[],[f1315,f866])).
fof(f15896,plain,(
  ( ! [X9] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X9)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X9)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X9)))) ) | $spl398),
  inference(cnf_transformation,[],[f15896_D])).
fof(f15896_D,plain,(
  ( ! [X9] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X9)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X9)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X9)))) ) <=> ~$spl398),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl398])])).
fof(f15895,plain,(
  ( ! [X10,X9] : (~p(s(bool,X10)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X9)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X9))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X9))))) ) | $spl48),
  inference(subsumption_resolution,[],[f15894,f860])).
fof(f15894,plain,(
  ( ! [X10,X9] : (~p(s(bool,X10)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X9)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X9))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X9)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))))) ) | $spl48),
  inference(subsumption_resolution,[],[f15868,f774])).
fof(f15868,plain,(
  ( ! [X10,X9] : (s(bool,X10) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X9))) | ~p(s(bool,X10)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X9)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X9))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X9)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))))) ) | $spl48),
  inference(superposition,[],[f1315,f866])).
fof(f14501,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,X6) != s(bool,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),sK14)))) ) | $spl396),
  inference(cnf_transformation,[],[f14501_D])).
fof(f14501_D,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,X6) != s(bool,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),sK14)))) ) <=> ~$spl396),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl396])])).
fof(f14462,plain,(
  ( ! [X216,X215] : (p(s(bool,X215)) | p(s(bool,X216)) | s(bool,X216) != s(bool,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),sK14)))) ) | $spl29),
  inference(resolution,[],[f12255,f905])).
fof(f14364,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),sK14))) | s(bool,X27) != s(bool,X28)) ) | $spl29),
  inference(resolution,[],[f12255,f630])).
fof(f14361,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X19) != s(bool,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),sK14)))) ) | $spl29),
  inference(resolution,[],[f12255,f629])).
fof(f14360,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),sK14)))) ) | $spl29),
  inference(resolution,[],[f12255,f628])).
fof(f14359,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),sK14))) | s(bool,X12) != s(bool,X13)) ) | $spl29),
  inference(resolution,[],[f12255,f626])).
fof(f14357,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK14,sK13)),s(fun(sK13,bool),sK14)))) ) | $spl29),
  inference(resolution,[],[f12255,f624])).
fof(f14348,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,X6) != s(bool,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK15,sK13)),s(fun(sK13,bool),sK15)))) ) | $spl394),
  inference(cnf_transformation,[],[f14348_D])).
fof(f14348_D,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,X6) != s(bool,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK15,sK13)),s(fun(sK13,bool),sK15)))) ) <=> ~$spl394),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl394])])).
fof(f14309,plain,(
  ( ! [X216,X215] : (p(s(bool,X215)) | p(s(bool,X216)) | s(bool,X216) != s(bool,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK15,sK13)),s(fun(sK13,bool),sK15)))) ) | $spl27),
  inference(resolution,[],[f12007,f905])).
fof(f14211,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK15,sK13)),s(fun(sK13,bool),sK15))) | s(bool,X27) != s(bool,X28)) ) | $spl27),
  inference(resolution,[],[f12007,f630])).
fof(f14208,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X19) != s(bool,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK15,sK13)),s(fun(sK13,bool),sK15)))) ) | $spl27),
  inference(resolution,[],[f12007,f629])).
fof(f14207,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK15,sK13)),s(fun(sK13,bool),sK15)))) ) | $spl27),
  inference(resolution,[],[f12007,f628])).
fof(f14206,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK15,sK13)),s(fun(sK13,bool),sK15))) | s(bool,X12) != s(bool,X13)) ) | $spl27),
  inference(resolution,[],[f12007,f626])).
fof(f14204,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,'const.bool.IN_2'(s(sK13,sK39('const.pred_set.EMPTY_0',sK15,sK13)),s(fun(sK13,bool),sK15)))) ) | $spl27),
  inference(resolution,[],[f12007,f624])).
fof(f14194,plain,(
  ( ! [X6] : (~p(s(bool,X6)) | s(bool,X6) != s(bool,'const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),sK14)))) ) | $spl392),
  inference(cnf_transformation,[],[f14194_D])).
fof(f14194_D,plain,(
  ( ! [X6] : (~p(s(bool,X6)) | s(bool,X6) != s(bool,'const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),sK14)))) ) <=> ~$spl392),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl392])])).
fof(f14157,plain,(
  ( ! [X218,X217] : (p(s(bool,X217)) | ~p(s(bool,X218)) | s(bool,X218) != s(bool,'const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),sK14)))) ) | $spl5),
  inference(resolution,[],[f11723,f905])).
fof(f14059,plain,(
  ( ! [X28,X26,X27] : (~p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),sK14))) | s(bool,X26) != s(bool,X28)) ) | $spl5),
  inference(resolution,[],[f11723,f630])).
fof(f14057,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),sK14)))) ) | $spl5),
  inference(resolution,[],[f11723,f629])).
fof(f14055,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),sK14)))) ) | $spl5),
  inference(resolution,[],[f11723,f628])).
fof(f14054,plain,(
  ( ! [X12,X13,X11] : (~p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X12) != s(bool,'const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),sK14))) | s(bool,X11) != s(bool,X13)) ) | $spl5),
  inference(resolution,[],[f11723,f626])).
fof(f14052,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,'const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),sK14)))) ) | $spl5),
  inference(resolution,[],[f11723,f624])).
fof(f14043,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,X6) != s(bool,'const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),sK15)))) ) | $spl390),
  inference(cnf_transformation,[],[f14043_D])).
fof(f14043_D,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,X6) != s(bool,'const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),sK15)))) ) <=> ~$spl390),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl390])])).
fof(f14004,plain,(
  ( ! [X216,X215] : (p(s(bool,X215)) | p(s(bool,X216)) | s(bool,X216) != s(bool,'const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),sK15)))) ) | $spl5),
  inference(resolution,[],[f11722,f905])).
fof(f13906,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,'const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),sK15))) | s(bool,X27) != s(bool,X28)) ) | $spl5),
  inference(resolution,[],[f11722,f630])).
fof(f13903,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X19) != s(bool,'const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),sK15)))) ) | $spl5),
  inference(resolution,[],[f11722,f629])).
fof(f13902,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),sK15)))) ) | $spl5),
  inference(resolution,[],[f11722,f628])).
fof(f13901,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,'const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),sK15))) | s(bool,X12) != s(bool,X13)) ) | $spl5),
  inference(resolution,[],[f11722,f626])).
fof(f13899,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,'const.bool.IN_2'(s(sK13,sK39(sK14,sK15,sK13)),s(fun(sK13,bool),sK15)))) ) | $spl5),
  inference(resolution,[],[f11722,f624])).
fof(f12416,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | s(bool,X3) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) | $spl388),
  inference(cnf_transformation,[],[f12416_D])).
fof(f12416_D,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | s(bool,X3) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) <=> ~$spl388),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl388])])).
fof(f12366,plain,(
  ( ! [X215,X214] : (p(s(bool,X214)) | ~p(s(bool,X215)) | s(bool,X215) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) | $spl29),
  inference(resolution,[],[f1105,f905])).
fof(f12268,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s(bool,X23) != s(bool,X25)) ) | $spl29),
  inference(resolution,[],[f1105,f630])).
fof(f12266,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) | $spl29),
  inference(resolution,[],[f1105,f629])).
fof(f12264,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) | $spl29),
  inference(resolution,[],[f1105,f628])).
fof(f12263,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X9) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s(bool,X8) != s(bool,X10)) ) | $spl29),
  inference(resolution,[],[f1105,f626])).
fof(f12261,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) | $spl29),
  inference(resolution,[],[f1105,f624])).
fof(f7552,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | s(bool,X3) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) | $spl288),
  inference(cnf_transformation,[],[f7552_D])).
fof(f7552_D,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | s(bool,X3) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) <=> ~$spl288),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl288])])).
fof(f12118,plain,(
  ( ! [X215,X214] : (p(s(bool,X214)) | ~p(s(bool,X215)) | s(bool,X215) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) | $spl27),
  inference(resolution,[],[f1104,f905])).
fof(f12020,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s(bool,X23) != s(bool,X25)) ) | $spl27),
  inference(resolution,[],[f1104,f630])).
fof(f12018,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) | $spl27),
  inference(resolution,[],[f1104,f629])).
fof(f12016,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) | $spl27),
  inference(resolution,[],[f1104,f628])).
fof(f12015,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X9) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s(bool,X8) != s(bool,X10)) ) | $spl27),
  inference(resolution,[],[f1104,f626])).
fof(f12013,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) | $spl27),
  inference(resolution,[],[f1104,f624])).
fof(f1104,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) | $spl27),
  inference(cnf_transformation,[],[f1104_D])).
fof(f1104_D,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f6457,plain,(
  ( ! [X12,X13] : (sP3(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | ~p(s(bool,X12))) ) | $spl272),
  inference(cnf_transformation,[],[f6457_D])).
fof(f6457_D,plain,(
  ( ! [X12,X13] : (sP3(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | ~p(s(bool,X12))) ) <=> ~$spl272),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl272])])).
fof(f11993,plain,(
  ( ! [X12,X13] : (sP3(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | ~p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl8),
  inference(superposition,[],[f7625,f864])).
fof(f864,plain,(
  ( ! [X0,X1] : (s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(cnf_transformation,[],[f611])).
fof(f611,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1)) & (s(fun(X0,bool),'const.pred_set.EMPTY_0') != s(fun(X0,bool),X1) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))))),
  inference(nnf_transformation,[],[f354])).
fof(f354,plain,(
  ! [X0,X1] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) <=> s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1))),
  inference(flattening,[],[f353])).
fof(f353,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) <=> s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1))),
  inference(rectify,[],[f96])).
fof(f96,axiom,(
  ! [X6] : ! [X7] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),'const.pred_set.EMPTY_0')))) <=> s(fun(X6,bool),'const.pred_set.EMPTY_0') = s(fun(X6,bool),X7))),
  file('thm.pred_set.SUBSET_DIFF_EMPTY','thm.pred_set.SUBSET_EMPTY')).
fof(f7625,plain,(
  ( ! [X88,X89] : (sP3(X88,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)),X89) | ~p(s(bool,X88))) ) | $spl8),
  inference(resolution,[],[f1089,f661])).
fof(f6435,plain,(
  ( ! [X12,X13] : (sP3('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | ~p(s(bool,X13))) ) | $spl270),
  inference(cnf_transformation,[],[f6435_D])).
fof(f6435_D,plain,(
  ( ! [X12,X13] : (sP3('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | ~p(s(bool,X13))) ) <=> ~$spl270),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl270])])).
fof(f11970,plain,(
  ( ! [X12,X13] : (sP3('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | ~p(s(bool,X13)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl8),
  inference(superposition,[],[f7624,f864])).
fof(f7624,plain,(
  ( ! [X87,X86] : (sP3('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)),X86,X87) | ~p(s(bool,X87))) ) | $spl8),
  inference(resolution,[],[f1089,f660])).
fof(f6412,plain,(
  ( ! [X12,X13] : (sP3(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X12))) ) | $spl268),
  inference(cnf_transformation,[],[f6412_D])).
fof(f6412_D,plain,(
  ( ! [X12,X13] : (sP3(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X12))) ) <=> ~$spl268),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl268])])).
fof(f11947,plain,(
  ( ! [X12,X13] : (sP3(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl8),
  inference(superposition,[],[f7623,f864])).
fof(f7623,plain,(
  ( ! [X85,X84] : (sP3(X84,X85,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) | ~p(s(bool,X84))) ) | $spl8),
  inference(resolution,[],[f1089,f660])).
fof(f6390,plain,(
  ( ! [X12,X13] : (~sP3(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | p(s(bool,X12))) ) | $spl266),
  inference(cnf_transformation,[],[f6390_D])).
fof(f6390_D,plain,(
  ( ! [X12,X13] : (~sP3(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | p(s(bool,X12))) ) <=> ~$spl266),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl266])])).
fof(f11924,plain,(
  ( ! [X12,X13] : (~sP3(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl8),
  inference(superposition,[],[f7617,f864])).
fof(f7617,plain,(
  ( ! [X72,X73] : (~sP3(X72,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)),X73) | p(s(bool,X72))) ) | $spl8),
  inference(resolution,[],[f1089,f649])).
fof(f6368,plain,(
  ( ! [X12,X13] : (~sP3(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X12))) ) | $spl264),
  inference(cnf_transformation,[],[f6368_D])).
fof(f6368_D,plain,(
  ( ! [X12,X13] : (~sP3(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X12))) ) <=> ~$spl264),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl264])])).
fof(f11900,plain,(
  ( ! [X12,X13] : (~sP3(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl8),
  inference(superposition,[],[f7616,f864])).
fof(f7616,plain,(
  ( ! [X70,X71] : (~sP3(X70,X71,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) | p(s(bool,X70))) ) | $spl8),
  inference(resolution,[],[f1089,f648])).
fof(f1105,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) | $spl29),
  inference(cnf_transformation,[],[f1105_D])).
fof(f1105_D,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f11852,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) | $spl5),
  inference(superposition,[],[f1086,f864])).
fof(f4554,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | s(bool,X3) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)))) ) | $spl198),
  inference(cnf_transformation,[],[f4554_D])).
fof(f4554_D,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | s(bool,X3) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)))) ) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f11833,plain,(
  ( ! [X215,X214] : (p(s(bool,X214)) | ~p(s(bool,X215)) | s(bool,X215) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)))) ) | $spl5),
  inference(resolution,[],[f1086,f905])).
fof(f11735,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))) | s(bool,X23) != s(bool,X25)) ) | $spl5),
  inference(resolution,[],[f1086,f630])).
fof(f11733,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)))) ) | $spl5),
  inference(resolution,[],[f1086,f629])).
fof(f11731,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)))) ) | $spl5),
  inference(resolution,[],[f1086,f628])).
fof(f11730,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X9) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))) | s(bool,X8) != s(bool,X10)) ) | $spl5),
  inference(resolution,[],[f1086,f626])).
fof(f11728,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)))) ) | $spl5),
  inference(resolution,[],[f1086,f624])).
fof(f6313,plain,(
  ( ! [X12,X13] : (sP1('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | ~p(s(bool,X12))) ) | $spl262),
  inference(cnf_transformation,[],[f6313_D])).
fof(f6313_D,plain,(
  ( ! [X12,X13] : (sP1('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | ~p(s(bool,X12))) ) <=> ~$spl262),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl262])])).
fof(f11708,plain,(
  ( ! [X12,X13] : (sP1('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | ~p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl8),
  inference(superposition,[],[f7612,f864])).
fof(f7612,plain,(
  ( ! [X62,X63] : (sP1('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)),X62,X63) | ~p(s(bool,X62))) ) | $spl8),
  inference(resolution,[],[f1089,f644])).
fof(f6291,plain,(
  ( ! [X12,X13] : (sP1(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | ~p(s(bool,X12))) ) | $spl260),
  inference(cnf_transformation,[],[f6291_D])).
fof(f6291_D,plain,(
  ( ! [X12,X13] : (sP1(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | ~p(s(bool,X12))) ) <=> ~$spl260),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl260])])).
fof(f11687,plain,(
  ( ! [X12,X13] : (sP1(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | ~p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl8),
  inference(superposition,[],[f7611,f864])).
fof(f7611,plain,(
  ( ! [X61,X60] : (sP1(X60,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)),X61) | ~p(s(bool,X60))) ) | $spl8),
  inference(resolution,[],[f1089,f644])).
fof(f6270,plain,(
  ( ! [X12,X13] : (sP1('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | p(s(bool,X13))) ) | $spl258),
  inference(cnf_transformation,[],[f6270_D])).
fof(f6270_D,plain,(
  ( ! [X12,X13] : (sP1('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | p(s(bool,X13))) ) <=> ~$spl258),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl258])])).
fof(f11666,plain,(
  ( ! [X12,X13] : (sP1('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | p(s(bool,X13)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl8),
  inference(superposition,[],[f7610,f864])).
fof(f7610,plain,(
  ( ! [X59,X58] : (sP1('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)),X58,X59) | p(s(bool,X59))) ) | $spl8),
  inference(resolution,[],[f1089,f643])).
fof(f4582,plain,(
  ( ! [X12,X13] : (~sP1(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | p(s(bool,X12))) ) | $spl200),
  inference(cnf_transformation,[],[f4582_D])).
fof(f4582_D,plain,(
  ( ! [X12,X13] : (~sP1(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | p(s(bool,X12))) ) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f11644,plain,(
  ( ! [X12,X13] : (~sP1(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl8),
  inference(superposition,[],[f7601,f864])).
fof(f7601,plain,(
  ( ! [X41,X40] : (~sP1(X40,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)),X41) | p(s(bool,X40))) ) | $spl8),
  inference(resolution,[],[f1089,f632])).
fof(f4202,plain,(
  ( ! [X12,X13] : (sP11(X12,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13)) ) | $spl178),
  inference(cnf_transformation,[],[f4202_D])).
fof(f4202_D,plain,(
  ( ! [X12,X13] : (sP11(X12,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13)) ) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f11617,plain,(
  ( ! [X12,X13] : (sP11(X12,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f1015,f864])).
fof(f1015,plain,(
  ( ! [X175,X176] : (sP11(X175,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X176)) ) | $spl2),
  inference(resolution,[],[f915,f727])).
fof(f4198,plain,(
  ( ! [X12,X13] : (sP11(X12,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),sK15)),X13)) ) | $spl176),
  inference(cnf_transformation,[],[f4198_D])).
fof(f4198_D,plain,(
  ( ! [X12,X13] : (sP11(X12,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),sK15)),X13)) ) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f11605,plain,(
  ( ! [X12,X13] : (sP11(X12,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),sK15)),X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f1015,f864])).
fof(f11593,plain,(
  ( ! [X17,X18,X16] : (sP11(X16,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X17)),X18)) ) | $spl386),
  inference(cnf_transformation,[],[f11593_D])).
fof(f11593_D,plain,(
  ( ! [X17,X18,X16] : (sP11(X16,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X17)),X18)) ) <=> ~$spl386),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl386])])).
fof(f11581,plain,(
  ( ! [X17,X18,X16] : (sP11(X16,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X17)),X18) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl204),
  inference(superposition,[],[f10836,f864])).
fof(f11569,plain,(
  ( ! [X17,X18,X16] : (sP10(happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X16)),X17,X18)) ) | $spl384),
  inference(cnf_transformation,[],[f11569_D])).
fof(f11569_D,plain,(
  ( ! [X17,X18,X16] : (sP10(happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X16)),X17,X18)) ) <=> ~$spl384),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl384])])).
fof(f11557,plain,(
  ( ! [X17,X18,X16] : (sP10(happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X16)),X17,X18) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl204),
  inference(superposition,[],[f10829,f864])).
fof(f11545,plain,(
  ( ! [X17,X18,X16] : (sP6(happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X16)),X17,X18)) ) | $spl382),
  inference(cnf_transformation,[],[f11545_D])).
fof(f11545_D,plain,(
  ( ! [X17,X18,X16] : (sP6(happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X16)),X17,X18)) ) <=> ~$spl382),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl382])])).
fof(f11530,plain,(
  ( ! [X17,X18,X16] : (sP6(happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X16)),X17,X18) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl204),
  inference(superposition,[],[f10818,f864])).
fof(f4058,plain,(
  ( ! [X12,X13] : (sP10('const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13)) ) | $spl164),
  inference(cnf_transformation,[],[f4058_D])).
fof(f4058_D,plain,(
  ( ! [X12,X13] : (sP10('const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13)) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f11508,plain,(
  ( ! [X12,X13] : (sP10('const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f1008,f864])).
fof(f1008,plain,(
  ( ! [X161,X162] : (sP10('const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X161,X162)) ) | $spl2),
  inference(resolution,[],[f915,f718])).
fof(f4054,plain,(
  ( ! [X12,X13] : (sP10('const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),sK15)),X12,X13)) ) | $spl162),
  inference(cnf_transformation,[],[f4054_D])).
fof(f4054_D,plain,(
  ( ! [X12,X13] : (sP10('const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),sK15)),X12,X13)) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f11496,plain,(
  ( ! [X12,X13] : (sP10('const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),sK15)),X12,X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f1008,f864])).
fof(f11484,plain,(
  ( ! [X17,X18,X16] : (sP6(X16,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X17)),X18)) ) | $spl380),
  inference(cnf_transformation,[],[f11484_D])).
fof(f11484_D,plain,(
  ( ! [X17,X18,X16] : (sP6(X16,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X17)),X18)) ) <=> ~$spl380),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl380])])).
fof(f11469,plain,(
  ( ! [X17,X18,X16] : (sP6(X16,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X17)),X18) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl204),
  inference(superposition,[],[f10817,f864])).
fof(f11457,plain,(
  ( ! [X17,X18,X16] : (sP6(X16,X17,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X18)))) ) | $spl378),
  inference(cnf_transformation,[],[f11457_D])).
fof(f11457_D,plain,(
  ( ! [X17,X18,X16] : (sP6(X16,X17,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X18)))) ) <=> ~$spl378),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl378])])).
fof(f11442,plain,(
  ( ! [X17,X18,X16] : (sP6(X16,X17,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X18))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl204),
  inference(superposition,[],[f10816,f864])).
fof(f11430,plain,(
  ( ! [X17,X18,X16] : (sP7(X16,X17,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X18)))) ) | $spl376),
  inference(cnf_transformation,[],[f11430_D])).
fof(f11430_D,plain,(
  ( ! [X17,X18,X16] : (sP7(X16,X17,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X18)))) ) <=> ~$spl376),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl376])])).
fof(f11416,plain,(
  ( ! [X17,X18,X16] : (sP7(X16,X17,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X18))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl204),
  inference(superposition,[],[f10815,f864])).
fof(f11404,plain,(
  ( ! [X17,X18,X16] : (sP8(X16,X17,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X18)))) ) | $spl374),
  inference(cnf_transformation,[],[f11404_D])).
fof(f11404_D,plain,(
  ( ! [X17,X18,X16] : (sP8(X16,X17,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X18)))) ) <=> ~$spl374),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl374])])).
fof(f11390,plain,(
  ( ! [X17,X18,X16] : (sP8(X16,X17,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X18))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl204),
  inference(superposition,[],[f10812,f864])).
fof(f11378,plain,(
  ( ! [X17,X18,X16] : (sP4(X16,X17,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X18)))) ) | $spl372),
  inference(cnf_transformation,[],[f11378_D])).
fof(f11378_D,plain,(
  ( ! [X17,X18,X16] : (sP4(X16,X17,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X18)))) ) <=> ~$spl372),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl372])])).
fof(f11364,plain,(
  ( ! [X17,X18,X16] : (sP4(X16,X17,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X18))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl204),
  inference(superposition,[],[f10799,f864])).
fof(f3975,plain,(
  ( ! [X12,X13] : (sP6('const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13)) ) | $spl156),
  inference(cnf_transformation,[],[f3975_D])).
fof(f3975_D,plain,(
  ( ! [X12,X13] : (sP6('const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13)) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f11338,plain,(
  ( ! [X12,X13] : (sP6('const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f997,f864])).
fof(f997,plain,(
  ( ! [X140,X139] : (sP6('const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X139,X140)) ) | $spl2),
  inference(resolution,[],[f915,f700])).
fof(f3971,plain,(
  ( ! [X12,X13] : (sP6('const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),sK15)),X12,X13)) ) | $spl154),
  inference(cnf_transformation,[],[f3971_D])).
fof(f3971_D,plain,(
  ( ! [X12,X13] : (sP6('const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),sK15)),X12,X13)) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f11326,plain,(
  ( ! [X12,X13] : (sP6('const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),sK15)),X12,X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f997,f864])).
fof(f11314,plain,(
  ( ! [X17,X18,X16] : (sP2(X16,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X17)),X18)) ) | $spl370),
  inference(cnf_transformation,[],[f11314_D])).
fof(f11314_D,plain,(
  ( ! [X17,X18,X16] : (sP2(X16,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X17)),X18)) ) <=> ~$spl370),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl370])])).
fof(f11301,plain,(
  ( ! [X17,X18,X16] : (sP2(X16,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X17)),X18) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl204),
  inference(superposition,[],[f10783,f864])).
fof(f11288,plain,(
  ( ! [X17,X18,X16] : (sP2(X16,X17,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X18)))) ) | $spl368),
  inference(cnf_transformation,[],[f11288_D])).
fof(f11288_D,plain,(
  ( ! [X17,X18,X16] : (sP2(X16,X17,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X18)))) ) <=> ~$spl368),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl368])])).
fof(f11274,plain,(
  ( ! [X17,X18,X16] : (sP2(X16,X17,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X18))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl204),
  inference(superposition,[],[f10782,f864])).
fof(f11261,plain,(
  ( ! [X17,X18,X16] : (sP0(X16,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X17)),X18)) ) | $spl366),
  inference(cnf_transformation,[],[f11261_D])).
fof(f11261_D,plain,(
  ( ! [X17,X18,X16] : (sP0(X16,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X17)),X18)) ) <=> ~$spl366),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl366])])).
fof(f11248,plain,(
  ( ! [X17,X18,X16] : (sP0(X16,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X17)),X18) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl204),
  inference(superposition,[],[f10769,f864])).
fof(f3870,plain,(
  ( ! [X12,X13] : (sP6(X12,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13)) ) | $spl146),
  inference(cnf_transformation,[],[f3870_D])).
fof(f3870_D,plain,(
  ( ! [X12,X13] : (sP6(X12,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13)) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f11222,plain,(
  ( ! [X12,X13] : (sP6(X12,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f996,f864])).
fof(f996,plain,(
  ( ! [X138,X137] : (sP6(X137,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X138)) ) | $spl2),
  inference(resolution,[],[f915,f699])).
fof(f3866,plain,(
  ( ! [X12,X13] : (sP6(X12,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),sK15)),X13)) ) | $spl144),
  inference(cnf_transformation,[],[f3866_D])).
fof(f3866_D,plain,(
  ( ! [X12,X13] : (sP6(X12,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),sK15)),X13)) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f11210,plain,(
  ( ! [X12,X13] : (sP6(X12,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),sK15)),X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f996,f864])).
fof(f3813,plain,(
  ( ! [X12,X13] : (sP6(X12,X13,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) | $spl140),
  inference(cnf_transformation,[],[f3813_D])).
fof(f3813_D,plain,(
  ( ! [X12,X13] : (sP6(X12,X13,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f11185,plain,(
  ( ! [X12,X13] : (sP6(X12,X13,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f995,f864])).
fof(f995,plain,(
  ( ! [X136,X135] : (sP6(X135,X136,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))) ) | $spl2),
  inference(resolution,[],[f915,f698])).
fof(f3809,plain,(
  ( ! [X12,X13] : (sP6(X12,X13,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),sK15)))) ) | $spl138),
  inference(cnf_transformation,[],[f3809_D])).
fof(f3809_D,plain,(
  ( ! [X12,X13] : (sP6(X12,X13,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),sK15)))) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f11173,plain,(
  ( ! [X12,X13] : (sP6(X12,X13,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),sK15))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f995,f864])).
fof(f3693,plain,(
  ( ! [X12,X13] : (sP7(X12,X13,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) | $spl128),
  inference(cnf_transformation,[],[f3693_D])).
fof(f3693_D,plain,(
  ( ! [X12,X13] : (sP7(X12,X13,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f11149,plain,(
  ( ! [X12,X13] : (sP7(X12,X13,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f994,f864])).
fof(f994,plain,(
  ( ! [X134,X133] : (sP7(X133,X134,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))) ) | $spl2),
  inference(resolution,[],[f915,f694])).
fof(f3689,plain,(
  ( ! [X12,X13] : (sP7(X12,X13,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),sK15)))) ) | $spl126),
  inference(cnf_transformation,[],[f3689_D])).
fof(f3689_D,plain,(
  ( ! [X12,X13] : (sP7(X12,X13,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),sK15)))) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f11137,plain,(
  ( ! [X12,X13] : (sP7(X12,X13,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),sK15))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f994,f864])).
fof(f3777,plain,(
  ( ! [X8] : (sP12(X8,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X8))) ) | $spl136),
  inference(cnf_transformation,[],[f3777_D])).
fof(f3777_D,plain,(
  ( ! [X8] : (sP12(X8,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X8))) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f11119,plain,(
  ( ! [X8] : (sP12(X8,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X8)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl8),
  inference(superposition,[],[f7681,f864])).
fof(f7681,plain,(
  ( ! [X197] : (sP12(X197,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) | p(s(bool,X197))) ) | $spl8),
  inference(resolution,[],[f1089,f756])).
fof(f3755,plain,(
  ( ! [X8] : (sP12('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X8) | p(s(bool,X8))) ) | $spl134),
  inference(cnf_transformation,[],[f3755_D])).
fof(f3755_D,plain,(
  ( ! [X8] : (sP12('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X8) | p(s(bool,X8))) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f11098,plain,(
  ( ! [X8] : (sP12('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X8) | p(s(bool,X8)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl8),
  inference(superposition,[],[f7680,f864])).
fof(f7680,plain,(
  ( ! [X196] : (sP12('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)),X196) | p(s(bool,X196))) ) | $spl8),
  inference(resolution,[],[f1089,f753])).
fof(f3733,plain,(
  ( ! [X8] : (~sP12(X8,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X8))) ) | $spl132),
  inference(cnf_transformation,[],[f3733_D])).
fof(f3733_D,plain,(
  ( ! [X8] : (~sP12(X8,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X8))) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f11077,plain,(
  ( ! [X8] : (~sP12(X8,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X8)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl8),
  inference(superposition,[],[f7679,f864])).
fof(f7679,plain,(
  ( ! [X195] : (~sP12(X195,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) | ~p(s(bool,X195))) ) | $spl8),
  inference(resolution,[],[f1089,f752])).
fof(f3713,plain,(
  ( ! [X8] : (~sP12('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X8) | ~p(s(bool,X8))) ) | $spl130),
  inference(cnf_transformation,[],[f3713_D])).
fof(f3713_D,plain,(
  ( ! [X8] : (~sP12('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X8) | ~p(s(bool,X8))) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f11058,plain,(
  ( ! [X8] : (~sP12('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X8) | ~p(s(bool,X8)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl8),
  inference(superposition,[],[f7678,f864])).
fof(f7678,plain,(
  ( ! [X194] : (~sP12('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)),X194) | ~p(s(bool,X194))) ) | $spl8),
  inference(resolution,[],[f1089,f752])).
fof(f11048,plain,(
  $false | ($spl204 | $spl214)),
  inference(subsumption_resolution,[],[f5502,f5495])).
fof(f3635,plain,(
  ( ! [X12,X13] : (sP8(X12,X13,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) | $spl122),
  inference(cnf_transformation,[],[f3635_D])).
fof(f3635_D,plain,(
  ( ! [X12,X13] : (sP8(X12,X13,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f11032,plain,(
  ( ! [X12,X13] : (sP8(X12,X13,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f991,f864])).
fof(f991,plain,(
  ( ! [X127,X128] : (sP8(X127,X128,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))) ) | $spl2),
  inference(resolution,[],[f915,f690])).
fof(f3631,plain,(
  ( ! [X12,X13] : (sP8(X12,X13,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),sK15)))) ) | $spl120),
  inference(cnf_transformation,[],[f3631_D])).
fof(f3631_D,plain,(
  ( ! [X12,X13] : (sP8(X12,X13,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),sK15)))) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f11020,plain,(
  ( ! [X12,X13] : (sP8(X12,X13,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),sK15))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f991,f864])).
fof(f8530,plain,(
  ( ! [X55] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X55))))) ) | $spl326),
  inference(cnf_transformation,[],[f8530_D])).
fof(f8530_D,plain,(
  ( ! [X55] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X55))))) ) <=> ~$spl326),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl326])])).
fof(f10868,plain,(
  ( ! [X8] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X8)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl204),
  inference(superposition,[],[f5495,f864])).
fof(f10897,plain,(
  ( ! [X47,X49] : (p(s(bool,X47)) | s(bool,X47) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X49)))) ) | $spl364),
  inference(cnf_transformation,[],[f10897_D])).
fof(f10897_D,plain,(
  ( ! [X47,X49] : (p(s(bool,X47)) | s(bool,X47) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X49)))) ) <=> ~$spl364),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl364])])).
fof(f10859,plain,(
  ( ! [X360,X362,X361] : (p(s(bool,X360)) | p(s(bool,X361)) | s(bool,X361) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X362)))) ) | $spl204),
  inference(resolution,[],[f5495,f905])).
fof(f10761,plain,(
  ( ! [X76,X74,X75,X73] : (p(s(bool,X73)) | ~p(s(bool,X74)) | s(bool,X73) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X75))) | s(bool,X74) != s(bool,X76)) ) | $spl204),
  inference(resolution,[],[f5495,f630])).
fof(f10758,plain,(
  ( ! [X64,X62,X65,X63] : (p(s(bool,X62)) | ~p(s(bool,X63)) | p(s(bool,X64)) | s(bool,X64) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X65)))) ) | $spl204),
  inference(resolution,[],[f5495,f629])).
fof(f10757,plain,(
  ( ! [X61,X59,X60,X58] : (p(s(bool,X58)) | p(s(bool,X59)) | s(bool,X58) != s(bool,X60) | s(bool,X59) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X61)))) ) | $spl204),
  inference(resolution,[],[f5495,f628])).
fof(f10756,plain,(
  ( ! [X57,X54,X56,X55] : (p(s(bool,X54)) | ~p(s(bool,X55)) | s(bool,X54) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X56))) | s(bool,X55) != s(bool,X57)) ) | $spl204),
  inference(resolution,[],[f5495,f626])).
fof(f10754,plain,(
  ( ! [X47,X48,X46,X49] : (p(s(bool,X46)) | p(s(bool,X47)) | s(bool,X46) != s(bool,X48) | s(bool,X47) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X49)))) ) | $spl204),
  inference(resolution,[],[f5495,f624])).
fof(f3535,plain,(
  ( ! [X12,X13] : (sP4(X12,X13,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) | $spl112),
  inference(cnf_transformation,[],[f3535_D])).
fof(f3535_D,plain,(
  ( ! [X12,X13] : (sP4(X12,X13,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f10707,plain,(
  ( ! [X12,X13] : (sP4(X12,X13,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f978,f864])).
fof(f978,plain,(
  ( ! [X101,X102] : (sP4(X101,X102,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))) ) | $spl2),
  inference(resolution,[],[f915,f673])).
fof(f3531,plain,(
  ( ! [X12,X13] : (sP4(X12,X13,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),sK15)))) ) | $spl110),
  inference(cnf_transformation,[],[f3531_D])).
fof(f3531_D,plain,(
  ( ! [X12,X13] : (sP4(X12,X13,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),sK15)))) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f10695,plain,(
  ( ! [X12,X13] : (sP4(X12,X13,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),sK15))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f978,f864])).
fof(f3499,plain,(
  ( ! [X12,X13] : (sP2(X12,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13)) ) | $spl108),
  inference(cnf_transformation,[],[f3499_D])).
fof(f3499_D,plain,(
  ( ! [X12,X13] : (sP2(X12,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13)) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f10671,plain,(
  ( ! [X12,X13] : (sP2(X12,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f962,f864])).
fof(f962,plain,(
  ( ! [X70,X69] : (sP2(X69,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X70)) ) | $spl2),
  inference(resolution,[],[f915,f657])).
fof(f3495,plain,(
  ( ! [X12,X13] : (sP2(X12,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),sK15)),X13)) ) | $spl106),
  inference(cnf_transformation,[],[f3495_D])).
fof(f3495_D,plain,(
  ( ! [X12,X13] : (sP2(X12,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),sK15)),X13)) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f10659,plain,(
  ( ! [X12,X13] : (sP2(X12,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),sK15)),X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f962,f864])).
fof(f3464,plain,(
  ( ! [X12,X13] : (sP2(X12,X13,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) | $spl104),
  inference(cnf_transformation,[],[f3464_D])).
fof(f3464_D,plain,(
  ( ! [X12,X13] : (sP2(X12,X13,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f10593,plain,(
  ( ! [X12,X13] : (sP2(X12,X13,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f961,f864])).
fof(f961,plain,(
  ( ! [X68,X67] : (sP2(X67,X68,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))) ) | $spl2),
  inference(resolution,[],[f915,f656])).
fof(f3460,plain,(
  ( ! [X12,X13] : (sP2(X12,X13,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),sK15)))) ) | $spl102),
  inference(cnf_transformation,[],[f3460_D])).
fof(f3460_D,plain,(
  ( ! [X12,X13] : (sP2(X12,X13,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),sK15)))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f10581,plain,(
  ( ! [X12,X13] : (sP2(X12,X13,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),sK15))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f961,f864])).
fof(f3395,plain,(
  ( ! [X12,X13] : (sP0(X12,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13)) ) | $spl96),
  inference(cnf_transformation,[],[f3395_D])).
fof(f3395_D,plain,(
  ( ! [X12,X13] : (sP0(X12,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13)) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f10485,plain,(
  ( ! [X12,X13] : (sP0(X12,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f948,f864])).
fof(f948,plain,(
  ( ! [X41,X42] : (sP0(X41,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X42)) ) | $spl2),
  inference(resolution,[],[f915,f640])).
fof(f3391,plain,(
  ( ! [X12,X13] : (sP0(X12,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),sK15)),X13)) ) | $spl94),
  inference(cnf_transformation,[],[f3391_D])).
fof(f3391_D,plain,(
  ( ! [X12,X13] : (sP0(X12,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),sK15)),X13)) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f10473,plain,(
  ( ! [X12,X13] : (sP0(X12,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),sK15)),X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f948,f864])).
fof(f1111,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) | $spl33),
  inference(cnf_transformation,[],[f1111_D])).
fof(f1111_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f3429,plain,(
  ( ! [X12,X13] : (sP11(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13)) ) | $spl100),
  inference(cnf_transformation,[],[f3429_D])).
fof(f3429_D,plain,(
  ( ! [X12,X13] : (sP11(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13)) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f10427,plain,(
  ( ! [X21,X20] : (sP11(X20,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X21) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15)))) ) | $spl8),
  inference(subsumption_resolution,[],[f10422,f1089])).
fof(f10422,plain,(
  ( ! [X21,X20] : (sP11(X20,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X21) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))))) ) | $spl8),
  inference(superposition,[],[f1688,f897])).
fof(f1688,plain,(
  ( ! [X189,X188] : (sP11(X188,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)),X189)) ) | $spl8),
  inference(resolution,[],[f1089,f727])).
fof(f10419,plain,(
  ( ! [X12,X13] : (sP11(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl8),
  inference(superposition,[],[f1688,f864])).
fof(f3412,plain,(
  ( ! [X12,X13] : (sP10('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13)) ) | $spl98),
  inference(cnf_transformation,[],[f3412_D])).
fof(f3412_D,plain,(
  ( ! [X12,X13] : (sP10('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13)) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f10410,plain,(
  ( ! [X21,X20] : (sP10('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X20,X21) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15)))) ) | $spl8),
  inference(subsumption_resolution,[],[f10405,f1089])).
fof(f10405,plain,(
  ( ! [X21,X20] : (sP10('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X20,X21) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))))) ) | $spl8),
  inference(superposition,[],[f1681,f897])).
fof(f1681,plain,(
  ( ! [X175,X174] : (sP10('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)),X174,X175)) ) | $spl8),
  inference(resolution,[],[f1089,f718])).
fof(f10402,plain,(
  ( ! [X12,X13] : (sP10('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl8),
  inference(superposition,[],[f1681,f864])).
fof(f10381,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))))) | $spl1),
  inference(trivial_inequality_removal,[],[f10376])).
fof(f10376,plain,(
  s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),'const.pred_set.EMPTY_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))))) | $spl1),
  inference(superposition,[],[f914,f897])).
fof(f10393,plain,(
  ( ! [X5] : (s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),X5) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))))))) ) | $spl1),
  inference(subsumption_resolution,[],[f10375,f860])).
fof(f10375,plain,(
  ( ! [X5] : (s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),X5) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X5))))) ) | $spl1),
  inference(superposition,[],[f914,f866])).
fof(f8463,plain,(
  s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | $spl297),
  inference(cnf_transformation,[],[f8463_D])).
fof(f8463_D,plain,(
  s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) <=> ~$spl297),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl297])])).
fof(f10391,plain,(
  s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) | ($spl1 | $spl8)),
  inference(subsumption_resolution,[],[f10364,f1089])).
fof(f10364,plain,(
  s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) | $spl1),
  inference(superposition,[],[f914,f897])).
fof(f3360,plain,(
  ( ! [X12,X13] : (sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13)) ) | $spl92),
  inference(cnf_transformation,[],[f3360_D])).
fof(f3360_D,plain,(
  ( ! [X12,X13] : (sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13)) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f10344,plain,(
  ( ! [X21,X20] : (sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X20,X21) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15)))) ) | $spl8),
  inference(subsumption_resolution,[],[f10336,f1089])).
fof(f10336,plain,(
  ( ! [X21,X20] : (sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X20,X21) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))))) ) | $spl8),
  inference(superposition,[],[f1670,f897])).
fof(f1670,plain,(
  ( ! [X152,X153] : (sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)),X152,X153)) ) | $spl8),
  inference(resolution,[],[f1089,f700])).
fof(f10333,plain,(
  ( ! [X12,X13] : (sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl8),
  inference(superposition,[],[f1670,f864])).
fof(f3340,plain,(
  ( ! [X12,X13] : (sP6(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13)) ) | $spl90),
  inference(cnf_transformation,[],[f3340_D])).
fof(f3340_D,plain,(
  ( ! [X12,X13] : (sP6(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13)) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f10324,plain,(
  ( ! [X21,X20] : (sP6(X20,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X21) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15)))) ) | $spl8),
  inference(subsumption_resolution,[],[f10316,f1089])).
fof(f10316,plain,(
  ( ! [X21,X20] : (sP6(X20,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X21) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))))) ) | $spl8),
  inference(superposition,[],[f1669,f897])).
fof(f1669,plain,(
  ( ! [X151,X150] : (sP6(X150,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)),X151)) ) | $spl8),
  inference(resolution,[],[f1089,f699])).
fof(f10313,plain,(
  ( ! [X12,X13] : (sP6(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl8),
  inference(superposition,[],[f1669,f864])).
fof(f3320,plain,(
  ( ! [X12,X13] : (sP6(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) | $spl88),
  inference(cnf_transformation,[],[f3320_D])).
fof(f3320_D,plain,(
  ( ! [X12,X13] : (sP6(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f10304,plain,(
  ( ! [X21,X20] : (sP6(X20,X21,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15)))) ) | $spl8),
  inference(subsumption_resolution,[],[f10296,f1089])).
fof(f10296,plain,(
  ( ! [X21,X20] : (sP6(X20,X21,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))))) ) | $spl8),
  inference(superposition,[],[f1668,f897])).
fof(f1668,plain,(
  ( ! [X149,X148] : (sP6(X148,X149,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | $spl8),
  inference(resolution,[],[f1089,f698])).
fof(f10293,plain,(
  ( ! [X12,X13] : (sP6(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl8),
  inference(superposition,[],[f1668,f864])).
fof(f3300,plain,(
  ( ! [X12,X13] : (sP7(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) | $spl86),
  inference(cnf_transformation,[],[f3300_D])).
fof(f3300_D,plain,(
  ( ! [X12,X13] : (sP7(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f10284,plain,(
  ( ! [X21,X20] : (sP7(X20,X21,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15)))) ) | $spl8),
  inference(subsumption_resolution,[],[f10277,f1089])).
fof(f10277,plain,(
  ( ! [X21,X20] : (sP7(X20,X21,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))))) ) | $spl8),
  inference(superposition,[],[f1667,f897])).
fof(f1667,plain,(
  ( ! [X146,X147] : (sP7(X146,X147,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | $spl8),
  inference(resolution,[],[f1089,f694])).
fof(f10274,plain,(
  ( ! [X12,X13] : (sP7(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl8),
  inference(superposition,[],[f1667,f864])).
fof(f3281,plain,(
  ( ! [X12,X13] : (sP8(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) | $spl84),
  inference(cnf_transformation,[],[f3281_D])).
fof(f3281_D,plain,(
  ( ! [X12,X13] : (sP8(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f10265,plain,(
  ( ! [X21,X20] : (sP8(X20,X21,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15)))) ) | $spl8),
  inference(subsumption_resolution,[],[f10258,f1089])).
fof(f10258,plain,(
  ( ! [X21,X20] : (sP8(X20,X21,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))))) ) | $spl8),
  inference(superposition,[],[f1664,f897])).
fof(f1664,plain,(
  ( ! [X140,X141] : (sP8(X140,X141,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | $spl8),
  inference(resolution,[],[f1089,f690])).
fof(f10255,plain,(
  ( ! [X12,X13] : (sP8(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl8),
  inference(superposition,[],[f1664,f864])).
fof(f1106,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) | $spl28),
  inference(cnf_transformation,[],[f1106_D])).
fof(f1106_D,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f10246,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) | ($spl2 | $spl8)),
  inference(subsumption_resolution,[],[f10215,f1089])).
fof(f10215,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) | $spl2),
  inference(superposition,[],[f915,f897])).
fof(f10212,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) | $spl2),
  inference(superposition,[],[f915,f864])).
fof(f10191,plain,(
  ( ! [X213,X214] : (p(s(bool,X213)) | p(s(bool,X214)) | s(bool,X214) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))) ) | $spl2),
  inference(resolution,[],[f915,f905])).
fof(f10093,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))) | s(bool,X25) != s(bool,X26)) ) | $spl2),
  inference(resolution,[],[f915,f630])).
fof(f10090,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X17) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))) ) | $spl2),
  inference(resolution,[],[f915,f629])).
fof(f10089,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))) ) | $spl2),
  inference(resolution,[],[f915,f628])).
fof(f10088,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))) | s(bool,X10) != s(bool,X11)) ) | $spl2),
  inference(resolution,[],[f915,f626])).
fof(f10086,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))) ) | $spl2),
  inference(resolution,[],[f915,f624])).
fof(f1094,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) | $spl13),
  inference(cnf_transformation,[],[f1094_D])).
fof(f1094_D,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f1088,plain,(
  s(fun(sK13,bool),sK14) = s(fun(sK13,bool),sK15) | $spl6),
  inference(cnf_transformation,[],[f1088_D])).
fof(f1088_D,plain,(
  s(fun(sK13,bool),sK14) = s(fun(sK13,bool),sK15) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f10239,plain,(
  s(fun(sK13,bool),sK14) = s(fun(sK13,bool),sK15) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) | ($spl2 | $spl8 | $spl10)),
  inference(subsumption_resolution,[],[f10238,f1092])).
fof(f10238,plain,(
  s(fun(sK13,bool),sK14) = s(fun(sK13,bool),sK15) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ($spl2 | $spl8)),
  inference(subsumption_resolution,[],[f10084,f1089])).
fof(f10084,plain,(
  s(fun(sK13,bool),sK14) = s(fun(sK13,bool),sK15) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | $spl2),
  inference(resolution,[],[f915,f866])).
fof(f1086,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)))) | $spl5),
  inference(cnf_transformation,[],[f1086_D])).
fof(f1086_D,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f10080,plain,(
  s(fun(sK13,bool),sK14) = s(fun(sK13,bool),sK15) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)))) | $spl2),
  inference(resolution,[],[f915,f859])).
fof(f859,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f439])).
fof(f439,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2))),
  inference(flattening,[],[f438])).
fof(f438,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2))),
  inference(ennf_transformation,[],[f346])).
fof(f346,plain,(
  ! [X0,X1,X2] : ((p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))))) => s(fun(X0,bool),X1) = s(fun(X0,bool),X2))),
  inference(flattening,[],[f345])).
fof(f345,plain,(
  ! [X0] : ! [X1,X2] : ((p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))))) => s(fun(X0,bool),X1) = s(fun(X0,bool),X2))),
  inference(rectify,[],[f108])).
fof(f108,axiom,(
  ! [X6] : ! [X7,X8] : ((p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X7))))) => s(fun(X6,bool),X7) = s(fun(X6,bool),X8))),
  file('thm.pred_set.SUBSET_DIFF_EMPTY','thm.pred_set.SUBSET_ANTISYM')).
fof(f10079,plain,(
  s(fun(sK13,bool),sK14) = s(fun(sK13,bool),sK15) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)))) | $spl2),
  inference(resolution,[],[f915,f859])).
fof(f10078,plain,(
  s(fun(sK13,bool),sK14) = s(fun(sK13,bool),sK15) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)))) | $spl2),
  inference(resolution,[],[f915,f858])).
fof(f858,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f610])).
fof(f610,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))))) & (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(flattening,[],[f609])).
fof(f609,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))))) & ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(nnf_transformation,[],[f344])).
fof(f344,plain,(
  ! [X0,X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))))))),
  inference(flattening,[],[f343])).
fof(f343,plain,(
  ! [X0] : ! [X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))))))),
  inference(rectify,[],[f117])).
fof(f117,axiom,(
  ! [X6] : ! [X12,X13] : (s(fun(X6,bool),X12) = s(fun(X6,bool),X13) <=> (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X12),s(fun(X6,bool),X13)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X13),s(fun(X6,bool),X12))))))),
  file('thm.pred_set.SUBSET_DIFF_EMPTY','thm.pred_set.SET_EQ_SUBSET')).
fof(f10077,plain,(
  s(fun(sK13,bool),sK14) = s(fun(sK13,bool),sK15) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)))) | $spl2),
  inference(resolution,[],[f915,f858])).
fof(f3241,plain,(
  ( ! [X12,X13] : (sP4(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) | $spl82),
  inference(cnf_transformation,[],[f3241_D])).
fof(f3241_D,plain,(
  ( ! [X12,X13] : (sP4(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f10069,plain,(
  ( ! [X21,X20] : (sP4(X20,X21,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15)))) ) | $spl8),
  inference(subsumption_resolution,[],[f10062,f1089])).
fof(f10062,plain,(
  ( ! [X21,X20] : (sP4(X20,X21,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))))) ) | $spl8),
  inference(superposition,[],[f1651,f897])).
fof(f1651,plain,(
  ( ! [X114,X115] : (sP4(X114,X115,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | $spl8),
  inference(resolution,[],[f1089,f673])).
fof(f10059,plain,(
  ( ! [X12,X13] : (sP4(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl8),
  inference(superposition,[],[f1651,f864])).
fof(f3220,plain,(
  ( ! [X12,X13] : (sP2(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13)) ) | $spl80),
  inference(cnf_transformation,[],[f3220_D])).
fof(f3220_D,plain,(
  ( ! [X12,X13] : (sP2(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13)) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f10049,plain,(
  ( ! [X21,X20] : (sP2(X20,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X21) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15)))) ) | $spl8),
  inference(subsumption_resolution,[],[f10043,f1089])).
fof(f10043,plain,(
  ( ! [X21,X20] : (sP2(X20,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X21) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))))) ) | $spl8),
  inference(superposition,[],[f1635,f897])).
fof(f1635,plain,(
  ( ! [X83,X82] : (sP2(X82,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)),X83)) ) | $spl8),
  inference(resolution,[],[f1089,f657])).
fof(f10040,plain,(
  ( ! [X12,X13] : (sP2(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl8),
  inference(superposition,[],[f1635,f864])).
fof(f3200,plain,(
  ( ! [X12,X13] : (sP2(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) | $spl78),
  inference(cnf_transformation,[],[f3200_D])).
fof(f3200_D,plain,(
  ( ! [X12,X13] : (sP2(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f10030,plain,(
  ( ! [X21,X20] : (sP2(X20,X21,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15)))) ) | $spl8),
  inference(subsumption_resolution,[],[f10024,f1089])).
fof(f10024,plain,(
  ( ! [X21,X20] : (sP2(X20,X21,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))))) ) | $spl8),
  inference(superposition,[],[f1634,f897])).
fof(f1634,plain,(
  ( ! [X80,X81] : (sP2(X80,X81,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | $spl8),
  inference(resolution,[],[f1089,f656])).
fof(f10021,plain,(
  ( ! [X12,X13] : (sP2(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl8),
  inference(superposition,[],[f1634,f864])).
fof(f3180,plain,(
  ( ! [X12,X13] : (sP0(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13)) ) | $spl76),
  inference(cnf_transformation,[],[f3180_D])).
fof(f3180_D,plain,(
  ( ! [X12,X13] : (sP0(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13)) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f10011,plain,(
  ( ! [X21,X20] : (sP0(X20,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X21) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15)))) ) | $spl8),
  inference(subsumption_resolution,[],[f10005,f1089])).
fof(f10005,plain,(
  ( ! [X21,X20] : (sP0(X20,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X21) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))))) ) | $spl8),
  inference(superposition,[],[f1621,f897])).
fof(f1621,plain,(
  ( ! [X54,X55] : (sP0(X54,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)),X55)) ) | $spl8),
  inference(resolution,[],[f1089,f640])).
fof(f10002,plain,(
  ( ! [X12,X13] : (sP0(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl8),
  inference(superposition,[],[f1621,f864])).
fof(f1748,plain,(
  ( ! [X17] : (s(bool,X17) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) | p(s(bool,X17))) ) | $spl52),
  inference(cnf_transformation,[],[f1748_D])).
fof(f1748_D,plain,(
  ( ! [X17] : (s(bool,X17) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) | p(s(bool,X17))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f9958,plain,(
  ( ! [X227,X226] : (p(s(bool,X226)) | p(s(bool,X227)) | s(bool,X227) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | $spl8),
  inference(resolution,[],[f1089,f905])).
fof(f9860,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) | s(bool,X38) != s(bool,X39)) ) | $spl8),
  inference(resolution,[],[f1089,f630])).
fof(f9857,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X28)) | ~p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X30) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | $spl8),
  inference(resolution,[],[f1089,f629])).
fof(f9856,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | $spl8),
  inference(resolution,[],[f1089,f628])).
fof(f9855,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) | s(bool,X23) != s(bool,X24)) ) | $spl8),
  inference(resolution,[],[f1089,f626])).
fof(f9853,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X16) != s(bool,X18) | s(bool,X17) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | $spl8),
  inference(resolution,[],[f1089,f624])).
fof(f1746,plain,(
  s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),sK15) | $spl51),
  inference(cnf_transformation,[],[f1746_D])).
fof(f1746_D,plain,(
  s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),sK15) <=> ~$spl51),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl51])])).
fof(f1110,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) | $spl32),
  inference(cnf_transformation,[],[f1110_D])).
fof(f1110_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f9852,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) | s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),sK15) | $spl8),
  inference(resolution,[],[f1089,f898])).
fof(f1747,plain,(
  s(fun(sK13,bool),'const.pred_set.EMPTY_0') = s(fun(sK13,bool),sK15) | $spl50),
  inference(cnf_transformation,[],[f1747_D])).
fof(f1747_D,plain,(
  s(fun(sK13,bool),'const.pred_set.EMPTY_0') = s(fun(sK13,bool),sK15) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f9851,plain,(
  s(fun(sK13,bool),'const.pred_set.EMPTY_0') = s(fun(sK13,bool),sK15) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) | $spl8),
  inference(resolution,[],[f1089,f897])).
fof(f9495,plain,(
  $false | ($spl1 | $spl6)),
  inference(subsumption_resolution,[],[f9142,f889])).
fof(f9142,plain,(
  s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl1 | $spl6)),
  inference(backward_demodulation,[],[f1088,f914])).
fof(f9129,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))))) | $spl1),
  inference(trivial_inequality_removal,[],[f9124])).
fof(f9124,plain,(
  s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),'const.pred_set.EMPTY_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))))) | $spl1),
  inference(superposition,[],[f914,f897])).
fof(f9141,plain,(
  ( ! [X5] : (s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),X5) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))))))) ) | $spl1),
  inference(subsumption_resolution,[],[f9123,f860])).
fof(f9123,plain,(
  ( ! [X5] : (s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),X5) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X5))))) ) | $spl1),
  inference(superposition,[],[f914,f866])).
fof(f9139,plain,(
  s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) | ($spl1 | $spl8)),
  inference(subsumption_resolution,[],[f9112,f1089])).
fof(f9112,plain,(
  s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) | $spl1),
  inference(superposition,[],[f914,f897])).
fof(f9092,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) | ($spl2 | $spl8)),
  inference(subsumption_resolution,[],[f9061,f1089])).
fof(f9061,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) | $spl2),
  inference(superposition,[],[f915,f897])).
fof(f9058,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) | $spl2),
  inference(superposition,[],[f915,f864])).
fof(f9037,plain,(
  ( ! [X213,X214] : (p(s(bool,X213)) | p(s(bool,X214)) | s(bool,X214) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))) ) | $spl2),
  inference(resolution,[],[f915,f905])).
fof(f8939,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))) | s(bool,X25) != s(bool,X26)) ) | $spl2),
  inference(resolution,[],[f915,f630])).
fof(f8936,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X17) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))) ) | $spl2),
  inference(resolution,[],[f915,f629])).
fof(f8935,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))) ) | $spl2),
  inference(resolution,[],[f915,f628])).
fof(f8934,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))) | s(bool,X10) != s(bool,X11)) ) | $spl2),
  inference(resolution,[],[f915,f626])).
fof(f8932,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))) ) | $spl2),
  inference(resolution,[],[f915,f624])).
fof(f9085,plain,(
  s(fun(sK13,bool),sK14) = s(fun(sK13,bool),sK15) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) | ($spl2 | $spl8 | $spl10)),
  inference(subsumption_resolution,[],[f9084,f1092])).
fof(f9084,plain,(
  s(fun(sK13,bool),sK14) = s(fun(sK13,bool),sK15) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ($spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8930,f1089])).
fof(f8930,plain,(
  s(fun(sK13,bool),sK14) = s(fun(sK13,bool),sK15) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | $spl2),
  inference(resolution,[],[f915,f866])).
fof(f8926,plain,(
  s(fun(sK13,bool),sK14) = s(fun(sK13,bool),sK15) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)))) | $spl2),
  inference(resolution,[],[f915,f859])).
fof(f8925,plain,(
  s(fun(sK13,bool),sK14) = s(fun(sK13,bool),sK15) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)))) | $spl2),
  inference(resolution,[],[f915,f859])).
fof(f8924,plain,(
  s(fun(sK13,bool),sK14) = s(fun(sK13,bool),sK15) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)))) | $spl2),
  inference(resolution,[],[f915,f858])).
fof(f8923,plain,(
  s(fun(sK13,bool),sK14) = s(fun(sK13,bool),sK15) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)))) | $spl2),
  inference(resolution,[],[f915,f858])).
fof(f8915,plain,(
  ( ! [X21,X20] : (sP11(X20,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X21) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15)))) ) | $spl8),
  inference(subsumption_resolution,[],[f8910,f1089])).
fof(f8910,plain,(
  ( ! [X21,X20] : (sP11(X20,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X21) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))))) ) | $spl8),
  inference(superposition,[],[f1688,f897])).
fof(f8907,plain,(
  ( ! [X12,X13] : (sP11(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl8),
  inference(superposition,[],[f1688,f864])).
fof(f8898,plain,(
  ( ! [X21,X20] : (sP10('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X20,X21) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15)))) ) | $spl8),
  inference(subsumption_resolution,[],[f8893,f1089])).
fof(f8893,plain,(
  ( ! [X21,X20] : (sP10('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X20,X21) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))))) ) | $spl8),
  inference(superposition,[],[f1681,f897])).
fof(f8890,plain,(
  ( ! [X12,X13] : (sP10('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl8),
  inference(superposition,[],[f1681,f864])).
fof(f8813,plain,(
  ( ! [X337] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X337))))) ) | $spl360),
  inference(cnf_transformation,[],[f8813_D])).
fof(f8813_D,plain,(
  ( ! [X337] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X337))))) ) <=> ~$spl360),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl360])])).
fof(f8821,plain,(
  ( ! [X339] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X339))))) ) | $spl0),
  inference(forward_demodulation,[],[f8820,f901])).
fof(f8820,plain,(
  ( ! [X339] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X339))))) ) | $spl0),
  inference(forward_demodulation,[],[f8819,f891])).
fof(f8819,plain,(
  ( ! [X339] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),X339))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X339))))) ) | $spl0),
  inference(forward_demodulation,[],[f8818,f901])).
fof(f8818,plain,(
  ( ! [X339] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),X339))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X339))))) ) | $spl0),
  inference(forward_demodulation,[],[f8817,f900])).
fof(f8817,plain,(
  ( ! [X339] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),X339))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),X339))))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X339))))) ) | $spl0),
  inference(forward_demodulation,[],[f8816,f913])).
fof(f8816,plain,(
  ( ! [X339] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))),s(fun(sK13,bool),X339))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))),s(fun(sK13,bool),X339))))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X339))))) ) | $spl0),
  inference(subsumption_resolution,[],[f8355,f902])).
fof(f8355,plain,(
  ( ! [X339] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))),s(fun(sK13,bool),X339))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))),s(fun(sK13,bool),X339))))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X339))))) ) | $spl0),
  inference(superposition,[],[f887,f913])).
fof(f8812,plain,(
  ( ! [X337] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X337))))) ) | $spl0),
  inference(subsumption_resolution,[],[f8811,f902])).
fof(f8811,plain,(
  ( ! [X337] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X337))))) ) | $spl0),
  inference(forward_demodulation,[],[f8810,f913])).
fof(f8810,plain,(
  ( ! [X337] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X337))))) ) | $spl0),
  inference(forward_demodulation,[],[f8809,f901])).
fof(f8809,plain,(
  ( ! [X337] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X337))))) ) | $spl0),
  inference(forward_demodulation,[],[f8808,f891])).
fof(f8808,plain,(
  ( ! [X337] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),X337))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X337))))) ) | $spl0),
  inference(forward_demodulation,[],[f8807,f901])).
fof(f8807,plain,(
  ( ! [X337] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),X337))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X337))))) ) | $spl0),
  inference(forward_demodulation,[],[f8353,f900])).
fof(f8353,plain,(
  ( ! [X337] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),X337))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.INTER_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),X337))))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X337))))) ) | $spl0),
  inference(superposition,[],[f887,f913])).
fof(f8790,plain,(
  ( ! [X331] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X331))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X331))))) ) | $spl358),
  inference(cnf_transformation,[],[f8790_D])).
fof(f8790_D,plain,(
  ( ! [X331] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X331))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X331))))) ) <=> ~$spl358),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl358])])).
fof(f8789,plain,(
  ( ! [X331] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X331))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X331))))) ) | $spl0),
  inference(subsumption_resolution,[],[f8788,f902])).
fof(f8788,plain,(
  ( ! [X331] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X331))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X331))))) ) | $spl0),
  inference(forward_demodulation,[],[f8787,f913])).
fof(f8787,plain,(
  ( ! [X331] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X331))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X331))))) ) | $spl0),
  inference(forward_demodulation,[],[f8786,f901])).
fof(f8786,plain,(
  ( ! [X331] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X331))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X331))))) ) | $spl0),
  inference(forward_demodulation,[],[f8785,f913])).
fof(f8785,plain,(
  ( ! [X331] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X331))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X331))))) ) | $spl0),
  inference(forward_demodulation,[],[f8784,f901])).
fof(f8784,plain,(
  ( ! [X331] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X331))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X331))))) ) | $spl0),
  inference(forward_demodulation,[],[f8347,f891])).
fof(f8347,plain,(
  ( ! [X331] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),X331)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X331))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X331))))) ) | $spl0),
  inference(superposition,[],[f886,f913])).
fof(f8577,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK38('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13))))) | $spl357),
  inference(cnf_transformation,[],[f8577_D])).
fof(f8577_D,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK38('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13))))) <=> ~$spl357),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl357])])).
fof(f8578,plain,(
  ( ! [X121] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK38('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13))))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X121))))) ) | $spl0),
  inference(forward_demodulation,[],[f8186,f913])).
fof(f8186,plain,(
  ( ! [X121] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X121)))) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))),s(sK13,sK38('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13)))))) ) | $spl0),
  inference(superposition,[],[f825,f913])).
fof(f8575,plain,(
  ( ! [X120] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X120)))) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK38('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13)))))) ) | $spl0),
  inference(forward_demodulation,[],[f8185,f913])).
fof(f8185,plain,(
  ( ! [X120] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK38('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13))))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))),s(sK13,X120))))) ) | $spl0),
  inference(superposition,[],[f825,f913])).
fof(f8573,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK37('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13))))) | $spl355),
  inference(cnf_transformation,[],[f8573_D])).
fof(f8573_D,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK37('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13))))) <=> ~$spl355),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl355])])).
fof(f8574,plain,(
  ( ! [X119] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK37('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13))))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X119))))) ) | $spl0),
  inference(forward_demodulation,[],[f8184,f913])).
fof(f8184,plain,(
  ( ! [X119] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X119)))) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))),s(sK13,sK37('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13)))))) ) | $spl0),
  inference(superposition,[],[f824,f913])).
fof(f8571,plain,(
  ( ! [X118] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X118)))) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK37('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13)))))) ) | $spl0),
  inference(forward_demodulation,[],[f8183,f913])).
fof(f8183,plain,(
  ( ! [X118] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK37('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13))))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))),s(sK13,X118))))) ) | $spl0),
  inference(superposition,[],[f824,f913])).
fof(f8569,plain,(
  ( ! [X114] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK36('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X114,sK13))))) | ~p(s(bool,X114))) ) | $spl352),
  inference(cnf_transformation,[],[f8569_D])).
fof(f8569_D,plain,(
  ( ! [X114] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK36('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X114,sK13))))) | ~p(s(bool,X114))) ) <=> ~$spl352),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl352])])).
fof(f8570,plain,(
  ( ! [X116,X117] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK36('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X117,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X116)))) | ~p(s(bool,X117))) ) | $spl0),
  inference(forward_demodulation,[],[f8182,f913])).
fof(f8182,plain,(
  ( ! [X116,X117] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X116)))) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))),s(sK13,sK36('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X117,sK13))))) | ~p(s(bool,X117))) ) | $spl0),
  inference(superposition,[],[f821,f913])).
fof(f8568,plain,(
  ( ! [X114,X115] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X115)))) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK36('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X114,sK13))))) | ~p(s(bool,X114))) ) | $spl0),
  inference(forward_demodulation,[],[f8181,f913])).
fof(f8181,plain,(
  ( ! [X114,X115] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK36('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X114,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))),s(sK13,X115)))) | ~p(s(bool,X114))) ) | $spl0),
  inference(superposition,[],[f821,f913])).
fof(f8566,plain,(
  ( ! [X110] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK35('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X110,sK13))))) | ~p(s(bool,X110))) ) | $spl350),
  inference(cnf_transformation,[],[f8566_D])).
fof(f8566_D,plain,(
  ( ! [X110] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK35('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X110,sK13))))) | ~p(s(bool,X110))) ) <=> ~$spl350),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl350])])).
fof(f8567,plain,(
  ( ! [X112,X113] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK35('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X113,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X112)))) | ~p(s(bool,X113))) ) | $spl0),
  inference(forward_demodulation,[],[f8180,f913])).
fof(f8180,plain,(
  ( ! [X112,X113] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X112)))) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))),s(sK13,sK35('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X113,sK13))))) | ~p(s(bool,X113))) ) | $spl0),
  inference(superposition,[],[f819,f913])).
fof(f8565,plain,(
  ( ! [X111,X110] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X111)))) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK35('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X110,sK13))))) | ~p(s(bool,X110))) ) | $spl0),
  inference(forward_demodulation,[],[f8179,f913])).
fof(f8179,plain,(
  ( ! [X111,X110] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK35('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X110,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))),s(sK13,X111)))) | ~p(s(bool,X110))) ) | $spl0),
  inference(superposition,[],[f819,f913])).
fof(f8501,plain,(
  ( ! [X8] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X8))))) ) | $spl308),
  inference(cnf_transformation,[],[f8501_D])).
fof(f8501_D,plain,(
  ( ! [X8] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X8))))) ) <=> ~$spl308),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl308])])).
fof(f8563,plain,(
  ( ! [X106] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK34('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X106,sK13))))) | p(s(bool,X106))) ) | $spl348),
  inference(cnf_transformation,[],[f8563_D])).
fof(f8563_D,plain,(
  ( ! [X106] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK34('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X106,sK13))))) | p(s(bool,X106))) ) <=> ~$spl348),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl348])])).
fof(f8564,plain,(
  ( ! [X109,X108] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK34('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X109,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X108)))) | p(s(bool,X109))) ) | $spl0),
  inference(forward_demodulation,[],[f8178,f913])).
fof(f8178,plain,(
  ( ! [X109,X108] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X108)))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))),s(sK13,sK34('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X109,sK13))))) | p(s(bool,X109))) ) | $spl0),
  inference(superposition,[],[f817,f913])).
fof(f8562,plain,(
  ( ! [X107,X106] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X107)))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK34('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X106,sK13))))) | p(s(bool,X106))) ) | $spl0),
  inference(forward_demodulation,[],[f8177,f913])).
fof(f8177,plain,(
  ( ! [X107,X106] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK34('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X106,sK13))))) | p(s(bool,X106)) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))),s(sK13,X107))))) ) | $spl0),
  inference(superposition,[],[f817,f913])).
fof(f8560,plain,(
  ( ! [X102] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK33('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X102,sK13))))) | p(s(bool,X102))) ) | $spl346),
  inference(cnf_transformation,[],[f8560_D])).
fof(f8560_D,plain,(
  ( ! [X102] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK33('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X102,sK13))))) | p(s(bool,X102))) ) <=> ~$spl346),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl346])])).
fof(f8561,plain,(
  ( ! [X105,X104] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK33('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X105,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X104)))) | p(s(bool,X105))) ) | $spl0),
  inference(forward_demodulation,[],[f8176,f913])).
fof(f8176,plain,(
  ( ! [X105,X104] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X104)))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))),s(sK13,sK33('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X105,sK13))))) | p(s(bool,X105))) ) | $spl0),
  inference(superposition,[],[f815,f913])).
fof(f8559,plain,(
  ( ! [X103,X102] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X103)))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK33('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X102,sK13))))) | p(s(bool,X102))) ) | $spl0),
  inference(forward_demodulation,[],[f8175,f913])).
fof(f8175,plain,(
  ( ! [X103,X102] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK33('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X102,sK13))))) | p(s(bool,X102)) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))),s(sK13,X103))))) ) | $spl0),
  inference(superposition,[],[f815,f913])).
fof(f8557,plain,(
  ( ! [X98] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK32(X98,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13))))) | p(s(bool,X98))) ) | $spl344),
  inference(cnf_transformation,[],[f8557_D])).
fof(f8557_D,plain,(
  ( ! [X98] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK32(X98,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13))))) | p(s(bool,X98))) ) <=> ~$spl344),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl344])])).
fof(f8558,plain,(
  ( ! [X101,X100] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK32(X101,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X100)))) | p(s(bool,X101))) ) | $spl0),
  inference(forward_demodulation,[],[f8174,f913])).
fof(f8174,plain,(
  ( ! [X101,X100] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X100)))) | p(s(bool,X101)) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))),s(sK13,sK32(X101,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13)))))) ) | $spl0),
  inference(superposition,[],[f812,f913])).
fof(f8556,plain,(
  ( ! [X99,X98] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X99)))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK32(X98,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13))))) | p(s(bool,X98))) ) | $spl0),
  inference(forward_demodulation,[],[f8173,f913])).
fof(f8173,plain,(
  ( ! [X99,X98] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK32(X98,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13))))) | p(s(bool,X98)) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))),s(sK13,X99))))) ) | $spl0),
  inference(superposition,[],[f812,f913])).
fof(f8554,plain,(
  ( ! [X94] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK31(X94,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13))))) | p(s(bool,X94))) ) | $spl342),
  inference(cnf_transformation,[],[f8554_D])).
fof(f8554_D,plain,(
  ( ! [X94] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK31(X94,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13))))) | p(s(bool,X94))) ) <=> ~$spl342),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl342])])).
fof(f8555,plain,(
  ( ! [X97,X96] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK31(X97,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X96)))) | p(s(bool,X97))) ) | $spl0),
  inference(forward_demodulation,[],[f8172,f913])).
fof(f8172,plain,(
  ( ! [X97,X96] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X96)))) | p(s(bool,X97)) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))),s(sK13,sK31(X97,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13)))))) ) | $spl0),
  inference(superposition,[],[f810,f913])).
fof(f8553,plain,(
  ( ! [X94,X95] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X95)))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK31(X94,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13))))) | p(s(bool,X94))) ) | $spl0),
  inference(forward_demodulation,[],[f8171,f913])).
fof(f8171,plain,(
  ( ! [X94,X95] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK31(X94,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13))))) | p(s(bool,X94)) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))),s(sK13,X95))))) ) | $spl0),
  inference(superposition,[],[f810,f913])).
fof(f8551,plain,(
  ( ! [X90] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK30('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X90,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X90),s(sK13,sK29('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X90,sK13)))))) ) | $spl340),
  inference(cnf_transformation,[],[f8551_D])).
fof(f8551_D,plain,(
  ( ! [X90] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK30('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X90,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X90),s(sK13,sK29('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X90,sK13)))))) ) <=> ~$spl340),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl340])])).
fof(f8552,plain,(
  ( ! [X92,X93] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK30('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X93,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X92)))) | ~p(s(bool,happ(s(fun(sK13,bool),X93),s(sK13,sK29('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X93,sK13)))))) ) | $spl0),
  inference(forward_demodulation,[],[f8170,f913])).
fof(f8170,plain,(
  ( ! [X92,X93] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X92)))) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))),s(sK13,sK30('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X93,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X93),s(sK13,sK29('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X93,sK13)))))) ) | $spl0),
  inference(superposition,[],[f809,f913])).
fof(f8550,plain,(
  ( ! [X90,X91] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X91)))) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK30('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X90,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X90),s(sK13,sK29('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X90,sK13)))))) ) | $spl0),
  inference(forward_demodulation,[],[f8169,f913])).
fof(f8169,plain,(
  ( ! [X90,X91] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK30('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X90,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))),s(sK13,X91)))) | ~p(s(bool,happ(s(fun(sK13,bool),X90),s(sK13,sK29('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X90,sK13)))))) ) | $spl0),
  inference(superposition,[],[f809,f913])).
fof(f8548,plain,(
  ( ! [X82] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK29(X82,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X82),s(sK13,sK30(X82,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13)))))) ) | $spl338),
  inference(cnf_transformation,[],[f8548_D])).
fof(f8548_D,plain,(
  ( ! [X82] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK29(X82,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X82),s(sK13,sK30(X82,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13)))))) ) <=> ~$spl338),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl338])])).
fof(f8549,plain,(
  ( ! [X87,X86] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK29(X87,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13))))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X86)))) | ~p(s(bool,happ(s(fun(sK13,bool),X87),s(sK13,sK30(X87,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13)))))) ) | $spl0),
  inference(forward_demodulation,[],[f8167,f913])).
fof(f8167,plain,(
  ( ! [X87,X86] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X86)))) | ~p(s(bool,happ(s(fun(sK13,bool),X87),s(sK13,sK30(X87,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))),s(sK13,sK29(X87,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13)))))) ) | $spl0),
  inference(superposition,[],[f808,f913])).
fof(f8547,plain,(
  ( ! [X83,X82] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X83)))) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK29(X82,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X82),s(sK13,sK30(X82,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13)))))) ) | $spl0),
  inference(forward_demodulation,[],[f8165,f913])).
fof(f8165,plain,(
  ( ! [X83,X82] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK29(X82,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13))))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))),s(sK13,X83)))) | ~p(s(bool,happ(s(fun(sK13,bool),X82),s(sK13,sK30(X82,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13)))))) ) | $spl0),
  inference(superposition,[],[f808,f913])).
fof(f8545,plain,(
  ( ! [X78] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK28('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X78,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X78),s(sK13,sK28('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X78,sK13)))))) ) | $spl336),
  inference(cnf_transformation,[],[f8545_D])).
fof(f8545_D,plain,(
  ( ! [X78] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK28('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X78,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X78),s(sK13,sK28('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X78,sK13)))))) ) <=> ~$spl336),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl336])])).
fof(f8546,plain,(
  ( ! [X80,X81] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK28('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X81,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X80)))) | ~p(s(bool,happ(s(fun(sK13,bool),X81),s(sK13,sK28('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X81,sK13)))))) ) | $spl0),
  inference(forward_demodulation,[],[f8164,f913])).
fof(f8164,plain,(
  ( ! [X80,X81] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X80)))) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))),s(sK13,sK28('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X81,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X81),s(sK13,sK28('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X81,sK13)))))) ) | $spl0),
  inference(superposition,[],[f807,f913])).
fof(f8544,plain,(
  ( ! [X78,X79] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X79)))) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK28('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X78,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X78),s(sK13,sK28('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X78,sK13)))))) ) | $spl0),
  inference(forward_demodulation,[],[f8163,f913])).
fof(f8163,plain,(
  ( ! [X78,X79] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK28('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X78,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))),s(sK13,X79)))) | ~p(s(bool,happ(s(fun(sK13,bool),X78),s(sK13,sK28('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X78,sK13)))))) ) | $spl0),
  inference(superposition,[],[f807,f913])).
fof(f8542,plain,(
  ( ! [X70] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK28(X70,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X70),s(sK13,sK28(X70,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13)))))) ) | $spl334),
  inference(cnf_transformation,[],[f8542_D])).
fof(f8542_D,plain,(
  ( ! [X70] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK28(X70,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X70),s(sK13,sK28(X70,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13)))))) ) <=> ~$spl334),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl334])])).
fof(f8543,plain,(
  ( ! [X74,X75] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK28(X75,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13))))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X74)))) | ~p(s(bool,happ(s(fun(sK13,bool),X75),s(sK13,sK28(X75,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13)))))) ) | $spl0),
  inference(forward_demodulation,[],[f8161,f913])).
fof(f8161,plain,(
  ( ! [X74,X75] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X74)))) | ~p(s(bool,happ(s(fun(sK13,bool),X75),s(sK13,sK28(X75,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))),s(sK13,sK28(X75,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13)))))) ) | $spl0),
  inference(superposition,[],[f806,f913])).
fof(f8541,plain,(
  ( ! [X70,X71] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X71)))) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK28(X70,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X70),s(sK13,sK28(X70,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13)))))) ) | $spl0),
  inference(forward_demodulation,[],[f8159,f913])).
fof(f8159,plain,(
  ( ! [X70,X71] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK28(X70,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13))))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))),s(sK13,X71)))) | ~p(s(bool,happ(s(fun(sK13,bool),X70),s(sK13,sK28(X70,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13)))))) ) | $spl0),
  inference(superposition,[],[f806,f913])).
fof(f8539,plain,(
  ( ! [X66] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK27('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X66,sK13))))) | p(s(bool,X66))) ) | $spl332),
  inference(cnf_transformation,[],[f8539_D])).
fof(f8539_D,plain,(
  ( ! [X66] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK27('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X66,sK13))))) | p(s(bool,X66))) ) <=> ~$spl332),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl332])])).
fof(f8540,plain,(
  ( ! [X68,X69] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK27('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X69,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X68)))) | p(s(bool,X69))) ) | $spl0),
  inference(forward_demodulation,[],[f8158,f913])).
fof(f8158,plain,(
  ( ! [X68,X69] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X68)))) | p(s(bool,X69)) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))),s(sK13,sK27('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X69,sK13)))))) ) | $spl0),
  inference(superposition,[],[f805,f913])).
fof(f8538,plain,(
  ( ! [X66,X67] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X67)))) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK27('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X66,sK13))))) | p(s(bool,X66))) ) | $spl0),
  inference(forward_demodulation,[],[f8157,f913])).
fof(f8157,plain,(
  ( ! [X66,X67] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK27('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X66,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))),s(sK13,X67)))) | p(s(bool,X66))) ) | $spl0),
  inference(superposition,[],[f805,f913])).
fof(f8536,plain,(
  ( ! [X62] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK26('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X62,sK13))))) | p(s(bool,X62))) ) | $spl330),
  inference(cnf_transformation,[],[f8536_D])).
fof(f8536_D,plain,(
  ( ! [X62] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK26('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X62,sK13))))) | p(s(bool,X62))) ) <=> ~$spl330),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl330])])).
fof(f8537,plain,(
  ( ! [X64,X65] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK26('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X65,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X64)))) | p(s(bool,X65))) ) | $spl0),
  inference(forward_demodulation,[],[f8156,f913])).
fof(f8156,plain,(
  ( ! [X64,X65] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X64)))) | p(s(bool,X65)) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))),s(sK13,sK26('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X65,sK13)))))) ) | $spl0),
  inference(superposition,[],[f803,f913])).
fof(f8535,plain,(
  ( ! [X62,X63] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X63)))) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK26('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X62,sK13))))) | p(s(bool,X62))) ) | $spl0),
  inference(forward_demodulation,[],[f8155,f913])).
fof(f8155,plain,(
  ( ! [X62,X63] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK26('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X62,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))),s(sK13,X63)))) | p(s(bool,X62))) ) | $spl0),
  inference(superposition,[],[f803,f913])).
fof(f8533,plain,(
  ( ! [X58] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK25('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X58,sK13))))) | p(s(bool,X58))) ) | $spl328),
  inference(cnf_transformation,[],[f8533_D])).
fof(f8533_D,plain,(
  ( ! [X58] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK25('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X58,sK13))))) | p(s(bool,X58))) ) <=> ~$spl328),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl328])])).
fof(f8534,plain,(
  ( ! [X61,X60] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK25('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X61,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X60)))) | p(s(bool,X61))) ) | $spl0),
  inference(forward_demodulation,[],[f8154,f913])).
fof(f8154,plain,(
  ( ! [X61,X60] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X60)))) | p(s(bool,X61)) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))),s(sK13,sK25('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X61,sK13)))))) ) | $spl0),
  inference(superposition,[],[f800,f913])).
fof(f8532,plain,(
  ( ! [X59,X58] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X59)))) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK25('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X58,sK13))))) | p(s(bool,X58))) ) | $spl0),
  inference(forward_demodulation,[],[f8153,f913])).
fof(f8153,plain,(
  ( ! [X59,X58] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK25('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X58,sK13))))) | p(s(bool,X58)) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))),s(sK13,X59))))) ) | $spl0),
  inference(superposition,[],[f800,f913])).
fof(f8529,plain,(
  ( ! [X54] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK24('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X54,sK13))))) | p(s(bool,X54))) ) | $spl324),
  inference(cnf_transformation,[],[f8529_D])).
fof(f8529_D,plain,(
  ( ! [X54] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK24('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X54,sK13))))) | p(s(bool,X54))) ) <=> ~$spl324),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl324])])).
fof(f8531,plain,(
  ( ! [X57,X56] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK24('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X57,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X56)))) | p(s(bool,X57))) ) | $spl0),
  inference(forward_demodulation,[],[f8152,f913])).
fof(f8152,plain,(
  ( ! [X57,X56] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X56)))) | p(s(bool,X57)) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))),s(sK13,sK24('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X57,sK13)))))) ) | $spl0),
  inference(superposition,[],[f798,f913])).
fof(f8528,plain,(
  ( ! [X54,X55] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X55)))) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK24('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X54,sK13))))) | p(s(bool,X54))) ) | $spl0),
  inference(forward_demodulation,[],[f8151,f913])).
fof(f8151,plain,(
  ( ! [X54,X55] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK24('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X54,sK13))))) | p(s(bool,X54)) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))),s(sK13,X55))))) ) | $spl0),
  inference(superposition,[],[f798,f913])).
fof(f8526,plain,(
  ( ! [X50] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK23('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X50,sK13))))) | ~p(s(bool,X50))) ) | $spl322),
  inference(cnf_transformation,[],[f8526_D])).
fof(f8526_D,plain,(
  ( ! [X50] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK23('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X50,sK13))))) | ~p(s(bool,X50))) ) <=> ~$spl322),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl322])])).
fof(f8527,plain,(
  ( ! [X52,X53] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK23('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X53,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X52)))) | ~p(s(bool,X53))) ) | $spl0),
  inference(forward_demodulation,[],[f8150,f913])).
fof(f8150,plain,(
  ( ! [X52,X53] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X52)))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))),s(sK13,sK23('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X53,sK13))))) | ~p(s(bool,X53))) ) | $spl0),
  inference(superposition,[],[f797,f913])).
fof(f8525,plain,(
  ( ! [X50,X51] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X51)))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK23('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X50,sK13))))) | ~p(s(bool,X50))) ) | $spl0),
  inference(forward_demodulation,[],[f8149,f913])).
fof(f8149,plain,(
  ( ! [X50,X51] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK23('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X50,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))),s(sK13,X51)))) | ~p(s(bool,X50))) ) | $spl0),
  inference(superposition,[],[f797,f913])).
fof(f8523,plain,(
  ( ! [X46] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK22('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X46,sK13))))) | ~p(s(bool,X46))) ) | $spl320),
  inference(cnf_transformation,[],[f8523_D])).
fof(f8523_D,plain,(
  ( ! [X46] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK22('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X46,sK13))))) | ~p(s(bool,X46))) ) <=> ~$spl320),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl320])])).
fof(f8524,plain,(
  ( ! [X48,X49] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK22('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X49,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X48)))) | ~p(s(bool,X49))) ) | $spl0),
  inference(forward_demodulation,[],[f8148,f913])).
fof(f8148,plain,(
  ( ! [X48,X49] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X48)))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))),s(sK13,sK22('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X49,sK13))))) | ~p(s(bool,X49))) ) | $spl0),
  inference(superposition,[],[f795,f913])).
fof(f8522,plain,(
  ( ! [X47,X46] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X47)))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK22('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X46,sK13))))) | ~p(s(bool,X46))) ) | $spl0),
  inference(forward_demodulation,[],[f8147,f913])).
fof(f8147,plain,(
  ( ! [X47,X46] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK22('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X46,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))),s(sK13,X47)))) | ~p(s(bool,X46))) ) | $spl0),
  inference(superposition,[],[f795,f913])).
fof(f8516,plain,(
  ( ! [X35] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK20(X35,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13))))) | ~p(s(bool,X35))) ) | $spl318),
  inference(cnf_transformation,[],[f8516_D])).
fof(f8516_D,plain,(
  ( ! [X35] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK20(X35,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13))))) | ~p(s(bool,X35))) ) <=> ~$spl318),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl318])])).
fof(f8517,plain,(
  ( ! [X37,X38] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK20(X38,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X37)))) | ~p(s(bool,X38))) ) | $spl0),
  inference(forward_demodulation,[],[f8141,f913])).
fof(f8141,plain,(
  ( ! [X37,X38] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X37)))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))),s(sK13,sK20(X38,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13))))) | ~p(s(bool,X38))) ) | $spl0),
  inference(superposition,[],[f789,f913])).
fof(f8515,plain,(
  ( ! [X35,X36] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X36)))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK20(X35,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13))))) | ~p(s(bool,X35))) ) | $spl0),
  inference(forward_demodulation,[],[f8140,f913])).
fof(f8140,plain,(
  ( ! [X35,X36] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK20(X35,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13))))) | ~p(s(bool,X35)) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))),s(sK13,X36))))) ) | $spl0),
  inference(superposition,[],[f789,f913])).
fof(f8513,plain,(
  ( ! [X31] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK19(X31,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13))))) | ~p(s(bool,X31))) ) | $spl316),
  inference(cnf_transformation,[],[f8513_D])).
fof(f8513_D,plain,(
  ( ! [X31] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK19(X31,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13))))) | ~p(s(bool,X31))) ) <=> ~$spl316),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl316])])).
fof(f8514,plain,(
  ( ! [X33,X34] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK19(X34,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X33)))) | ~p(s(bool,X34))) ) | $spl0),
  inference(forward_demodulation,[],[f8139,f913])).
fof(f8139,plain,(
  ( ! [X33,X34] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X33)))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))),s(sK13,sK19(X34,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13))))) | ~p(s(bool,X34))) ) | $spl0),
  inference(superposition,[],[f787,f913])).
fof(f8512,plain,(
  ( ! [X31,X32] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X32)))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK19(X31,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13))))) | ~p(s(bool,X31))) ) | $spl0),
  inference(forward_demodulation,[],[f8138,f913])).
fof(f8138,plain,(
  ( ! [X31,X32] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK19(X31,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13))))) | ~p(s(bool,X31)) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))),s(sK13,X32))))) ) | $spl0),
  inference(superposition,[],[f787,f913])).
fof(f8510,plain,(
  ( ! [X27] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK18('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X27,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X27),s(sK13,sK18('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X27,sK13)))))) ) | $spl314),
  inference(cnf_transformation,[],[f8510_D])).
fof(f8510_D,plain,(
  ( ! [X27] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK18('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X27,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X27),s(sK13,sK18('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X27,sK13)))))) ) <=> ~$spl314),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl314])])).
fof(f8511,plain,(
  ( ! [X30,X29] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK18('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X30,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X29)))) | p(s(bool,happ(s(fun(sK13,bool),X30),s(sK13,sK18('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X30,sK13)))))) ) | $spl0),
  inference(forward_demodulation,[],[f8137,f913])).
fof(f8137,plain,(
  ( ! [X30,X29] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X29)))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))),s(sK13,sK18('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X30,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X30),s(sK13,sK18('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X30,sK13)))))) ) | $spl0),
  inference(superposition,[],[f786,f913])).
fof(f8509,plain,(
  ( ! [X28,X27] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X28)))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK18('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X27,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X27),s(sK13,sK18('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X27,sK13)))))) ) | $spl0),
  inference(forward_demodulation,[],[f8136,f913])).
fof(f8136,plain,(
  ( ! [X28,X27] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK18('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X27,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X27),s(sK13,sK18('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X27,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))),s(sK13,X28))))) ) | $spl0),
  inference(superposition,[],[f786,f913])).
fof(f8507,plain,(
  ( ! [X19] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK18(X19,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X19),s(sK13,sK18(X19,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13)))))) ) | $spl312),
  inference(cnf_transformation,[],[f8507_D])).
fof(f8507_D,plain,(
  ( ! [X19] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK18(X19,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X19),s(sK13,sK18(X19,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13)))))) ) <=> ~$spl312),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl312])])).
fof(f8508,plain,(
  ( ! [X24,X23] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK18(X24,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X23)))) | p(s(bool,happ(s(fun(sK13,bool),X24),s(sK13,sK18(X24,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13)))))) ) | $spl0),
  inference(forward_demodulation,[],[f8134,f913])).
fof(f8134,plain,(
  ( ! [X24,X23] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X23)))) | p(s(bool,happ(s(fun(sK13,bool),X24),s(sK13,sK18(X24,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13))))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))),s(sK13,sK18(X24,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13)))))) ) | $spl0),
  inference(superposition,[],[f785,f913])).
fof(f8506,plain,(
  ( ! [X19,X20] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X20)))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK18(X19,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X19),s(sK13,sK18(X19,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13)))))) ) | $spl0),
  inference(forward_demodulation,[],[f8132,f913])).
fof(f8132,plain,(
  ( ! [X19,X20] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK18(X19,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X19),s(sK13,sK18(X19,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))),s(sK13,X20))))) ) | $spl0),
  inference(superposition,[],[f785,f913])).
fof(f8504,plain,(
  ( ! [X15] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK17('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X15,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X15),s(sK13,sK16('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X15,sK13)))))) ) | $spl310),
  inference(cnf_transformation,[],[f8504_D])).
fof(f8504_D,plain,(
  ( ! [X15] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK17('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X15,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X15),s(sK13,sK16('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X15,sK13)))))) ) <=> ~$spl310),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl310])])).
fof(f8505,plain,(
  ( ! [X17,X18] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK17('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X18,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X17)))) | p(s(bool,happ(s(fun(sK13,bool),X18),s(sK13,sK16('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X18,sK13)))))) ) | $spl0),
  inference(forward_demodulation,[],[f8131,f913])).
fof(f8131,plain,(
  ( ! [X17,X18] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X17)))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))),s(sK13,sK17('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X18,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X18),s(sK13,sK16('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X18,sK13)))))) ) | $spl0),
  inference(superposition,[],[f784,f913])).
fof(f8503,plain,(
  ( ! [X15,X16] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X16)))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK17('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X15,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X15),s(sK13,sK16('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X15,sK13)))))) ) | $spl0),
  inference(forward_demodulation,[],[f8130,f913])).
fof(f8130,plain,(
  ( ! [X15,X16] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK17('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X15,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X15),s(sK13,sK16('const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),X15,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))),s(sK13,X16))))) ) | $spl0),
  inference(superposition,[],[f784,f913])).
fof(f8500,plain,(
  ( ! [X7] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK16(X7,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X7),s(sK13,sK17(X7,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13)))))) ) | $spl306),
  inference(cnf_transformation,[],[f8500_D])).
fof(f8500_D,plain,(
  ( ! [X7] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK16(X7,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X7),s(sK13,sK17(X7,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13)))))) ) <=> ~$spl306),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl306])])).
fof(f8502,plain,(
  ( ! [X12,X11] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK16(X12,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X11)))) | p(s(bool,happ(s(fun(sK13,bool),X12),s(sK13,sK17(X12,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13)))))) ) | $spl0),
  inference(forward_demodulation,[],[f8128,f913])).
fof(f8128,plain,(
  ( ! [X12,X11] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X11)))) | p(s(bool,happ(s(fun(sK13,bool),X12),s(sK13,sK17(X12,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13))))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))),s(sK13,sK16(X12,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13)))))) ) | $spl0),
  inference(superposition,[],[f783,f913])).
fof(f8499,plain,(
  ( ! [X8,X7] : (~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,X8)))) | p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK16(X7,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X7),s(sK13,sK17(X7,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13)))))) ) | $spl0),
  inference(forward_demodulation,[],[f8126,f913])).
fof(f8126,plain,(
  ( ! [X8,X7] : (p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(sK13,sK16(X7,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X7),s(sK13,sK17(X7,'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)),sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))),s(sK13,X8))))) ) | $spl0),
  inference(superposition,[],[f783,f913])).
fof(f8496,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14)))))) | $spl303),
  inference(cnf_transformation,[],[f8496_D])).
fof(f8496_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14)))))) <=> ~$spl303),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl303])])).
fof(f8494,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14)))))) | ($spl0 | $spl8 | $spl10)),
  inference(forward_demodulation,[],[f8493,f901])).
fof(f8493,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14)))))) | ($spl0 | $spl8 | $spl10)),
  inference(subsumption_resolution,[],[f8492,f1089])).
fof(f8492,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14)))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) | ($spl0 | $spl10)),
  inference(subsumption_resolution,[],[f8124,f1092])).
fof(f8124,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14)))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) | $spl0),
  inference(superposition,[],[f886,f913])).
fof(f8487,plain,(
  s(fun(sK13,bool),sK14) = s(fun(sK13,bool),'const.pred_set.UNION_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),sK15))) | $spl300),
  inference(cnf_transformation,[],[f8487_D])).
fof(f8487_D,plain,(
  s(fun(sK13,bool),sK14) = s(fun(sK13,bool),'const.pred_set.UNION_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),sK15))) <=> ~$spl300),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl300])])).
fof(f8112,plain,(
  s(fun(sK13,bool),sK14) = s(fun(sK13,bool),'const.pred_set.UNION_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),sK15))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)))) | $spl0),
  inference(superposition,[],[f862,f913])).
fof(f862,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X2))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f442])).
fof(f442,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))) | (s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X2),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) & s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X2)))))),
  inference(ennf_transformation,[],[f350])).
fof(f350,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))) => (s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X2),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) & s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X2)))))),
  inference(flattening,[],[f349])).
fof(f349,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))) => (s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X2),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) & s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X2)))))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X6] : ! [X8,X7] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))) => (s(fun(X6,bool),X8) = s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),X7),s(fun(X6,bool),'const.pred_set.DIFF_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X7))))) & s(fun(X6,bool),X8) = s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),'const.pred_set.DIFF_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X7))),s(fun(X6,bool),X7)))))),
  file('thm.pred_set.SUBSET_DIFF_EMPTY','thm.pred_set.UNION_DIFF')).
fof(f8485,plain,(
  s(fun(sK13,bool),sK14) = s(fun(sK13,bool),'const.pred_set.UNION_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | $spl298),
  inference(cnf_transformation,[],[f8485_D])).
fof(f8485_D,plain,(
  s(fun(sK13,bool),sK14) = s(fun(sK13,bool),'const.pred_set.UNION_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) <=> ~$spl298),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl298])])).
fof(f8111,plain,(
  s(fun(sK13,bool),sK14) = s(fun(sK13,bool),'const.pred_set.UNION_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)))) | $spl0),
  inference(superposition,[],[f861,f913])).
fof(f861,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X2),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f442])).
fof(f8464,plain,(
  s(fun(sK13,bool),'const.pred_set.EMPTY_0') = s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | $spl296),
  inference(cnf_transformation,[],[f8464_D])).
fof(f8464_D,plain,(
  s(fun(sK13,bool),'const.pred_set.EMPTY_0') = s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) <=> ~$spl296),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl296])])).
fof(f8462,plain,(
  s(fun(sK13,bool),'const.pred_set.EMPTY_0') = s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f8096,f1089])).
fof(f8096,plain,(
  s(fun(sK13,bool),'const.pred_set.EMPTY_0') = s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) | $spl0),
  inference(superposition,[],[f913,f897])).
fof(f8076,plain,(
  ( ! [X21,X20] : (sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X20,X21) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15)))) ) | $spl8),
  inference(subsumption_resolution,[],[f8068,f1089])).
fof(f8068,plain,(
  ( ! [X21,X20] : (sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X20,X21) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))))) ) | $spl8),
  inference(superposition,[],[f1670,f897])).
fof(f8065,plain,(
  ( ! [X12,X13] : (sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl8),
  inference(superposition,[],[f1670,f864])).
fof(f8056,plain,(
  ( ! [X21,X20] : (sP6(X20,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X21) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15)))) ) | $spl8),
  inference(subsumption_resolution,[],[f8048,f1089])).
fof(f8048,plain,(
  ( ! [X21,X20] : (sP6(X20,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X21) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))))) ) | $spl8),
  inference(superposition,[],[f1669,f897])).
fof(f8045,plain,(
  ( ! [X12,X13] : (sP6(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl8),
  inference(superposition,[],[f1669,f864])).
fof(f8036,plain,(
  ( ! [X21,X20] : (sP6(X20,X21,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15)))) ) | $spl8),
  inference(subsumption_resolution,[],[f8028,f1089])).
fof(f8028,plain,(
  ( ! [X21,X20] : (sP6(X20,X21,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))))) ) | $spl8),
  inference(superposition,[],[f1668,f897])).
fof(f8025,plain,(
  ( ! [X12,X13] : (sP6(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl8),
  inference(superposition,[],[f1668,f864])).
fof(f8016,plain,(
  ( ! [X21,X20] : (sP7(X20,X21,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15)))) ) | $spl8),
  inference(subsumption_resolution,[],[f8009,f1089])).
fof(f8009,plain,(
  ( ! [X21,X20] : (sP7(X20,X21,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))))) ) | $spl8),
  inference(superposition,[],[f1667,f897])).
fof(f8006,plain,(
  ( ! [X12,X13] : (sP7(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl8),
  inference(superposition,[],[f1667,f864])).
fof(f7997,plain,(
  ( ! [X21,X20] : (sP8(X20,X21,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15)))) ) | $spl8),
  inference(subsumption_resolution,[],[f7990,f1089])).
fof(f7990,plain,(
  ( ! [X21,X20] : (sP8(X20,X21,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))))) ) | $spl8),
  inference(superposition,[],[f1664,f897])).
fof(f7987,plain,(
  ( ! [X12,X13] : (sP8(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl8),
  inference(superposition,[],[f1664,f864])).
fof(f7978,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f7947,f1089])).
fof(f7947,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) | $spl3),
  inference(superposition,[],[f916,f897])).
fof(f7944,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) | $spl3),
  inference(superposition,[],[f916,f864])).
fof(f7925,plain,(
  ( ! [X215,X214] : (p(s(bool,X214)) | ~p(s(bool,X215)) | s(bool,X215) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))) ) | $spl3),
  inference(resolution,[],[f916,f905])).
fof(f7827,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))) | s(bool,X23) != s(bool,X25)) ) | $spl3),
  inference(resolution,[],[f916,f630])).
fof(f7825,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))) ) | $spl3),
  inference(resolution,[],[f916,f629])).
fof(f7823,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))) ) | $spl3),
  inference(resolution,[],[f916,f628])).
fof(f7822,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X9) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))) | s(bool,X8) != s(bool,X10)) ) | $spl3),
  inference(resolution,[],[f916,f626])).
fof(f7820,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))) ) | $spl3),
  inference(resolution,[],[f916,f624])).
fof(f7809,plain,(
  ( ! [X21,X20] : (sP4(X20,X21,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15)))) ) | $spl8),
  inference(subsumption_resolution,[],[f7802,f1089])).
fof(f7802,plain,(
  ( ! [X21,X20] : (sP4(X20,X21,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))))) ) | $spl8),
  inference(superposition,[],[f1651,f897])).
fof(f7799,plain,(
  ( ! [X12,X13] : (sP4(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl8),
  inference(superposition,[],[f1651,f864])).
fof(f7789,plain,(
  ( ! [X21,X20] : (sP2(X20,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X21) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15)))) ) | $spl8),
  inference(subsumption_resolution,[],[f7783,f1089])).
fof(f7783,plain,(
  ( ! [X21,X20] : (sP2(X20,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X21) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))))) ) | $spl8),
  inference(superposition,[],[f1635,f897])).
fof(f7780,plain,(
  ( ! [X12,X13] : (sP2(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl8),
  inference(superposition,[],[f1635,f864])).
fof(f7770,plain,(
  ( ! [X21,X20] : (sP2(X20,X21,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15)))) ) | $spl8),
  inference(subsumption_resolution,[],[f7764,f1089])).
fof(f7764,plain,(
  ( ! [X21,X20] : (sP2(X20,X21,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))))) ) | $spl8),
  inference(superposition,[],[f1634,f897])).
fof(f7761,plain,(
  ( ! [X12,X13] : (sP2(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl8),
  inference(superposition,[],[f1634,f864])).
fof(f7751,plain,(
  ( ! [X21,X20] : (sP0(X20,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X21) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15)))) ) | $spl8),
  inference(subsumption_resolution,[],[f7745,f1089])).
fof(f7745,plain,(
  ( ! [X21,X20] : (sP0(X20,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X21) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))))) ) | $spl8),
  inference(superposition,[],[f1621,f897])).
fof(f7742,plain,(
  ( ! [X12,X13] : (sP0(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl8),
  inference(superposition,[],[f1621,f864])).
fof(f7698,plain,(
  ( ! [X227,X226] : (p(s(bool,X226)) | p(s(bool,X227)) | s(bool,X227) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | $spl8),
  inference(resolution,[],[f1089,f905])).
fof(f7600,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) | s(bool,X38) != s(bool,X39)) ) | $spl8),
  inference(resolution,[],[f1089,f630])).
fof(f7597,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X28)) | ~p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X30) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | $spl8),
  inference(resolution,[],[f1089,f629])).
fof(f7596,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | $spl8),
  inference(resolution,[],[f1089,f628])).
fof(f7595,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) | s(bool,X23) != s(bool,X24)) ) | $spl8),
  inference(resolution,[],[f1089,f626])).
fof(f7593,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X16) != s(bool,X18) | s(bool,X17) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | $spl8),
  inference(resolution,[],[f1089,f624])).
fof(f7592,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) | s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),sK15) | $spl8),
  inference(resolution,[],[f1089,f898])).
fof(f7591,plain,(
  s(fun(sK13,bool),'const.pred_set.EMPTY_0') = s(fun(sK13,bool),sK15) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) | $spl8),
  inference(resolution,[],[f1089,f897])).
fof(f1090,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) | $spl9),
  inference(cnf_transformation,[],[f1090_D])).
fof(f1090_D,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f7573,plain,(
  ( ! [X5] : (s(fun(sK13,bool),sK14) != s(fun(sK13,bool),X5) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK15))))) ) | $spl294),
  inference(cnf_transformation,[],[f7573_D])).
fof(f7573_D,plain,(
  ( ! [X5] : (s(fun(sK13,bool),sK14) != s(fun(sK13,bool),X5) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK15))))) ) <=> ~$spl294),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl294])])).
fof(f7572,plain,(
  ( ! [X5] : (s(fun(sK13,bool),sK14) != s(fun(sK13,bool),X5) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK15)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))))) ) | $spl7),
  inference(subsumption_resolution,[],[f6058,f860])).
fof(f6058,plain,(
  ( ! [X5] : (s(fun(sK13,bool),sK14) != s(fun(sK13,bool),X5) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK15)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X5))))) ) | $spl7),
  inference(superposition,[],[f1087,f866])).
fof(f1087,plain,(
  s(fun(sK13,bool),sK14) != s(fun(sK13,bool),sK15) | $spl7),
  inference(cnf_transformation,[],[f1087_D])).
fof(f1087_D,plain,(
  s(fun(sK13,bool),sK14) != s(fun(sK13,bool),sK15) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f5523,plain,(
  ( ! [X5] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK15))))) ) | $spl256),
  inference(cnf_transformation,[],[f5523_D])).
fof(f5523_D,plain,(
  ( ! [X5] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK15))))) ) <=> ~$spl256),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl256])])).
fof(f7570,plain,(
  ( ! [X5] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK15)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))))) ) | $spl33),
  inference(subsumption_resolution,[],[f6070,f860])).
fof(f6070,plain,(
  ( ! [X5] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK15)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X5))))) ) | $spl33),
  inference(superposition,[],[f1111,f866])).
fof(f7568,plain,(
  ( ! [X5] : (s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),X5) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK15))))) ) | $spl292),
  inference(cnf_transformation,[],[f7568_D])).
fof(f7568_D,plain,(
  ( ! [X5] : (s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),X5) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK15))))) ) <=> ~$spl292),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl292])])).
fof(f7567,plain,(
  ( ! [X5] : (s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),X5) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK15)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))))) ) | $spl51),
  inference(subsumption_resolution,[],[f6087,f860])).
fof(f6087,plain,(
  ( ! [X5] : (s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),X5) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK15)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X5))))) ) | $spl51),
  inference(superposition,[],[f1746,f866])).
fof(f7565,plain,(
  ( ! [X5] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK15))))) ) | $spl290),
  inference(cnf_transformation,[],[f7565_D])).
fof(f7565_D,plain,(
  ( ! [X5] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK15))))) ) <=> ~$spl290),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl290])])).
fof(f7564,plain,(
  ( ! [X5] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK15)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))))) ) | $spl5),
  inference(subsumption_resolution,[],[f7302,f860])).
fof(f7302,plain,(
  ( ! [X5] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK15)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X5))))) ) | $spl5),
  inference(superposition,[],[f1086,f866])).
fof(f7502,plain,(
  ( ! [X215,X214] : (p(s(bool,X214)) | ~p(s(bool,X215)) | s(bool,X215) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) | $spl27),
  inference(resolution,[],[f1104,f905])).
fof(f7404,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s(bool,X23) != s(bool,X25)) ) | $spl27),
  inference(resolution,[],[f1104,f630])).
fof(f7402,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) | $spl27),
  inference(resolution,[],[f1104,f629])).
fof(f7400,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) | $spl27),
  inference(resolution,[],[f1104,f628])).
fof(f7399,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X9) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s(bool,X8) != s(bool,X10)) ) | $spl27),
  inference(resolution,[],[f1104,f626])).
fof(f7397,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) | $spl27),
  inference(resolution,[],[f1104,f624])).
fof(f4307,plain,(
  ( ! [X12,X13] : (~sP10('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | p(s(bool,X12))) ) | $spl188),
  inference(cnf_transformation,[],[f4307_D])).
fof(f4307_D,plain,(
  ( ! [X12,X13] : (~sP10('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | p(s(bool,X12))) ) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f7376,plain,(
  ( ! [X12,X13] : (~sP10('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f5598,f864])).
fof(f5598,plain,(
  ( ! [X152,X153] : (~sP10('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)),X152,X153) | p(s(bool,X152))) ) | $spl9),
  inference(resolution,[],[f1090,f712])).
fof(f712,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | ~sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f521])).
fof(f4288,plain,(
  ( ! [X12,X13] : (~sP10(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | p(s(bool,X12))) ) | $spl186),
  inference(cnf_transformation,[],[f4288_D])).
fof(f4288_D,plain,(
  ( ! [X12,X13] : (~sP10(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | p(s(bool,X12))) ) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f7356,plain,(
  ( ! [X12,X13] : (~sP10(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f5597,f864])).
fof(f5597,plain,(
  ( ! [X151,X150] : (~sP10(X150,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)),X151) | p(s(bool,X150))) ) | $spl9),
  inference(resolution,[],[f1090,f712])).
fof(f7312,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) | $spl5),
  inference(superposition,[],[f1086,f864])).
fof(f7293,plain,(
  ( ! [X215,X214] : (p(s(bool,X214)) | ~p(s(bool,X215)) | s(bool,X215) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)))) ) | $spl5),
  inference(resolution,[],[f1086,f905])).
fof(f7195,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))) | s(bool,X23) != s(bool,X25)) ) | $spl5),
  inference(resolution,[],[f1086,f630])).
fof(f7193,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)))) ) | $spl5),
  inference(resolution,[],[f1086,f629])).
fof(f7191,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)))) ) | $spl5),
  inference(resolution,[],[f1086,f628])).
fof(f7190,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X9) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))) | s(bool,X8) != s(bool,X10)) ) | $spl5),
  inference(resolution,[],[f1086,f626])).
fof(f7188,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)))) ) | $spl5),
  inference(resolution,[],[f1086,f624])).
fof(f4267,plain,(
  ( ! [X12,X13] : (~sP10(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X12))) ) | $spl184),
  inference(cnf_transformation,[],[f4267_D])).
fof(f4267_D,plain,(
  ( ! [X12,X13] : (~sP10(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X12))) ) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f7170,plain,(
  ( ! [X12,X13] : (~sP10(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f5596,f864])).
fof(f5596,plain,(
  ( ! [X149,X148] : (~sP10(X148,X149,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) | p(s(bool,X148))) ) | $spl9),
  inference(resolution,[],[f1090,f711])).
fof(f711,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X2)) | ~sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f521])).
fof(f4244,plain,(
  ( ! [X12,X13] : (~sP10('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | p(s(bool,X13))) ) | $spl182),
  inference(cnf_transformation,[],[f4244_D])).
fof(f4244_D,plain,(
  ( ! [X12,X13] : (~sP10('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | p(s(bool,X13))) ) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f7147,plain,(
  ( ! [X12,X13] : (~sP10('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | p(s(bool,X13)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f5595,f864])).
fof(f5595,plain,(
  ( ! [X146,X147] : (~sP10('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)),X147,X146) | p(s(bool,X146))) ) | $spl9),
  inference(resolution,[],[f1090,f711])).
fof(f4225,plain,(
  ( ! [X12,X13] : (sP5('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | p(s(bool,X12))) ) | $spl180),
  inference(cnf_transformation,[],[f4225_D])).
fof(f4225_D,plain,(
  ( ! [X12,X13] : (sP5('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | p(s(bool,X12))) ) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f7125,plain,(
  ( ! [X12,X13] : (sP5('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f5577,f864])).
fof(f5577,plain,(
  ( ! [X111,X110] : (sP5('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)),X110,X111) | p(s(bool,X110))) ) | $spl9),
  inference(resolution,[],[f1090,f677])).
fof(f677,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f504])).
fof(f504,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP5(X2,X0,X1)) & (~sP5(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f503])).
fof(f503,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP5(X2,X0,X1)) & (~sP5(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f466])).
fof(f466,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP5(X2,X0,X1))),
  inference(definition_folding,[],[f270,f465,f464])).
fof(f270,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f269])).
fof(f269,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f73])).
fof(f73,axiom,(
  ! [X25,X23,X24] : ((p(s(bool,X24)) <=> (p(s(bool,X23)) & p(s(bool,X25)))) <=> ((p(s(bool,X24)) | ~p(s(bool,X23)) | ~p(s(bool,X25))) & (p(s(bool,X23)) | ~p(s(bool,X24))) & (p(s(bool,X25)) | ~p(s(bool,X24)))))),
  file('thm.pred_set.SUBSET_DIFF_EMPTY','thm.sat.sat_23_prove_ltb')).
fof(f7102,plain,(
  ( ! [X12,X13] : (sP11(X12,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f1015,f864])).
fof(f7090,plain,(
  ( ! [X12,X13] : (sP11(X12,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),sK15)),X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f1015,f864])).
fof(f4169,plain,(
  ( ! [X12,X13] : (sP5(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | p(s(bool,X12))) ) | $spl174),
  inference(cnf_transformation,[],[f4169_D])).
fof(f4169_D,plain,(
  ( ! [X12,X13] : (sP5(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | p(s(bool,X12))) ) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f7073,plain,(
  ( ! [X12,X13] : (sP5(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f5576,f864])).
fof(f5576,plain,(
  ( ! [X109,X108] : (sP5(X108,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)),X109) | p(s(bool,X108))) ) | $spl9),
  inference(resolution,[],[f1090,f677])).
fof(f4147,plain,(
  ( ! [X12,X13] : (sP5('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | p(s(bool,X13))) ) | $spl172),
  inference(cnf_transformation,[],[f4147_D])).
fof(f4147_D,plain,(
  ( ! [X12,X13] : (sP5('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | p(s(bool,X13))) ) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f7049,plain,(
  ( ! [X12,X13] : (sP5('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | p(s(bool,X13)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f5575,f864])).
fof(f5575,plain,(
  ( ! [X107,X106] : (sP5('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)),X106,X107) | p(s(bool,X107))) ) | $spl9),
  inference(resolution,[],[f1090,f676])).
fof(f676,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f504])).
fof(f4124,plain,(
  ( ! [X12,X13] : (sP5(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X12))) ) | $spl170),
  inference(cnf_transformation,[],[f4124_D])).
fof(f4124_D,plain,(
  ( ! [X12,X13] : (sP5(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X12))) ) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f7028,plain,(
  ( ! [X12,X13] : (sP5(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f5574,f864])).
fof(f5574,plain,(
  ( ! [X105,X104] : (sP5(X104,X105,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) | p(s(bool,X104))) ) | $spl9),
  inference(resolution,[],[f1090,f676])).
fof(f4102,plain,(
  ( ! [X12,X13] : (~sP5(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | ~p(s(bool,X12))) ) | $spl168),
  inference(cnf_transformation,[],[f4102_D])).
fof(f4102_D,plain,(
  ( ! [X12,X13] : (~sP5(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | ~p(s(bool,X12))) ) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f7004,plain,(
  ( ! [X12,X13] : (~sP5(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | ~p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f5568,f864])).
fof(f5568,plain,(
  ( ! [X92,X93] : (~sP5(X92,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)),X93) | ~p(s(bool,X92))) ) | $spl9),
  inference(resolution,[],[f1090,f667])).
fof(f4080,plain,(
  ( ! [X12,X13] : (~sP5(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X12))) ) | $spl166),
  inference(cnf_transformation,[],[f4080_D])).
fof(f4080_D,plain,(
  ( ! [X12,X13] : (~sP5(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X12))) ) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f6983,plain,(
  ( ! [X12,X13] : (~sP5(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f5567,f864])).
fof(f5567,plain,(
  ( ! [X90,X91] : (~sP5(X90,X91,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) | ~p(s(bool,X90))) ) | $spl9),
  inference(resolution,[],[f1090,f666])).
fof(f6961,plain,(
  ( ! [X12,X13] : (sP10('const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f1008,f864])).
fof(f6949,plain,(
  ( ! [X12,X13] : (sP10('const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),sK15)),X12,X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f1008,f864])).
fof(f4023,plain,(
  ( ! [X12,X13] : (~sP3('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | ~p(s(bool,X12))) ) | $spl160),
  inference(cnf_transformation,[],[f4023_D])).
fof(f4023_D,plain,(
  ( ! [X12,X13] : (~sP3('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | ~p(s(bool,X12))) ) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f6928,plain,(
  ( ! [X12,X13] : (~sP3('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | ~p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f5552,f864])).
fof(f5552,plain,(
  ( ! [X61,X60] : (~sP3('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)),X60,X61) | ~p(s(bool,X60))) ) | $spl9),
  inference(resolution,[],[f1090,f649])).
fof(f3998,plain,(
  ( ! [X12,X13] : (~sP3('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | ~p(s(bool,X13))) ) | $spl158),
  inference(cnf_transformation,[],[f3998_D])).
fof(f3998_D,plain,(
  ( ! [X12,X13] : (~sP3('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | ~p(s(bool,X13))) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f6902,plain,(
  ( ! [X12,X13] : (~sP3('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | ~p(s(bool,X13)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f5551,f864])).
fof(f5551,plain,(
  ( ! [X59,X58] : (~sP3('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)),X59,X58) | ~p(s(bool,X58))) ) | $spl9),
  inference(resolution,[],[f1090,f648])).
fof(f6875,plain,(
  ( ! [X12,X13] : (sP6('const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f997,f864])).
fof(f6863,plain,(
  ( ! [X12,X13] : (sP6('const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),sK15)),X12,X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f997,f864])).
fof(f3939,plain,(
  ( ! [X12,X13] : (sP1(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X12))) ) | $spl152),
  inference(cnf_transformation,[],[f3939_D])).
fof(f3939_D,plain,(
  ( ! [X12,X13] : (sP1(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X12))) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f6846,plain,(
  ( ! [X12,X13] : (sP1(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f5546,f864])).
fof(f5546,plain,(
  ( ! [X48,X49] : (sP1(X48,X49,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) | ~p(s(bool,X48))) ) | $spl9),
  inference(resolution,[],[f1090,f643])).
fof(f3914,plain,(
  ( ! [X12,X13] : (~sP1('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | ~p(s(bool,X12))) ) | $spl150),
  inference(cnf_transformation,[],[f3914_D])).
fof(f3914_D,plain,(
  ( ! [X12,X13] : (~sP1('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | ~p(s(bool,X12))) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f6819,plain,(
  ( ! [X12,X13] : (~sP1('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | ~p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f5538,f864])).
fof(f5538,plain,(
  ( ! [X33,X32] : (~sP1('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)),X32,X33) | ~p(s(bool,X32))) ) | $spl9),
  inference(resolution,[],[f1090,f632])).
fof(f3891,plain,(
  ( ! [X12,X13] : (~sP1(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X12))) ) | $spl148),
  inference(cnf_transformation,[],[f3891_D])).
fof(f3891_D,plain,(
  ( ! [X12,X13] : (~sP1(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X12))) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f6798,plain,(
  ( ! [X12,X13] : (~sP1(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f5537,f864])).
fof(f5537,plain,(
  ( ! [X30,X31] : (~sP1(X30,X31,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) | p(s(bool,X30))) ) | $spl9),
  inference(resolution,[],[f1090,f631])).
fof(f631,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | p(s(bool,X0)) | ~sP1(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f484])).
fof(f6770,plain,(
  ( ! [X12,X13] : (sP6(X12,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f996,f864])).
fof(f6758,plain,(
  ( ! [X12,X13] : (sP6(X12,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),sK15)),X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f996,f864])).
fof(f3834,plain,(
  ( ! [X12,X13] : (~sP1('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | p(s(bool,X13))) ) | $spl142),
  inference(cnf_transformation,[],[f3834_D])).
fof(f3834_D,plain,(
  ( ! [X12,X13] : (~sP1('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | p(s(bool,X13))) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f6741,plain,(
  ( ! [X12,X13] : (~sP1('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | p(s(bool,X13)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f5536,f864])).
fof(f5536,plain,(
  ( ! [X28,X29] : (~sP1('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)),X29,X28) | p(s(bool,X28))) ) | $spl9),
  inference(resolution,[],[f1090,f631])).
fof(f6728,plain,(
  ( ! [X12,X13] : (sP11(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X12))) ) | $spl286),
  inference(cnf_transformation,[],[f6728_D])).
fof(f6728_D,plain,(
  ( ! [X12,X13] : (sP11(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X12))) ) <=> ~$spl286),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl286])])).
fof(f6720,plain,(
  ( ! [X12,X13] : (sP11(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl10),
  inference(superposition,[],[f2300,f864])).
fof(f2300,plain,(
  ( ! [X187,X186] : (sP11(X186,X187,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | ~p(s(bool,X186))) ) | $spl10),
  inference(resolution,[],[f1092,f726])).
fof(f6711,plain,(
  ( ! [X12,X13] : (sP11('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | ~p(s(bool,X13))) ) | $spl284),
  inference(cnf_transformation,[],[f6711_D])).
fof(f6711_D,plain,(
  ( ! [X12,X13] : (sP11('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | ~p(s(bool,X13))) ) <=> ~$spl284),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl284])])).
fof(f6703,plain,(
  ( ! [X12,X13] : (sP11('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | ~p(s(bool,X13)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl10),
  inference(superposition,[],[f2299,f864])).
fof(f2299,plain,(
  ( ! [X185,X184] : (sP11('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X184,X185) | ~p(s(bool,X185))) ) | $spl10),
  inference(resolution,[],[f1092,f726])).
fof(f6681,plain,(
  ( ! [X12,X13] : (sP6(X12,X13,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f995,f864])).
fof(f6669,plain,(
  ( ! [X12,X13] : (sP6(X12,X13,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),sK15))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f995,f864])).
fof(f6648,plain,(
  ( ! [X12,X13] : (sP7(X12,X13,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f994,f864])).
fof(f6636,plain,(
  ( ! [X12,X13] : (sP7(X12,X13,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),sK15))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f994,f864])).
fof(f6627,plain,(
  ( ! [X12,X13] : (sP10(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X13))) ) | $spl282),
  inference(cnf_transformation,[],[f6627_D])).
fof(f6627_D,plain,(
  ( ! [X12,X13] : (sP10(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X13))) ) <=> ~$spl282),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl282])])).
fof(f6619,plain,(
  ( ! [X12,X13] : (sP10(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X13)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl10),
  inference(superposition,[],[f2287,f864])).
fof(f2287,plain,(
  ( ! [X161,X160] : (sP10(X160,X161,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | ~p(s(bool,X161))) ) | $spl10),
  inference(resolution,[],[f1092,f713])).
fof(f6610,plain,(
  ( ! [X12,X13] : (sP10(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | ~p(s(bool,X13))) ) | $spl280),
  inference(cnf_transformation,[],[f6610_D])).
fof(f6610_D,plain,(
  ( ! [X12,X13] : (sP10(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | ~p(s(bool,X13))) ) <=> ~$spl280),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl280])])).
fof(f6602,plain,(
  ( ! [X12,X13] : (sP10(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | ~p(s(bool,X13)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl10),
  inference(superposition,[],[f2286,f864])).
fof(f2286,plain,(
  ( ! [X158,X159] : (sP10(X158,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X159) | ~p(s(bool,X159))) ) | $spl10),
  inference(resolution,[],[f1092,f713])).
fof(f6581,plain,(
  ( ! [X12,X13] : (sP8(X12,X13,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f991,f864])).
fof(f6569,plain,(
  ( ! [X12,X13] : (sP8(X12,X13,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),sK15))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f991,f864])).
fof(f6560,plain,(
  ( ! [X12,X13] : (~sP5('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | p(s(bool,X12))) ) | $spl278),
  inference(cnf_transformation,[],[f6560_D])).
fof(f6560_D,plain,(
  ( ! [X12,X13] : (~sP5('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | p(s(bool,X12))) ) <=> ~$spl278),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl278])])).
fof(f6549,plain,(
  ( ! [X12,X13] : (~sP5('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl10),
  inference(superposition,[],[f2257,f864])).
fof(f2257,plain,(
  ( ! [X101,X100] : (~sP5('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X100,X101) | p(s(bool,X100))) ) | $spl10),
  inference(resolution,[],[f1092,f667])).
fof(f6537,plain,(
  ( ! [X12,X13] : (~sP5('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | p(s(bool,X13))) ) | $spl276),
  inference(cnf_transformation,[],[f6537_D])).
fof(f6537_D,plain,(
  ( ! [X12,X13] : (~sP5('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | p(s(bool,X13))) ) <=> ~$spl276),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl276])])).
fof(f6526,plain,(
  ( ! [X12,X13] : (~sP5('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | p(s(bool,X13)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl10),
  inference(superposition,[],[f2256,f864])).
fof(f2256,plain,(
  ( ! [X99,X98] : (~sP5('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X99,X98) | p(s(bool,X98))) ) | $spl10),
  inference(resolution,[],[f1092,f666])).
fof(f6514,plain,(
  ( ! [X12,X13] : (sP3('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | ~p(s(bool,X12))) ) | $spl274),
  inference(cnf_transformation,[],[f6514_D])).
fof(f6514_D,plain,(
  ( ! [X12,X13] : (sP3('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | ~p(s(bool,X12))) ) <=> ~$spl274),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl274])])).
fof(f6504,plain,(
  ( ! [X12,X13] : (sP3('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | ~p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl10),
  inference(superposition,[],[f2252,f864])).
fof(f2252,plain,(
  ( ! [X90,X91] : (sP3('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X90,X91) | ~p(s(bool,X90))) ) | $spl10),
  inference(resolution,[],[f1092,f661])).
fof(f6479,plain,(
  ( ! [X12,X13] : (sP4(X12,X13,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f978,f864])).
fof(f6467,plain,(
  ( ! [X12,X13] : (sP4(X12,X13,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),sK15))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f978,f864])).
fof(f6448,plain,(
  ( ! [X12,X13] : (sP3(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | ~p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl10),
  inference(superposition,[],[f2251,f864])).
fof(f2251,plain,(
  ( ! [X88,X89] : (sP3(X88,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X89) | ~p(s(bool,X88))) ) | $spl10),
  inference(resolution,[],[f1092,f661])).
fof(f6425,plain,(
  ( ! [X12,X13] : (sP3('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | ~p(s(bool,X13)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl10),
  inference(superposition,[],[f2250,f864])).
fof(f2250,plain,(
  ( ! [X87,X86] : (sP3('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X86,X87) | ~p(s(bool,X87))) ) | $spl10),
  inference(resolution,[],[f1092,f660])).
fof(f6403,plain,(
  ( ! [X12,X13] : (sP3(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl10),
  inference(superposition,[],[f2249,f864])).
fof(f2249,plain,(
  ( ! [X85,X84] : (sP3(X84,X85,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | ~p(s(bool,X84))) ) | $spl10),
  inference(resolution,[],[f1092,f660])).
fof(f6380,plain,(
  ( ! [X12,X13] : (~sP3(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl10),
  inference(superposition,[],[f2243,f864])).
fof(f2243,plain,(
  ( ! [X72,X73] : (~sP3(X72,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X73) | p(s(bool,X72))) ) | $spl10),
  inference(resolution,[],[f1092,f649])).
fof(f6358,plain,(
  ( ! [X12,X13] : (~sP3(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl10),
  inference(superposition,[],[f2242,f864])).
fof(f2242,plain,(
  ( ! [X70,X71] : (~sP3(X70,X71,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | p(s(bool,X70))) ) | $spl10),
  inference(resolution,[],[f1092,f648])).
fof(f6335,plain,(
  ( ! [X12,X13] : (sP2(X12,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f962,f864])).
fof(f6323,plain,(
  ( ! [X12,X13] : (sP2(X12,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),sK15)),X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f962,f864])).
fof(f6303,plain,(
  ( ! [X12,X13] : (sP1('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | ~p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl10),
  inference(superposition,[],[f2238,f864])).
fof(f2238,plain,(
  ( ! [X62,X63] : (sP1('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X62,X63) | ~p(s(bool,X62))) ) | $spl10),
  inference(resolution,[],[f1092,f644])).
fof(f6282,plain,(
  ( ! [X12,X13] : (sP1(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | ~p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl10),
  inference(superposition,[],[f2237,f864])).
fof(f2237,plain,(
  ( ! [X61,X60] : (sP1(X60,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X61) | ~p(s(bool,X60))) ) | $spl10),
  inference(resolution,[],[f1092,f644])).
fof(f6260,plain,(
  ( ! [X12,X13] : (sP1('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | p(s(bool,X13)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl10),
  inference(superposition,[],[f2236,f864])).
fof(f2236,plain,(
  ( ! [X59,X58] : (sP1('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X58,X59) | p(s(bool,X59))) ) | $spl10),
  inference(resolution,[],[f1092,f643])).
fof(f6237,plain,(
  ( ! [X12,X13] : (sP2(X12,X13,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f961,f864])).
fof(f6225,plain,(
  ( ! [X12,X13] : (sP2(X12,X13,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),sK15))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f961,f864])).
fof(f3656,plain,(
  ( ! [X8] : (sP12('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X8) | ~p(s(bool,X8))) ) | $spl124),
  inference(cnf_transformation,[],[f3656_D])).
fof(f3656_D,plain,(
  ( ! [X8] : (sP12('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X8) | ~p(s(bool,X8))) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f6204,plain,(
  ( ! [X8] : (sP12('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X8) | ~p(s(bool,X8)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f5614,f864])).
fof(f5614,plain,(
  ( ! [X181] : (sP12('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)),X181) | ~p(s(bool,X181))) ) | $spl9),
  inference(resolution,[],[f1090,f756])).
fof(f3598,plain,(
  ( ! [X8] : (sP12(X8,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X8))) ) | $spl118),
  inference(cnf_transformation,[],[f3598_D])).
fof(f3598_D,plain,(
  ( ! [X8] : (sP12(X8,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X8))) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f6181,plain,(
  ( ! [X8] : (sP12(X8,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X8)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f5613,f864])).
fof(f5613,plain,(
  ( ! [X180] : (sP12(X180,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) | ~p(s(bool,X180))) ) | $spl9),
  inference(resolution,[],[f1090,f753])).
fof(f6158,plain,(
  ( ! [X12,X13] : (sP0(X12,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f948,f864])).
fof(f6146,plain,(
  ( ! [X12,X13] : (sP0(X12,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),sK15)),X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f948,f864])).
fof(f3575,plain,(
  ( ! [X8] : (~sP12('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X8) | p(s(bool,X8))) ) | $spl116),
  inference(cnf_transformation,[],[f3575_D])).
fof(f3575_D,plain,(
  ( ! [X8] : (~sP12('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X8) | p(s(bool,X8))) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f6125,plain,(
  ( ! [X8] : (~sP12('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X8) | p(s(bool,X8)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f5612,f864])).
fof(f5612,plain,(
  ( ! [X179] : (~sP12('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)),X179) | p(s(bool,X179))) ) | $spl9),
  inference(resolution,[],[f1090,f751])).
fof(f3554,plain,(
  ( ! [X8] : (~sP12(X8,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X8))) ) | $spl114),
  inference(cnf_transformation,[],[f3554_D])).
fof(f3554_D,plain,(
  ( ! [X8] : (~sP12(X8,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X8))) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f6104,plain,(
  ( ! [X8] : (~sP12(X8,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X8)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f5611,f864])).
fof(f5611,plain,(
  ( ! [X178] : (~sP12(X178,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) | p(s(bool,X178))) ) | $spl9),
  inference(resolution,[],[f1090,f751])).
fof(f3126,plain,(
  ( ! [X12,X13] : (sP7('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13)) ) | $spl74),
  inference(cnf_transformation,[],[f3126_D])).
fof(f3126_D,plain,(
  ( ! [X12,X13] : (sP7('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13)) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f6041,plain,(
  ( ! [X12,X13] : (sP7('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f1236,f864])).
fof(f1236,plain,(
  ( ! [X134,X135] : (sP7('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)),X134,X135)) ) | $spl9),
  inference(resolution,[],[f1090,f696])).
fof(f6024,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))))) | $spl1),
  inference(trivial_inequality_removal,[],[f6019])).
fof(f6019,plain,(
  s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),'const.pred_set.EMPTY_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))))) | $spl1),
  inference(superposition,[],[f914,f897])).
fof(f6032,plain,(
  ( ! [X5] : (s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),X5) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))))))) ) | $spl1),
  inference(subsumption_resolution,[],[f6018,f860])).
fof(f6018,plain,(
  ( ! [X5] : (s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),X5) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X5))))) ) | $spl1),
  inference(superposition,[],[f914,f866])).
fof(f2822,plain,(
  ( ! [X12,X13] : (sP7(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13)) ) | $spl72),
  inference(cnf_transformation,[],[f2822_D])).
fof(f2822_D,plain,(
  ( ! [X12,X13] : (sP7(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13)) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f5977,plain,(
  ( ! [X12,X13] : (sP7(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f1235,f864])).
fof(f1235,plain,(
  ( ! [X132,X133] : (sP7(X132,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)),X133)) ) | $spl9),
  inference(resolution,[],[f1090,f695])).
fof(f2802,plain,(
  ( ! [X12,X13] : (sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13)) ) | $spl70),
  inference(cnf_transformation,[],[f2802_D])).
fof(f2802_D,plain,(
  ( ! [X12,X13] : (sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13)) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f5958,plain,(
  ( ! [X12,X13] : (sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f1233,f864])).
fof(f1233,plain,(
  ( ! [X128,X129] : (sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)),X128,X129)) ) | $spl9),
  inference(resolution,[],[f1090,f692])).
fof(f2782,plain,(
  ( ! [X12,X13] : (sP8(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13)) ) | $spl68),
  inference(cnf_transformation,[],[f2782_D])).
fof(f2782_D,plain,(
  ( ! [X12,X13] : (sP8(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13)) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f5940,plain,(
  ( ! [X12,X13] : (sP8(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f1232,f864])).
fof(f1232,plain,(
  ( ! [X127,X126] : (sP8(X126,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)),X127)) ) | $spl9),
  inference(resolution,[],[f1090,f691])).
fof(f2763,plain,(
  ( ! [X12,X13] : (sP4('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13)) ) | $spl66),
  inference(cnf_transformation,[],[f2763_D])).
fof(f2763_D,plain,(
  ( ! [X12,X13] : (sP4('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13)) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f5923,plain,(
  ( ! [X12,X13] : (sP4('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f1220,f864])).
fof(f1220,plain,(
  ( ! [X103,X102] : (sP4('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)),X102,X103)) ) | $spl9),
  inference(resolution,[],[f1090,f675])).
fof(f2435,plain,(
  ( ! [X12,X13] : (sP4(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13)) ) | $spl64),
  inference(cnf_transformation,[],[f2435_D])).
fof(f2435_D,plain,(
  ( ! [X12,X13] : (sP4(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13)) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f5905,plain,(
  ( ! [X12,X13] : (sP4(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f1219,f864])).
fof(f1219,plain,(
  ( ! [X101,X100] : (sP4(X100,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)),X101)) ) | $spl9),
  inference(resolution,[],[f1090,f674])).
fof(f5867,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) | $spl2),
  inference(superposition,[],[f915,f864])).
fof(f5846,plain,(
  ( ! [X213,X214] : (p(s(bool,X213)) | p(s(bool,X214)) | s(bool,X214) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))) ) | $spl2),
  inference(resolution,[],[f915,f905])).
fof(f5748,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))) | s(bool,X25) != s(bool,X26)) ) | $spl2),
  inference(resolution,[],[f915,f630])).
fof(f5745,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X17) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))) ) | $spl2),
  inference(resolution,[],[f915,f629])).
fof(f5744,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))) ) | $spl2),
  inference(resolution,[],[f915,f628])).
fof(f5743,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))) | s(bool,X10) != s(bool,X11)) ) | $spl2),
  inference(resolution,[],[f915,f626])).
fof(f5741,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))) ) | $spl2),
  inference(resolution,[],[f915,f624])).
fof(f5735,plain,(
  s(fun(sK13,bool),sK14) = s(fun(sK13,bool),sK15) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)))) | $spl2),
  inference(resolution,[],[f915,f859])).
fof(f5734,plain,(
  s(fun(sK13,bool),sK14) = s(fun(sK13,bool),sK15) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)))) | $spl2),
  inference(resolution,[],[f915,f859])).
fof(f5733,plain,(
  s(fun(sK13,bool),sK14) = s(fun(sK13,bool),sK15) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)))) | $spl2),
  inference(resolution,[],[f915,f858])).
fof(f5732,plain,(
  s(fun(sK13,bool),sK14) = s(fun(sK13,bool),sK15) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)))) | $spl2),
  inference(resolution,[],[f915,f858])).
fof(f2416,plain,(
  ( ! [X12,X13] : (sP2('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13)) ) | $spl62),
  inference(cnf_transformation,[],[f2416_D])).
fof(f2416_D,plain,(
  ( ! [X12,X13] : (sP2('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13)) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f5715,plain,(
  ( ! [X12,X13] : (sP2('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f1206,f864])).
fof(f1206,plain,(
  ( ! [X74,X75] : (sP2('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)),X74,X75)) ) | $spl9),
  inference(resolution,[],[f1090,f658])).
fof(f2396,plain,(
  ( ! [X12,X13] : (sP0('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13)) ) | $spl60),
  inference(cnf_transformation,[],[f2396_D])).
fof(f2396_D,plain,(
  ( ! [X12,X13] : (sP0('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13)) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f5697,plain,(
  ( ! [X12,X13] : (sP0('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f1190,f864])).
fof(f1190,plain,(
  ( ! [X43,X42] : (sP0('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)),X42,X43)) ) | $spl9),
  inference(resolution,[],[f1090,f641])).
fof(f2377,plain,(
  ( ! [X12,X13] : (sP0(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) | $spl58),
  inference(cnf_transformation,[],[f2377_D])).
fof(f2377_D,plain,(
  ( ! [X12,X13] : (sP0(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f5679,plain,(
  ( ! [X12,X13] : (sP0(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f1189,f864])).
fof(f1189,plain,(
  ( ! [X41,X40] : (sP0(X40,X41,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | $spl9),
  inference(resolution,[],[f1090,f639])).
fof(f1315,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) | ~p(s(bool,X5))) ) | $spl48),
  inference(cnf_transformation,[],[f1315_D])).
fof(f1315_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) | ~p(s(bool,X5))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f5633,plain,(
  ( ! [X216,X217] : (p(s(bool,X216)) | ~p(s(bool,X217)) | s(bool,X217) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | $spl9),
  inference(resolution,[],[f1090,f905])).
fof(f5535,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X26) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) | s(bool,X25) != s(bool,X27)) ) | $spl9),
  inference(resolution,[],[f1090,f630])).
fof(f5533,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X20) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | $spl9),
  inference(resolution,[],[f1090,f629])).
fof(f5531,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | $spl9),
  inference(resolution,[],[f1090,f628])).
fof(f5530,plain,(
  ( ! [X12,X10,X11] : (~p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X11) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) | s(bool,X10) != s(bool,X12)) ) | $spl9),
  inference(resolution,[],[f1090,f626])).
fof(f5528,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | $spl9),
  inference(resolution,[],[f1090,f624])).
fof(f3160,plain,(
  s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),sK15) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) | $spl33),
  inference(trivial_inequality_removal,[],[f3157])).
fof(f3157,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.|48|_0') | s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),sK15) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) | $spl33),
  inference(superposition,[],[f1111,f898])).
fof(f5472,plain,(
  ( ! [X5] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK15)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))))) ) | $spl33),
  inference(subsumption_resolution,[],[f3151,f860])).
fof(f3151,plain,(
  ( ! [X5] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK15)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X5))))) ) | $spl33),
  inference(superposition,[],[f1111,f866])).
fof(f5522,plain,(
  ( ! [X5] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK16(X5,sK15,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X5),s(sK13,sK17(X5,sK15,sK13)))))) ) | $spl254),
  inference(cnf_transformation,[],[f5522_D])).
fof(f5522_D,plain,(
  ( ! [X5] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK16(X5,sK15,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X5),s(sK13,sK17(X5,sK15,sK13)))))) ) <=> ~$spl254),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl254])])).
fof(f5156,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X1)))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK16(X0,sK15,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X0),s(sK13,sK17(X0,sK15,sK13)))))) ) | $spl6),
  inference(forward_demodulation,[],[f4796,f1088])).
fof(f4796,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK16(X0,sK15,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X0),s(sK13,sK17(X0,sK15,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),sK15),s(sK13,X1))))) ) | $spl6),
  inference(superposition,[],[f783,f1088])).
fof(f5157,plain,(
  ( ! [X4,X5] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK16(X5,sK15,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X4)))) | p(s(bool,happ(s(fun(sK13,bool),X5),s(sK13,sK17(X5,sK15,sK13)))))) ) | $spl6),
  inference(forward_demodulation,[],[f4798,f1088])).
fof(f4798,plain,(
  ( ! [X4,X5] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X4)))) | p(s(bool,happ(s(fun(sK13,bool),X5),s(sK13,sK17(X5,sK15,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK15),s(sK13,sK16(X5,sK15,sK13)))))) ) | $spl6),
  inference(superposition,[],[f783,f1088])).
fof(f5521,plain,(
  ( ! [X11] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK17(sK15,X11,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X11),s(sK13,sK16(sK15,X11,sK13)))))) ) | $spl252),
  inference(cnf_transformation,[],[f5521_D])).
fof(f5521_D,plain,(
  ( ! [X11] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK17(sK15,X11,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X11),s(sK13,sK16(sK15,X11,sK13)))))) ) <=> ~$spl252),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl252])])).
fof(f5158,plain,(
  ( ! [X8,X9] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X9)))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK17(sK15,X8,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X8),s(sK13,sK16(sK15,X8,sK13)))))) ) | $spl6),
  inference(forward_demodulation,[],[f4800,f1088])).
fof(f4800,plain,(
  ( ! [X8,X9] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK17(sK15,X8,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X8),s(sK13,sK16(sK15,X8,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),sK15),s(sK13,X9))))) ) | $spl6),
  inference(superposition,[],[f784,f1088])).
fof(f5159,plain,(
  ( ! [X10,X11] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK17(sK15,X11,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X10)))) | p(s(bool,happ(s(fun(sK13,bool),X11),s(sK13,sK16(sK15,X11,sK13)))))) ) | $spl6),
  inference(forward_demodulation,[],[f4801,f1088])).
fof(f4801,plain,(
  ( ! [X10,X11] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X10)))) | p(s(bool,happ(s(fun(sK13,bool),sK15),s(sK13,sK17(sK15,X11,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X11),s(sK13,sK16(sK15,X11,sK13)))))) ) | $spl6),
  inference(superposition,[],[f784,f1088])).
fof(f5520,plain,(
  ( ! [X17] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK18(X17,sK15,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X17),s(sK13,sK18(X17,sK15,sK13)))))) ) | $spl250),
  inference(cnf_transformation,[],[f5520_D])).
fof(f5520_D,plain,(
  ( ! [X17] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK18(X17,sK15,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X17),s(sK13,sK18(X17,sK15,sK13)))))) ) <=> ~$spl250),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl250])])).
fof(f5160,plain,(
  ( ! [X12,X13] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X13)))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK18(X12,sK15,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X12),s(sK13,sK18(X12,sK15,sK13)))))) ) | $spl6),
  inference(forward_demodulation,[],[f4802,f1088])).
fof(f4802,plain,(
  ( ! [X12,X13] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK18(X12,sK15,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X12),s(sK13,sK18(X12,sK15,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),sK15),s(sK13,X13))))) ) | $spl6),
  inference(superposition,[],[f785,f1088])).
fof(f5161,plain,(
  ( ! [X17,X16] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK18(X17,sK15,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X16)))) | p(s(bool,happ(s(fun(sK13,bool),X17),s(sK13,sK18(X17,sK15,sK13)))))) ) | $spl6),
  inference(forward_demodulation,[],[f4804,f1088])).
fof(f4804,plain,(
  ( ! [X17,X16] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X16)))) | p(s(bool,happ(s(fun(sK13,bool),X17),s(sK13,sK18(X17,sK15,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK15),s(sK13,sK18(X17,sK15,sK13)))))) ) | $spl6),
  inference(superposition,[],[f785,f1088])).
fof(f5519,plain,(
  ( ! [X23] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK18(sK15,X23,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X23),s(sK13,sK18(sK15,X23,sK13)))))) ) | $spl248),
  inference(cnf_transformation,[],[f5519_D])).
fof(f5519_D,plain,(
  ( ! [X23] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK18(sK15,X23,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X23),s(sK13,sK18(sK15,X23,sK13)))))) ) <=> ~$spl248),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl248])])).
fof(f5162,plain,(
  ( ! [X21,X20] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X21)))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK18(sK15,X20,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X20),s(sK13,sK18(sK15,X20,sK13)))))) ) | $spl6),
  inference(forward_demodulation,[],[f4806,f1088])).
fof(f4806,plain,(
  ( ! [X21,X20] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK18(sK15,X20,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X20),s(sK13,sK18(sK15,X20,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),sK15),s(sK13,X21))))) ) | $spl6),
  inference(superposition,[],[f786,f1088])).
fof(f5163,plain,(
  ( ! [X23,X22] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK18(sK15,X23,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X22)))) | p(s(bool,happ(s(fun(sK13,bool),X23),s(sK13,sK18(sK15,X23,sK13)))))) ) | $spl6),
  inference(forward_demodulation,[],[f4807,f1088])).
fof(f4807,plain,(
  ( ! [X23,X22] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X22)))) | p(s(bool,happ(s(fun(sK13,bool),sK15),s(sK13,sK18(sK15,X23,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X23),s(sK13,sK18(sK15,X23,sK13)))))) ) | $spl6),
  inference(superposition,[],[f786,f1088])).
fof(f5518,plain,(
  ( ! [X27] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK19(X27,sK15,sK13))))) | ~p(s(bool,X27))) ) | $spl246),
  inference(cnf_transformation,[],[f5518_D])).
fof(f5518_D,plain,(
  ( ! [X27] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK19(X27,sK15,sK13))))) | ~p(s(bool,X27))) ) <=> ~$spl246),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl246])])).
fof(f5164,plain,(
  ( ! [X24,X25] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X25)))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK19(X24,sK15,sK13))))) | ~p(s(bool,X24))) ) | $spl6),
  inference(forward_demodulation,[],[f4808,f1088])).
fof(f4808,plain,(
  ( ! [X24,X25] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK19(X24,sK15,sK13))))) | ~p(s(bool,X24)) | ~p(s(bool,happ(s(fun(sK13,bool),sK15),s(sK13,X25))))) ) | $spl6),
  inference(superposition,[],[f787,f1088])).
fof(f5165,plain,(
  ( ! [X26,X27] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK19(X27,sK15,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X26)))) | ~p(s(bool,X27))) ) | $spl6),
  inference(forward_demodulation,[],[f4809,f1088])).
fof(f4809,plain,(
  ( ! [X26,X27] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X26)))) | p(s(bool,happ(s(fun(sK13,bool),sK15),s(sK13,sK19(X27,sK15,sK13))))) | ~p(s(bool,X27))) ) | $spl6),
  inference(superposition,[],[f787,f1088])).
fof(f5517,plain,(
  ( ! [X31] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK20(X31,sK15,sK13))))) | ~p(s(bool,X31))) ) | $spl244),
  inference(cnf_transformation,[],[f5517_D])).
fof(f5517_D,plain,(
  ( ! [X31] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK20(X31,sK15,sK13))))) | ~p(s(bool,X31))) ) <=> ~$spl244),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl244])])).
fof(f5166,plain,(
  ( ! [X28,X29] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X29)))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK20(X28,sK15,sK13))))) | ~p(s(bool,X28))) ) | $spl6),
  inference(forward_demodulation,[],[f4810,f1088])).
fof(f4810,plain,(
  ( ! [X28,X29] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK20(X28,sK15,sK13))))) | ~p(s(bool,X28)) | ~p(s(bool,happ(s(fun(sK13,bool),sK15),s(sK13,X29))))) ) | $spl6),
  inference(superposition,[],[f789,f1088])).
fof(f5167,plain,(
  ( ! [X30,X31] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK20(X31,sK15,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X30)))) | ~p(s(bool,X31))) ) | $spl6),
  inference(forward_demodulation,[],[f4811,f1088])).
fof(f4811,plain,(
  ( ! [X30,X31] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X30)))) | p(s(bool,happ(s(fun(sK13,bool),sK15),s(sK13,sK20(X31,sK15,sK13))))) | ~p(s(bool,X31))) ) | $spl6),
  inference(superposition,[],[f789,f1088])).
fof(f5516,plain,(
  ( ! [X42] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK22(sK15,X42,sK13))))) | ~p(s(bool,X42))) ) | $spl242),
  inference(cnf_transformation,[],[f5516_D])).
fof(f5516_D,plain,(
  ( ! [X42] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK22(sK15,X42,sK13))))) | ~p(s(bool,X42))) ) <=> ~$spl242),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl242])])).
fof(f5172,plain,(
  ( ! [X39,X40] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X40)))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK22(sK15,X39,sK13))))) | ~p(s(bool,X39))) ) | $spl6),
  inference(forward_demodulation,[],[f4817,f1088])).
fof(f4817,plain,(
  ( ! [X39,X40] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK22(sK15,X39,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),sK15),s(sK13,X40)))) | ~p(s(bool,X39))) ) | $spl6),
  inference(superposition,[],[f795,f1088])).
fof(f5173,plain,(
  ( ! [X41,X42] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK22(sK15,X42,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X41)))) | ~p(s(bool,X42))) ) | $spl6),
  inference(forward_demodulation,[],[f4818,f1088])).
fof(f4818,plain,(
  ( ! [X41,X42] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X41)))) | p(s(bool,happ(s(fun(sK13,bool),sK15),s(sK13,sK22(sK15,X42,sK13))))) | ~p(s(bool,X42))) ) | $spl6),
  inference(superposition,[],[f795,f1088])).
fof(f5515,plain,(
  ( ! [X46] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(sK15,X46,sK13))))) | ~p(s(bool,X46))) ) | $spl240),
  inference(cnf_transformation,[],[f5515_D])).
fof(f5515_D,plain,(
  ( ! [X46] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(sK15,X46,sK13))))) | ~p(s(bool,X46))) ) <=> ~$spl240),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl240])])).
fof(f5174,plain,(
  ( ! [X43,X44] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X44)))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(sK15,X43,sK13))))) | ~p(s(bool,X43))) ) | $spl6),
  inference(forward_demodulation,[],[f4819,f1088])).
fof(f4819,plain,(
  ( ! [X43,X44] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(sK15,X43,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),sK15),s(sK13,X44)))) | ~p(s(bool,X43))) ) | $spl6),
  inference(superposition,[],[f797,f1088])).
fof(f5175,plain,(
  ( ! [X45,X46] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(sK15,X46,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X45)))) | ~p(s(bool,X46))) ) | $spl6),
  inference(forward_demodulation,[],[f4820,f1088])).
fof(f4820,plain,(
  ( ! [X45,X46] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X45)))) | p(s(bool,happ(s(fun(sK13,bool),sK15),s(sK13,sK23(sK15,X46,sK13))))) | ~p(s(bool,X46))) ) | $spl6),
  inference(superposition,[],[f797,f1088])).
fof(f5514,plain,(
  ( ! [X50] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK24(sK15,X50,sK13))))) | p(s(bool,X50))) ) | $spl238),
  inference(cnf_transformation,[],[f5514_D])).
fof(f5514_D,plain,(
  ( ! [X50] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK24(sK15,X50,sK13))))) | p(s(bool,X50))) ) <=> ~$spl238),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl238])])).
fof(f5176,plain,(
  ( ! [X47,X48] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X48)))) | ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK24(sK15,X47,sK13))))) | p(s(bool,X47))) ) | $spl6),
  inference(forward_demodulation,[],[f4821,f1088])).
fof(f4821,plain,(
  ( ! [X47,X48] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK24(sK15,X47,sK13))))) | p(s(bool,X47)) | p(s(bool,happ(s(fun(sK13,bool),sK15),s(sK13,X48))))) ) | $spl6),
  inference(superposition,[],[f798,f1088])).
fof(f5177,plain,(
  ( ! [X50,X49] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK24(sK15,X50,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X49)))) | p(s(bool,X50))) ) | $spl6),
  inference(forward_demodulation,[],[f4822,f1088])).
fof(f4822,plain,(
  ( ! [X50,X49] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X49)))) | p(s(bool,X50)) | ~p(s(bool,happ(s(fun(sK13,bool),sK15),s(sK13,sK24(sK15,X50,sK13)))))) ) | $spl6),
  inference(superposition,[],[f798,f1088])).
fof(f5513,plain,(
  ( ! [X54] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK25(sK15,X54,sK13))))) | p(s(bool,X54))) ) | $spl236),
  inference(cnf_transformation,[],[f5513_D])).
fof(f5513_D,plain,(
  ( ! [X54] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK25(sK15,X54,sK13))))) | p(s(bool,X54))) ) <=> ~$spl236),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl236])])).
fof(f5178,plain,(
  ( ! [X52,X51] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X52)))) | ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK25(sK15,X51,sK13))))) | p(s(bool,X51))) ) | $spl6),
  inference(forward_demodulation,[],[f4823,f1088])).
fof(f4823,plain,(
  ( ! [X52,X51] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK25(sK15,X51,sK13))))) | p(s(bool,X51)) | p(s(bool,happ(s(fun(sK13,bool),sK15),s(sK13,X52))))) ) | $spl6),
  inference(superposition,[],[f800,f1088])).
fof(f5179,plain,(
  ( ! [X54,X53] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK25(sK15,X54,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X53)))) | p(s(bool,X54))) ) | $spl6),
  inference(forward_demodulation,[],[f4824,f1088])).
fof(f4824,plain,(
  ( ! [X54,X53] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X53)))) | p(s(bool,X54)) | ~p(s(bool,happ(s(fun(sK13,bool),sK15),s(sK13,sK25(sK15,X54,sK13)))))) ) | $spl6),
  inference(superposition,[],[f800,f1088])).
fof(f5512,plain,(
  ( ! [X58] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK26(sK15,X58,sK13))))) | p(s(bool,X58))) ) | $spl234),
  inference(cnf_transformation,[],[f5512_D])).
fof(f5512_D,plain,(
  ( ! [X58] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK26(sK15,X58,sK13))))) | p(s(bool,X58))) ) <=> ~$spl234),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl234])])).
fof(f5180,plain,(
  ( ! [X56,X55] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X56)))) | ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK26(sK15,X55,sK13))))) | p(s(bool,X55))) ) | $spl6),
  inference(forward_demodulation,[],[f4825,f1088])).
fof(f4825,plain,(
  ( ! [X56,X55] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK26(sK15,X55,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK15),s(sK13,X56)))) | p(s(bool,X55))) ) | $spl6),
  inference(superposition,[],[f803,f1088])).
fof(f5181,plain,(
  ( ! [X57,X58] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK26(sK15,X58,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X57)))) | p(s(bool,X58))) ) | $spl6),
  inference(forward_demodulation,[],[f4826,f1088])).
fof(f4826,plain,(
  ( ! [X57,X58] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X57)))) | p(s(bool,X58)) | ~p(s(bool,happ(s(fun(sK13,bool),sK15),s(sK13,sK26(sK15,X58,sK13)))))) ) | $spl6),
  inference(superposition,[],[f803,f1088])).
fof(f5511,plain,(
  ( ! [X62] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK27(sK15,X62,sK13))))) | p(s(bool,X62))) ) | $spl232),
  inference(cnf_transformation,[],[f5511_D])).
fof(f5511_D,plain,(
  ( ! [X62] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK27(sK15,X62,sK13))))) | p(s(bool,X62))) ) <=> ~$spl232),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl232])])).
fof(f5182,plain,(
  ( ! [X59,X60] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X60)))) | ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK27(sK15,X59,sK13))))) | p(s(bool,X59))) ) | $spl6),
  inference(forward_demodulation,[],[f4827,f1088])).
fof(f4827,plain,(
  ( ! [X59,X60] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK27(sK15,X59,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK15),s(sK13,X60)))) | p(s(bool,X59))) ) | $spl6),
  inference(superposition,[],[f805,f1088])).
fof(f5183,plain,(
  ( ! [X61,X62] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK27(sK15,X62,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X61)))) | p(s(bool,X62))) ) | $spl6),
  inference(forward_demodulation,[],[f4828,f1088])).
fof(f4828,plain,(
  ( ! [X61,X62] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X61)))) | p(s(bool,X62)) | ~p(s(bool,happ(s(fun(sK13,bool),sK15),s(sK13,sK27(sK15,X62,sK13)))))) ) | $spl6),
  inference(superposition,[],[f805,f1088])).
fof(f5510,plain,(
  ( ! [X68] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK28(X68,sK15,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X68),s(sK13,sK28(X68,sK15,sK13)))))) ) | $spl230),
  inference(cnf_transformation,[],[f5510_D])).
fof(f5510_D,plain,(
  ( ! [X68] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK28(X68,sK15,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X68),s(sK13,sK28(X68,sK15,sK13)))))) ) <=> ~$spl230),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl230])])).
fof(f5184,plain,(
  ( ! [X64,X63] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X64)))) | ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK28(X63,sK15,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X63),s(sK13,sK28(X63,sK15,sK13)))))) ) | $spl6),
  inference(forward_demodulation,[],[f4829,f1088])).
fof(f4829,plain,(
  ( ! [X64,X63] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK28(X63,sK15,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK15),s(sK13,X64)))) | ~p(s(bool,happ(s(fun(sK13,bool),X63),s(sK13,sK28(X63,sK15,sK13)))))) ) | $spl6),
  inference(superposition,[],[f806,f1088])).
fof(f5185,plain,(
  ( ! [X68,X67] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK28(X68,sK15,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X67)))) | ~p(s(bool,happ(s(fun(sK13,bool),X68),s(sK13,sK28(X68,sK15,sK13)))))) ) | $spl6),
  inference(forward_demodulation,[],[f4831,f1088])).
fof(f4831,plain,(
  ( ! [X68,X67] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X67)))) | ~p(s(bool,happ(s(fun(sK13,bool),X68),s(sK13,sK28(X68,sK15,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),sK15),s(sK13,sK28(X68,sK15,sK13)))))) ) | $spl6),
  inference(superposition,[],[f806,f1088])).
fof(f5509,plain,(
  ( ! [X74] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK28(sK15,X74,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X74),s(sK13,sK28(sK15,X74,sK13)))))) ) | $spl228),
  inference(cnf_transformation,[],[f5509_D])).
fof(f5509_D,plain,(
  ( ! [X74] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK28(sK15,X74,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X74),s(sK13,sK28(sK15,X74,sK13)))))) ) <=> ~$spl228),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl228])])).
fof(f5186,plain,(
  ( ! [X72,X71] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X72)))) | ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK28(sK15,X71,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X71),s(sK13,sK28(sK15,X71,sK13)))))) ) | $spl6),
  inference(forward_demodulation,[],[f4833,f1088])).
fof(f4833,plain,(
  ( ! [X72,X71] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK28(sK15,X71,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK15),s(sK13,X72)))) | ~p(s(bool,happ(s(fun(sK13,bool),X71),s(sK13,sK28(sK15,X71,sK13)))))) ) | $spl6),
  inference(superposition,[],[f807,f1088])).
fof(f5187,plain,(
  ( ! [X74,X73] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK28(sK15,X74,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X73)))) | ~p(s(bool,happ(s(fun(sK13,bool),X74),s(sK13,sK28(sK15,X74,sK13)))))) ) | $spl6),
  inference(forward_demodulation,[],[f4834,f1088])).
fof(f4834,plain,(
  ( ! [X74,X73] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X73)))) | ~p(s(bool,happ(s(fun(sK13,bool),sK15),s(sK13,sK28(sK15,X74,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X74),s(sK13,sK28(sK15,X74,sK13)))))) ) | $spl6),
  inference(superposition,[],[f807,f1088])).
fof(f5508,plain,(
  ( ! [X80] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK29(X80,sK15,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X80),s(sK13,sK30(X80,sK15,sK13)))))) ) | $spl226),
  inference(cnf_transformation,[],[f5508_D])).
fof(f5508_D,plain,(
  ( ! [X80] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK29(X80,sK15,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X80),s(sK13,sK30(X80,sK15,sK13)))))) ) <=> ~$spl226),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl226])])).
fof(f5188,plain,(
  ( ! [X76,X75] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X76)))) | ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK29(X75,sK15,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X75),s(sK13,sK30(X75,sK15,sK13)))))) ) | $spl6),
  inference(forward_demodulation,[],[f4835,f1088])).
fof(f4835,plain,(
  ( ! [X76,X75] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK29(X75,sK15,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK15),s(sK13,X76)))) | ~p(s(bool,happ(s(fun(sK13,bool),X75),s(sK13,sK30(X75,sK15,sK13)))))) ) | $spl6),
  inference(superposition,[],[f808,f1088])).
fof(f5189,plain,(
  ( ! [X80,X79] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK29(X80,sK15,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X79)))) | ~p(s(bool,happ(s(fun(sK13,bool),X80),s(sK13,sK30(X80,sK15,sK13)))))) ) | $spl6),
  inference(forward_demodulation,[],[f4837,f1088])).
fof(f4837,plain,(
  ( ! [X80,X79] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X79)))) | ~p(s(bool,happ(s(fun(sK13,bool),X80),s(sK13,sK30(X80,sK15,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),sK15),s(sK13,sK29(X80,sK15,sK13)))))) ) | $spl6),
  inference(superposition,[],[f808,f1088])).
fof(f5507,plain,(
  ( ! [X86] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK30(sK15,X86,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X86),s(sK13,sK29(sK15,X86,sK13)))))) ) | $spl224),
  inference(cnf_transformation,[],[f5507_D])).
fof(f5507_D,plain,(
  ( ! [X86] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK30(sK15,X86,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X86),s(sK13,sK29(sK15,X86,sK13)))))) ) <=> ~$spl224),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl224])])).
fof(f5190,plain,(
  ( ! [X83,X84] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X84)))) | ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK30(sK15,X83,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X83),s(sK13,sK29(sK15,X83,sK13)))))) ) | $spl6),
  inference(forward_demodulation,[],[f4839,f1088])).
fof(f4839,plain,(
  ( ! [X83,X84] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK30(sK15,X83,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK15),s(sK13,X84)))) | ~p(s(bool,happ(s(fun(sK13,bool),X83),s(sK13,sK29(sK15,X83,sK13)))))) ) | $spl6),
  inference(superposition,[],[f809,f1088])).
fof(f5191,plain,(
  ( ! [X85,X86] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK30(sK15,X86,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X85)))) | ~p(s(bool,happ(s(fun(sK13,bool),X86),s(sK13,sK29(sK15,X86,sK13)))))) ) | $spl6),
  inference(forward_demodulation,[],[f4840,f1088])).
fof(f4840,plain,(
  ( ! [X85,X86] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X85)))) | ~p(s(bool,happ(s(fun(sK13,bool),sK15),s(sK13,sK30(sK15,X86,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X86),s(sK13,sK29(sK15,X86,sK13)))))) ) | $spl6),
  inference(superposition,[],[f809,f1088])).
fof(f5506,plain,(
  ( ! [X90] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK31(X90,sK15,sK13))))) | p(s(bool,X90))) ) | $spl222),
  inference(cnf_transformation,[],[f5506_D])).
fof(f5506_D,plain,(
  ( ! [X90] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK31(X90,sK15,sK13))))) | p(s(bool,X90))) ) <=> ~$spl222),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl222])])).
fof(f5192,plain,(
  ( ! [X88,X87] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X88)))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK31(X87,sK15,sK13))))) | p(s(bool,X87))) ) | $spl6),
  inference(forward_demodulation,[],[f4841,f1088])).
fof(f4841,plain,(
  ( ! [X88,X87] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK31(X87,sK15,sK13))))) | p(s(bool,X87)) | ~p(s(bool,happ(s(fun(sK13,bool),sK15),s(sK13,X88))))) ) | $spl6),
  inference(superposition,[],[f810,f1088])).
fof(f5193,plain,(
  ( ! [X90,X89] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK31(X90,sK15,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X89)))) | p(s(bool,X90))) ) | $spl6),
  inference(forward_demodulation,[],[f4842,f1088])).
fof(f4842,plain,(
  ( ! [X90,X89] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X89)))) | p(s(bool,X90)) | p(s(bool,happ(s(fun(sK13,bool),sK15),s(sK13,sK31(X90,sK15,sK13)))))) ) | $spl6),
  inference(superposition,[],[f810,f1088])).
fof(f5505,plain,(
  ( ! [X94] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK32(X94,sK15,sK13))))) | p(s(bool,X94))) ) | $spl220),
  inference(cnf_transformation,[],[f5505_D])).
fof(f5505_D,plain,(
  ( ! [X94] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK32(X94,sK15,sK13))))) | p(s(bool,X94))) ) <=> ~$spl220),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl220])])).
fof(f5194,plain,(
  ( ! [X92,X91] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X92)))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK32(X91,sK15,sK13))))) | p(s(bool,X91))) ) | $spl6),
  inference(forward_demodulation,[],[f4843,f1088])).
fof(f4843,plain,(
  ( ! [X92,X91] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK32(X91,sK15,sK13))))) | p(s(bool,X91)) | ~p(s(bool,happ(s(fun(sK13,bool),sK15),s(sK13,X92))))) ) | $spl6),
  inference(superposition,[],[f812,f1088])).
fof(f5195,plain,(
  ( ! [X94,X93] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK32(X94,sK15,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X93)))) | p(s(bool,X94))) ) | $spl6),
  inference(forward_demodulation,[],[f4844,f1088])).
fof(f4844,plain,(
  ( ! [X94,X93] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X93)))) | p(s(bool,X94)) | p(s(bool,happ(s(fun(sK13,bool),sK15),s(sK13,sK32(X94,sK15,sK13)))))) ) | $spl6),
  inference(superposition,[],[f812,f1088])).
fof(f5504,plain,(
  ( ! [X98] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK33(sK15,X98,sK13))))) | p(s(bool,X98))) ) | $spl218),
  inference(cnf_transformation,[],[f5504_D])).
fof(f5504_D,plain,(
  ( ! [X98] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK33(sK15,X98,sK13))))) | p(s(bool,X98))) ) <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f5196,plain,(
  ( ! [X95,X96] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X96)))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK33(sK15,X95,sK13))))) | p(s(bool,X95))) ) | $spl6),
  inference(forward_demodulation,[],[f4845,f1088])).
fof(f4845,plain,(
  ( ! [X95,X96] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK33(sK15,X95,sK13))))) | p(s(bool,X95)) | ~p(s(bool,happ(s(fun(sK13,bool),sK15),s(sK13,X96))))) ) | $spl6),
  inference(superposition,[],[f815,f1088])).
fof(f5197,plain,(
  ( ! [X97,X98] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK33(sK15,X98,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X97)))) | p(s(bool,X98))) ) | $spl6),
  inference(forward_demodulation,[],[f4846,f1088])).
fof(f4846,plain,(
  ( ! [X97,X98] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X97)))) | p(s(bool,happ(s(fun(sK13,bool),sK15),s(sK13,sK33(sK15,X98,sK13))))) | p(s(bool,X98))) ) | $spl6),
  inference(superposition,[],[f815,f1088])).
fof(f5503,plain,(
  ( ! [X102] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK34(sK15,X102,sK13))))) | p(s(bool,X102))) ) | $spl216),
  inference(cnf_transformation,[],[f5503_D])).
fof(f5503_D,plain,(
  ( ! [X102] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK34(sK15,X102,sK13))))) | p(s(bool,X102))) ) <=> ~$spl216),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl216])])).
fof(f5198,plain,(
  ( ! [X99,X100] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X100)))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK34(sK15,X99,sK13))))) | p(s(bool,X99))) ) | $spl6),
  inference(forward_demodulation,[],[f4847,f1088])).
fof(f4847,plain,(
  ( ! [X99,X100] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK34(sK15,X99,sK13))))) | p(s(bool,X99)) | ~p(s(bool,happ(s(fun(sK13,bool),sK15),s(sK13,X100))))) ) | $spl6),
  inference(superposition,[],[f817,f1088])).
fof(f5199,plain,(
  ( ! [X101,X102] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK34(sK15,X102,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X101)))) | p(s(bool,X102))) ) | $spl6),
  inference(forward_demodulation,[],[f4848,f1088])).
fof(f4848,plain,(
  ( ! [X101,X102] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X101)))) | p(s(bool,happ(s(fun(sK13,bool),sK15),s(sK13,sK34(sK15,X102,sK13))))) | p(s(bool,X102))) ) | $spl6),
  inference(superposition,[],[f817,f1088])).
fof(f5501,plain,(
  ( ! [X106] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK35(sK15,X106,sK13))))) | ~p(s(bool,X106))) ) | $spl212),
  inference(cnf_transformation,[],[f5501_D])).
fof(f5501_D,plain,(
  ( ! [X106] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK35(sK15,X106,sK13))))) | ~p(s(bool,X106))) ) <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f5200,plain,(
  ( ! [X103,X104] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X104)))) | ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK35(sK15,X103,sK13))))) | ~p(s(bool,X103))) ) | $spl6),
  inference(forward_demodulation,[],[f4849,f1088])).
fof(f4849,plain,(
  ( ! [X103,X104] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK35(sK15,X103,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK15),s(sK13,X104)))) | ~p(s(bool,X103))) ) | $spl6),
  inference(superposition,[],[f819,f1088])).
fof(f5201,plain,(
  ( ! [X105,X106] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK35(sK15,X106,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X105)))) | ~p(s(bool,X106))) ) | $spl6),
  inference(forward_demodulation,[],[f4850,f1088])).
fof(f4850,plain,(
  ( ! [X105,X106] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X105)))) | ~p(s(bool,happ(s(fun(sK13,bool),sK15),s(sK13,sK35(sK15,X106,sK13))))) | ~p(s(bool,X106))) ) | $spl6),
  inference(superposition,[],[f819,f1088])).
fof(f5500,plain,(
  ( ! [X110] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK36(sK15,X110,sK13))))) | ~p(s(bool,X110))) ) | $spl210),
  inference(cnf_transformation,[],[f5500_D])).
fof(f5500_D,plain,(
  ( ! [X110] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK36(sK15,X110,sK13))))) | ~p(s(bool,X110))) ) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f5202,plain,(
  ( ! [X107,X108] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X108)))) | ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK36(sK15,X107,sK13))))) | ~p(s(bool,X107))) ) | $spl6),
  inference(forward_demodulation,[],[f4851,f1088])).
fof(f4851,plain,(
  ( ! [X107,X108] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK36(sK15,X107,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK15),s(sK13,X108)))) | ~p(s(bool,X107))) ) | $spl6),
  inference(superposition,[],[f821,f1088])).
fof(f5203,plain,(
  ( ! [X109,X110] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK36(sK15,X110,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X109)))) | ~p(s(bool,X110))) ) | $spl6),
  inference(forward_demodulation,[],[f4852,f1088])).
fof(f4852,plain,(
  ( ! [X109,X110] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X109)))) | ~p(s(bool,happ(s(fun(sK13,bool),sK15),s(sK13,sK36(sK15,X110,sK13))))) | ~p(s(bool,X110))) ) | $spl6),
  inference(superposition,[],[f821,f1088])).
fof(f5499,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK37(sK15,sK13))))) | $spl209),
  inference(cnf_transformation,[],[f5499_D])).
fof(f5499_D,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK37(sK15,sK13))))) <=> ~$spl209),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl209])])).
fof(f5204,plain,(
  ( ! [X111] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X111)))) | ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK37(sK15,sK13)))))) ) | $spl6),
  inference(forward_demodulation,[],[f4853,f1088])).
fof(f4853,plain,(
  ( ! [X111] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK37(sK15,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK15),s(sK13,X111))))) ) | $spl6),
  inference(superposition,[],[f824,f1088])).
fof(f5205,plain,(
  ( ! [X112] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK37(sK15,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X112))))) ) | $spl6),
  inference(forward_demodulation,[],[f4854,f1088])).
fof(f4854,plain,(
  ( ! [X112] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X112)))) | ~p(s(bool,happ(s(fun(sK13,bool),sK15),s(sK13,sK37(sK15,sK13)))))) ) | $spl6),
  inference(superposition,[],[f824,f1088])).
fof(f5497,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK38(sK15,sK13))))) | $spl207),
  inference(cnf_transformation,[],[f5497_D])).
fof(f5497_D,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK38(sK15,sK13))))) <=> ~$spl207),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl207])])).
fof(f5206,plain,(
  ( ! [X113] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X113)))) | ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK38(sK15,sK13)))))) ) | $spl6),
  inference(forward_demodulation,[],[f4855,f1088])).
fof(f4855,plain,(
  ( ! [X113] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK38(sK15,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK15),s(sK13,X113))))) ) | $spl6),
  inference(superposition,[],[f825,f1088])).
fof(f5207,plain,(
  ( ! [X114] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK38(sK15,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X114))))) ) | $spl6),
  inference(forward_demodulation,[],[f4856,f1088])).
fof(f4856,plain,(
  ( ! [X114] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X114)))) | ~p(s(bool,happ(s(fun(sK13,bool),sK15),s(sK13,sK38(sK15,sK13)))))) ) | $spl6),
  inference(superposition,[],[f825,f1088])).
fof(f5123,plain,(
  $false | ($spl6 | $spl9 | $spl10)),
  inference(subsumption_resolution,[],[f4749,f1092])).
fof(f4749,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ($spl6 | $spl9)),
  inference(backward_demodulation,[],[f1088,f1090])).
fof(f4717,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) | $spl3),
  inference(superposition,[],[f916,f864])).
fof(f4698,plain,(
  ( ! [X215,X214] : (p(s(bool,X214)) | ~p(s(bool,X215)) | s(bool,X215) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))) ) | $spl3),
  inference(resolution,[],[f916,f905])).
fof(f4600,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))) | s(bool,X23) != s(bool,X25)) ) | $spl3),
  inference(resolution,[],[f916,f630])).
fof(f4598,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))) ) | $spl3),
  inference(resolution,[],[f916,f629])).
fof(f4596,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))) ) | $spl3),
  inference(resolution,[],[f916,f628])).
fof(f4595,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X9) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))) | s(bool,X8) != s(bool,X10)) ) | $spl3),
  inference(resolution,[],[f916,f626])).
fof(f4593,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))) ) | $spl3),
  inference(resolution,[],[f916,f624])).
fof(f4572,plain,(
  ( ! [X12,X13] : (~sP1(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl10),
  inference(superposition,[],[f2227,f864])).
fof(f2227,plain,(
  ( ! [X41,X40] : (~sP1(X40,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X41) | p(s(bool,X40))) ) | $spl10),
  inference(resolution,[],[f1092,f632])).
fof(f4528,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) | $spl5),
  inference(superposition,[],[f1086,f864])).
fof(f4509,plain,(
  ( ! [X215,X214] : (p(s(bool,X214)) | ~p(s(bool,X215)) | s(bool,X215) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)))) ) | $spl5),
  inference(resolution,[],[f1086,f905])).
fof(f4411,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))) | s(bool,X23) != s(bool,X25)) ) | $spl5),
  inference(resolution,[],[f1086,f630])).
fof(f4409,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)))) ) | $spl5),
  inference(resolution,[],[f1086,f629])).
fof(f4407,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)))) ) | $spl5),
  inference(resolution,[],[f1086,f628])).
fof(f4406,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X9) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14))) | s(bool,X8) != s(bool,X10)) ) | $spl5),
  inference(resolution,[],[f1086,f626])).
fof(f4404,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)))) ) | $spl5),
  inference(resolution,[],[f1086,f624])).
fof(f4393,plain,(
  ( ! [X12,X13] : (~sP11('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | p(s(bool,X12))) ) | $spl196),
  inference(cnf_transformation,[],[f4393_D])).
fof(f4393_D,plain,(
  ( ! [X12,X13] : (~sP11('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | p(s(bool,X12))) ) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f4385,plain,(
  ( ! [X12,X13] : (~sP11('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f1831,f864])).
fof(f1831,plain,(
  ( ! [X169,X168] : (~sP11('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)),X168,X169) | p(s(bool,X168))) ) | $spl9),
  inference(resolution,[],[f1090,f722])).
fof(f722,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | ~sP11(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f526])).
fof(f4370,plain,(
  ( ! [X12,X13] : (~sP11(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | p(s(bool,X12))) ) | $spl194),
  inference(cnf_transformation,[],[f4370_D])).
fof(f4370_D,plain,(
  ( ! [X12,X13] : (~sP11(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | p(s(bool,X12))) ) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f4362,plain,(
  ( ! [X12,X13] : (~sP11(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f1830,f864])).
fof(f1830,plain,(
  ( ! [X167,X166] : (~sP11(X166,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)),X167) | p(s(bool,X166))) ) | $spl9),
  inference(resolution,[],[f1090,f722])).
fof(f4351,plain,(
  ( ! [X12,X13] : (~sP11(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X13))) ) | $spl192),
  inference(cnf_transformation,[],[f4351_D])).
fof(f4351_D,plain,(
  ( ! [X12,X13] : (~sP11(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X13))) ) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f4343,plain,(
  ( ! [X12,X13] : (~sP11(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X13)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f1829,f864])).
fof(f1829,plain,(
  ( ! [X165,X164] : (~sP11(X165,X164,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) | p(s(bool,X164))) ) | $spl9),
  inference(resolution,[],[f1090,f721])).
fof(f721,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | p(s(bool,X1)) | ~sP11(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f526])).
fof(f4328,plain,(
  ( ! [X12,X13] : (~sP11(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | p(s(bool,X13))) ) | $spl190),
  inference(cnf_transformation,[],[f4328_D])).
fof(f4328_D,plain,(
  ( ! [X12,X13] : (~sP11(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | p(s(bool,X13))) ) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f4320,plain,(
  ( ! [X12,X13] : (~sP11(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | p(s(bool,X13)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f1828,f864])).
fof(f1828,plain,(
  ( ! [X163,X162] : (~sP11(X163,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)),X162) | p(s(bool,X162))) ) | $spl9),
  inference(resolution,[],[f1090,f721])).
fof(f4299,plain,(
  ( ! [X12,X13] : (~sP10('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f1823,f864])).
fof(f1823,plain,(
  ( ! [X152,X153] : (~sP10('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)),X152,X153) | p(s(bool,X152))) ) | $spl9),
  inference(resolution,[],[f1090,f712])).
fof(f4280,plain,(
  ( ! [X12,X13] : (~sP10(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f1822,f864])).
fof(f1822,plain,(
  ( ! [X151,X150] : (~sP10(X150,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)),X151) | p(s(bool,X150))) ) | $spl9),
  inference(resolution,[],[f1090,f712])).
fof(f4259,plain,(
  ( ! [X12,X13] : (~sP10(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f1821,f864])).
fof(f1821,plain,(
  ( ! [X149,X148] : (~sP10(X148,X149,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) | p(s(bool,X148))) ) | $spl9),
  inference(resolution,[],[f1090,f711])).
fof(f4236,plain,(
  ( ! [X12,X13] : (~sP10('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | p(s(bool,X13)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f1820,f864])).
fof(f1820,plain,(
  ( ! [X146,X147] : (~sP10('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)),X147,X146) | p(s(bool,X146))) ) | $spl9),
  inference(resolution,[],[f1090,f711])).
fof(f4215,plain,(
  ( ! [X12,X13] : (sP5('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f1802,f864])).
fof(f1802,plain,(
  ( ! [X111,X110] : (sP5('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)),X110,X111) | p(s(bool,X110))) ) | $spl9),
  inference(resolution,[],[f1090,f677])).
fof(f4190,plain,(
  ( ! [X12,X13] : (sP11(X12,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f1015,f864])).
fof(f4178,plain,(
  ( ! [X12,X13] : (sP11(X12,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),sK15)),X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f1015,f864])).
fof(f4160,plain,(
  ( ! [X12,X13] : (sP5(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f1801,f864])).
fof(f1801,plain,(
  ( ! [X109,X108] : (sP5(X108,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)),X109) | p(s(bool,X108))) ) | $spl9),
  inference(resolution,[],[f1090,f677])).
fof(f4137,plain,(
  ( ! [X12,X13] : (sP5('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | p(s(bool,X13)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f1800,f864])).
fof(f1800,plain,(
  ( ! [X107,X106] : (sP5('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)),X106,X107) | p(s(bool,X107))) ) | $spl9),
  inference(resolution,[],[f1090,f676])).
fof(f4115,plain,(
  ( ! [X12,X13] : (sP5(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f1799,f864])).
fof(f1799,plain,(
  ( ! [X105,X104] : (sP5(X104,X105,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) | p(s(bool,X104))) ) | $spl9),
  inference(resolution,[],[f1090,f676])).
fof(f4092,plain,(
  ( ! [X12,X13] : (~sP5(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | ~p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f1793,f864])).
fof(f1793,plain,(
  ( ! [X92,X93] : (~sP5(X92,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)),X93) | ~p(s(bool,X92))) ) | $spl9),
  inference(resolution,[],[f1090,f667])).
fof(f4070,plain,(
  ( ! [X12,X13] : (~sP5(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f1792,f864])).
fof(f1792,plain,(
  ( ! [X90,X91] : (~sP5(X90,X91,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) | ~p(s(bool,X90))) ) | $spl9),
  inference(resolution,[],[f1090,f666])).
fof(f4046,plain,(
  ( ! [X12,X13] : (sP10('const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f1008,f864])).
fof(f4034,plain,(
  ( ! [X12,X13] : (sP10('const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),sK15)),X12,X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f1008,f864])).
fof(f4012,plain,(
  ( ! [X12,X13] : (~sP3('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | ~p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f1777,f864])).
fof(f1777,plain,(
  ( ! [X61,X60] : (~sP3('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)),X60,X61) | ~p(s(bool,X60))) ) | $spl9),
  inference(resolution,[],[f1090,f649])).
fof(f3987,plain,(
  ( ! [X12,X13] : (~sP3('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | ~p(s(bool,X13)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f1776,f864])).
fof(f1776,plain,(
  ( ! [X59,X58] : (~sP3('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)),X59,X58) | ~p(s(bool,X58))) ) | $spl9),
  inference(resolution,[],[f1090,f648])).
fof(f3960,plain,(
  ( ! [X12,X13] : (sP6('const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f997,f864])).
fof(f3948,plain,(
  ( ! [X12,X13] : (sP6('const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),sK15)),X12,X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f997,f864])).
fof(f3930,plain,(
  ( ! [X12,X13] : (sP1(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f1771,f864])).
fof(f1771,plain,(
  ( ! [X48,X49] : (sP1(X48,X49,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) | ~p(s(bool,X48))) ) | $spl9),
  inference(resolution,[],[f1090,f643])).
fof(f3903,plain,(
  ( ! [X12,X13] : (~sP1('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | ~p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f1763,f864])).
fof(f1763,plain,(
  ( ! [X33,X32] : (~sP1('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)),X32,X33) | ~p(s(bool,X32))) ) | $spl9),
  inference(resolution,[],[f1090,f632])).
fof(f3882,plain,(
  ( ! [X12,X13] : (~sP1(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f1762,f864])).
fof(f1762,plain,(
  ( ! [X30,X31] : (~sP1(X30,X31,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) | p(s(bool,X30))) ) | $spl9),
  inference(resolution,[],[f1090,f631])).
fof(f3855,plain,(
  ( ! [X12,X13] : (sP6(X12,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f996,f864])).
fof(f3843,plain,(
  ( ! [X12,X13] : (sP6(X12,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),sK15)),X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f996,f864])).
fof(f3825,plain,(
  ( ! [X12,X13] : (~sP1('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | p(s(bool,X13)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f1761,f864])).
fof(f1761,plain,(
  ( ! [X28,X29] : (~sP1('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)),X29,X28) | p(s(bool,X28))) ) | $spl9),
  inference(resolution,[],[f1090,f631])).
fof(f3798,plain,(
  ( ! [X12,X13] : (sP6(X12,X13,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f995,f864])).
fof(f3786,plain,(
  ( ! [X12,X13] : (sP6(X12,X13,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),sK15))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f995,f864])).
fof(f3767,plain,(
  ( ! [X8] : (sP12(X8,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X8)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl10),
  inference(superposition,[],[f2307,f864])).
fof(f2307,plain,(
  ( ! [X197] : (sP12(X197,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | p(s(bool,X197))) ) | $spl10),
  inference(resolution,[],[f1092,f756])).
fof(f3745,plain,(
  ( ! [X8] : (sP12('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X8) | p(s(bool,X8)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl10),
  inference(superposition,[],[f2306,f864])).
fof(f2306,plain,(
  ( ! [X196] : (sP12('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X196) | p(s(bool,X196))) ) | $spl10),
  inference(resolution,[],[f1092,f753])).
fof(f3723,plain,(
  ( ! [X8] : (~sP12(X8,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X8)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl10),
  inference(superposition,[],[f2305,f864])).
fof(f2305,plain,(
  ( ! [X195] : (~sP12(X195,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | ~p(s(bool,X195))) ) | $spl10),
  inference(resolution,[],[f1092,f752])).
fof(f3703,plain,(
  ( ! [X8] : (~sP12('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X8) | ~p(s(bool,X8)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl10),
  inference(superposition,[],[f2304,f864])).
fof(f2304,plain,(
  ( ! [X194] : (~sP12('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X194) | ~p(s(bool,X194))) ) | $spl10),
  inference(resolution,[],[f1092,f752])).
fof(f3679,plain,(
  ( ! [X12,X13] : (sP7(X12,X13,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f994,f864])).
fof(f3667,plain,(
  ( ! [X12,X13] : (sP7(X12,X13,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),sK15))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f994,f864])).
fof(f3646,plain,(
  ( ! [X8] : (sP12('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X8) | ~p(s(bool,X8)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f1839,f864])).
fof(f1839,plain,(
  ( ! [X181] : (sP12('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)),X181) | ~p(s(bool,X181))) ) | $spl9),
  inference(resolution,[],[f1090,f756])).
fof(f3621,plain,(
  ( ! [X12,X13] : (sP8(X12,X13,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f991,f864])).
fof(f3609,plain,(
  ( ! [X12,X13] : (sP8(X12,X13,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),sK15))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f991,f864])).
fof(f3588,plain,(
  ( ! [X8] : (sP12(X8,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X8)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f1838,f864])).
fof(f1838,plain,(
  ( ! [X180] : (sP12(X180,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) | ~p(s(bool,X180))) ) | $spl9),
  inference(resolution,[],[f1090,f753])).
fof(f3565,plain,(
  ( ! [X8] : (~sP12('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X8) | p(s(bool,X8)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f1837,f864])).
fof(f1837,plain,(
  ( ! [X179] : (~sP12('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)),X179) | p(s(bool,X179))) ) | $spl9),
  inference(resolution,[],[f1090,f751])).
fof(f3544,plain,(
  ( ! [X8] : (~sP12(X8,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X8)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f1836,f864])).
fof(f1836,plain,(
  ( ! [X178] : (~sP12(X178,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) | p(s(bool,X178))) ) | $spl9),
  inference(resolution,[],[f1090,f751])).
fof(f3521,plain,(
  ( ! [X12,X13] : (sP4(X12,X13,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f978,f864])).
fof(f3509,plain,(
  ( ! [X12,X13] : (sP4(X12,X13,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),sK15))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f978,f864])).
fof(f3486,plain,(
  ( ! [X12,X13] : (sP2(X12,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f962,f864])).
fof(f3474,plain,(
  ( ! [X12,X13] : (sP2(X12,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),sK15)),X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f962,f864])).
fof(f3451,plain,(
  ( ! [X12,X13] : (sP2(X12,X13,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f961,f864])).
fof(f3439,plain,(
  ( ! [X12,X13] : (sP2(X12,X13,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),sK15))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f961,f864])).
fof(f3421,plain,(
  ( ! [X12,X13] : (sP11(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl10),
  inference(superposition,[],[f1995,f864])).
fof(f1995,plain,(
  ( ! [X189,X188] : (sP11(X188,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X189)) ) | $spl10),
  inference(resolution,[],[f1092,f727])).
fof(f3404,plain,(
  ( ! [X12,X13] : (sP10('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl10),
  inference(superposition,[],[f1988,f864])).
fof(f1988,plain,(
  ( ! [X175,X174] : (sP10('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X174,X175)) ) | $spl10),
  inference(resolution,[],[f1092,f718])).
fof(f3382,plain,(
  ( ! [X12,X13] : (sP0(X12,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f948,f864])).
fof(f3370,plain,(
  ( ! [X12,X13] : (sP0(X12,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),sK15)),X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f948,f864])).
fof(f3349,plain,(
  ( ! [X12,X13] : (sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl10),
  inference(superposition,[],[f1977,f864])).
fof(f1977,plain,(
  ( ! [X152,X153] : (sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X152,X153)) ) | $spl10),
  inference(resolution,[],[f1092,f700])).
fof(f3329,plain,(
  ( ! [X12,X13] : (sP6(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl10),
  inference(superposition,[],[f1976,f864])).
fof(f1976,plain,(
  ( ! [X151,X150] : (sP6(X150,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X151)) ) | $spl10),
  inference(resolution,[],[f1092,f699])).
fof(f3309,plain,(
  ( ! [X12,X13] : (sP6(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl10),
  inference(superposition,[],[f1975,f864])).
fof(f1975,plain,(
  ( ! [X149,X148] : (sP6(X148,X149,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl10),
  inference(resolution,[],[f1092,f698])).
fof(f3290,plain,(
  ( ! [X12,X13] : (sP7(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl10),
  inference(superposition,[],[f1974,f864])).
fof(f1974,plain,(
  ( ! [X146,X147] : (sP7(X146,X147,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl10),
  inference(resolution,[],[f1092,f694])).
fof(f3271,plain,(
  ( ! [X12,X13] : (sP8(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl10),
  inference(superposition,[],[f1971,f864])).
fof(f1971,plain,(
  ( ! [X140,X141] : (sP8(X140,X141,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl10),
  inference(resolution,[],[f1092,f690])).
fof(f1156,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) | $spl37),
  inference(cnf_transformation,[],[f1156_D])).
fof(f1156_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) <=> ~$spl37),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl37])])).
fof(f3242,plain,(
  ( ! [X21,X20] : (sP4(X20,X21,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14)))) ) | $spl10),
  inference(subsumption_resolution,[],[f3234,f1092])).
fof(f3234,plain,(
  ( ! [X21,X20] : (sP4(X20,X21,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl10),
  inference(superposition,[],[f1958,f897])).
fof(f1958,plain,(
  ( ! [X114,X115] : (sP4(X114,X115,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl10),
  inference(resolution,[],[f1092,f673])).
fof(f3231,plain,(
  ( ! [X12,X13] : (sP4(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl10),
  inference(superposition,[],[f1958,f864])).
fof(f3221,plain,(
  ( ! [X21,X20] : (sP2(X20,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X21) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14)))) ) | $spl10),
  inference(subsumption_resolution,[],[f3214,f1092])).
fof(f3214,plain,(
  ( ! [X21,X20] : (sP2(X20,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X21) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl10),
  inference(superposition,[],[f1942,f897])).
fof(f1942,plain,(
  ( ! [X83,X82] : (sP2(X82,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X83)) ) | $spl10),
  inference(resolution,[],[f1092,f657])).
fof(f3211,plain,(
  ( ! [X12,X13] : (sP2(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl10),
  inference(superposition,[],[f1942,f864])).
fof(f3201,plain,(
  ( ! [X21,X20] : (sP2(X20,X21,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14)))) ) | $spl10),
  inference(subsumption_resolution,[],[f3194,f1092])).
fof(f3194,plain,(
  ( ! [X21,X20] : (sP2(X20,X21,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl10),
  inference(superposition,[],[f1941,f897])).
fof(f1941,plain,(
  ( ! [X80,X81] : (sP2(X80,X81,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl10),
  inference(resolution,[],[f1092,f656])).
fof(f3191,plain,(
  ( ! [X12,X13] : (sP2(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl10),
  inference(superposition,[],[f1941,f864])).
fof(f3181,plain,(
  ( ! [X21,X20] : (sP0(X20,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X21) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14)))) ) | $spl10),
  inference(subsumption_resolution,[],[f3174,f1092])).
fof(f3174,plain,(
  ( ! [X21,X20] : (sP0(X20,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X21) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl10),
  inference(superposition,[],[f1928,f897])).
fof(f1928,plain,(
  ( ! [X54,X55] : (sP0(X54,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X55)) ) | $spl10),
  inference(resolution,[],[f1092,f640])).
fof(f3171,plain,(
  ( ! [X12,X13] : (sP0(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl10),
  inference(superposition,[],[f1928,f864])).
fof(f3115,plain,(
  ( ! [X12,X13] : (sP7('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f1236,f864])).
fof(f3106,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2992,f2857])).
fof(f2857,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl18),
  inference(equality_resolution,[],[f1097])).
fof(f2992,plain,(
  ( ! [X391,X390,X392] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X390,bool),X391),s(fun(X390,bool),X392))))) ) | $spl18),
  inference(resolution,[],[f2857,f850])).
fof(f3105,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2974,f2857])).
fof(f2974,plain,(
  ( ! [X339,X340,X338] : (p(s(bool,'const.bool.IN_2'(s(X338,sK39(X339,X340,X338)),s(fun(X338,bool),X340))))) ) | $spl18),
  inference(resolution,[],[f2857,f850])).
fof(f2983,plain,(
  $false | $spl18),
  inference(resolution,[],[f2857,f863])).
fof(f2984,plain,(
  $false | $spl18),
  inference(resolution,[],[f2857,f868])).
fof(f2986,plain,(
  $false | ($spl2 | $spl18)),
  inference(resolution,[],[f2857,f915])).
fof(f2999,plain,(
  $false | $spl18),
  inference(resolution,[],[f2857,f902])).
fof(f3001,plain,(
  $false | ($spl10 | $spl18)),
  inference(resolution,[],[f2857,f1092])).
fof(f2811,plain,(
  ( ! [X12,X13] : (sP7(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f1235,f864])).
fof(f2791,plain,(
  ( ! [X12,X13] : (sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f1233,f864])).
fof(f2772,plain,(
  ( ! [X12,X13] : (sP8(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f1232,f864])).
fof(f2754,plain,(
  ( ! [X12,X13] : (sP4('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f1220,f864])).
fof(f2739,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2613,f2470])).
fof(f2470,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl16),
  inference(equality_resolution,[],[f1096])).
fof(f2613,plain,(
  ( ! [X422,X421,X423] : (~p(s(bool,'const.bool.IN_2'(s(X421,X422),s(fun(X421,bool),X423))))) ) | $spl16),
  inference(resolution,[],[f2470,f873])).
fof(f873,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))) )),
  inference(cnf_transformation,[],[f614])).
fof(f2738,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2611,f2470])).
fof(f2611,plain,(
  ( ! [X414,X416,X415,X417] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X414,bool),X415),s(fun(X414,bool),'const.pred_set.DELETE_2'(s(fun(X414,bool),X416),s(X414,X417))))))) ) | $spl16),
  inference(resolution,[],[f2470,f839])).
fof(f839,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X3),s(X0,X1))))))) )),
  inference(cnf_transformation,[],[f599])).
fof(f2737,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2607,f2470])).
fof(f2607,plain,(
  ( ! [X399,X401,X398,X400] : (~p(s(bool,'const.bool.IN_2'(s(X398,X399),s(fun(X398,bool),'const.pred_set.DIFF_2'(s(fun(X398,bool),X400),s(fun(X398,bool),X401))))))) ) | $spl16),
  inference(resolution,[],[f2470,f873])).
fof(f2736,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2603,f2470])).
fof(f2603,plain,(
  ( ! [X387,X388,X386] : (~p(s(bool,'const.bool.IN_2'(s(X386,X387),s(fun(X386,bool),X388))))) ) | $spl16),
  inference(resolution,[],[f2470,f839])).
fof(f2614,plain,(
  $false | $spl16),
  inference(resolution,[],[f2470,f892])).
fof(f2636,plain,(
  $false | ($spl9 | $spl16)),
  inference(resolution,[],[f2470,f1090])).
fof(f2426,plain,(
  ( ! [X12,X13] : (sP4(X12,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f1219,f864])).
fof(f2406,plain,(
  ( ! [X12,X13] : (sP2('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f1206,f864])).
fof(f2387,plain,(
  ( ! [X12,X13] : (sP0('const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X12,X13) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f1190,f864])).
fof(f2368,plain,(
  ( ! [X12,X13] : (sP0(X12,X13,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | $spl9),
  inference(superposition,[],[f1189,f864])).
fof(f2324,plain,(
  ( ! [X227,X226] : (p(s(bool,X226)) | p(s(bool,X227)) | s(bool,X227) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl10),
  inference(resolution,[],[f1092,f905])).
fof(f2226,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | s(bool,X38) != s(bool,X39)) ) | $spl10),
  inference(resolution,[],[f1092,f630])).
fof(f2223,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X28)) | ~p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X30) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl10),
  inference(resolution,[],[f1092,f629])).
fof(f2222,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl10),
  inference(resolution,[],[f1092,f628])).
fof(f2221,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | s(bool,X23) != s(bool,X24)) ) | $spl10),
  inference(resolution,[],[f1092,f626])).
fof(f2219,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X16) != s(bool,X18) | s(bool,X17) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl10),
  inference(resolution,[],[f1092,f624])).
fof(f1159,plain,(
  s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),sK14) | $spl39),
  inference(cnf_transformation,[],[f1159_D])).
fof(f1159_D,plain,(
  s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),sK14) <=> ~$spl39),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl39])])).
fof(f1155,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) | $spl36),
  inference(cnf_transformation,[],[f1155_D])).
fof(f1155_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f2218,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) | s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),sK14) | $spl10),
  inference(resolution,[],[f1092,f898])).
fof(f1158,plain,(
  s(fun(sK13,bool),'const.pred_set.EMPTY_0') = s(fun(sK13,bool),sK14) | $spl38),
  inference(cnf_transformation,[],[f1158_D])).
fof(f1158_D,plain,(
  s(fun(sK13,bool),'const.pred_set.EMPTY_0') = s(fun(sK13,bool),sK14) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f2217,plain,(
  s(fun(sK13,bool),'const.pred_set.EMPTY_0') = s(fun(sK13,bool),sK14) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) | $spl10),
  inference(resolution,[],[f1092,f897])).
fof(f2163,plain,(
  ( ! [X216,X217] : (p(s(bool,X216)) | ~p(s(bool,X217)) | s(bool,X217) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl11),
  inference(resolution,[],[f1091,f905])).
fof(f2065,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X26) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | s(bool,X25) != s(bool,X27)) ) | $spl11),
  inference(resolution,[],[f1091,f630])).
fof(f2063,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X20) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl11),
  inference(resolution,[],[f1091,f629])).
fof(f2061,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl11),
  inference(resolution,[],[f1091,f628])).
fof(f2060,plain,(
  ( ! [X12,X10,X11] : (~p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X11) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | s(bool,X10) != s(bool,X12)) ) | $spl11),
  inference(resolution,[],[f1091,f626])).
fof(f2058,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl11),
  inference(resolution,[],[f1091,f624])).
fof(f2018,plain,(
  ( ! [X227,X226] : (p(s(bool,X226)) | p(s(bool,X227)) | s(bool,X227) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl10),
  inference(resolution,[],[f1092,f905])).
fof(f1920,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | s(bool,X38) != s(bool,X39)) ) | $spl10),
  inference(resolution,[],[f1092,f630])).
fof(f1917,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X28)) | ~p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X30) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl10),
  inference(resolution,[],[f1092,f629])).
fof(f1916,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl10),
  inference(resolution,[],[f1092,f628])).
fof(f1915,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | s(bool,X23) != s(bool,X24)) ) | $spl10),
  inference(resolution,[],[f1092,f626])).
fof(f1913,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X16) != s(bool,X18) | s(bool,X17) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl10),
  inference(resolution,[],[f1092,f624])).
fof(f1912,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) | s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),sK14) | $spl10),
  inference(resolution,[],[f1092,f898])).
fof(f1911,plain,(
  s(fun(sK13,bool),'const.pred_set.EMPTY_0') = s(fun(sK13,bool),sK14) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) | $spl10),
  inference(resolution,[],[f1092,f897])).
fof(f1858,plain,(
  ( ! [X216,X217] : (p(s(bool,X216)) | ~p(s(bool,X217)) | s(bool,X217) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | $spl9),
  inference(resolution,[],[f1090,f905])).
fof(f1760,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X26) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) | s(bool,X25) != s(bool,X27)) ) | $spl9),
  inference(resolution,[],[f1090,f630])).
fof(f1758,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X20) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | $spl9),
  inference(resolution,[],[f1090,f629])).
fof(f1756,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | $spl9),
  inference(resolution,[],[f1090,f628])).
fof(f1755,plain,(
  ( ! [X12,X10,X11] : (~p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X11) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) | s(bool,X10) != s(bool,X12)) ) | $spl9),
  inference(resolution,[],[f1090,f626])).
fof(f1753,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | $spl9),
  inference(resolution,[],[f1090,f624])).
fof(f1711,plain,(
  ( ! [X227,X226] : (p(s(bool,X226)) | p(s(bool,X227)) | s(bool,X227) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | $spl8),
  inference(resolution,[],[f1089,f905])).
fof(f1613,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) | s(bool,X38) != s(bool,X39)) ) | $spl8),
  inference(resolution,[],[f1089,f630])).
fof(f1610,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X28)) | ~p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X30) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | $spl8),
  inference(resolution,[],[f1089,f629])).
fof(f1609,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | $spl8),
  inference(resolution,[],[f1089,f628])).
fof(f1608,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) | s(bool,X23) != s(bool,X24)) ) | $spl8),
  inference(resolution,[],[f1089,f626])).
fof(f1606,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X16) != s(bool,X18) | s(bool,X17) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | $spl8),
  inference(resolution,[],[f1089,f624])).
fof(f1605,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) | s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),sK15) | $spl8),
  inference(resolution,[],[f1089,f898])).
fof(f1604,plain,(
  s(fun(sK13,bool),'const.pred_set.EMPTY_0') = s(fun(sK13,bool),sK15) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) | $spl8),
  inference(resolution,[],[f1089,f897])).
fof(f1582,plain,(
  $false | $spl22),
  inference(subsumption_resolution,[],[f1456,f1099])).
fof(f1456,plain,(
  ( ! [X422,X421,X423] : (~p(s(bool,'const.bool.IN_2'(s(X421,X422),s(fun(X421,bool),X423))))) ) | $spl22),
  inference(resolution,[],[f1099,f873])).
fof(f1581,plain,(
  $false | $spl22),
  inference(subsumption_resolution,[],[f1454,f1099])).
fof(f1454,plain,(
  ( ! [X414,X416,X415,X417] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X414,bool),X415),s(fun(X414,bool),'const.pred_set.DELETE_2'(s(fun(X414,bool),X416),s(X414,X417))))))) ) | $spl22),
  inference(resolution,[],[f1099,f839])).
fof(f1580,plain,(
  $false | $spl22),
  inference(subsumption_resolution,[],[f1450,f1099])).
fof(f1450,plain,(
  ( ! [X399,X401,X398,X400] : (~p(s(bool,'const.bool.IN_2'(s(X398,X399),s(fun(X398,bool),'const.pred_set.DIFF_2'(s(fun(X398,bool),X400),s(fun(X398,bool),X401))))))) ) | $spl22),
  inference(resolution,[],[f1099,f873])).
fof(f1579,plain,(
  $false | $spl22),
  inference(subsumption_resolution,[],[f1446,f1099])).
fof(f1446,plain,(
  ( ! [X387,X388,X386] : (~p(s(bool,'const.bool.IN_2'(s(X386,X387),s(fun(X386,bool),X388))))) ) | $spl22),
  inference(resolution,[],[f1099,f839])).
fof(f1457,plain,(
  $false | $spl22),
  inference(resolution,[],[f1099,f892])).
fof(f1479,plain,(
  $false | ($spl9 | $spl22)),
  inference(resolution,[],[f1099,f1090])).
fof(f1280,plain,(
  ( ! [X216,X217] : (p(s(bool,X216)) | ~p(s(bool,X217)) | s(bool,X217) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | $spl9),
  inference(resolution,[],[f1090,f905])).
fof(f1182,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X26) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) | s(bool,X25) != s(bool,X27)) ) | $spl9),
  inference(resolution,[],[f1090,f630])).
fof(f1180,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X20) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | $spl9),
  inference(resolution,[],[f1090,f629])).
fof(f1178,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | $spl9),
  inference(resolution,[],[f1090,f628])).
fof(f1177,plain,(
  ( ! [X12,X10,X11] : (~p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X11) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))) | s(bool,X10) != s(bool,X12)) ) | $spl9),
  inference(resolution,[],[f1090,f626])).
fof(f1175,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) ) | $spl9),
  inference(resolution,[],[f1090,f624])).
fof(f1148,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))))) | $spl1),
  inference(trivial_inequality_removal,[],[f1143])).
fof(f1143,plain,(
  s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),'const.pred_set.EMPTY_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))))) | $spl1),
  inference(superposition,[],[f914,f897])).
fof(f1165,plain,(
  ( ! [X5] : (s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),X5) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))))))) ) | $spl1),
  inference(subsumption_resolution,[],[f1142,f860])).
fof(f1142,plain,(
  ( ! [X5] : (s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),X5) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X5))))) ) | $spl1),
  inference(superposition,[],[f914,f866])).
fof(f1163,plain,(
  s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),sK14) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) | $spl1),
  inference(forward_demodulation,[],[f1131,f890])).
fof(f890,plain,(
  ( ! [X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(cnf_transformation,[],[f398])).
fof(f398,plain,(
  ! [X0,X1] : s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))),
  inference(flattening,[],[f397])).
fof(f397,plain,(
  ! [X0] : ! [X1] : s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X6] : ! [X7] : s(fun(X6,bool),X7) = s(fun(X6,bool),'const.pred_set.DIFF_2'(s(fun(X6,bool),X7),s(fun(X6,bool),'const.pred_set.EMPTY_0')))),
  file('thm.pred_set.SUBSET_DIFF_EMPTY','thm.pred_set.DIFF_EMPTY')).
fof(f1131,plain,(
  s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) | $spl1),
  inference(superposition,[],[f914,f897])).
fof(f1162,plain,(
  ( ! [X5] : (s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),X5))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK15))))) ) | $spl40),
  inference(cnf_transformation,[],[f1162_D])).
fof(f1162_D,plain,(
  ( ! [X5] : (s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),X5))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK15))))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f1161,plain,(
  ( ! [X5] : (s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK15)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))))) ) | $spl1),
  inference(subsumption_resolution,[],[f1130,f860])).
fof(f1130,plain,(
  ( ! [X5] : (s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK15)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X5))))) ) | $spl1),
  inference(superposition,[],[f914,f866])).
fof(f1157,plain,(
  s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),sK14) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) | $spl1),
  inference(forward_demodulation,[],[f1128,f890])).
fof(f1128,plain,(
  s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) | $spl1),
  inference(superposition,[],[f914,f864])).
fof(f1154,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | $spl1),
  inference(subsumption_resolution,[],[f1119,f891])).
fof(f1119,plain,(
  s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),'const.pred_set.EMPTY_0'),s(fun(sK13,bool),sK15))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | $spl1),
  inference(superposition,[],[f914,f897])).
fof(f1153,plain,(
  ( ! [X5] : (s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK15))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK14))))) ) | $spl34),
  inference(cnf_transformation,[],[f1153_D])).
fof(f1153_D,plain,(
  ( ! [X5] : (s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK15))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK14))))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f1152,plain,(
  ( ! [X5] : (s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK15))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl1),
  inference(subsumption_resolution,[],[f1118,f860])).
fof(f1118,plain,(
  ( ! [X5] : (s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK15))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X5))))) ) | $spl1),
  inference(superposition,[],[f914,f866])).
fof(f1062,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) | $spl2),
  inference(superposition,[],[f915,f897])).
fof(f1109,plain,(
  ( ! [X5] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),X5)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK15))))) ) | $spl30),
  inference(cnf_transformation,[],[f1109_D])).
fof(f1109_D,plain,(
  ( ! [X5] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),X5)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK15))))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f1108,plain,(
  ( ! [X5] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),X5)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK15)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15))))) ) | $spl2),
  inference(subsumption_resolution,[],[f1061,f860])).
fof(f1061,plain,(
  ( ! [X5] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),X5)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK15)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X5))))) ) | $spl2),
  inference(superposition,[],[f915,f866])).
fof(f1059,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) | $spl2),
  inference(superposition,[],[f915,f864])).
fof(f1102,plain,(
  ( ! [X5] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK15)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK14))))) ) | $spl24),
  inference(cnf_transformation,[],[f1102_D])).
fof(f1102_D,plain,(
  ( ! [X5] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK15)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK14))))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f1101,plain,(
  ( ! [X5] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK15)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl2),
  inference(subsumption_resolution,[],[f1049,f860])).
fof(f1049,plain,(
  ( ! [X5] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK15)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X5),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),X5))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X5))))) ) | $spl2),
  inference(superposition,[],[f915,f866])).
fof(f1038,plain,(
  ( ! [X213,X214] : (p(s(bool,X213)) | p(s(bool,X214)) | s(bool,X214) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))) ) | $spl2),
  inference(resolution,[],[f915,f905])).
fof(f940,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))) | s(bool,X25) != s(bool,X26)) ) | $spl2),
  inference(resolution,[],[f915,f630])).
fof(f937,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X17) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))) ) | $spl2),
  inference(resolution,[],[f915,f629])).
fof(f936,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))) ) | $spl2),
  inference(resolution,[],[f915,f628])).
fof(f935,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))) | s(bool,X10) != s(bool,X11)) ) | $spl2),
  inference(resolution,[],[f915,f626])).
fof(f933,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))) ) | $spl2),
  inference(resolution,[],[f915,f624])).
fof(f932,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) | $spl2),
  inference(resolution,[],[f915,f867])).
fof(f867,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f445])).
fof(f445,plain,(
  ! [X0,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) | ! [X2] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))) | p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2))))))),
  inference(ennf_transformation,[],[f358])).
fof(f358,plain,(
  ! [X0,X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) => ! [X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))) => p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2))))))),
  inference(flattening,[],[f357])).
fof(f357,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) => ! [X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))) => p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f247])).
fof(f247,axiom,(
  ! [X6] : ! [X7] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X7)))) => ! [X8] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X7)))) => p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X8))))))),
  file('thm.pred_set.SUBSET_DIFF_EMPTY','thm.pred_set.SUBSET_FINITE')).
fof(f931,plain,(
  s(fun(sK13,bool),sK14) = s(fun(sK13,bool),sK15) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK13,bool),sK15))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | $spl2),
  inference(resolution,[],[f915,f866])).
fof(f928,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK15)))) | $spl2),
  inference(resolution,[],[f915,f860])).
fof(f927,plain,(
  s(fun(sK13,bool),sK14) = s(fun(sK13,bool),sK15) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)))) | $spl2),
  inference(resolution,[],[f915,f859])).
fof(f926,plain,(
  s(fun(sK13,bool),sK14) = s(fun(sK13,bool),sK15) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)))) | $spl2),
  inference(resolution,[],[f915,f859])).
fof(f925,plain,(
  s(fun(sK13,bool),sK14) = s(fun(sK13,bool),sK15) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)))) | $spl2),
  inference(resolution,[],[f915,f858])).
fof(f924,plain,(
  s(fun(sK13,bool),sK14) = s(fun(sK13,bool),sK15) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),sK14)))) | $spl2),
  inference(resolution,[],[f915,f858])).
fof(f915,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))) | $spl2),
  inference(cnf_transformation,[],[f915_D])).
fof(f915_D,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f621,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))) | s(fun(sK13,bool),'const.pred_set.EMPTY_0') = s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))),
  inference(cnf_transformation,[],[f479])).
fof(f479,plain,(
  (s(fun(sK13,bool),'const.pred_set.EMPTY_0') = s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))))) & (s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13,sK14,sK15])],[f478])).
fof(f478,plain,(
  ? [X0,X1,X2] : ((s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) & (s(fun(X0,bool),'const.pred_set.EMPTY_0') != s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(nnf_transformation,[],[f415])).
fof(f415,plain,(
  ? [X0,X1,X2] : (s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) <~> p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))),
  inference(ennf_transformation,[],[f263])).
fof(f263,plain,(
  ~! [X0,X1,X2] : (s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) <=> p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))),
  inference(flattening,[],[f262])).
fof(f262,plain,(
  ~! [X0] : ! [X1,X2] : (s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) <=> p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : (s(fun(X6,bool),'const.pred_set.EMPTY_0') = s(fun(X6,bool),'const.pred_set.DIFF_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8))) <=> p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : (s(fun(X6,bool),'const.pred_set.EMPTY_0') = s(fun(X6,bool),'const.pred_set.DIFF_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8))) <=> p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))))),
  file('thm.pred_set.SUBSET_DIFF_EMPTY','thm.pred_set.SUBSET_DIFF_EMPTY')).
fof(f916,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))) | $spl3),
  inference(cnf_transformation,[],[f916_D])).
fof(f916_D,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f622,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))) | s(fun(sK13,bool),'const.pred_set.EMPTY_0') != s(fun(sK13,bool),'const.pred_set.DIFF_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK15)))),
  inference(cnf_transformation,[],[f479])).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 1166
% Passive clauses: 19456
% Generated clauses: 53465
% Final active clauses: 832
% Final passive clauses: 11291
% Input formulas: 260
% Initial clauses: 284
% 
% Selected by SInE selection: 84
% SInE iterations: 1
% 
% Duplicate literals: 1515
% Trivial inequalities: 57
% Fw subsumption resolutions: 6244
% Fw demodulations: 4123
% Bw demodulations: 74
% 
% Simple tautologies: 2884
% Equational tautologies: 42
% Forward subsumptions: 18079
% 
% Binary resolution: 16751
% Forward superposition: 16405
% Backward superposition: 7520
% Self superposition: 10
% Equality resolution: 13
% 
% Split clauses: 945
% Split components: 2038
% SAT solver clauses: 1011
% SAT solver unit clauses: 22
% SAT solver binary clauses: 25
% 
% Sat splits: 945
% Sat splitting refutations: 47
% 
% Memory used [KB]: 21364
% Time elapsed: 1.123 s
% ------------------------------
% ------------------------------
% Success in time 12.88 s
