Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.ringNorm
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.ringNorm
% SZS output start Proof for thm.ringNorm
fof(f9245,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f9216])).
fof(f9216,plain,(
  s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0') != s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0')),
  inference(superposition,[],[f65,f9121])).
fof(f9121,plain,(
  ( ! [X6,X4,X5] : (s('type.canonical.canonical_sum'(X4),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'(X4),'const.ringNorm.r_canonical_sum_scalar_3'(s('type.ring.ring'(X4),X5),s(X4,X6),s('type.canonical.canonical_sum'(X4),'const.canonical.Nil_monom_0')))) )),
  inference(superposition,[],[f4062,f71])).
fof(f71,plain,(
  ( ! [X2,X0,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar_3'(s('type.semi_ring.semi_ring'(X0),X1),s(X0,X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))) )),
  inference(cnf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0,X1,X2] : s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar_3'(s('type.semi_ring.semi_ring'(X0),X1),s(X0,X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))),
  inference(flattening,[],[f30])).
fof(f30,plain,(
  ! [X0] : ! [X1,X2] : s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar_3'(s('type.semi_ring.semi_ring'(X0),X1),s(X0,X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6] : ! [X11,X8] : s('type.canonical.canonical_sum'(X6),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'(X6),'const.canonical.canonical_sum_scalar_3'(s('type.semi_ring.semi_ring'(X6),X11),s(X6,X8),s('type.canonical.canonical_sum'(X6),'const.canonical.Nil_monom_0')))),
  file('thm.ringNorm.canonical_sum_scalar_def|split|2','thm.canonical.canonical_sum_scalar_def|split|2')).
fof(f4062,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_scalar_3'(s('type.ring.ring'(X0),X1),s(X0,X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar_3'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s(X0,X2),s('type.canonical.canonical_sum'(X0),X3)))) )),
  inference(forward_demodulation,[],[f3997,f78])).
fof(f78,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_scalar_3'(s('type.ring.ring'(X0),X1),s(X0,X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun(X0,fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_scalar_1'(s('type.ring.ring'(X0),X1))),s(X0,X2))),s('type.canonical.canonical_sum'(X0),X3)))) )),
  inference(cnf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0,X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_scalar_3'(s('type.ring.ring'(X0),X1),s(X0,X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun(X0,fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_scalar_1'(s('type.ring.ring'(X0),X1))),s(X0,X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(flattening,[],[f42])).
fof(f42,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_scalar_3'(s('type.ring.ring'(X0),X1),s(X0,X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun(X0,fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_scalar_1'(s('type.ring.ring'(X0),X1))),s(X0,X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X6] : ! [X19,X20,X21] : s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_canonical_sum_scalar_3'(s('type.ring.ring'(X6),X19),s(X6,X20),s('type.canonical.canonical_sum'(X6),X21))) = s('type.canonical.canonical_sum'(X6),happ(s(fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6)),happ(s(fun(X6,fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6))),'const.ringNorm.r_canonical_sum_scalar_1'(s('type.ring.ring'(X6),X19))),s(X6,X20))),s('type.canonical.canonical_sum'(X6),X21)))),
  file('thm.ringNorm.canonical_sum_scalar_def|split|2','ARITY_1_3_const.ringNorm.r_canonical_sum_scalar')).
fof(f3997,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar_3'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s(X0,X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun(X0,fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_scalar_1'(s('type.ring.ring'(X0),X1))),s(X0,X2))),s('type.canonical.canonical_sum'(X0),X3)))) )),
  inference(superposition,[],[f79,f69])).
fof(f69,plain,(
  ( ! [X0,X1] : (s(fun(X0,fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_scalar_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_scalar_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f27])).
fof(f27,plain,(
  ! [X0,X1] : s(fun(X0,fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_scalar_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_scalar_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f26])).
fof(f26,plain,(
  ! [X0] : ! [X1] : s(fun(X0,fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_scalar_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_scalar_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6] : ! [X7] : s(fun(X6,fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6))),'const.ringNorm.r_canonical_sum_scalar_1'(s('type.ring.ring'(X6),X7))) = s(fun(X6,fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6))),'const.canonical.canonical_sum_scalar_1'(s('type.semi_ring.semi_ring'(X6),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X6),X7)))))),
  file('thm.ringNorm.canonical_sum_scalar_def|split|2','thm.ringNorm.r_canonical_sum_scalar_def')).
fof(f79,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar_3'(s('type.semi_ring.semi_ring'(X0),X1),s(X0,X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun(X0,fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_scalar_1'(s('type.semi_ring.semi_ring'(X0),X1))),s(X0,X2))),s('type.canonical.canonical_sum'(X0),X3)))) )),
  inference(cnf_transformation,[],[f45])).
fof(f45,plain,(
  ! [X0,X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar_3'(s('type.semi_ring.semi_ring'(X0),X1),s(X0,X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun(X0,fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_scalar_1'(s('type.semi_ring.semi_ring'(X0),X1))),s(X0,X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(flattening,[],[f44])).
fof(f44,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar_3'(s('type.semi_ring.semi_ring'(X0),X1),s(X0,X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun(X0,fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_scalar_1'(s('type.semi_ring.semi_ring'(X0),X1))),s(X0,X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X6] : ! [X19,X20,X21] : s('type.canonical.canonical_sum'(X6),'const.canonical.canonical_sum_scalar_3'(s('type.semi_ring.semi_ring'(X6),X19),s(X6,X20),s('type.canonical.canonical_sum'(X6),X21))) = s('type.canonical.canonical_sum'(X6),happ(s(fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6)),happ(s(fun(X6,fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6))),'const.canonical.canonical_sum_scalar_1'(s('type.semi_ring.semi_ring'(X6),X19))),s(X6,X20))),s('type.canonical.canonical_sum'(X6),X21)))),
  file('thm.ringNorm.canonical_sum_scalar_def|split|2','ARITY_1_3_const.canonical.canonical_sum_scalar')).
fof(f65,plain,(
  s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0') != s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_canonical_sum_scalar_3'(s('type.ring.ring'(sK0),sK1),s(sK0,sK2),s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0')))),
  inference(cnf_transformation,[],[f62])).
fof(f62,plain,(
  s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0') != s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_canonical_sum_scalar_3'(s('type.ring.ring'(sK0),sK1),s(sK0,sK2),s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0')))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f60])).
fof(f60,plain,(
  ? [X0,X1,X2] : s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0') != s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_scalar_3'(s('type.ring.ring'(X0),X1),s(X0,X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))),
  inference(ennf_transformation,[],[f24])).
fof(f24,plain,(
  ~! [X0,X1,X2] : s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_scalar_3'(s('type.ring.ring'(X0),X1),s(X0,X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))),
  inference(flattening,[],[f23])).
fof(f23,plain,(
  ~! [X0] : ! [X1,X2] : s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_scalar_3'(s('type.ring.ring'(X0),X1),s(X0,X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : s('type.canonical.canonical_sum'(X6),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_canonical_sum_scalar_3'(s('type.ring.ring'(X6),X7),s(X6,X8),s('type.canonical.canonical_sum'(X6),'const.canonical.Nil_monom_0')))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : s('type.canonical.canonical_sum'(X6),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_canonical_sum_scalar_3'(s('type.ring.ring'(X6),X7),s(X6,X8),s('type.canonical.canonical_sum'(X6),'const.canonical.Nil_monom_0')))),
  file('thm.ringNorm.canonical_sum_scalar_def|split|2','thm.ringNorm.canonical_sum_scalar_def|split|2')).
% SZS output end Proof for thm.ringNorm
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 85
% Passive clauses: 694
% Generated clauses: 8991
% Final active clauses: 81
% Final passive clauses: 568
% Input formulas: 21
% Initial clauses: 22
% 
% Duplicate literals: 732
% Trivial inequalities: 23
% Fw subsumption resolutions: 774
% Fw demodulations: 1211
% Bw demodulations: 2
% 
% Simple tautologies: 119
% Equational tautologies: 17
% Forward subsumptions: 5056
% 
% Binary resolution: 10
% Forward superposition: 3687
% Backward superposition: 2370
% Self superposition: 113
% Equality factoring: 32
% Equality resolution: 2
% 
% Split clauses: 197
% Split components: 399
% SAT solver clauses: 374
% SAT solver unit clauses: 1
% SAT solver binary clauses: 10
% 
% Sat splits: 197
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 374
% TWLsolver calls for satisfiability: 49
% 
% Memory used [KB]: 7164
% Time elapsed: 0.153 s
% ------------------------------
% ------------------------------
% Success in time 0.155 s
