Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.divides
% (69779)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1078
% Passive clauses: 5957
% Generated clauses: 33576
% Final active clauses: 808
% Final passive clauses: 4224
% Input formulas: 108
% Initial clauses: 259
% 
% Duplicate literals: 250
% Trivial inequalities: 66
% Fw subsumption resolutions: 2041
% Fw demodulations: 3405
% Bw demodulations: 250
% 
% Simple tautologies: 221
% Equational tautologies: 27
% Forward subsumptions: 21731
% Fw demodulations to eq. taut.: 21
% 
% Binary resolution: 6482
% Factoring: 1362
% Forward superposition: 9697
% Backward superposition: 9566
% Self superposition: 40
% Equality factoring: 2
% Equality resolution: 11
% 
% Split clauses: 249
% Split components: 499
% SAT solver clauses: 443
% SAT solver unit clauses: 12
% SAT solver binary clauses: 65
% 
% Sat splits: 249
% Sat splitting refutations: 11
% 
% TWLsolver clauses: 443
% TWLsolver calls for satisfiability: 391
% 
% Memory used [KB]: 17526
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.divides
% (69806)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 621
% Passive clauses: 3200
% Generated clauses: 23711
% Final active clauses: 588
% Final passive clauses: 2284
% Input formulas: 108
% Initial clauses: 259
% 
% Duplicate literals: 207
% Trivial inequalities: 58
% Fw subsumption resolutions: 1407
% Fw demodulations: 3640
% Bw demodulations: 304
% Condensations: 177
% Global subsumptions: 19
% 
% Simple tautologies: 289
% Equational tautologies: 19
% Forward subsumptions: 10836
% Fw demodulations to eq. taut.: 21
% 
% Binary resolution: 3128
% Factoring: 264
% Forward superposition: 6331
% Backward superposition: 4542
% Self superposition: 68
% Equality factoring: 26
% Equality resolution: 13
% 
% Split clauses: 54
% Split components: 108
% SAT solver clauses: 4048
% SAT solver unit clauses: 317
% SAT solver binary clauses: 88
% SAT solver learnt clauses: 38
% SAT solver learnt literals: 12
% 
% Sat splits: 54
% Sat splitting refutations: 10
% 
% TWLsolver clauses: 3972
% TWLsolver calls for satisfiability: 1874
% 
% Memory used [KB]: 18677
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.divides
% (69843)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 3273
% Passive clauses: 25168
% Generated clauses: 297986
% Final active clauses: 2028
% Final passive clauses: 6658
% Input formulas: 108
% Initial clauses: 259
% Discarded non-redundant clauses: 41266
% 
% Duplicate literals: 222
% Trivial inequalities: 166
% Fw subsumption resolutions: 6449
% Fw demodulations: 61017
% Bw demodulations: 7366
% Condensations: 311
% Global subsumptions: 264
% 
% Simple tautologies: 151
% Equational tautologies: 120
% Forward subsumptions: 199742
% Fw demodulations to eq. taut.: 727
% Bw demodulations to eq. taut.: 48
% 
% Unit resulting resolution: 46336
% Factoring: 4
% Forward superposition: 97950
% Backward superposition: 71269
% Self superposition: 144
% Equality factoring: 26
% Equality resolution: 53
% 
% Split clauses: 142
% Split components: 285
% SAT solver clauses: 28682
% SAT solver unit clauses: 16117
% SAT solver binary clauses: 7190
% 
% Sat splits: 142
% Sat splitting refutations: 2964
% 
% Memory used [KB]: 169250
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.divides
% (70117)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 2846
% Passive clauses: 59423
% Generated clauses: 191879
% Final active clauses: 2766
% Final passive clauses: 53186
% Input formulas: 108
% Initial clauses: 147
% 
% Selected by SInE selection: 37
% SInE iterations: 1
% 
% Duplicate literals: 3498
% Trivial inequalities: 195
% Fw subsumption resolutions: 15287
% Fw demodulations: 8906
% Bw demodulations: 27
% 
% Simple tautologies: 5841
% Equational tautologies: 473
% Forward subsumptions: 94523
% Fw demodulations to eq. taut.: 4
% 
% Binary resolution: 139751
% Factoring: 401
% Forward superposition: 18068
% Backward superposition: 4846
% Self superposition: 14
% Equality resolution: 53
% 
% Split clauses: 5570
% Split components: 13453
% SAT solver clauses: 6116
% SAT solver unit clauses: 6
% SAT solver binary clauses: 151
% 
% Sat splits: 5570
% Sat splitting refutations: 37
% 
% Memory used [KB]: 58335
% Time elapsed: 4.0000 s
% ------------------------------
% ------------------------------
% remaining time: 142 next slice time: 6
lrs+1002_2:3_br=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:sas=minisat:stl=30:sac=on:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none:sp=reverse_arity:urr=on_4 on thm.divides
% (70244)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 644
% Passive clauses: 4240
% Generated clauses: 28816
% Final active clauses: 542
% Final passive clauses: 2589
% Input formulas: 108
% Initial clauses: 259
% 
% Duplicate literals: 144
% Trivial inequalities: 67
% Fw subsumption resolutions: 659
% Fw demodulations: 3852
% Bw demodulations: 1036
% 
% Simple tautologies: 193
% Equational tautologies: 28
% Forward subsumptions: 19395
% Fw demodulations to eq. taut.: 109
% Bw demodulations to eq. taut.: 4
% 
% Unit resulting resolution: 5729
% Factoring: 3
% Forward superposition: 10703
% Backward superposition: 6298
% Self superposition: 89
% Equality factoring: 2
% Equality resolution: 7
% 
% Split clauses: 5
% Split components: 10
% SAT solver clauses: 4376
% SAT solver unit clauses: 2154
% SAT solver binary clauses: 819
% 
% Sat splits: 5
% Sat splitting refutations: 129
% 
% Memory used [KB]: 15607
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 135 next slice time: 15
lrs+1011_12_bs=on:bsr=unit_only:cond=on:gs=on:gsaa=from_current:gsssp=full:nwc=1.1:sas=minisat:stl=60:sos=all:sac=on:sdd=off:sser=off:sfr=on:ssfp=100000:ssfq=1.2:ssnc=none:sp=reverse_arity:updr=off_14 on thm.divides
% (70262)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 786
% Passive clauses: 9428
% Generated clauses: 44020
% Final active clauses: 386
% Final passive clauses: 4153
% Input formulas: 108
% Initial clauses: 259
% 
% Duplicate literals: 1407
% Trivial inequalities: 133
% Fw subsumption resolutions: 2392
% Bw subsumption resolutions: 69
% Fw demodulations: 9664
% Bw demodulations: 1361
% Condensations: 3161
% Global subsumptions: 3
% 
% Simple tautologies: 2461
% Equational tautologies: 26
% Forward subsumptions: 14642
% Backward subsumptions: 774
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 8291
% Factoring: 3
% Forward superposition: 9291
% Backward superposition: 2250
% Equality resolution: 8
% 
% Split clauses: 6
% Split components: 12
% SAT solver clauses: 9440
% SAT solver unit clauses: 162
% SAT solver binary clauses: 1279
% 
% Sat splits: 6
% Sat splitting refutations: 1
% 
% Memory used [KB]: 19189
% Time elapsed: 1.600 s
% ------------------------------
% ------------------------------
% remaining time: 119 next slice time: 9
lrs+10_8:1_bd=preordered:bs=on:ccuc=first:fsr=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:nicw=on:sas=minisat:stl=120:sos=on:sscc=on:sser=off:sfr=on:ssfp=4000:ssfq=1.0:ssnc=none:sp=reverse_arity:urr=on_7 on thm.divides
% (70311)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 653
% Passive clauses: 12733
% Generated clauses: 30937
% Final active clauses: 227
% Final passive clauses: 1837
% Input formulas: 108
% Initial clauses: 259
% 
% Duplicate literals: 34
% Trivial inequalities: 8
% Fw demodulations: 3727
% Bw demodulations: 5607
% Global subsumptions: 16
% 
% Simple tautologies: 210
% Equational tautologies: 16
% Forward subsumptions: 9849
% Backward subsumptions: 1992
% Fw demodulations to eq. taut.: 26
% Bw demodulations to eq. taut.: 4
% 
% Binary resolution: 1853
% Unit resulting resolution: 8899
% Factoring: 3
% Forward superposition: 1096
% Backward superposition: 1320
% Self superposition: 5
% Equality resolution: 6
% 
% Split clauses: 141
% Split components: 346
% SAT solver clauses: 21500
% SAT solver unit clauses: 38
% SAT solver binary clauses: 15463
% 
% Sat splits: 141
% Sat splitting refutations: 12
% 
% Memory used [KB]: 19445
% Time elapsed: 1.0000 s
% ------------------------------
% ------------------------------
% remaining time: 109 next slice time: 85
lrs+1_5:4_cond=on:fsr=off:fde=none:gs=on:gsem=on:lwlo=on:nm=64:nwc=1:stl=60:sos=all:spl=off_80 on thm.divides
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.divides
% SZS output start Proof for thm.divides
fof(f5457,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f5397])).
fof(f5397,plain,(
  s(bool,'F_0') != s(bool,'F_0')),
  inference(resolution,[],[f5274,f399])).
fof(f399,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f280])).
fof(f280,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f157])).
fof(f157,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f156])).
fof(f156,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X10] : (s(bool,'F_0') = s(bool,X10) <=> ~p(s(bool,X10)))),
  file('thm.divides.PRIMES_|49||49|','thm.bool.EQ_CLAUSES|split|2')).
fof(f5274,plain,(
  p(s(bool,'F_0'))),
  inference(duplicate_literal_removal,[],[f5273])).
fof(f5273,plain,(
  p(s(bool,'F_0')) | p(s(bool,'F_0'))),
  inference(forward_demodulation,[],[f5193,f2181])).
fof(f2181,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.divides.PRIMES_1'(s('type.num.num',sK1))),s('type.num.num','const.divides.PRIMES_1'(s('type.num.num',sK1)))))),
  inference(resolution,[],[f1948,f398])).
fof(f398,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f279])).
fof(f279,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f155])).
fof(f155,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f154])).
fof(f154,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X10] : (s(bool,'F_0') = s(bool,X10) <=> ~p(s(bool,X10)))),
  file('thm.divides.PRIMES_|49||49|','thm.bool.EQ_CLAUSES|split|3')).
fof(f1948,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.divides.PRIMES_1'(s('type.num.num',sK1))),s('type.num.num','const.divides.PRIMES_1'(s('type.num.num',sK1))))))),
  inference(trivial_inequality_removal,[],[f1947])).
fof(f1947,plain,(
  s('type.num.num','const.divides.PRIMES_1'(s('type.num.num',sK1))) != s('type.num.num','const.divides.PRIMES_1'(s('type.num.num',sK1))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.divides.PRIMES_1'(s('type.num.num',sK1))),s('type.num.num','const.divides.PRIMES_1'(s('type.num.num',sK1))))))),
  inference(forward_demodulation,[],[f1868,f342])).
fof(f342,plain,(
  s('type.num.num','const.divides.PRIMES_1'(s('type.num.num',sK1))) = s('type.num.num','const.divides.PRIMES_1'(s('type.num.num',sK2)))),
  inference(cnf_transformation,[],[f274])).
fof(f274,plain,(
  s('type.num.num','const.divides.PRIMES_1'(s('type.num.num',sK1))) = s('type.num.num','const.divides.PRIMES_1'(s('type.num.num',sK2))) & s('type.num.num',sK1) != s('type.num.num',sK2)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1,sK2])],[f238])).
fof(f238,plain,(
  ? [X0,X1] : (s('type.num.num','const.divides.PRIMES_1'(s('type.num.num',X0))) = s('type.num.num','const.divides.PRIMES_1'(s('type.num.num',X1))) & s('type.num.num',X0) != s('type.num.num',X1))),
  inference(ennf_transformation,[],[f110])).
fof(f110,plain,(
  ~! [X0,X1] : (s('type.num.num','const.divides.PRIMES_1'(s('type.num.num',X0))) = s('type.num.num','const.divides.PRIMES_1'(s('type.num.num',X1))) => s('type.num.num',X0) = s('type.num.num',X1))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (s('type.num.num','const.divides.PRIMES_1'(s('type.num.num',X6))) = s('type.num.num','const.divides.PRIMES_1'(s('type.num.num',X7))) => s('type.num.num',X6) = s('type.num.num',X7))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (s('type.num.num','const.divides.PRIMES_1'(s('type.num.num',X6))) = s('type.num.num','const.divides.PRIMES_1'(s('type.num.num',X7))) => s('type.num.num',X6) = s('type.num.num',X7))),
  file('thm.divides.PRIMES_|49||49|','thm.divides.PRIMES_|49||49|')).
fof(f1868,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.divides.PRIMES_1'(s('type.num.num',sK1))),s('type.num.num','const.divides.PRIMES_1'(s('type.num.num',sK1)))))) | s('type.num.num','const.divides.PRIMES_1'(s('type.num.num',sK1))) != s('type.num.num','const.divides.PRIMES_1'(s('type.num.num',sK2)))),
  inference(superposition,[],[f1175,f342])).
fof(f1175,plain,(
  ( ! [X76] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.divides.PRIMES_1'(s('type.num.num',sK1))),s('type.num.num',X76)))) | s('type.num.num',X76) != s('type.num.num','const.divides.PRIMES_1'(s('type.num.num',sK1)))) )),
  inference(forward_demodulation,[],[f1174,f342])).
fof(f1174,plain,(
  ( ! [X76] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.divides.PRIMES_1'(s('type.num.num',sK1))),s('type.num.num',X76)))) | s('type.num.num',X76) != s('type.num.num','const.divides.PRIMES_1'(s('type.num.num',sK2)))) )),
  inference(forward_demodulation,[],[f1052,f448])).
fof(f448,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f181])).
fof(f181,plain,(
  ! [X0,X1] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6,X7] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num',X7))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))),s('type.num.num',X7)))),
  file('thm.divides.PRIMES_|49||49|','thm.arithmetic.LESS_EQ')).
fof(f1052,plain,(
  ( ! [X76] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.divides.PRIMES_1'(s('type.num.num',sK1))))),s('type.num.num',X76)))) | s('type.num.num',X76) != s('type.num.num','const.divides.PRIMES_1'(s('type.num.num',sK2)))) )),
  inference(superposition,[],[f508,f342])).
fof(f508,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) != s('type.num.num',X1) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f323])).
fof(f323,plain,(
  ! [X0,X1] : ((s('type.num.num',X0) = s('type.num.num',X1) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0))))) & ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))) & ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0))))) | s('type.num.num',X0) != s('type.num.num',X1)))),
  inference(flattening,[],[f322])).
fof(f322,plain,(
  ! [X0,X1] : ((s('type.num.num',X0) = s('type.num.num',X1) | (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0)))))) & ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))) & ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0))))) | s('type.num.num',X0) != s('type.num.num',X1)))),
  inference(nnf_transformation,[],[f218])).
fof(f218,plain,(
  ! [X0,X1] : (s('type.num.num',X0) != s('type.num.num',X1) <=> (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0))))))),
  inference(flattening,[],[f217])).
fof(f217,plain,(
  ! [X0,X1] : (~s('type.num.num',X0) = s('type.num.num',X1) <=> (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0))))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X6,X7] : (~s('type.num.num',X6) = s('type.num.num',X7) <=> (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))),s('type.num.num',X7)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))),s('type.num.num',X6))))))),
  file('thm.divides.PRIMES_|49||49|','thm.arithmetic.NOT_NUM_EQ')).
fof(f5193,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.divides.PRIMES_1'(s('type.num.num',sK1))),s('type.num.num','const.divides.PRIMES_1'(s('type.num.num',sK1))))))),
  inference(resolution,[],[f5176,f933])).
fof(f933,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',sK2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.divides.PRIMES_1'(s('type.num.num',X0))),s('type.num.num','const.divides.PRIMES_1'(s('type.num.num',sK1))))))) )),
  inference(superposition,[],[f464,f342])).
fof(f464,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.divides.PRIMES_1'(s('type.num.num',X0))),s('type.num.num','const.divides.PRIMES_1'(s('type.num.num',X1)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f250])).
fof(f250,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.divides.PRIMES_1'(s('type.num.num',X0))),s('type.num.num','const.divides.PRIMES_1'(s('type.num.num',X1)))))))),
  inference(ennf_transformation,[],[f197])).
fof(f197,plain,(
  ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) => p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.divides.PRIMES_1'(s('type.num.num',X0))),s('type.num.num','const.divides.PRIMES_1'(s('type.num.num',X1)))))))),
  inference(rectify,[],[f51])).
fof(f51,axiom,(
  ! [X6,X7] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num',X7)))) => p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.divides.PRIMES_1'(s('type.num.num',X6))),s('type.num.num','const.divides.PRIMES_1'(s('type.num.num',X7)))))))),
  file('thm.divides.PRIMES_|49||49|','thm.divides.LT_PRIMES')).
fof(f5176,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)))) | p(s(bool,'F_0'))),
  inference(forward_demodulation,[],[f5102,f2181])).
fof(f5102,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.divides.PRIMES_1'(s('type.num.num',sK1))),s('type.num.num','const.divides.PRIMES_1'(s('type.num.num',sK1))))))),
  inference(resolution,[],[f3286,f934])).
fof(f934,plain,(
  ( ! [X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK2),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.divides.PRIMES_1'(s('type.num.num',sK1))),s('type.num.num','const.divides.PRIMES_1'(s('type.num.num',X1))))))) )),
  inference(superposition,[],[f464,f342])).
fof(f3286,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK2),s('type.num.num',sK1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2))))),
  inference(equality_resolution,[],[f912])).
fof(f912,plain,(
  ( ! [X4] : (s('type.num.num',sK1) != s('type.num.num',X4) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num',sK2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK2),s('type.num.num',X4))))) )),
  inference(forward_demodulation,[],[f911,f448])).
fof(f911,plain,(
  ( ! [X4] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num',sK2)))) | s('type.num.num',sK1) != s('type.num.num',X4) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',X4))))) )),
  inference(forward_demodulation,[],[f909,f448])).
fof(f909,plain,(
  ( ! [X4] : (s('type.num.num',sK1) != s('type.num.num',X4) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))),s('type.num.num',sK2)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',X4))))) )),
  inference(superposition,[],[f343,f507])).
fof(f507,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))) | s('type.num.num',X0) = s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f323])).
fof(f343,plain,(
  s('type.num.num',sK1) != s('type.num.num',sK2)),
  inference(cnf_transformation,[],[f274])).
% SZS output end Proof for thm.divides
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 168
% Passive clauses: 689
% Generated clauses: 4543
% Final active clauses: 165
% Final passive clauses: 649
% Input formulas: 108
% Initial clauses: 388
% 
% Duplicate literals: 354
% Trivial inequalities: 32
% Fw demodulations: 1295
% Bw demodulations: 3
% Condensations: 77
% Global subsumptions: 6
% 
% Simple tautologies: 393
% Equational tautologies: 6
% Forward subsumptions: 2137
% 
% Binary resolution: 908
% Factoring: 2
% Forward superposition: 1710
% Backward superposition: 391
% Equality resolution: 3
% 
% SAT solver clauses: 713
% SAT solver unit clauses: 23
% SAT solver binary clauses: 112
% SAT solver learnt clauses: 12
% SAT solver learnt literals: 28
% 
% TWLsolver clauses: 695
% TWLsolver calls for satisfiability: 13431
% 
% Memory used [KB]: 3709
% Time elapsed: 0.119 s
% ------------------------------
% ------------------------------
% Success in time 19.158 s
