Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.frac
% (76267)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 137
% Passive clauses: 2349
% Generated clauses: 18203
% Final active clauses: 112
% Final passive clauses: 1914
% Input formulas: 29
% Initial clauses: 77
% 
% Duplicate literals: 695
% Trivial inequalities: 294
% Fw subsumption resolutions: 2539
% Fw demodulations: 2636
% Bw demodulations: 22
% 
% Simple tautologies: 351
% Equational tautologies: 43
% Forward subsumptions: 8668
% Fw demodulations to eq. taut.: 5
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 466
% Factoring: 3
% Forward superposition: 5828
% Backward superposition: 5310
% Self superposition: 129
% Equality factoring: 86
% Equality resolution: 36
% 
% Split clauses: 306
% Split components: 632
% SAT solver clauses: 766
% SAT solver unit clauses: 16
% SAT solver binary clauses: 147
% SAT solver learnt clauses: 1
% SAT solver learnt literals: 1
% 
% Sat splits: 306
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 761
% TWLsolver calls for satisfiability: 179
% 
% Memory used [KB]: 9594
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.frac
% (76279)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 99
% Passive clauses: 529
% Generated clauses: 11284
% Final active clauses: 85
% Final passive clauses: 400
% Input formulas: 29
% Initial clauses: 77
% 
% Duplicate literals: 578
% Trivial inequalities: 107
% Fw subsumption resolutions: 896
% Fw demodulations: 1020
% Bw demodulations: 29
% Condensations: 1428
% Global subsumptions: 8
% 
% Simple tautologies: 184
% Equational tautologies: 38
% Forward subsumptions: 5467
% Fw demodulations to eq. taut.: 1
% 
% Binary resolution: 137
% Factoring: 1
% Forward superposition: 2940
% Backward superposition: 3341
% Self superposition: 115
% Equality factoring: 80
% Equality resolution: 31
% 
% Split clauses: 98
% Split components: 205
% SAT solver clauses: 1091
% SAT solver unit clauses: 82
% SAT solver binary clauses: 48
% SAT solver learnt clauses: 17
% SAT solver learnt literals: 1
% 
% Sat splits: 98
% Sat splitting refutations: 13
% 
% TWLsolver clauses: 1058
% TWLsolver calls for satisfiability: 360
% 
% Memory used [KB]: 8955
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.frac
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.frac
% SZS output start Proof for thm.frac
fof(f12982,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f120,f216,f214,f212,f119,f215,f118,f213,f211,f220,f884,f885,f886,f887,f981,f982,f983,f1270,f984,f1272,f2168,f2167,f2161,f2152,f2146,f1837,f1839,f2196,f2098,f2198,f2366,f2368,f2369,f2370,f2371,f2373,f2375,f2376,f2377,f2586,f2601,f2602,f2659,f2674,f2675,f2868,f1271,f2870,f1273,f2869,f4061,f4084,f4620,f4621,f4622,f4623,f4624,f4639,f4625,f4641,f4643,f4644,f4642,f4627,f4645,f4635,f4628,f4629,f4630,f4649,f4648,f4631,f4638,f4636,f4646,f4723,f4726,f4724,f4727,f4728,f4732,f4736,f4737,f4757,f4761,f4763,f4768,f4772,f4774,f4811,f4814,f4820,f4823,f4827,f5154,f5155,f5156,f5181,f5184,f6471,f6512,f6749,f6755,f7224,f7283,f7478,f7481,f7522,f7515,f7483,f7523,f7528,f7488,f7489,f7531,f7532,f7517,f7529,f9421,f9424,f9422,f9426,f9423,f9427,f9438,f9440,f9442,f9446,f9448,f9450,f9452,f9456,f9458,f9460,f9462,f9465,f9467,f9470,f9472,f9474,f9486,f9488,f9502,f9504,f9509,f9513,f9611,f9617,f9619,f9621,f9643,f9644,f9645,f9646,f10863,f10873,f11029,f11027,f4725,f11116,f11338,f11348,f11692,f11769,f11777,f11778,f11781,f11789,f11782,f2197,f12636,f12968])).
fof(f12968,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f12967,f12659])).
fof(f12659,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.integer.int',X0),'const.pair.,_2'(s('type.integer.int',sK0),s(X0,X1))) != s('type.pair.prod'('type.integer.int',X0),'const.pair.,_2'(s('type.integer.int',sK2),s(X0,X2)))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f214,f176])).
fof(f176,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) != s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))) | s(X0,X2) = s(X0,X4)) )),
  inference(cnf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) != s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))) | (s(X0,X2) = s(X0,X4) & s(X1,X3) = s(X1,X5))) & (s(X0,X2) != s(X0,X4) | s(X1,X3) != s(X1,X5) | s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5)))))),
  inference(flattening,[],[f110])).
fof(f110,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) != s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))) | (s(X0,X2) = s(X0,X4) & s(X1,X3) = s(X1,X5))) & ((s(X0,X2) != s(X0,X4) | s(X1,X3) != s(X1,X5)) | s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5)))))),
  inference(nnf_transformation,[],[f62])).
fof(f62,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))) <=> (s(X0,X2) = s(X0,X4) & s(X1,X3) = s(X1,X5)))),
  inference(flattening,[],[f61])).
fof(f61,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5] : (s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))) <=> (s(X0,X2) = s(X0,X4) & s(X1,X3) = s(X1,X5)))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X13,X21] : ! [X14,X15,X19,X22] : (s('type.pair.prod'(X13,X21),'const.pair.,_2'(s(X13,X14),s(X21,X15))) = s('type.pair.prod'(X13,X21),'const.pair.,_2'(s(X13,X19),s(X21,X22))) <=> (s(X13,X14) = s(X13,X19) & s(X21,X15) = s(X21,X22)))),
  file('thm.frac.FRAC_EQ','thm.pair.CLOSED_PAIR_EQ')).
fof(f12967,plain,(
  s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))) = s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | $spl0),
  inference(forward_demodulation,[],[f12833,f9425])).
fof(f9425,plain,(
  ( ! [X5] : (s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X5),s('type.integer.int',sK1))) = s('type.pair.prod'('type.integer.int','type.integer.int'),'const.frac.rep_frac_1'(s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X5),s('type.integer.int',sK1)))))))) )),
  inference(subsumption_resolution,[],[f9399,f121])).
fof(f121,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.frac.FRAC_EQ',aHL_TRUTH)).
fof(f9399,plain,(
  ( ! [X5] : (~p(s(bool,'T_0')) | s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X5),s('type.integer.int',sK1))) = s('type.pair.prod'('type.integer.int','type.integer.int'),'const.frac.rep_frac_1'(s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X5),s('type.integer.int',sK1)))))))) )),
  inference(superposition,[],[f584,f224])).
fof(f224,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))),
  inference(unit_resulting_resolution,[],[f116,f127])).
fof(f127,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f34])).
fof(f34,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X12] : (s(bool,'T_0') = s(bool,X12) <=> p(s(bool,X12)))),
  file('thm.frac.FRAC_EQ','thm.bool.EQ_CLAUSES|split|1')).
fof(f116,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))))),
  inference(cnf_transformation,[],[f81])).
fof(f81,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3)))) & (s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))))) = s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))))) | (s('type.integer.int',sK0) = s('type.integer.int',sK2) & s('type.integer.int',sK1) = s('type.integer.int',sK3))) & (s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))))) != s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))))) | s('type.integer.int',sK0) != s('type.integer.int',sK2) | s('type.integer.int',sK1) != s('type.integer.int',sK3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f80])).
fof(f80,plain,(
  ? [X0,X1,X2,X3] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X3)))) & (s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) = s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X2),s('type.integer.int',X3))))) | (s('type.integer.int',X0) = s('type.integer.int',X2) & s('type.integer.int',X1) = s('type.integer.int',X3))) & (s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) != s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X2),s('type.integer.int',X3))))) | s('type.integer.int',X0) != s('type.integer.int',X2) | s('type.integer.int',X1) != s('type.integer.int',X3)))),
  inference(flattening,[],[f79])).
fof(f79,plain,(
  ? [X0,X1,X2,X3] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X3)))) & ((s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) = s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X2),s('type.integer.int',X3))))) | (s('type.integer.int',X0) = s('type.integer.int',X2) & s('type.integer.int',X1) = s('type.integer.int',X3))) & (s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) != s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X2),s('type.integer.int',X3))))) | (s('type.integer.int',X0) != s('type.integer.int',X2) | s('type.integer.int',X1) != s('type.integer.int',X3)))))),
  inference(nnf_transformation,[],[f66])).
fof(f66,plain,(
  ? [X0,X1,X2,X3] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X3)))) & (s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) = s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X2),s('type.integer.int',X3))))) <~> (s('type.integer.int',X0) = s('type.integer.int',X2) & s('type.integer.int',X1) = s('type.integer.int',X3))))),
  inference(flattening,[],[f65])).
fof(f65,plain,(
  ? [X0,X1,X2,X3] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))) & (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X3)))) & (s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) = s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X2),s('type.integer.int',X3))))) <~> (s('type.integer.int',X0) = s('type.integer.int',X2) & s('type.integer.int',X1) = s('type.integer.int',X3)))))),
  inference(ennf_transformation,[],[f31])).
fof(f31,plain,(
  ~! [X0,X1,X2,X3] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))) => (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X3)))) => (s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) = s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X2),s('type.integer.int',X3))))) <=> (s('type.integer.int',X0) = s('type.integer.int',X2) & s('type.integer.int',X1) = s('type.integer.int',X3)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8,X9] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X7)))) => (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X9)))) => (s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X6),s('type.integer.int',X7))))) = s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X8),s('type.integer.int',X9))))) <=> (s('type.integer.int',X6) = s('type.integer.int',X8) & s('type.integer.int',X7) = s('type.integer.int',X9)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8,X9] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X7)))) => (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X9)))) => (s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X6),s('type.integer.int',X7))))) = s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X8),s('type.integer.int',X9))))) <=> (s('type.integer.int',X6) = s('type.integer.int',X8) & s('type.integer.int',X7) = s('type.integer.int',X9)))))),
  file('thm.frac.FRAC_EQ','thm.frac.FRAC_EQ')).
fof(f584,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))) | s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X0),s('type.integer.int',X1))) = s('type.pair.prod'('type.integer.int','type.integer.int'),'const.frac.rep_frac_1'(s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))))) )),
  inference(superposition,[],[f134,f174])).
fof(f174,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))) )),
  inference(cnf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0,X1,X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(flattening,[],[f57])).
fof(f57,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X13,X21] : ! [X14,X15] : s(X21,X15) = s(X21,'const.pair.SND_1'(s('type.pair.prod'(X13,X21),'const.pair.,_2'(s(X13,X14),s(X21,X15)))))),
  file('thm.frac.FRAC_EQ','thm.pair.SND')).
fof(f134,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.pair.SND_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),X0)))))) | s('type.pair.prod'('type.integer.int','type.integer.int'),X0) = s('type.pair.prod'('type.integer.int','type.integer.int'),'const.frac.rep_frac_1'(s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),X0)))))) )),
  inference(cnf_transformation,[],[f86])).
fof(f86,plain,(
  ! [X0] : ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.pair.SND_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),X0)))))) | s('type.pair.prod'('type.integer.int','type.integer.int'),X0) = s('type.pair.prod'('type.integer.int','type.integer.int'),'const.frac.rep_frac_1'(s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),X0)))))) & (s('type.pair.prod'('type.integer.int','type.integer.int'),X0) != s('type.pair.prod'('type.integer.int','type.integer.int'),'const.frac.rep_frac_1'(s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),X0))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.pair.SND_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),X0))))))))),
  inference(nnf_transformation,[],[f40])).
fof(f40,plain,(
  ! [X0] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.pair.SND_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),X0)))))) <=> s('type.pair.prod'('type.integer.int','type.integer.int'),X0) = s('type.pair.prod'('type.integer.int','type.integer.int'),'const.frac.rep_frac_1'(s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),X0))))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X20] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.pair.SND_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),X20)))))) <=> s('type.pair.prod'('type.integer.int','type.integer.int'),X20) = s('type.pair.prod'('type.integer.int','type.integer.int'),'const.frac.rep_frac_1'(s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),X20))))))),
  file('thm.frac.FRAC_EQ','thm.frac.frac_tybij|split|1')).
fof(f12833,plain,(
  s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) = s('type.pair.prod'('type.integer.int','type.integer.int'),'const.frac.rep_frac_1'(s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))))))) | $spl0),
  inference(superposition,[],[f9428,f211])).
fof(f9428,plain,(
  ( ! [X8] : (s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X8),s('type.integer.int',sK3))) = s('type.pair.prod'('type.integer.int','type.integer.int'),'const.frac.rep_frac_1'(s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X8),s('type.integer.int',sK3)))))))) )),
  inference(subsumption_resolution,[],[f9401,f121])).
fof(f9401,plain,(
  ( ! [X8] : (~p(s(bool,'T_0')) | s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X8),s('type.integer.int',sK3))) = s('type.pair.prod'('type.integer.int','type.integer.int'),'const.frac.rep_frac_1'(s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X8),s('type.integer.int',sK3)))))))) )),
  inference(superposition,[],[f584,f225])).
fof(f225,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3)))),
  inference(unit_resulting_resolution,[],[f117,f127])).
fof(f117,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))))),
  inference(cnf_transformation,[],[f81])).
fof(f12636,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f12635,f468])).
fof(f468,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(X0,'type.integer.int'),'const.pair.,_2'(s(X0,X1),s('type.integer.int',sK3))) != s('type.pair.prod'(X0,'type.integer.int'),'const.pair.,_2'(s(X0,X2),s('type.integer.int',sK1)))) ) | $spl5),
  inference(unit_resulting_resolution,[],[f216,f177])).
fof(f177,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) != s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))) | s(X1,X3) = s(X1,X5)) )),
  inference(cnf_transformation,[],[f111])).
fof(f12635,plain,(
  s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))) = s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',sK0),s('type.integer.int',sK3))) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f12481,f9425])).
fof(f12481,plain,(
  s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',sK0),s('type.integer.int',sK3))) = s('type.pair.prod'('type.integer.int','type.integer.int'),'const.frac.rep_frac_1'(s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))))))) | ($spl0 | $spl2)),
  inference(superposition,[],[f9428,f221])).
fof(f221,plain,(
  s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))))) = s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',sK0),s('type.integer.int',sK3))))) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f211,f213])).
fof(f2197,plain,(
  ( ! [X15] : (s(bool,'T_0') = s(bool,X15)) ) | $spl16),
  inference(cnf_transformation,[],[f2197_D])).
fof(f2197_D,plain,(
  ( ! [X15] : (s(bool,'T_0') = s(bool,X15)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f11782,plain,(
  ( ! [X61,X62,X60,X58,X63] : (s(X60,X61) != s(X60,X62) | s(bool,'F_0') = s(bool,X63) | s(bool,X58) = s(bool,X63) | s('type.pair.prod'(bool,X60),'const.pair.,_2'(s(bool,'F_0'),s(X60,X61))) = s('type.pair.prod'(bool,X60),'const.pair.,_2'(s(bool,X58),s(X60,X62)))) ) | $spl84),
  inference(cnf_transformation,[],[f11782_D])).
fof(f11782_D,plain,(
  ( ! [X61,X62,X60,X58,X63] : (s(X60,X61) != s(X60,X62) | s(bool,'F_0') = s(bool,X63) | s(bool,X58) = s(bool,X63) | s('type.pair.prod'(bool,X60),'const.pair.,_2'(s(bool,'F_0'),s(X60,X61))) = s('type.pair.prod'(bool,X60),'const.pair.,_2'(s(bool,X58),s(X60,X62)))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f11789,plain,(
  ( ! [X80,X78,X76,X81,X79,X77] : (s(X78,X79) != s(X78,X80) | s('type.pair.prod'(bool,X78),'const.pair.,_2'(s(bool,'F_0'),s(X78,X79))) = s('type.pair.prod'(bool,X78),'const.pair.,_2'(s(bool,X76),s(X78,X80))) | s(bool,'T_0') = s(bool,X77) | s(bool,X76) = s(bool,X81) | s(bool,'F_0') = s(bool,X81)) ) | $spl6),
  inference(subsumption_resolution,[],[f11663,f125])).
fof(f125,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.frac.FRAC_EQ',aHL_BOOL_CASES)).
fof(f11663,plain,(
  ( ! [X80,X78,X76,X81,X79,X77] : (s(bool,'F_0') != s(bool,X77) | s(X78,X79) != s(X78,X80) | s('type.pair.prod'(bool,X78),'const.pair.,_2'(s(bool,'F_0'),s(X78,X79))) = s('type.pair.prod'(bool,X78),'const.pair.,_2'(s(bool,X76),s(X78,X80))) | s(bool,'T_0') = s(bool,X77) | s(bool,X76) = s(bool,X81) | s(bool,'F_0') = s(bool,X81)) ) | $spl6),
  inference(superposition,[],[f1437,f2694])).
fof(f2694,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | $spl6),
  inference(superposition,[],[f2097,f2097])).
fof(f2097,plain,(
  ( ! [X99] : (s(bool,X99) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,'F_0') = s(bool,X99)) ) | $spl6),
  inference(trivial_inequality_removal,[],[f1851])).
fof(f1851,plain,(
  ( ! [X99] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X99) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,'F_0') = s(bool,X99)) ) | $spl6),
  inference(superposition,[],[f222,f1270])).
fof(f222,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f123,f126])).
fof(f126,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f82])).
fof(f123,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f32])).
fof(f32,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.frac.FRAC_EQ',aHL_FALSITY)).
fof(f1437,plain,(
  ( ! [X52,X50,X53,X51,X49] : (s(bool,X49) != s(bool,X50) | s(X51,X52) != s(X51,X53) | s('type.pair.prod'(bool,X51),'const.pair.,_2'(s(bool,'F_0'),s(X51,X52))) = s('type.pair.prod'(bool,X51),'const.pair.,_2'(s(bool,X50),s(X51,X53))) | s(bool,'T_0') = s(bool,X49)) )),
  inference(superposition,[],[f178,f125])).
fof(f178,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,X2) != s(X0,X4) | s(X1,X3) != s(X1,X5) | s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5)))) )),
  inference(cnf_transformation,[],[f111])).
fof(f11781,plain,(
  ( ! [X61,X59,X62,X60,X58,X63] : (s(X60,X61) != s(X60,X62) | s('type.pair.prod'(bool,X60),'const.pair.,_2'(s(bool,'F_0'),s(X60,X61))) = s('type.pair.prod'(bool,X60),'const.pair.,_2'(s(bool,X58),s(X60,X62))) | s(bool,'T_0') = s(bool,X59) | s(bool,X58) = s(bool,X63) | s(bool,'F_0') = s(bool,X63)) ) | $spl6),
  inference(subsumption_resolution,[],[f11660,f125])).
fof(f11660,plain,(
  ( ! [X61,X59,X62,X60,X58,X63] : (s(bool,'F_0') != s(bool,X59) | s(X60,X61) != s(X60,X62) | s('type.pair.prod'(bool,X60),'const.pair.,_2'(s(bool,'F_0'),s(X60,X61))) = s('type.pair.prod'(bool,X60),'const.pair.,_2'(s(bool,X58),s(X60,X62))) | s(bool,'T_0') = s(bool,X59) | s(bool,X58) = s(bool,X63) | s(bool,'F_0') = s(bool,X63)) ) | $spl6),
  inference(superposition,[],[f1437,f2694])).
fof(f11778,plain,(
  ( ! [X4,X2,X3,X1] : (s(X1,X2) != s(X1,X3) | s(bool,X4) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'F_0'),s(X1,X2))) = s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,X4),s(X1,X3)))) ) | $spl82),
  inference(cnf_transformation,[],[f11778_D])).
fof(f11778_D,plain,(
  ( ! [X4,X2,X3,X1] : (s(X1,X2) != s(X1,X3) | s(bool,X4) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'F_0'),s(X1,X2))) = s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,X4),s(X1,X3)))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f11777,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(X1,X2) != s(X1,X3) | s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'F_0'),s(X1,X2))) = s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,X4),s(X1,X3))) | s(bool,X4) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) ) | $spl6),
  inference(condensation,[],[f11776])).
fof(f11776,plain,(
  ( ! [X45,X43,X41,X46,X44,X42] : (s(X44,X45) != s(X44,X46) | s('type.pair.prod'(bool,X44),'const.pair.,_2'(s(bool,'F_0'),s(X44,X45))) = s('type.pair.prod'(bool,X44),'const.pair.,_2'(s(bool,X41),s(X44,X46))) | s(bool,'T_0') = s(bool,X43) | s(bool,X41) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X42)) ) | $spl6),
  inference(subsumption_resolution,[],[f11657,f247])).
fof(f247,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f125,f125])).
fof(f11657,plain,(
  ( ! [X45,X43,X41,X46,X44,X42] : (s(bool,X42) != s(bool,X43) | s(X44,X45) != s(X44,X46) | s('type.pair.prod'(bool,X44),'const.pair.,_2'(s(bool,'F_0'),s(X44,X45))) = s('type.pair.prod'(bool,X44),'const.pair.,_2'(s(bool,X41),s(X44,X46))) | s(bool,'T_0') = s(bool,X43) | s(bool,X41) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X42)) ) | $spl6),
  inference(superposition,[],[f1437,f1270])).
fof(f11769,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s(bool,X48) != s(bool,X49) | s(bool,X48) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s('type.pair.prod'(bool,X50),'const.pair.,_2'(s(bool,'F_0'),s(X50,X51))) = s('type.pair.prod'(bool,X50),'const.pair.,_2'(s(bool,X49),s(X50,X52))) | s(X50,X51) != s(X50,X52)) ) | $spl80),
  inference(cnf_transformation,[],[f11769_D])).
fof(f11769_D,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s(bool,X48) != s(bool,X49) | s(bool,X48) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s('type.pair.prod'(bool,X50),'const.pair.,_2'(s(bool,'F_0'),s(X50,X51))) = s('type.pair.prod'(bool,X50),'const.pair.,_2'(s(bool,X49),s(X50,X52))) | s(X50,X51) != s(X50,X52)) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f11692,plain,(
  ( ! [X47,X52,X50,X48,X51,X49] : (s(bool,X48) != s(bool,X49) | s(X50,X51) != s(X50,X52) | s('type.pair.prod'(bool,X50),'const.pair.,_2'(s(bool,'F_0'),s(X50,X51))) = s('type.pair.prod'(bool,X50),'const.pair.,_2'(s(bool,X49),s(X50,X52))) | s(bool,'T_0') = s(bool,X47) | s(bool,X48) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f11618])).
fof(f11618,plain,(
  ( ! [X47,X52,X50,X48,X51,X49] : (s(bool,X48) != s(bool,X49) | s(X50,X51) != s(X50,X52) | s('type.pair.prod'(bool,X50),'const.pair.,_2'(s(bool,'F_0'),s(X50,X51))) = s('type.pair.prod'(bool,X50),'const.pair.,_2'(s(bool,X49),s(X50,X52))) | s(bool,'T_0') = s(bool,X47) | s(bool,X48) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X47)) ) | $spl6),
  inference(superposition,[],[f1437,f1270])).
fof(f11348,plain,(
  ( ! [X16] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,X16))))) | s(bool,X16) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) ) | $spl78),
  inference(cnf_transformation,[],[f11348_D])).
fof(f11348_D,plain,(
  ( ! [X16] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,X16))))) | s(bool,X16) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f11338,plain,(
  ( ! [X15,X16] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,X16))))) | s(bool,'T_0') = s(bool,X15) | s(bool,X16) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) ) | ($spl6 | $spl20 | $spl44)),
  inference(duplicate_literal_removal,[],[f11132])).
fof(f11132,plain,(
  ( ! [X15,X16] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,X16))))) | s(bool,'T_0') = s(bool,X15) | s(bool,X16) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X15)) ) | ($spl6 | $spl20 | $spl44)),
  inference(superposition,[],[f6688,f1270])).
fof(f6688,plain,(
  ( ! [X102] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,X102))))) | s(bool,'T_0') = s(bool,X102)) ) | ($spl20 | $spl44)),
  inference(subsumption_resolution,[],[f6613,f123])).
fof(f6613,plain,(
  ( ! [X102] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,X102))))) | s(bool,'T_0') = s(bool,X102)) ) | ($spl20 | $spl44)),
  inference(superposition,[],[f5936,f5933])).
fof(f5933,plain,(
  ( ! [X28] : (s(bool,'F_0') = s(bool,f(s(bool,X28))) | s(bool,'T_0') = s(bool,X28)) ) | ($spl20 | $spl44)),
  inference(subsumption_resolution,[],[f5838,f369])).
fof(f369,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(subsumption_resolution,[],[f363,f121])).
fof(f363,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | s(bool,X1) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f251,f225])).
fof(f251,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f130,f125])).
fof(f130,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f37])).
fof(f37,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f36])).
fof(f36,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X12] : (s(bool,'F_0') = s(bool,X12) <=> ~p(s(bool,X12)))),
  file('thm.frac.FRAC_EQ','thm.bool.EQ_CLAUSES|split|2')).
fof(f5838,plain,(
  ( ! [X28] : (s(bool,X28) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) | s(bool,'T_0') = s(bool,X28) | s(bool,'F_0') = s(bool,f(s(bool,X28)))) ) | ($spl20 | $spl44)),
  inference(superposition,[],[f5290,f2869])).
fof(f5290,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,f(s(bool,X4))) | s(bool,'T_0') = s(bool,X4)) ) | $spl44),
  inference(superposition,[],[f4704,f125])).
fof(f4704,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | $spl44),
  inference(equality_resolution,[],[f4649])).
fof(f5936,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,'F_0') = s(bool,f(s(bool,X6)))) ) | $spl44),
  inference(subsumption_resolution,[],[f5846,f126])).
fof(f5846,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,'F_0') = s(bool,f(s(bool,X6))) | s(bool,'T_0') = s(bool,X6)) ) | $spl44),
  inference(superposition,[],[f131,f5290])).
fof(f131,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f84])).
fof(f11116,plain,(
  ( ! [X4,X2,X3,X1] : (s(X1,X2) != s(X1,X3) | s(bool,'T_0') = s(bool,X4) | s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'F_0'),s(X1,X3))) = s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,X4),s(X1,X2)))) ) | $spl76),
  inference(cnf_transformation,[],[f11116_D])).
fof(f11116_D,plain,(
  ( ! [X4,X2,X3,X1] : (s(X1,X2) != s(X1,X3) | s(bool,'T_0') = s(bool,X4) | s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'F_0'),s(X1,X3))) = s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,X4),s(X1,X2)))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f4725,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,X1)) ) | $spl48),
  inference(cnf_transformation,[],[f4725_D])).
fof(f4725_D,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,X1)) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f11027,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,X0) | s(X1,X2) != s(X1,X3) | s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'F_0'),s(X1,X3))) = s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,X4),s(X1,X2))) | s(bool,'T_0') = s(bool,X4)) ) | $spl6),
  inference(condensation,[],[f11026])).
fof(f11026,plain,(
  ( ! [X111,X116,X114,X112,X115,X113] : (s(X112,X113) != s(X112,X114) | s('type.pair.prod'(bool,X112),'const.pair.,_2'(s(bool,'F_0'),s(X112,X114))) = s('type.pair.prod'(bool,X112),'const.pair.,_2'(s(bool,X115),s(X112,X113))) | s(bool,'T_0') = s(bool,X115) | s(bool,X111) = s(bool,X116) | s(bool,'F_0') = s(bool,X116)) ) | $spl6),
  inference(trivial_inequality_removal,[],[f11004])).
fof(f11004,plain,(
  ( ! [X111,X116,X114,X112,X115,X113] : (s(bool,X111) != s(bool,X111) | s(X112,X113) != s(X112,X114) | s('type.pair.prod'(bool,X112),'const.pair.,_2'(s(bool,'F_0'),s(X112,X114))) = s('type.pair.prod'(bool,X112),'const.pair.,_2'(s(bool,X115),s(X112,X113))) | s(bool,'T_0') = s(bool,X115) | s(bool,X111) = s(bool,X116) | s(bool,'F_0') = s(bool,X116)) ) | $spl6),
  inference(superposition,[],[f1434,f2694])).
fof(f1434,plain,(
  ( ! [X35,X33,X36,X34,X32] : (s(bool,'F_0') != s(bool,X33) | s(X34,X35) != s(X34,X36) | s('type.pair.prod'(bool,X34),'const.pair.,_2'(s(bool,X32),s(X34,X35))) = s('type.pair.prod'(bool,X34),'const.pair.,_2'(s(bool,X33),s(X34,X36))) | s(bool,'T_0') = s(bool,X32)) )),
  inference(superposition,[],[f178,f125])).
fof(f11029,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,X0) | s(X1,X2) != s(X1,X3) | s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'F_0'),s(X1,X3))) = s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,X4),s(X1,X2))) | s(bool,'T_0') = s(bool,X4)) ) | $spl6),
  inference(condensation,[],[f11028])).
fof(f11028,plain,(
  ( ! [X109,X107,X105,X110,X108,X106] : (s(X106,X107) != s(X106,X108) | s('type.pair.prod'(bool,X106),'const.pair.,_2'(s(bool,'F_0'),s(X106,X108))) = s('type.pair.prod'(bool,X106),'const.pair.,_2'(s(bool,X109),s(X106,X107))) | s(bool,'T_0') = s(bool,X109) | s(bool,X105) = s(bool,X110) | s(bool,'F_0') = s(bool,X110)) ) | $spl6),
  inference(trivial_inequality_removal,[],[f11003])).
fof(f11003,plain,(
  ( ! [X109,X107,X105,X110,X108,X106] : (s(bool,X105) != s(bool,X105) | s(X106,X107) != s(X106,X108) | s('type.pair.prod'(bool,X106),'const.pair.,_2'(s(bool,'F_0'),s(X106,X108))) = s('type.pair.prod'(bool,X106),'const.pair.,_2'(s(bool,X109),s(X106,X107))) | s(bool,'T_0') = s(bool,X109) | s(bool,X105) = s(bool,X110) | s(bool,'F_0') = s(bool,X110)) ) | $spl6),
  inference(superposition,[],[f1434,f2694])).
fof(f10873,plain,(
  ( ! [X16] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,X16))))) | s(bool,X16) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) ) | $spl74),
  inference(cnf_transformation,[],[f10873_D])).
fof(f10873_D,plain,(
  ( ! [X16] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,X16))))) | s(bool,X16) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f10863,plain,(
  ( ! [X15,X16] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,X16))))) | s(bool,'T_0') = s(bool,X15) | s(bool,X16) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) ) | ($spl6 | $spl20 | $spl44)),
  inference(duplicate_literal_removal,[],[f10782])).
fof(f10782,plain,(
  ( ! [X15,X16] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,X16))))) | s(bool,'T_0') = s(bool,X15) | s(bool,X16) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X15)) ) | ($spl6 | $spl20 | $spl44)),
  inference(superposition,[],[f6685,f1270])).
fof(f6685,plain,(
  ( ! [X100] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,X100))))) | s(bool,'T_0') = s(bool,X100)) ) | ($spl20 | $spl44)),
  inference(subsumption_resolution,[],[f6611,f123])).
fof(f6611,plain,(
  ( ! [X100] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,X100))))) | s(bool,'T_0') = s(bool,X100)) ) | ($spl20 | $spl44)),
  inference(superposition,[],[f5934,f5933])).
fof(f5934,plain,(
  ( ! [X3] : (p(s(bool,X3)) | s(bool,'T_0') != s(bool,f(s(bool,X3)))) ) | $spl44),
  inference(subsumption_resolution,[],[f5843,f126])).
fof(f5843,plain,(
  ( ! [X3] : (p(s(bool,X3)) | s(bool,'T_0') != s(bool,f(s(bool,X3))) | s(bool,'T_0') = s(bool,X3)) ) | $spl44),
  inference(superposition,[],[f126,f5290])).
fof(f9646,plain,(
  $false | $spl70),
  inference(subsumption_resolution,[],[f9449,f9423])).
fof(f9449,plain,(
  p(s(bool,'F_0')) | $spl70),
  inference(unit_resulting_resolution,[],[f222,f9423,f170])).
fof(f170,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f72])).
fof(f72,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f56])).
fof(f56,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X10,X11] : ((p(s(bool,X10)) => p(s(bool,X11))) => ((p(s(bool,X11)) => p(s(bool,X10))) => s(bool,X10) = s(bool,X11)))),
  file('thm.frac.FRAC_EQ','thm.bool.IMP_ANTISYM_AX')).
fof(f9645,plain,(
  $false | ($spl33 | $spl70)),
  inference(subsumption_resolution,[],[f9451,f9423])).
fof(f9451,plain,(
  p(s(bool,f(s(bool,'F_0')))) | ($spl33 | $spl70)),
  inference(unit_resulting_resolution,[],[f4641,f9423,f170])).
fof(f9644,plain,(
  $false | $spl70),
  inference(subsumption_resolution,[],[f9459,f9423])).
fof(f9459,plain,(
  p(s(bool,'F_0')) | $spl70),
  inference(unit_resulting_resolution,[],[f222,f9423,f170])).
fof(f9643,plain,(
  $false | ($spl33 | $spl70)),
  inference(subsumption_resolution,[],[f9461,f9423])).
fof(f9461,plain,(
  p(s(bool,f(s(bool,'F_0')))) | ($spl33 | $spl70)),
  inference(unit_resulting_resolution,[],[f4641,f9423,f170])).
fof(f9621,plain,(
  $false | ($spl6 | $spl70)),
  inference(subsumption_resolution,[],[f9620,f222])).
fof(f9620,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl6 | $spl70)),
  inference(forward_demodulation,[],[f9495,f224])).
fof(f9495,plain,(
  s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | ($spl6 | $spl70)),
  inference(unit_resulting_resolution,[],[f224,f9423,f2854])).
fof(f2854,plain,(
  ( ! [X12,X11] : (p(s(bool,X12)) | s(bool,X11) != s(bool,X12) | s(bool,'F_0') = s(bool,X11)) ) | $spl6),
  inference(subsumption_resolution,[],[f2853,f121])).
fof(f2853,plain,(
  ( ! [X12,X11] : (~p(s(bool,'T_0')) | s(bool,X11) != s(bool,X12) | p(s(bool,X12)) | s(bool,'F_0') = s(bool,X11)) ) | $spl6),
  inference(forward_demodulation,[],[f2715,f224])).
fof(f2715,plain,(
  ( ! [X12,X11] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | s(bool,X11) != s(bool,X12) | p(s(bool,X12)) | s(bool,'F_0') = s(bool,X11)) ) | $spl6),
  inference(superposition,[],[f204,f2097])).
fof(f204,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,X0) != s(bool,X1) | p(s(bool,X1))) )),
  inference(condensation,[],[f203])).
fof(f203,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2)) )),
  inference(condensation,[],[f192])).
fof(f192,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f115])).
fof(f115,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f114])).
fof(f114,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f77])).
fof(f77,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f64])).
fof(f64,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X14,X17,X15,X18] : ((s(bool,X14) = s(bool,X17) & (p(s(bool,X17)) => s(bool,X15) = s(bool,X18))) => ((p(s(bool,X14)) => p(s(bool,X15))) <=> (p(s(bool,X17)) => p(s(bool,X18)))))),
  file('thm.frac.FRAC_EQ','thm.bool.IMP_CONG')).
fof(f9619,plain,(
  $false | ($spl6 | $spl70)),
  inference(subsumption_resolution,[],[f9618,f222])).
fof(f9618,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl6 | $spl70)),
  inference(forward_demodulation,[],[f9496,f225])).
fof(f9496,plain,(
  s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) | ($spl6 | $spl70)),
  inference(unit_resulting_resolution,[],[f225,f9423,f2854])).
fof(f9617,plain,(
  $false | ($spl6 | $spl24 | $spl70)),
  inference(subsumption_resolution,[],[f9616,f222])).
fof(f9616,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl6 | $spl24 | $spl70)),
  inference(forward_demodulation,[],[f9499,f4635])).
fof(f9499,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | ($spl6 | $spl24 | $spl70)),
  inference(unit_resulting_resolution,[],[f4635,f9423,f2854])).
fof(f9611,plain,(
  $false | ($spl33 | $spl44 | $spl70)),
  inference(subsumption_resolution,[],[f9512,f9423])).
fof(f9512,plain,(
  p(s(bool,'F_0')) | ($spl33 | $spl44 | $spl70)),
  inference(unit_resulting_resolution,[],[f4641,f9423,f5937])).
fof(f5937,plain,(
  ( ! [X8,X7] : (p(s(bool,X7)) | p(s(bool,X8)) | s(bool,X8) = s(bool,f(s(bool,X7)))) ) | $spl44),
  inference(subsumption_resolution,[],[f5847,f126])).
fof(f5847,plain,(
  ( ! [X8,X7] : (p(s(bool,X7)) | p(s(bool,X8)) | s(bool,X8) = s(bool,f(s(bool,X7))) | s(bool,'T_0') = s(bool,X7)) ) | $spl44),
  inference(superposition,[],[f170,f5290])).
fof(f9513,plain,(
  $false | ($spl33 | $spl44 | $spl70)),
  inference(unit_resulting_resolution,[],[f123,f4641,f9423,f5937])).
fof(f9509,plain,(
  $false | ($spl33 | $spl44 | $spl70)),
  inference(unit_resulting_resolution,[],[f9423,f4641,f9423,f5937])).
fof(f9504,plain,(
  $false | ($spl6 | $spl44 | $spl61 | $spl70)),
  inference(unit_resulting_resolution,[],[f8684,f9423,f5934])).
fof(f8684,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | ($spl6 | $spl61)),
  inference(forward_demodulation,[],[f8647,f224])).
fof(f8647,plain,(
  s(bool,f(s(bool,'T_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | ($spl6 | $spl61)),
  inference(unit_resulting_resolution,[],[f222,f7517,f1270])).
fof(f9502,plain,(
  $false | ($spl6 | $spl70)),
  inference(unit_resulting_resolution,[],[f224,f222,f9423,f2854])).
fof(f9488,plain,(
  $false | ($spl33 | $spl70)),
  inference(unit_resulting_resolution,[],[f4641,f4641,f9423,f286])).
fof(f286,plain,(
  ( ! [X4,X5] : (p(s(bool,X5)) | s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f281,f123])).
fof(f281,plain,(
  ( ! [X4,X5] : (p(s(bool,'F_0')) | p(s(bool,X5)) | s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f170,f125])).
fof(f9486,plain,(
  $false | $spl70),
  inference(unit_resulting_resolution,[],[f222,f222,f9423,f286])).
fof(f9474,plain,(
  $false | $spl70),
  inference(unit_resulting_resolution,[],[f121,f174,f9423,f204])).
fof(f9472,plain,(
  $false | ($spl24 | $spl70)),
  inference(unit_resulting_resolution,[],[f121,f4635,f9423,f204])).
fof(f9470,plain,(
  $false | ($spl6 | $spl61 | $spl70)),
  inference(unit_resulting_resolution,[],[f121,f8684,f9423,f204])).
fof(f9467,plain,(
  $false | $spl70),
  inference(unit_resulting_resolution,[],[f121,f225,f9423,f204])).
fof(f9465,plain,(
  $false | $spl70),
  inference(unit_resulting_resolution,[],[f121,f224,f9423,f204])).
fof(f9462,plain,(
  $false | ($spl33 | $spl70)),
  inference(unit_resulting_resolution,[],[f9423,f4641,f9423,f170])).
fof(f9460,plain,(
  $false | $spl70),
  inference(unit_resulting_resolution,[],[f123,f222,f9423,f170])).
fof(f9458,plain,(
  $false | ($spl33 | $spl70)),
  inference(unit_resulting_resolution,[],[f9423,f4641,f9423,f170])).
fof(f9456,plain,(
  $false | $spl70),
  inference(unit_resulting_resolution,[],[f9423,f222,f9423,f170])).
fof(f9452,plain,(
  $false | ($spl33 | $spl70)),
  inference(unit_resulting_resolution,[],[f9423,f4641,f9423,f170])).
fof(f9450,plain,(
  $false | $spl70),
  inference(unit_resulting_resolution,[],[f123,f222,f9423,f170])).
fof(f9448,plain,(
  $false | ($spl33 | $spl70)),
  inference(unit_resulting_resolution,[],[f9423,f4641,f9423,f170])).
fof(f9446,plain,(
  $false | $spl70),
  inference(unit_resulting_resolution,[],[f9423,f222,f9423,f170])).
fof(f9442,plain,(
  $false | $spl70),
  inference(unit_resulting_resolution,[],[f222,f9423,f131])).
fof(f9440,plain,(
  $false | $spl70),
  inference(unit_resulting_resolution,[],[f224,f9423,f126])).
fof(f9438,plain,(
  $false | $spl70),
  inference(unit_resulting_resolution,[],[f121,f9423])).
fof(f9427,plain,(
  ( ! [X7] : (s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X7),s('type.integer.int',sK3))) = s('type.pair.prod'('type.integer.int','type.integer.int'),'const.frac.rep_frac_1'(s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X7),s('type.integer.int',sK3)))))))) ) | $spl72),
  inference(cnf_transformation,[],[f9427_D])).
fof(f9427_D,plain,(
  ( ! [X7] : (s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X7),s('type.integer.int',sK3))) = s('type.pair.prod'('type.integer.int','type.integer.int'),'const.frac.rep_frac_1'(s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X7),s('type.integer.int',sK3)))))))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f9423,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl70),
  inference(cnf_transformation,[],[f9423_D])).
fof(f9423_D,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f9426,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X7),s('type.integer.int',sK3))) = s('type.pair.prod'('type.integer.int','type.integer.int'),'const.frac.rep_frac_1'(s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X7),s('type.integer.int',sK3)))))))) ) | $spl20),
  inference(subsumption_resolution,[],[f9400,f130])).
fof(f9400,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X7),s('type.integer.int',sK3))) = s('type.pair.prod'('type.integer.int','type.integer.int'),'const.frac.rep_frac_1'(s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X7),s('type.integer.int',sK3))))))) | s(bool,'F_0') = s(bool,X6)) ) | $spl20),
  inference(superposition,[],[f584,f2869])).
fof(f9422,plain,(
  ( ! [X1] : (s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X1),s('type.integer.int',sK1))) = s('type.pair.prod'('type.integer.int','type.integer.int'),'const.frac.rep_frac_1'(s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X1),s('type.integer.int',sK1)))))))) ) | $spl68),
  inference(cnf_transformation,[],[f9422_D])).
fof(f9422_D,plain,(
  ( ! [X1] : (s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X1),s('type.integer.int',sK1))) = s('type.pair.prod'('type.integer.int','type.integer.int'),'const.frac.rep_frac_1'(s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X1),s('type.integer.int',sK1)))))))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f9424,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,X2)) | s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X3),s('type.integer.int',sK1))) = s('type.pair.prod'('type.integer.int','type.integer.int'),'const.frac.rep_frac_1'(s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X3),s('type.integer.int',sK1))))))) | s(bool,'T_0') = s(bool,X4)) ) | $spl6),
  inference(subsumption_resolution,[],[f9398,f251])).
fof(f9398,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,X2)) | s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X3),s('type.integer.int',sK1))) = s('type.pair.prod'('type.integer.int','type.integer.int'),'const.frac.rep_frac_1'(s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X3),s('type.integer.int',sK1))))))) | s(bool,'T_0') = s(bool,X4) | s(bool,X2) = s(bool,X4)) ) | $spl6),
  inference(superposition,[],[f584,f1270])).
fof(f9421,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X1),s('type.integer.int',sK1))) = s('type.pair.prod'('type.integer.int','type.integer.int'),'const.frac.rep_frac_1'(s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X1),s('type.integer.int',sK1)))))))) ) | $spl6),
  inference(subsumption_resolution,[],[f9397,f130])).
fof(f9397,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X1),s('type.integer.int',sK1))) = s('type.pair.prod'('type.integer.int','type.integer.int'),'const.frac.rep_frac_1'(s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X1),s('type.integer.int',sK1))))))) | s(bool,'F_0') = s(bool,X0)) ) | $spl6),
  inference(superposition,[],[f584,f2097])).
fof(f7529,plain,(
  ( ! [X26,X27] : (s(fun(bool,bool),X26) = s(fun(bool,bool),X27) | s(bool,'F_0') = s(bool,sK4(X26,X27,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK4(X26,X27,bool,bool))))) ) | $spl64),
  inference(cnf_transformation,[],[f7529_D])).
fof(f7529_D,plain,(
  ( ! [X26,X27] : (s(fun(bool,bool),X26) = s(fun(bool,bool),X27) | s(bool,'F_0') = s(bool,sK4(X26,X27,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK4(X26,X27,bool,bool))))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f7517,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | $spl61),
  inference(cnf_transformation,[],[f7517_D])).
fof(f7517_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl61),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl61])])).
fof(f7532,plain,(
  ( ! [X41,X40] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X40) = s(fun(bool,bool),X41) | s(bool,'T_0') = s(bool,g(s(bool,sK4(X40,X41,bool,bool)))) | s(bool,'F_0') = s(bool,sK4(X40,X41,bool,bool))) ) | $spl20),
  inference(forward_demodulation,[],[f7491,f225])).
fof(f7491,plain,(
  ( ! [X41,X40] : (s(bool,'F_0') != s(bool,f(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))))) | s(fun(bool,bool),X40) = s(fun(bool,bool),X41) | s(bool,'T_0') = s(bool,g(s(bool,sK4(X40,X41,bool,bool)))) | s(bool,'F_0') = s(bool,sK4(X40,X41,bool,bool))) ) | $spl20),
  inference(superposition,[],[f548,f2869])).
fof(f548,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK4(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK4(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f175,f125])).
fof(f175,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK4(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK4(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK4(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK4(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f74])).
fof(f74,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f59])).
fof(f59,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.frac.FRAC_EQ',aHL_EXT)).
fof(f7531,plain,(
  ( ! [X33] : (s(bool,'F_0') != s(bool,f(s(bool,X33))) | s(bool,'F_0') = s(bool,X33)) ) | $spl66),
  inference(cnf_transformation,[],[f7531_D])).
fof(f7531_D,plain,(
  ( ! [X33] : (s(bool,'F_0') != s(bool,f(s(bool,X33))) | s(bool,'F_0') = s(bool,X33)) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f7489,plain,(
  ( ! [X35,X36,X34] : (s(bool,'F_0') != s(bool,f(s(bool,X36))) | s(fun(bool,bool),X34) = s(fun(bool,bool),X35) | s(bool,'T_0') = s(bool,g(s(bool,sK4(X34,X35,bool,bool)))) | s(bool,'F_0') = s(bool,sK4(X34,X35,bool,bool)) | s(bool,'F_0') = s(bool,X36)) ) | $spl6),
  inference(superposition,[],[f548,f2694])).
fof(f7488,plain,(
  ( ! [X33,X31,X32] : (s(bool,'F_0') != s(bool,f(s(bool,X33))) | s(fun(bool,bool),X31) = s(fun(bool,bool),X32) | s(bool,'T_0') = s(bool,g(s(bool,sK4(X31,X32,bool,bool)))) | s(bool,'F_0') = s(bool,X33) | s(bool,'F_0') = s(bool,sK4(X31,X32,bool,bool))) ) | $spl6),
  inference(superposition,[],[f548,f2694])).
fof(f7528,plain,(
  ( ! [X26,X27] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X26) = s(fun(bool,bool),X27) | s(bool,'T_0') = s(bool,g(s(bool,sK4(X26,X27,bool,bool)))) | s(bool,'F_0') = s(bool,sK4(X26,X27,bool,bool))) ) | $spl6),
  inference(forward_demodulation,[],[f7486,f224])).
fof(f7486,plain,(
  ( ! [X26,X27] : (s(bool,'F_0') != s(bool,f(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))))) | s(fun(bool,bool),X26) = s(fun(bool,bool),X27) | s(bool,'T_0') = s(bool,g(s(bool,sK4(X26,X27,bool,bool)))) | s(bool,'F_0') = s(bool,sK4(X26,X27,bool,bool))) ) | $spl6),
  inference(superposition,[],[f548,f2097])).
fof(f7523,plain,(
  ( ! [X19,X17,X18] : (s(fun(bool,bool),X17) = s(fun(bool,bool),X18) | s(bool,X19) = s(bool,sK4(X17,X18,bool,bool)) | s(bool,X19) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,g(s(bool,sK4(X17,X18,bool,bool))))) ) | $spl62),
  inference(cnf_transformation,[],[f7523_D])).
fof(f7523_D,plain,(
  ( ! [X19,X17,X18] : (s(fun(bool,bool),X17) = s(fun(bool,bool),X18) | s(bool,X19) = s(bool,sK4(X17,X18,bool,bool)) | s(bool,X19) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,g(s(bool,sK4(X17,X18,bool,bool))))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f7483,plain,(
  ( ! [X19,X17,X18] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X17) = s(fun(bool,bool),X18) | s(bool,'T_0') = s(bool,g(s(bool,sK4(X17,X18,bool,bool)))) | s(bool,X19) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,X19) = s(bool,sK4(X17,X18,bool,bool))) ) | $spl6),
  inference(superposition,[],[f548,f1270])).
fof(f7515,plain,(
  ( ! [X4,X2,X3] : (s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,X4) | s(bool,X4) = s(bool,sK4(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK4(X2,X3,bool,bool))))) ) | $spl58),
  inference(cnf_transformation,[],[f7515_D])).
fof(f7515_D,plain,(
  ( ! [X4,X2,X3] : (s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,X4) | s(bool,X4) = s(bool,sK4(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK4(X2,X3,bool,bool))))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f7522,plain,(
  ( ! [X14,X15,X16] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X14) = s(fun(bool,bool),X15) | s(bool,'T_0') = s(bool,g(s(bool,sK4(X14,X15,bool,bool)))) | s(bool,'T_0') = s(bool,X16) | s(bool,X16) = s(bool,sK4(X14,X15,bool,bool))) ) | $spl6),
  inference(forward_demodulation,[],[f7482,f224])).
fof(f7482,plain,(
  ( ! [X14,X15,X16] : (s(bool,'F_0') != s(bool,f(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))))) | s(fun(bool,bool),X14) = s(fun(bool,bool),X15) | s(bool,'T_0') = s(bool,g(s(bool,sK4(X14,X15,bool,bool)))) | s(bool,'T_0') = s(bool,X16) | s(bool,X16) = s(bool,sK4(X14,X15,bool,bool))) ) | $spl6),
  inference(superposition,[],[f548,f1270])).
fof(f7481,plain,(
  ( ! [X12,X13,X11] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X11) = s(fun(bool,bool),X12) | s(bool,'T_0') = s(bool,g(s(bool,sK4(X11,X12,bool,bool)))) | s(bool,X13) = s(bool,sK4(X11,X12,bool,bool)) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f548,f247])).
fof(f7478,plain,(
  ( ! [X4,X2,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,g(s(bool,sK4(X2,X3,bool,bool)))) | s(bool,X4) = s(bool,sK4(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f548,f247])).
fof(f7283,plain,(
  ( ! [X94,X92,X97,X95,X93] : (s('type.pair.prod'(X93,bool),'const.pair.,_2'(s(X93,X94),s(bool,'F_0'))) != s('type.pair.prod'(X93,bool),'const.pair.,_2'(s(X93,X95),s(bool,X92))) | s(bool,X92) = s(bool,X97) | s(bool,X97) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) ) | $spl56),
  inference(cnf_transformation,[],[f7283_D])).
fof(f7283_D,plain,(
  ( ! [X94,X92,X97,X95,X93] : (s('type.pair.prod'(X93,bool),'const.pair.,_2'(s(X93,X94),s(bool,'F_0'))) != s('type.pair.prod'(X93,bool),'const.pair.,_2'(s(X93,X95),s(bool,X92))) | s(bool,X92) = s(bool,X97) | s(bool,X97) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f7224,plain,(
  ( ! [X94,X92,X97,X95,X93,X96] : (s('type.pair.prod'(X93,bool),'const.pair.,_2'(s(X93,X94),s(bool,'F_0'))) != s('type.pair.prod'(X93,bool),'const.pair.,_2'(s(X93,X95),s(bool,X92))) | s(bool,'T_0') = s(bool,X96) | s(bool,X97) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,X92) = s(bool,X97)) ) | $spl6),
  inference(duplicate_literal_removal,[],[f7202])).
fof(f7202,plain,(
  ( ! [X94,X92,X97,X95,X93,X96] : (s('type.pair.prod'(X93,bool),'const.pair.,_2'(s(X93,X94),s(bool,'F_0'))) != s('type.pair.prod'(X93,bool),'const.pair.,_2'(s(X93,X95),s(bool,X92))) | s(bool,'T_0') = s(bool,X96) | s(bool,'T_0') = s(bool,X96) | s(bool,X97) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,X92) = s(bool,X97)) ) | $spl6),
  inference(superposition,[],[f483,f1270])).
fof(f483,plain,(
  ( ! [X24,X23,X21,X25,X22] : (s('type.pair.prod'(X22,bool),'const.pair.,_2'(s(X22,X23),s(bool,'F_0'))) != s('type.pair.prod'(X22,bool),'const.pair.,_2'(s(X22,X24),s(bool,X25))) | s(bool,X21) = s(bool,X25) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f177,f125])).
fof(f6755,plain,(
  ( ! [X16] : (s(bool,'T_0') != s(bool,f(s(bool,X16))) | s(bool,X16) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) ) | $spl54),
  inference(cnf_transformation,[],[f6755_D])).
fof(f6755_D,plain,(
  ( ! [X16] : (s(bool,'T_0') != s(bool,f(s(bool,X16))) | s(bool,X16) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f6749,plain,(
  ( ! [X15,X16] : (s(bool,'T_0') != s(bool,f(s(bool,X16))) | s(bool,'T_0') = s(bool,X15) | s(bool,X16) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) ) | ($spl6 | $spl44)),
  inference(duplicate_literal_removal,[],[f6698])).
fof(f6698,plain,(
  ( ! [X15,X16] : (s(bool,'T_0') != s(bool,f(s(bool,X16))) | s(bool,'T_0') = s(bool,X15) | s(bool,X16) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X15)) ) | ($spl6 | $spl44)),
  inference(superposition,[],[f5975,f1270])).
fof(f5975,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl44),
  inference(subsumption_resolution,[],[f5943,f123])).
fof(f5943,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl44),
  inference(superposition,[],[f5934,f125])).
fof(f6512,plain,(
  ( ! [X94,X92,X97,X95,X93] : (s('type.pair.prod'(bool,X93),'const.pair.,_2'(s(bool,'F_0'),s(X93,X94))) != s('type.pair.prod'(bool,X93),'const.pair.,_2'(s(bool,X92),s(X93,X95))) | s(bool,X92) = s(bool,X97) | s(bool,X97) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) ) | $spl52),
  inference(cnf_transformation,[],[f6512_D])).
fof(f6512_D,plain,(
  ( ! [X94,X92,X97,X95,X93] : (s('type.pair.prod'(bool,X93),'const.pair.,_2'(s(bool,'F_0'),s(X93,X94))) != s('type.pair.prod'(bool,X93),'const.pair.,_2'(s(bool,X92),s(X93,X95))) | s(bool,X92) = s(bool,X97) | s(bool,X97) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f6471,plain,(
  ( ! [X94,X92,X97,X95,X93,X96] : (s('type.pair.prod'(bool,X93),'const.pair.,_2'(s(bool,'F_0'),s(X93,X94))) != s('type.pair.prod'(bool,X93),'const.pair.,_2'(s(bool,X92),s(X93,X95))) | s(bool,'T_0') = s(bool,X96) | s(bool,X97) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,X92) = s(bool,X97)) ) | $spl6),
  inference(duplicate_literal_removal,[],[f6399])).
fof(f6399,plain,(
  ( ! [X94,X92,X97,X95,X93,X96] : (s('type.pair.prod'(bool,X93),'const.pair.,_2'(s(bool,'F_0'),s(X93,X94))) != s('type.pair.prod'(bool,X93),'const.pair.,_2'(s(bool,X92),s(X93,X95))) | s(bool,'T_0') = s(bool,X96) | s(bool,'T_0') = s(bool,X96) | s(bool,X97) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,X92) = s(bool,X97)) ) | $spl6),
  inference(superposition,[],[f397,f1270])).
fof(f397,plain,(
  ( ! [X24,X23,X21,X25,X22] : (s('type.pair.prod'(bool,X22),'const.pair.,_2'(s(bool,'F_0'),s(X22,X23))) != s('type.pair.prod'(bool,X22),'const.pair.,_2'(s(bool,X24),s(X22,X25))) | s(bool,X21) = s(bool,X24) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f176,f125])).
fof(f5184,plain,(
  $false | $spl48),
  inference(subsumption_resolution,[],[f5183,f4861])).
fof(f4861,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1)) ) | $spl48),
  inference(superposition,[],[f4725,f4725])).
fof(f5183,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,X1)) ) | $spl48),
  inference(subsumption_resolution,[],[f4965,f4826])).
fof(f4826,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl48),
  inference(unit_resulting_resolution,[],[f123,f4725,f204])).
fof(f4965,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,X1)) ) | $spl48),
  inference(superposition,[],[f126,f4725])).
fof(f5181,plain,(
  $false | $spl48),
  inference(subsumption_resolution,[],[f4900,f4861])).
fof(f4900,plain,(
  ( ! [X5] : (s(bool,'T_0') != s(bool,X5)) ) | $spl48),
  inference(superposition,[],[f222,f4725])).
fof(f5156,plain,(
  $false | $spl48),
  inference(subsumption_resolution,[],[f4813,f4826])).
fof(f4813,plain,(
  p(s(bool,'F_0')) | $spl48),
  inference(unit_resulting_resolution,[],[f4725,f126])).
fof(f5155,plain,(
  $false | ($spl6 | $spl48)),
  inference(subsumption_resolution,[],[f4822,f4826])).
fof(f4822,plain,(
  p(s(bool,'F_0')) | ($spl6 | $spl48)),
  inference(unit_resulting_resolution,[],[f222,f4725,f2854])).
fof(f5154,plain,(
  $false | $spl48),
  inference(subsumption_resolution,[],[f4825,f4826])).
fof(f4825,plain,(
  p(s(bool,'F_0')) | $spl48),
  inference(unit_resulting_resolution,[],[f121,f4725,f204])).
fof(f4827,plain,(
  $false | $spl48),
  inference(unit_resulting_resolution,[],[f121,f123,f4725,f204])).
fof(f4823,plain,(
  $false | ($spl6 | $spl48)),
  inference(unit_resulting_resolution,[],[f123,f222,f4725,f2854])).
fof(f4820,plain,(
  $false | ($spl6 | $spl48)),
  inference(unit_resulting_resolution,[],[f222,f222,f4725,f2864])).
fof(f2864,plain,(
  ( ! [X21,X22] : (s(bool,X21) != s(bool,X22) | s(bool,'T_0') = s(bool,X22) | s(bool,'F_0') = s(bool,X21)) ) | $spl6),
  inference(subsumption_resolution,[],[f2863,f121])).
fof(f2863,plain,(
  ( ! [X21,X22] : (~p(s(bool,'T_0')) | s(bool,X21) != s(bool,X22) | s(bool,'T_0') = s(bool,X22) | s(bool,'F_0') = s(bool,X21)) ) | $spl6),
  inference(forward_demodulation,[],[f2720,f224])).
fof(f2720,plain,(
  ( ! [X21,X22] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | s(bool,X21) != s(bool,X22) | s(bool,'T_0') = s(bool,X22) | s(bool,'F_0') = s(bool,X21)) ) | $spl6),
  inference(superposition,[],[f251,f2097])).
fof(f4814,plain,(
  $false | $spl48),
  inference(unit_resulting_resolution,[],[f123,f4725,f126])).
fof(f4811,plain,(
  $false | ($spl6 | $spl48)),
  inference(unit_resulting_resolution,[],[f222,f222,f4725,f3294])).
fof(f3294,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') != s(bool,X8) | s(bool,X8) = s(bool,X9) | s(bool,'F_0') = s(bool,X9)) ) | $spl6),
  inference(superposition,[],[f222,f2694])).
fof(f4774,plain,(
  $false | $spl48),
  inference(unit_resulting_resolution,[],[f222,f4725,f369])).
fof(f4772,plain,(
  $false | $spl48),
  inference(unit_resulting_resolution,[],[f222,f4725,f368])).
fof(f368,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f362,f121])).
fof(f362,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,X0) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f251,f224])).
fof(f4768,plain,(
  $false | $spl48),
  inference(unit_resulting_resolution,[],[f121,f222,f4725,f251])).
fof(f4763,plain,(
  $false | $spl48),
  inference(unit_resulting_resolution,[],[f121,f4725,f130])).
fof(f4761,plain,(
  $false | ($spl6 | $spl48)),
  inference(unit_resulting_resolution,[],[f222,f222,f4725,f2234])).
fof(f2234,plain,(
  ( ! [X12,X11] : (s(bool,'F_0') != s(bool,X11) | s(bool,'T_0') = s(bool,X12) | s(bool,X11) = s(bool,X12)) ) | $spl6),
  inference(subsumption_resolution,[],[f2233,f121])).
fof(f2233,plain,(
  ( ! [X12,X11] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X11) | s(bool,'T_0') = s(bool,X12) | s(bool,X11) = s(bool,X12)) ) | $spl6),
  inference(forward_demodulation,[],[f1732,f224])).
fof(f1732,plain,(
  ( ! [X12,X11] : (s(bool,'F_0') != s(bool,X11) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | s(bool,'T_0') = s(bool,X12) | s(bool,X11) = s(bool,X12)) ) | $spl6),
  inference(superposition,[],[f130,f1270])).
fof(f4757,plain,(
  $false | $spl48),
  inference(unit_resulting_resolution,[],[f222,f4725])).
fof(f4737,plain,(
  $false | ($spl29 | $spl44)),
  inference(subsumption_resolution,[],[f4704,f4638])).
fof(f4736,plain,(
  $false | ($spl6 | $spl29 | $spl44)),
  inference(subsumption_resolution,[],[f4735,f4638])).
fof(f4735,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | ($spl6 | $spl44)),
  inference(condensation,[],[f4734])).
fof(f4734,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X0)) ) | ($spl6 | $spl44)),
  inference(condensation,[],[f4733])).
fof(f4733,plain,(
  ( ! [X19,X20] : (s(bool,'F_0') = s(bool,X20) | s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X19)) ) | ($spl6 | $spl44)),
  inference(subsumption_resolution,[],[f4702,f2694])).
fof(f4702,plain,(
  ( ! [X19,X20] : (s(bool,X19) != s(bool,X20) | s(bool,'F_0') = s(bool,X20) | s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X19)) ) | ($spl6 | $spl44)),
  inference(superposition,[],[f4649,f2694])).
fof(f4732,plain,(
  $false | ($spl6 | $spl29 | $spl44)),
  inference(subsumption_resolution,[],[f4731,f4638])).
fof(f4731,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | ($spl6 | $spl44)),
  inference(condensation,[],[f4730])).
fof(f4730,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'F_0') = s(bool,f(s(bool,'F_0')))) ) | ($spl6 | $spl44)),
  inference(condensation,[],[f4729])).
fof(f4729,plain,(
  ( ! [X17,X18] : (s(bool,'F_0') = s(bool,X18) | s(bool,'F_0') = s(bool,X17) | s(bool,'F_0') = s(bool,f(s(bool,'F_0')))) ) | ($spl6 | $spl44)),
  inference(subsumption_resolution,[],[f4701,f2694])).
fof(f4701,plain,(
  ( ! [X17,X18] : (s(bool,X17) != s(bool,X18) | s(bool,'F_0') = s(bool,X18) | s(bool,'F_0') = s(bool,X17) | s(bool,'F_0') = s(bool,f(s(bool,'F_0')))) ) | ($spl6 | $spl44)),
  inference(superposition,[],[f4649,f2694])).
fof(f4728,plain,(
  ( ! [X10] : (s(bool,X10) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,X10) = s(bool,f(s(bool,'F_0')))) ) | $spl50),
  inference(cnf_transformation,[],[f4728_D])).
fof(f4728_D,plain,(
  ( ! [X10] : (s(bool,X10) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,X10) = s(bool,f(s(bool,'F_0')))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f4727,plain,(
  ( ! [X10,X9] : (s(bool,'F_0') = s(bool,X9) | s(bool,X10) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,X10) = s(bool,f(s(bool,'F_0')))) ) | ($spl6 | $spl44)),
  inference(subsumption_resolution,[],[f4697,f125])).
fof(f4697,plain,(
  ( ! [X10,X9] : (s(bool,'T_0') != s(bool,X9) | s(bool,'F_0') = s(bool,X9) | s(bool,X10) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,X10) = s(bool,f(s(bool,'F_0')))) ) | ($spl6 | $spl44)),
  inference(superposition,[],[f4649,f1270])).
fof(f4724,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl46),
  inference(cnf_transformation,[],[f4724_D])).
fof(f4724_D,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f4726,plain,(
  ( ! [X8,X7] : (s(bool,'F_0') = s(bool,X7) | s(bool,X8) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X8)) ) | $spl44),
  inference(subsumption_resolution,[],[f4696,f125])).
fof(f4696,plain,(
  ( ! [X8,X7] : (s(bool,'T_0') != s(bool,X7) | s(bool,'F_0') = s(bool,X7) | s(bool,X8) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X8)) ) | $spl44),
  inference(superposition,[],[f4649,f247])).
fof(f4723,plain,(
  ( ! [X2,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X2) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl44),
  inference(subsumption_resolution,[],[f4693,f125])).
fof(f4693,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,'F_0') = s(bool,X1) | s(bool,X2) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl44),
  inference(superposition,[],[f4649,f247])).
fof(f4646,plain,(
  ( ! [X25] : (s(bool,X25) = s(bool,g(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X25)) ) | $spl40),
  inference(cnf_transformation,[],[f4646_D])).
fof(f4646_D,plain,(
  ( ! [X25] : (s(bool,X25) = s(bool,g(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X25)) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f4636,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK4(X0,X1,bool,bool))) ) | $spl26),
  inference(cnf_transformation,[],[f4636_D])).
fof(f4636_D,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK4(X0,X1,bool,bool))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f4638,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl29),
  inference(cnf_transformation,[],[f4638_D])).
fof(f4638_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f4631,plain,(
  ( ! [X33,X34,X32] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X32) = s(fun(bool,bool),X33) | s(bool,'T_0') = s(bool,sK4(X32,X33,bool,bool)) | s(bool,X34) = s(bool,g(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X34)) ) | $spl6),
  inference(superposition,[],[f547,f2694])).
fof(f547,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK4(X0,X1,X2,bool))) )),
  inference(superposition,[],[f175,f125])).
fof(f4648,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | $spl42),
  inference(cnf_transformation,[],[f4648_D])).
fof(f4648_D,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f4649,plain,(
  ( ! [X26] : (s(bool,X26) != s(bool,f(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X26)) ) | $spl44),
  inference(cnf_transformation,[],[f4649_D])).
fof(f4649_D,plain,(
  ( ! [X26] : (s(bool,X26) != s(bool,f(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X26)) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f4630,plain,(
  ( ! [X30,X31,X29] : (s(bool,X29) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X30) = s(fun(bool,bool),X31) | s(bool,'T_0') = s(bool,sK4(X30,X31,bool,bool)) | s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X29)) ) | $spl6),
  inference(superposition,[],[f547,f2694])).
fof(f4629,plain,(
  ( ! [X28,X26,X27] : (s(bool,X26) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X27) = s(fun(bool,bool),X28) | s(bool,'T_0') = s(bool,sK4(X27,X28,bool,bool)) | s(bool,'F_0') = s(bool,X26) | s(bool,'F_0') = s(bool,g(s(bool,'F_0')))) ) | $spl6),
  inference(superposition,[],[f547,f2694])).
fof(f4628,plain,(
  ( ! [X24,X23,X25] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X23) = s(fun(bool,bool),X24) | s(bool,'T_0') = s(bool,sK4(X23,X24,bool,bool)) | s(bool,X25) = s(bool,g(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X25)) ) | $spl6),
  inference(superposition,[],[f547,f2694])).
fof(f4635,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl24),
  inference(cnf_transformation,[],[f4635_D])).
fof(f4635_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f4645,plain,(
  ( ! [X20] : (s(bool,X20) != s(bool,f(s(bool,'F_0'))) | s(bool,X20) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) ) | $spl38),
  inference(cnf_transformation,[],[f4645_D])).
fof(f4645_D,plain,(
  ( ! [X20] : (s(bool,X20) != s(bool,f(s(bool,'F_0'))) | s(bool,X20) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f4627,plain,(
  ( ! [X21,X22,X20] : (s(bool,X20) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X21) = s(fun(bool,bool),X22) | s(bool,'T_0') = s(bool,sK4(X21,X22,bool,bool)) | s(bool,X20) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) ) | $spl6),
  inference(superposition,[],[f547,f1270])).
fof(f4642,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) ) | $spl34),
  inference(cnf_transformation,[],[f4642_D])).
fof(f4642_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f4644,plain,(
  ( ! [X19,X17,X18] : (s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,X17) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X18) = s(fun(bool,bool),X19) | s(bool,'T_0') = s(bool,sK4(X18,X19,bool,bool)) | s(bool,'T_0') = s(bool,X17)) ) | $spl6),
  inference(forward_demodulation,[],[f4626,f224])).
fof(f4626,plain,(
  ( ! [X19,X17,X18] : (s(bool,X17) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X18) = s(fun(bool,bool),X19) | s(bool,'T_0') = s(bool,sK4(X18,X19,bool,bool)) | s(bool,g(s(bool,'F_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X17)) ) | $spl6),
  inference(superposition,[],[f547,f1270])).
fof(f4643,plain,(
  ( ! [X16] : (s(bool,X16) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,X16) = s(bool,g(s(bool,'F_0')))) ) | $spl36),
  inference(cnf_transformation,[],[f4643_D])).
fof(f4643_D,plain,(
  ( ! [X16] : (s(bool,X16) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,X16) = s(bool,g(s(bool,'F_0')))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f4641,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl33),
  inference(cnf_transformation,[],[f4641_D])).
fof(f4641_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f4625,plain,(
  ( ! [X14,X15,X16] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X14) = s(fun(bool,bool),X15) | s(bool,'T_0') = s(bool,sK4(X14,X15,bool,bool)) | s(bool,X16) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,X16) = s(bool,g(s(bool,'F_0')))) ) | $spl6),
  inference(superposition,[],[f547,f1270])).
fof(f4639,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) ) | $spl30),
  inference(cnf_transformation,[],[f4639_D])).
fof(f4639_D,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f4624,plain,(
  ( ! [X12,X13,X11] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X11) = s(fun(bool,bool),X12) | s(bool,'T_0') = s(bool,sK4(X11,X12,bool,bool)) | s(bool,X13) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f547,f247])).
fof(f4623,plain,(
  ( ! [X10,X8,X9] : (s(bool,X8) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X9) = s(fun(bool,bool),X10) | s(bool,'T_0') = s(bool,sK4(X9,X10,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f547,f247])).
fof(f4622,plain,(
  ( ! [X6,X7,X5] : (s(bool,X5) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X6) = s(fun(bool,bool),X7) | s(bool,'T_0') = s(bool,sK4(X6,X7,bool,bool)) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f547,f247])).
fof(f4621,plain,(
  ( ! [X4,X2,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,sK4(X2,X3,bool,bool)) | s(bool,X4) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f547,f247])).
fof(f4620,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK4(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f547,f125])).
fof(f4084,plain,(
  ( ! [X24,X25] : (s(bool,X24) != s(bool,X25) | s(bool,X24) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,'F_0') = s(bool,X25)) ) | $spl22),
  inference(cnf_transformation,[],[f4084_D])).
fof(f4084_D,plain,(
  ( ! [X24,X25] : (s(bool,X24) != s(bool,X25) | s(bool,X24) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,'F_0') = s(bool,X25)) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f4061,plain,(
  ( ! [X24,X23,X25] : (s(bool,X24) != s(bool,X25) | s(bool,'T_0') = s(bool,X23) | s(bool,'F_0') = s(bool,X25) | s(bool,X24) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f4042])).
fof(f4042,plain,(
  ( ! [X24,X23,X25] : (s(bool,X24) != s(bool,X25) | s(bool,'T_0') = s(bool,X23) | s(bool,'F_0') = s(bool,X25) | s(bool,X24) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X23)) ) | $spl6),
  inference(superposition,[],[f2864,f1270])).
fof(f2869,plain,(
  ( ! [X34] : (s(bool,X34) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) | s(bool,'F_0') = s(bool,X34)) ) | $spl20),
  inference(cnf_transformation,[],[f2869_D])).
fof(f2869_D,plain,(
  ( ! [X34] : (s(bool,X34) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) | s(bool,'F_0') = s(bool,X34)) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f1273,plain,(
  ( ! [X47,X48,X49] : (s('type.pair.prod'(X47,bool),'const.pair.,_2'(s(X47,X48),s(bool,'T_0'))) != s('type.pair.prod'(X47,bool),'const.pair.,_2'(s(X47,X49),s(bool,'T_0')))) ) | $spl12),
  inference(cnf_transformation,[],[f1273_D])).
fof(f1273_D,plain,(
  ( ! [X47,X48,X49] : (s('type.pair.prod'(X47,bool),'const.pair.,_2'(s(X47,X48),s(bool,'T_0'))) != s('type.pair.prod'(X47,bool),'const.pair.,_2'(s(X47,X49),s(bool,'T_0')))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f2870,plain,(
  ( ! [X45,X43,X44,X42] : (s('type.pair.prod'(X43,bool),'const.pair.,_2'(s(X43,X44),s(bool,'T_0'))) != s('type.pair.prod'(X43,bool),'const.pair.,_2'(s(X43,X45),s(bool,'T_0'))) | s(bool,X42) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) | s(bool,'F_0') = s(bool,X42)) ) | $spl6),
  inference(forward_demodulation,[],[f2728,f224])).
fof(f2728,plain,(
  ( ! [X45,X43,X44,X42] : (s('type.pair.prod'(X43,bool),'const.pair.,_2'(s(X43,X44),s(bool,'T_0'))) != s('type.pair.prod'(X43,bool),'const.pair.,_2'(s(X43,X45),s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))))) | s(bool,X42) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) | s(bool,'F_0') = s(bool,X42)) ) | $spl6),
  inference(superposition,[],[f481,f2097])).
fof(f481,plain,(
  ( ! [X14,X12,X15,X13] : (s('type.pair.prod'(X12,bool),'const.pair.,_2'(s(X12,X13),s(bool,'T_0'))) != s('type.pair.prod'(X12,bool),'const.pair.,_2'(s(X12,X14),s(bool,X15))) | s(bool,X15) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3)))) )),
  inference(superposition,[],[f177,f225])).
fof(f1271,plain,(
  ( ! [X39,X37,X38] : (s('type.pair.prod'(bool,X37),'const.pair.,_2'(s(bool,'T_0'),s(X37,X38))) != s('type.pair.prod'(bool,X37),'const.pair.,_2'(s(bool,'T_0'),s(X37,X39)))) ) | $spl8),
  inference(cnf_transformation,[],[f1271_D])).
fof(f1271_D,plain,(
  ( ! [X39,X37,X38] : (s('type.pair.prod'(bool,X37),'const.pair.,_2'(s(bool,'T_0'),s(X37,X38))) != s('type.pair.prod'(bool,X37),'const.pair.,_2'(s(bool,'T_0'),s(X37,X39)))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f2868,plain,(
  ( ! [X37,X35,X36,X34] : (s('type.pair.prod'(bool,X35),'const.pair.,_2'(s(bool,'T_0'),s(X35,X36))) != s('type.pair.prod'(bool,X35),'const.pair.,_2'(s(bool,'T_0'),s(X35,X37))) | s(bool,X34) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) | s(bool,'F_0') = s(bool,X34)) ) | $spl6),
  inference(forward_demodulation,[],[f2726,f224])).
fof(f2726,plain,(
  ( ! [X37,X35,X36,X34] : (s('type.pair.prod'(bool,X35),'const.pair.,_2'(s(bool,'T_0'),s(X35,X36))) != s('type.pair.prod'(bool,X35),'const.pair.,_2'(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))),s(X35,X37))) | s(bool,X34) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) | s(bool,'F_0') = s(bool,X34)) ) | $spl6),
  inference(superposition,[],[f395,f2097])).
fof(f395,plain,(
  ( ! [X14,X12,X15,X13] : (s('type.pair.prod'(bool,X12),'const.pair.,_2'(s(bool,'T_0'),s(X12,X13))) != s('type.pair.prod'(bool,X12),'const.pair.,_2'(s(bool,X14),s(X12,X15))) | s(bool,X14) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3)))) )),
  inference(superposition,[],[f176,f225])).
fof(f2675,plain,(
  $false | $spl12),
  inference(trivial_inequality_removal,[],[f2606])).
fof(f2606,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl12),
  inference(unit_resulting_resolution,[],[f174,f1273,f178])).
fof(f2674,plain,(
  $false | $spl12),
  inference(trivial_inequality_removal,[],[f2603])).
fof(f2603,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl12),
  inference(unit_resulting_resolution,[],[f174,f1273,f178])).
fof(f2659,plain,(
  $false | $spl12),
  inference(equality_resolution,[],[f1273])).
fof(f2602,plain,(
  $false | $spl8),
  inference(trivial_inequality_removal,[],[f2533])).
fof(f2533,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f174,f1271,f178])).
fof(f2601,plain,(
  $false | $spl8),
  inference(trivial_inequality_removal,[],[f2530])).
fof(f2530,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f174,f1271,f178])).
fof(f2586,plain,(
  $false | $spl8),
  inference(equality_resolution,[],[f1271])).
fof(f2377,plain,(
  $false | ($spl8 | $spl16)),
  inference(global_subsumption,[],[f1271,f120,f119,f118,f116,f121,f123,f124,f125,f127,f126,f131,f130,f135,f134,f173,f170,f174,f175,f178,f177,f176,f204,f117,f222,f225,f224,f247,f251,f286,f334,f368,f369,f386,f392,f393,f394,f395,f396,f397,f469,f479,f480,f481,f482,f483,f547,f548,f584,f638,f679,f691,f701,f749,f752,f768,f770,f830,f831,f832,f833,f834,f1267,f1268,f1269,f884,f885,f886,f887,f1275,f904,f907,f923,f925,f981,f982,f983,f984,f1430,f1432,f1434,f1435,f1437,f1457,f2197,f2367,f2345,f2348,f2372,f2374])).
fof(f2374,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.integer.int','type.integer.int'),X0) = s('type.pair.prod'('type.integer.int','type.integer.int'),'const.frac.rep_frac_1'(s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),X0)))))) ) | $spl16),
  inference(subsumption_resolution,[],[f2361,f121])).
fof(f2361,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.pair.prod'('type.integer.int','type.integer.int'),X0) = s('type.pair.prod'('type.integer.int','type.integer.int'),'const.frac.rep_frac_1'(s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),X0)))))) ) | $spl16),
  inference(backward_demodulation,[],[f2197,f134])).
fof(f2372,plain,(
  ( ! [X0,X1] : (s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X0),s('type.integer.int',X1))) = s('type.pair.prod'('type.integer.int','type.integer.int'),'const.frac.rep_frac_1'(s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))))) ) | $spl16),
  inference(subsumption_resolution,[],[f2360,f121])).
fof(f2360,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X0),s('type.integer.int',X1))) = s('type.pair.prod'('type.integer.int','type.integer.int'),'const.frac.rep_frac_1'(s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))))) ) | $spl16),
  inference(backward_demodulation,[],[f2197,f584])).
fof(f2348,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X1),s(bool,'T_0'))) != s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X2),s(bool,'T_0')))) ) | $spl16),
  inference(backward_demodulation,[],[f2197,f469])).
fof(f2345,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,'T_0'),s(X0,X1))) != s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,'T_0'),s(X0,X2)))) ) | $spl16),
  inference(backward_demodulation,[],[f2197,f392])).
fof(f2367,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl16),
  inference(subsumption_resolution,[],[f2342,f2197])).
fof(f2342,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl16),
  inference(backward_demodulation,[],[f2197,f130])).
fof(f1457,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,X1) != s(X0,X2) | s('type.pair.prod'(X3,X0),'const.pair.,_2'(s(X3,X4),s(X0,X1))) = s('type.pair.prod'(X3,X0),'const.pair.,_2'(s(X3,X4),s(X0,X2)))) )),
  inference(equality_resolution,[],[f178])).
fof(f1435,plain,(
  ( ! [X39,X37,X41,X38,X42,X40] : (s(bool,X37) != s(bool,X38) | s(X39,X40) != s(X39,X41) | s('type.pair.prod'(bool,X39),'const.pair.,_2'(s(bool,'T_0'),s(X39,X40))) = s('type.pair.prod'(bool,X39),'const.pair.,_2'(s(bool,X38),s(X39,X41))) | s(bool,X37) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f178,f247])).
fof(f1432,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (s(bool,X21) != s(bool,X22) | s(X23,X24) != s(X23,X25) | s('type.pair.prod'(bool,X23),'const.pair.,_2'(s(bool,X20),s(X23,X24))) = s('type.pair.prod'(bool,X23),'const.pair.,_2'(s(bool,X22),s(X23,X25))) | s(bool,'T_0') = s(bool,X20) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f178,f247])).
fof(f1430,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (s(bool,'T_0') != s(bool,X9) | s(X10,X11) != s(X10,X12) | s('type.pair.prod'(bool,X10),'const.pair.,_2'(s(bool,X8),s(X10,X11))) = s('type.pair.prod'(bool,X10),'const.pair.,_2'(s(bool,X9),s(X10,X12))) | s(bool,X8) = s(bool,X13) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f178,f247])).
fof(f925,plain,(
  ( ! [X206,X204,X207,X205] : (s(X206,f(s(bool,'T_0'))) != s(X206,g(s(bool,'T_0'))) | s(fun(bool,X206),X204) = s(fun(bool,X206),X205) | s(bool,X207) = s(bool,sK4(X204,X205,X206,bool)) | s(bool,'T_0') = s(bool,X207)) )),
  inference(superposition,[],[f175,f247])).
fof(f923,plain,(
  ( ! [X198,X196,X199,X197] : (s(bool,'T_0') != s(bool,f(s(X196,sK4(X197,X198,bool,X196)))) | s(fun(X196,bool),X197) = s(fun(X196,bool),X198) | s(bool,X199) = s(bool,g(s(X196,sK4(X197,X198,bool,X196)))) | s(bool,'T_0') = s(bool,X199)) )),
  inference(superposition,[],[f175,f247])).
fof(f907,plain,(
  ( ! [X127,X125,X130,X128,X126,X129] : (s('type.pair.prod'(X126,bool),'const.pair.,_2'(s(X126,X129),s(bool,'T_0'))) != s('type.pair.prod'(X126,bool),'const.pair.,_2'(s(X126,X127),s(bool,X128))) | s(bool,X125) = s(bool,X128) | s(bool,X125) = s(bool,X130) | s(bool,'T_0') = s(bool,X130)) )),
  inference(superposition,[],[f177,f247])).
fof(f904,plain,(
  ( ! [X111,X109,X107,X112,X110,X108] : (s('type.pair.prod'(bool,X108),'const.pair.,_2'(s(bool,'T_0'),s(X108,X109))) != s('type.pair.prod'(bool,X108),'const.pair.,_2'(s(bool,X110),s(X108,X111))) | s(bool,X107) = s(bool,X110) | s(bool,X107) = s(bool,X112) | s(bool,'T_0') = s(bool,X112)) )),
  inference(superposition,[],[f176,f247])).
fof(f1275,plain,(
  ( ! [X92,X91] : (s('type.pair.prod'('type.integer.int','type.integer.int'),X91) = s('type.pair.prod'('type.integer.int','type.integer.int'),'const.frac.rep_frac_1'(s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),X91))))) | s(bool,X92) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.pair.SND_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),X91))))) | s(bool,'T_0') = s(bool,X92)) )),
  inference(subsumption_resolution,[],[f900,f121])).
fof(f900,plain,(
  ( ! [X92,X91] : (~p(s(bool,'T_0')) | s('type.pair.prod'('type.integer.int','type.integer.int'),X91) = s('type.pair.prod'('type.integer.int','type.integer.int'),'const.frac.rep_frac_1'(s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),X91))))) | s(bool,X92) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.pair.SND_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),X91))))) | s(bool,'T_0') = s(bool,X92)) )),
  inference(superposition,[],[f134,f247])).
fof(f1269,plain,(
  ( ! [X24,X23,X22] : (s(bool,X22) != s(bool,X23) | s(bool,'T_0') = s(bool,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f879,f121])).
fof(f879,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | s(bool,X22) != s(bool,X23) | s(bool,'T_0') = s(bool,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f251,f247])).
fof(f1268,plain,(
  ( ! [X21,X19,X20] : (s(bool,X19) != s(bool,X20) | p(s(bool,X20)) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(subsumption_resolution,[],[f878,f121])).
fof(f878,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | s(bool,X19) != s(bool,X20) | p(s(bool,X20)) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f204,f247])).
fof(f1267,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f877,f121])).
fof(f877,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f173,f247])).
fof(f834,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (s('type.pair.prod'(X27,bool),'const.pair.,_2'(s(X27,X28),s(bool,X26))) != s('type.pair.prod'(X27,bool),'const.pair.,_2'(s(X27,X29),s(bool,X30))) | s(bool,X30) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) | s(bool,X26) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(superposition,[],[f481,f247])).
fof(f833,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (s('type.pair.prod'(X21,bool),'const.pair.,_2'(s(X21,X22),s(bool,X20))) != s('type.pair.prod'(X21,bool),'const.pair.,_2'(s(X21,X23),s(bool,X24))) | s(bool,X24) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,X20) = s(bool,X25) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f480,f247])).
fof(f832,plain,(
  ( ! [X14,X19,X17,X15,X18,X16] : (s('type.pair.prod'(bool,X15),'const.pair.,_2'(s(bool,X14),s(X15,X16))) != s('type.pair.prod'(bool,X15),'const.pair.,_2'(s(bool,X17),s(X15,X18))) | s(bool,X17) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) | s(bool,X14) = s(bool,X19) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f395,f247])).
fof(f831,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (s('type.pair.prod'(bool,X9),'const.pair.,_2'(s(bool,X8),s(X9,X10))) != s('type.pair.prod'(bool,X9),'const.pair.,_2'(s(bool,X11),s(X9,X12))) | s(bool,X11) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,X8) = s(bool,X13) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f394,f247])).
fof(f830,plain,(
  ( ! [X6,X4,X7,X5,X3] : (s(bool,X3) != s(bool,X4) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X5,bool),'const.pair.,_2'(s(X5,X6),s(bool,X4))))) | s(bool,X3) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f386,f247])).
fof(f770,plain,(
  ( ! [X210,X208,X211,X209] : (s(X210,f(s(bool,X211))) != s(X210,g(s(bool,X211))) | s(fun(bool,X210),X208) = s(fun(bool,X210),X209) | s(bool,'T_0') = s(bool,X211) | s(bool,'T_0') = s(bool,sK4(X208,X209,X210,bool))) )),
  inference(superposition,[],[f175,f247])).
fof(f768,plain,(
  ( ! [X202,X200,X203,X201] : (s(bool,X203) != s(bool,f(s(X200,sK4(X201,X202,bool,X200)))) | s(fun(X200,bool),X201) = s(fun(X200,bool),X202) | s(bool,'T_0') = s(bool,X203) | s(bool,'T_0') = s(bool,g(s(X200,sK4(X201,X202,bool,X200))))) )),
  inference(superposition,[],[f175,f247])).
fof(f752,plain,(
  ( ! [X132,X130,X128,X133,X131,X129] : (s('type.pair.prod'(X130,bool),'const.pair.,_2'(s(X130,X131),s(bool,X132))) != s('type.pair.prod'(X130,bool),'const.pair.,_2'(s(X130,X133),s(bool,X129))) | s(bool,X128) = s(bool,X132) | s(bool,'T_0') = s(bool,X129) | s(bool,'T_0') = s(bool,X128)) )),
  inference(superposition,[],[f177,f247])).
fof(f749,plain,(
  ( ! [X111,X114,X112,X110,X115,X113] : (s('type.pair.prod'(bool,X112),'const.pair.,_2'(s(bool,X111),s(X112,X113))) != s('type.pair.prod'(bool,X112),'const.pair.,_2'(s(bool,X114),s(X112,X115))) | s(bool,X110) = s(bool,X114) | s(bool,'T_0') = s(bool,X111) | s(bool,'T_0') = s(bool,X110)) )),
  inference(superposition,[],[f176,f247])).
fof(f701,plain,(
  ( ! [X14,X17,X15,X16] : (s(bool,'T_0') = s(bool,X16) | s(bool,X17) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X14,bool),'const.pair.,_2'(s(X14,X15),s(bool,X16))))) | s(bool,'T_0') = s(bool,X17)) )),
  inference(superposition,[],[f247,f174])).
fof(f691,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f247,f247])).
fof(f679,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,X24) = s(bool,X25) | s(bool,'T_0') = s(bool,X25) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X22,bool),'const.pair.,_2'(s(X22,X23),s(bool,X24)))))) )),
  inference(superposition,[],[f247,f174])).
fof(f638,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.pair.SND_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.frac.rep_frac_1'(s('type.frac.frac',X0))))))))) )),
  inference(trivial_inequality_removal,[],[f637])).
fof(f637,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.integer.int','type.integer.int'),'const.frac.rep_frac_1'(s('type.frac.frac',X0))) != s('type.pair.prod'('type.integer.int','type.integer.int'),'const.frac.rep_frac_1'(s('type.frac.frac',X0))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.pair.SND_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.frac.rep_frac_1'(s('type.frac.frac',X0))))))))) )),
  inference(superposition,[],[f135,f124])).
fof(f482,plain,(
  ( ! [X19,X17,X20,X18,X16] : (s('type.pair.prod'(X17,bool),'const.pair.,_2'(s(X17,X18),s(bool,X16))) != s('type.pair.prod'(X17,bool),'const.pair.,_2'(s(X17,X19),s(bool,X20))) | s(bool,'F_0') = s(bool,X20) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f177,f125])).
fof(f480,plain,(
  ( ! [X10,X8,X11,X9] : (s('type.pair.prod'(X8,bool),'const.pair.,_2'(s(X8,X9),s(bool,'T_0'))) != s('type.pair.prod'(X8,bool),'const.pair.,_2'(s(X8,X10),s(bool,X11))) | s(bool,X11) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) )),
  inference(superposition,[],[f177,f224])).
fof(f479,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s('type.pair.prod'(X4,X0),'const.pair.,_2'(s(X4,X5),s(X0,X3))) != s('type.pair.prod'(X4,X0),'const.pair.,_2'(s(X4,X6),s(X0,X7))) | s(X0,X7) = s(X0,'const.pair.SND_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))) )),
  inference(superposition,[],[f177,f174])).
fof(f469,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X1),s(bool,'F_0'))) != s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X2),s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f222,f177])).
fof(f396,plain,(
  ( ! [X19,X17,X20,X18,X16] : (s('type.pair.prod'(bool,X17),'const.pair.,_2'(s(bool,X16),s(X17,X18))) != s('type.pair.prod'(bool,X17),'const.pair.,_2'(s(bool,X19),s(X17,X20))) | s(bool,'F_0') = s(bool,X19) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f176,f125])).
fof(f394,plain,(
  ( ! [X10,X8,X11,X9] : (s('type.pair.prod'(bool,X8),'const.pair.,_2'(s(bool,'T_0'),s(X8,X9))) != s('type.pair.prod'(bool,X8),'const.pair.,_2'(s(bool,X10),s(X8,X11))) | s(bool,X10) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) )),
  inference(superposition,[],[f176,f224])).
fof(f393,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s('type.pair.prod'(X0,X4),'const.pair.,_2'(s(X0,X3),s(X4,X5))) != s('type.pair.prod'(X0,X4),'const.pair.,_2'(s(X0,X6),s(X4,X7))) | s(X0,X6) = s(X0,'const.pair.SND_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))) )),
  inference(superposition,[],[f176,f174])).
fof(f392,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,'T_0'),s(X0,X1))) != s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,'F_0'),s(X0,X2)))) )),
  inference(unit_resulting_resolution,[],[f222,f176])).
fof(f386,plain,(
  ( ! [X4,X2,X3] : (s(bool,'T_0') != s(bool,X4) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X2,bool),'const.pair.,_2'(s(X2,X3),s(bool,X4)))))) )),
  inference(forward_demodulation,[],[f377,f224])).
fof(f377,plain,(
  ( ! [X4,X2,X3] : (s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X2,bool),'const.pair.,_2'(s(X2,X3),s(bool,X4)))))) )),
  inference(superposition,[],[f368,f174])).
fof(f334,plain,(
  ( ! [X6,X4,X5] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X4,bool),'const.pair.,_2'(s(X4,X5),s(bool,X6))))) | s(bool,'F_0') = s(bool,X6)) )),
  inference(superposition,[],[f174,f125])).
fof(f173,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f73])).
fof(f135,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.integer.int','type.integer.int'),X0) != s('type.pair.prod'('type.integer.int','type.integer.int'),'const.frac.rep_frac_1'(s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),X0))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.pair.SND_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),X0))))))) )),
  inference(cnf_transformation,[],[f86])).
fof(f124,plain,(
  ( ! [X0] : (s('type.frac.frac',X0) = s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.frac.rep_frac_1'(s('type.frac.frac',X0)))))) )),
  inference(cnf_transformation,[],[f33])).
fof(f33,plain,(
  ! [X0] : s('type.frac.frac',X0) = s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.frac.rep_frac_1'(s('type.frac.frac',X0)))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X19] : s('type.frac.frac',X19) = s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.frac.rep_frac_1'(s('type.frac.frac',X19)))))),
  file('thm.frac.FRAC_EQ','thm.frac.frac_tybij|split|0')).
fof(f2376,plain,(
  $false | ($spl12 | $spl16)),
  inference(global_subsumption,[],[f1273,f120,f119,f118,f116,f121,f123,f124,f125,f127,f126,f131,f130,f135,f134,f173,f170,f174,f175,f178,f177,f176,f204,f117,f222,f225,f224,f247,f251,f286,f334,f368,f369,f386,f392,f393,f394,f395,f396,f397,f469,f479,f480,f481,f482,f483,f547,f548,f584,f638,f679,f691,f701,f749,f752,f768,f770,f830,f831,f832,f833,f834,f1267,f1268,f1269,f884,f885,f886,f887,f1275,f904,f907,f923,f925,f981,f982,f983,f984,f1430,f1432,f1434,f1435,f1437,f1457,f2197,f2367,f2345,f2348,f2372,f2374])).
fof(f2375,plain,(
  $false | $spl16),
  inference(global_subsumption,[],[f120,f119,f118,f116,f121,f123,f124,f125,f127,f126,f131,f130,f135,f134,f173,f170,f174,f175,f178,f177,f176,f204,f117,f222,f225,f224,f247,f251,f286,f334,f368,f369,f386,f392,f393,f394,f395,f396,f397,f469,f479,f480,f481,f482,f483,f547,f548,f584,f638,f679,f691,f701,f749,f752,f768,f770,f830,f831,f832,f833,f834,f1267,f1268,f1269,f884,f885,f886,f887,f1275,f904,f907,f923,f925,f981,f982,f983,f984,f1430,f1432,f1434,f1435,f1437,f1457,f2197,f2367,f2345,f2348,f2372,f2374])).
fof(f2373,plain,(
  $false | $spl16),
  inference(global_subsumption,[],[f120,f119,f118,f116,f121,f123,f124,f125,f127,f126,f131,f130,f135,f134,f173,f170,f174,f175,f178,f177,f176,f204,f117,f222,f225,f224,f247,f251,f286,f334,f368,f369,f386,f392,f393,f394,f395,f396,f397,f469,f479,f480,f481,f482,f483,f547,f548,f584,f638,f679,f691,f701,f749,f752,f768,f770,f830,f831,f832,f833,f834,f1267,f1268,f1269,f884,f885,f886,f887,f1275,f904,f907,f923,f925,f981,f982,f983,f984,f1430,f1432,f1434,f1435,f1437,f1457,f2197,f2367,f2345,f2348,f2372])).
fof(f2371,plain,(
  $false | ($spl6 | $spl16)),
  inference(global_subsumption,[],[f2365,f120,f119,f118,f116,f121,f123,f124,f125,f127,f126,f131,f130,f135,f134,f173,f170,f174,f175,f178,f177,f176,f204,f117,f222,f225,f224,f247,f251,f286,f334,f368,f369,f386,f392,f393,f394,f395,f396,f397,f469,f479,f480,f481,f482,f483,f547,f548,f584,f638,f679,f691,f701,f749,f752,f768,f770,f830,f831,f832,f833,f834,f1267,f1268,f1269,f884,f885,f886,f887,f1275,f904,f907,f923,f925,f981,f982,f983,f984,f1430,f1432,f1434,f1435,f1437,f1457,f2197,f2367,f2345,f2348])).
fof(f2365,plain,(
  ( ! [X222,X225,X223,X224] : (s(fun(X222,bool),X223) = s(fun(X222,bool),X224) | s(bool,X225) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,X225) = s(bool,g(s(X222,sK4(X223,X224,bool,X222))))) ) | ($spl6 | $spl16)),
  inference(trivial_inequality_removal,[],[f2357])).
fof(f2357,plain,(
  ( ! [X222,X225,X223,X224] : (s(bool,'T_0') != s(bool,'T_0') | s(fun(X222,bool),X223) = s(fun(X222,bool),X224) | s(bool,X225) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,X225) = s(bool,g(s(X222,sK4(X223,X224,bool,X222))))) ) | ($spl6 | $spl16)),
  inference(backward_demodulation,[],[f2197,f1880])).
fof(f1880,plain,(
  ( ! [X222,X225,X223,X224] : (s(bool,'T_0') != s(bool,f(s(X222,sK4(X223,X224,bool,X222)))) | s(fun(X222,bool),X223) = s(fun(X222,bool),X224) | s(bool,X225) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,X225) = s(bool,g(s(X222,sK4(X223,X224,bool,X222))))) ) | $spl6),
  inference(superposition,[],[f175,f1270])).
fof(f2370,plain,(
  $false | $spl16),
  inference(global_subsumption,[],[f120,f119,f118,f116,f121,f123,f124,f125,f127,f126,f131,f130,f135,f134,f173,f170,f174,f175,f178,f177,f176,f204,f117,f222,f225,f224,f247,f251,f286,f334,f368,f369,f386,f392,f393,f394,f395,f396,f397,f469,f479,f480,f481,f482,f483,f547,f548,f584,f638,f679,f691,f701,f749,f752,f768,f770,f830,f831,f832,f833,f834,f1267,f1268,f1269,f884,f885,f886,f887,f1275,f904,f907,f923,f925,f981,f982,f983,f984,f1430,f1432,f1434,f1435,f1437,f1457,f2197,f2367,f2345,f2348])).
fof(f2369,plain,(
  $false | $spl16),
  inference(global_subsumption,[],[f120,f119,f118,f116,f121,f123,f124,f125,f127,f126,f131,f130,f135,f134,f173,f170,f174,f175,f178,f177,f176,f204,f117,f222,f225,f224,f247,f251,f286,f334,f368,f369,f386,f392,f393,f394,f395,f396,f397,f469,f479,f480,f481,f482,f483,f547,f548,f584,f638,f679,f691,f701,f749,f752,f768,f770,f830,f831,f832,f833,f834,f1267,f1268,f1269,f884,f885,f886,f887,f1275,f904,f907,f923,f925,f981,f982,f983,f984,f1430,f1432,f1434,f1435,f1437,f1457,f2197,f2367,f2345])).
fof(f2368,plain,(
  $false | $spl16),
  inference(global_subsumption,[],[f120,f119,f118,f116,f121,f123,f124,f125,f127,f126,f131,f130,f135,f134,f173,f170,f174,f175,f178,f177,f176,f204,f117,f222,f225,f224,f247,f251,f286,f334,f368,f369,f386,f392,f393,f394,f395,f396,f397,f469,f479,f480,f481,f482,f483,f547,f548,f584,f638,f679,f691,f701,f749,f752,f768,f770,f830,f831,f832,f833,f834,f1267,f1268,f1269,f884,f885,f886,f887,f1275,f904,f907,f923,f925,f981,f982,f983,f984,f1430,f1432,f1434,f1435,f1437,f1457,f2197,f2367])).
fof(f2366,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2341,f121])).
fof(f2341,plain,(
  ~p(s(bool,'T_0')) | $spl16),
  inference(backward_demodulation,[],[f2197,f123])).
fof(f2198,plain,(
  ( ! [X24] : (s(bool,X24) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,'F_0') = s(bool,X24)) ) | $spl18),
  inference(cnf_transformation,[],[f2198_D])).
fof(f2198_D,plain,(
  ( ! [X24] : (s(bool,X24) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,'F_0') = s(bool,X24)) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f2098,plain,(
  ( ! [X95,X96] : (s(bool,'T_0') = s(bool,X95) | s(bool,X96) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,'F_0') = s(bool,X96)) ) | $spl6),
  inference(duplicate_literal_removal,[],[f1849])).
fof(f1849,plain,(
  ( ! [X95,X96] : (s(bool,'T_0') = s(bool,X95) | s(bool,'T_0') = s(bool,X95) | s(bool,X96) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,'F_0') = s(bool,X96)) ) | $spl6),
  inference(superposition,[],[f125,f1270])).
fof(f2196,plain,(
  ( ! [X14,X16] : (s(bool,X16) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,X14) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,X14) = s(bool,X16)) ) | $spl14),
  inference(cnf_transformation,[],[f2196_D])).
fof(f2196_D,plain,(
  ( ! [X14,X16] : (s(bool,X16) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,X14) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,X14) = s(bool,X16)) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f1839,plain,(
  ( ! [X57,X54,X58,X56,X55] : (s('type.pair.prod'(X55,bool),'const.pair.,_2'(s(X55,X56),s(bool,'T_0'))) != s('type.pair.prod'(X55,bool),'const.pair.,_2'(s(X55,X57),s(bool,'T_0'))) | s(bool,X54) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,X58) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,X54) = s(bool,X58)) ) | $spl6),
  inference(superposition,[],[f480,f1270])).
fof(f1837,plain,(
  ( ! [X47,X45,X48,X46,X44] : (s('type.pair.prod'(bool,X45),'const.pair.,_2'(s(bool,'T_0'),s(X45,X46))) != s('type.pair.prod'(bool,X45),'const.pair.,_2'(s(bool,'T_0'),s(X45,X47))) | s(bool,X44) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,X48) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,X44) = s(bool,X48)) ) | $spl6),
  inference(superposition,[],[f394,f1270])).
fof(f2146,plain,(
  ( ! [X24,X23,X25] : (s(bool,X24) = s(bool,X25) | s(bool,X25) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X23) | s(bool,X24) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f1664])).
fof(f1664,plain,(
  ( ! [X24,X23,X25] : (s(bool,X24) = s(bool,X25) | s(bool,X25) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X23) | s(bool,X24) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X23)) ) | $spl6),
  inference(superposition,[],[f1270,f1270])).
fof(f2152,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X0)) ) | $spl6),
  inference(duplicate_literal_removal,[],[f1656])).
fof(f1656,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl6),
  inference(superposition,[],[f1270,f125])).
fof(f2161,plain,(
  ( ! [X19,X17,X18] : (s(bool,'T_0') = s(bool,X18) | s(bool,X17) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,X19) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,X17) = s(bool,X19)) ) | $spl6),
  inference(duplicate_literal_removal,[],[f1645])).
fof(f1645,plain,(
  ( ! [X19,X17,X18] : (s(bool,'T_0') = s(bool,X18) | s(bool,X17) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X18) | s(bool,X19) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,X17) = s(bool,X19)) ) | $spl6),
  inference(superposition,[],[f1270,f1270])).
fof(f2167,plain,(
  ( ! [X24,X23] : (s(bool,'T_0') = s(bool,X23) | s(bool,X24) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,'F_0') = s(bool,X24)) ) | $spl6),
  inference(duplicate_literal_removal,[],[f1636])).
fof(f1636,plain,(
  ( ! [X24,X23] : (s(bool,'T_0') = s(bool,X23) | s(bool,X24) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,'F_0') = s(bool,X24) | s(bool,'T_0') = s(bool,X23)) ) | $spl6),
  inference(superposition,[],[f1270,f125])).
fof(f2168,plain,(
  ( ! [X14,X15,X16] : (s(bool,'T_0') = s(bool,X15) | s(bool,X16) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,X14) = s(bool,X16) | s(bool,X14) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f1634])).
fof(f1634,plain,(
  ( ! [X14,X15,X16] : (s(bool,'T_0') = s(bool,X15) | s(bool,X16) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,X14) = s(bool,X16) | s(bool,X14) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X15)) ) | $spl6),
  inference(superposition,[],[f1270,f1270])).
fof(f1272,plain,(
  ( ! [X45,X41] : (s(bool,X41) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) | s(bool,'T_0') = s(bool,X45) | s(bool,X41) = s(bool,X45)) ) | $spl10),
  inference(cnf_transformation,[],[f1272_D])).
fof(f1272_D,plain,(
  ( ! [X45,X41] : (s(bool,X41) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) | s(bool,'T_0') = s(bool,X45) | s(bool,X41) = s(bool,X45)) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f984,plain,(
  ( ! [X54,X52,X55,X53,X51] : (s('type.pair.prod'(X52,bool),'const.pair.,_2'(s(X52,X53),s(bool,'T_0'))) != s('type.pair.prod'(X52,bool),'const.pair.,_2'(s(X52,X54),s(bool,'T_0'))) | s(bool,X51) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) | s(bool,X51) = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) )),
  inference(superposition,[],[f481,f247])).
fof(f1270,plain,(
  ( ! [X36,X40] : (s(bool,X36) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X40) | s(bool,X36) = s(bool,X40)) ) | $spl6),
  inference(cnf_transformation,[],[f1270_D])).
fof(f1270_D,plain,(
  ( ! [X36,X40] : (s(bool,X36) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X40) | s(bool,X36) = s(bool,X40)) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f983,plain,(
  ( ! [X47,X50,X48,X46,X49] : (s('type.pair.prod'(X47,bool),'const.pair.,_2'(s(X47,X48),s(bool,'T_0'))) != s('type.pair.prod'(X47,bool),'const.pair.,_2'(s(X47,X49),s(bool,'T_0'))) | s(bool,X46) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,X46) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f480,f247])).
fof(f982,plain,(
  ( ! [X45,X43,X41,X44,X42] : (s('type.pair.prod'(bool,X42),'const.pair.,_2'(s(bool,'T_0'),s(X42,X43))) != s('type.pair.prod'(bool,X42),'const.pair.,_2'(s(bool,'T_0'),s(X42,X44))) | s(bool,X41) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) | s(bool,X41) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) )),
  inference(superposition,[],[f395,f247])).
fof(f981,plain,(
  ( ! [X39,X37,X38,X36,X40] : (s('type.pair.prod'(bool,X37),'const.pair.,_2'(s(bool,'T_0'),s(X37,X38))) != s('type.pair.prod'(bool,X37),'const.pair.,_2'(s(bool,'T_0'),s(X37,X39))) | s(bool,X36) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,X36) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f394,f247])).
fof(f887,plain,(
  ( ! [X54,X52,X55,X53,X51] : (s('type.pair.prod'(X52,bool),'const.pair.,_2'(s(X52,X53),s(bool,'T_0'))) != s('type.pair.prod'(X52,bool),'const.pair.,_2'(s(X52,X54),s(bool,'T_0'))) | s(bool,X51) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) | s(bool,X51) = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) )),
  inference(superposition,[],[f481,f247])).
fof(f886,plain,(
  ( ! [X47,X50,X48,X46,X49] : (s('type.pair.prod'(X47,bool),'const.pair.,_2'(s(X47,X48),s(bool,'T_0'))) != s('type.pair.prod'(X47,bool),'const.pair.,_2'(s(X47,X49),s(bool,'T_0'))) | s(bool,X46) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,X46) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f480,f247])).
fof(f885,plain,(
  ( ! [X45,X43,X41,X44,X42] : (s('type.pair.prod'(bool,X42),'const.pair.,_2'(s(bool,'T_0'),s(X42,X43))) != s('type.pair.prod'(bool,X42),'const.pair.,_2'(s(bool,'T_0'),s(X42,X44))) | s(bool,X41) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) | s(bool,X41) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) )),
  inference(superposition,[],[f395,f247])).
fof(f884,plain,(
  ( ! [X39,X37,X38,X36,X40] : (s('type.pair.prod'(bool,X37),'const.pair.,_2'(s(bool,'T_0'),s(X37,X38))) != s('type.pair.prod'(bool,X37),'const.pair.,_2'(s(bool,'T_0'),s(X37,X39))) | s(bool,X36) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s(bool,X36) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f394,f247])).
fof(f220,plain,(
  $false | ($spl1 | $spl2 | $spl4)),
  inference(trivial_inequality_removal,[],[f219])).
fof(f219,plain,(
  s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))))) != s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))))) | ($spl1 | $spl2 | $spl4)),
  inference(forward_demodulation,[],[f218,f213])).
fof(f218,plain,(
  s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))))) != s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',sK2),s('type.integer.int',sK1))))) | ($spl1 | $spl4)),
  inference(forward_demodulation,[],[f212,f215])).
fof(f211,plain,(
  s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))))) = s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))))) | $spl0),
  inference(cnf_transformation,[],[f211_D])).
fof(f211_D,plain,(
  s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))))) = s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f213,plain,(
  s('type.integer.int',sK0) = s('type.integer.int',sK2) | $spl2),
  inference(cnf_transformation,[],[f213_D])).
fof(f213_D,plain,(
  s('type.integer.int',sK0) = s('type.integer.int',sK2) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f118,plain,(
  s('type.integer.int',sK0) = s('type.integer.int',sK2) | s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))))) = s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',sK2),s('type.integer.int',sK3)))))),
  inference(cnf_transformation,[],[f81])).
fof(f215,plain,(
  s('type.integer.int',sK1) = s('type.integer.int',sK3) | $spl4),
  inference(cnf_transformation,[],[f215_D])).
fof(f215_D,plain,(
  s('type.integer.int',sK1) = s('type.integer.int',sK3) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f119,plain,(
  s('type.integer.int',sK1) = s('type.integer.int',sK3) | s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))))) = s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',sK2),s('type.integer.int',sK3)))))),
  inference(cnf_transformation,[],[f81])).
fof(f212,plain,(
  s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))))) != s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))))) | $spl1),
  inference(cnf_transformation,[],[f212_D])).
fof(f212_D,plain,(
  s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))))) != s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f214,plain,(
  s('type.integer.int',sK0) != s('type.integer.int',sK2) | $spl3),
  inference(cnf_transformation,[],[f214_D])).
fof(f214_D,plain,(
  s('type.integer.int',sK0) != s('type.integer.int',sK2) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f216,plain,(
  s('type.integer.int',sK1) != s('type.integer.int',sK3) | $spl5),
  inference(cnf_transformation,[],[f216_D])).
fof(f216_D,plain,(
  s('type.integer.int',sK1) != s('type.integer.int',sK3) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f120,plain,(
  s('type.integer.int',sK1) != s('type.integer.int',sK3) | s('type.integer.int',sK0) != s('type.integer.int',sK2) | s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))))) != s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',sK2),s('type.integer.int',sK3)))))),
  inference(cnf_transformation,[],[f81])).
% SZS output end Proof for thm.frac
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 116
% Passive clauses: 525
% Generated clauses: 12746
% Final active clauses: 88
% Final passive clauses: 280
% Input formulas: 29
% Initial clauses: 77
% 
% Duplicate literals: 577
% Trivial inequalities: 128
% Fw subsumption resolutions: 755
% Fw demodulations: 1883
% Bw demodulations: 103
% Condensations: 1159
% Global subsumptions: 9
% 
% Simple tautologies: 185
% Equational tautologies: 97
% Forward subsumptions: 6506
% Fw demodulations to eq. taut.: 803
% Bw demodulations to eq. taut.: 3
% 
% Unit resulting resolution: 1355
% Factoring: 3
% Forward superposition: 3662
% Backward superposition: 2702
% Self superposition: 87
% Equality factoring: 80
% Equality resolution: 29
% 
% Split clauses: 59
% Split components: 132
% SAT solver clauses: 721
% SAT solver unit clauses: 55
% SAT solver binary clauses: 126
% 
% Sat splits: 59
% Sat splitting refutations: 67
% 
% Memory used [KB]: 6652
% Time elapsed: 0.454 s
% ------------------------------
% ------------------------------
% Success in time 1.77 s
