Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.arithmetic
% (88365)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 636
% Passive clauses: 9231
% Generated clauses: 20065
% Final active clauses: 597
% Final passive clauses: 8122
% Input formulas: 69
% Initial clauses: 243
% 
% Duplicate literals: 397
% Trivial inequalities: 57
% Fw subsumption resolutions: 1336
% Fw demodulations: 756
% Bw demodulations: 38
% 
% Simple tautologies: 269
% Equational tautologies: 60
% Forward subsumptions: 7447
% Fw demodulations to eq. taut.: 9
% 
% Binary resolution: 9115
% Factoring: 940
% Forward superposition: 2545
% Backward superposition: 4351
% Self superposition: 63
% Equality factoring: 26
% Equality resolution: 26
% 
% Split clauses: 340
% Split components: 681
% SAT solver clauses: 622
% SAT solver unit clauses: 14
% SAT solver binary clauses: 146
% 
% Sat splits: 340
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 619
% TWLsolver calls for satisfiability: 297
% 
% Memory used [KB]: 10618
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 144 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.arithmetic
% (88383)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 408
% Passive clauses: 3063
% Generated clauses: 14777
% Final active clauses: 402
% Final passive clauses: 2629
% Input formulas: 69
% Initial clauses: 243
% 
% Duplicate literals: 233
% Trivial inequalities: 14
% Fw subsumption resolutions: 935
% Fw demodulations: 706
% Bw demodulations: 28
% Condensations: 424
% Global subsumptions: 22
% 
% Simple tautologies: 285
% Equational tautologies: 22
% Forward subsumptions: 5436
% Fw demodulations to eq. taut.: 9
% 
% Binary resolution: 3969
% Factoring: 330
% Forward superposition: 2019
% Backward superposition: 2645
% Self superposition: 64
% Equality factoring: 26
% Equality resolution: 7
% 
% Split clauses: 42
% Split components: 84
% SAT solver clauses: 3872
% SAT solver unit clauses: 246
% SAT solver binary clauses: 85
% SAT solver learnt clauses: 45
% SAT solver learnt literals: 11
% 
% Sat splits: 42
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 3783
% TWLsolver calls for satisfiability: 2022
% 
% Memory used [KB]: 16630
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 137 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.arithmetic
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.arithmetic
% SZS output start Proof for thm.arithmetic
fof(f60739,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f334,f591,f589,f333,f590,f588,f2841,f2843,f2997,f2998,f2999,f3000,f3001,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3013,f3015,f3016,f3017,f3018,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3038,f3039,f3040,f3047,f3048,f3065,f3066,f3068,f3070,f3072,f3075,f3078,f3093,f3096,f3099,f3102,f3105,f3108,f3110,f3112,f3116,f3128,f3132,f3136,f3140,f5137,f5081,f25806,f25824,f25808,f25830,f25837,f25840,f25846,f25849,f25852,f25855,f25858,f25861,f25864,f25920,f25831,f25926,f25838,f25928,f25841,f25933,f25847,f25935,f25850,f25937,f25853,f25939,f25856,f25941,f25859,f25943,f25862,f25945,f25865,f26015,f26018,f26016,f26041,f26044,f26042,f26070,f26073,f26071,f26107,f26110,f26108,f26135,f26138,f26136,f26195,f26198,f26196,f26277,f26280,f26278,f26364,f26367,f26365,f27276,f27278,f27277,f28092,f28093,f28145,f28148,f28150,f28152,f28154,f28156,f28158,f28160,f28162,f28164,f28166,f28168,f28170,f28172,f28174,f28176,f28179,f28182,f28185,f28188,f28191,f28193,f28195,f28197,f28199,f28207,f28211,f28215,f28236,f28238,f28240,f28242,f28250,f28254,f28258,f27533,f28362,f27532,f28363,f27531,f28364,f27530,f28365,f27529,f27528,f28367,f27512,f28369,f28368,f27506,f28371,f27431,f28378,f27429,f28380,f27428,f28382,f27425,f28384,f27424,f28386,f27420,f28387,f27417,f28388,f28370,f27412,f28389,f27411,f28390,f27398,f28395,f27396,f28396,f27393,f28397,f27391,f28398,f27357,f28412,f28416,f28417,f28366,f28361,f46819,f46821,f46820,f46951,f46953,f46952,f47089,f47091,f47090,f49121,f49123,f49122,f49309,f49311,f49310,f49657,f49660,f49658,f49750,f49753,f2842,f49751,f50046,f50050,f50113,f50111,f57075,f60640,f60642,f25807,f60641,f60738])).
fof(f60738,plain,(
  $false | ($spl3 | $spl115 | $spl116)),
  inference(subsumption_resolution,[],[f60737,f591])).
fof(f60737,plain,(
  s('type.num.num',sK12) = s('type.num.num',sK13) | ($spl115 | $spl116)),
  inference(forward_demodulation,[],[f60651,f50113])).
fof(f60651,plain,(
  s('type.num.num',sK13) = s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',sK12),s('type.num.num',sK13))) | $spl115),
  inference(unit_resulting_resolution,[],[f50110,f1931])).
fof(f1931,plain,(
  ( ! [X0,X1] : (s('type.num.num',X1) = s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X0),s('type.num.num',X1))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) )),
  inference(forward_demodulation,[],[f1898,f414])).
fof(f414,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0,X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f117])).
fof(f117,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X12] : ! [X10,X11] : s(X12,X11) = s(X12,'const.bool.COND_3'(s(bool,'F_0'),s(X12,X10),s(X12,X11)))),
  file('thm.arithmetic.MIN_MAX_EQ','thm.bool.COND_CLAUSES|split|1')).
fof(f1898,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X0),s('type.num.num',X1))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) )),
  inference(superposition,[],[f431,f343])).
fof(f343,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X9] : (s(bool,'T_0') = s(bool,X9) | s(bool,'F_0') = s(bool,X9))),
  file('thm.arithmetic.MIN_MAX_EQ','thm.bool.BOOL_CASES_AX')).
fof(f431,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X0),s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f129])).
fof(f129,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X0),s('type.num.num',X1)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6,X7] : s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X6),s('type.num.num',X7))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num',X7))),s('type.num.num',X6),s('type.num.num',X7)))),
  file('thm.arithmetic.MIN_MAX_EQ','thm.arithmetic.MIN_DEF')).
fof(f50110,plain,(
  s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK12),s('type.num.num',sK13))) | $spl115),
  inference(cnf_transformation,[],[f50110_D])).
fof(f50110_D,plain,(
  s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK12),s('type.num.num',sK13))) <=> ~$spl115),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl115])])).
fof(f60641,plain,(
  ( ! [X14,X12,X17,X15,X16] : (s(X14,X15) = s(X14,'const.bool.COND_3'(s(bool,X12),s(X14,X16),s(X14,X17))) | s(bool,'T_0') = s(bool,X12) | s(X14,X15) != s(X14,X17)) ) | $spl118),
  inference(cnf_transformation,[],[f60641_D])).
fof(f60641_D,plain,(
  ( ! [X14,X12,X17,X15,X16] : (s(X14,X15) = s(X14,'const.bool.COND_3'(s(bool,X12),s(X14,X16),s(X14,X17))) | s(bool,'T_0') = s(bool,X12) | s(X14,X15) != s(X14,X17)) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f25807,plain,(
  ( ! [X292] : (p(s(bool,X292))) ) | $spl8),
  inference(cnf_transformation,[],[f25807_D])).
fof(f25807_D,plain,(
  ( ! [X292] : (p(s(bool,X292))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f60642,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (p(s(bool,X19)) | s(X20,X21) = s(X20,'const.bool.COND_3'(s(bool,X18),s(X20,X22),s(X20,X23))) | s(X20,X21) != s(X20,X23) | s(bool,'T_0') = s(bool,X18)) ) | $spl6),
  inference(subsumption_resolution,[],[f60605,f354])).
fof(f354,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f226])).
fof(f226,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X9] : (s(bool,'T_0') = s(bool,X9) <=> p(s(bool,X9)))),
  file('thm.arithmetic.MIN_MAX_EQ','thm.bool.EQ_CLAUSES|split|1')).
fof(f60605,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (p(s(bool,X19)) | s(X20,X21) = s(X20,'const.bool.COND_3'(s(bool,X18),s(X20,X22),s(X20,X23))) | s(X20,X21) != s(X20,X23) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,X18)) ) | $spl6),
  inference(superposition,[],[f2843,f701])).
fof(f701,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f343,f343])).
fof(f60640,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (p(s(bool,X13)) | s(X14,X15) = s(X14,'const.bool.COND_3'(s(bool,X12),s(X14,X16),s(X14,X17))) | s(X14,X15) != s(X14,X17) | s(bool,'T_0') = s(bool,X12)) ) | $spl6),
  inference(subsumption_resolution,[],[f60604,f354])).
fof(f60604,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (p(s(bool,X13)) | s(X14,X15) = s(X14,'const.bool.COND_3'(s(bool,X12),s(X14,X16),s(X14,X17))) | s(X14,X15) != s(X14,X17) | s(bool,'T_0') = s(bool,X12) | s(bool,'T_0') = s(bool,X13)) ) | $spl6),
  inference(superposition,[],[f2843,f701])).
fof(f57075,plain,(
  $false | ($spl0 | $spl3 | $spl114)),
  inference(subsumption_resolution,[],[f57074,f591])).
fof(f57074,plain,(
  s('type.num.num',sK12) = s('type.num.num',sK13) | ($spl0 | $spl114)),
  inference(backward_demodulation,[],[f57073,f57072])).
fof(f57072,plain,(
  s('type.num.num',sK13) = s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',sK12),s('type.num.num',sK13))) | ($spl0 | $spl114)),
  inference(backward_demodulation,[],[f57071,f588])).
fof(f57071,plain,(
  s('type.num.num',sK13) = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',sK12),s('type.num.num',sK13))) | $spl114),
  inference(forward_demodulation,[],[f52182,f413])).
fof(f413,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f116])).
fof(f116,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f115])).
fof(f115,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X12] : ! [X10,X11] : s(X12,X10) = s(X12,'const.bool.COND_3'(s(bool,'T_0'),s(X12,X10),s(X12,X11)))),
  file('thm.arithmetic.MIN_MAX_EQ','thm.bool.COND_CLAUSES|split|0')).
fof(f52182,plain,(
  s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',sK12),s('type.num.num',sK13))) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',sK13),s('type.num.num',sK12))) | $spl114),
  inference(superposition,[],[f432,f50111])).
fof(f432,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X1),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X1),s('type.num.num',X0)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X6,X7] : s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X6),s('type.num.num',X7))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num',X7))),s('type.num.num',X7),s('type.num.num',X6)))),
  file('thm.arithmetic.MIN_MAX_EQ','thm.arithmetic.MAX_DEF')).
fof(f57073,plain,(
  s('type.num.num',sK12) = s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',sK12),s('type.num.num',sK13))) | $spl114),
  inference(forward_demodulation,[],[f52183,f413])).
fof(f52183,plain,(
  s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',sK12),s('type.num.num',sK13))) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',sK12),s('type.num.num',sK13))) | $spl114),
  inference(superposition,[],[f431,f50111])).
fof(f50111,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK12),s('type.num.num',sK13))) | $spl114),
  inference(cnf_transformation,[],[f50111_D])).
fof(f50111_D,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK12),s('type.num.num',sK13))) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f50113,plain,(
  s('type.num.num',sK12) = s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',sK12),s('type.num.num',sK13))) | $spl116),
  inference(cnf_transformation,[],[f50113_D])).
fof(f50113_D,plain,(
  s('type.num.num',sK12) = s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',sK12),s('type.num.num',sK13))) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f50050,plain,(
  s('type.num.num',sK12) = s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',sK12),s('type.num.num',sK13))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK12),s('type.num.num',sK13))) | $spl0),
  inference(superposition,[],[f588,f2001])).
fof(f2001,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X0),s('type.num.num',X1))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) )),
  inference(forward_demodulation,[],[f1962,f414])).
fof(f1962,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X1),s('type.num.num',X0))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) )),
  inference(superposition,[],[f432,f343])).
fof(f50046,plain,(
  s('type.num.num',sK12) = s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',sK12),s('type.num.num',sK13))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK12),s('type.num.num',sK13))) | $spl0),
  inference(superposition,[],[f2001,f588])).
fof(f49751,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,'T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,X1)))) ) | $spl112),
  inference(cnf_transformation,[],[f49751_D])).
fof(f49751_D,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,'T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,X1)))) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f2842,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl4),
  inference(cnf_transformation,[],[f2842_D])).
fof(f2842_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f49753,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,'T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,X1)))) )),
  inference(condensation,[],[f49752])).
fof(f49752,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X30)) | p(s(bool,X31)) | ~sP7(X31,'T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X32),s(bool,X30))) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f49738,f747])).
fof(f747,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f744,f592])).
fof(f592,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f338,f335])).
fof(f335,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.arithmetic.MIN_MAX_EQ',aHL_TRUTH)).
fof(f338,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f224])).
fof(f224,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f72])).
fof(f72,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f52])).
fof(f52,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.arithmetic.MIN_MAX_EQ','thm.bool.NOT_CLAUSES|split|1')).
fof(f744,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f434,f343])).
fof(f434,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f188])).
fof(f188,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f187])).
fof(f187,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X10,X11] : ((p(s(bool,X10)) => p(s(bool,X11))) => ((p(s(bool,X11)) => p(s(bool,X10))) => s(bool,X10) = s(bool,X11)))),
  file('thm.arithmetic.MIN_MAX_EQ','thm.bool.IMP_ANTISYM_AX')).
fof(f49738,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | ~p(s(bool,X30)) | p(s(bool,X31)) | ~sP7(X31,'T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X32),s(bool,X30))) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f1048,f701])).
fof(f1048,plain,(
  ( ! [X12,X10,X13,X11] : (~p(s(bool,X12)) | ~p(s(bool,X11)) | p(s(bool,X13)) | ~sP7(X13,X12,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X10),s(bool,X11)))) )),
  inference(superposition,[],[f541,f414])).
fof(f541,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP7(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f316])).
fof(f316,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f315])).
fof(f315,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f215])).
fof(f215,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP7(X2,X0,X1))),
  inference(definition_folding,[],[f163,f214,f213])).
fof(f213,plain,(
  ! [X0,X1,X2] : (sP6(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f214,plain,(
  ! [X2,X0,X1] : (sP7(X2,X0,X1) <=> (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f163,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f162])).
fof(f162,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X14,X15,X8] : ((p(s(bool,X8)) <=> (p(s(bool,X15)) & p(s(bool,X14)))) <=> ((p(s(bool,X8)) | ~p(s(bool,X15)) | ~p(s(bool,X14))) & (p(s(bool,X15)) | ~p(s(bool,X8))) & (p(s(bool,X14)) | ~p(s(bool,X8)))))),
  file('thm.arithmetic.MIN_MAX_EQ','thm.sat.sat_23_prove_ltb')).
fof(f49750,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,'T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,X1)))) )),
  inference(condensation,[],[f49749])).
fof(f49749,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X25)) | p(s(bool,X26)) | ~sP7(X26,'T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X27),s(bool,X25))) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f49737,f747])).
fof(f49737,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | p(s(bool,X26)) | ~sP7(X26,'T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X27),s(bool,X25))) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f1048,f701])).
fof(f49658,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,'T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X3)))) ) | $spl110),
  inference(cnf_transformation,[],[f49658_D])).
fof(f49658_D,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,'T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X3)))) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f49660,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,'T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X3)))) )),
  inference(condensation,[],[f49659])).
fof(f49659,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X30)) | p(s(bool,X31)) | ~sP7(X31,'T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X30),s(bool,X32))) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f49645,f747])).
fof(f49645,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | ~p(s(bool,X30)) | p(s(bool,X31)) | ~sP7(X31,'T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X30),s(bool,X32))) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f1047,f701])).
fof(f1047,plain,(
  ( ! [X6,X8,X7,X9] : (~p(s(bool,X8)) | ~p(s(bool,X6)) | p(s(bool,X9)) | ~sP7(X9,X8,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X6),s(bool,X7)))) )),
  inference(superposition,[],[f541,f413])).
fof(f49657,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,'T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X3)))) )),
  inference(condensation,[],[f49656])).
fof(f49656,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X25)) | p(s(bool,X26)) | ~sP7(X26,'T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X25),s(bool,X27))) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f49644,f747])).
fof(f49644,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | p(s(bool,X26)) | ~sP7(X26,'T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X25),s(bool,X27))) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f1047,f701])).
fof(f49310,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,'const.bool.COND_3'(s(bool,X10),s(bool,X11),s(bool,X9)))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,X10)) ) | $spl108),
  inference(cnf_transformation,[],[f49310_D])).
fof(f49310_D,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,'const.bool.COND_3'(s(bool,X10),s(bool,X11),s(bool,X9)))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,X10)) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f49311,plain,(
  ( ! [X14,X12,X15,X13] : (~p(s(bool,'const.bool.COND_3'(s(bool,X14),s(bool,X15),s(bool,X13)))) | p(s(bool,X12)) | s(bool,'T_0') = s(bool,X14) | s(bool,'T_0') = s(bool,X13)) )),
  inference(subsumption_resolution,[],[f49263,f354])).
fof(f49263,plain,(
  ( ! [X14,X12,X15,X13] : (~p(s(bool,'const.bool.COND_3'(s(bool,X14),s(bool,X15),s(bool,X13)))) | p(s(bool,X12)) | s(bool,'T_0') = s(bool,X14) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f1486,f701])).
fof(f1486,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X2)))) | p(s(bool,X2)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f1469,f592])).
fof(f1469,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X2)))) | p(s(bool,'F_0')) | p(s(bool,X2)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f485,f343])).
fof(f485,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X2)))) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f293])).
fof(f293,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X2)))) | ((~p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2))))) & ((p(s(bool,X0)) & ~p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X2))) | p(s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X2))))))),
  inference(flattening,[],[f292])).
fof(f292,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X2)))) | ((~p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2))))) & (((p(s(bool,X0)) & ~p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X2)))) | p(s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X2))))))),
  inference(nnf_transformation,[],[f157])).
fof(f157,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X2)))) <=> ((~p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f156])).
fof(f156,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X2)))) <=> ((~p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X22,X10,X11] : (p(s(bool,'const.bool.COND_3'(s(bool,X22),s(bool,X10),s(bool,X11)))) <=> ((~p(s(bool,X22)) | p(s(bool,X10))) & (p(s(bool,X22)) | p(s(bool,X11)))))),
  file('thm.arithmetic.MIN_MAX_EQ','thm.bool.COND_EXPAND')).
fof(f49309,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,'const.bool.COND_3'(s(bool,X10),s(bool,X11),s(bool,X9)))) | p(s(bool,X8)) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X9)) )),
  inference(subsumption_resolution,[],[f49262,f354])).
fof(f49262,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,'const.bool.COND_3'(s(bool,X10),s(bool,X11),s(bool,X9)))) | p(s(bool,X8)) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f1486,f701])).
fof(f49122,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',X6))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num',X7)))) ) | $spl106),
  inference(cnf_transformation,[],[f49122_D])).
fof(f49122_D,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',X6))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num',X7)))) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f49123,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X11)) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X10),s('type.num.num',X9))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num',X10)))) )),
  inference(subsumption_resolution,[],[f49101,f354])).
fof(f49101,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X11)) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X10),s('type.num.num',X9))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num',X10)))) )),
  inference(superposition,[],[f1230,f701])).
fof(f1230,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) )),
  inference(subsumption_resolution,[],[f1223,f592])).
fof(f1223,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) )),
  inference(superposition,[],[f467,f343])).
fof(f467,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f286])).
fof(f286,plain,(
  ! [X0,X1] : ((p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) & (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(nnf_transformation,[],[f141])).
fof(f141,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(flattening,[],[f140])).
fof(f140,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X6,X7] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num',X7)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',X6)))))),
  file('thm.arithmetic.MIN_MAX_EQ','thm.arithmetic.NOT_LESS')).
fof(f49121,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X8)) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',X6))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num',X7)))) )),
  inference(subsumption_resolution,[],[f49100,f354])).
fof(f49100,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X8)) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',X6))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num',X7))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f1230,f701])).
fof(f47090,plain,(
  ( ! [X8,X7] : (sP1('const.bool.COND_3'(s(bool,'F_0'),s(bool,X8),s(bool,X7)),'T_0') | s(bool,'T_0') = s(bool,X7)) ) | $spl104),
  inference(cnf_transformation,[],[f47090_D])).
fof(f47090_D,plain,(
  ( ! [X8,X7] : (sP1('const.bool.COND_3'(s(bool,'F_0'),s(bool,X8),s(bool,X7)),'T_0') | s(bool,'T_0') = s(bool,X7)) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f47091,plain,(
  ( ! [X10,X11,X9] : (sP1('const.bool.COND_3'(s(bool,'F_0'),s(bool,X11),s(bool,X10)),'T_0') | p(s(bool,X9)) | s(bool,'T_0') = s(bool,X10)) ) | $spl16),
  inference(subsumption_resolution,[],[f47073,f354])).
fof(f47073,plain,(
  ( ! [X10,X11,X9] : (sP1('const.bool.COND_3'(s(bool,'F_0'),s(bool,X11),s(bool,X10)),'T_0') | p(s(bool,X9)) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X9)) ) | $spl16),
  inference(superposition,[],[f25991,f701])).
fof(f25991,plain,(
  ( ! [X17,X18] : (sP1('const.bool.COND_3'(s(bool,'F_0'),s(bool,X17),s(bool,X18)),'T_0') | p(s(bool,X18))) ) | $spl16),
  inference(superposition,[],[f25841,f414])).
fof(f47089,plain,(
  ( ! [X6,X8,X7] : (sP1('const.bool.COND_3'(s(bool,'F_0'),s(bool,X8),s(bool,X7)),'T_0') | p(s(bool,X6)) | s(bool,'T_0') = s(bool,X7)) ) | $spl16),
  inference(subsumption_resolution,[],[f47072,f354])).
fof(f47072,plain,(
  ( ! [X6,X8,X7] : (sP1('const.bool.COND_3'(s(bool,'F_0'),s(bool,X8),s(bool,X7)),'T_0') | p(s(bool,X6)) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,X7)) ) | $spl16),
  inference(superposition,[],[f25991,f701])).
fof(f46952,plain,(
  ( ! [X8,X7] : (sP1('const.bool.COND_3'(s(bool,'T_0'),s(bool,X7),s(bool,X8)),'T_0') | s(bool,'T_0') = s(bool,X7)) ) | $spl102),
  inference(cnf_transformation,[],[f46952_D])).
fof(f46952_D,plain,(
  ( ! [X8,X7] : (sP1('const.bool.COND_3'(s(bool,'T_0'),s(bool,X7),s(bool,X8)),'T_0') | s(bool,'T_0') = s(bool,X7)) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f46953,plain,(
  ( ! [X10,X11,X9] : (sP1('const.bool.COND_3'(s(bool,'T_0'),s(bool,X10),s(bool,X11)),'T_0') | p(s(bool,X9)) | s(bool,'T_0') = s(bool,X10)) ) | $spl16),
  inference(subsumption_resolution,[],[f46918,f354])).
fof(f46918,plain,(
  ( ! [X10,X11,X9] : (sP1('const.bool.COND_3'(s(bool,'T_0'),s(bool,X10),s(bool,X11)),'T_0') | p(s(bool,X9)) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X9)) ) | $spl16),
  inference(superposition,[],[f25990,f701])).
fof(f25990,plain,(
  ( ! [X15,X16] : (sP1('const.bool.COND_3'(s(bool,'T_0'),s(bool,X15),s(bool,X16)),'T_0') | p(s(bool,X15))) ) | $spl16),
  inference(superposition,[],[f25841,f413])).
fof(f46951,plain,(
  ( ! [X6,X8,X7] : (sP1('const.bool.COND_3'(s(bool,'T_0'),s(bool,X7),s(bool,X8)),'T_0') | p(s(bool,X6)) | s(bool,'T_0') = s(bool,X7)) ) | $spl16),
  inference(subsumption_resolution,[],[f46917,f354])).
fof(f46917,plain,(
  ( ! [X6,X8,X7] : (sP1('const.bool.COND_3'(s(bool,'T_0'),s(bool,X7),s(bool,X8)),'T_0') | p(s(bool,X6)) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,X7)) ) | $spl16),
  inference(superposition,[],[f25990,f701])).
fof(f46820,plain,(
  ( ! [X8,X7] : (sP1('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X8),s(bool,X7))) | s(bool,'T_0') = s(bool,X7)) ) | $spl100),
  inference(cnf_transformation,[],[f46820_D])).
fof(f46820_D,plain,(
  ( ! [X8,X7] : (sP1('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X8),s(bool,X7))) | s(bool,'T_0') = s(bool,X7)) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f46821,plain,(
  ( ! [X10,X11,X9] : (sP1('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X11),s(bool,X10))) | p(s(bool,X9)) | s(bool,'T_0') = s(bool,X10)) ) | $spl14),
  inference(subsumption_resolution,[],[f46803,f354])).
fof(f46803,plain,(
  ( ! [X10,X11,X9] : (sP1('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X11),s(bool,X10))) | p(s(bool,X9)) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X9)) ) | $spl14),
  inference(superposition,[],[f25971,f701])).
fof(f25971,plain,(
  ( ! [X17,X18] : (sP1('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X17),s(bool,X18))) | p(s(bool,X18))) ) | $spl14),
  inference(superposition,[],[f25838,f414])).
fof(f46819,plain,(
  ( ! [X6,X8,X7] : (sP1('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X8),s(bool,X7))) | p(s(bool,X6)) | s(bool,'T_0') = s(bool,X7)) ) | $spl14),
  inference(subsumption_resolution,[],[f46802,f354])).
fof(f46802,plain,(
  ( ! [X6,X8,X7] : (sP1('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X8),s(bool,X7))) | p(s(bool,X6)) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,X7)) ) | $spl14),
  inference(superposition,[],[f25971,f701])).
fof(f28361,plain,(
  ( ! [X642] : (s(bool,'F_0') = s(bool,X642)) ) | $spl50),
  inference(cnf_transformation,[],[f28361_D])).
fof(f28361_D,plain,(
  ( ! [X642] : (s(bool,'F_0') = s(bool,X642)) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f28366,plain,(
  ( ! [X628] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X628),s(bool,'T_0')))) ) | $spl60),
  inference(cnf_transformation,[],[f28366_D])).
fof(f28366_D,plain,(
  ( ! [X628] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X628),s(bool,'T_0')))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f28417,plain,(
  ( ! [X26,X24] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X26),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X24)) )),
  inference(global_subsumption,[],[f334,f333,f335,f343,f355,f354,f413,f428,f427,f426,f429,f430,f437,f434,f444,f443,f593,f462,f459,f458,f457,f478,f479,f480,f481,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f819,f825,f828,f830,f840,f888,f898,f901,f917,f923,f933,f939,f946,f954,f960,f966,f971,f975,f980,f984,f994,f1000,f1007,f1012,f1024,f1032,f1047,f1050,f1062,f1068,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1239,f1275,f1339,f1332,f1373,f1385,f1402,f1422,f1433,f1437,f1461,f1501,f1512,f1516,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1662,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1981,f2015,f2030,f2073,f2082,f2119,f2118,f2128,f2157,f2170,f2252,f2222,f2262,f2270,f2617,f2622,f2841,f2844,f2751,f2544,f2504,f2465,f2432,f2400,f2371,f2351,f2317,f1837,f1818,f424,f423,f422,f421,f420,f419,f418,f417,f412,f411,f1854,f1861,f1852,f1859,f2052,f2050,f2043,f2041,f482,f433,f1542,f1492,f1486,f489,f488,f487,f485,f484,f1219,f1249,f1217,f1247,f1212,f1256,f1285,f1210,f1254,f1283,f2099,f2097,f2090,f2088,f2001,f1931,f1230,f483,f477,f468,f467,f432,f431,f2848,f2847,f2770,f2765,f2545,f2505,f2466,f2433,f2401,f2398,f2373,f2372,f2352,f2349,f2319,f2318,f2272,f2271,f2269,f2251,f1838,f1835,f1819,f1816,f1663,f1661,f1494,f1411,f1410,f1403,f1366,f1365,f1333,f1267,f1231,f1069,f1059,f1057,f1048,f1035,f1033,f1025,f1023,f1001,f991,f961,f942,f940,f930,f926,f924,f904,f902,f891,f889,f873,f871,f870,f869,f868,f867,f866,f865,f864,f863,f862,f861,f852,f839,f838,f837,f836,f829,f818,f787,f592,f414,f359,f358,f4838,f4843,f4844,f4853,f4855,f4857,f4858,f4859,f4860,f4861,f4862,f4863,f4864,f4644,f4646,f4648,f4650,f4870,f4655,f4658,f4660,f4661,f4663,f4669,f4977,f4986,f5268,f5276,f7384,f7394,f7492,f7497,f7506,f7666,f7671,f7680,f7351,f11643,f7360,f11657,f11665,f12884,f12883,f12097,f12096,f13263,f13282,f13283,f13284,f13285,f13298,f13312,f13313,f13314,f13315,f13323,f13324,f13325,f13326,f13327,f13328,f13329,f13331,f13333,f13335,f13337,f13343,f13346,f13348,f13349,f13351,f13356,f13358,f13361,f16694,f16695,f13454,f16696,f13456,f16697,f16698,f16699,f16700,f16701,f16717,f16728,f18264,f18263,f17257,f17256,f18761,f18782,f18783,f18784,f18785,f18786,f18802,f18816,f18817,f18818,f18819,f18828,f18829,f18830,f18831,f18832,f18833,f18834,f18836,f18838,f18840,f18842,f18848,f18851,f18853,f18854,f18856,f18861,f18863,f18866,f23086,f23087,f18961,f23088,f18963,f23089,f23090,f23091,f23092,f23093,f23129,f23139,f23156,f23157,f23158,f23159,f23173,f23183,f23184,f23185,f23186,f25796,f23209,f23235,f23237,f23239,f23241,f23247,f23250,f23252,f23253,f23255,f23261,f23263,f23266,f25806,f23476,f23477,f23508,f23513,f23514,f23515,f23516,f23517,f23518,f23519,f23520,f23521,f23522,f23523,f25824,f23848,f23849,f23850,f23851,f23852,f23853,f23854,f23855,f23856,f23857,f23858,f23859,f23860,f23861,f23862,f23863,f23864,f23865,f23866,f23867,f23868,f23869,f23870,f23871,f23872,f23873,f23874,f23875,f23876,f23877,f23878,f23879,f23880,f23881,f23882,f23883,f23884,f23885,f23886,f23887,f23888,f23889,f23890,f23891,f23892,f23893,f23894,f23895,f23896,f23897,f23898,f23899,f23900,f23901,f23902,f23903,f23904,f23905,f23906,f23907,f23908,f23909,f23910,f23911,f23912,f23913,f23914,f23915,f23916,f23917,f23918,f23919,f23920,f23921,f23922,f23923,f23924,f23925,f23926,f23927,f23928,f23929,f23930,f23931,f23932,f23933,f23934,f23935,f23936,f23937,f23938,f23939,f25830,f23947,f25837,f25840,f25846,f25849,f25852,f25855,f25858,f25861,f25864,f23973,f23975,f23977,f23979,f23985,f23988,f23990,f23991,f23993,f24001,f25872,f25873,f24118,f25875,f25876,f25877,f25878,f25879,f25880,f25883,f25884,f25885,f25886,f25887,f25888,f25889,f25890,f25891,f24144,f24146,f24148,f24150,f25892,f25893,f25894,f25895,f24156,f24159,f24161,f24162,f24164,f25896,f24170,f24172,f24175,f25897,f25898,f25900,f25901,f25902,f25903,f25904,f25905,f25906,f25907,f25908,f25909,f25910,f25911,f25912,f25913,f25914,f25915,f25916,f25917,f25918,f24416,f24421,f24422,f24423,f24424,f24425,f24426,f24427,f24428,f24429,f24430,f24431,f25920,f25926,f25928,f25933,f25935,f25937,f25939,f25941,f25943,f25945,f25998,f26013,f26015,f26018,f26019,f26009,f26010,f26011,f26012,f26024,f26039,f26041,f26044,f26045,f26035,f26036,f26037,f26038,f26053,f26068,f26070,f26073,f26075,f26090,f26105,f26107,f26110,f26111,f26101,f26102,f26103,f26104,f26118,f26133,f26135,f26138,f26139,f26129,f26130,f26131,f26132,f26141,f26160,f26178,f26193,f26195,f26198,f26199,f26189,f26190,f26191,f26192,f26203,f26233,f26260,f26275,f26277,f26280,f26282,f26290,f26320,f26347,f26362,f26364,f26367,f26368,f26358,f26359,f26360,f26361,f26372,f26391,f26408,f26427,f26442,f26458,f26459,f26460,f26461,f26709,f26711,f26714,f26717,f26718,f26719,f26720,f26721,f26723,f26731,f26733,f26735,f26736,f26737,f26738,f26739,f26740,f26741,f26742,f26521,f26523,f26525,f26527,f26751,f26752,f26755,f26537,f26540,f26542,f26543,f26545,f26757,f26759,f26761,f26763,f26764,f26765,f26767,f26768,f26555,f26557,f26560,f26772,f26777,f26783,f26803,f26820,f26840,f26865,f26866,f26867,f26868,f26869,f26921,f26923,f26925,f26927,f26928,f26948,f26949,f26950,f26951,f26952,f26953,f26954,f26955,f26957,f26959,f26960,f26961,f26962,f26963,f26966,f26967,f26968,f26970,f26971,f26972,f26973,f26974,f26975,f26976,f26977,f26978,f26979,f26984,f26985,f26987,f26989,f26991,f26993,f26995,f26996,f26997,f26998,f26999,f27000,f27001,f27002,f27003,f27006,f27008,f27009,f27011,f27013,f27014,f27015,f27016,f27017,f27018,f27019,f27020,f27021,f27023,f27026,f27030,f27033,f27043,f27106,f27107,f27109,f27119,f27132,f27133,f27134,f27135,f27136,f27137,f27138,f27139,f27140,f27141,f27155,f27156,f27157,f27158,f27159,f27160,f27161,f27162,f27166,f27167,f27168,f27169,f27199,f27298,f27533,f27532,f27531,f27530,f27529,f27528,f27524,f27523,f27506,f28372,f28376,f28377,f27431,f27429,f27428,f27425,f27424,f27420,f27417,f27413,f27412,f27411,f27410,f27408,f28347,f28391,f28344,f28341,f28392,f28393,f28338,f28335,f27399,f28394,f27398,f27397,f27396,f27395,f27394,f27393,f27392,f27391,f27389,f28399,f28400,f28401,f28402,f28403,f28404,f28405,f28406,f27539,f27377,f27375,f27373,f27371,f28310,f28307,f28304,f28407,f28301,f28298,f28408,f28409,f28410,f28411,f27357,f28413,f28414,f28415,f28296,f28294,f28290,f27344,f28280,f28278,f27324,f27323,f27322,f27321,f27320,f28416,f27318])).
fof(f27318,plain,(
  ( ! [X26,X24,X25] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X26),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X24),s(bool,X25))) | s(bool,'F_0') = s(bool,X24)) )),
  inference(superposition,[],[f852,f787])).
fof(f27320,plain,(
  ( ! [X30,X31,X32] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X32),s(bool,X30))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X30),s(bool,X31)))) )),
  inference(superposition,[],[f852,f413])).
fof(f27321,plain,(
  ( ! [X35,X33,X34] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X35),s(bool,X34))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X33),s(bool,X34)))) )),
  inference(superposition,[],[f852,f414])).
fof(f27322,plain,(
  ( ! [X39,X37,X41,X38,X36,X40] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X41),s(bool,happ(s(fun(X37,bool),X38),s(X37,'const.bool.COND_3'(s(bool,X36),s(X37,X39),s(X37,X40))))))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X36),s(bool,happ(s(fun(X37,bool),X38),s(X37,X39))),s(bool,happ(s(fun(X37,bool),X38),s(X37,X40)))))) )),
  inference(superposition,[],[f852,f479])).
fof(f27323,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X47),s(bool,happ(s(fun(X43,bool),'const.bool.COND_3'(s(bool,X42),s(fun(X43,bool),X44),s(fun(X43,bool),X46))),s(X43,X45))))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X42),s(bool,happ(s(fun(X43,bool),X44),s(X43,X45))),s(bool,happ(s(fun(X43,bool),X46),s(X43,X45)))))) )),
  inference(superposition,[],[f852,f480])).
fof(f27324,plain,(
  ( ! [X50,X48,X51,X49] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X51),s(bool,X50))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X48),s(bool,X49),s(bool,X50))) | s(bool,'T_0') = s(bool,X48)) )),
  inference(superposition,[],[f852,f840])).
fof(f28278,plain,(
  ( ! [X10,X11] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X10),s(bool,X11))) | s(bool,'F_0') = s(bool,X11)) )),
  inference(subsumption_resolution,[],[f27338,f335])).
fof(f27338,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X10),s(bool,X11))) | s(bool,'F_0') = s(bool,X11)) )),
  inference(superposition,[],[f358,f852])).
fof(f28280,plain,(
  ( ! [X14,X21,X19,X17,X15,X20,X18,X16] : (s(X16,'const.bool.COND_3'(s(bool,X17),s(X16,X18),s(X16,X19))) = s(X16,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X14),s(bool,X15))),s(X16,X20),s(X16,X21))) | s(X16,X18) != s(X16,X20) | s(bool,X17) != s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X14),s(bool,X15))) | s(bool,'F_0') = s(bool,X15)) )),
  inference(subsumption_resolution,[],[f27340,f335])).
fof(f27340,plain,(
  ( ! [X14,X21,X19,X17,X15,X20,X18,X16] : (~p(s(bool,'T_0')) | s(X16,'const.bool.COND_3'(s(bool,X17),s(X16,X18),s(X16,X19))) = s(X16,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X14),s(bool,X15))),s(X16,X20),s(X16,X21))) | s(X16,X18) != s(X16,X20) | s(bool,X17) != s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X14),s(bool,X15))) | s(bool,'F_0') = s(bool,X15)) )),
  inference(superposition,[],[f426,f852])).
fof(f27344,plain,(
  ( ! [X47,X52,X50,X48,X46,X53,X51,X49] : (s(bool,'T_0') != s(bool,X48) | s(X49,X50) != s(X49,X51) | s(X49,X52) != s(X49,X53) | s(X49,'const.bool.COND_3'(s(bool,X48),s(X49,X53),s(X49,X51))) = s(X49,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X46),s(bool,X47))),s(X49,X52),s(X49,X50))) | s(bool,'F_0') = s(bool,X47)) )),
  inference(superposition,[],[f428,f852])).
fof(f28290,plain,(
  ( ! [X61,X62,X60] : (~p(s(bool,X62)) | s(bool,X62) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X60),s(bool,X61))) | s(bool,'F_0') = s(bool,X61)) )),
  inference(subsumption_resolution,[],[f27347,f335])).
fof(f27347,plain,(
  ( ! [X61,X62,X60] : (~p(s(bool,'T_0')) | ~p(s(bool,X62)) | s(bool,X62) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X60),s(bool,X61))) | s(bool,'F_0') = s(bool,X61)) )),
  inference(superposition,[],[f437,f852])).
fof(f28294,plain,(
  ( ! [X64,X65,X63] : (sP0(X65,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X63),s(bool,X64))) | s(bool,'F_0') = s(bool,X64)) )),
  inference(subsumption_resolution,[],[f27348,f335])).
fof(f27348,plain,(
  ( ! [X64,X65,X63] : (~p(s(bool,'T_0')) | sP0(X65,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X63),s(bool,X64))) | s(bool,'F_0') = s(bool,X64)) )),
  inference(superposition,[],[f443,f852])).
fof(f28296,plain,(
  ( ! [X68,X66,X67] : (sP0('const.bool.COND_3'(s(bool,'F_0'),s(bool,X66),s(bool,X67)),X68) | s(bool,'F_0') = s(bool,X67)) )),
  inference(subsumption_resolution,[],[f27349,f335])).
fof(f27349,plain,(
  ( ! [X68,X66,X67] : (~p(s(bool,'T_0')) | sP0('const.bool.COND_3'(s(bool,'F_0'),s(bool,X66),s(bool,X67)),X68) | s(bool,'F_0') = s(bool,X67)) )),
  inference(superposition,[],[f444,f852])).
fof(f28415,plain,(
  ( ! [X70,X71,X69] : (p(s(bool,X71)) | sP1('const.bool.COND_3'(s(bool,'F_0'),s(bool,X69),s(bool,X70)),X71) | s(bool,'F_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f27350,f335])).
fof(f27350,plain,(
  ( ! [X70,X71,X69] : (~p(s(bool,'T_0')) | p(s(bool,X71)) | sP1('const.bool.COND_3'(s(bool,'F_0'),s(bool,X69),s(bool,X70)),X71) | s(bool,'F_0') = s(bool,X70)) )),
  inference(superposition,[],[f459,f852])).
fof(f28414,plain,(
  ( ! [X74,X72,X73] : (p(s(bool,X74)) | sP1(X74,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X72),s(bool,X73))) | s(bool,'F_0') = s(bool,X73)) )),
  inference(subsumption_resolution,[],[f27351,f335])).
fof(f27351,plain,(
  ( ! [X74,X72,X73] : (~p(s(bool,'T_0')) | p(s(bool,X74)) | sP1(X74,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X72),s(bool,X73))) | s(bool,'F_0') = s(bool,X73)) )),
  inference(superposition,[],[f462,f852])).
fof(f28413,plain,(
  ( ! [X92,X90,X91,X89] : (p(s(bool,X91)) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X89),s(bool,X90))),s(bool,X91),s(bool,X92)))) | s(bool,'F_0') = s(bool,X90)) )),
  inference(subsumption_resolution,[],[f27354,f335])).
fof(f27354,plain,(
  ( ! [X92,X90,X91,X89] : (~p(s(bool,'T_0')) | p(s(bool,X91)) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X89),s(bool,X90))),s(bool,X91),s(bool,X92)))) | s(bool,'F_0') = s(bool,X90)) )),
  inference(superposition,[],[f484,f852])).
fof(f28411,plain,(
  ( ! [X107,X105,X108,X106] : (~p(s(bool,X107)) | p(s(bool,'const.bool.COND_3'(s(bool,X107),s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X105),s(bool,X106))),s(bool,X108)))) | s(bool,'F_0') = s(bool,X106)) )),
  inference(subsumption_resolution,[],[f27358,f335])).
fof(f27358,plain,(
  ( ! [X107,X105,X108,X106] : (~p(s(bool,'T_0')) | ~p(s(bool,X107)) | p(s(bool,'const.bool.COND_3'(s(bool,X107),s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X105),s(bool,X106))),s(bool,X108)))) | s(bool,'F_0') = s(bool,X106)) )),
  inference(superposition,[],[f487,f852])).
fof(f28410,plain,(
  ( ! [X111,X109,X112,X110] : (p(s(bool,'const.bool.COND_3'(s(bool,X111),s(bool,X112),s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X109),s(bool,X110)))))) | p(s(bool,X111)) | s(bool,'F_0') = s(bool,X110)) )),
  inference(subsumption_resolution,[],[f27359,f335])).
fof(f27359,plain,(
  ( ! [X111,X109,X112,X110] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.COND_3'(s(bool,X111),s(bool,X112),s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X109),s(bool,X110)))))) | p(s(bool,X111)) | s(bool,'F_0') = s(bool,X110)) )),
  inference(superposition,[],[f488,f852])).
fof(f28409,plain,(
  ( ! [X116,X114,X115,X113] : (~p(s(bool,X115)) | p(s(bool,'const.bool.COND_3'(s(bool,X116),s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X113),s(bool,X114))),s(bool,X115)))) | s(bool,'F_0') = s(bool,X114)) )),
  inference(subsumption_resolution,[],[f27360,f335])).
fof(f27360,plain,(
  ( ! [X116,X114,X115,X113] : (~p(s(bool,'T_0')) | ~p(s(bool,X115)) | p(s(bool,'const.bool.COND_3'(s(bool,X116),s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X113),s(bool,X114))),s(bool,X115)))) | s(bool,'F_0') = s(bool,X114)) )),
  inference(superposition,[],[f489,f852])).
fof(f28408,plain,(
  ( ! [X118,X120,X119,X117] : (~p(s(bool,X119)) | p(s(bool,X120)) | ~sP7(X120,X119,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X117),s(bool,X118))) | s(bool,'F_0') = s(bool,X118)) )),
  inference(subsumption_resolution,[],[f27361,f335])).
fof(f27361,plain,(
  ( ! [X118,X120,X119,X117] : (~p(s(bool,'T_0')) | ~p(s(bool,X119)) | p(s(bool,X120)) | ~sP7(X120,X119,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X117),s(bool,X118))) | s(bool,'F_0') = s(bool,X118)) )),
  inference(superposition,[],[f541,f852])).
fof(f28298,plain,(
  ( ! [X123,X121,X124,X122] : (~p(s(bool,X123)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X121),s(bool,X122)),X124,X123) | s(bool,'F_0') = s(bool,X122)) )),
  inference(subsumption_resolution,[],[f27362,f335])).
fof(f27362,plain,(
  ( ! [X123,X121,X124,X122] : (~p(s(bool,'T_0')) | ~p(s(bool,X123)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X121),s(bool,X122)),X124,X123) | s(bool,'F_0') = s(bool,X122)) )),
  inference(superposition,[],[f554,f852])).
fof(f28301,plain,(
  ( ! [X127,X125,X128,X126] : (~p(s(bool,X127)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X125),s(bool,X126)),X127,X128) | s(bool,'F_0') = s(bool,X126)) )),
  inference(subsumption_resolution,[],[f27363,f335])).
fof(f27363,plain,(
  ( ! [X127,X125,X128,X126] : (~p(s(bool,'T_0')) | ~p(s(bool,X127)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X125),s(bool,X126)),X127,X128) | s(bool,'F_0') = s(bool,X126)) )),
  inference(superposition,[],[f555,f852])).
fof(f28407,plain,(
  ( ! [X132,X130,X131,X129] : (p(s(bool,X131)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X129),s(bool,X130)),X132,X131) | s(bool,'F_0') = s(bool,X130)) )),
  inference(subsumption_resolution,[],[f27364,f335])).
fof(f27364,plain,(
  ( ! [X132,X130,X131,X129] : (~p(s(bool,'T_0')) | p(s(bool,X131)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X129),s(bool,X130)),X132,X131) | s(bool,'F_0') = s(bool,X130)) )),
  inference(superposition,[],[f571,f852])).
fof(f28304,plain,(
  ( ! [X134,X136,X135,X133] : (~p(s(bool,X135)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X133),s(bool,X134)),X135,X136) | s(bool,'F_0') = s(bool,X134)) )),
  inference(subsumption_resolution,[],[f27365,f335])).
fof(f27365,plain,(
  ( ! [X134,X136,X135,X133] : (~p(s(bool,'T_0')) | ~p(s(bool,X135)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X133),s(bool,X134)),X135,X136) | s(bool,'F_0') = s(bool,X134)) )),
  inference(superposition,[],[f572,f852])).
fof(f28307,plain,(
  ( ! [X140,X138,X139,X137] : (~sP10('const.bool.COND_3'(s(bool,'F_0'),s(bool,X137),s(bool,X138)),X139,X140) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X137),s(bool,X138)),X139,X140) | s(bool,'F_0') = s(bool,X138)) )),
  inference(subsumption_resolution,[],[f27366,f335])).
fof(f27366,plain,(
  ( ! [X140,X138,X139,X137] : (~p(s(bool,'T_0')) | ~sP10('const.bool.COND_3'(s(bool,'F_0'),s(bool,X137),s(bool,X138)),X139,X140) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X137),s(bool,X138)),X139,X140) | s(bool,'F_0') = s(bool,X138)) )),
  inference(superposition,[],[f576,f852])).
fof(f28310,plain,(
  ( ! [X144,X142,X143,X141] : (~sP8('const.bool.COND_3'(s(bool,'F_0'),s(bool,X141),s(bool,X142)),X143,X144) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X141),s(bool,X142)),X143,X144) | s(bool,'F_0') = s(bool,X142)) )),
  inference(subsumption_resolution,[],[f27367,f335])).
fof(f27367,plain,(
  ( ! [X144,X142,X143,X141] : (~p(s(bool,'T_0')) | ~sP8('const.bool.COND_3'(s(bool,'F_0'),s(bool,X141),s(bool,X142)),X143,X144) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X141),s(bool,X142)),X143,X144) | s(bool,'F_0') = s(bool,X142)) )),
  inference(superposition,[],[f577,f852])).
fof(f27371,plain,(
  ( ! [X158,X159,X160] : (~sP5('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X160)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X158),s(bool,X159))) | s(bool,'F_0') = s(bool,X159)) )),
  inference(superposition,[],[f775,f852])).
fof(f27373,plain,(
  ( ! [X165,X166,X164] : (sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X166)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X164),s(bool,X165))) | s(bool,'F_0') = s(bool,X165)) )),
  inference(superposition,[],[f776,f852])).
fof(f27375,plain,(
  ( ! [X171,X172,X170] : (~sP5('F_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X170),s(bool,X171)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X172))) | s(bool,'F_0') = s(bool,X171)) )),
  inference(superposition,[],[f777,f852])).
fof(f27377,plain,(
  ( ! [X177,X178,X176] : (sP5('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X176),s(bool,X177)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X178))) | s(bool,'F_0') = s(bool,X177)) )),
  inference(superposition,[],[f778,f852])).
fof(f27539,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,'T_0')))) )),
  inference(condensation,[],[f27380])).
fof(f27380,plain,(
  ( ! [X187,X186] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X187),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X187) | s(bool,'F_0') = s(bool,X186)) )),
  inference(superposition,[],[f787,f852])).
fof(f28406,plain,(
  ( ! [X191,X189,X190,X188] : (sP0(X190,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X188),s(bool,X189))),s(bool,X191))) | s(bool,'F_0') = s(bool,X189)) )),
  inference(subsumption_resolution,[],[f27381,f335])).
fof(f27381,plain,(
  ( ! [X191,X189,X190,X188] : (~p(s(bool,'T_0')) | sP0(X190,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X188),s(bool,X189))),s(bool,X191))) | s(bool,'F_0') = s(bool,X189)) )),
  inference(superposition,[],[f796,f852])).
fof(f28405,plain,(
  ( ! [X194,X192,X195,X193] : (sP0('const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X192),s(bool,X193))),s(bool,X194)),X195) | s(bool,'F_0') = s(bool,X193)) )),
  inference(subsumption_resolution,[],[f27382,f335])).
fof(f27382,plain,(
  ( ! [X194,X192,X195,X193] : (~p(s(bool,'T_0')) | sP0('const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X192),s(bool,X193))),s(bool,X194)),X195) | s(bool,'F_0') = s(bool,X193)) )),
  inference(superposition,[],[f797,f852])).
fof(f28404,plain,(
  ( ! [X198,X196,X199,X197] : (p(s(bool,X198)) | sP1('const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X196),s(bool,X197))),s(bool,X199)),X198) | s(bool,'F_0') = s(bool,X197)) )),
  inference(subsumption_resolution,[],[f27383,f335])).
fof(f27383,plain,(
  ( ! [X198,X196,X199,X197] : (~p(s(bool,'T_0')) | p(s(bool,X198)) | sP1('const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X196),s(bool,X197))),s(bool,X199)),X198) | s(bool,'F_0') = s(bool,X197)) )),
  inference(superposition,[],[f798,f852])).
fof(f28403,plain,(
  ( ! [X202,X200,X203,X201] : (p(s(bool,X202)) | sP1(X202,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X200),s(bool,X201))),s(bool,X203))) | s(bool,'F_0') = s(bool,X201)) )),
  inference(subsumption_resolution,[],[f27384,f335])).
fof(f27384,plain,(
  ( ! [X202,X200,X203,X201] : (~p(s(bool,'T_0')) | p(s(bool,X202)) | sP1(X202,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X200),s(bool,X201))),s(bool,X203))) | s(bool,'F_0') = s(bool,X201)) )),
  inference(superposition,[],[f799,f852])).
fof(f28402,plain,(
  ( ! [X208,X206,X204,X207,X205] : (~p(s(bool,X206)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X206),s(bool,X207)),X208,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X204),s(bool,X205))) | s(bool,'F_0') = s(bool,X205)) )),
  inference(subsumption_resolution,[],[f27385,f335])).
fof(f27385,plain,(
  ( ! [X208,X206,X204,X207,X205] : (~p(s(bool,'T_0')) | ~p(s(bool,X206)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X206),s(bool,X207)),X208,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X204),s(bool,X205))) | s(bool,'F_0') = s(bool,X205)) )),
  inference(superposition,[],[f800,f852])).
fof(f28401,plain,(
  ( ! [X212,X210,X213,X211,X209] : (~p(s(bool,X211)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X211),s(bool,X212)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X209),s(bool,X210)),X213) | s(bool,'F_0') = s(bool,X210)) )),
  inference(subsumption_resolution,[],[f27386,f335])).
fof(f27386,plain,(
  ( ! [X212,X210,X213,X211,X209] : (~p(s(bool,'T_0')) | ~p(s(bool,X211)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X211),s(bool,X212)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X209),s(bool,X210)),X213) | s(bool,'F_0') = s(bool,X210)) )),
  inference(superposition,[],[f801,f852])).
fof(f28400,plain,(
  ( ! [X218,X216,X215,X217,X214] : (p(s(bool,X216)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X214),s(bool,X215))),s(bool,X217)),X218,X216) | s(bool,'F_0') = s(bool,X215)) )),
  inference(subsumption_resolution,[],[f27387,f335])).
fof(f27387,plain,(
  ( ! [X218,X216,X215,X217,X214] : (~p(s(bool,'T_0')) | p(s(bool,X216)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X214),s(bool,X215))),s(bool,X217)),X218,X216) | s(bool,'F_0') = s(bool,X215)) )),
  inference(superposition,[],[f802,f852])).
fof(f28399,plain,(
  ( ! [X222,X220,X219,X223,X221] : (~p(s(bool,X221)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X221),s(bool,X222)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X219),s(bool,X220)),X223) | s(bool,'F_0') = s(bool,X220)) )),
  inference(subsumption_resolution,[],[f27388,f335])).
fof(f27388,plain,(
  ( ! [X222,X220,X219,X223,X221] : (~p(s(bool,'T_0')) | ~p(s(bool,X221)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X221),s(bool,X222)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X219),s(bool,X220)),X223) | s(bool,'F_0') = s(bool,X220)) )),
  inference(superposition,[],[f803,f852])).
fof(f27389,plain,(
  ( ! [X225,X226,X224] : (sP4('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X224),s(bool,X225)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X226))) | s(bool,'F_0') = s(bool,X225)) )),
  inference(superposition,[],[f825,f852])).
fof(f27392,plain,(
  ( ! [X235,X233,X234] : (~sP5('F_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X235),s(bool,'T_0')),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X233),s(bool,X234))) | s(bool,'F_0') = s(bool,X234)) )),
  inference(superposition,[],[f836,f852])).
fof(f27394,plain,(
  ( ! [X241,X239,X240] : (sP5('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X241),s(bool,'T_0')),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X239),s(bool,X240))) | s(bool,'F_0') = s(bool,X240)) )),
  inference(superposition,[],[f837,f852])).
fof(f27395,plain,(
  ( ! [X243,X244,X242] : (~sP5('F_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X242),s(bool,X243)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X244),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X243)) )),
  inference(superposition,[],[f838,f852])).
fof(f27397,plain,(
  ( ! [X249,X250,X248] : (sP5('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X248),s(bool,X249)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X250),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X249)) )),
  inference(superposition,[],[f839,f852])).
fof(f28394,plain,(
  ( ! [X255,X254] : (s(bool,'F_0') = s(bool,X255) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X254),s(bool,X255)))) )),
  inference(global_subsumption,[],[f334,f333,f335,f343,f355,f354,f413,f428,f427,f426,f429,f430,f437,f434,f444,f443,f593,f462,f459,f458,f457,f478,f479,f480,f481,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f819,f825,f828,f830,f840,f888,f898,f901,f917,f923,f933,f939,f946,f954,f960,f966,f971,f975,f980,f984,f994,f1000,f1007,f1012,f1024,f1032,f1047,f1050,f1062,f1068,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1239,f1275,f1339,f1332,f1373,f1385,f1402,f1422,f1433,f1437,f1461,f1501,f1512,f1516,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1662,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1981,f2015,f2030,f2073,f2082,f2119,f2118,f2128,f2157,f2170,f2252,f2222,f2262,f2270,f2617,f2622,f2841,f2844,f2751,f2544,f2504,f2465,f2432,f2400,f2371,f2351,f2317,f1837,f1818,f424,f423,f422,f421,f420,f419,f418,f417,f412,f411,f1854,f1861,f1852,f1859,f2052,f2050,f2043,f2041,f482,f433,f1542,f1492,f1486,f489,f488,f487,f485,f484,f1219,f1249,f1217,f1247,f1212,f1256,f1285,f1210,f1254,f1283,f2099,f2097,f2090,f2088,f2001,f1931,f1230,f483,f477,f468,f467,f432,f431,f2848,f2847,f2770,f2765,f2545,f2505,f2466,f2433,f2401,f2398,f2373,f2372,f2352,f2349,f2319,f2318,f2272,f2271,f2269,f2251,f1838,f1835,f1819,f1816,f1663,f1661,f1494,f1411,f1410,f1403,f1366,f1365,f1333,f1267,f1231,f1069,f1059,f1057,f1048,f1035,f1033,f1025,f1023,f1001,f991,f961,f942,f940,f930,f926,f924,f904,f902,f891,f889,f873,f871,f870,f869,f868,f867,f866,f865,f864,f863,f862,f861,f852,f839,f838,f837,f836,f829,f818,f787,f592,f414,f359,f358,f4838,f4843,f4844,f4853,f4855,f4857,f4858,f4859,f4860,f4861,f4862,f4863,f4864,f4644,f4646,f4648,f4650,f4870,f4655,f4658,f4660,f4661,f4663,f4669,f4977,f4986,f5268,f5276,f7384,f7394,f7492,f7497,f7506,f7666,f7671,f7680,f7351,f11643,f7360,f11657,f11665,f12884,f12883,f12097,f12096,f13263,f13282,f13283,f13284,f13285,f13298,f13312,f13313,f13314,f13315,f13323,f13324,f13325,f13326,f13327,f13328,f13329,f13331,f13333,f13335,f13337,f13343,f13346,f13348,f13349,f13351,f13356,f13358,f13361,f16694,f16695,f13454,f16696,f13456,f16697,f16698,f16699,f16700,f16701,f16717,f16728,f18264,f18263,f17257,f17256,f18761,f18782,f18783,f18784,f18785,f18786,f18802,f18816,f18817,f18818,f18819,f18828,f18829,f18830,f18831,f18832,f18833,f18834,f18836,f18838,f18840,f18842,f18848,f18851,f18853,f18854,f18856,f18861,f18863,f18866,f23086,f23087,f18961,f23088,f18963,f23089,f23090,f23091,f23092,f23093,f23129,f23139,f23156,f23157,f23158,f23159,f23173,f23183,f23184,f23185,f23186,f25796,f23209,f23235,f23237,f23239,f23241,f23247,f23250,f23252,f23253,f23255,f23261,f23263,f23266,f25806,f23476,f23477,f23508,f23513,f23514,f23515,f23516,f23517,f23518,f23519,f23520,f23521,f23522,f23523,f25824,f23848,f23849,f23850,f23851,f23852,f23853,f23854,f23855,f23856,f23857,f23858,f23859,f23860,f23861,f23862,f23863,f23864,f23865,f23866,f23867,f23868,f23869,f23870,f23871,f23872,f23873,f23874,f23875,f23876,f23877,f23878,f23879,f23880,f23881,f23882,f23883,f23884,f23885,f23886,f23887,f23888,f23889,f23890,f23891,f23892,f23893,f23894,f23895,f23896,f23897,f23898,f23899,f23900,f23901,f23902,f23903,f23904,f23905,f23906,f23907,f23908,f23909,f23910,f23911,f23912,f23913,f23914,f23915,f23916,f23917,f23918,f23919,f23920,f23921,f23922,f23923,f23924,f23925,f23926,f23927,f23928,f23929,f23930,f23931,f23932,f23933,f23934,f23935,f23936,f23937,f23938,f23939,f25830,f23947,f25837,f25840,f25846,f25849,f25852,f25855,f25858,f25861,f25864,f23973,f23975,f23977,f23979,f23985,f23988,f23990,f23991,f23993,f24001,f25872,f25873,f24118,f25875,f25876,f25877,f25878,f25879,f25880,f25883,f25884,f25885,f25886,f25887,f25888,f25889,f25890,f25891,f24144,f24146,f24148,f24150,f25892,f25893,f25894,f25895,f24156,f24159,f24161,f24162,f24164,f25896,f24170,f24172,f24175,f25897,f25898,f25900,f25901,f25902,f25903,f25904,f25905,f25906,f25907,f25908,f25909,f25910,f25911,f25912,f25913,f25914,f25915,f25916,f25917,f25918,f24416,f24421,f24422,f24423,f24424,f24425,f24426,f24427,f24428,f24429,f24430,f24431,f25920,f25926,f25928,f25933,f25935,f25937,f25939,f25941,f25943,f25945,f25998,f26013,f26015,f26018,f26019,f26009,f26010,f26011,f26012,f26024,f26039,f26041,f26044,f26045,f26035,f26036,f26037,f26038,f26053,f26068,f26070,f26073,f26075,f26090,f26105,f26107,f26110,f26111,f26101,f26102,f26103,f26104,f26118,f26133,f26135,f26138,f26139,f26129,f26130,f26131,f26132,f26141,f26160,f26178,f26193,f26195,f26198,f26199,f26189,f26190,f26191,f26192,f26203,f26233,f26260,f26275,f26277,f26280,f26282,f26290,f26320,f26347,f26362,f26364,f26367,f26368,f26358,f26359,f26360,f26361,f26372,f26391,f26408,f26427,f26442,f26458,f26459,f26460,f26461,f26709,f26711,f26714,f26717,f26718,f26719,f26720,f26721,f26723,f26731,f26733,f26735,f26736,f26737,f26738,f26739,f26740,f26741,f26742,f26521,f26523,f26525,f26527,f26751,f26752,f26755,f26537,f26540,f26542,f26543,f26545,f26757,f26759,f26761,f26763,f26764,f26765,f26767,f26768,f26555,f26557,f26560,f26772,f26777,f26783,f26803,f26820,f26840,f26865,f26866,f26867,f26868,f26869,f26921,f26923,f26925,f26927,f26928,f26948,f26949,f26950,f26951,f26952,f26953,f26954,f26955,f26957,f26959,f26960,f26961,f26962,f26963,f26966,f26967,f26968,f26970,f26971,f26972,f26973,f26974,f26975,f26976,f26977,f26978,f26979,f26984,f26985,f26987,f26989,f26991,f26993,f26995,f26996,f26997,f26998,f26999,f27000,f27001,f27002,f27003,f27006,f27008,f27009,f27011,f27013,f27014,f27015,f27016,f27017,f27018,f27019,f27020,f27021,f27023,f27026,f27030,f27033,f27043,f27106,f27107,f27109,f27119,f27132,f27133,f27134,f27135,f27136,f27137,f27138,f27139,f27140,f27141,f27155,f27156,f27157,f27158,f27159,f27160,f27161,f27162,f27166,f27167,f27168,f27169,f27199,f27298,f27533,f27532,f27531,f27530,f27529,f27528,f27524,f27523,f27506,f28372,f28376,f28377,f27431,f27429,f27428,f27425,f27424,f27420,f27417,f27413,f27412,f27411,f27410,f27408,f28347,f28391,f28344,f28341,f28392,f28393,f28338,f28335,f27399])).
fof(f27399,plain,(
  ( ! [X255,X257,X254,X256,X258] : (s(X256,X257) = s(X256,'const.bool.COND_3'(s(bool,'T_0'),s(X256,X258),s(X256,X257))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X254),s(bool,X255))) | s(bool,'F_0') = s(bool,X255)) )),
  inference(superposition,[],[f840,f852])).
fof(f28335,plain,(
  ( ! [X261,X259,X260,X262] : (sP0(X261,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X262),s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X259),s(bool,X260))))) | s(bool,'F_0') = s(bool,X260)) )),
  inference(subsumption_resolution,[],[f27400,f335])).
fof(f27400,plain,(
  ( ! [X261,X259,X260,X262] : (~p(s(bool,'T_0')) | sP0(X261,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X262),s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X259),s(bool,X260))))) | s(bool,'F_0') = s(bool,X260)) )),
  inference(superposition,[],[f861,f852])).
fof(f28338,plain,(
  ( ! [X263,X265,X264,X266] : (sP0('const.bool.COND_3'(s(bool,'F_0'),s(bool,X265),s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X263),s(bool,X264)))),X266) | s(bool,'F_0') = s(bool,X264)) )),
  inference(subsumption_resolution,[],[f27401,f335])).
fof(f27401,plain,(
  ( ! [X263,X265,X264,X266] : (~p(s(bool,'T_0')) | sP0('const.bool.COND_3'(s(bool,'F_0'),s(bool,X265),s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X263),s(bool,X264)))),X266) | s(bool,'F_0') = s(bool,X264)) )),
  inference(superposition,[],[f862,f852])).
fof(f28393,plain,(
  ( ! [X269,X267,X268,X270] : (p(s(bool,X269)) | sP1('const.bool.COND_3'(s(bool,'F_0'),s(bool,X270),s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X267),s(bool,X268)))),X269) | s(bool,'F_0') = s(bool,X268)) )),
  inference(subsumption_resolution,[],[f27402,f335])).
fof(f27402,plain,(
  ( ! [X269,X267,X268,X270] : (~p(s(bool,'T_0')) | p(s(bool,X269)) | sP1('const.bool.COND_3'(s(bool,'F_0'),s(bool,X270),s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X267),s(bool,X268)))),X269) | s(bool,'F_0') = s(bool,X268)) )),
  inference(superposition,[],[f863,f852])).
fof(f28392,plain,(
  ( ! [X271,X273,X272,X274] : (p(s(bool,X273)) | sP1(X273,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X274),s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X271),s(bool,X272))))) | s(bool,'F_0') = s(bool,X272)) )),
  inference(subsumption_resolution,[],[f27403,f335])).
fof(f27403,plain,(
  ( ! [X271,X273,X272,X274] : (~p(s(bool,'T_0')) | p(s(bool,X273)) | sP1(X273,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X274),s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X271),s(bool,X272))))) | s(bool,'F_0') = s(bool,X272)) )),
  inference(superposition,[],[f864,f852])).
fof(f28341,plain,(
  ( ! [X277,X279,X275,X276,X278] : (~p(s(bool,X277)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X278),s(bool,X277)),X279,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X275),s(bool,X276))) | s(bool,'F_0') = s(bool,X276)) )),
  inference(subsumption_resolution,[],[f27404,f335])).
fof(f27404,plain,(
  ( ! [X277,X279,X275,X276,X278] : (~p(s(bool,'T_0')) | ~p(s(bool,X277)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X278),s(bool,X277)),X279,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X275),s(bool,X276))) | s(bool,'F_0') = s(bool,X276)) )),
  inference(superposition,[],[f865,f852])).
fof(f28344,plain,(
  ( ! [X280,X282,X281,X283,X284] : (~p(s(bool,X282)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X283),s(bool,X282)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X280),s(bool,X281)),X284) | s(bool,'F_0') = s(bool,X281)) )),
  inference(subsumption_resolution,[],[f27405,f335])).
fof(f27405,plain,(
  ( ! [X280,X282,X281,X283,X284] : (~p(s(bool,'T_0')) | ~p(s(bool,X282)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X283),s(bool,X282)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X280),s(bool,X281)),X284) | s(bool,'F_0') = s(bool,X281)) )),
  inference(superposition,[],[f866,f852])).
fof(f28391,plain,(
  ( ! [X288,X285,X287,X289,X286] : (p(s(bool,X287)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X288),s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X285),s(bool,X286)))),X289,X287) | s(bool,'F_0') = s(bool,X286)) )),
  inference(subsumption_resolution,[],[f27406,f335])).
fof(f27406,plain,(
  ( ! [X288,X285,X287,X289,X286] : (~p(s(bool,'T_0')) | p(s(bool,X287)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X288),s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X285),s(bool,X286)))),X289,X287) | s(bool,'F_0') = s(bool,X286)) )),
  inference(superposition,[],[f867,f852])).
fof(f28347,plain,(
  ( ! [X292,X294,X290,X293,X291] : (~p(s(bool,X292)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X293),s(bool,X292)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X290),s(bool,X291)),X294) | s(bool,'F_0') = s(bool,X291)) )),
  inference(subsumption_resolution,[],[f27407,f335])).
fof(f27407,plain,(
  ( ! [X292,X294,X290,X293,X291] : (~p(s(bool,'T_0')) | ~p(s(bool,X292)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X293),s(bool,X292)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X290),s(bool,X291)),X294) | s(bool,'F_0') = s(bool,X291)) )),
  inference(superposition,[],[f868,f852])).
fof(f27408,plain,(
  ( ! [X296,X295,X297] : (sP4('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X297)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X295),s(bool,X296))) | s(bool,'F_0') = s(bool,X296)) )),
  inference(superposition,[],[f898,f852])).
fof(f27410,plain,(
  ( ! [X302,X301,X303] : (sP4('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X301),s(bool,X302)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X303),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X302)) )),
  inference(superposition,[],[f930,f852])).
fof(f27413,plain,(
  ( ! [X310,X312,X311] : (sP4('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X312),s(bool,'T_0')),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X310),s(bool,X311))) | s(bool,'F_0') = s(bool,X311)) )),
  inference(superposition,[],[f991,f852])).
fof(f28377,plain,(
  ( ! [X520,X521,X519] : (p(s(bool,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X519),s(bool,X520))),s(bool,'T_0'),s(bool,X521)))) | s(bool,'F_0') = s(bool,X520)) )),
  inference(subsumption_resolution,[],[f27493,f335])).
fof(f27493,plain,(
  ( ! [X520,X521,X519] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X519),s(bool,X520))),s(bool,'T_0'),s(bool,X521)))) | s(bool,'F_0') = s(bool,X520)) )),
  inference(superposition,[],[f4853,f852])).
fof(f28376,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,X1)))) | s(bool,'F_0') = s(bool,X1)) )),
  inference(condensation,[],[f28375])).
fof(f28375,plain,(
  ( ! [X524,X522,X523] : (p(s(bool,X524)) | p(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X522),s(bool,X523)))) | s(bool,'F_0') = s(bool,X523)) )),
  inference(forward_demodulation,[],[f27494,f413])).
fof(f27494,plain,(
  ( ! [X524,X522,X523] : (p(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X524),s(bool,'T_0')))) | p(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X522),s(bool,X523)))) | s(bool,'F_0') = s(bool,X523)) )),
  inference(superposition,[],[f4855,f852])).
fof(f28372,plain,(
  ( ! [X528,X530,X529] : (p(s(bool,'const.bool.COND_3'(s(bool,X530),s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X528),s(bool,X529)))))) | s(bool,'F_0') = s(bool,X529)) )),
  inference(subsumption_resolution,[],[f27496,f335])).
fof(f27496,plain,(
  ( ! [X528,X530,X529] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.COND_3'(s(bool,X530),s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X528),s(bool,X529)))))) | s(bool,'F_0') = s(bool,X529)) )),
  inference(superposition,[],[f4857,f852])).
fof(f27523,plain,(
  ( ! [X609,X607,X608,X610] : (s(bool,'T_0') = s(bool,happ(s(fun(X607,bool),X608),s(X607,'const.bool.COND_3'(s(bool,'F_0'),s(X607,X609),s(X607,X610))))) | s(bool,'F_0') = s(bool,happ(s(fun(X607,bool),X608),s(X607,X610)))) )),
  inference(superposition,[],[f479,f852])).
fof(f27524,plain,(
  ( ! [X613,X611,X612,X614] : (s(bool,'T_0') = s(bool,happ(s(fun(X611,bool),'const.bool.COND_3'(s(bool,'F_0'),s(fun(X611,bool),X612),s(fun(X611,bool),X614))),s(X611,X613))) | s(bool,'F_0') = s(bool,happ(s(fun(X611,bool),X614),s(X611,X613)))) )),
  inference(superposition,[],[f480,f852])).
fof(f27298,plain,(
  ( ! [X14,X15,X16] : (s(bool,X14) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X14),s(bool,X15))) | s(bool,'F_0') = s(bool,X15) | s(bool,X14) = s(bool,X16) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f852,f701])).
fof(f27199,plain,(
  ( ! [X2,X0,X1] : (sP1(X1,'const.bool.COND_3'(s(bool,X0),s(bool,X2),s(bool,'T_0'))) | p(s(bool,X1)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f23139,f343])).
fof(f27169,plain,(
  ( ! [X1050,X1054,X1052,X1051,X1053] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1053),s('type.num.num',X1052)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.bool.COND_3'(s(bool,X1050),s('type.num.num',X1051),s('type.num.num',X1052))),s('type.num.num',X1054)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1053),s('type.num.num',X1054)))) | s(bool,'T_0') = s(bool,X1050)) )),
  inference(superposition,[],[f483,f840])).
fof(f27168,plain,(
  ( ! [X1048,X1046,X1049,X1047,X1045] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1047),s('type.num.num',X1048)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1048),s('type.num.num',X1049)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.bool.COND_3'(s(bool,X1045),s('type.num.num',X1046),s('type.num.num',X1047))),s('type.num.num',X1049)))) | s(bool,'T_0') = s(bool,X1045)) )),
  inference(superposition,[],[f483,f840])).
fof(f27167,plain,(
  ( ! [X1042,X1040,X1044,X1043,X1041] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1043),s('type.num.num',X1042)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bool.COND_3'(s(bool,X1040),s('type.num.num',X1041),s('type.num.num',X1042))),s('type.num.num',X1044)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1043),s('type.num.num',X1044)))) | s(bool,'T_0') = s(bool,X1040)) )),
  inference(superposition,[],[f482,f840])).
fof(f27166,plain,(
  ( ! [X1035,X1039,X1037,X1038,X1036] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1037),s('type.num.num',X1038)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1038),s('type.num.num',X1039)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bool.COND_3'(s(bool,X1035),s('type.num.num',X1036),s('type.num.num',X1037))),s('type.num.num',X1039)))) | s(bool,'T_0') = s(bool,X1035)) )),
  inference(superposition,[],[f482,f840])).
fof(f27162,plain,(
  ( ! [X1019,X1017,X1018,X1016] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1019),s('type.num.num',X1018)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.bool.COND_3'(s(bool,X1016),s('type.num.num',X1017),s('type.num.num',X1018))),s('type.num.num',X1019)))) | s(bool,'T_0') = s(bool,X1016)) )),
  inference(superposition,[],[f477,f840])).
fof(f27161,plain,(
  ( ! [X1015,X1013,X1014,X1012] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1014),s('type.num.num',X1015)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1015),s('type.num.num','const.bool.COND_3'(s(bool,X1012),s('type.num.num',X1013),s('type.num.num',X1014)))))) | s(bool,'T_0') = s(bool,X1012)) )),
  inference(superposition,[],[f477,f840])).
fof(f27160,plain,(
  ( ! [X1011,X1009,X1010,X1008] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1011),s('type.num.num',X1010)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bool.COND_3'(s(bool,X1008),s('type.num.num',X1009),s('type.num.num',X1010))),s('type.num.num',X1011)))) | s(bool,'T_0') = s(bool,X1008)) )),
  inference(superposition,[],[f468,f840])).
fof(f27159,plain,(
  ( ! [X1006,X1004,X1007,X1005] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1006),s('type.num.num',X1007)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1007),s('type.num.num','const.bool.COND_3'(s(bool,X1004),s('type.num.num',X1005),s('type.num.num',X1006)))))) | s(bool,'T_0') = s(bool,X1004)) )),
  inference(superposition,[],[f468,f840])).
fof(f27158,plain,(
  ( ! [X1002,X1000,X1003,X1001] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1002),s('type.num.num',X1003)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1003),s('type.num.num','const.bool.COND_3'(s(bool,X1000),s('type.num.num',X1001),s('type.num.num',X1002)))))) | s(bool,'T_0') = s(bool,X1000)) )),
  inference(superposition,[],[f467,f840])).
fof(f27157,plain,(
  ( ! [X998,X996,X999,X997] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X999),s('type.num.num',X998)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.bool.COND_3'(s(bool,X996),s('type.num.num',X997),s('type.num.num',X998))),s('type.num.num',X999)))) | s(bool,'T_0') = s(bool,X996)) )),
  inference(superposition,[],[f467,f840])).
fof(f27156,plain,(
  ( ! [X994,X992,X995,X993] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X995),s('type.num.num',X994)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bool.COND_3'(s(bool,X992),s('type.num.num',X993),s('type.num.num',X994))),s('type.num.num',X995)))) | s('type.num.num',X995) = s('type.num.num','const.bool.COND_3'(s(bool,X992),s('type.num.num',X993),s('type.num.num',X994))) | s(bool,'T_0') = s(bool,X992)) )),
  inference(superposition,[],[f433,f840])).
fof(f27155,plain,(
  ( ! [X990,X988,X991,X989] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X990),s('type.num.num',X991)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X991),s('type.num.num','const.bool.COND_3'(s(bool,X988),s('type.num.num',X989),s('type.num.num',X990)))))) | s('type.num.num',X991) = s('type.num.num','const.bool.COND_3'(s(bool,X988),s('type.num.num',X989),s('type.num.num',X990))) | s(bool,'T_0') = s(bool,X988)) )),
  inference(superposition,[],[f433,f840])).
fof(f27141,plain,(
  ( ! [X922,X926,X924,X923,X927,X925] : (~p(s(bool,happ(s(fun(X922,bool),X925),s(X922,X926)))) | p(s(bool,happ(s(fun(X922,bool),X927),s(X922,sK21('const.bool.COND_3'(s(bool,X923),s(fun(X922,bool),X924),s(fun(X922,bool),X925)),X927,X922))))) | p(s(bool,happ(s(fun(X922,bool),'const.bool.COND_3'(s(bool,X923),s(fun(X922,bool),X924),s(fun(X922,bool),X925))),s(X922,sK21('const.bool.COND_3'(s(bool,X923),s(fun(X922,bool),X924),s(fun(X922,bool),X925)),X927,X922))))) | s(bool,'T_0') = s(bool,X923)) )),
  inference(superposition,[],[f424,f840])).
fof(f27140,plain,(
  ( ! [X920,X918,X916,X921,X919,X917] : (~p(s(bool,happ(s(fun(X916,bool),X919),s(X916,X920)))) | p(s(bool,happ(s(fun(X916,bool),'const.bool.COND_3'(s(bool,X917),s(fun(X916,bool),X918),s(fun(X916,bool),X919))),s(X916,sK21(X921,'const.bool.COND_3'(s(bool,X917),s(fun(X916,bool),X918),s(fun(X916,bool),X919)),X916))))) | p(s(bool,happ(s(fun(X916,bool),X921),s(X916,sK21(X921,'const.bool.COND_3'(s(bool,X917),s(fun(X916,bool),X918),s(fun(X916,bool),X919)),X916))))) | s(bool,'T_0') = s(bool,X917)) )),
  inference(superposition,[],[f423,f840])).
fof(f27139,plain,(
  ( ! [X912,X914,X911,X915,X913,X910] : (~p(s(bool,happ(s(fun(X910,bool),X913),s(X910,X914)))) | p(s(bool,happ(s(fun(X910,bool),X915),s(X910,sK19('const.bool.COND_3'(s(bool,X911),s(fun(X910,bool),X912),s(fun(X910,bool),X913)),X915,X910))))) | p(s(bool,happ(s(fun(X910,bool),'const.bool.COND_3'(s(bool,X911),s(fun(X910,bool),X912),s(fun(X910,bool),X913))),s(X910,sK20('const.bool.COND_3'(s(bool,X911),s(fun(X910,bool),X912),s(fun(X910,bool),X913)),X915,X910))))) | s(bool,'T_0') = s(bool,X911)) )),
  inference(superposition,[],[f422,f840])).
fof(f27138,plain,(
  ( ! [X907,X905,X909,X906,X904,X908] : (~p(s(bool,happ(s(fun(X904,bool),X907),s(X904,X908)))) | p(s(bool,happ(s(fun(X904,bool),'const.bool.COND_3'(s(bool,X905),s(fun(X904,bool),X906),s(fun(X904,bool),X907))),s(X904,sK19(X909,'const.bool.COND_3'(s(bool,X905),s(fun(X904,bool),X906),s(fun(X904,bool),X907)),X904))))) | p(s(bool,happ(s(fun(X904,bool),X909),s(X904,sK20(X909,'const.bool.COND_3'(s(bool,X905),s(fun(X904,bool),X906),s(fun(X904,bool),X907)),X904))))) | s(bool,'T_0') = s(bool,X905)) )),
  inference(superposition,[],[f421,f840])).
fof(f27137,plain,(
  ( ! [X899,X903,X901,X898,X902,X900] : (~p(s(bool,happ(s(fun(X898,bool),X901),s(X898,sK18('const.bool.COND_3'(s(bool,X899),s(fun(X898,bool),X900),s(fun(X898,bool),X901)),X902,X898))))) | p(s(bool,happ(s(fun(X898,bool),'const.bool.COND_3'(s(bool,X899),s(fun(X898,bool),X900),s(fun(X898,bool),X901))),s(X898,X903)))) | ~p(s(bool,happ(s(fun(X898,bool),X902),s(X898,sK17('const.bool.COND_3'(s(bool,X899),s(fun(X898,bool),X900),s(fun(X898,bool),X901)),X902,X898))))) | s(bool,'T_0') = s(bool,X899)) )),
  inference(superposition,[],[f420,f840])).
fof(f27136,plain,(
  ( ! [X895,X893,X897,X894,X892,X896] : (~p(s(bool,happ(s(fun(X892,bool),X895),s(X892,sK18('const.bool.COND_3'(s(bool,X893),s(fun(X892,bool),X894),s(fun(X892,bool),X895)),X896,X892))))) | p(s(bool,happ(s(fun(X892,bool),X896),s(X892,X897)))) | ~p(s(bool,happ(s(fun(X892,bool),X896),s(X892,sK17('const.bool.COND_3'(s(bool,X893),s(fun(X892,bool),X894),s(fun(X892,bool),X895)),X896,X892))))) | s(bool,'T_0') = s(bool,X893)) )),
  inference(superposition,[],[f419,f840])).
fof(f27135,plain,(
  ( ! [X887,X890,X888,X886,X891,X889] : (~p(s(bool,happ(s(fun(X886,bool),X889),s(X886,sK16('const.bool.COND_3'(s(bool,X887),s(fun(X886,bool),X888),s(fun(X886,bool),X889)),X890,X886))))) | p(s(bool,happ(s(fun(X886,bool),'const.bool.COND_3'(s(bool,X887),s(fun(X886,bool),X888),s(fun(X886,bool),X889))),s(X886,X891)))) | ~p(s(bool,happ(s(fun(X886,bool),X890),s(X886,sK16('const.bool.COND_3'(s(bool,X887),s(fun(X886,bool),X888),s(fun(X886,bool),X889)),X890,X886))))) | s(bool,'T_0') = s(bool,X887)) )),
  inference(superposition,[],[f418,f840])).
fof(f27134,plain,(
  ( ! [X883,X881,X885,X882,X880,X884] : (~p(s(bool,happ(s(fun(X880,bool),X883),s(X880,sK16('const.bool.COND_3'(s(bool,X881),s(fun(X880,bool),X882),s(fun(X880,bool),X883)),X884,X880))))) | p(s(bool,happ(s(fun(X880,bool),X884),s(X880,X885)))) | ~p(s(bool,happ(s(fun(X880,bool),X884),s(X880,sK16('const.bool.COND_3'(s(bool,X881),s(fun(X880,bool),X882),s(fun(X880,bool),X883)),X884,X880))))) | s(bool,'T_0') = s(bool,X881)) )),
  inference(superposition,[],[f417,f840])).
fof(f27133,plain,(
  ( ! [X878,X876,X875,X879,X877] : (~p(s(bool,happ(s(fun(X875,bool),X878),s(X875,sK15('const.bool.COND_3'(s(bool,X876),s(fun(X875,bool),X877),s(fun(X875,bool),X878)),X875))))) | p(s(bool,happ(s(fun(X875,bool),'const.bool.COND_3'(s(bool,X876),s(fun(X875,bool),X877),s(fun(X875,bool),X878))),s(X875,X879)))) | s(bool,'T_0') = s(bool,X876)) )),
  inference(superposition,[],[f412,f840])).
fof(f27132,plain,(
  ( ! [X874,X872,X870,X873,X871] : (~p(s(bool,happ(s(fun(X870,bool),X873),s(X870,sK14('const.bool.COND_3'(s(bool,X871),s(fun(X870,bool),X872),s(fun(X870,bool),X873)),X870))))) | p(s(bool,happ(s(fun(X870,bool),'const.bool.COND_3'(s(bool,X871),s(fun(X870,bool),X872),s(fun(X870,bool),X873))),s(X870,X874)))) | s(bool,'T_0') = s(bool,X871)) )),
  inference(superposition,[],[f411,f840])).
fof(f27119,plain,(
  ( ! [X823,X821,X824,X822] : (sP1('const.bool.COND_3'(s(bool,X823),s(bool,X824),s(bool,'T_0')),'F_0') | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X821),s(bool,X822),s(bool,X823))) | s(bool,'T_0') = s(bool,X821)) )),
  inference(superposition,[],[f16717,f840])).
fof(f27109,plain,(
  ( ! [X779,X781,X782,X780] : (~p(s(bool,X781)) | p(s(bool,'const.bool.COND_3'(s(bool,X782),s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,X779),s(bool,X780),s(bool,X781)))))) | s(bool,'T_0') = s(bool,X779)) )),
  inference(superposition,[],[f4857,f840])).
fof(f27107,plain,(
  ( ! [X771,X773,X774,X772] : (p(s(bool,'const.bool.COND_3'(s(bool,X773),s(bool,X774),s(bool,'T_0')))) | p(s(bool,'const.bool.COND_3'(s(bool,X771),s(bool,X772),s(bool,X773)))) | s(bool,'T_0') = s(bool,X771)) )),
  inference(superposition,[],[f4855,f840])).
fof(f27106,plain,(
  ( ! [X767,X769,X770,X768] : (~p(s(bool,X769)) | p(s(bool,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,X767),s(bool,X768),s(bool,X769))),s(bool,'T_0'),s(bool,X770)))) | s(bool,'T_0') = s(bool,X767)) )),
  inference(superposition,[],[f4853,f840])).
fof(f27043,plain,(
  ( ! [X542,X545,X543,X544] : (~p(s(bool,'const.bool.COND_3'(s(bool,X544),s(bool,X545),s(bool,'F_0')))) | p(s(bool,'const.bool.COND_3'(s(bool,X542),s(bool,X543),s(bool,X544)))) | s(bool,'T_0') = s(bool,X542)) )),
  inference(superposition,[],[f1494,f840])).
fof(f27033,plain,(
  ( ! [X509,X511,X512,X508,X510] : (sP0('const.bool.COND_3'(s(bool,X510),s(bool,X511),s(bool,'T_0')),X512) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X508),s(bool,X509),s(bool,X510))) | s(bool,'T_0') = s(bool,X508)) )),
  inference(superposition,[],[f1275,f840])).
fof(f27030,plain,(
  ( ! [X494,X496,X498,X495,X497] : (sP0(X497,'const.bool.COND_3'(s(bool,X496),s(bool,X498),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X494),s(bool,X495),s(bool,X496))) | s(bool,'T_0') = s(bool,X494)) )),
  inference(superposition,[],[f1239,f840])).
fof(f27026,plain,(
  ( ! [X478,X480,X479,X481] : (sP4('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X481),s(bool,X480)),'const.bool.COND_3'(s(bool,X478),s(bool,X479),s(bool,X480))) | s(bool,'T_0') = s(bool,X478)) )),
  inference(superposition,[],[f991,f840])).
fof(f27023,plain,(
  ( ! [X466,X469,X467,X468] : (sP4('T_0','const.bool.COND_3'(s(bool,X466),s(bool,X467),s(bool,X468)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X469),s(bool,X468))) | s(bool,'T_0') = s(bool,X466)) )),
  inference(superposition,[],[f930,f840])).
fof(f27021,plain,(
  ( ! [X461,X459,X460,X458] : (sP4('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X460),s(bool,X461)),'const.bool.COND_3'(s(bool,X458),s(bool,X459),s(bool,X460))) | s(bool,'T_0') = s(bool,X458)) )),
  inference(superposition,[],[f898,f840])).
fof(f27020,plain,(
  ( ! [X453,X455,X457,X452,X454,X456] : (~p(s(bool,X454)) | ~p(s(bool,X455)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X456),s(bool,X455)),'const.bool.COND_3'(s(bool,X452),s(bool,X453),s(bool,X454)),X457) | s(bool,'T_0') = s(bool,X452)) )),
  inference(superposition,[],[f868,f840])).
fof(f27019,plain,(
  ( ! [X447,X449,X451,X446,X448,X450] : (~p(s(bool,X448)) | p(s(bool,X449)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X450),s(bool,'const.bool.COND_3'(s(bool,X446),s(bool,X447),s(bool,X448)))),X451,X449) | s(bool,'T_0') = s(bool,X446)) )),
  inference(superposition,[],[f867,f840])).
fof(f27018,plain,(
  ( ! [X445,X441,X443,X444,X440,X442] : (~p(s(bool,X442)) | ~p(s(bool,X443)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X444),s(bool,X443)),'const.bool.COND_3'(s(bool,X440),s(bool,X441),s(bool,X442)),X445) | s(bool,'T_0') = s(bool,X440)) )),
  inference(superposition,[],[f866,f840])).
fof(f27017,plain,(
  ( ! [X436,X438,X434,X437,X439,X435] : (~p(s(bool,X436)) | ~p(s(bool,X437)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X438),s(bool,X437)),X439,'const.bool.COND_3'(s(bool,X434),s(bool,X435),s(bool,X436))) | s(bool,'T_0') = s(bool,X434)) )),
  inference(superposition,[],[f865,f840])).
fof(f27016,plain,(
  ( ! [X430,X432,X429,X431,X433] : (~p(s(bool,X431)) | p(s(bool,X432)) | sP1(X432,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X433),s(bool,'const.bool.COND_3'(s(bool,X429),s(bool,X430),s(bool,X431))))) | s(bool,'T_0') = s(bool,X429)) )),
  inference(superposition,[],[f864,f840])).
fof(f27015,plain,(
  ( ! [X428,X424,X426,X425,X427] : (~p(s(bool,X426)) | p(s(bool,X427)) | sP1('const.bool.COND_3'(s(bool,'F_0'),s(bool,X428),s(bool,'const.bool.COND_3'(s(bool,X424),s(bool,X425),s(bool,X426)))),X427) | s(bool,'T_0') = s(bool,X424)) )),
  inference(superposition,[],[f863,f840])).
fof(f27014,plain,(
  ( ! [X420,X422,X421,X423,X419] : (~p(s(bool,X421)) | sP0('const.bool.COND_3'(s(bool,'F_0'),s(bool,X422),s(bool,'const.bool.COND_3'(s(bool,X419),s(bool,X420),s(bool,X421)))),X423) | s(bool,'T_0') = s(bool,X419)) )),
  inference(superposition,[],[f862,f840])).
fof(f27013,plain,(
  ( ! [X414,X416,X418,X415,X417] : (~p(s(bool,X416)) | sP0(X417,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X418),s(bool,'const.bool.COND_3'(s(bool,X414),s(bool,X415),s(bool,X416))))) | s(bool,'T_0') = s(bool,X414)) )),
  inference(superposition,[],[f861,f840])).
fof(f27011,plain,(
  ( ! [X408,X407,X409,X406] : (sP5('T_0','const.bool.COND_3'(s(bool,X406),s(bool,X407),s(bool,X408)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X409),s(bool,X408))) | s(bool,'T_0') = s(bool,X406)) )),
  inference(superposition,[],[f839,f840])).
fof(f27009,plain,(
  ( ! [X399,X401,X398,X400] : (~sP5('F_0','const.bool.COND_3'(s(bool,X398),s(bool,X399),s(bool,X400)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X401),s(bool,X400))) | s(bool,'T_0') = s(bool,X398)) )),
  inference(superposition,[],[f838,f840])).
fof(f27008,plain,(
  ( ! [X397,X395,X396,X394] : (sP5('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X397),s(bool,X396)),'const.bool.COND_3'(s(bool,X394),s(bool,X395),s(bool,X396))) | s(bool,'T_0') = s(bool,X394)) )),
  inference(superposition,[],[f837,f840])).
fof(f27006,plain,(
  ( ! [X389,X387,X388,X386] : (~sP5('F_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X389),s(bool,X388)),'const.bool.COND_3'(s(bool,X386),s(bool,X387),s(bool,X388))) | s(bool,'T_0') = s(bool,X386)) )),
  inference(superposition,[],[f836,f840])).
fof(f27003,plain,(
  ( ! [X377,X374,X376,X375] : (sP4('T_0','const.bool.COND_3'(s(bool,X374),s(bool,X375),s(bool,X376)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X376),s(bool,X377))) | s(bool,'T_0') = s(bool,X374)) )),
  inference(superposition,[],[f825,f840])).
fof(f27002,plain,(
  ( ! [X372,X368,X370,X373,X369,X371] : (~p(s(bool,X370)) | ~p(s(bool,X371)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X371),s(bool,X372)),'const.bool.COND_3'(s(bool,X368),s(bool,X369),s(bool,X370)),X373) | s(bool,'T_0') = s(bool,X368)) )),
  inference(superposition,[],[f803,f840])).
fof(f27001,plain,(
  ( ! [X364,X366,X362,X365,X367,X363] : (~p(s(bool,X364)) | p(s(bool,X365)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,X362),s(bool,X363),s(bool,X364))),s(bool,X366)),X367,X365) | s(bool,'T_0') = s(bool,X362)) )),
  inference(superposition,[],[f802,f840])).
fof(f27000,plain,(
  ( ! [X356,X358,X360,X357,X359,X361] : (~p(s(bool,X358)) | ~p(s(bool,X359)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X359),s(bool,X360)),'const.bool.COND_3'(s(bool,X356),s(bool,X357),s(bool,X358)),X361) | s(bool,'T_0') = s(bool,X356)) )),
  inference(superposition,[],[f801,f840])).
fof(f26999,plain,(
  ( ! [X350,X352,X354,X351,X353,X355] : (~p(s(bool,X352)) | ~p(s(bool,X353)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X353),s(bool,X354)),X355,'const.bool.COND_3'(s(bool,X350),s(bool,X351),s(bool,X352))) | s(bool,'T_0') = s(bool,X350)) )),
  inference(superposition,[],[f800,f840])).
fof(f26998,plain,(
  ( ! [X346,X349,X345,X347,X348] : (~p(s(bool,X347)) | p(s(bool,X348)) | sP1(X348,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,X345),s(bool,X346),s(bool,X347))),s(bool,X349))) | s(bool,'T_0') = s(bool,X345)) )),
  inference(superposition,[],[f799,f840])).
fof(f26997,plain,(
  ( ! [X344,X341,X343,X340,X342] : (~p(s(bool,X342)) | p(s(bool,X343)) | sP1('const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,X340),s(bool,X341),s(bool,X342))),s(bool,X344)),X343) | s(bool,'T_0') = s(bool,X340)) )),
  inference(superposition,[],[f798,f840])).
fof(f26996,plain,(
  ( ! [X335,X337,X339,X336,X338] : (~p(s(bool,X337)) | sP0('const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,X335),s(bool,X336),s(bool,X337))),s(bool,X338)),X339) | s(bool,'T_0') = s(bool,X335)) )),
  inference(superposition,[],[f797,f840])).
fof(f26995,plain,(
  ( ! [X333,X331,X332,X334,X330] : (~p(s(bool,X332)) | sP0(X333,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,X330),s(bool,X331),s(bool,X332))),s(bool,X334))) | s(bool,'T_0') = s(bool,X330)) )),
  inference(superposition,[],[f796,f840])).
fof(f26993,plain,(
  ( ! [X325,X323,X324,X322] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X324),s(bool,X325))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X322),s(bool,X323),s(bool,X324))) | s(bool,'T_0') = s(bool,X322)) )),
  inference(superposition,[],[f787,f840])).
fof(f26991,plain,(
  ( ! [X315,X316,X314,X317] : (sP5('T_0','const.bool.COND_3'(s(bool,X314),s(bool,X315),s(bool,X316)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X316),s(bool,X317))) | s(bool,'T_0') = s(bool,X314)) )),
  inference(superposition,[],[f778,f840])).
fof(f26989,plain,(
  ( ! [X308,X306,X309,X307] : (~sP5('F_0','const.bool.COND_3'(s(bool,X306),s(bool,X307),s(bool,X308)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X308),s(bool,X309))) | s(bool,'T_0') = s(bool,X306)) )),
  inference(superposition,[],[f777,f840])).
fof(f26987,plain,(
  ( ! [X300,X298,X301,X299] : (sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X300),s(bool,X301)),'const.bool.COND_3'(s(bool,X298),s(bool,X299),s(bool,X300))) | s(bool,'T_0') = s(bool,X298)) )),
  inference(superposition,[],[f776,f840])).
fof(f26985,plain,(
  ( ! [X292,X290,X293,X291] : (~sP5('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X292),s(bool,X293)),'const.bool.COND_3'(s(bool,X290),s(bool,X291),s(bool,X292))) | s(bool,'T_0') = s(bool,X290)) )),
  inference(superposition,[],[f775,f840])).
fof(f26984,plain,(
  ( ! [X288,X287,X289,X286] : (p(s(bool,X288)) | s(bool,X289) = s(bool,'const.bool.COND_3'(s(bool,X286),s(bool,X287),s(bool,X288))) | s(bool,'T_0') = s(bool,X289) | s(bool,'T_0') = s(bool,X286)) )),
  inference(superposition,[],[f747,f840])).
fof(f26979,plain,(
  ( ! [X269,X265,X267,X268,X266] : (~p(s(bool,X267)) | ~sP8('const.bool.COND_3'(s(bool,X265),s(bool,X266),s(bool,X267)),X268,X269) | sP9('const.bool.COND_3'(s(bool,X265),s(bool,X266),s(bool,X267)),X268,X269) | s(bool,'T_0') = s(bool,X265)) )),
  inference(superposition,[],[f577,f840])).
fof(f26978,plain,(
  ( ! [X261,X263,X260,X262,X264] : (~p(s(bool,X262)) | ~sP10('const.bool.COND_3'(s(bool,X260),s(bool,X261),s(bool,X262)),X263,X264) | sP11('const.bool.COND_3'(s(bool,X260),s(bool,X261),s(bool,X262)),X263,X264) | s(bool,'T_0') = s(bool,X260)) )),
  inference(superposition,[],[f576,f840])).
fof(f26977,plain,(
  ( ! [X255,X257,X259,X256,X258] : (~p(s(bool,X257)) | ~p(s(bool,X258)) | sP11('const.bool.COND_3'(s(bool,X255),s(bool,X256),s(bool,X257)),X258,X259) | s(bool,'T_0') = s(bool,X255)) )),
  inference(superposition,[],[f572,f840])).
fof(f26976,plain,(
  ( ! [X253,X251,X254,X252,X250] : (~p(s(bool,X252)) | p(s(bool,X253)) | sP11('const.bool.COND_3'(s(bool,X250),s(bool,X251),s(bool,X252)),X254,X253) | s(bool,'T_0') = s(bool,X250)) )),
  inference(superposition,[],[f571,f840])).
fof(f26975,plain,(
  ( ! [X249,X246,X248,X247,X245] : (~p(s(bool,X247)) | ~p(s(bool,X248)) | sP9('const.bool.COND_3'(s(bool,X245),s(bool,X246),s(bool,X247)),X248,X249) | s(bool,'T_0') = s(bool,X245)) )),
  inference(superposition,[],[f555,f840])).
fof(f26974,plain,(
  ( ! [X243,X241,X244,X242,X240] : (~p(s(bool,X242)) | ~p(s(bool,X243)) | sP9('const.bool.COND_3'(s(bool,X240),s(bool,X241),s(bool,X242)),X244,X243) | s(bool,'T_0') = s(bool,X240)) )),
  inference(superposition,[],[f554,f840])).
fof(f26973,plain,(
  ( ! [X239,X237,X235,X238,X236] : (~p(s(bool,X237)) | ~p(s(bool,X238)) | p(s(bool,X239)) | ~sP7(X239,X238,'const.bool.COND_3'(s(bool,X235),s(bool,X236),s(bool,X237))) | s(bool,'T_0') = s(bool,X235)) )),
  inference(superposition,[],[f541,f840])).
fof(f26972,plain,(
  ( ! [X231,X233,X230,X234,X232] : (~p(s(bool,X232)) | ~p(s(bool,X233)) | p(s(bool,'const.bool.COND_3'(s(bool,X234),s(bool,'const.bool.COND_3'(s(bool,X230),s(bool,X231),s(bool,X232))),s(bool,X233)))) | s(bool,'T_0') = s(bool,X230)) )),
  inference(superposition,[],[f489,f840])).
fof(f26971,plain,(
  ( ! [X229,X227,X225,X228,X226] : (~p(s(bool,X227)) | p(s(bool,'const.bool.COND_3'(s(bool,X228),s(bool,X229),s(bool,'const.bool.COND_3'(s(bool,X225),s(bool,X226),s(bool,X227)))))) | p(s(bool,X228)) | s(bool,'T_0') = s(bool,X225)) )),
  inference(superposition,[],[f488,f840])).
fof(f26970,plain,(
  ( ! [X222,X220,X223,X221,X224] : (~p(s(bool,X222)) | ~p(s(bool,X223)) | p(s(bool,'const.bool.COND_3'(s(bool,X223),s(bool,'const.bool.COND_3'(s(bool,X220),s(bool,X221),s(bool,X222))),s(bool,X224)))) | s(bool,'T_0') = s(bool,X220)) )),
  inference(superposition,[],[f487,f840])).
fof(f26968,plain,(
  ( ! [X212,X210,X213,X211,X214] : (~p(s(bool,'const.bool.COND_3'(s(bool,X213),s(bool,X214),s(bool,X212)))) | p(s(bool,X213)) | p(s(bool,'const.bool.COND_3'(s(bool,X210),s(bool,X211),s(bool,X212)))) | s(bool,'T_0') = s(bool,X210)) )),
  inference(superposition,[],[f485,f840])).
fof(f26967,plain,(
  ( ! [X208,X206,X209,X207,X205] : (~p(s(bool,'const.bool.COND_3'(s(bool,X207),s(bool,X208),s(bool,X209)))) | p(s(bool,'const.bool.COND_3'(s(bool,X205),s(bool,X206),s(bool,X207)))) | p(s(bool,X209)) | s(bool,'T_0') = s(bool,X205)) )),
  inference(superposition,[],[f485,f840])).
fof(f26966,plain,(
  ( ! [X204,X202,X200,X203,X201] : (~p(s(bool,X202)) | p(s(bool,X203)) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,X200),s(bool,X201),s(bool,X202))),s(bool,X203),s(bool,X204)))) | s(bool,'T_0') = s(bool,X200)) )),
  inference(superposition,[],[f484,f840])).
fof(f26963,plain,(
  ( ! [X182,X180,X183,X181] : (~p(s(bool,X182)) | p(s(bool,X183)) | sP1(X183,'const.bool.COND_3'(s(bool,X180),s(bool,X181),s(bool,X182))) | s(bool,'T_0') = s(bool,X180)) )),
  inference(superposition,[],[f462,f840])).
fof(f26962,plain,(
  ( ! [X177,X178,X176,X179] : (~p(s(bool,X178)) | p(s(bool,X179)) | sP1('const.bool.COND_3'(s(bool,X176),s(bool,X177),s(bool,X178)),X179) | s(bool,'T_0') = s(bool,X176)) )),
  inference(superposition,[],[f459,f840])).
fof(f26961,plain,(
  ( ! [X175,X173,X174,X172] : (~p(s(bool,X174)) | sP0('const.bool.COND_3'(s(bool,X172),s(bool,X173),s(bool,X174)),X175) | s(bool,'T_0') = s(bool,X172)) )),
  inference(superposition,[],[f444,f840])).
fof(f26960,plain,(
  ( ! [X171,X169,X170,X168] : (~p(s(bool,X170)) | sP0(X171,'const.bool.COND_3'(s(bool,X168),s(bool,X169),s(bool,X170))) | s(bool,'T_0') = s(bool,X168)) )),
  inference(superposition,[],[f443,f840])).
fof(f26959,plain,(
  ( ! [X167,X165,X166,X164] : (~p(s(bool,X166)) | ~p(s(bool,X167)) | s(bool,X167) = s(bool,'const.bool.COND_3'(s(bool,X164),s(bool,X165),s(bool,X166))) | s(bool,'T_0') = s(bool,X164)) )),
  inference(superposition,[],[f437,f840])).
fof(f26957,plain,(
  ( ! [X158,X156,X159,X157] : (p(s(bool,X158)) | p(s(bool,X159)) | s(bool,X159) = s(bool,'const.bool.COND_3'(s(bool,X156),s(bool,X157),s(bool,X158))) | s(bool,'T_0') = s(bool,X156)) )),
  inference(superposition,[],[f434,f840])).
fof(f26955,plain,(
  ( ! [X146,X144,X142,X140,X138,X145,X143,X141,X139] : (s(bool,X140) != s(bool,X141) | s(X142,X143) != s(X142,X144) | s(X142,X145) != s(X142,X146) | s(X142,'const.bool.COND_3'(s(bool,X141),s(X142,X145),s(X142,X143))) = s(X142,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,X138),s(bool,X139),s(bool,X140))),s(X142,X146),s(X142,X144))) | s(bool,'T_0') = s(bool,X138)) )),
  inference(superposition,[],[f428,f840])).
fof(f26954,plain,(
  ( ! [X134,X132,X130,X136,X135,X133,X131,X129,X137] : (s(bool,X131) != s(bool,X132) | s(X133,X134) != s(X133,X135) | p(s(bool,X132)) | s(X133,'const.bool.COND_3'(s(bool,X132),s(X133,X137),s(X133,X135))) = s(X133,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,X129),s(bool,X130),s(bool,X131))),s(X133,X136),s(X133,X134))) | s(bool,'T_0') = s(bool,X129)) )),
  inference(superposition,[],[f427,f840])).
fof(f26953,plain,(
  ( ! [X127,X125,X123,X121,X128,X126,X124,X122,X120] : (s(bool,X122) != s(bool,X123) | s(X124,X125) != s(X124,X126) | p(s(bool,'const.bool.COND_3'(s(bool,X120),s(bool,X121),s(bool,X122)))) | s(X124,'const.bool.COND_3'(s(bool,X123),s(X124,X127),s(X124,X125))) = s(X124,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,X120),s(bool,X121),s(bool,X122))),s(X124,X128),s(X124,X126))) | s(bool,'T_0') = s(bool,X120)) )),
  inference(superposition,[],[f427,f840])).
fof(f26952,plain,(
  ( ! [X111,X118,X116,X114,X112,X119,X117,X115,X113] : (~p(s(bool,X113)) | s(X114,'const.bool.COND_3'(s(bool,X115),s(X114,X116),s(X114,X117))) = s(X114,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,X111),s(bool,X112),s(bool,X113))),s(X114,X118),s(X114,X119))) | s(X114,X116) != s(X114,X118) | s(bool,X115) != s(bool,'const.bool.COND_3'(s(bool,X111),s(bool,X112),s(bool,X113))) | s(bool,'T_0') = s(bool,X111)) )),
  inference(superposition,[],[f426,f840])).
fof(f26951,plain,(
  ( ! [X109,X110,X108] : (p(s(bool,X110)) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X108),s(bool,X109),s(bool,X110))) | s(bool,'T_0') = s(bool,X108)) )),
  inference(superposition,[],[f359,f840])).
fof(f26950,plain,(
  ( ! [X107,X105,X106] : (~p(s(bool,X107)) | s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X105),s(bool,X106),s(bool,X107))) | s(bool,'T_0') = s(bool,X105)) )),
  inference(superposition,[],[f358,f840])).
fof(f26949,plain,(
  ( ! [X103,X102,X104] : (~p(s(bool,X104)) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X102),s(bool,X103),s(bool,X104))) | s(bool,'T_0') = s(bool,X102)) )),
  inference(superposition,[],[f355,f840])).
fof(f26948,plain,(
  ( ! [X101,X99,X100] : (s(bool,'T_0') != s(bool,X101) | p(s(bool,'const.bool.COND_3'(s(bool,X99),s(bool,X100),s(bool,X101)))) | s(bool,'T_0') = s(bool,X99)) )),
  inference(superposition,[],[f354,f840])).
fof(f26928,plain,(
  ( ! [X39,X43,X41,X38,X42,X40] : (s(X38,happ(s(fun(X40,X38),X43),s(X40,X42))) = s(X38,happ(s(fun(X40,X38),'const.bool.COND_3'(s(bool,X39),s(fun(X40,X38),X41),s(fun(X40,X38),X43))),s(X40,X42))) | s(bool,'T_0') = s(bool,X39)) )),
  inference(superposition,[],[f840,f480])).
fof(f26927,plain,(
  ( ! [X37,X35,X33,X36,X34,X32] : (s(X32,happ(s(fun(X34,X32),X35),s(X34,X37))) = s(X32,happ(s(fun(X34,X32),X35),s(X34,'const.bool.COND_3'(s(bool,X33),s(X34,X36),s(X34,X37))))) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f840,f479])).
fof(f26925,plain,(
  ( ! [X24,X23,X25] : (s(bool,'F_0') = s(bool,X25) | s(bool,'T_0') = s(bool,X23) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X23),s(bool,X24),s(bool,X25)))) )),
  inference(superposition,[],[f840,f343])).
fof(f26923,plain,(
  ( ! [X17,X15,X18,X16] : (s(bool,X17) = s(bool,X18) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X15),s(bool,X16),s(bool,X17))) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f840,f701])).
fof(f26921,plain,(
  ( ! [X10,X8,X7,X9] : (s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,X7) | s(bool,X10) = s(bool,'const.bool.COND_3'(s(bool,X7),s(bool,X8),s(bool,X9))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f840,f701])).
fof(f26869,plain,(
  ( ! [X76,X74,X72,X71,X75,X73] : (s(X74,X75) = s(X74,'const.bool.COND_3'(s(bool,X73),s(X74,X76),s(X74,X75))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X71),s(bool,X72),s(bool,X73))) | s(bool,'T_0') = s(bool,X71)) )),
  inference(superposition,[],[f840,f840])).
fof(f26868,plain,(
  ( ! [X70,X68,X66,X64,X69,X67,X65,X63] : (s(X68,X69) = s(X68,'const.bool.COND_3'(s(bool,happ(s(fun(X64,bool),'const.bool.COND_3'(s(bool,X63),s(fun(X64,bool),X65),s(fun(X64,bool),X67))),s(X64,X66))),s(X68,X70),s(X68,X69))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X63),s(bool,happ(s(fun(X64,bool),X65),s(X64,X66))),s(bool,happ(s(fun(X64,bool),X67),s(X64,X66)))))) )),
  inference(superposition,[],[f840,f480])).
fof(f26867,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X56,X55] : (s(X60,X61) = s(X60,'const.bool.COND_3'(s(bool,happ(s(fun(X56,bool),X57),s(X56,'const.bool.COND_3'(s(bool,X55),s(X56,X58),s(X56,X59))))),s(X60,X62),s(X60,X61))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X55),s(bool,happ(s(fun(X56,bool),X57),s(X56,X58))),s(bool,happ(s(fun(X56,bool),X57),s(X56,X59)))))) )),
  inference(superposition,[],[f840,f479])).
fof(f26866,plain,(
  ( ! [X54,X52,X50,X53,X51] : (s(X52,X53) = s(X52,'const.bool.COND_3'(s(bool,X51),s(X52,X54),s(X52,X53))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X50),s(bool,X51)))) )),
  inference(superposition,[],[f840,f414])).
fof(f26865,plain,(
  ( ! [X47,X45,X48,X46,X49] : (s(X47,X48) = s(X47,'const.bool.COND_3'(s(bool,X45),s(X47,X49),s(X47,X48))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X45),s(bool,X46)))) )),
  inference(superposition,[],[f840,f413])).
fof(f26840,plain,(
  ( ! [X2,X0,X3,X1] : (sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,X0)),'T_0',X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f26347,f701])).
fof(f26820,plain,(
  ( ! [X2,X0,X1] : (sP11('const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,'T_0')),'T_0',X2) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f26347,f343])).
fof(f26803,plain,(
  ( ! [X2,X0,X3,X1] : (sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,X0)),X2,'F_0') | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f26260,f701])).
fof(f26783,plain,(
  ( ! [X2,X0,X1] : (sP11('const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,'T_0')),X2,'F_0') | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f26260,f343])).
fof(f26777,plain,(
  ( ! [X517,X519] : (p(s(bool,'const.bool.COND_3'(s(bool,X519),s(bool,'T_0'),s(bool,X517)))) | s(bool,'F_0') = s(bool,X517)) )),
  inference(forward_demodulation,[],[f26776,f413])).
fof(f26776,plain,(
  ( ! [X518,X517,X519] : (p(s(bool,'const.bool.COND_3'(s(bool,X519),s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X517),s(bool,X518)))))) | s(bool,'F_0') = s(bool,X517)) )),
  inference(subsumption_resolution,[],[f26643,f335])).
fof(f26643,plain,(
  ( ! [X518,X517,X519] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.COND_3'(s(bool,X519),s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X517),s(bool,X518)))))) | s(bool,'F_0') = s(bool,X517)) )),
  inference(superposition,[],[f4857,f787])).
fof(f26772,plain,(
  ( ! [X508,X510] : (p(s(bool,'const.bool.COND_3'(s(bool,X508),s(bool,'T_0'),s(bool,X510)))) | s(bool,'F_0') = s(bool,X508)) )),
  inference(forward_demodulation,[],[f26771,f413])).
fof(f26771,plain,(
  ( ! [X509,X508,X510] : (p(s(bool,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X508),s(bool,X509))),s(bool,'T_0'),s(bool,X510)))) | s(bool,'F_0') = s(bool,X508)) )),
  inference(subsumption_resolution,[],[f26640,f335])).
fof(f26640,plain,(
  ( ! [X509,X508,X510] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X508),s(bool,X509))),s(bool,'T_0'),s(bool,X510)))) | s(bool,'F_0') = s(bool,X508)) )),
  inference(superposition,[],[f4853,f787])).
fof(f26560,plain,(
  ( ! [X300,X301,X299] : (sP4('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X301),s(bool,'T_0')),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X299),s(bool,X300))) | s(bool,'F_0') = s(bool,X299)) )),
  inference(superposition,[],[f991,f787])).
fof(f26557,plain,(
  ( ! [X292,X290,X291] : (sP4('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X290),s(bool,X291)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X292),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X290)) )),
  inference(superposition,[],[f930,f787])).
fof(f26555,plain,(
  ( ! [X285,X284,X286] : (sP4('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X286)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X284),s(bool,X285))) | s(bool,'F_0') = s(bool,X284)) )),
  inference(superposition,[],[f898,f787])).
fof(f26768,plain,(
  ( ! [X280,X282,X279,X281,X283] : (~p(s(bool,X281)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X282),s(bool,X281)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X279),s(bool,X280)),X283) | s(bool,'F_0') = s(bool,X279)) )),
  inference(subsumption_resolution,[],[f26554,f335])).
fof(f26554,plain,(
  ( ! [X280,X282,X279,X281,X283] : (~p(s(bool,'T_0')) | ~p(s(bool,X281)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X282),s(bool,X281)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X279),s(bool,X280)),X283) | s(bool,'F_0') = s(bool,X279)) )),
  inference(superposition,[],[f868,f787])).
fof(f26767,plain,(
  ( ! [X277,X276,X278,X274] : (sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X277),s(bool,X274)),X278,X276) | p(s(bool,X276)) | s(bool,'F_0') = s(bool,X274)) )),
  inference(forward_demodulation,[],[f26766,f413])).
fof(f26766,plain,(
  ( ! [X277,X275,X276,X278,X274] : (p(s(bool,X276)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X277),s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X274),s(bool,X275)))),X278,X276) | s(bool,'F_0') = s(bool,X274)) )),
  inference(subsumption_resolution,[],[f26553,f335])).
fof(f26553,plain,(
  ( ! [X277,X275,X276,X278,X274] : (~p(s(bool,'T_0')) | p(s(bool,X276)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X277),s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X274),s(bool,X275)))),X278,X276) | s(bool,'F_0') = s(bool,X274)) )),
  inference(superposition,[],[f867,f787])).
fof(f26765,plain,(
  ( ! [X269,X271,X273,X270,X272] : (~p(s(bool,X271)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X272),s(bool,X271)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X269),s(bool,X270)),X273) | s(bool,'F_0') = s(bool,X269)) )),
  inference(subsumption_resolution,[],[f26552,f335])).
fof(f26552,plain,(
  ( ! [X269,X271,X273,X270,X272] : (~p(s(bool,'T_0')) | ~p(s(bool,X271)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X272),s(bool,X271)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X269),s(bool,X270)),X273) | s(bool,'F_0') = s(bool,X269)) )),
  inference(superposition,[],[f866,f787])).
fof(f26764,plain,(
  ( ! [X265,X267,X268,X264,X266] : (~p(s(bool,X266)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X267),s(bool,X266)),X268,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X264),s(bool,X265))) | s(bool,'F_0') = s(bool,X264)) )),
  inference(subsumption_resolution,[],[f26551,f335])).
fof(f26551,plain,(
  ( ! [X265,X267,X268,X264,X266] : (~p(s(bool,'T_0')) | ~p(s(bool,X266)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X267),s(bool,X266)),X268,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X264),s(bool,X265))) | s(bool,'F_0') = s(bool,X264)) )),
  inference(superposition,[],[f865,f787])).
fof(f26763,plain,(
  ( ! [X263,X260,X262] : (sP1(X262,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X263),s(bool,X260))) | p(s(bool,X262)) | s(bool,'F_0') = s(bool,X260)) )),
  inference(forward_demodulation,[],[f26762,f413])).
fof(f26762,plain,(
  ( ! [X261,X263,X260,X262] : (p(s(bool,X262)) | sP1(X262,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X263),s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X260),s(bool,X261))))) | s(bool,'F_0') = s(bool,X260)) )),
  inference(subsumption_resolution,[],[f26550,f335])).
fof(f26550,plain,(
  ( ! [X261,X263,X260,X262] : (~p(s(bool,'T_0')) | p(s(bool,X262)) | sP1(X262,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X263),s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X260),s(bool,X261))))) | s(bool,'F_0') = s(bool,X260)) )),
  inference(superposition,[],[f864,f787])).
fof(f26761,plain,(
  ( ! [X259,X256,X258] : (sP1('const.bool.COND_3'(s(bool,'F_0'),s(bool,X259),s(bool,X256)),X258) | p(s(bool,X258)) | s(bool,'F_0') = s(bool,X256)) )),
  inference(forward_demodulation,[],[f26760,f413])).
fof(f26760,plain,(
  ( ! [X257,X259,X256,X258] : (p(s(bool,X258)) | sP1('const.bool.COND_3'(s(bool,'F_0'),s(bool,X259),s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X256),s(bool,X257)))),X258) | s(bool,'F_0') = s(bool,X256)) )),
  inference(subsumption_resolution,[],[f26549,f335])).
fof(f26549,plain,(
  ( ! [X257,X259,X256,X258] : (~p(s(bool,'T_0')) | p(s(bool,X258)) | sP1('const.bool.COND_3'(s(bool,'F_0'),s(bool,X259),s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X256),s(bool,X257)))),X258) | s(bool,'F_0') = s(bool,X256)) )),
  inference(superposition,[],[f863,f787])).
fof(f26759,plain,(
  ( ! [X255,X254,X252] : (sP0('const.bool.COND_3'(s(bool,'F_0'),s(bool,X254),s(bool,X252)),X255) | s(bool,'F_0') = s(bool,X252)) )),
  inference(forward_demodulation,[],[f26758,f413])).
fof(f26758,plain,(
  ( ! [X255,X253,X254,X252] : (sP0('const.bool.COND_3'(s(bool,'F_0'),s(bool,X254),s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X252),s(bool,X253)))),X255) | s(bool,'F_0') = s(bool,X252)) )),
  inference(subsumption_resolution,[],[f26548,f335])).
fof(f26548,plain,(
  ( ! [X255,X253,X254,X252] : (~p(s(bool,'T_0')) | sP0('const.bool.COND_3'(s(bool,'F_0'),s(bool,X254),s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X252),s(bool,X253)))),X255) | s(bool,'F_0') = s(bool,X252)) )),
  inference(superposition,[],[f862,f787])).
fof(f26757,plain,(
  ( ! [X251,X250,X248] : (sP0(X250,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X251),s(bool,X248))) | s(bool,'F_0') = s(bool,X248)) )),
  inference(forward_demodulation,[],[f26756,f413])).
fof(f26756,plain,(
  ( ! [X251,X249,X250,X248] : (sP0(X250,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X251),s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X248),s(bool,X249))))) | s(bool,'F_0') = s(bool,X248)) )),
  inference(subsumption_resolution,[],[f26547,f335])).
fof(f26547,plain,(
  ( ! [X251,X249,X250,X248] : (~p(s(bool,'T_0')) | sP0(X250,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X251),s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X248),s(bool,X249))))) | s(bool,'F_0') = s(bool,X248)) )),
  inference(superposition,[],[f861,f787])).
fof(f26545,plain,(
  ( ! [X243,X244,X242] : (sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X242),s(bool,X243)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X244),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X242)) )),
  inference(superposition,[],[f839,f787])).
fof(f26543,plain,(
  ( ! [X237,X238,X236] : (~sP5('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X236),s(bool,X237)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X238),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X236)) )),
  inference(superposition,[],[f838,f787])).
fof(f26542,plain,(
  ( ! [X235,X233,X234] : (sP5('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X235),s(bool,'T_0')),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X233),s(bool,X234))) | s(bool,'F_0') = s(bool,X233)) )),
  inference(superposition,[],[f837,f787])).
fof(f26540,plain,(
  ( ! [X229,X227,X228] : (~sP5('F_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X229),s(bool,'T_0')),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X227),s(bool,X228))) | s(bool,'F_0') = s(bool,X227)) )),
  inference(superposition,[],[f836,f787])).
fof(f26537,plain,(
  ( ! [X218,X220,X219] : (sP4('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X218),s(bool,X219)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X220))) | s(bool,'F_0') = s(bool,X218)) )),
  inference(superposition,[],[f825,f787])).
fof(f26755,plain,(
  ( ! [X216,X215,X213,X217,X214] : (~p(s(bool,X215)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X215),s(bool,X216)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X213),s(bool,X214)),X217) | s(bool,'F_0') = s(bool,X213)) )),
  inference(subsumption_resolution,[],[f26536,f335])).
fof(f26536,plain,(
  ( ! [X216,X215,X213,X217,X214] : (~p(s(bool,'T_0')) | ~p(s(bool,X215)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X215),s(bool,X216)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X213),s(bool,X214)),X217) | s(bool,'F_0') = s(bool,X213)) )),
  inference(superposition,[],[f803,f787])).
fof(f26752,plain,(
  ( ! [X206,X204,X207,X205,X203] : (~p(s(bool,X205)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X205),s(bool,X206)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X203),s(bool,X204)),X207) | s(bool,'F_0') = s(bool,X203)) )),
  inference(subsumption_resolution,[],[f26534,f335])).
fof(f26534,plain,(
  ( ! [X206,X204,X207,X205,X203] : (~p(s(bool,'T_0')) | ~p(s(bool,X205)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X205),s(bool,X206)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X203),s(bool,X204)),X207) | s(bool,'F_0') = s(bool,X203)) )),
  inference(superposition,[],[f801,f787])).
fof(f26751,plain,(
  ( ! [X198,X202,X200,X199,X201] : (~p(s(bool,X200)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X200),s(bool,X201)),X202,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X198),s(bool,X199))) | s(bool,'F_0') = s(bool,X198)) )),
  inference(subsumption_resolution,[],[f26533,f335])).
fof(f26533,plain,(
  ( ! [X198,X202,X200,X199,X201] : (~p(s(bool,'T_0')) | ~p(s(bool,X200)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X200),s(bool,X201)),X202,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X198),s(bool,X199))) | s(bool,'F_0') = s(bool,X198)) )),
  inference(superposition,[],[f800,f787])).
fof(f26527,plain,(
  ( ! [X177,X178,X176] : (sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X176),s(bool,X177)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X178))) | s(bool,'F_0') = s(bool,X176)) )),
  inference(superposition,[],[f778,f787])).
fof(f26525,plain,(
  ( ! [X171,X172,X170] : (~sP5('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X170),s(bool,X171)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X172))) | s(bool,'F_0') = s(bool,X170)) )),
  inference(superposition,[],[f777,f787])).
fof(f26523,plain,(
  ( ! [X165,X166,X164] : (sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X166)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X164),s(bool,X165))) | s(bool,'F_0') = s(bool,X164)) )),
  inference(superposition,[],[f776,f787])).
fof(f26521,plain,(
  ( ! [X158,X159,X160] : (~sP5('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X160)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X158),s(bool,X159))) | s(bool,'F_0') = s(bool,X158)) )),
  inference(superposition,[],[f775,f787])).
fof(f26742,plain,(
  ( ! [X144,X142,X143,X141] : (~sP8('const.bool.COND_3'(s(bool,'T_0'),s(bool,X141),s(bool,X142)),X143,X144) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X141),s(bool,X142)),X143,X144) | s(bool,'F_0') = s(bool,X141)) )),
  inference(subsumption_resolution,[],[f26517,f335])).
fof(f26517,plain,(
  ( ! [X144,X142,X143,X141] : (~p(s(bool,'T_0')) | ~sP8('const.bool.COND_3'(s(bool,'T_0'),s(bool,X141),s(bool,X142)),X143,X144) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X141),s(bool,X142)),X143,X144) | s(bool,'F_0') = s(bool,X141)) )),
  inference(superposition,[],[f577,f787])).
fof(f26741,plain,(
  ( ! [X140,X138,X139,X137] : (~sP10('const.bool.COND_3'(s(bool,'T_0'),s(bool,X137),s(bool,X138)),X139,X140) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X137),s(bool,X138)),X139,X140) | s(bool,'F_0') = s(bool,X137)) )),
  inference(subsumption_resolution,[],[f26516,f335])).
fof(f26516,plain,(
  ( ! [X140,X138,X139,X137] : (~p(s(bool,'T_0')) | ~sP10('const.bool.COND_3'(s(bool,'T_0'),s(bool,X137),s(bool,X138)),X139,X140) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X137),s(bool,X138)),X139,X140) | s(bool,'F_0') = s(bool,X137)) )),
  inference(superposition,[],[f576,f787])).
fof(f26740,plain,(
  ( ! [X134,X136,X135,X133] : (~p(s(bool,X135)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X133),s(bool,X134)),X135,X136) | s(bool,'F_0') = s(bool,X133)) )),
  inference(subsumption_resolution,[],[f26515,f335])).
fof(f26515,plain,(
  ( ! [X134,X136,X135,X133] : (~p(s(bool,'T_0')) | ~p(s(bool,X135)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X133),s(bool,X134)),X135,X136) | s(bool,'F_0') = s(bool,X133)) )),
  inference(superposition,[],[f572,f787])).
fof(f26739,plain,(
  ( ! [X132,X130,X131,X129] : (p(s(bool,X131)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X129),s(bool,X130)),X132,X131) | s(bool,'F_0') = s(bool,X129)) )),
  inference(subsumption_resolution,[],[f26514,f335])).
fof(f26514,plain,(
  ( ! [X132,X130,X131,X129] : (~p(s(bool,'T_0')) | p(s(bool,X131)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X129),s(bool,X130)),X132,X131) | s(bool,'F_0') = s(bool,X129)) )),
  inference(superposition,[],[f571,f787])).
fof(f26738,plain,(
  ( ! [X127,X125,X128,X126] : (~p(s(bool,X127)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X125),s(bool,X126)),X127,X128) | s(bool,'F_0') = s(bool,X125)) )),
  inference(subsumption_resolution,[],[f26513,f335])).
fof(f26513,plain,(
  ( ! [X127,X125,X128,X126] : (~p(s(bool,'T_0')) | ~p(s(bool,X127)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X125),s(bool,X126)),X127,X128) | s(bool,'F_0') = s(bool,X125)) )),
  inference(superposition,[],[f555,f787])).
fof(f26737,plain,(
  ( ! [X123,X121,X124,X122] : (~p(s(bool,X123)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X121),s(bool,X122)),X124,X123) | s(bool,'F_0') = s(bool,X121)) )),
  inference(subsumption_resolution,[],[f26512,f335])).
fof(f26512,plain,(
  ( ! [X123,X121,X124,X122] : (~p(s(bool,'T_0')) | ~p(s(bool,X123)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X121),s(bool,X122)),X124,X123) | s(bool,'F_0') = s(bool,X121)) )),
  inference(superposition,[],[f554,f787])).
fof(f26736,plain,(
  ( ! [X118,X120,X119,X117] : (~p(s(bool,X119)) | p(s(bool,X120)) | ~sP7(X120,X119,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X117),s(bool,X118))) | s(bool,'F_0') = s(bool,X117)) )),
  inference(subsumption_resolution,[],[f26511,f335])).
fof(f26511,plain,(
  ( ! [X118,X120,X119,X117] : (~p(s(bool,'T_0')) | ~p(s(bool,X119)) | p(s(bool,X120)) | ~sP7(X120,X119,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X117),s(bool,X118))) | s(bool,'F_0') = s(bool,X117)) )),
  inference(superposition,[],[f541,f787])).
fof(f26735,plain,(
  ( ! [X116,X115,X113] : (p(s(bool,'const.bool.COND_3'(s(bool,X116),s(bool,X113),s(bool,X115)))) | ~p(s(bool,X115)) | s(bool,'F_0') = s(bool,X113)) )),
  inference(forward_demodulation,[],[f26734,f413])).
fof(f26734,plain,(
  ( ! [X116,X114,X115,X113] : (~p(s(bool,X115)) | p(s(bool,'const.bool.COND_3'(s(bool,X116),s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X113),s(bool,X114))),s(bool,X115)))) | s(bool,'F_0') = s(bool,X113)) )),
  inference(subsumption_resolution,[],[f26510,f335])).
fof(f26510,plain,(
  ( ! [X116,X114,X115,X113] : (~p(s(bool,'T_0')) | ~p(s(bool,X115)) | p(s(bool,'const.bool.COND_3'(s(bool,X116),s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X113),s(bool,X114))),s(bool,X115)))) | s(bool,'F_0') = s(bool,X113)) )),
  inference(superposition,[],[f489,f787])).
fof(f26733,plain,(
  ( ! [X111,X109,X112] : (p(s(bool,'const.bool.COND_3'(s(bool,X111),s(bool,X112),s(bool,X109)))) | p(s(bool,X111)) | s(bool,'F_0') = s(bool,X109)) )),
  inference(forward_demodulation,[],[f26732,f413])).
fof(f26732,plain,(
  ( ! [X111,X109,X112,X110] : (p(s(bool,'const.bool.COND_3'(s(bool,X111),s(bool,X112),s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X109),s(bool,X110)))))) | p(s(bool,X111)) | s(bool,'F_0') = s(bool,X109)) )),
  inference(subsumption_resolution,[],[f26509,f335])).
fof(f26509,plain,(
  ( ! [X111,X109,X112,X110] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.COND_3'(s(bool,X111),s(bool,X112),s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X109),s(bool,X110)))))) | p(s(bool,X111)) | s(bool,'F_0') = s(bool,X109)) )),
  inference(superposition,[],[f488,f787])).
fof(f26731,plain,(
  ( ! [X107,X105,X108] : (p(s(bool,'const.bool.COND_3'(s(bool,X107),s(bool,X105),s(bool,X108)))) | ~p(s(bool,X107)) | s(bool,'F_0') = s(bool,X105)) )),
  inference(forward_demodulation,[],[f26730,f413])).
fof(f26730,plain,(
  ( ! [X107,X105,X108,X106] : (~p(s(bool,X107)) | p(s(bool,'const.bool.COND_3'(s(bool,X107),s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X105),s(bool,X106))),s(bool,X108)))) | s(bool,'F_0') = s(bool,X105)) )),
  inference(subsumption_resolution,[],[f26508,f335])).
fof(f26508,plain,(
  ( ! [X107,X105,X108,X106] : (~p(s(bool,'T_0')) | ~p(s(bool,X107)) | p(s(bool,'const.bool.COND_3'(s(bool,X107),s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X105),s(bool,X106))),s(bool,X108)))) | s(bool,'F_0') = s(bool,X105)) )),
  inference(superposition,[],[f487,f787])).
fof(f26723,plain,(
  ( ! [X92,X91,X89] : (~p(s(bool,'const.bool.COND_3'(s(bool,X89),s(bool,X91),s(bool,X92)))) | p(s(bool,X91)) | s(bool,'F_0') = s(bool,X89)) )),
  inference(forward_demodulation,[],[f26722,f413])).
fof(f26722,plain,(
  ( ! [X92,X90,X91,X89] : (p(s(bool,X91)) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X89),s(bool,X90))),s(bool,X91),s(bool,X92)))) | s(bool,'F_0') = s(bool,X89)) )),
  inference(subsumption_resolution,[],[f26504,f335])).
fof(f26504,plain,(
  ( ! [X92,X90,X91,X89] : (~p(s(bool,'T_0')) | p(s(bool,X91)) | ~p(s(bool,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X89),s(bool,X90))),s(bool,X91),s(bool,X92)))) | s(bool,'F_0') = s(bool,X89)) )),
  inference(superposition,[],[f484,f787])).
fof(f26721,plain,(
  ( ! [X74,X72,X73] : (p(s(bool,X74)) | sP1(X74,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X72),s(bool,X73))) | s(bool,'F_0') = s(bool,X72)) )),
  inference(subsumption_resolution,[],[f26501,f335])).
fof(f26501,plain,(
  ( ! [X74,X72,X73] : (~p(s(bool,'T_0')) | p(s(bool,X74)) | sP1(X74,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X72),s(bool,X73))) | s(bool,'F_0') = s(bool,X72)) )),
  inference(superposition,[],[f462,f787])).
fof(f26720,plain,(
  ( ! [X70,X71,X69] : (p(s(bool,X71)) | sP1('const.bool.COND_3'(s(bool,'T_0'),s(bool,X69),s(bool,X70)),X71) | s(bool,'F_0') = s(bool,X69)) )),
  inference(subsumption_resolution,[],[f26500,f335])).
fof(f26500,plain,(
  ( ! [X70,X71,X69] : (~p(s(bool,'T_0')) | p(s(bool,X71)) | sP1('const.bool.COND_3'(s(bool,'T_0'),s(bool,X69),s(bool,X70)),X71) | s(bool,'F_0') = s(bool,X69)) )),
  inference(superposition,[],[f459,f787])).
fof(f26719,plain,(
  ( ! [X68,X66,X67] : (sP0('const.bool.COND_3'(s(bool,'T_0'),s(bool,X66),s(bool,X67)),X68) | s(bool,'F_0') = s(bool,X66)) )),
  inference(subsumption_resolution,[],[f26499,f335])).
fof(f26499,plain,(
  ( ! [X68,X66,X67] : (~p(s(bool,'T_0')) | sP0('const.bool.COND_3'(s(bool,'T_0'),s(bool,X66),s(bool,X67)),X68) | s(bool,'F_0') = s(bool,X66)) )),
  inference(superposition,[],[f444,f787])).
fof(f26718,plain,(
  ( ! [X64,X65,X63] : (sP0(X65,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X63),s(bool,X64))) | s(bool,'F_0') = s(bool,X63)) )),
  inference(subsumption_resolution,[],[f26498,f335])).
fof(f26498,plain,(
  ( ! [X64,X65,X63] : (~p(s(bool,'T_0')) | sP0(X65,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X63),s(bool,X64))) | s(bool,'F_0') = s(bool,X63)) )),
  inference(superposition,[],[f443,f787])).
fof(f26717,plain,(
  ( ! [X62,X60] : (~p(s(bool,X62)) | s(bool,X60) = s(bool,X62) | s(bool,'F_0') = s(bool,X60)) )),
  inference(forward_demodulation,[],[f26716,f413])).
fof(f26716,plain,(
  ( ! [X61,X62,X60] : (~p(s(bool,X62)) | s(bool,X62) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X60),s(bool,X61))) | s(bool,'F_0') = s(bool,X60)) )),
  inference(subsumption_resolution,[],[f26497,f335])).
fof(f26497,plain,(
  ( ! [X61,X62,X60] : (~p(s(bool,'T_0')) | ~p(s(bool,X62)) | s(bool,X62) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X60),s(bool,X61))) | s(bool,'F_0') = s(bool,X60)) )),
  inference(superposition,[],[f437,f787])).
fof(f26714,plain,(
  ( ! [X45,X43,X41,X38,X44,X42,X40] : (s(X41,'const.bool.COND_3'(s(bool,X38),s(X41,X45),s(X41,X43))) = s(X41,'const.bool.COND_3'(s(bool,X40),s(X41,X44),s(X41,X42))) | s(bool,'T_0') != s(bool,X40) | s(X41,X42) != s(X41,X43) | s(X41,X44) != s(X41,X45) | s(bool,'F_0') = s(bool,X38)) )),
  inference(forward_demodulation,[],[f26493,f413])).
fof(f26493,plain,(
  ( ! [X39,X45,X43,X41,X38,X44,X42,X40] : (s(bool,'T_0') != s(bool,X40) | s(X41,X42) != s(X41,X43) | s(X41,X44) != s(X41,X45) | s(X41,'const.bool.COND_3'(s(bool,X40),s(X41,X44),s(X41,X42))) = s(X41,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X38),s(bool,X39))),s(X41,X45),s(X41,X43))) | s(bool,'F_0') = s(bool,X38)) )),
  inference(superposition,[],[f428,f787])).
fof(f26711,plain,(
  ( ! [X14,X21,X19,X17,X15,X20,X18,X16] : (s(X16,'const.bool.COND_3'(s(bool,X14),s(X16,X20),s(X16,X21))) = s(X16,'const.bool.COND_3'(s(bool,X17),s(X16,X18),s(X16,X19))) | s(X16,X18) != s(X16,X20) | s(bool,X17) != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X14),s(bool,X15))) | s(bool,'F_0') = s(bool,X14)) )),
  inference(forward_demodulation,[],[f26710,f413])).
fof(f26710,plain,(
  ( ! [X14,X21,X19,X17,X15,X20,X18,X16] : (s(X16,'const.bool.COND_3'(s(bool,X17),s(X16,X18),s(X16,X19))) = s(X16,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X14),s(bool,X15))),s(X16,X20),s(X16,X21))) | s(X16,X18) != s(X16,X20) | s(bool,X17) != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X14),s(bool,X15))) | s(bool,'F_0') = s(bool,X14)) )),
  inference(subsumption_resolution,[],[f26490,f335])).
fof(f26490,plain,(
  ( ! [X14,X21,X19,X17,X15,X20,X18,X16] : (~p(s(bool,'T_0')) | s(X16,'const.bool.COND_3'(s(bool,X17),s(X16,X18),s(X16,X19))) = s(X16,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X14),s(bool,X15))),s(X16,X20),s(X16,X21))) | s(X16,X18) != s(X16,X20) | s(bool,X17) != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X14),s(bool,X15))) | s(bool,'F_0') = s(bool,X14)) )),
  inference(superposition,[],[f426,f787])).
fof(f26709,plain,(
  ( ! [X10,X11] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X10),s(bool,X11))) | s(bool,'F_0') = s(bool,X10)) )),
  inference(subsumption_resolution,[],[f26488,f335])).
fof(f26488,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X10),s(bool,X11))) | s(bool,'F_0') = s(bool,X10)) )),
  inference(superposition,[],[f358,f787])).
fof(f26461,plain,(
  ( ! [X39,X43,X41,X44,X42,X40] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,happ(s(fun(X40,bool),'const.bool.COND_3'(s(bool,X39),s(fun(X40,bool),X41),s(fun(X40,bool),X43))),s(X40,X42))),s(bool,X44))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X39),s(bool,happ(s(fun(X40,bool),X41),s(X40,X42))),s(bool,happ(s(fun(X40,bool),X43),s(X40,X42)))))) )),
  inference(superposition,[],[f787,f480])).
fof(f26460,plain,(
  ( ! [X37,X35,X33,X38,X36,X34] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,happ(s(fun(X34,bool),X35),s(X34,'const.bool.COND_3'(s(bool,X33),s(X34,X36),s(X34,X37))))),s(bool,X38))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X33),s(bool,happ(s(fun(X34,bool),X35),s(X34,X36))),s(bool,happ(s(fun(X34,bool),X35),s(X34,X37)))))) )),
  inference(superposition,[],[f787,f479])).
fof(f26459,plain,(
  ( ! [X30,X31,X32] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X31),s(bool,X32))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X30),s(bool,X31)))) )),
  inference(superposition,[],[f787,f414])).
fof(f26458,plain,(
  ( ! [X28,X29,X27] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X27),s(bool,X29))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X27),s(bool,X28)))) )),
  inference(superposition,[],[f787,f413])).
fof(f26442,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X2))) | s(bool,'F_0') = s(bool,X1) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f787,f701])).
fof(f26427,plain,(
  ( ! [X2,X0,X3,X1] : (sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,X0)),'T_0',X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f26178,f701])).
fof(f26408,plain,(
  ( ! [X2,X0,X1] : (sP9('const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,'T_0')),'T_0',X2) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f26178,f343])).
fof(f26391,plain,(
  ( ! [X2,X0,X3,X1] : (sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,X0)),X2,'T_0') | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f26118,f701])).
fof(f26372,plain,(
  ( ! [X2,X0,X1] : (sP9('const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,'T_0')),X2,'T_0') | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f26118,f343])).
fof(f26361,plain,(
  ( ! [X61,X66,X64,X62,X60,X67,X65,X63] : (~p(s(bool,happ(s(fun(X61,bool),'const.bool.COND_3'(s(bool,X60),s(fun(X61,bool),X62),s(fun(X61,bool),X64))),s(X61,X63)))) | ~p(s(bool,X65)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X66),s(bool,X65)),'const.bool.COND_3'(s(bool,X60),s(bool,happ(s(fun(X61,bool),X62),s(X61,X63))),s(bool,happ(s(fun(X61,bool),X64),s(X61,X63)))),X67)) )),
  inference(superposition,[],[f868,f480])).
fof(f26360,plain,(
  ( ! [X59,X57,X54,X52,X58,X56,X55,X53] : (~p(s(bool,happ(s(fun(X53,bool),X54),s(X53,'const.bool.COND_3'(s(bool,X52),s(X53,X55),s(X53,X56)))))) | ~p(s(bool,X57)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X58),s(bool,X57)),'const.bool.COND_3'(s(bool,X52),s(bool,happ(s(fun(X53,bool),X54),s(X53,X55))),s(bool,happ(s(fun(X53,bool),X54),s(X53,X56)))),X59)) )),
  inference(superposition,[],[f868,f479])).
fof(f26359,plain,(
  ( ! [X47,X50,X48,X51,X49] : (~p(s(bool,X48)) | ~p(s(bool,X49)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X50),s(bool,X49)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X47),s(bool,X48)),X51)) )),
  inference(superposition,[],[f868,f414])).
fof(f26358,plain,(
  ( ! [X45,X43,X46,X44,X42] : (~p(s(bool,X42)) | ~p(s(bool,X44)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X45),s(bool,X44)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X42),s(bool,X43)),X46)) )),
  inference(superposition,[],[f868,f413])).
fof(f26368,plain,(
  ( ! [X39,X41,X38,X40] : (~p(s(bool,X39)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X40),s(bool,X39)),'const.bool.COND_3'(s(bool,X38),s(bool,'T_0'),s(bool,'T_0')),X41)) )),
  inference(subsumption_resolution,[],[f26357,f335])).
fof(f26357,plain,(
  ( ! [X39,X41,X38,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X39)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X40),s(bool,X39)),'const.bool.COND_3'(s(bool,X38),s(bool,'T_0'),s(bool,'T_0')),X41)) )),
  inference(superposition,[],[f868,f1595])).
fof(f26362,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X1)),X0,X3) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f26349,f335])).
fof(f26349,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X1)),X0,X3) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f868,f701])).
fof(f26347,plain,(
  ( ! [X0,X1] : (sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')),'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f335,f335,f868])).
fof(f26320,plain,(
  ( ! [X2,X0,X3,X1] : (sP11('const.bool.COND_3'(s(bool,X0),s(bool,X0),s(bool,X1)),'T_0',X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f26090,f701])).
fof(f26290,plain,(
  ( ! [X2,X0,X3,X1] : (sP11('const.bool.COND_3'(s(bool,X0),s(bool,X0),s(bool,X1)),X2,'F_0') | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f26053,f701])).
fof(f26282,plain,(
  ( ! [X39,X41,X40] : (sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X40),s(bool,'T_0')),X41,X39) | p(s(bool,X39))) )),
  inference(forward_demodulation,[],[f26281,f1595])).
fof(f26281,plain,(
  ( ! [X39,X41,X38,X40] : (p(s(bool,X39)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X40),s(bool,'const.bool.COND_3'(s(bool,X38),s(bool,'T_0'),s(bool,'T_0')))),X41,X39)) )),
  inference(subsumption_resolution,[],[f26270,f335])).
fof(f26270,plain,(
  ( ! [X39,X41,X38,X40] : (~p(s(bool,'T_0')) | p(s(bool,X39)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X40),s(bool,'const.bool.COND_3'(s(bool,X38),s(bool,'T_0'),s(bool,'T_0')))),X41,X39)) )),
  inference(superposition,[],[f867,f1595])).
fof(f26275,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X0)),X3,X1) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f26262,f335])).
fof(f26262,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X0)),X3,X1) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f867,f701])).
fof(f26260,plain,(
  ( ! [X0,X1] : (sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')),X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f335,f592,f867])).
fof(f26233,plain,(
  ( ! [X2,X0,X3,X1] : (sP9('const.bool.COND_3'(s(bool,X0),s(bool,X0),s(bool,X1)),'T_0',X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f26024,f701])).
fof(f26203,plain,(
  ( ! [X2,X0,X3,X1] : (sP9('const.bool.COND_3'(s(bool,X0),s(bool,X0),s(bool,X1)),X2,'T_0') | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f25998,f701])).
fof(f26192,plain,(
  ( ! [X61,X66,X64,X62,X60,X67,X65,X63] : (~p(s(bool,happ(s(fun(X61,bool),'const.bool.COND_3'(s(bool,X60),s(fun(X61,bool),X62),s(fun(X61,bool),X64))),s(X61,X63)))) | ~p(s(bool,X65)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X66),s(bool,X65)),'const.bool.COND_3'(s(bool,X60),s(bool,happ(s(fun(X61,bool),X62),s(X61,X63))),s(bool,happ(s(fun(X61,bool),X64),s(X61,X63)))),X67)) )),
  inference(superposition,[],[f866,f480])).
fof(f26191,plain,(
  ( ! [X59,X57,X54,X52,X58,X56,X55,X53] : (~p(s(bool,happ(s(fun(X53,bool),X54),s(X53,'const.bool.COND_3'(s(bool,X52),s(X53,X55),s(X53,X56)))))) | ~p(s(bool,X57)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X58),s(bool,X57)),'const.bool.COND_3'(s(bool,X52),s(bool,happ(s(fun(X53,bool),X54),s(X53,X55))),s(bool,happ(s(fun(X53,bool),X54),s(X53,X56)))),X59)) )),
  inference(superposition,[],[f866,f479])).
fof(f26190,plain,(
  ( ! [X47,X50,X48,X51,X49] : (~p(s(bool,X48)) | ~p(s(bool,X49)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X50),s(bool,X49)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X47),s(bool,X48)),X51)) )),
  inference(superposition,[],[f866,f414])).
fof(f26189,plain,(
  ( ! [X45,X43,X46,X44,X42] : (~p(s(bool,X42)) | ~p(s(bool,X44)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X45),s(bool,X44)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X42),s(bool,X43)),X46)) )),
  inference(superposition,[],[f866,f413])).
fof(f26199,plain,(
  ( ! [X39,X41,X38,X40] : (~p(s(bool,X39)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X40),s(bool,X39)),'const.bool.COND_3'(s(bool,X38),s(bool,'T_0'),s(bool,'T_0')),X41)) )),
  inference(subsumption_resolution,[],[f26188,f335])).
fof(f26188,plain,(
  ( ! [X39,X41,X38,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X39)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X40),s(bool,X39)),'const.bool.COND_3'(s(bool,X38),s(bool,'T_0'),s(bool,'T_0')),X41)) )),
  inference(superposition,[],[f866,f1595])).
fof(f26193,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X1)),X0,X3) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f26180,f335])).
fof(f26180,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X1)),X0,X3) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f866,f701])).
fof(f26178,plain,(
  ( ! [X0,X1] : (sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')),'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f335,f335,f866])).
fof(f26160,plain,(
  ( ! [X2,X0,X1] : (sP1('F_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,X0))) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f23129,f701])).
fof(f26141,plain,(
  ( ! [X0,X1] : (sP1('F_0','const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f23129,f343])).
fof(f26132,plain,(
  ( ! [X61,X66,X64,X62,X60,X67,X65,X63] : (~p(s(bool,happ(s(fun(X61,bool),'const.bool.COND_3'(s(bool,X60),s(fun(X61,bool),X62),s(fun(X61,bool),X64))),s(X61,X63)))) | ~p(s(bool,X65)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X66),s(bool,X65)),X67,'const.bool.COND_3'(s(bool,X60),s(bool,happ(s(fun(X61,bool),X62),s(X61,X63))),s(bool,happ(s(fun(X61,bool),X64),s(X61,X63)))))) )),
  inference(superposition,[],[f865,f480])).
fof(f26131,plain,(
  ( ! [X59,X57,X54,X52,X58,X56,X55,X53] : (~p(s(bool,happ(s(fun(X53,bool),X54),s(X53,'const.bool.COND_3'(s(bool,X52),s(X53,X55),s(X53,X56)))))) | ~p(s(bool,X57)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X58),s(bool,X57)),X59,'const.bool.COND_3'(s(bool,X52),s(bool,happ(s(fun(X53,bool),X54),s(X53,X55))),s(bool,happ(s(fun(X53,bool),X54),s(X53,X56)))))) )),
  inference(superposition,[],[f865,f479])).
fof(f26130,plain,(
  ( ! [X47,X50,X48,X51,X49] : (~p(s(bool,X48)) | ~p(s(bool,X49)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X50),s(bool,X49)),X51,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X47),s(bool,X48)))) )),
  inference(superposition,[],[f865,f414])).
fof(f26129,plain,(
  ( ! [X45,X43,X46,X44,X42] : (~p(s(bool,X42)) | ~p(s(bool,X44)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X45),s(bool,X44)),X46,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X42),s(bool,X43)))) )),
  inference(superposition,[],[f865,f413])).
fof(f26139,plain,(
  ( ! [X39,X41,X38,X40] : (~p(s(bool,X39)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X40),s(bool,X39)),X41,'const.bool.COND_3'(s(bool,X38),s(bool,'T_0'),s(bool,'T_0')))) )),
  inference(subsumption_resolution,[],[f26128,f335])).
fof(f26128,plain,(
  ( ! [X39,X41,X38,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X39)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X40),s(bool,X39)),X41,'const.bool.COND_3'(s(bool,X38),s(bool,'T_0'),s(bool,'T_0')))) )),
  inference(superposition,[],[f865,f1595])).
fof(f26133,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X1)),X3,X0) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f26120,f335])).
fof(f26120,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X1)),X3,X0) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f865,f701])).
fof(f26118,plain,(
  ( ! [X0,X1] : (sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')),X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f335,f335,f865])).
fof(f26104,plain,(
  ( ! [X61,X66,X64,X62,X60,X67,X65,X63] : (~p(s(bool,happ(s(fun(X61,bool),'const.bool.COND_3'(s(bool,X60),s(fun(X61,bool),X62),s(fun(X61,bool),X64))),s(X61,X63)))) | ~p(s(bool,X65)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X65),s(bool,X66)),'const.bool.COND_3'(s(bool,X60),s(bool,happ(s(fun(X61,bool),X62),s(X61,X63))),s(bool,happ(s(fun(X61,bool),X64),s(X61,X63)))),X67)) )),
  inference(superposition,[],[f803,f480])).
fof(f26103,plain,(
  ( ! [X59,X57,X54,X52,X58,X56,X55,X53] : (~p(s(bool,happ(s(fun(X53,bool),X54),s(X53,'const.bool.COND_3'(s(bool,X52),s(X53,X55),s(X53,X56)))))) | ~p(s(bool,X57)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X57),s(bool,X58)),'const.bool.COND_3'(s(bool,X52),s(bool,happ(s(fun(X53,bool),X54),s(X53,X55))),s(bool,happ(s(fun(X53,bool),X54),s(X53,X56)))),X59)) )),
  inference(superposition,[],[f803,f479])).
fof(f26102,plain,(
  ( ! [X47,X50,X48,X51,X49] : (~p(s(bool,X48)) | ~p(s(bool,X49)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X49),s(bool,X50)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X47),s(bool,X48)),X51)) )),
  inference(superposition,[],[f803,f414])).
fof(f26101,plain,(
  ( ! [X45,X43,X46,X44,X42] : (~p(s(bool,X42)) | ~p(s(bool,X44)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X44),s(bool,X45)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X42),s(bool,X43)),X46)) )),
  inference(superposition,[],[f803,f413])).
fof(f26111,plain,(
  ( ! [X39,X41,X38,X40] : (~p(s(bool,X39)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X39),s(bool,X40)),'const.bool.COND_3'(s(bool,X38),s(bool,'T_0'),s(bool,'T_0')),X41)) )),
  inference(subsumption_resolution,[],[f26100,f335])).
fof(f26100,plain,(
  ( ! [X39,X41,X38,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X39)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X39),s(bool,X40)),'const.bool.COND_3'(s(bool,X38),s(bool,'T_0'),s(bool,'T_0')),X41)) )),
  inference(superposition,[],[f803,f1595])).
fof(f26105,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2)),X0,X3) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f26092,f335])).
fof(f26092,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2)),X0,X3) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f803,f701])).
fof(f26090,plain,(
  ( ! [X0,X1] : (sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)),'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f335,f335,f803])).
fof(f26075,plain,(
  ( ! [X39,X41,X40] : (sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X40)),X41,X39) | p(s(bool,X39))) )),
  inference(forward_demodulation,[],[f26074,f1595])).
fof(f26074,plain,(
  ( ! [X39,X41,X38,X40] : (p(s(bool,X39)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,X38),s(bool,'T_0'),s(bool,'T_0'))),s(bool,X40)),X41,X39)) )),
  inference(subsumption_resolution,[],[f26063,f335])).
fof(f26063,plain,(
  ( ! [X39,X41,X38,X40] : (~p(s(bool,'T_0')) | p(s(bool,X39)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,X38),s(bool,'T_0'),s(bool,'T_0'))),s(bool,X40)),X41,X39)) )),
  inference(superposition,[],[f802,f1595])).
fof(f26068,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X2)),X3,X1) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f26055,f335])).
fof(f26055,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X2)),X3,X1) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f802,f701])).
fof(f26053,plain,(
  ( ! [X0,X1] : (sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)),X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f335,f592,f802])).
fof(f26038,plain,(
  ( ! [X61,X66,X64,X62,X60,X67,X65,X63] : (~p(s(bool,happ(s(fun(X61,bool),'const.bool.COND_3'(s(bool,X60),s(fun(X61,bool),X62),s(fun(X61,bool),X64))),s(X61,X63)))) | ~p(s(bool,X65)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X65),s(bool,X66)),'const.bool.COND_3'(s(bool,X60),s(bool,happ(s(fun(X61,bool),X62),s(X61,X63))),s(bool,happ(s(fun(X61,bool),X64),s(X61,X63)))),X67)) )),
  inference(superposition,[],[f801,f480])).
fof(f26037,plain,(
  ( ! [X59,X57,X54,X52,X58,X56,X55,X53] : (~p(s(bool,happ(s(fun(X53,bool),X54),s(X53,'const.bool.COND_3'(s(bool,X52),s(X53,X55),s(X53,X56)))))) | ~p(s(bool,X57)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X57),s(bool,X58)),'const.bool.COND_3'(s(bool,X52),s(bool,happ(s(fun(X53,bool),X54),s(X53,X55))),s(bool,happ(s(fun(X53,bool),X54),s(X53,X56)))),X59)) )),
  inference(superposition,[],[f801,f479])).
fof(f26036,plain,(
  ( ! [X47,X50,X48,X51,X49] : (~p(s(bool,X48)) | ~p(s(bool,X49)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X49),s(bool,X50)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X47),s(bool,X48)),X51)) )),
  inference(superposition,[],[f801,f414])).
fof(f26035,plain,(
  ( ! [X45,X43,X46,X44,X42] : (~p(s(bool,X42)) | ~p(s(bool,X44)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X44),s(bool,X45)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X42),s(bool,X43)),X46)) )),
  inference(superposition,[],[f801,f413])).
fof(f26045,plain,(
  ( ! [X39,X41,X38,X40] : (~p(s(bool,X39)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X39),s(bool,X40)),'const.bool.COND_3'(s(bool,X38),s(bool,'T_0'),s(bool,'T_0')),X41)) )),
  inference(subsumption_resolution,[],[f26034,f335])).
fof(f26034,plain,(
  ( ! [X39,X41,X38,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X39)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X39),s(bool,X40)),'const.bool.COND_3'(s(bool,X38),s(bool,'T_0'),s(bool,'T_0')),X41)) )),
  inference(superposition,[],[f801,f1595])).
fof(f26039,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2)),X0,X3) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f26026,f335])).
fof(f26026,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2)),X0,X3) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f801,f701])).
fof(f26024,plain,(
  ( ! [X0,X1] : (sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)),'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f335,f335,f801])).
fof(f26012,plain,(
  ( ! [X61,X66,X64,X62,X60,X67,X65,X63] : (~p(s(bool,happ(s(fun(X61,bool),'const.bool.COND_3'(s(bool,X60),s(fun(X61,bool),X62),s(fun(X61,bool),X64))),s(X61,X63)))) | ~p(s(bool,X65)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X65),s(bool,X66)),X67,'const.bool.COND_3'(s(bool,X60),s(bool,happ(s(fun(X61,bool),X62),s(X61,X63))),s(bool,happ(s(fun(X61,bool),X64),s(X61,X63)))))) )),
  inference(superposition,[],[f800,f480])).
fof(f26011,plain,(
  ( ! [X59,X57,X54,X52,X58,X56,X55,X53] : (~p(s(bool,happ(s(fun(X53,bool),X54),s(X53,'const.bool.COND_3'(s(bool,X52),s(X53,X55),s(X53,X56)))))) | ~p(s(bool,X57)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X57),s(bool,X58)),X59,'const.bool.COND_3'(s(bool,X52),s(bool,happ(s(fun(X53,bool),X54),s(X53,X55))),s(bool,happ(s(fun(X53,bool),X54),s(X53,X56)))))) )),
  inference(superposition,[],[f800,f479])).
fof(f26010,plain,(
  ( ! [X47,X50,X48,X51,X49] : (~p(s(bool,X48)) | ~p(s(bool,X49)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X49),s(bool,X50)),X51,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X47),s(bool,X48)))) )),
  inference(superposition,[],[f800,f414])).
fof(f26009,plain,(
  ( ! [X45,X43,X46,X44,X42] : (~p(s(bool,X42)) | ~p(s(bool,X44)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X44),s(bool,X45)),X46,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X42),s(bool,X43)))) )),
  inference(superposition,[],[f800,f413])).
fof(f26019,plain,(
  ( ! [X39,X41,X38,X40] : (~p(s(bool,X39)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X39),s(bool,X40)),X41,'const.bool.COND_3'(s(bool,X38),s(bool,'T_0'),s(bool,'T_0')))) )),
  inference(subsumption_resolution,[],[f26008,f335])).
fof(f26008,plain,(
  ( ! [X39,X41,X38,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X39)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X39),s(bool,X40)),X41,'const.bool.COND_3'(s(bool,X38),s(bool,'T_0'),s(bool,'T_0')))) )),
  inference(superposition,[],[f800,f1595])).
fof(f26013,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2)),X3,X0) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f26000,f335])).
fof(f26000,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2)),X3,X0) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f800,f701])).
fof(f25998,plain,(
  ( ! [X0,X1] : (sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)),X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f335,f335,f800])).
fof(f24431,plain,(
  ( ! [X1006,X1004,X1003,X1005] : (s(X1005,f(s(bool,'T_0'))) != s(X1005,g(s(bool,'T_0'))) | s(fun(bool,X1005),X1003) = s(fun(bool,X1005),X1004) | s(bool,X1006) = s(bool,sK22(X1003,X1004,X1005,bool)) | s(bool,'T_0') = s(bool,X1006)) )),
  inference(superposition,[],[f478,f701])).
fof(f24430,plain,(
  ( ! [X1002,X1000,X1001,X999] : (~p(s(bool,happ(s(fun(bool,bool),X999),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X1000),s(bool,X1001)))) | ~p(s(bool,happ(s(fun(bool,bool),X1000),s(bool,sK17(X999,X1000,bool))))) | s(bool,X1002) = s(bool,sK18(X999,X1000,bool)) | s(bool,'T_0') = s(bool,X1002)) )),
  inference(superposition,[],[f419,f701])).
fof(f24429,plain,(
  ( ! [X998,X996,X995,X997] : (~p(s(bool,happ(s(fun(bool,bool),X995),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X995),s(bool,X997)))) | ~p(s(bool,happ(s(fun(bool,bool),X996),s(bool,sK17(X995,X996,bool))))) | s(bool,X998) = s(bool,sK18(X995,X996,bool)) | s(bool,'T_0') = s(bool,X998)) )),
  inference(superposition,[],[f420,f701])).
fof(f24428,plain,(
  ( ! [X994,X992,X991,X993] : (~p(s(bool,happ(s(fun(bool,bool),X991),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X992),s(bool,X993)))) | ~p(s(bool,happ(s(fun(bool,bool),X992),s(bool,sK16(X991,X992,bool))))) | s(bool,X994) = s(bool,sK16(X991,X992,bool)) | s(bool,'T_0') = s(bool,X994)) )),
  inference(superposition,[],[f417,f701])).
fof(f24427,plain,(
  ( ! [X990,X988,X987,X989] : (~p(s(bool,happ(s(fun(bool,bool),X987),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X987),s(bool,X989)))) | ~p(s(bool,happ(s(fun(bool,bool),X988),s(bool,sK16(X987,X988,bool))))) | s(bool,X990) = s(bool,sK16(X987,X988,bool)) | s(bool,'T_0') = s(bool,X990)) )),
  inference(superposition,[],[f418,f701])).
fof(f24426,plain,(
  ( ! [X984,X985,X986] : (~p(s(bool,happ(s(fun(bool,bool),X984),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X984),s(bool,X985)))) | s(bool,X986) = s(bool,sK15(X984,bool)) | s(bool,'T_0') = s(bool,X986)) )),
  inference(superposition,[],[f412,f701])).
fof(f24425,plain,(
  ( ! [X982,X983,X981] : (~p(s(bool,happ(s(fun(bool,bool),X981),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X981),s(bool,X982)))) | s(bool,X983) = s(bool,sK14(X981,bool)) | s(bool,'T_0') = s(bool,X983)) )),
  inference(superposition,[],[f411,f701])).
fof(f24424,plain,(
  ( ! [X978,X976,X980,X975,X979,X977] : (s(bool,'const.bool.COND_3'(s(bool,X978),s(bool,'T_0'),s(bool,happ(s(fun(X975,bool),X976),s(X975,X979))))) = s(bool,happ(s(fun(X975,bool),X976),s(X975,'const.bool.COND_3'(s(bool,X978),s(X975,X977),s(X975,X979))))) | s(bool,X980) = s(bool,happ(s(fun(X975,bool),X976),s(X975,X977))) | s(bool,'T_0') = s(bool,X980)) )),
  inference(superposition,[],[f479,f701])).
fof(f24423,plain,(
  ( ! [X971,X969,X973,X970,X974,X972] : (s(bool,'const.bool.COND_3'(s(bool,X972),s(bool,happ(s(fun(X969,bool),X970),s(X969,X973))),s(bool,'T_0'))) = s(bool,happ(s(fun(X969,bool),X970),s(X969,'const.bool.COND_3'(s(bool,X972),s(X969,X973),s(X969,X971))))) | s(bool,X974) = s(bool,happ(s(fun(X969,bool),X970),s(X969,X971))) | s(bool,'T_0') = s(bool,X974)) )),
  inference(superposition,[],[f479,f701])).
fof(f24422,plain,(
  ( ! [X963,X967,X965,X968,X966,X964] : (s(bool,'const.bool.COND_3'(s(bool,X966),s(bool,'T_0'),s(bool,happ(s(fun(X963,bool),X967),s(X963,X965))))) = s(bool,happ(s(fun(X963,bool),'const.bool.COND_3'(s(bool,X966),s(fun(X963,bool),X964),s(fun(X963,bool),X967))),s(X963,X965))) | s(bool,X968) = s(bool,happ(s(fun(X963,bool),X964),s(X963,X965))) | s(bool,'T_0') = s(bool,X968)) )),
  inference(superposition,[],[f480,f701])).
fof(f24421,plain,(
  ( ! [X957,X961,X958,X962,X960,X959] : (s(bool,'const.bool.COND_3'(s(bool,X960),s(bool,happ(s(fun(X957,bool),X961),s(X957,X959))),s(bool,'T_0'))) = s(bool,happ(s(fun(X957,bool),'const.bool.COND_3'(s(bool,X960),s(fun(X957,bool),X961),s(fun(X957,bool),X958))),s(X957,X959))) | s(bool,X962) = s(bool,happ(s(fun(X957,bool),X958),s(X957,X959))) | s(bool,'T_0') = s(bool,X962)) )),
  inference(superposition,[],[f480,f701])).
fof(f24416,plain,(
  ( ! [X938,X936,X937,X935] : (s(bool,'T_0') != s(bool,f(s(X935,sK22(X936,X937,bool,X935)))) | s(fun(X935,bool),X936) = s(fun(X935,bool),X937) | s(bool,X938) = s(bool,g(s(X935,sK22(X936,X937,bool,X935)))) | s(bool,'T_0') = s(bool,X938)) )),
  inference(superposition,[],[f478,f701])).
fof(f25918,plain,(
  ( ! [X883,X881,X885,X882,X884] : (p(s(bool,happ(s(fun(X881,bool),X883),s(X881,X884)))) | ~p(s(bool,happ(s(fun(X881,bool),X883),s(X881,sK17(X882,X883,X881))))) | s(bool,X885) = s(bool,happ(s(fun(X881,bool),X882),s(X881,sK18(X882,X883,X881)))) | s(bool,'T_0') = s(bool,X885)) )),
  inference(subsumption_resolution,[],[f24404,f335])).
fof(f24404,plain,(
  ( ! [X883,X881,X885,X882,X884] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X881,bool),X883),s(X881,X884)))) | ~p(s(bool,happ(s(fun(X881,bool),X883),s(X881,sK17(X882,X883,X881))))) | s(bool,X885) = s(bool,happ(s(fun(X881,bool),X882),s(X881,sK18(X882,X883,X881)))) | s(bool,'T_0') = s(bool,X885)) )),
  inference(superposition,[],[f419,f701])).
fof(f25917,plain,(
  ( ! [X878,X876,X880,X879,X877] : (p(s(bool,happ(s(fun(X876,bool),X877),s(X876,X879)))) | ~p(s(bool,happ(s(fun(X876,bool),X878),s(X876,sK17(X877,X878,X876))))) | s(bool,X880) = s(bool,happ(s(fun(X876,bool),X877),s(X876,sK18(X877,X878,X876)))) | s(bool,'T_0') = s(bool,X880)) )),
  inference(subsumption_resolution,[],[f24403,f335])).
fof(f24403,plain,(
  ( ! [X878,X876,X880,X879,X877] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X876,bool),X877),s(X876,X879)))) | ~p(s(bool,happ(s(fun(X876,bool),X878),s(X876,sK17(X877,X878,X876))))) | s(bool,X880) = s(bool,happ(s(fun(X876,bool),X877),s(X876,sK18(X877,X878,X876)))) | s(bool,'T_0') = s(bool,X880)) )),
  inference(superposition,[],[f420,f701])).
fof(f25916,plain,(
  ( ! [X874,X872,X875,X873,X871] : (p(s(bool,happ(s(fun(X871,bool),X873),s(X871,X874)))) | ~p(s(bool,happ(s(fun(X871,bool),X873),s(X871,sK16(X872,X873,X871))))) | s(bool,X875) = s(bool,happ(s(fun(X871,bool),X872),s(X871,sK16(X872,X873,X871)))) | s(bool,'T_0') = s(bool,X875)) )),
  inference(subsumption_resolution,[],[f24402,f335])).
fof(f24402,plain,(
  ( ! [X874,X872,X875,X873,X871] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X871,bool),X873),s(X871,X874)))) | ~p(s(bool,happ(s(fun(X871,bool),X873),s(X871,sK16(X872,X873,X871))))) | s(bool,X875) = s(bool,happ(s(fun(X871,bool),X872),s(X871,sK16(X872,X873,X871)))) | s(bool,'T_0') = s(bool,X875)) )),
  inference(superposition,[],[f417,f701])).
fof(f25915,plain,(
  ( ! [X866,X870,X868,X867,X869] : (p(s(bool,happ(s(fun(X866,bool),X867),s(X866,X869)))) | ~p(s(bool,happ(s(fun(X866,bool),X868),s(X866,sK16(X867,X868,X866))))) | s(bool,X870) = s(bool,happ(s(fun(X866,bool),X867),s(X866,sK16(X867,X868,X866)))) | s(bool,'T_0') = s(bool,X870)) )),
  inference(subsumption_resolution,[],[f24401,f335])).
fof(f24401,plain,(
  ( ! [X866,X870,X868,X867,X869] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X866,bool),X867),s(X866,X869)))) | ~p(s(bool,happ(s(fun(X866,bool),X868),s(X866,sK16(X867,X868,X866))))) | s(bool,X870) = s(bool,happ(s(fun(X866,bool),X867),s(X866,sK16(X867,X868,X866)))) | s(bool,'T_0') = s(bool,X870)) )),
  inference(superposition,[],[f418,f701])).
fof(f25914,plain,(
  ( ! [X864,X863,X865,X862] : (p(s(bool,happ(s(fun(X862,bool),X863),s(X862,X864)))) | s(bool,X865) = s(bool,happ(s(fun(X862,bool),X863),s(X862,sK15(X863,X862)))) | s(bool,'T_0') = s(bool,X865)) )),
  inference(subsumption_resolution,[],[f24400,f335])).
fof(f24400,plain,(
  ( ! [X864,X863,X865,X862] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X862,bool),X863),s(X862,X864)))) | s(bool,X865) = s(bool,happ(s(fun(X862,bool),X863),s(X862,sK15(X863,X862)))) | s(bool,'T_0') = s(bool,X865)) )),
  inference(superposition,[],[f412,f701])).
fof(f25913,plain,(
  ( ! [X860,X859,X861,X858] : (p(s(bool,happ(s(fun(X858,bool),X859),s(X858,X860)))) | s(bool,X861) = s(bool,happ(s(fun(X858,bool),X859),s(X858,sK14(X859,X858)))) | s(bool,'T_0') = s(bool,X861)) )),
  inference(subsumption_resolution,[],[f24399,f335])).
fof(f24399,plain,(
  ( ! [X860,X859,X861,X858] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X858,bool),X859),s(X858,X860)))) | s(bool,X861) = s(bool,happ(s(fun(X858,bool),X859),s(X858,sK14(X859,X858)))) | s(bool,'T_0') = s(bool,X861)) )),
  inference(superposition,[],[f411,f701])).
fof(f25912,plain,(
  ( ! [X854,X857,X855,X853,X856] : (p(s(bool,happ(s(fun(X853,bool),X854),s(X853,sK19(X856,X854,X853))))) | p(s(bool,happ(s(fun(X853,bool),X856),s(X853,sK20(X856,X854,X853))))) | s(bool,X857) = s(bool,happ(s(fun(X853,bool),X854),s(X853,X855))) | s(bool,'T_0') = s(bool,X857)) )),
  inference(subsumption_resolution,[],[f24398,f335])).
fof(f24398,plain,(
  ( ! [X854,X857,X855,X853,X856] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X853,bool),X854),s(X853,sK19(X856,X854,X853))))) | p(s(bool,happ(s(fun(X853,bool),X856),s(X853,sK20(X856,X854,X853))))) | s(bool,X857) = s(bool,happ(s(fun(X853,bool),X854),s(X853,X855))) | s(bool,'T_0') = s(bool,X857)) )),
  inference(superposition,[],[f421,f701])).
fof(f25911,plain,(
  ( ! [X850,X848,X852,X851,X849] : (p(s(bool,happ(s(fun(X848,bool),X851),s(X848,sK19(X849,X851,X848))))) | p(s(bool,happ(s(fun(X848,bool),X849),s(X848,sK20(X849,X851,X848))))) | s(bool,X852) = s(bool,happ(s(fun(X848,bool),X849),s(X848,X850))) | s(bool,'T_0') = s(bool,X852)) )),
  inference(subsumption_resolution,[],[f24397,f335])).
fof(f24397,plain,(
  ( ! [X850,X848,X852,X851,X849] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X848,bool),X851),s(X848,sK19(X849,X851,X848))))) | p(s(bool,happ(s(fun(X848,bool),X849),s(X848,sK20(X849,X851,X848))))) | s(bool,X852) = s(bool,happ(s(fun(X848,bool),X849),s(X848,X850))) | s(bool,'T_0') = s(bool,X852)) )),
  inference(superposition,[],[f422,f701])).
fof(f25910,plain,(
  ( ! [X843,X847,X845,X846,X844] : (p(s(bool,happ(s(fun(X843,bool),X844),s(X843,sK21(X846,X844,X843))))) | p(s(bool,happ(s(fun(X843,bool),X846),s(X843,sK21(X846,X844,X843))))) | s(bool,X847) = s(bool,happ(s(fun(X843,bool),X844),s(X843,X845))) | s(bool,'T_0') = s(bool,X847)) )),
  inference(subsumption_resolution,[],[f24396,f335])).
fof(f24396,plain,(
  ( ! [X843,X847,X845,X846,X844] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X843,bool),X844),s(X843,sK21(X846,X844,X843))))) | p(s(bool,happ(s(fun(X843,bool),X846),s(X843,sK21(X846,X844,X843))))) | s(bool,X847) = s(bool,happ(s(fun(X843,bool),X844),s(X843,X845))) | s(bool,'T_0') = s(bool,X847)) )),
  inference(superposition,[],[f423,f701])).
fof(f25909,plain,(
  ( ! [X841,X839,X842,X840,X838] : (p(s(bool,happ(s(fun(X838,bool),X841),s(X838,sK21(X839,X841,X838))))) | p(s(bool,happ(s(fun(X838,bool),X839),s(X838,sK21(X839,X841,X838))))) | s(bool,X842) = s(bool,happ(s(fun(X838,bool),X839),s(X838,X840))) | s(bool,'T_0') = s(bool,X842)) )),
  inference(subsumption_resolution,[],[f24395,f335])).
fof(f24395,plain,(
  ( ! [X841,X839,X842,X840,X838] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X838,bool),X841),s(X838,sK21(X839,X841,X838))))) | p(s(bool,happ(s(fun(X838,bool),X839),s(X838,sK21(X839,X841,X838))))) | s(bool,X842) = s(bool,happ(s(fun(X838,bool),X839),s(X838,X840))) | s(bool,'T_0') = s(bool,X842)) )),
  inference(superposition,[],[f424,f701])).
fof(f25908,plain,(
  ( ! [X833,X834,X832] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X833),s('type.num.num',X832)))) | s('type.num.num',X832) = s('type.num.num',X833) | s(bool,X834) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X832),s('type.num.num',X833))) | s(bool,'T_0') = s(bool,X834)) )),
  inference(subsumption_resolution,[],[f24393,f335])).
fof(f24393,plain,(
  ( ! [X833,X834,X832] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X833),s('type.num.num',X832)))) | s('type.num.num',X832) = s('type.num.num',X833) | s(bool,X834) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X832),s('type.num.num',X833))) | s(bool,'T_0') = s(bool,X834)) )),
  inference(superposition,[],[f433,f701])).
fof(f25907,plain,(
  ( ! [X831,X829,X830,X828] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X829),s('type.num.num',X830)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X828),s('type.num.num',X830)))) | s(bool,X831) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X828),s('type.num.num',X829))) | s(bool,'T_0') = s(bool,X831)) )),
  inference(subsumption_resolution,[],[f24392,f335])).
fof(f24392,plain,(
  ( ! [X831,X829,X830,X828] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X829),s('type.num.num',X830)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X828),s('type.num.num',X830)))) | s(bool,X831) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X828),s('type.num.num',X829))) | s(bool,'T_0') = s(bool,X831)) )),
  inference(superposition,[],[f482,f701])).
fof(f25906,plain,(
  ( ! [X823,X821,X822] : (p(s(bool,X821)) | s(bool,X823) = s(bool,'const.bool.COND_3'(s(bool,X821),s(bool,X822),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X823)) )),
  inference(subsumption_resolution,[],[f24390,f335])).
fof(f24390,plain,(
  ( ! [X823,X821,X822] : (~p(s(bool,'T_0')) | p(s(bool,X821)) | s(bool,X823) = s(bool,'const.bool.COND_3'(s(bool,X821),s(bool,X822),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X823)) )),
  inference(superposition,[],[f1494,f701])).
fof(f25905,plain,(
  ( ! [X819,X817,X818,X820] : (p(s(bool,X817)) | p(s(bool,X819)) | s(bool,X820) = s(bool,'const.bool.COND_3'(s(bool,X817),s(bool,X818),s(bool,X819))) | s(bool,'T_0') = s(bool,X820)) )),
  inference(subsumption_resolution,[],[f24389,f335])).
fof(f24389,plain,(
  ( ! [X819,X817,X818,X820] : (~p(s(bool,'T_0')) | p(s(bool,X817)) | p(s(bool,X819)) | s(bool,X820) = s(bool,'const.bool.COND_3'(s(bool,X817),s(bool,X818),s(bool,X819))) | s(bool,'T_0') = s(bool,X820)) )),
  inference(superposition,[],[f485,f701])).
fof(f25904,plain,(
  ( ! [X814,X816,X815] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X815),s('type.num.num',X814)))) | s(bool,X816) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X814),s('type.num.num',X815))) | s(bool,'T_0') = s(bool,X816)) )),
  inference(subsumption_resolution,[],[f24388,f335])).
fof(f24388,plain,(
  ( ! [X814,X816,X815] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X815),s('type.num.num',X814)))) | s(bool,X816) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X814),s('type.num.num',X815))) | s(bool,'T_0') = s(bool,X816)) )),
  inference(superposition,[],[f468,f701])).
fof(f25903,plain,(
  ( ! [X812,X811,X813] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X812),s('type.num.num',X811)))) | s(bool,X813) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X811),s('type.num.num',X812))) | s(bool,'T_0') = s(bool,X813)) )),
  inference(subsumption_resolution,[],[f24387,f335])).
fof(f24387,plain,(
  ( ! [X812,X811,X813] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X812),s('type.num.num',X811)))) | s(bool,X813) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X811),s('type.num.num',X812))) | s(bool,'T_0') = s(bool,X813)) )),
  inference(superposition,[],[f477,f701])).
fof(f25902,plain,(
  ( ! [X810,X808,X809] : (s('type.num.num',X808) = s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X808),s('type.num.num',X809))) | s(bool,X810) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X808),s('type.num.num',X809))) | s(bool,'T_0') = s(bool,X810)) )),
  inference(forward_demodulation,[],[f24386,f413])).
fof(f24386,plain,(
  ( ! [X810,X808,X809] : (s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X808),s('type.num.num',X809))) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',X808),s('type.num.num',X809))) | s(bool,X810) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X808),s('type.num.num',X809))) | s(bool,'T_0') = s(bool,X810)) )),
  inference(superposition,[],[f431,f701])).
fof(f25901,plain,(
  ( ! [X806,X807,X805] : (s('type.num.num',X806) = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X805),s('type.num.num',X806))) | s(bool,X807) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X805),s('type.num.num',X806))) | s(bool,'T_0') = s(bool,X807)) )),
  inference(forward_demodulation,[],[f24385,f413])).
fof(f24385,plain,(
  ( ! [X806,X807,X805] : (s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X805),s('type.num.num',X806))) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',X806),s('type.num.num',X805))) | s(bool,X807) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X805),s('type.num.num',X806))) | s(bool,'T_0') = s(bool,X807)) )),
  inference(superposition,[],[f432,f701])).
fof(f25900,plain,(
  ( ! [X802,X804,X803,X801] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X802),s('type.num.num',X803)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X801),s('type.num.num',X803)))) | s(bool,X804) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X801),s('type.num.num',X802))) | s(bool,'T_0') = s(bool,X804)) )),
  inference(subsumption_resolution,[],[f24384,f335])).
fof(f24384,plain,(
  ( ! [X802,X804,X803,X801] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X802),s('type.num.num',X803)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X801),s('type.num.num',X803)))) | s(bool,X804) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X801),s('type.num.num',X802))) | s(bool,'T_0') = s(bool,X804)) )),
  inference(superposition,[],[f483,f701])).
fof(f25898,plain,(
  ( ! [X463,X465,X464] : (p(s(bool,'const.bool.COND_3'(s(bool,X464),s(bool,'T_0'),s(bool,X463)))) | s(bool,X463) = s(bool,X465) | s(bool,'T_0') = s(bool,X465)) )),
  inference(subsumption_resolution,[],[f24258,f335])).
fof(f24258,plain,(
  ( ! [X463,X465,X464] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.COND_3'(s(bool,X464),s(bool,'T_0'),s(bool,X463)))) | s(bool,X463) = s(bool,X465) | s(bool,'T_0') = s(bool,X465)) )),
  inference(superposition,[],[f4857,f701])).
fof(f25897,plain,(
  ( ! [X455,X454,X456] : (p(s(bool,'const.bool.COND_3'(s(bool,X454),s(bool,'T_0'),s(bool,X455)))) | s(bool,X454) = s(bool,X456) | s(bool,'T_0') = s(bool,X456)) )),
  inference(subsumption_resolution,[],[f24255,f335])).
fof(f24255,plain,(
  ( ! [X455,X454,X456] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.COND_3'(s(bool,X454),s(bool,'T_0'),s(bool,X455)))) | s(bool,X454) = s(bool,X456) | s(bool,'T_0') = s(bool,X456)) )),
  inference(superposition,[],[f4853,f701])).
fof(f24175,plain,(
  ( ! [X246,X247,X245] : (sP4('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X246),s(bool,'T_0')),X245) | s(bool,X245) = s(bool,X247) | s(bool,'T_0') = s(bool,X247)) )),
  inference(superposition,[],[f991,f701])).
fof(f24172,plain,(
  ( ! [X237,X238,X236] : (sP4('T_0',X236,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X237),s(bool,'T_0'))) | s(bool,X236) = s(bool,X238) | s(bool,'T_0') = s(bool,X238)) )),
  inference(superposition,[],[f930,f701])).
fof(f24170,plain,(
  ( ! [X231,X230,X232] : (sP4('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X231)),X230) | s(bool,X230) = s(bool,X232) | s(bool,'T_0') = s(bool,X232)) )),
  inference(superposition,[],[f898,f701])).
fof(f25896,plain,(
  ( ! [X229,X227,X228,X226] : (p(s(bool,X227)) | sP1(X227,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X228),s(bool,X226))) | s(bool,X226) = s(bool,X229) | s(bool,'T_0') = s(bool,X229)) )),
  inference(subsumption_resolution,[],[f24169,f335])).
fof(f24169,plain,(
  ( ! [X229,X227,X228,X226] : (~p(s(bool,'T_0')) | p(s(bool,X227)) | sP1(X227,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X228),s(bool,X226))) | s(bool,X226) = s(bool,X229) | s(bool,'T_0') = s(bool,X229)) )),
  inference(superposition,[],[f864,f701])).
fof(f24164,plain,(
  ( ! [X210,X208,X209] : (sP5('T_0',X208,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X209),s(bool,'T_0'))) | s(bool,X208) = s(bool,X210) | s(bool,'T_0') = s(bool,X210)) )),
  inference(superposition,[],[f839,f701])).
fof(f24162,plain,(
  ( ! [X204,X202,X203] : (~sP5('F_0',X202,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X203),s(bool,'T_0'))) | s(bool,X202) = s(bool,X204) | s(bool,'T_0') = s(bool,X204)) )),
  inference(superposition,[],[f838,f701])).
fof(f24161,plain,(
  ( ! [X200,X199,X201] : (sP5('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X200),s(bool,'T_0')),X199) | s(bool,X199) = s(bool,X201) | s(bool,'T_0') = s(bool,X201)) )),
  inference(superposition,[],[f837,f701])).
fof(f24159,plain,(
  ( ! [X194,X195,X193] : (~sP5('F_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X194),s(bool,'T_0')),X193) | s(bool,X193) = s(bool,X195) | s(bool,'T_0') = s(bool,X195)) )),
  inference(superposition,[],[f836,f701])).
fof(f24156,plain,(
  ( ! [X185,X186,X184] : (sP4('T_0',X184,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X185))) | s(bool,X184) = s(bool,X186) | s(bool,'T_0') = s(bool,X186)) )),
  inference(superposition,[],[f825,f701])).
fof(f25895,plain,(
  ( ! [X182,X180,X183,X181] : (p(s(bool,X181)) | sP1(X181,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X180),s(bool,X182))) | s(bool,X180) = s(bool,X183) | s(bool,'T_0') = s(bool,X183)) )),
  inference(subsumption_resolution,[],[f24155,f335])).
fof(f24155,plain,(
  ( ! [X182,X180,X183,X181] : (~p(s(bool,'T_0')) | p(s(bool,X181)) | sP1(X181,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X180),s(bool,X182))) | s(bool,X180) = s(bool,X183) | s(bool,'T_0') = s(bool,X183)) )),
  inference(superposition,[],[f799,f701])).
fof(f25894,plain,(
  ( ! [X177,X178,X176,X179] : (p(s(bool,X177)) | sP1('const.bool.COND_3'(s(bool,'T_0'),s(bool,X176),s(bool,X178)),X177) | s(bool,X176) = s(bool,X179) | s(bool,'T_0') = s(bool,X179)) )),
  inference(subsumption_resolution,[],[f24154,f335])).
fof(f24154,plain,(
  ( ! [X177,X178,X176,X179] : (~p(s(bool,'T_0')) | p(s(bool,X177)) | sP1('const.bool.COND_3'(s(bool,'T_0'),s(bool,X176),s(bool,X178)),X177) | s(bool,X176) = s(bool,X179) | s(bool,'T_0') = s(bool,X179)) )),
  inference(superposition,[],[f798,f701])).
fof(f25893,plain,(
  ( ! [X175,X173,X174,X172] : (sP0('const.bool.COND_3'(s(bool,'T_0'),s(bool,X172),s(bool,X173)),X174) | s(bool,X172) = s(bool,X175) | s(bool,'T_0') = s(bool,X175)) )),
  inference(subsumption_resolution,[],[f24153,f335])).
fof(f24153,plain,(
  ( ! [X175,X173,X174,X172] : (~p(s(bool,'T_0')) | sP0('const.bool.COND_3'(s(bool,'T_0'),s(bool,X172),s(bool,X173)),X174) | s(bool,X172) = s(bool,X175) | s(bool,'T_0') = s(bool,X175)) )),
  inference(superposition,[],[f797,f701])).
fof(f25892,plain,(
  ( ! [X171,X169,X170,X168] : (sP0(X169,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X168),s(bool,X170))) | s(bool,X168) = s(bool,X171) | s(bool,'T_0') = s(bool,X171)) )),
  inference(subsumption_resolution,[],[f24152,f335])).
fof(f24152,plain,(
  ( ! [X171,X169,X170,X168] : (~p(s(bool,'T_0')) | sP0(X169,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X168),s(bool,X170))) | s(bool,X168) = s(bool,X171) | s(bool,'T_0') = s(bool,X171)) )),
  inference(superposition,[],[f796,f701])).
fof(f24150,plain,(
  ( ! [X163,X164,X162] : (sP5('T_0',X162,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X163))) | s(bool,X162) = s(bool,X164) | s(bool,'T_0') = s(bool,X164)) )),
  inference(superposition,[],[f778,f701])).
fof(f24148,plain,(
  ( ! [X158,X156,X157] : (~sP5('F_0',X156,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X157))) | s(bool,X156) = s(bool,X158) | s(bool,'T_0') = s(bool,X158)) )),
  inference(superposition,[],[f777,f701])).
fof(f24146,plain,(
  ( ! [X152,X151,X150] : (sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X151)),X150) | s(bool,X150) = s(bool,X152) | s(bool,'T_0') = s(bool,X152)) )),
  inference(superposition,[],[f776,f701])).
fof(f24144,plain,(
  ( ! [X146,X144,X145] : (~sP5('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X145)),X144) | s(bool,X144) = s(bool,X146) | s(bool,'T_0') = s(bool,X146)) )),
  inference(superposition,[],[f775,f701])).
fof(f25891,plain,(
  ( ! [X140,X138,X139,X137] : (~sP8(X137,X138,X139) | sP9(X137,X138,X139) | s(bool,X137) = s(bool,X140) | s(bool,'T_0') = s(bool,X140)) )),
  inference(subsumption_resolution,[],[f24142,f335])).
fof(f24142,plain,(
  ( ! [X140,X138,X139,X137] : (~p(s(bool,'T_0')) | ~sP8(X137,X138,X139) | sP9(X137,X138,X139) | s(bool,X137) = s(bool,X140) | s(bool,'T_0') = s(bool,X140)) )),
  inference(superposition,[],[f577,f701])).
fof(f25890,plain,(
  ( ! [X134,X136,X135,X133] : (~sP10(X133,X134,X135) | sP11(X133,X134,X135) | s(bool,X133) = s(bool,X136) | s(bool,'T_0') = s(bool,X136)) )),
  inference(subsumption_resolution,[],[f24141,f335])).
fof(f24141,plain,(
  ( ! [X134,X136,X135,X133] : (~p(s(bool,'T_0')) | ~sP10(X133,X134,X135) | sP11(X133,X134,X135) | s(bool,X133) = s(bool,X136) | s(bool,'T_0') = s(bool,X136)) )),
  inference(superposition,[],[f576,f701])).
fof(f25889,plain,(
  ( ! [X132,X130,X131,X129] : (~p(s(bool,X130)) | sP11(X129,X130,X131) | s(bool,X129) = s(bool,X132) | s(bool,'T_0') = s(bool,X132)) )),
  inference(subsumption_resolution,[],[f24140,f335])).
fof(f24140,plain,(
  ( ! [X132,X130,X131,X129] : (~p(s(bool,'T_0')) | ~p(s(bool,X130)) | sP11(X129,X130,X131) | s(bool,X129) = s(bool,X132) | s(bool,'T_0') = s(bool,X132)) )),
  inference(superposition,[],[f572,f701])).
fof(f25888,plain,(
  ( ! [X127,X125,X128,X126] : (p(s(bool,X126)) | sP11(X125,X127,X126) | s(bool,X125) = s(bool,X128) | s(bool,'T_0') = s(bool,X128)) )),
  inference(subsumption_resolution,[],[f24139,f335])).
fof(f24139,plain,(
  ( ! [X127,X125,X128,X126] : (~p(s(bool,'T_0')) | p(s(bool,X126)) | sP11(X125,X127,X126) | s(bool,X125) = s(bool,X128) | s(bool,'T_0') = s(bool,X128)) )),
  inference(superposition,[],[f571,f701])).
fof(f25887,plain,(
  ( ! [X123,X121,X124,X122] : (~p(s(bool,X122)) | sP9(X121,X122,X123) | s(bool,X121) = s(bool,X124) | s(bool,'T_0') = s(bool,X124)) )),
  inference(subsumption_resolution,[],[f24138,f335])).
fof(f24138,plain,(
  ( ! [X123,X121,X124,X122] : (~p(s(bool,'T_0')) | ~p(s(bool,X122)) | sP9(X121,X122,X123) | s(bool,X121) = s(bool,X124) | s(bool,'T_0') = s(bool,X124)) )),
  inference(superposition,[],[f555,f701])).
fof(f25886,plain,(
  ( ! [X118,X120,X119,X117] : (~p(s(bool,X118)) | sP9(X117,X119,X118) | s(bool,X117) = s(bool,X120) | s(bool,'T_0') = s(bool,X120)) )),
  inference(subsumption_resolution,[],[f24137,f335])).
fof(f24137,plain,(
  ( ! [X118,X120,X119,X117] : (~p(s(bool,'T_0')) | ~p(s(bool,X118)) | sP9(X117,X119,X118) | s(bool,X117) = s(bool,X120) | s(bool,'T_0') = s(bool,X120)) )),
  inference(superposition,[],[f554,f701])).
fof(f25885,plain,(
  ( ! [X116,X114,X115,X113] : (~p(s(bool,X114)) | p(s(bool,X115)) | ~sP7(X115,X114,X113) | s(bool,X113) = s(bool,X116) | s(bool,'T_0') = s(bool,X116)) )),
  inference(subsumption_resolution,[],[f24136,f335])).
fof(f24136,plain,(
  ( ! [X116,X114,X115,X113] : (~p(s(bool,'T_0')) | ~p(s(bool,X114)) | p(s(bool,X115)) | ~sP7(X115,X114,X113) | s(bool,X113) = s(bool,X116) | s(bool,'T_0') = s(bool,X116)) )),
  inference(superposition,[],[f541,f701])).
fof(f25884,plain,(
  ( ! [X111,X109,X112,X110] : (~p(s(bool,X110)) | p(s(bool,'const.bool.COND_3'(s(bool,X111),s(bool,X109),s(bool,X110)))) | s(bool,X109) = s(bool,X112) | s(bool,'T_0') = s(bool,X112)) )),
  inference(subsumption_resolution,[],[f24135,f335])).
fof(f24135,plain,(
  ( ! [X111,X109,X112,X110] : (~p(s(bool,'T_0')) | ~p(s(bool,X110)) | p(s(bool,'const.bool.COND_3'(s(bool,X111),s(bool,X109),s(bool,X110)))) | s(bool,X109) = s(bool,X112) | s(bool,'T_0') = s(bool,X112)) )),
  inference(superposition,[],[f489,f701])).
fof(f25883,plain,(
  ( ! [X103,X101,X102,X104] : (~p(s(bool,X102)) | p(s(bool,'const.bool.COND_3'(s(bool,X102),s(bool,X101),s(bool,X103)))) | s(bool,X101) = s(bool,X104) | s(bool,'T_0') = s(bool,X104)) )),
  inference(subsumption_resolution,[],[f24133,f335])).
fof(f24133,plain,(
  ( ! [X103,X101,X102,X104] : (~p(s(bool,'T_0')) | ~p(s(bool,X102)) | p(s(bool,'const.bool.COND_3'(s(bool,X102),s(bool,X101),s(bool,X103)))) | s(bool,X101) = s(bool,X104) | s(bool,'T_0') = s(bool,X104)) )),
  inference(superposition,[],[f487,f701])).
fof(f25880,plain,(
  ( ! [X88,X87,X85,X86] : (p(s(bool,X86)) | ~p(s(bool,'const.bool.COND_3'(s(bool,X85),s(bool,X86),s(bool,X87)))) | s(bool,X85) = s(bool,X88) | s(bool,'T_0') = s(bool,X88)) )),
  inference(subsumption_resolution,[],[f24129,f335])).
fof(f24129,plain,(
  ( ! [X88,X87,X85,X86] : (~p(s(bool,'T_0')) | p(s(bool,X86)) | ~p(s(bool,'const.bool.COND_3'(s(bool,X85),s(bool,X86),s(bool,X87)))) | s(bool,X85) = s(bool,X88) | s(bool,'T_0') = s(bool,X88)) )),
  inference(superposition,[],[f484,f701])).
fof(f25879,plain,(
  ( ! [X70,X68,X69] : (p(s(bool,X69)) | sP1(X69,X68) | s(bool,X68) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f24126,f335])).
fof(f24126,plain,(
  ( ! [X70,X68,X69] : (~p(s(bool,'T_0')) | p(s(bool,X69)) | sP1(X69,X68) | s(bool,X68) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f462,f701])).
fof(f25878,plain,(
  ( ! [X66,X67,X65] : (p(s(bool,X66)) | sP1(X65,X66) | s(bool,X65) = s(bool,X67) | s(bool,'T_0') = s(bool,X67)) )),
  inference(subsumption_resolution,[],[f24125,f335])).
fof(f24125,plain,(
  ( ! [X66,X67,X65] : (~p(s(bool,'T_0')) | p(s(bool,X66)) | sP1(X65,X66) | s(bool,X65) = s(bool,X67) | s(bool,'T_0') = s(bool,X67)) )),
  inference(superposition,[],[f459,f701])).
fof(f25877,plain,(
  ( ! [X64,X62,X63] : (sP0(X62,X63) | s(bool,X62) = s(bool,X64) | s(bool,'T_0') = s(bool,X64)) )),
  inference(subsumption_resolution,[],[f24124,f335])).
fof(f24124,plain,(
  ( ! [X64,X62,X63] : (~p(s(bool,'T_0')) | sP0(X62,X63) | s(bool,X62) = s(bool,X64) | s(bool,'T_0') = s(bool,X64)) )),
  inference(superposition,[],[f444,f701])).
fof(f25876,plain,(
  ( ! [X61,X59,X60] : (sP0(X60,X59) | s(bool,X59) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) )),
  inference(subsumption_resolution,[],[f24123,f335])).
fof(f24123,plain,(
  ( ! [X61,X59,X60] : (~p(s(bool,'T_0')) | sP0(X60,X59) | s(bool,X59) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) )),
  inference(superposition,[],[f443,f701])).
fof(f25875,plain,(
  ( ! [X57,X58,X56] : (~p(s(bool,X57)) | s(bool,X56) = s(bool,X57) | s(bool,X56) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(subsumption_resolution,[],[f24122,f335])).
fof(f24122,plain,(
  ( ! [X57,X58,X56] : (~p(s(bool,'T_0')) | ~p(s(bool,X57)) | s(bool,X56) = s(bool,X57) | s(bool,X56) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f437,f701])).
fof(f24118,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X40] : (s(bool,'T_0') != s(bool,X35) | s(X36,X37) != s(X36,X38) | s(X36,X39) != s(X36,X40) | s(X36,'const.bool.COND_3'(s(bool,X34),s(X36,X40),s(X36,X38))) = s(X36,'const.bool.COND_3'(s(bool,X35),s(X36,X39),s(X36,X37))) | s(bool,X34) = s(bool,X41) | s(bool,'T_0') = s(bool,X41)) )),
  inference(superposition,[],[f428,f701])).
fof(f25873,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X16] : (s(X11,'const.bool.COND_3'(s(bool,X10),s(X11,X15),s(X11,X16))) = s(X11,'const.bool.COND_3'(s(bool,X12),s(X11,X13),s(X11,X14))) | s(X11,X13) != s(X11,X15) | s(bool,X10) != s(bool,X12) | s(bool,X10) = s(bool,X17) | s(bool,'T_0') = s(bool,X17)) )),
  inference(subsumption_resolution,[],[f24115,f335])).
fof(f24115,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X16] : (~p(s(bool,'T_0')) | s(X11,'const.bool.COND_3'(s(bool,X10),s(X11,X15),s(X11,X16))) = s(X11,'const.bool.COND_3'(s(bool,X12),s(X11,X13),s(X11,X14))) | s(X11,X13) != s(X11,X15) | s(bool,X10) != s(bool,X12) | s(bool,X10) = s(bool,X17) | s(bool,'T_0') = s(bool,X17)) )),
  inference(superposition,[],[f426,f701])).
fof(f25872,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f24113,f335])).
fof(f24113,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f358,f701])).
fof(f24001,plain,(
  ( ! [X576,X578,X577] : (sP4('T_0','T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X577),s(bool,X576))) | s(bool,X576) = s(bool,X578) | s(bool,'T_0') = s(bool,X578)) )),
  inference(superposition,[],[f930,f701])).
fof(f23993,plain,(
  ( ! [X549,X548,X550] : (sP5('T_0','T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X549),s(bool,X548))) | s(bool,X548) = s(bool,X550) | s(bool,'T_0') = s(bool,X550)) )),
  inference(superposition,[],[f839,f701])).
fof(f23991,plain,(
  ( ! [X542,X543,X544] : (~sP5('F_0','T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X543),s(bool,X542))) | s(bool,X542) = s(bool,X544) | s(bool,'T_0') = s(bool,X544)) )),
  inference(superposition,[],[f838,f701])).
fof(f23990,plain,(
  ( ! [X540,X541,X539] : (sP5('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X540),s(bool,X539)),'T_0') | s(bool,X539) = s(bool,X541) | s(bool,'T_0') = s(bool,X541)) )),
  inference(superposition,[],[f837,f701])).
fof(f23988,plain,(
  ( ! [X534,X533,X535] : (~sP5('F_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X534),s(bool,X533)),'T_0') | s(bool,X533) = s(bool,X535) | s(bool,'T_0') = s(bool,X535)) )),
  inference(superposition,[],[f836,f701])).
fof(f23985,plain,(
  ( ! [X524,X526,X525] : (sP4('T_0','T_0','const.bool.COND_3'(s(bool,X524),s(bool,X524),s(bool,X525))) | s(bool,X524) = s(bool,X526) | s(bool,'T_0') = s(bool,X526)) )),
  inference(superposition,[],[f825,f701])).
fof(f23979,plain,(
  ( ! [X502,X504,X503] : (sP5('T_0','T_0','const.bool.COND_3'(s(bool,X502),s(bool,X502),s(bool,X503))) | s(bool,X502) = s(bool,X504) | s(bool,'T_0') = s(bool,X504)) )),
  inference(superposition,[],[f778,f701])).
fof(f23977,plain,(
  ( ! [X496,X498,X497] : (~sP5('F_0','T_0','const.bool.COND_3'(s(bool,X496),s(bool,X496),s(bool,X497))) | s(bool,X496) = s(bool,X498) | s(bool,'T_0') = s(bool,X498)) )),
  inference(superposition,[],[f777,f701])).
fof(f23975,plain,(
  ( ! [X492,X490,X491] : (sP5('T_0','const.bool.COND_3'(s(bool,X490),s(bool,X490),s(bool,X491)),'T_0') | s(bool,X490) = s(bool,X492) | s(bool,'T_0') = s(bool,X492)) )),
  inference(superposition,[],[f776,f701])).
fof(f23973,plain,(
  ( ! [X484,X486,X485] : (~sP5('F_0','const.bool.COND_3'(s(bool,X484),s(bool,X484),s(bool,X485)),'T_0') | s(bool,X484) = s(bool,X486) | s(bool,'T_0') = s(bool,X486)) )),
  inference(superposition,[],[f775,f701])).
fof(f23947,plain,(
  ( ! [X377,X379,X374,X380,X376,X378,X375,X381] : (s(bool,X374) != s(bool,X375) | s(X376,X377) != s(X376,X378) | s(X376,X379) != s(X376,X380) | s(X376,'const.bool.COND_3'(s(bool,'T_0'),s(X376,X380),s(X376,X378))) = s(X376,'const.bool.COND_3'(s(bool,X375),s(X376,X379),s(X376,X377))) | s(bool,X374) = s(bool,X381) | s(bool,'T_0') = s(bool,X381)) )),
  inference(superposition,[],[f428,f701])).
fof(f23939,plain,(
  ( ! [X341,X339,X340,X338] : (sP1('const.bool.COND_3'(s(bool,X339),s(bool,X340),s(bool,X338)),'F_0') | s(bool,'T_0') = s(bool,X339) | s(bool,X338) = s(bool,X341) | s(bool,'T_0') = s(bool,X341)) )),
  inference(superposition,[],[f16717,f701])).
fof(f23938,plain,(
  ( ! [X335,X337,X334,X336] : (sP1('const.bool.COND_3'(s(bool,'F_0'),s(bool,X335),s(bool,X334)),X336) | p(s(bool,X336)) | s(bool,X334) = s(bool,X337) | s(bool,'T_0') = s(bool,X337)) )),
  inference(superposition,[],[f11665,f701])).
fof(f23937,plain,(
  ( ! [X333,X331,X332] : (sP1('const.bool.COND_3'(s(bool,'F_0'),s(bool,X332),s(bool,X331)),'F_0') | s(bool,X331) = s(bool,X333) | s(bool,'T_0') = s(bool,X333)) )),
  inference(superposition,[],[f11657,f701])).
fof(f23936,plain,(
  ( ! [X327,X329,X328,X330] : (sP1(X328,'const.bool.COND_3'(s(bool,X327),s(bool,X327),s(bool,X329))) | p(s(bool,X328)) | s(bool,X327) = s(bool,X330) | s(bool,'T_0') = s(bool,X330)) )),
  inference(superposition,[],[f5276,f701])).
fof(f23935,plain,(
  ( ! [X325,X324,X326] : (sP1('F_0','const.bool.COND_3'(s(bool,X324),s(bool,X324),s(bool,X325))) | s(bool,X324) = s(bool,X326) | s(bool,'T_0') = s(bool,X326)) )),
  inference(superposition,[],[f5268,f701])).
fof(f23934,plain,(
  ( ! [X321,X323,X320,X322] : (sP1('const.bool.COND_3'(s(bool,X320),s(bool,X320),s(bool,X321)),X322) | p(s(bool,X322)) | s(bool,X320) = s(bool,X323) | s(bool,'T_0') = s(bool,X323)) )),
  inference(superposition,[],[f4870,f701])).
fof(f23933,plain,(
  ( ! [X315,X316,X318,X317,X319] : (sP11('const.bool.COND_3'(s(bool,X316),s(bool,X315),s(bool,X315)),X317,X318) | ~p(s(bool,X317)) | s(bool,X315) = s(bool,X319) | s(bool,'T_0') = s(bool,X319)) )),
  inference(superposition,[],[f4862,f701])).
fof(f23932,plain,(
  ( ! [X313,X310,X312,X314,X311] : (sP11('const.bool.COND_3'(s(bool,X311),s(bool,X310),s(bool,X310)),X312,X313) | p(s(bool,X313)) | s(bool,X310) = s(bool,X314) | s(bool,'T_0') = s(bool,X314)) )),
  inference(superposition,[],[f4861,f701])).
fof(f23931,plain,(
  ( ! [X308,X306,X309,X305,X307] : (sP9('const.bool.COND_3'(s(bool,X306),s(bool,X305),s(bool,X305)),X307,X308) | ~p(s(bool,X307)) | s(bool,X305) = s(bool,X309) | s(bool,'T_0') = s(bool,X309)) )),
  inference(superposition,[],[f4860,f701])).
fof(f23930,plain,(
  ( ! [X300,X302,X304,X301,X303] : (sP9('const.bool.COND_3'(s(bool,X301),s(bool,X300),s(bool,X300)),X302,X303) | ~p(s(bool,X303)) | s(bool,X300) = s(bool,X304) | s(bool,'T_0') = s(bool,X304)) )),
  inference(superposition,[],[f4859,f701])).
fof(f23929,plain,(
  ( ! [X296,X298,X297,X299] : (p(s(bool,'const.bool.COND_3'(s(bool,X297),s(bool,X298),s(bool,X296)))) | p(s(bool,X297)) | s(bool,X296) = s(bool,X299) | s(bool,'T_0') = s(bool,X299)) )),
  inference(superposition,[],[f4855,f701])).
fof(f23928,plain,(
  ( ! [X292,X294,X293,X295] : (sP1(X293,'const.bool.COND_3'(s(bool,X294),s(bool,X292),s(bool,X292))) | p(s(bool,X293)) | s(bool,X292) = s(bool,X295) | s(bool,'T_0') = s(bool,X295)) )),
  inference(superposition,[],[f4844,f701])).
fof(f23927,plain,(
  ( ! [X288,X290,X289,X291] : (sP1('const.bool.COND_3'(s(bool,X289),s(bool,X288),s(bool,X288)),X290) | p(s(bool,X290)) | s(bool,X288) = s(bool,X291) | s(bool,'T_0') = s(bool,X291)) )),
  inference(superposition,[],[f4843,f701])).
fof(f23926,plain,(
  ( ! [X285,X287,X286] : (sP5('F_0','F_0','const.bool.COND_3'(s(bool,X286),s(bool,X285),s(bool,X285))) | s(bool,X285) = s(bool,X287) | s(bool,'T_0') = s(bool,X287)) )),
  inference(superposition,[],[f2262,f701])).
fof(f23925,plain,(
  ( ! [X282,X283,X284] : (sP5('F_0','const.bool.COND_3'(s(bool,X283),s(bool,X282),s(bool,X282)),'F_0') | s(bool,X282) = s(bool,X284) | s(bool,'T_0') = s(bool,X284)) )),
  inference(superposition,[],[f2222,f701])).
fof(f23924,plain,(
  ( ! [X280,X279,X281] : (sP5('const.bool.COND_3'(s(bool,X280),s(bool,X279),s(bool,X279)),'F_0','F_0') | s(bool,X279) = s(bool,X281) | s(bool,'T_0') = s(bool,X281)) )),
  inference(superposition,[],[f2128,f701])).
fof(f23923,plain,(
  ( ! [X277,X276,X278] : (sP5('const.bool.COND_3'(s(bool,X277),s(bool,X276),s(bool,X276)),'T_0','T_0') | s(bool,X276) = s(bool,X278) | s(bool,'T_0') = s(bool,X278)) )),
  inference(superposition,[],[f2118,f701])).
fof(f23922,plain,(
  ( ! [X273,X275,X274] : (~sP5('F_0','T_0','const.bool.COND_3'(s(bool,X274),s(bool,X273),s(bool,X273))) | s(bool,X273) = s(bool,X275) | s(bool,'T_0') = s(bool,X275)) )),
  inference(superposition,[],[f2082,f701])).
fof(f23921,plain,(
  ( ! [X271,X270,X272] : (~sP5('F_0','const.bool.COND_3'(s(bool,X271),s(bool,X270),s(bool,X270)),'T_0') | s(bool,X270) = s(bool,X272) | s(bool,'T_0') = s(bool,X272)) )),
  inference(superposition,[],[f2073,f701])).
fof(f23920,plain,(
  ( ! [X269,X267,X268] : (sP5('T_0','const.bool.COND_3'(s(bool,X268),s(bool,X267),s(bool,X267)),'T_0') | s(bool,X267) = s(bool,X269) | s(bool,'T_0') = s(bool,X269)) )),
  inference(superposition,[],[f2030,f701])).
fof(f23919,plain,(
  ( ! [X265,X264,X266] : (~sP5('T_0','F_0','const.bool.COND_3'(s(bool,X265),s(bool,X264),s(bool,X264))) | s(bool,X264) = s(bool,X266) | s(bool,'T_0') = s(bool,X266)) )),
  inference(superposition,[],[f2015,f701])).
fof(f23918,plain,(
  ( ! [X261,X263,X260,X262] : (sP9('const.bool.COND_3'(s(bool,X261),s(bool,X260),s(bool,X260)),'T_0',X262) | s(bool,X260) = s(bool,X263) | s(bool,'T_0') = s(bool,X263)) )),
  inference(superposition,[],[f1762,f701])).
fof(f23917,plain,(
  ( ! [X257,X259,X256,X258] : (sP9('const.bool.COND_3'(s(bool,X257),s(bool,X256),s(bool,X256)),X258,'T_0') | s(bool,X256) = s(bool,X259) | s(bool,'T_0') = s(bool,X259)) )),
  inference(superposition,[],[f1761,f701])).
fof(f23916,plain,(
  ( ! [X255,X253,X254,X252] : (sP11('const.bool.COND_3'(s(bool,X253),s(bool,X252),s(bool,X252)),'T_0',X254) | s(bool,X252) = s(bool,X255) | s(bool,'T_0') = s(bool,X255)) )),
  inference(superposition,[],[f1760,f701])).
fof(f23915,plain,(
  ( ! [X251,X249,X250,X248] : (sP11('const.bool.COND_3'(s(bool,X249),s(bool,X248),s(bool,X248)),X250,'F_0') | s(bool,X248) = s(bool,X251) | s(bool,'T_0') = s(bool,X251)) )),
  inference(superposition,[],[f1759,f701])).
fof(f23914,plain,(
  ( ! [X246,X247,X245] : (~sP11('T_0','F_0','const.bool.COND_3'(s(bool,X246),s(bool,X245),s(bool,X245))) | s(bool,X245) = s(bool,X247) | s(bool,'T_0') = s(bool,X247)) )),
  inference(superposition,[],[f1756,f701])).
fof(f23913,plain,(
  ( ! [X243,X244,X242] : (~sP11('const.bool.COND_3'(s(bool,X243),s(bool,X242),s(bool,X242)),'F_0','T_0') | s(bool,X242) = s(bool,X244) | s(bool,'T_0') = s(bool,X244)) )),
  inference(superposition,[],[f1754,f701])).
fof(f23912,plain,(
  ( ! [X241,X239,X240,X238] : (sP11('T_0','const.bool.COND_3'(s(bool,X239),s(bool,X238),s(bool,X238)),X240) | s(bool,X238) = s(bool,X241) | s(bool,'T_0') = s(bool,X241)) )),
  inference(superposition,[],[f1749,f701])).
fof(f23911,plain,(
  ( ! [X237,X235,X236] : (sP11('F_0','F_0','const.bool.COND_3'(s(bool,X236),s(bool,X235),s(bool,X235))) | s(bool,X235) = s(bool,X237) | s(bool,'T_0') = s(bool,X237)) )),
  inference(superposition,[],[f1746,f701])).
fof(f23910,plain,(
  ( ! [X231,X233,X230,X234,X232] : (sP10(X231,'const.bool.COND_3'(s(bool,X232),s(bool,X230),s(bool,X230)),X233) | s(bool,X230) = s(bool,X234) | s(bool,'T_0') = s(bool,X234)) )),
  inference(superposition,[],[f1745,f701])).
fof(f23909,plain,(
  ( ! [X229,X227,X228] : (~sP10('T_0','F_0','const.bool.COND_3'(s(bool,X228),s(bool,X227),s(bool,X227))) | s(bool,X227) = s(bool,X229) | s(bool,'T_0') = s(bool,X229)) )),
  inference(superposition,[],[f1741,f701])).
fof(f23908,plain,(
  ( ! [X225,X226,X224] : (~sP10('const.bool.COND_3'(s(bool,X225),s(bool,X224),s(bool,X224)),'F_0','T_0') | s(bool,X224) = s(bool,X226) | s(bool,'T_0') = s(bool,X226)) )),
  inference(superposition,[],[f1737,f701])).
fof(f23907,plain,(
  ( ! [X222,X220,X223,X221] : (~sP11('F_0','const.bool.COND_3'(s(bool,X221),s(bool,X220),s(bool,X220)),X222) | s(bool,X220) = s(bool,X223) | s(bool,'T_0') = s(bool,X223)) )),
  inference(superposition,[],[f1725,f701])).
fof(f23906,plain,(
  ( ! [X218,X219,X217] : (~sP9('const.bool.COND_3'(s(bool,X218),s(bool,X217),s(bool,X217)),'F_0','F_0') | s(bool,X217) = s(bool,X219) | s(bool,'T_0') = s(bool,X219)) )),
  inference(superposition,[],[f1724,f701])).
fof(f23905,plain,(
  ( ! [X216,X215,X213,X214] : (sP9('T_0','const.bool.COND_3'(s(bool,X214),s(bool,X213),s(bool,X213)),X215) | s(bool,X213) = s(bool,X216) | s(bool,'T_0') = s(bool,X216)) )),
  inference(superposition,[],[f1719,f701])).
fof(f23904,plain,(
  ( ! [X212,X210,X211,X209] : (sP9('T_0',X210,'const.bool.COND_3'(s(bool,X211),s(bool,X209),s(bool,X209))) | s(bool,X209) = s(bool,X212) | s(bool,'T_0') = s(bool,X212)) )),
  inference(superposition,[],[f1716,f701])).
fof(f23903,plain,(
  ( ! [X208,X206,X204,X207,X205] : (sP8(X205,'const.bool.COND_3'(s(bool,X206),s(bool,X204),s(bool,X204)),X207) | s(bool,X204) = s(bool,X208) | s(bool,'T_0') = s(bool,X208)) )),
  inference(superposition,[],[f1713,f701])).
fof(f23902,plain,(
  ( ! [X202,X200,X199,X203,X201] : (sP8(X200,X201,'const.bool.COND_3'(s(bool,X202),s(bool,X199),s(bool,X199))) | s(bool,X199) = s(bool,X203) | s(bool,'T_0') = s(bool,X203)) )),
  inference(superposition,[],[f1712,f701])).
fof(f23901,plain,(
  ( ! [X198,X196,X197] : (~sP8('const.bool.COND_3'(s(bool,X197),s(bool,X196),s(bool,X196)),'F_0','F_0') | s(bool,X196) = s(bool,X198) | s(bool,'T_0') = s(bool,X198)) )),
  inference(superposition,[],[f1709,f701])).
fof(f23900,plain,(
  ( ! [X194,X192,X195,X193] : (~sP9('F_0','const.bool.COND_3'(s(bool,X193),s(bool,X192),s(bool,X192)),X194) | s(bool,X192) = s(bool,X195) | s(bool,'T_0') = s(bool,X195)) )),
  inference(superposition,[],[f1705,f701])).
fof(f23899,plain,(
  ( ! [X191,X189,X190,X188] : (~sP9('F_0',X189,'const.bool.COND_3'(s(bool,X190),s(bool,X188),s(bool,X188))) | s(bool,X188) = s(bool,X191) | s(bool,'T_0') = s(bool,X191)) )),
  inference(superposition,[],[f1703,f701])).
fof(f23898,plain,(
  ( ! [X187,X185,X186] : (~sP7('F_0','T_0','const.bool.COND_3'(s(bool,X186),s(bool,X185),s(bool,X185))) | s(bool,X185) = s(bool,X187) | s(bool,'T_0') = s(bool,X187)) )),
  inference(superposition,[],[f1700,f701])).
fof(f23897,plain,(
  ( ! [X182,X184,X183] : (~sP7('F_0','const.bool.COND_3'(s(bool,X183),s(bool,X182),s(bool,X182)),'T_0') | s(bool,X182) = s(bool,X184) | s(bool,'T_0') = s(bool,X184)) )),
  inference(superposition,[],[f1697,f701])).
fof(f23896,plain,(
  ( ! [X180,X181,X179] : (sP7('T_0','T_0','const.bool.COND_3'(s(bool,X180),s(bool,X179),s(bool,X179))) | s(bool,X179) = s(bool,X181) | s(bool,'T_0') = s(bool,X181)) )),
  inference(superposition,[],[f1692,f701])).
fof(f23895,plain,(
  ( ! [X177,X178,X176] : (sP7('const.bool.COND_3'(s(bool,X177),s(bool,X176),s(bool,X176)),'T_0','T_0') | s(bool,X176) = s(bool,X178) | s(bool,'T_0') = s(bool,X178)) )),
  inference(superposition,[],[f1688,f701])).
fof(f23894,plain,(
  ( ! [X175,X173,X174] : (sP7('T_0','const.bool.COND_3'(s(bool,X174),s(bool,X173),s(bool,X173)),'T_0') | s(bool,X173) = s(bool,X175) | s(bool,'T_0') = s(bool,X175)) )),
  inference(superposition,[],[f1683,f701])).
fof(f23893,plain,(
  ( ! [X171,X169,X172,X170,X168] : (sP6(X169,X170,'const.bool.COND_3'(s(bool,X171),s(bool,X168),s(bool,X168))) | s(bool,X168) = s(bool,X172) | s(bool,'T_0') = s(bool,X172)) )),
  inference(superposition,[],[f1682,f701])).
fof(f23892,plain,(
  ( ! [X167,X165,X166] : (~sP6('T_0','const.bool.COND_3'(s(bool,X166),s(bool,X165),s(bool,X165)),'F_0') | s(bool,X165) = s(bool,X167) | s(bool,'T_0') = s(bool,X167)) )),
  inference(superposition,[],[f1678,f701])).
fof(f23891,plain,(
  ( ! [X163,X164,X162] : (~sP6('const.bool.COND_3'(s(bool,X163),s(bool,X162),s(bool,X162)),'T_0','F_0') | s(bool,X162) = s(bool,X164) | s(bool,'T_0') = s(bool,X164)) )),
  inference(superposition,[],[f1674,f701])).
fof(f23890,plain,(
  ( ! [X158,X161,X159,X160] : (~sP7('const.bool.COND_3'(s(bool,X159),s(bool,X158),s(bool,X158)),'F_0',X160) | s(bool,X158) = s(bool,X161) | s(bool,'T_0') = s(bool,X161)) )),
  inference(superposition,[],[f1667,f701])).
fof(f23889,plain,(
  ( ! [X156,X154,X157,X155] : (~sP7('const.bool.COND_3'(s(bool,X155),s(bool,X154),s(bool,X154)),X156,'F_0') | s(bool,X154) = s(bool,X157) | s(bool,'T_0') = s(bool,X157)) )),
  inference(superposition,[],[f1666,f701])).
fof(f23888,plain,(
  ( ! [X152,X151,X153] : (~sP5('const.bool.COND_3'(s(bool,X152),s(bool,X151),s(bool,X151)),'F_0','T_0') | s(bool,X151) = s(bool,X153) | s(bool,'T_0') = s(bool,X153)) )),
  inference(superposition,[],[f1665,f701])).
fof(f23887,plain,(
  ( ! [X149,X150,X148] : (~sP5('const.bool.COND_3'(s(bool,X149),s(bool,X148),s(bool,X148)),'T_0','F_0') | s(bool,X148) = s(bool,X150) | s(bool,'T_0') = s(bool,X150)) )),
  inference(superposition,[],[f1664,f701])).
fof(f23886,plain,(
  ( ! [X146,X144,X147,X145,X143] : (sP2('const.bool.COND_3'(s(bool,X144),s(bool,X143),s(bool,X143)),X145,X146) | s(bool,X143) = s(bool,X147) | s(bool,'T_0') = s(bool,X147)) )),
  inference(superposition,[],[f1659,f701])).
fof(f23885,plain,(
  ( ! [X142,X140,X138,X141,X139] : (sP2(X139,'const.bool.COND_3'(s(bool,X140),s(bool,X138),s(bool,X138)),X141) | s(bool,X138) = s(bool,X142) | s(bool,'T_0') = s(bool,X142)) )),
  inference(superposition,[],[f1658,f701])).
fof(f23884,plain,(
  ( ! [X134,X136,X135,X133,X137] : (sP2(X134,X135,'const.bool.COND_3'(s(bool,X136),s(bool,X133),s(bool,X133))) | s(bool,X133) = s(bool,X137) | s(bool,'T_0') = s(bool,X137)) )),
  inference(superposition,[],[f1657,f701])).
fof(f23883,plain,(
  ( ! [X132,X130,X128,X131,X129] : (sP3(X129,X130,'const.bool.COND_3'(s(bool,X131),s(bool,X128),s(bool,X128))) | s(bool,X128) = s(bool,X132) | s(bool,'T_0') = s(bool,X132)) )),
  inference(superposition,[],[f1656,f701])).
fof(f23882,plain,(
  ( ! [X127,X125,X126] : (~sP3('T_0','const.bool.COND_3'(s(bool,X126),s(bool,X125),s(bool,X125)),'F_0') | s(bool,X125) = s(bool,X127) | s(bool,'T_0') = s(bool,X127)) )),
  inference(superposition,[],[f1652,f701])).
fof(f23881,plain,(
  ( ! [X123,X124,X122] : (~sP3('const.bool.COND_3'(s(bool,X123),s(bool,X122),s(bool,X122)),'T_0','F_0') | s(bool,X122) = s(bool,X124) | s(bool,'T_0') = s(bool,X124)) )),
  inference(superposition,[],[f1648,f701])).
fof(f23880,plain,(
  ( ! [X121,X118,X120,X119,X117] : (sP4(X118,X119,'const.bool.COND_3'(s(bool,X120),s(bool,X117),s(bool,X117))) | s(bool,X117) = s(bool,X121) | s(bool,'T_0') = s(bool,X121)) )),
  inference(superposition,[],[f1647,f701])).
fof(f23879,plain,(
  ( ! [X116,X114,X115] : (~sP4('T_0','const.bool.COND_3'(s(bool,X115),s(bool,X114),s(bool,X114)),'F_0') | s(bool,X114) = s(bool,X116) | s(bool,'T_0') = s(bool,X116)) )),
  inference(superposition,[],[f1639,f701])).
fof(f23878,plain,(
  ( ! [X111,X112,X113] : (~sP4('const.bool.COND_3'(s(bool,X112),s(bool,X111),s(bool,X111)),'T_0','F_0') | s(bool,X111) = s(bool,X113) | s(bool,'T_0') = s(bool,X113)) )),
  inference(superposition,[],[f1635,f701])).
fof(f23877,plain,(
  ( ! [X109,X110,X108] : (sP5('T_0','T_0','const.bool.COND_3'(s(bool,X109),s(bool,X108),s(bool,X108))) | s(bool,X108) = s(bool,X110) | s(bool,'T_0') = s(bool,X110)) )),
  inference(superposition,[],[f1634,f701])).
fof(f23876,plain,(
  ( ! [X107,X105,X106] : (~sP5('T_0','const.bool.COND_3'(s(bool,X106),s(bool,X105),s(bool,X105)),'F_0') | s(bool,X105) = s(bool,X107) | s(bool,'T_0') = s(bool,X107)) )),
  inference(superposition,[],[f1619,f701])).
fof(f23875,plain,(
  ( ! [X103,X102,X104] : (sP1('F_0','const.bool.COND_3'(s(bool,X103),s(bool,X102),s(bool,X102))) | s(bool,X102) = s(bool,X104) | s(bool,'T_0') = s(bool,X104)) )),
  inference(superposition,[],[f1605,f701])).
fof(f23874,plain,(
  ( ! [X101,X99,X100] : (sP1('const.bool.COND_3'(s(bool,X100),s(bool,X99),s(bool,X99)),'F_0') | s(bool,X99) = s(bool,X101) | s(bool,'T_0') = s(bool,X101)) )),
  inference(superposition,[],[f1604,f701])).
fof(f23873,plain,(
  ( ! [X97,X98,X96] : (~sP1('T_0','const.bool.COND_3'(s(bool,X97),s(bool,X96),s(bool,X96))) | s(bool,X96) = s(bool,X98) | s(bool,'T_0') = s(bool,X98)) )),
  inference(superposition,[],[f1602,f701])).
fof(f23872,plain,(
  ( ! [X94,X95,X93] : (~sP1('const.bool.COND_3'(s(bool,X94),s(bool,X93),s(bool,X93)),'T_0') | s(bool,X93) = s(bool,X95) | s(bool,'T_0') = s(bool,X95)) )),
  inference(superposition,[],[f1599,f701])).
fof(f23871,plain,(
  ( ! [X92,X90,X91,X89] : (sP0('const.bool.COND_3'(s(bool,X90),s(bool,X89),s(bool,X89)),X91) | s(bool,X89) = s(bool,X92) | s(bool,'T_0') = s(bool,X92)) )),
  inference(superposition,[],[f1598,f701])).
fof(f23870,plain,(
  ( ! [X88,X87,X85,X86] : (sP0(X86,'const.bool.COND_3'(s(bool,X87),s(bool,X85),s(bool,X85))) | s(bool,X85) = s(bool,X88) | s(bool,'T_0') = s(bool,X88)) )),
  inference(superposition,[],[f1597,f701])).
fof(f23869,plain,(
  ( ! [X83,X84,X82] : (s(bool,X82) = s(bool,'const.bool.COND_3'(s(bool,X83),s(bool,X82),s(bool,X82))) | s(bool,X82) = s(bool,X84) | s(bool,'T_0') = s(bool,X84)) )),
  inference(superposition,[],[f1595,f701])).
fof(f23868,plain,(
  ( ! [X80,X81,X79] : (sP1('const.bool.COND_3'(s(bool,X79),s(bool,X79),s(bool,X80)),'F_0') | s(bool,X79) = s(bool,X81) | s(bool,'T_0') = s(bool,X81)) )),
  inference(superposition,[],[f1569,f701])).
fof(f23867,plain,(
  ( ! [X78,X76,X77] : (~sP2('F_0','const.bool.COND_3'(s(bool,X76),s(bool,'F_0'),s(bool,X77)),'F_0') | s(bool,X76) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f1410,f701])).
fof(f23866,plain,(
  ( ! [X74,X75,X73] : (~sP3('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X74),s(bool,X73)),'F_0') | s(bool,X73) = s(bool,X75) | s(bool,'T_0') = s(bool,X75)) )),
  inference(superposition,[],[f1403,f701])).
fof(f23865,plain,(
  ( ! [X70,X72,X71] : (~sP3('T_0','const.bool.COND_3'(s(bool,X70),s(bool,X70),s(bool,X71)),'F_0') | s(bool,X70) = s(bool,X72) | s(bool,'T_0') = s(bool,X72)) )),
  inference(superposition,[],[f1402,f701])).
fof(f23864,plain,(
  ( ! [X68,X69,X67] : (~sP2('const.bool.COND_3'(s(bool,X67),s(bool,'F_0'),s(bool,X68)),'F_0','F_0') | s(bool,X67) = s(bool,X69) | s(bool,'T_0') = s(bool,X69)) )),
  inference(superposition,[],[f1365,f701])).
fof(f23863,plain,(
  ( ! [X66,X64,X65] : (~sP3('const.bool.COND_3'(s(bool,'F_0'),s(bool,X65),s(bool,X64)),'T_0','F_0') | s(bool,X64) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f1333,f701])).
fof(f23862,plain,(
  ( ! [X61,X62,X63] : (~sP3('const.bool.COND_3'(s(bool,X61),s(bool,X61),s(bool,X62)),'T_0','F_0') | s(bool,X61) = s(bool,X63) | s(bool,'T_0') = s(bool,X63)) )),
  inference(superposition,[],[f1332,f701])).
fof(f23861,plain,(
  ( ! [X59,X57,X60,X58,X56] : (sP0('const.bool.COND_3'(s(bool,X57),s(bool,X58),s(bool,X56)),X59) | s(bool,'T_0') = s(bool,X57) | s(bool,X56) = s(bool,X60) | s(bool,'T_0') = s(bool,X60)) )),
  inference(superposition,[],[f1275,f701])).
fof(f23860,plain,(
  ( ! [X54,X52,X55,X53] : (sP0('const.bool.COND_3'(s(bool,'F_0'),s(bool,X53),s(bool,X52)),X54) | s(bool,X52) = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) )),
  inference(superposition,[],[f1267,f701])).
fof(f23859,plain,(
  ( ! [X47,X50,X48,X51,X49] : (sP0(X48,'const.bool.COND_3'(s(bool,X49),s(bool,X50),s(bool,X47))) | s(bool,'T_0') = s(bool,X49) | s(bool,X47) = s(bool,X51) | s(bool,'T_0') = s(bool,X51)) )),
  inference(superposition,[],[f1239,f701])).
fof(f23858,plain,(
  ( ! [X45,X43,X46,X44] : (sP0(X44,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X45),s(bool,X43))) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f1231,f701])).
fof(f23857,plain,(
  ( ! [X39,X41,X42,X40] : (sP0('const.bool.COND_3'(s(bool,X39),s(bool,X39),s(bool,X40)),X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f1196,f701])).
fof(f23856,plain,(
  ( ! [X37,X35,X38,X36] : (sP0(X36,'const.bool.COND_3'(s(bool,X35),s(bool,X35),s(bool,X37))) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f1157,f701])).
fof(f23855,plain,(
  ( ! [X33,X31,X34,X32] : (sP4('T_0','const.bool.COND_3'(s(bool,X31),s(bool,X32),s(bool,X33)),X32) | s(bool,X31) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f898,f701])).
fof(f23854,plain,(
  ( ! [X30,X28,X29,X27] : (sP4('T_0',X28,'const.bool.COND_3'(s(bool,X27),s(bool,X28),s(bool,X29))) | s(bool,X27) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f825,f701])).
fof(f23853,plain,(
  ( ! [X26,X24,X23,X25] : (sP5('T_0',X24,'const.bool.COND_3'(s(bool,X23),s(bool,X24),s(bool,X25))) | s(bool,X23) = s(bool,X26) | s(bool,'T_0') = s(bool,X26)) )),
  inference(superposition,[],[f778,f701])).
fof(f23852,plain,(
  ( ! [X21,X19,X22,X20] : (~sP5('F_0',X20,'const.bool.COND_3'(s(bool,X19),s(bool,X20),s(bool,X21))) | s(bool,X19) = s(bool,X22) | s(bool,'T_0') = s(bool,X22)) )),
  inference(superposition,[],[f777,f701])).
fof(f23851,plain,(
  ( ! [X17,X15,X18,X16] : (sP5('T_0','const.bool.COND_3'(s(bool,X15),s(bool,X16),s(bool,X17)),X16) | s(bool,X15) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f776,f701])).
fof(f23850,plain,(
  ( ! [X14,X12,X13,X11] : (~sP5('F_0','const.bool.COND_3'(s(bool,X11),s(bool,X12),s(bool,X13)),X12) | s(bool,X11) = s(bool,X14) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f775,f701])).
fof(f23849,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(X6,X7) = s(X6,'const.bool.COND_3'(s(bool,X5),s(X6,X7),s(X6,X8))) | s(bool,X5) = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f413,f701])).
fof(f23848,plain,(
  ( ! [X4,X2,X3] : (s(bool,X2) != s(bool,X3) | p(s(bool,X3)) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f354,f701])).
fof(f23523,plain,(
  ( ! [X1009,X1006,X1008,X1007] : (s(X1008,f(s(bool,X1009))) != s(X1008,g(s(bool,X1009))) | s(fun(bool,X1008),X1006) = s(fun(bool,X1008),X1007) | s(bool,'T_0') = s(bool,X1009) | s(bool,'T_0') = s(bool,sK22(X1006,X1007,X1008,bool))) )),
  inference(superposition,[],[f478,f701])).
fof(f23522,plain,(
  ( ! [X1002,X1004,X1003,X1005] : (~p(s(bool,happ(s(fun(bool,bool),X1002),s(bool,X1004)))) | p(s(bool,happ(s(fun(bool,bool),X1003),s(bool,X1005)))) | ~p(s(bool,happ(s(fun(bool,bool),X1003),s(bool,sK17(X1002,X1003,bool))))) | s(bool,'T_0') = s(bool,X1004) | s(bool,'T_0') = s(bool,sK18(X1002,X1003,bool))) )),
  inference(superposition,[],[f419,f701])).
fof(f23521,plain,(
  ( ! [X1000,X998,X1001,X999] : (~p(s(bool,happ(s(fun(bool,bool),X998),s(bool,X1000)))) | p(s(bool,happ(s(fun(bool,bool),X998),s(bool,X1001)))) | ~p(s(bool,happ(s(fun(bool,bool),X999),s(bool,sK17(X998,X999,bool))))) | s(bool,'T_0') = s(bool,X1000) | s(bool,'T_0') = s(bool,sK18(X998,X999,bool))) )),
  inference(superposition,[],[f420,f701])).
fof(f23520,plain,(
  ( ! [X994,X996,X995,X997] : (~p(s(bool,happ(s(fun(bool,bool),X994),s(bool,X996)))) | p(s(bool,happ(s(fun(bool,bool),X995),s(bool,X997)))) | ~p(s(bool,happ(s(fun(bool,bool),X995),s(bool,sK16(X994,X995,bool))))) | s(bool,'T_0') = s(bool,X996) | s(bool,'T_0') = s(bool,sK16(X994,X995,bool))) )),
  inference(superposition,[],[f417,f701])).
fof(f23519,plain,(
  ( ! [X990,X992,X991,X993] : (~p(s(bool,happ(s(fun(bool,bool),X990),s(bool,X992)))) | p(s(bool,happ(s(fun(bool,bool),X990),s(bool,X993)))) | ~p(s(bool,happ(s(fun(bool,bool),X991),s(bool,sK16(X990,X991,bool))))) | s(bool,'T_0') = s(bool,X992) | s(bool,'T_0') = s(bool,sK16(X990,X991,bool))) )),
  inference(superposition,[],[f418,f701])).
fof(f23518,plain,(
  ( ! [X988,X987,X989] : (~p(s(bool,happ(s(fun(bool,bool),X987),s(bool,X988)))) | p(s(bool,happ(s(fun(bool,bool),X987),s(bool,X989)))) | s(bool,'T_0') = s(bool,X988) | s(bool,'T_0') = s(bool,sK15(X987,bool))) )),
  inference(superposition,[],[f412,f701])).
fof(f23517,plain,(
  ( ! [X984,X985,X986] : (~p(s(bool,happ(s(fun(bool,bool),X984),s(bool,X985)))) | p(s(bool,happ(s(fun(bool,bool),X984),s(bool,X986)))) | s(bool,'T_0') = s(bool,X985) | s(bool,'T_0') = s(bool,sK14(X984,bool))) )),
  inference(superposition,[],[f411,f701])).
fof(f23516,plain,(
  ( ! [X978,X982,X980,X979,X983,X981] : (s(bool,'const.bool.COND_3'(s(bool,X982),s(bool,X981),s(bool,happ(s(fun(X978,bool),X979),s(X978,X983))))) = s(bool,happ(s(fun(X978,bool),X979),s(X978,'const.bool.COND_3'(s(bool,X982),s(X978,X980),s(X978,X983))))) | s(bool,'T_0') = s(bool,X981) | s(bool,'T_0') = s(bool,happ(s(fun(X978,bool),X979),s(X978,X980)))) )),
  inference(superposition,[],[f479,f701])).
fof(f23515,plain,(
  ( ! [X976,X975,X973,X977,X974,X972] : (s(bool,'const.bool.COND_3'(s(bool,X976),s(bool,happ(s(fun(X972,bool),X973),s(X972,X977))),s(bool,X975))) = s(bool,happ(s(fun(X972,bool),X973),s(X972,'const.bool.COND_3'(s(bool,X976),s(X972,X977),s(X972,X974))))) | s(bool,'T_0') = s(bool,X975) | s(bool,'T_0') = s(bool,happ(s(fun(X972,bool),X973),s(X972,X974)))) )),
  inference(superposition,[],[f479,f701])).
fof(f23514,plain,(
  ( ! [X971,X969,X967,X970,X968,X966] : (s(bool,'const.bool.COND_3'(s(bool,X970),s(bool,X969),s(bool,happ(s(fun(X966,bool),X971),s(X966,X968))))) = s(bool,happ(s(fun(X966,bool),'const.bool.COND_3'(s(bool,X970),s(fun(X966,bool),X967),s(fun(X966,bool),X971))),s(X966,X968))) | s(bool,'T_0') = s(bool,X969) | s(bool,'T_0') = s(bool,happ(s(fun(X966,bool),X967),s(X966,X968)))) )),
  inference(superposition,[],[f480,f701])).
fof(f23513,plain,(
  ( ! [X963,X961,X965,X962,X960,X964] : (s(bool,'const.bool.COND_3'(s(bool,X964),s(bool,happ(s(fun(X960,bool),X965),s(X960,X962))),s(bool,X963))) = s(bool,happ(s(fun(X960,bool),'const.bool.COND_3'(s(bool,X964),s(fun(X960,bool),X965),s(fun(X960,bool),X961))),s(X960,X962))) | s(bool,'T_0') = s(bool,X963) | s(bool,'T_0') = s(bool,happ(s(fun(X960,bool),X961),s(X960,X962)))) )),
  inference(superposition,[],[f480,f701])).
fof(f23508,plain,(
  ( ! [X938,X940,X939,X941] : (s(bool,X941) != s(bool,f(s(X938,sK22(X939,X940,bool,X938)))) | s(fun(X938,bool),X939) = s(fun(X938,bool),X940) | s(bool,'T_0') = s(bool,X941) | s(bool,'T_0') = s(bool,g(s(X938,sK22(X939,X940,bool,X938))))) )),
  inference(superposition,[],[f478,f701])).
fof(f23477,plain,(
  ( ! [X810,X811,X809] : (s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X809),s('type.num.num',X810))) = s('type.num.num','const.bool.COND_3'(s(bool,X811),s('type.num.num',X809),s('type.num.num',X810))) | s(bool,'T_0') = s(bool,X811) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X809),s('type.num.num',X810)))) )),
  inference(superposition,[],[f431,f701])).
fof(f23476,plain,(
  ( ! [X808,X806,X807] : (s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X806),s('type.num.num',X807))) = s('type.num.num','const.bool.COND_3'(s(bool,X808),s('type.num.num',X807),s('type.num.num',X806))) | s(bool,'T_0') = s(bool,X808) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X806),s('type.num.num',X807)))) )),
  inference(superposition,[],[f432,f701])).
fof(f23266,plain,(
  ( ! [X246,X247,X245] : (sP4('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X247),s(bool,X246)),X245) | s(bool,'T_0') = s(bool,X246) | s(bool,'T_0') = s(bool,X245)) )),
  inference(superposition,[],[f991,f701])).
fof(f23263,plain,(
  ( ! [X237,X238,X236] : (sP4('T_0',X236,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X238),s(bool,X237))) | s(bool,'T_0') = s(bool,X237) | s(bool,'T_0') = s(bool,X236)) )),
  inference(superposition,[],[f930,f701])).
fof(f23261,plain,(
  ( ! [X231,X230,X232] : (sP4('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X231),s(bool,X232)),X230) | s(bool,'T_0') = s(bool,X231) | s(bool,'T_0') = s(bool,X230)) )),
  inference(superposition,[],[f898,f701])).
fof(f23255,plain,(
  ( ! [X210,X208,X209] : (sP5('T_0',X208,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X210),s(bool,X209))) | s(bool,'T_0') = s(bool,X209) | s(bool,'T_0') = s(bool,X208)) )),
  inference(superposition,[],[f839,f701])).
fof(f23253,plain,(
  ( ! [X204,X202,X203] : (~sP5('F_0',X202,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X204),s(bool,X203))) | s(bool,'T_0') = s(bool,X203) | s(bool,'T_0') = s(bool,X202)) )),
  inference(superposition,[],[f838,f701])).
fof(f23252,plain,(
  ( ! [X200,X199,X201] : (sP5('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X201),s(bool,X200)),X199) | s(bool,'T_0') = s(bool,X200) | s(bool,'T_0') = s(bool,X199)) )),
  inference(superposition,[],[f837,f701])).
fof(f23250,plain,(
  ( ! [X194,X195,X193] : (~sP5('F_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X195),s(bool,X194)),X193) | s(bool,'T_0') = s(bool,X194) | s(bool,'T_0') = s(bool,X193)) )),
  inference(superposition,[],[f836,f701])).
fof(f23247,plain,(
  ( ! [X185,X186,X184] : (sP4('T_0',X184,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X185),s(bool,X186))) | s(bool,'T_0') = s(bool,X185) | s(bool,'T_0') = s(bool,X184)) )),
  inference(superposition,[],[f825,f701])).
fof(f23241,plain,(
  ( ! [X163,X164,X162] : (sP5('T_0',X162,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X163),s(bool,X164))) | s(bool,'T_0') = s(bool,X163) | s(bool,'T_0') = s(bool,X162)) )),
  inference(superposition,[],[f778,f701])).
fof(f23239,plain,(
  ( ! [X158,X156,X157] : (~sP5('F_0',X156,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X157),s(bool,X158))) | s(bool,'T_0') = s(bool,X157) | s(bool,'T_0') = s(bool,X156)) )),
  inference(superposition,[],[f777,f701])).
fof(f23237,plain,(
  ( ! [X152,X151,X150] : (sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X151),s(bool,X152)),X150) | s(bool,'T_0') = s(bool,X151) | s(bool,'T_0') = s(bool,X150)) )),
  inference(superposition,[],[f776,f701])).
fof(f23235,plain,(
  ( ! [X146,X144,X145] : (~sP5('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X145),s(bool,X146)),X144) | s(bool,'T_0') = s(bool,X145) | s(bool,'T_0') = s(bool,X144)) )),
  inference(superposition,[],[f775,f701])).
fof(f23209,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X40] : (s(bool,X35) != s(bool,X36) | s(X37,X38) != s(X37,X39) | s(X37,X40) != s(X37,X41) | s(X37,'const.bool.COND_3'(s(bool,X34),s(X37,X41),s(X37,X39))) = s(X37,'const.bool.COND_3'(s(bool,X36),s(X37,X40),s(X37,X38))) | s(bool,'T_0') = s(bool,X35) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f428,f701])).
fof(f25796,plain,(
  ( ! [X24,X23,X21,X19,X25,X22,X20,X18] : (s(bool,X19) != s(bool,X20) | s(X21,X22) != s(X21,X23) | p(s(bool,X18)) | s(X21,'const.bool.COND_3'(s(bool,X18),s(X21,X25),s(X21,X23))) = s(X21,'const.bool.COND_3'(s(bool,X20),s(X21,X24),s(X21,X22))) | s(bool,'T_0') = s(bool,X19)) )),
  inference(subsumption_resolution,[],[f23207,f354])).
fof(f23207,plain,(
  ( ! [X24,X23,X21,X19,X25,X22,X20,X18] : (s(bool,X19) != s(bool,X20) | s(X21,X22) != s(X21,X23) | p(s(bool,X18)) | s(X21,'const.bool.COND_3'(s(bool,X18),s(X21,X25),s(X21,X23))) = s(X21,'const.bool.COND_3'(s(bool,X20),s(X21,X24),s(X21,X22))) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f427,f701])).
fof(f23186,plain,(
  ( ! [X30,X28,X31,X29,X27,X32] : (s(bool,'T_0') = s(bool,happ(s(fun(X28,bool),'const.bool.COND_3'(s(bool,X27),s(fun(X28,bool),X29),s(fun(X28,bool),X31))),s(X28,X30))) | s(bool,X32) = s(bool,'const.bool.COND_3'(s(bool,X27),s(bool,happ(s(fun(X28,bool),X29),s(X28,X30))),s(bool,happ(s(fun(X28,bool),X31),s(X28,X30))))) | s(bool,'T_0') = s(bool,X32)) )),
  inference(superposition,[],[f701,f480])).
fof(f23185,plain,(
  ( ! [X26,X24,X23,X21,X25,X22] : (s(bool,'T_0') = s(bool,happ(s(fun(X22,bool),X23),s(X22,'const.bool.COND_3'(s(bool,X21),s(X22,X24),s(X22,X25))))) | s(bool,X26) = s(bool,'const.bool.COND_3'(s(bool,X21),s(bool,happ(s(fun(X22,bool),X23),s(X22,X24))),s(bool,happ(s(fun(X22,bool),X23),s(X22,X25))))) | s(bool,'T_0') = s(bool,X26)) )),
  inference(superposition,[],[f701,f479])).
fof(f23184,plain,(
  ( ! [X19,X20,X18] : (s(bool,'T_0') = s(bool,X19) | s(bool,X20) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X18),s(bool,X19))) | s(bool,'T_0') = s(bool,X20)) )),
  inference(superposition,[],[f701,f414])).
fof(f23183,plain,(
  ( ! [X17,X15,X16] : (s(bool,'T_0') = s(bool,X15) | s(bool,X17) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X15),s(bool,X16))) | s(bool,'T_0') = s(bool,X17)) )),
  inference(superposition,[],[f701,f413])).
fof(f23173,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f701,f701])).
fof(f23159,plain,(
  ( ! [X39,X37,X35,X38,X36,X34] : (s(bool,X39) = s(bool,happ(s(fun(X35,bool),'const.bool.COND_3'(s(bool,X34),s(fun(X35,bool),X36),s(fun(X35,bool),X38))),s(X35,X37))) | s(bool,'T_0') = s(bool,X39) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X34),s(bool,happ(s(fun(X35,bool),X36),s(X35,X37))),s(bool,happ(s(fun(X35,bool),X38),s(X35,X37)))))) )),
  inference(superposition,[],[f701,f480])).
fof(f23158,plain,(
  ( ! [X30,X28,X33,X31,X29,X32] : (s(bool,X33) = s(bool,happ(s(fun(X29,bool),X30),s(X29,'const.bool.COND_3'(s(bool,X28),s(X29,X31),s(X29,X32))))) | s(bool,'T_0') = s(bool,X33) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X28),s(bool,happ(s(fun(X29,bool),X30),s(X29,X31))),s(bool,happ(s(fun(X29,bool),X30),s(X29,X32)))))) )),
  inference(superposition,[],[f701,f479])).
fof(f23157,plain,(
  ( ! [X26,X27,X25] : (s(bool,X26) = s(bool,X27) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X25),s(bool,X26)))) )),
  inference(superposition,[],[f701,f414])).
fof(f23156,plain,(
  ( ! [X24,X23,X22] : (s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X22),s(bool,X23)))) )),
  inference(superposition,[],[f701,f413])).
fof(f23139,plain,(
  ( ! [X8,X7] : (sP1(X7,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X8),s(bool,'T_0'))) | p(s(bool,X7))) )),
  inference(forward_demodulation,[],[f23138,f1595])).
fof(f23138,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X7)) | sP1(X7,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X8),s(bool,'const.bool.COND_3'(s(bool,X6),s(bool,'T_0'),s(bool,'T_0')))))) )),
  inference(subsumption_resolution,[],[f23133,f335])).
fof(f23133,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'T_0')) | p(s(bool,X7)) | sP1(X7,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X8),s(bool,'const.bool.COND_3'(s(bool,X6),s(bool,'T_0'),s(bool,'T_0')))))) )),
  inference(superposition,[],[f864,f1595])).
fof(f23129,plain,(
  ( ! [X0] : (sP1('F_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f335,f592,f864])).
fof(f23093,plain,(
  ( ! [X1128,X1122,X1126,X1124,X1129,X1123,X1127,X1125] : (p(s(bool,happ(s(fun(X1122,bool),happ(s(fun(X1124,fun(X1122,bool)),'const.bool.COND_3'(s(bool,X1123),s(fun(X1124,fun(X1122,bool)),X1125),s(fun(X1124,fun(X1122,bool)),X1127))),s(X1124,X1126))),s(X1122,sK21('const.bool.COND_3'(s(bool,X1123),s(fun(X1122,bool),happ(s(fun(X1124,fun(X1122,bool)),X1125),s(X1124,X1126))),s(fun(X1122,bool),happ(s(fun(X1124,fun(X1122,bool)),X1127),s(X1124,X1126)))),X1129,X1122))))) | ~p(s(bool,happ(s(fun(X1122,bool),happ(s(fun(X1124,fun(X1122,bool)),'const.bool.COND_3'(s(bool,X1123),s(fun(X1124,fun(X1122,bool)),X1125),s(fun(X1124,fun(X1122,bool)),X1127))),s(X1124,X1126))),s(X1122,X1128)))) | p(s(bool,happ(s(fun(X1122,bool),X1129),s(X1122,sK21('const.bool.COND_3'(s(bool,X1123),s(fun(X1122,bool),happ(s(fun(X1124,fun(X1122,bool)),X1125),s(X1124,X1126))),s(fun(X1122,bool),happ(s(fun(X1124,fun(X1122,bool)),X1127),s(X1124,X1126)))),X1129,X1122)))))) )),
  inference(forward_demodulation,[],[f18968,f480])).
fof(f18968,plain,(
  ( ! [X1128,X1122,X1126,X1124,X1129,X1123,X1127,X1125] : (~p(s(bool,happ(s(fun(X1122,bool),happ(s(fun(X1124,fun(X1122,bool)),'const.bool.COND_3'(s(bool,X1123),s(fun(X1124,fun(X1122,bool)),X1125),s(fun(X1124,fun(X1122,bool)),X1127))),s(X1124,X1126))),s(X1122,X1128)))) | p(s(bool,happ(s(fun(X1122,bool),X1129),s(X1122,sK21('const.bool.COND_3'(s(bool,X1123),s(fun(X1122,bool),happ(s(fun(X1124,fun(X1122,bool)),X1125),s(X1124,X1126))),s(fun(X1122,bool),happ(s(fun(X1124,fun(X1122,bool)),X1127),s(X1124,X1126)))),X1129,X1122))))) | p(s(bool,happ(s(fun(X1122,bool),'const.bool.COND_3'(s(bool,X1123),s(fun(X1122,bool),happ(s(fun(X1124,fun(X1122,bool)),X1125),s(X1124,X1126))),s(fun(X1122,bool),happ(s(fun(X1124,fun(X1122,bool)),X1127),s(X1124,X1126))))),s(X1122,sK21('const.bool.COND_3'(s(bool,X1123),s(fun(X1122,bool),happ(s(fun(X1124,fun(X1122,bool)),X1125),s(X1124,X1126))),s(fun(X1122,bool),happ(s(fun(X1124,fun(X1122,bool)),X1127),s(X1124,X1126)))),X1129,X1122)))))) )),
  inference(superposition,[],[f424,f480])).
fof(f23092,plain,(
  ( ! [X1118,X1116,X1120,X1115,X1119,X1117,X1121,X1114] : (p(s(bool,happ(s(fun(X1114,bool),happ(s(fun(X1116,fun(X1114,bool)),'const.bool.COND_3'(s(bool,X1115),s(fun(X1116,fun(X1114,bool)),X1117),s(fun(X1116,fun(X1114,bool)),X1119))),s(X1116,X1118))),s(X1114,sK21(X1121,'const.bool.COND_3'(s(bool,X1115),s(fun(X1114,bool),happ(s(fun(X1116,fun(X1114,bool)),X1117),s(X1116,X1118))),s(fun(X1114,bool),happ(s(fun(X1116,fun(X1114,bool)),X1119),s(X1116,X1118)))),X1114))))) | ~p(s(bool,happ(s(fun(X1114,bool),happ(s(fun(X1116,fun(X1114,bool)),'const.bool.COND_3'(s(bool,X1115),s(fun(X1116,fun(X1114,bool)),X1117),s(fun(X1116,fun(X1114,bool)),X1119))),s(X1116,X1118))),s(X1114,X1120)))) | p(s(bool,happ(s(fun(X1114,bool),X1121),s(X1114,sK21(X1121,'const.bool.COND_3'(s(bool,X1115),s(fun(X1114,bool),happ(s(fun(X1116,fun(X1114,bool)),X1117),s(X1116,X1118))),s(fun(X1114,bool),happ(s(fun(X1116,fun(X1114,bool)),X1119),s(X1116,X1118)))),X1114)))))) )),
  inference(forward_demodulation,[],[f18967,f480])).
fof(f18967,plain,(
  ( ! [X1118,X1116,X1120,X1115,X1119,X1117,X1121,X1114] : (~p(s(bool,happ(s(fun(X1114,bool),happ(s(fun(X1116,fun(X1114,bool)),'const.bool.COND_3'(s(bool,X1115),s(fun(X1116,fun(X1114,bool)),X1117),s(fun(X1116,fun(X1114,bool)),X1119))),s(X1116,X1118))),s(X1114,X1120)))) | p(s(bool,happ(s(fun(X1114,bool),'const.bool.COND_3'(s(bool,X1115),s(fun(X1114,bool),happ(s(fun(X1116,fun(X1114,bool)),X1117),s(X1116,X1118))),s(fun(X1114,bool),happ(s(fun(X1116,fun(X1114,bool)),X1119),s(X1116,X1118))))),s(X1114,sK21(X1121,'const.bool.COND_3'(s(bool,X1115),s(fun(X1114,bool),happ(s(fun(X1116,fun(X1114,bool)),X1117),s(X1116,X1118))),s(fun(X1114,bool),happ(s(fun(X1116,fun(X1114,bool)),X1119),s(X1116,X1118)))),X1114))))) | p(s(bool,happ(s(fun(X1114,bool),X1121),s(X1114,sK21(X1121,'const.bool.COND_3'(s(bool,X1115),s(fun(X1114,bool),happ(s(fun(X1116,fun(X1114,bool)),X1117),s(X1116,X1118))),s(fun(X1114,bool),happ(s(fun(X1116,fun(X1114,bool)),X1119),s(X1116,X1118)))),X1114)))))) )),
  inference(superposition,[],[f423,f480])).
fof(f23091,plain,(
  ( ! [X1112,X1106,X1110,X1108,X1113,X1107,X1111,X1109] : (p(s(bool,happ(s(fun(X1106,bool),happ(s(fun(X1108,fun(X1106,bool)),'const.bool.COND_3'(s(bool,X1107),s(fun(X1108,fun(X1106,bool)),X1109),s(fun(X1108,fun(X1106,bool)),X1111))),s(X1108,X1110))),s(X1106,sK20('const.bool.COND_3'(s(bool,X1107),s(fun(X1106,bool),happ(s(fun(X1108,fun(X1106,bool)),X1109),s(X1108,X1110))),s(fun(X1106,bool),happ(s(fun(X1108,fun(X1106,bool)),X1111),s(X1108,X1110)))),X1113,X1106))))) | ~p(s(bool,happ(s(fun(X1106,bool),happ(s(fun(X1108,fun(X1106,bool)),'const.bool.COND_3'(s(bool,X1107),s(fun(X1108,fun(X1106,bool)),X1109),s(fun(X1108,fun(X1106,bool)),X1111))),s(X1108,X1110))),s(X1106,X1112)))) | p(s(bool,happ(s(fun(X1106,bool),X1113),s(X1106,sK19('const.bool.COND_3'(s(bool,X1107),s(fun(X1106,bool),happ(s(fun(X1108,fun(X1106,bool)),X1109),s(X1108,X1110))),s(fun(X1106,bool),happ(s(fun(X1108,fun(X1106,bool)),X1111),s(X1108,X1110)))),X1113,X1106)))))) )),
  inference(forward_demodulation,[],[f18966,f480])).
fof(f18966,plain,(
  ( ! [X1112,X1106,X1110,X1108,X1113,X1107,X1111,X1109] : (~p(s(bool,happ(s(fun(X1106,bool),happ(s(fun(X1108,fun(X1106,bool)),'const.bool.COND_3'(s(bool,X1107),s(fun(X1108,fun(X1106,bool)),X1109),s(fun(X1108,fun(X1106,bool)),X1111))),s(X1108,X1110))),s(X1106,X1112)))) | p(s(bool,happ(s(fun(X1106,bool),X1113),s(X1106,sK19('const.bool.COND_3'(s(bool,X1107),s(fun(X1106,bool),happ(s(fun(X1108,fun(X1106,bool)),X1109),s(X1108,X1110))),s(fun(X1106,bool),happ(s(fun(X1108,fun(X1106,bool)),X1111),s(X1108,X1110)))),X1113,X1106))))) | p(s(bool,happ(s(fun(X1106,bool),'const.bool.COND_3'(s(bool,X1107),s(fun(X1106,bool),happ(s(fun(X1108,fun(X1106,bool)),X1109),s(X1108,X1110))),s(fun(X1106,bool),happ(s(fun(X1108,fun(X1106,bool)),X1111),s(X1108,X1110))))),s(X1106,sK20('const.bool.COND_3'(s(bool,X1107),s(fun(X1106,bool),happ(s(fun(X1108,fun(X1106,bool)),X1109),s(X1108,X1110))),s(fun(X1106,bool),happ(s(fun(X1108,fun(X1106,bool)),X1111),s(X1108,X1110)))),X1113,X1106)))))) )),
  inference(superposition,[],[f422,f480])).
fof(f23090,plain,(
  ( ! [X1104,X1099,X1103,X1101,X1105,X1098,X1102,X1100] : (p(s(bool,happ(s(fun(X1098,bool),happ(s(fun(X1100,fun(X1098,bool)),'const.bool.COND_3'(s(bool,X1099),s(fun(X1100,fun(X1098,bool)),X1101),s(fun(X1100,fun(X1098,bool)),X1103))),s(X1100,X1102))),s(X1098,sK19(X1105,'const.bool.COND_3'(s(bool,X1099),s(fun(X1098,bool),happ(s(fun(X1100,fun(X1098,bool)),X1101),s(X1100,X1102))),s(fun(X1098,bool),happ(s(fun(X1100,fun(X1098,bool)),X1103),s(X1100,X1102)))),X1098))))) | ~p(s(bool,happ(s(fun(X1098,bool),happ(s(fun(X1100,fun(X1098,bool)),'const.bool.COND_3'(s(bool,X1099),s(fun(X1100,fun(X1098,bool)),X1101),s(fun(X1100,fun(X1098,bool)),X1103))),s(X1100,X1102))),s(X1098,X1104)))) | p(s(bool,happ(s(fun(X1098,bool),X1105),s(X1098,sK20(X1105,'const.bool.COND_3'(s(bool,X1099),s(fun(X1098,bool),happ(s(fun(X1100,fun(X1098,bool)),X1101),s(X1100,X1102))),s(fun(X1098,bool),happ(s(fun(X1100,fun(X1098,bool)),X1103),s(X1100,X1102)))),X1098)))))) )),
  inference(forward_demodulation,[],[f18965,f480])).
fof(f18965,plain,(
  ( ! [X1104,X1099,X1103,X1101,X1105,X1098,X1102,X1100] : (~p(s(bool,happ(s(fun(X1098,bool),happ(s(fun(X1100,fun(X1098,bool)),'const.bool.COND_3'(s(bool,X1099),s(fun(X1100,fun(X1098,bool)),X1101),s(fun(X1100,fun(X1098,bool)),X1103))),s(X1100,X1102))),s(X1098,X1104)))) | p(s(bool,happ(s(fun(X1098,bool),'const.bool.COND_3'(s(bool,X1099),s(fun(X1098,bool),happ(s(fun(X1100,fun(X1098,bool)),X1101),s(X1100,X1102))),s(fun(X1098,bool),happ(s(fun(X1100,fun(X1098,bool)),X1103),s(X1100,X1102))))),s(X1098,sK19(X1105,'const.bool.COND_3'(s(bool,X1099),s(fun(X1098,bool),happ(s(fun(X1100,fun(X1098,bool)),X1101),s(X1100,X1102))),s(fun(X1098,bool),happ(s(fun(X1100,fun(X1098,bool)),X1103),s(X1100,X1102)))),X1098))))) | p(s(bool,happ(s(fun(X1098,bool),X1105),s(X1098,sK20(X1105,'const.bool.COND_3'(s(bool,X1099),s(fun(X1098,bool),happ(s(fun(X1100,fun(X1098,bool)),X1101),s(X1100,X1102))),s(fun(X1098,bool),happ(s(fun(X1100,fun(X1098,bool)),X1103),s(X1100,X1102)))),X1098)))))) )),
  inference(superposition,[],[f421,f480])).
fof(f23089,plain,(
  ( ! [X1097,X1091,X1095,X1093,X1096,X1090,X1094,X1092] : (p(s(bool,happ(s(fun(X1090,bool),happ(s(fun(X1092,fun(X1090,bool)),'const.bool.COND_3'(s(bool,X1091),s(fun(X1092,fun(X1090,bool)),X1093),s(fun(X1092,fun(X1090,bool)),X1095))),s(X1092,X1094))),s(X1090,X1097)))) | ~p(s(bool,happ(s(fun(X1090,bool),happ(s(fun(X1092,fun(X1090,bool)),'const.bool.COND_3'(s(bool,X1091),s(fun(X1092,fun(X1090,bool)),X1093),s(fun(X1092,fun(X1090,bool)),X1095))),s(X1092,X1094))),s(X1090,sK18('const.bool.COND_3'(s(bool,X1091),s(fun(X1090,bool),happ(s(fun(X1092,fun(X1090,bool)),X1093),s(X1092,X1094))),s(fun(X1090,bool),happ(s(fun(X1092,fun(X1090,bool)),X1095),s(X1092,X1094)))),X1096,X1090))))) | ~p(s(bool,happ(s(fun(X1090,bool),X1096),s(X1090,sK17('const.bool.COND_3'(s(bool,X1091),s(fun(X1090,bool),happ(s(fun(X1092,fun(X1090,bool)),X1093),s(X1092,X1094))),s(fun(X1090,bool),happ(s(fun(X1092,fun(X1090,bool)),X1095),s(X1092,X1094)))),X1096,X1090)))))) )),
  inference(forward_demodulation,[],[f18964,f480])).
fof(f18964,plain,(
  ( ! [X1097,X1091,X1095,X1093,X1096,X1090,X1094,X1092] : (~p(s(bool,happ(s(fun(X1090,bool),happ(s(fun(X1092,fun(X1090,bool)),'const.bool.COND_3'(s(bool,X1091),s(fun(X1092,fun(X1090,bool)),X1093),s(fun(X1092,fun(X1090,bool)),X1095))),s(X1092,X1094))),s(X1090,sK18('const.bool.COND_3'(s(bool,X1091),s(fun(X1090,bool),happ(s(fun(X1092,fun(X1090,bool)),X1093),s(X1092,X1094))),s(fun(X1090,bool),happ(s(fun(X1092,fun(X1090,bool)),X1095),s(X1092,X1094)))),X1096,X1090))))) | p(s(bool,happ(s(fun(X1090,bool),'const.bool.COND_3'(s(bool,X1091),s(fun(X1090,bool),happ(s(fun(X1092,fun(X1090,bool)),X1093),s(X1092,X1094))),s(fun(X1090,bool),happ(s(fun(X1092,fun(X1090,bool)),X1095),s(X1092,X1094))))),s(X1090,X1097)))) | ~p(s(bool,happ(s(fun(X1090,bool),X1096),s(X1090,sK17('const.bool.COND_3'(s(bool,X1091),s(fun(X1090,bool),happ(s(fun(X1092,fun(X1090,bool)),X1093),s(X1092,X1094))),s(fun(X1090,bool),happ(s(fun(X1092,fun(X1090,bool)),X1095),s(X1092,X1094)))),X1096,X1090)))))) )),
  inference(superposition,[],[f420,f480])).
fof(f18963,plain,(
  ( ! [X1087,X1085,X1089,X1082,X1086,X1084,X1088,X1083] : (~p(s(bool,happ(s(fun(X1082,bool),happ(s(fun(X1084,fun(X1082,bool)),'const.bool.COND_3'(s(bool,X1083),s(fun(X1084,fun(X1082,bool)),X1085),s(fun(X1084,fun(X1082,bool)),X1087))),s(X1084,X1086))),s(X1082,sK18('const.bool.COND_3'(s(bool,X1083),s(fun(X1082,bool),happ(s(fun(X1084,fun(X1082,bool)),X1085),s(X1084,X1086))),s(fun(X1082,bool),happ(s(fun(X1084,fun(X1082,bool)),X1087),s(X1084,X1086)))),X1088,X1082))))) | p(s(bool,happ(s(fun(X1082,bool),X1088),s(X1082,X1089)))) | ~p(s(bool,happ(s(fun(X1082,bool),X1088),s(X1082,sK17('const.bool.COND_3'(s(bool,X1083),s(fun(X1082,bool),happ(s(fun(X1084,fun(X1082,bool)),X1085),s(X1084,X1086))),s(fun(X1082,bool),happ(s(fun(X1084,fun(X1082,bool)),X1087),s(X1084,X1086)))),X1088,X1082)))))) )),
  inference(superposition,[],[f419,f480])).
fof(f23088,plain,(
  ( ! [X1081,X1075,X1079,X1077,X1080,X1074,X1078,X1076] : (p(s(bool,happ(s(fun(X1074,bool),happ(s(fun(X1076,fun(X1074,bool)),'const.bool.COND_3'(s(bool,X1075),s(fun(X1076,fun(X1074,bool)),X1077),s(fun(X1076,fun(X1074,bool)),X1079))),s(X1076,X1078))),s(X1074,X1081)))) | ~p(s(bool,happ(s(fun(X1074,bool),happ(s(fun(X1076,fun(X1074,bool)),'const.bool.COND_3'(s(bool,X1075),s(fun(X1076,fun(X1074,bool)),X1077),s(fun(X1076,fun(X1074,bool)),X1079))),s(X1076,X1078))),s(X1074,sK16('const.bool.COND_3'(s(bool,X1075),s(fun(X1074,bool),happ(s(fun(X1076,fun(X1074,bool)),X1077),s(X1076,X1078))),s(fun(X1074,bool),happ(s(fun(X1076,fun(X1074,bool)),X1079),s(X1076,X1078)))),X1080,X1074))))) | ~p(s(bool,happ(s(fun(X1074,bool),X1080),s(X1074,sK16('const.bool.COND_3'(s(bool,X1075),s(fun(X1074,bool),happ(s(fun(X1076,fun(X1074,bool)),X1077),s(X1076,X1078))),s(fun(X1074,bool),happ(s(fun(X1076,fun(X1074,bool)),X1079),s(X1076,X1078)))),X1080,X1074)))))) )),
  inference(forward_demodulation,[],[f18962,f480])).
fof(f18962,plain,(
  ( ! [X1081,X1075,X1079,X1077,X1080,X1074,X1078,X1076] : (~p(s(bool,happ(s(fun(X1074,bool),happ(s(fun(X1076,fun(X1074,bool)),'const.bool.COND_3'(s(bool,X1075),s(fun(X1076,fun(X1074,bool)),X1077),s(fun(X1076,fun(X1074,bool)),X1079))),s(X1076,X1078))),s(X1074,sK16('const.bool.COND_3'(s(bool,X1075),s(fun(X1074,bool),happ(s(fun(X1076,fun(X1074,bool)),X1077),s(X1076,X1078))),s(fun(X1074,bool),happ(s(fun(X1076,fun(X1074,bool)),X1079),s(X1076,X1078)))),X1080,X1074))))) | p(s(bool,happ(s(fun(X1074,bool),'const.bool.COND_3'(s(bool,X1075),s(fun(X1074,bool),happ(s(fun(X1076,fun(X1074,bool)),X1077),s(X1076,X1078))),s(fun(X1074,bool),happ(s(fun(X1076,fun(X1074,bool)),X1079),s(X1076,X1078))))),s(X1074,X1081)))) | ~p(s(bool,happ(s(fun(X1074,bool),X1080),s(X1074,sK16('const.bool.COND_3'(s(bool,X1075),s(fun(X1074,bool),happ(s(fun(X1076,fun(X1074,bool)),X1077),s(X1076,X1078))),s(fun(X1074,bool),happ(s(fun(X1076,fun(X1074,bool)),X1079),s(X1076,X1078)))),X1080,X1074)))))) )),
  inference(superposition,[],[f418,f480])).
fof(f18961,plain,(
  ( ! [X1073,X1066,X1070,X1068,X1072,X1067,X1071,X1069] : (~p(s(bool,happ(s(fun(X1066,bool),happ(s(fun(X1068,fun(X1066,bool)),'const.bool.COND_3'(s(bool,X1067),s(fun(X1068,fun(X1066,bool)),X1069),s(fun(X1068,fun(X1066,bool)),X1071))),s(X1068,X1070))),s(X1066,sK16('const.bool.COND_3'(s(bool,X1067),s(fun(X1066,bool),happ(s(fun(X1068,fun(X1066,bool)),X1069),s(X1068,X1070))),s(fun(X1066,bool),happ(s(fun(X1068,fun(X1066,bool)),X1071),s(X1068,X1070)))),X1072,X1066))))) | p(s(bool,happ(s(fun(X1066,bool),X1072),s(X1066,X1073)))) | ~p(s(bool,happ(s(fun(X1066,bool),X1072),s(X1066,sK16('const.bool.COND_3'(s(bool,X1067),s(fun(X1066,bool),happ(s(fun(X1068,fun(X1066,bool)),X1069),s(X1068,X1070))),s(fun(X1066,bool),happ(s(fun(X1068,fun(X1066,bool)),X1071),s(X1068,X1070)))),X1072,X1066)))))) )),
  inference(superposition,[],[f417,f480])).
fof(f23087,plain,(
  ( ! [X1064,X1062,X1060,X1065,X1059,X1063,X1061] : (p(s(bool,happ(s(fun(X1059,bool),happ(s(fun(X1061,fun(X1059,bool)),'const.bool.COND_3'(s(bool,X1060),s(fun(X1061,fun(X1059,bool)),X1062),s(fun(X1061,fun(X1059,bool)),X1064))),s(X1061,X1063))),s(X1059,X1065)))) | ~p(s(bool,happ(s(fun(X1059,bool),happ(s(fun(X1061,fun(X1059,bool)),'const.bool.COND_3'(s(bool,X1060),s(fun(X1061,fun(X1059,bool)),X1062),s(fun(X1061,fun(X1059,bool)),X1064))),s(X1061,X1063))),s(X1059,sK15('const.bool.COND_3'(s(bool,X1060),s(fun(X1059,bool),happ(s(fun(X1061,fun(X1059,bool)),X1062),s(X1061,X1063))),s(fun(X1059,bool),happ(s(fun(X1061,fun(X1059,bool)),X1064),s(X1061,X1063)))),X1059)))))) )),
  inference(forward_demodulation,[],[f18960,f480])).
fof(f18960,plain,(
  ( ! [X1064,X1062,X1060,X1065,X1059,X1063,X1061] : (~p(s(bool,happ(s(fun(X1059,bool),happ(s(fun(X1061,fun(X1059,bool)),'const.bool.COND_3'(s(bool,X1060),s(fun(X1061,fun(X1059,bool)),X1062),s(fun(X1061,fun(X1059,bool)),X1064))),s(X1061,X1063))),s(X1059,sK15('const.bool.COND_3'(s(bool,X1060),s(fun(X1059,bool),happ(s(fun(X1061,fun(X1059,bool)),X1062),s(X1061,X1063))),s(fun(X1059,bool),happ(s(fun(X1061,fun(X1059,bool)),X1064),s(X1061,X1063)))),X1059))))) | p(s(bool,happ(s(fun(X1059,bool),'const.bool.COND_3'(s(bool,X1060),s(fun(X1059,bool),happ(s(fun(X1061,fun(X1059,bool)),X1062),s(X1061,X1063))),s(fun(X1059,bool),happ(s(fun(X1061,fun(X1059,bool)),X1064),s(X1061,X1063))))),s(X1059,X1065))))) )),
  inference(superposition,[],[f412,f480])).
fof(f23086,plain,(
  ( ! [X1054,X1052,X1058,X1056,X1055,X1053,X1057] : (p(s(bool,happ(s(fun(X1052,bool),happ(s(fun(X1054,fun(X1052,bool)),'const.bool.COND_3'(s(bool,X1053),s(fun(X1054,fun(X1052,bool)),X1055),s(fun(X1054,fun(X1052,bool)),X1057))),s(X1054,X1056))),s(X1052,X1058)))) | ~p(s(bool,happ(s(fun(X1052,bool),happ(s(fun(X1054,fun(X1052,bool)),'const.bool.COND_3'(s(bool,X1053),s(fun(X1054,fun(X1052,bool)),X1055),s(fun(X1054,fun(X1052,bool)),X1057))),s(X1054,X1056))),s(X1052,sK14('const.bool.COND_3'(s(bool,X1053),s(fun(X1052,bool),happ(s(fun(X1054,fun(X1052,bool)),X1055),s(X1054,X1056))),s(fun(X1052,bool),happ(s(fun(X1054,fun(X1052,bool)),X1057),s(X1054,X1056)))),X1052)))))) )),
  inference(forward_demodulation,[],[f18959,f480])).
fof(f18959,plain,(
  ( ! [X1054,X1052,X1058,X1056,X1055,X1053,X1057] : (~p(s(bool,happ(s(fun(X1052,bool),happ(s(fun(X1054,fun(X1052,bool)),'const.bool.COND_3'(s(bool,X1053),s(fun(X1054,fun(X1052,bool)),X1055),s(fun(X1054,fun(X1052,bool)),X1057))),s(X1054,X1056))),s(X1052,sK14('const.bool.COND_3'(s(bool,X1053),s(fun(X1052,bool),happ(s(fun(X1054,fun(X1052,bool)),X1055),s(X1054,X1056))),s(fun(X1052,bool),happ(s(fun(X1054,fun(X1052,bool)),X1057),s(X1054,X1056)))),X1052))))) | p(s(bool,happ(s(fun(X1052,bool),'const.bool.COND_3'(s(bool,X1053),s(fun(X1052,bool),happ(s(fun(X1054,fun(X1052,bool)),X1055),s(X1054,X1056))),s(fun(X1052,bool),happ(s(fun(X1054,fun(X1052,bool)),X1057),s(X1054,X1056))))),s(X1052,X1058))))) )),
  inference(superposition,[],[f411,f480])).
fof(f18866,plain,(
  ( ! [X528,X524,X526,X525,X527,X523] : (sP4('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X528),s(bool,happ(s(fun(X524,bool),'const.bool.COND_3'(s(bool,X523),s(fun(X524,bool),X525),s(fun(X524,bool),X527))),s(X524,X526)))),'const.bool.COND_3'(s(bool,X523),s(bool,happ(s(fun(X524,bool),X525),s(X524,X526))),s(bool,happ(s(fun(X524,bool),X527),s(X524,X526)))))) )),
  inference(superposition,[],[f991,f480])).
fof(f18863,plain,(
  ( ! [X509,X505,X507,X508,X510,X506] : (sP4('T_0','const.bool.COND_3'(s(bool,X505),s(bool,happ(s(fun(X506,bool),X507),s(X506,X508))),s(bool,happ(s(fun(X506,bool),X509),s(X506,X508)))),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X510),s(bool,happ(s(fun(X506,bool),'const.bool.COND_3'(s(bool,X505),s(fun(X506,bool),X507),s(fun(X506,bool),X509))),s(X506,X508)))))) )),
  inference(superposition,[],[f930,f480])).
fof(f18861,plain,(
  ( ! [X494,X496,X498,X493,X495,X497] : (sP4('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,happ(s(fun(X494,bool),'const.bool.COND_3'(s(bool,X493),s(fun(X494,bool),X495),s(fun(X494,bool),X497))),s(X494,X496))),s(bool,X498)),'const.bool.COND_3'(s(bool,X493),s(bool,happ(s(fun(X494,bool),X495),s(X494,X496))),s(bool,happ(s(fun(X494,bool),X497),s(X494,X496)))))) )),
  inference(superposition,[],[f898,f480])).
fof(f18856,plain,(
  ( ! [X461,X463,X465,X460,X462,X464] : (sP5('T_0','const.bool.COND_3'(s(bool,X460),s(bool,happ(s(fun(X461,bool),X462),s(X461,X463))),s(bool,happ(s(fun(X461,bool),X464),s(X461,X463)))),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X465),s(bool,happ(s(fun(X461,bool),'const.bool.COND_3'(s(bool,X460),s(fun(X461,bool),X462),s(fun(X461,bool),X464))),s(X461,X463)))))) )),
  inference(superposition,[],[f839,f480])).
fof(f18854,plain,(
  ( ! [X453,X449,X451,X452,X448,X450] : (~sP5('F_0','const.bool.COND_3'(s(bool,X448),s(bool,happ(s(fun(X449,bool),X450),s(X449,X451))),s(bool,happ(s(fun(X449,bool),X452),s(X449,X451)))),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X453),s(bool,happ(s(fun(X449,bool),'const.bool.COND_3'(s(bool,X448),s(fun(X449,bool),X450),s(fun(X449,bool),X452))),s(X449,X451)))))) )),
  inference(superposition,[],[f838,f480])).
fof(f18853,plain,(
  ( ! [X445,X447,X443,X444,X446,X442] : (sP5('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X447),s(bool,happ(s(fun(X443,bool),'const.bool.COND_3'(s(bool,X442),s(fun(X443,bool),X444),s(fun(X443,bool),X446))),s(X443,X445)))),'const.bool.COND_3'(s(bool,X442),s(bool,happ(s(fun(X443,bool),X444),s(X443,X445))),s(bool,happ(s(fun(X443,bool),X446),s(X443,X445)))))) )),
  inference(superposition,[],[f837,f480])).
fof(f18851,plain,(
  ( ! [X430,X432,X434,X431,X433,X435] : (~sP5('F_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X435),s(bool,happ(s(fun(X431,bool),'const.bool.COND_3'(s(bool,X430),s(fun(X431,bool),X432),s(fun(X431,bool),X434))),s(X431,X433)))),'const.bool.COND_3'(s(bool,X430),s(bool,happ(s(fun(X431,bool),X432),s(X431,X433))),s(bool,happ(s(fun(X431,bool),X434),s(X431,X433)))))) )),
  inference(superposition,[],[f836,f480])).
fof(f18848,plain,(
  ( ! [X412,X414,X416,X413,X415,X417] : (sP4('T_0','const.bool.COND_3'(s(bool,X412),s(bool,happ(s(fun(X413,bool),X414),s(X413,X415))),s(bool,happ(s(fun(X413,bool),X416),s(X413,X415)))),'const.bool.COND_3'(s(bool,'T_0'),s(bool,happ(s(fun(X413,bool),'const.bool.COND_3'(s(bool,X412),s(fun(X413,bool),X414),s(fun(X413,bool),X416))),s(X413,X415))),s(bool,X417)))) )),
  inference(superposition,[],[f825,f480])).
fof(f18842,plain,(
  ( ! [X377,X372,X374,X376,X373,X375] : (sP5('T_0','const.bool.COND_3'(s(bool,X372),s(bool,happ(s(fun(X373,bool),X374),s(X373,X375))),s(bool,happ(s(fun(X373,bool),X376),s(X373,X375)))),'const.bool.COND_3'(s(bool,'T_0'),s(bool,happ(s(fun(X373,bool),'const.bool.COND_3'(s(bool,X372),s(fun(X373,bool),X374),s(fun(X373,bool),X376))),s(X373,X375))),s(bool,X377)))) )),
  inference(superposition,[],[f778,f480])).
fof(f18840,plain,(
  ( ! [X364,X360,X362,X365,X361,X363] : (~sP5('F_0','const.bool.COND_3'(s(bool,X360),s(bool,happ(s(fun(X361,bool),X362),s(X361,X363))),s(bool,happ(s(fun(X361,bool),X364),s(X361,X363)))),'const.bool.COND_3'(s(bool,'T_0'),s(bool,happ(s(fun(X361,bool),'const.bool.COND_3'(s(bool,X360),s(fun(X361,bool),X362),s(fun(X361,bool),X364))),s(X361,X363))),s(bool,X365)))) )),
  inference(superposition,[],[f777,f480])).
fof(f18838,plain,(
  ( ! [X350,X352,X349,X351,X353,X348] : (sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,happ(s(fun(X349,bool),'const.bool.COND_3'(s(bool,X348),s(fun(X349,bool),X350),s(fun(X349,bool),X352))),s(X349,X351))),s(bool,X353)),'const.bool.COND_3'(s(bool,X348),s(bool,happ(s(fun(X349,bool),X350),s(X349,X351))),s(bool,happ(s(fun(X349,bool),X352),s(X349,X351)))))) )),
  inference(superposition,[],[f776,f480])).
fof(f18836,plain,(
  ( ! [X341,X337,X339,X340,X336,X338] : (~sP5('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,happ(s(fun(X337,bool),'const.bool.COND_3'(s(bool,X336),s(fun(X337,bool),X338),s(fun(X337,bool),X340))),s(X337,X339))),s(bool,X341)),'const.bool.COND_3'(s(bool,X336),s(bool,happ(s(fun(X337,bool),X338),s(X337,X339))),s(bool,happ(s(fun(X337,bool),X340),s(X337,X339)))))) )),
  inference(superposition,[],[f775,f480])).
fof(f18834,plain,(
  ( ! [X325,X327,X323,X329,X324,X326,X328] : (~p(s(bool,happ(s(fun(X324,bool),'const.bool.COND_3'(s(bool,X323),s(fun(X324,bool),X325),s(fun(X324,bool),X327))),s(X324,X326)))) | ~sP8('const.bool.COND_3'(s(bool,X323),s(bool,happ(s(fun(X324,bool),X325),s(X324,X326))),s(bool,happ(s(fun(X324,bool),X327),s(X324,X326)))),X328,X329) | sP9('const.bool.COND_3'(s(bool,X323),s(bool,happ(s(fun(X324,bool),X325),s(X324,X326))),s(bool,happ(s(fun(X324,bool),X327),s(X324,X326)))),X328,X329)) )),
  inference(superposition,[],[f577,f480])).
fof(f18833,plain,(
  ( ! [X321,X316,X318,X320,X322,X317,X319] : (~p(s(bool,happ(s(fun(X317,bool),'const.bool.COND_3'(s(bool,X316),s(fun(X317,bool),X318),s(fun(X317,bool),X320))),s(X317,X319)))) | ~sP10('const.bool.COND_3'(s(bool,X316),s(bool,happ(s(fun(X317,bool),X318),s(X317,X319))),s(bool,happ(s(fun(X317,bool),X320),s(X317,X319)))),X321,X322) | sP11('const.bool.COND_3'(s(bool,X316),s(bool,happ(s(fun(X317,bool),X318),s(X317,X319))),s(bool,happ(s(fun(X317,bool),X320),s(X317,X319)))),X321,X322)) )),
  inference(superposition,[],[f576,f480])).
fof(f18832,plain,(
  ( ! [X313,X315,X310,X312,X314,X309,X311] : (~p(s(bool,happ(s(fun(X310,bool),'const.bool.COND_3'(s(bool,X309),s(fun(X310,bool),X311),s(fun(X310,bool),X313))),s(X310,X312)))) | ~p(s(bool,X314)) | sP11('const.bool.COND_3'(s(bool,X309),s(bool,happ(s(fun(X310,bool),X311),s(X310,X312))),s(bool,happ(s(fun(X310,bool),X313),s(X310,X312)))),X314,X315)) )),
  inference(superposition,[],[f572,f480])).
fof(f18831,plain,(
  ( ! [X302,X308,X304,X306,X303,X305,X307] : (~p(s(bool,happ(s(fun(X303,bool),'const.bool.COND_3'(s(bool,X302),s(fun(X303,bool),X304),s(fun(X303,bool),X306))),s(X303,X305)))) | p(s(bool,X307)) | sP11('const.bool.COND_3'(s(bool,X302),s(bool,happ(s(fun(X303,bool),X304),s(X303,X305))),s(bool,happ(s(fun(X303,bool),X306),s(X303,X305)))),X308,X307)) )),
  inference(superposition,[],[f571,f480])).
fof(f18830,plain,(
  ( ! [X300,X296,X298,X295,X301,X297,X299] : (~p(s(bool,happ(s(fun(X296,bool),'const.bool.COND_3'(s(bool,X295),s(fun(X296,bool),X297),s(fun(X296,bool),X299))),s(X296,X298)))) | ~p(s(bool,X300)) | sP9('const.bool.COND_3'(s(bool,X295),s(bool,happ(s(fun(X296,bool),X297),s(X296,X298))),s(bool,happ(s(fun(X296,bool),X299),s(X296,X298)))),X300,X301)) )),
  inference(superposition,[],[f555,f480])).
fof(f18829,plain,(
  ( ! [X292,X294,X288,X290,X293,X289,X291] : (~p(s(bool,happ(s(fun(X289,bool),'const.bool.COND_3'(s(bool,X288),s(fun(X289,bool),X290),s(fun(X289,bool),X292))),s(X289,X291)))) | ~p(s(bool,X293)) | sP9('const.bool.COND_3'(s(bool,X288),s(bool,happ(s(fun(X289,bool),X290),s(X289,X291))),s(bool,happ(s(fun(X289,bool),X292),s(X289,X291)))),X294,X293)) )),
  inference(superposition,[],[f554,f480])).
fof(f18828,plain,(
  ( ! [X282,X285,X287,X281,X283,X284,X286] : (~p(s(bool,happ(s(fun(X282,bool),'const.bool.COND_3'(s(bool,X281),s(fun(X282,bool),X283),s(fun(X282,bool),X285))),s(X282,X284)))) | ~p(s(bool,X286)) | p(s(bool,X287)) | ~sP7(X287,X286,'const.bool.COND_3'(s(bool,X281),s(bool,happ(s(fun(X282,bool),X283),s(X282,X284))),s(bool,happ(s(fun(X282,bool),X285),s(X282,X284)))))) )),
  inference(superposition,[],[f541,f480])).
fof(f18819,plain,(
  ( ! [X218,X216,X220,X219,X217,X221] : (~p(s(bool,happ(s(fun(X217,bool),'const.bool.COND_3'(s(bool,X216),s(fun(X217,bool),X218),s(fun(X217,bool),X220))),s(X217,X219)))) | p(s(bool,X221)) | sP1(X221,'const.bool.COND_3'(s(bool,X216),s(bool,happ(s(fun(X217,bool),X218),s(X217,X219))),s(bool,happ(s(fun(X217,bool),X220),s(X217,X219)))))) )),
  inference(superposition,[],[f462,f480])).
fof(f18818,plain,(
  ( ! [X212,X210,X215,X213,X211,X214] : (~p(s(bool,happ(s(fun(X211,bool),'const.bool.COND_3'(s(bool,X210),s(fun(X211,bool),X212),s(fun(X211,bool),X214))),s(X211,X213)))) | p(s(bool,X215)) | sP1('const.bool.COND_3'(s(bool,X210),s(bool,happ(s(fun(X211,bool),X212),s(X211,X213))),s(bool,happ(s(fun(X211,bool),X214),s(X211,X213)))),X215)) )),
  inference(superposition,[],[f459,f480])).
fof(f18817,plain,(
  ( ! [X208,X206,X204,X209,X207,X205] : (~p(s(bool,happ(s(fun(X205,bool),'const.bool.COND_3'(s(bool,X204),s(fun(X205,bool),X206),s(fun(X205,bool),X208))),s(X205,X207)))) | sP0('const.bool.COND_3'(s(bool,X204),s(bool,happ(s(fun(X205,bool),X206),s(X205,X207))),s(bool,happ(s(fun(X205,bool),X208),s(X205,X207)))),X209)) )),
  inference(superposition,[],[f444,f480])).
fof(f18816,plain,(
  ( ! [X198,X202,X200,X199,X203,X201] : (~p(s(bool,happ(s(fun(X199,bool),'const.bool.COND_3'(s(bool,X198),s(fun(X199,bool),X200),s(fun(X199,bool),X202))),s(X199,X201)))) | sP0(X203,'const.bool.COND_3'(s(bool,X198),s(bool,happ(s(fun(X199,bool),X200),s(X199,X201))),s(bool,happ(s(fun(X199,bool),X202),s(X199,X201)))))) )),
  inference(superposition,[],[f443,f480])).
fof(f18802,plain,(
  ( ! [X99,X97,X95,X98,X96] : (~p(s(bool,happ(s(fun(X96,bool),'const.bool.COND_3'(s(bool,X95),s(fun(X96,bool),X97),s(fun(X96,bool),X99))),s(X96,X98)))) | p(s(bool,X95)) | p(s(bool,happ(s(fun(X96,bool),X99),s(X96,X98))))) )),
  inference(superposition,[],[f485,f480])).
fof(f18786,plain,(
  ( ! [X14,X17,X15,X13,X16] : (s(X13,happ(s(fun(X15,X13),X16),s(X15,'const.bool.COND_3'(s(bool,X14),s(X15,X17),s(X15,X17))))) = s(X13,happ(s(fun(X15,X13),'const.bool.COND_3'(s(bool,X14),s(fun(X15,X13),X16),s(fun(X15,X13),X16))),s(X15,X17)))) )),
  inference(superposition,[],[f480,f479])).
fof(f18785,plain,(
  ( ! [X12,X10,X8,X11,X9] : (s(bool,'F_0') = s(bool,happ(s(fun(X9,bool),'const.bool.COND_3'(s(bool,X8),s(fun(X9,bool),X10),s(fun(X9,bool),X12))),s(X9,X11))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X8),s(bool,happ(s(fun(X9,bool),X10),s(X9,X11))),s(bool,happ(s(fun(X9,bool),X12),s(X9,X11)))))) )),
  inference(superposition,[],[f480,f343])).
fof(f18784,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',happ(s(fun(X4,'type.num.num'),X5),s(X4,X6))),s('type.num.num',happ(s(fun(X4,'type.num.num'),X7),s(X4,X6))))) = s('type.num.num',happ(s(fun(X4,'type.num.num'),'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun(X4,'type.num.num'),X5),s(X4,X6))),s('type.num.num',happ(s(fun(X4,'type.num.num'),X7),s(X4,X6))))),s(fun(X4,'type.num.num'),X7),s(fun(X4,'type.num.num'),X5))),s(X4,X6)))) )),
  inference(superposition,[],[f480,f432])).
fof(f18783,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,X2))),s('type.num.num',happ(s(fun(X0,'type.num.num'),X3),s(X0,X2))))) = s('type.num.num',happ(s(fun(X0,'type.num.num'),'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,X2))),s('type.num.num',happ(s(fun(X0,'type.num.num'),X3),s(X0,X2))))),s(fun(X0,'type.num.num'),X1),s(fun(X0,'type.num.num'),X3))),s(X0,X2)))) )),
  inference(superposition,[],[f480,f431])).
fof(f18782,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'const.bool.COND_3'(s(bool,X3),s(bool,happ(s(fun(X0,bool),X4),s(X0,X2))),s(bool,'F_0'))) = s(bool,happ(s(fun(X0,bool),'const.bool.COND_3'(s(bool,X3),s(fun(X0,bool),X4),s(fun(X0,bool),X1))),s(X0,X2))) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) )),
  inference(superposition,[],[f480,f343])).
fof(f18761,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'const.bool.COND_3'(s(bool,X3),s(bool,'F_0'),s(bool,happ(s(fun(X0,bool),X4),s(X0,X2))))) = s(bool,happ(s(fun(X0,bool),'const.bool.COND_3'(s(bool,X3),s(fun(X0,bool),X1),s(fun(X0,bool),X4))),s(X0,X2))) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) )),
  inference(superposition,[],[f480,f343])).
fof(f17256,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP5('F_0','const.bool.COND_3'(s(bool,X0),s(bool,happ(s(fun(X1,bool),X2),s(X1,X3))),s(bool,happ(s(fun(X1,bool),X4),s(X1,X3)))),happ(s(fun(X1,bool),'const.bool.COND_3'(s(bool,X0),s(fun(X1,bool),X2),s(fun(X1,bool),X4))),s(X1,X3)))) )),
  inference(unit_resulting_resolution,[],[f592,f480,f524])).
fof(f17257,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP5('T_0','const.bool.COND_3'(s(bool,X0),s(bool,happ(s(fun(X1,bool),X2),s(X1,X3))),s(bool,happ(s(fun(X1,bool),X4),s(X1,X3)))),happ(s(fun(X1,bool),'const.bool.COND_3'(s(bool,X0),s(fun(X1,bool),X2),s(fun(X1,bool),X4))),s(X1,X3)))) )),
  inference(unit_resulting_resolution,[],[f335,f480,f522])).
fof(f18263,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP5('F_0',happ(s(fun(X0,bool),'const.bool.COND_3'(s(bool,X1),s(fun(X0,bool),X2),s(fun(X0,bool),X3))),s(X0,X4)),'const.bool.COND_3'(s(bool,X1),s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))),s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))))) )),
  inference(unit_resulting_resolution,[],[f592,f480,f524])).
fof(f18264,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP5('T_0',happ(s(fun(X0,bool),'const.bool.COND_3'(s(bool,X1),s(fun(X0,bool),X2),s(fun(X0,bool),X3))),s(X0,X4)),'const.bool.COND_3'(s(bool,X1),s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))),s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))))) )),
  inference(unit_resulting_resolution,[],[f335,f480,f522])).
fof(f16728,plain,(
  ( ! [X2,X0,X1] : (sP1('const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,'T_0')),X2) | p(s(bool,X2)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f11665,f343])).
fof(f16717,plain,(
  ( ! [X0,X1] : (sP1('const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,'T_0')),'F_0') | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f11657,f343])).
fof(f16701,plain,(
  ( ! [X1081,X1079,X1077,X1082,X1080,X1078,X1076,X1083] : (p(s(bool,happ(s(fun(X1076,bool),happ(s(fun(X1078,fun(X1076,bool)),X1079),s(X1078,'const.bool.COND_3'(s(bool,X1077),s(X1078,X1080),s(X1078,X1081))))),s(X1076,sK21('const.bool.COND_3'(s(bool,X1077),s(fun(X1076,bool),happ(s(fun(X1078,fun(X1076,bool)),X1079),s(X1078,X1080))),s(fun(X1076,bool),happ(s(fun(X1078,fun(X1076,bool)),X1079),s(X1078,X1081)))),X1083,X1076))))) | ~p(s(bool,happ(s(fun(X1076,bool),happ(s(fun(X1078,fun(X1076,bool)),X1079),s(X1078,'const.bool.COND_3'(s(bool,X1077),s(X1078,X1080),s(X1078,X1081))))),s(X1076,X1082)))) | p(s(bool,happ(s(fun(X1076,bool),X1083),s(X1076,sK21('const.bool.COND_3'(s(bool,X1077),s(fun(X1076,bool),happ(s(fun(X1078,fun(X1076,bool)),X1079),s(X1078,X1080))),s(fun(X1076,bool),happ(s(fun(X1078,fun(X1076,bool)),X1079),s(X1078,X1081)))),X1083,X1076)))))) )),
  inference(forward_demodulation,[],[f13461,f479])).
fof(f13461,plain,(
  ( ! [X1081,X1079,X1077,X1082,X1080,X1078,X1076,X1083] : (~p(s(bool,happ(s(fun(X1076,bool),happ(s(fun(X1078,fun(X1076,bool)),X1079),s(X1078,'const.bool.COND_3'(s(bool,X1077),s(X1078,X1080),s(X1078,X1081))))),s(X1076,X1082)))) | p(s(bool,happ(s(fun(X1076,bool),X1083),s(X1076,sK21('const.bool.COND_3'(s(bool,X1077),s(fun(X1076,bool),happ(s(fun(X1078,fun(X1076,bool)),X1079),s(X1078,X1080))),s(fun(X1076,bool),happ(s(fun(X1078,fun(X1076,bool)),X1079),s(X1078,X1081)))),X1083,X1076))))) | p(s(bool,happ(s(fun(X1076,bool),'const.bool.COND_3'(s(bool,X1077),s(fun(X1076,bool),happ(s(fun(X1078,fun(X1076,bool)),X1079),s(X1078,X1080))),s(fun(X1076,bool),happ(s(fun(X1078,fun(X1076,bool)),X1079),s(X1078,X1081))))),s(X1076,sK21('const.bool.COND_3'(s(bool,X1077),s(fun(X1076,bool),happ(s(fun(X1078,fun(X1076,bool)),X1079),s(X1078,X1080))),s(fun(X1076,bool),happ(s(fun(X1078,fun(X1076,bool)),X1079),s(X1078,X1081)))),X1083,X1076)))))) )),
  inference(superposition,[],[f424,f479])).
fof(f16700,plain,(
  ( ! [X1075,X1073,X1070,X1068,X1074,X1072,X1071,X1069] : (p(s(bool,happ(s(fun(X1068,bool),happ(s(fun(X1070,fun(X1068,bool)),X1071),s(X1070,'const.bool.COND_3'(s(bool,X1069),s(X1070,X1072),s(X1070,X1073))))),s(X1068,sK21(X1075,'const.bool.COND_3'(s(bool,X1069),s(fun(X1068,bool),happ(s(fun(X1070,fun(X1068,bool)),X1071),s(X1070,X1072))),s(fun(X1068,bool),happ(s(fun(X1070,fun(X1068,bool)),X1071),s(X1070,X1073)))),X1068))))) | ~p(s(bool,happ(s(fun(X1068,bool),happ(s(fun(X1070,fun(X1068,bool)),X1071),s(X1070,'const.bool.COND_3'(s(bool,X1069),s(X1070,X1072),s(X1070,X1073))))),s(X1068,X1074)))) | p(s(bool,happ(s(fun(X1068,bool),X1075),s(X1068,sK21(X1075,'const.bool.COND_3'(s(bool,X1069),s(fun(X1068,bool),happ(s(fun(X1070,fun(X1068,bool)),X1071),s(X1070,X1072))),s(fun(X1068,bool),happ(s(fun(X1070,fun(X1068,bool)),X1071),s(X1070,X1073)))),X1068)))))) )),
  inference(forward_demodulation,[],[f13460,f479])).
fof(f13460,plain,(
  ( ! [X1075,X1073,X1070,X1068,X1074,X1072,X1071,X1069] : (~p(s(bool,happ(s(fun(X1068,bool),happ(s(fun(X1070,fun(X1068,bool)),X1071),s(X1070,'const.bool.COND_3'(s(bool,X1069),s(X1070,X1072),s(X1070,X1073))))),s(X1068,X1074)))) | p(s(bool,happ(s(fun(X1068,bool),'const.bool.COND_3'(s(bool,X1069),s(fun(X1068,bool),happ(s(fun(X1070,fun(X1068,bool)),X1071),s(X1070,X1072))),s(fun(X1068,bool),happ(s(fun(X1070,fun(X1068,bool)),X1071),s(X1070,X1073))))),s(X1068,sK21(X1075,'const.bool.COND_3'(s(bool,X1069),s(fun(X1068,bool),happ(s(fun(X1070,fun(X1068,bool)),X1071),s(X1070,X1072))),s(fun(X1068,bool),happ(s(fun(X1070,fun(X1068,bool)),X1071),s(X1070,X1073)))),X1068))))) | p(s(bool,happ(s(fun(X1068,bool),X1075),s(X1068,sK21(X1075,'const.bool.COND_3'(s(bool,X1069),s(fun(X1068,bool),happ(s(fun(X1070,fun(X1068,bool)),X1071),s(X1070,X1072))),s(fun(X1068,bool),happ(s(fun(X1070,fun(X1068,bool)),X1071),s(X1070,X1073)))),X1068)))))) )),
  inference(superposition,[],[f423,f479])).
fof(f16699,plain,(
  ( ! [X1066,X1064,X1062,X1060,X1067,X1065,X1063,X1061] : (p(s(bool,happ(s(fun(X1060,bool),happ(s(fun(X1062,fun(X1060,bool)),X1063),s(X1062,'const.bool.COND_3'(s(bool,X1061),s(X1062,X1064),s(X1062,X1065))))),s(X1060,sK20('const.bool.COND_3'(s(bool,X1061),s(fun(X1060,bool),happ(s(fun(X1062,fun(X1060,bool)),X1063),s(X1062,X1064))),s(fun(X1060,bool),happ(s(fun(X1062,fun(X1060,bool)),X1063),s(X1062,X1065)))),X1067,X1060))))) | ~p(s(bool,happ(s(fun(X1060,bool),happ(s(fun(X1062,fun(X1060,bool)),X1063),s(X1062,'const.bool.COND_3'(s(bool,X1061),s(X1062,X1064),s(X1062,X1065))))),s(X1060,X1066)))) | p(s(bool,happ(s(fun(X1060,bool),X1067),s(X1060,sK19('const.bool.COND_3'(s(bool,X1061),s(fun(X1060,bool),happ(s(fun(X1062,fun(X1060,bool)),X1063),s(X1062,X1064))),s(fun(X1060,bool),happ(s(fun(X1062,fun(X1060,bool)),X1063),s(X1062,X1065)))),X1067,X1060)))))) )),
  inference(forward_demodulation,[],[f13459,f479])).
fof(f13459,plain,(
  ( ! [X1066,X1064,X1062,X1060,X1067,X1065,X1063,X1061] : (~p(s(bool,happ(s(fun(X1060,bool),happ(s(fun(X1062,fun(X1060,bool)),X1063),s(X1062,'const.bool.COND_3'(s(bool,X1061),s(X1062,X1064),s(X1062,X1065))))),s(X1060,X1066)))) | p(s(bool,happ(s(fun(X1060,bool),X1067),s(X1060,sK19('const.bool.COND_3'(s(bool,X1061),s(fun(X1060,bool),happ(s(fun(X1062,fun(X1060,bool)),X1063),s(X1062,X1064))),s(fun(X1060,bool),happ(s(fun(X1062,fun(X1060,bool)),X1063),s(X1062,X1065)))),X1067,X1060))))) | p(s(bool,happ(s(fun(X1060,bool),'const.bool.COND_3'(s(bool,X1061),s(fun(X1060,bool),happ(s(fun(X1062,fun(X1060,bool)),X1063),s(X1062,X1064))),s(fun(X1060,bool),happ(s(fun(X1062,fun(X1060,bool)),X1063),s(X1062,X1065))))),s(X1060,sK20('const.bool.COND_3'(s(bool,X1061),s(fun(X1060,bool),happ(s(fun(X1062,fun(X1060,bool)),X1063),s(X1062,X1064))),s(fun(X1060,bool),happ(s(fun(X1062,fun(X1060,bool)),X1063),s(X1062,X1065)))),X1067,X1060)))))) )),
  inference(superposition,[],[f422,f479])).
fof(f16698,plain,(
  ( ! [X1054,X1052,X1058,X1056,X1055,X1053,X1059,X1057] : (p(s(bool,happ(s(fun(X1052,bool),happ(s(fun(X1054,fun(X1052,bool)),X1055),s(X1054,'const.bool.COND_3'(s(bool,X1053),s(X1054,X1056),s(X1054,X1057))))),s(X1052,sK19(X1059,'const.bool.COND_3'(s(bool,X1053),s(fun(X1052,bool),happ(s(fun(X1054,fun(X1052,bool)),X1055),s(X1054,X1056))),s(fun(X1052,bool),happ(s(fun(X1054,fun(X1052,bool)),X1055),s(X1054,X1057)))),X1052))))) | ~p(s(bool,happ(s(fun(X1052,bool),happ(s(fun(X1054,fun(X1052,bool)),X1055),s(X1054,'const.bool.COND_3'(s(bool,X1053),s(X1054,X1056),s(X1054,X1057))))),s(X1052,X1058)))) | p(s(bool,happ(s(fun(X1052,bool),X1059),s(X1052,sK20(X1059,'const.bool.COND_3'(s(bool,X1053),s(fun(X1052,bool),happ(s(fun(X1054,fun(X1052,bool)),X1055),s(X1054,X1056))),s(fun(X1052,bool),happ(s(fun(X1054,fun(X1052,bool)),X1055),s(X1054,X1057)))),X1052)))))) )),
  inference(forward_demodulation,[],[f13458,f479])).
fof(f13458,plain,(
  ( ! [X1054,X1052,X1058,X1056,X1055,X1053,X1059,X1057] : (~p(s(bool,happ(s(fun(X1052,bool),happ(s(fun(X1054,fun(X1052,bool)),X1055),s(X1054,'const.bool.COND_3'(s(bool,X1053),s(X1054,X1056),s(X1054,X1057))))),s(X1052,X1058)))) | p(s(bool,happ(s(fun(X1052,bool),'const.bool.COND_3'(s(bool,X1053),s(fun(X1052,bool),happ(s(fun(X1054,fun(X1052,bool)),X1055),s(X1054,X1056))),s(fun(X1052,bool),happ(s(fun(X1054,fun(X1052,bool)),X1055),s(X1054,X1057))))),s(X1052,sK19(X1059,'const.bool.COND_3'(s(bool,X1053),s(fun(X1052,bool),happ(s(fun(X1054,fun(X1052,bool)),X1055),s(X1054,X1056))),s(fun(X1052,bool),happ(s(fun(X1054,fun(X1052,bool)),X1055),s(X1054,X1057)))),X1052))))) | p(s(bool,happ(s(fun(X1052,bool),X1059),s(X1052,sK20(X1059,'const.bool.COND_3'(s(bool,X1053),s(fun(X1052,bool),happ(s(fun(X1054,fun(X1052,bool)),X1055),s(X1054,X1056))),s(fun(X1052,bool),happ(s(fun(X1054,fun(X1052,bool)),X1055),s(X1054,X1057)))),X1052)))))) )),
  inference(superposition,[],[f421,f479])).
fof(f16697,plain,(
  ( ! [X1050,X1048,X1046,X1044,X1051,X1049,X1047,X1045] : (p(s(bool,happ(s(fun(X1044,bool),happ(s(fun(X1046,fun(X1044,bool)),X1047),s(X1046,'const.bool.COND_3'(s(bool,X1045),s(X1046,X1048),s(X1046,X1049))))),s(X1044,X1051)))) | ~p(s(bool,happ(s(fun(X1044,bool),happ(s(fun(X1046,fun(X1044,bool)),X1047),s(X1046,'const.bool.COND_3'(s(bool,X1045),s(X1046,X1048),s(X1046,X1049))))),s(X1044,sK18('const.bool.COND_3'(s(bool,X1045),s(fun(X1044,bool),happ(s(fun(X1046,fun(X1044,bool)),X1047),s(X1046,X1048))),s(fun(X1044,bool),happ(s(fun(X1046,fun(X1044,bool)),X1047),s(X1046,X1049)))),X1050,X1044))))) | ~p(s(bool,happ(s(fun(X1044,bool),X1050),s(X1044,sK17('const.bool.COND_3'(s(bool,X1045),s(fun(X1044,bool),happ(s(fun(X1046,fun(X1044,bool)),X1047),s(X1046,X1048))),s(fun(X1044,bool),happ(s(fun(X1046,fun(X1044,bool)),X1047),s(X1046,X1049)))),X1050,X1044)))))) )),
  inference(forward_demodulation,[],[f13457,f479])).
fof(f13457,plain,(
  ( ! [X1050,X1048,X1046,X1044,X1051,X1049,X1047,X1045] : (~p(s(bool,happ(s(fun(X1044,bool),happ(s(fun(X1046,fun(X1044,bool)),X1047),s(X1046,'const.bool.COND_3'(s(bool,X1045),s(X1046,X1048),s(X1046,X1049))))),s(X1044,sK18('const.bool.COND_3'(s(bool,X1045),s(fun(X1044,bool),happ(s(fun(X1046,fun(X1044,bool)),X1047),s(X1046,X1048))),s(fun(X1044,bool),happ(s(fun(X1046,fun(X1044,bool)),X1047),s(X1046,X1049)))),X1050,X1044))))) | p(s(bool,happ(s(fun(X1044,bool),'const.bool.COND_3'(s(bool,X1045),s(fun(X1044,bool),happ(s(fun(X1046,fun(X1044,bool)),X1047),s(X1046,X1048))),s(fun(X1044,bool),happ(s(fun(X1046,fun(X1044,bool)),X1047),s(X1046,X1049))))),s(X1044,X1051)))) | ~p(s(bool,happ(s(fun(X1044,bool),X1050),s(X1044,sK17('const.bool.COND_3'(s(bool,X1045),s(fun(X1044,bool),happ(s(fun(X1046,fun(X1044,bool)),X1047),s(X1046,X1048))),s(fun(X1044,bool),happ(s(fun(X1046,fun(X1044,bool)),X1047),s(X1046,X1049)))),X1050,X1044)))))) )),
  inference(superposition,[],[f420,f479])).
fof(f13456,plain,(
  ( ! [X1042,X1040,X1039,X1037,X1043,X1041,X1038,X1036] : (~p(s(bool,happ(s(fun(X1036,bool),happ(s(fun(X1038,fun(X1036,bool)),X1039),s(X1038,'const.bool.COND_3'(s(bool,X1037),s(X1038,X1040),s(X1038,X1041))))),s(X1036,sK18('const.bool.COND_3'(s(bool,X1037),s(fun(X1036,bool),happ(s(fun(X1038,fun(X1036,bool)),X1039),s(X1038,X1040))),s(fun(X1036,bool),happ(s(fun(X1038,fun(X1036,bool)),X1039),s(X1038,X1041)))),X1042,X1036))))) | p(s(bool,happ(s(fun(X1036,bool),X1042),s(X1036,X1043)))) | ~p(s(bool,happ(s(fun(X1036,bool),X1042),s(X1036,sK17('const.bool.COND_3'(s(bool,X1037),s(fun(X1036,bool),happ(s(fun(X1038,fun(X1036,bool)),X1039),s(X1038,X1040))),s(fun(X1036,bool),happ(s(fun(X1038,fun(X1036,bool)),X1039),s(X1038,X1041)))),X1042,X1036)))))) )),
  inference(superposition,[],[f419,f479])).
fof(f16696,plain,(
  ( ! [X1035,X1033,X1031,X1029,X1034,X1032,X1030,X1028] : (p(s(bool,happ(s(fun(X1028,bool),happ(s(fun(X1030,fun(X1028,bool)),X1031),s(X1030,'const.bool.COND_3'(s(bool,X1029),s(X1030,X1032),s(X1030,X1033))))),s(X1028,X1035)))) | ~p(s(bool,happ(s(fun(X1028,bool),happ(s(fun(X1030,fun(X1028,bool)),X1031),s(X1030,'const.bool.COND_3'(s(bool,X1029),s(X1030,X1032),s(X1030,X1033))))),s(X1028,sK16('const.bool.COND_3'(s(bool,X1029),s(fun(X1028,bool),happ(s(fun(X1030,fun(X1028,bool)),X1031),s(X1030,X1032))),s(fun(X1028,bool),happ(s(fun(X1030,fun(X1028,bool)),X1031),s(X1030,X1033)))),X1034,X1028))))) | ~p(s(bool,happ(s(fun(X1028,bool),X1034),s(X1028,sK16('const.bool.COND_3'(s(bool,X1029),s(fun(X1028,bool),happ(s(fun(X1030,fun(X1028,bool)),X1031),s(X1030,X1032))),s(fun(X1028,bool),happ(s(fun(X1030,fun(X1028,bool)),X1031),s(X1030,X1033)))),X1034,X1028)))))) )),
  inference(forward_demodulation,[],[f13455,f479])).
fof(f13455,plain,(
  ( ! [X1035,X1033,X1031,X1029,X1034,X1032,X1030,X1028] : (~p(s(bool,happ(s(fun(X1028,bool),happ(s(fun(X1030,fun(X1028,bool)),X1031),s(X1030,'const.bool.COND_3'(s(bool,X1029),s(X1030,X1032),s(X1030,X1033))))),s(X1028,sK16('const.bool.COND_3'(s(bool,X1029),s(fun(X1028,bool),happ(s(fun(X1030,fun(X1028,bool)),X1031),s(X1030,X1032))),s(fun(X1028,bool),happ(s(fun(X1030,fun(X1028,bool)),X1031),s(X1030,X1033)))),X1034,X1028))))) | p(s(bool,happ(s(fun(X1028,bool),'const.bool.COND_3'(s(bool,X1029),s(fun(X1028,bool),happ(s(fun(X1030,fun(X1028,bool)),X1031),s(X1030,X1032))),s(fun(X1028,bool),happ(s(fun(X1030,fun(X1028,bool)),X1031),s(X1030,X1033))))),s(X1028,X1035)))) | ~p(s(bool,happ(s(fun(X1028,bool),X1034),s(X1028,sK16('const.bool.COND_3'(s(bool,X1029),s(fun(X1028,bool),happ(s(fun(X1030,fun(X1028,bool)),X1031),s(X1030,X1032))),s(fun(X1028,bool),happ(s(fun(X1030,fun(X1028,bool)),X1031),s(X1030,X1033)))),X1034,X1028)))))) )),
  inference(superposition,[],[f418,f479])).
fof(f13454,plain,(
  ( ! [X1023,X1021,X1027,X1025,X1022,X1020,X1026,X1024] : (~p(s(bool,happ(s(fun(X1020,bool),happ(s(fun(X1022,fun(X1020,bool)),X1023),s(X1022,'const.bool.COND_3'(s(bool,X1021),s(X1022,X1024),s(X1022,X1025))))),s(X1020,sK16('const.bool.COND_3'(s(bool,X1021),s(fun(X1020,bool),happ(s(fun(X1022,fun(X1020,bool)),X1023),s(X1022,X1024))),s(fun(X1020,bool),happ(s(fun(X1022,fun(X1020,bool)),X1023),s(X1022,X1025)))),X1026,X1020))))) | p(s(bool,happ(s(fun(X1020,bool),X1026),s(X1020,X1027)))) | ~p(s(bool,happ(s(fun(X1020,bool),X1026),s(X1020,sK16('const.bool.COND_3'(s(bool,X1021),s(fun(X1020,bool),happ(s(fun(X1022,fun(X1020,bool)),X1023),s(X1022,X1024))),s(fun(X1020,bool),happ(s(fun(X1022,fun(X1020,bool)),X1023),s(X1022,X1025)))),X1026,X1020)))))) )),
  inference(superposition,[],[f417,f479])).
fof(f16695,plain,(
  ( ! [X1019,X1017,X1015,X1013,X1018,X1016,X1014] : (p(s(bool,happ(s(fun(X1013,bool),happ(s(fun(X1015,fun(X1013,bool)),X1016),s(X1015,'const.bool.COND_3'(s(bool,X1014),s(X1015,X1017),s(X1015,X1018))))),s(X1013,X1019)))) | ~p(s(bool,happ(s(fun(X1013,bool),happ(s(fun(X1015,fun(X1013,bool)),X1016),s(X1015,'const.bool.COND_3'(s(bool,X1014),s(X1015,X1017),s(X1015,X1018))))),s(X1013,sK15('const.bool.COND_3'(s(bool,X1014),s(fun(X1013,bool),happ(s(fun(X1015,fun(X1013,bool)),X1016),s(X1015,X1017))),s(fun(X1013,bool),happ(s(fun(X1015,fun(X1013,bool)),X1016),s(X1015,X1018)))),X1013)))))) )),
  inference(forward_demodulation,[],[f13453,f479])).
fof(f13453,plain,(
  ( ! [X1019,X1017,X1015,X1013,X1018,X1016,X1014] : (~p(s(bool,happ(s(fun(X1013,bool),happ(s(fun(X1015,fun(X1013,bool)),X1016),s(X1015,'const.bool.COND_3'(s(bool,X1014),s(X1015,X1017),s(X1015,X1018))))),s(X1013,sK15('const.bool.COND_3'(s(bool,X1014),s(fun(X1013,bool),happ(s(fun(X1015,fun(X1013,bool)),X1016),s(X1015,X1017))),s(fun(X1013,bool),happ(s(fun(X1015,fun(X1013,bool)),X1016),s(X1015,X1018)))),X1013))))) | p(s(bool,happ(s(fun(X1013,bool),'const.bool.COND_3'(s(bool,X1014),s(fun(X1013,bool),happ(s(fun(X1015,fun(X1013,bool)),X1016),s(X1015,X1017))),s(fun(X1013,bool),happ(s(fun(X1015,fun(X1013,bool)),X1016),s(X1015,X1018))))),s(X1013,X1019))))) )),
  inference(superposition,[],[f412,f479])).
fof(f16694,plain,(
  ( ! [X1011,X1009,X1006,X1010,X1008,X1012,X1007] : (p(s(bool,happ(s(fun(X1006,bool),happ(s(fun(X1008,fun(X1006,bool)),X1009),s(X1008,'const.bool.COND_3'(s(bool,X1007),s(X1008,X1010),s(X1008,X1011))))),s(X1006,X1012)))) | ~p(s(bool,happ(s(fun(X1006,bool),happ(s(fun(X1008,fun(X1006,bool)),X1009),s(X1008,'const.bool.COND_3'(s(bool,X1007),s(X1008,X1010),s(X1008,X1011))))),s(X1006,sK14('const.bool.COND_3'(s(bool,X1007),s(fun(X1006,bool),happ(s(fun(X1008,fun(X1006,bool)),X1009),s(X1008,X1010))),s(fun(X1006,bool),happ(s(fun(X1008,fun(X1006,bool)),X1009),s(X1008,X1011)))),X1006)))))) )),
  inference(forward_demodulation,[],[f13452,f479])).
fof(f13452,plain,(
  ( ! [X1011,X1009,X1006,X1010,X1008,X1012,X1007] : (~p(s(bool,happ(s(fun(X1006,bool),happ(s(fun(X1008,fun(X1006,bool)),X1009),s(X1008,'const.bool.COND_3'(s(bool,X1007),s(X1008,X1010),s(X1008,X1011))))),s(X1006,sK14('const.bool.COND_3'(s(bool,X1007),s(fun(X1006,bool),happ(s(fun(X1008,fun(X1006,bool)),X1009),s(X1008,X1010))),s(fun(X1006,bool),happ(s(fun(X1008,fun(X1006,bool)),X1009),s(X1008,X1011)))),X1006))))) | p(s(bool,happ(s(fun(X1006,bool),'const.bool.COND_3'(s(bool,X1007),s(fun(X1006,bool),happ(s(fun(X1008,fun(X1006,bool)),X1009),s(X1008,X1010))),s(fun(X1006,bool),happ(s(fun(X1008,fun(X1006,bool)),X1009),s(X1008,X1011))))),s(X1006,X1012))))) )),
  inference(superposition,[],[f411,f479])).
fof(f13361,plain,(
  ( ! [X492,X488,X490,X493,X489,X491] : (sP4('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X493),s(bool,happ(s(fun(X489,bool),X490),s(X489,'const.bool.COND_3'(s(bool,X488),s(X489,X491),s(X489,X492)))))),'const.bool.COND_3'(s(bool,X488),s(bool,happ(s(fun(X489,bool),X490),s(X489,X491))),s(bool,happ(s(fun(X489,bool),X490),s(X489,X492)))))) )),
  inference(superposition,[],[f991,f479])).
fof(f13358,plain,(
  ( ! [X472,X474,X471,X475,X473,X470] : (sP4('T_0','const.bool.COND_3'(s(bool,X470),s(bool,happ(s(fun(X471,bool),X472),s(X471,X473))),s(bool,happ(s(fun(X471,bool),X472),s(X471,X474)))),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X475),s(bool,happ(s(fun(X471,bool),X472),s(X471,'const.bool.COND_3'(s(bool,X470),s(X471,X473),s(X471,X474)))))))) )),
  inference(superposition,[],[f930,f479])).
fof(f13356,plain,(
  ( ! [X461,X463,X459,X460,X462,X458] : (sP4('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,happ(s(fun(X459,bool),X460),s(X459,'const.bool.COND_3'(s(bool,X458),s(X459,X461),s(X459,X462))))),s(bool,X463)),'const.bool.COND_3'(s(bool,X458),s(bool,happ(s(fun(X459,bool),X460),s(X459,X461))),s(bool,happ(s(fun(X459,bool),X460),s(X459,X462)))))) )),
  inference(superposition,[],[f898,f479])).
fof(f13351,plain,(
  ( ! [X428,X430,X426,X429,X425,X427] : (sP5('T_0','const.bool.COND_3'(s(bool,X425),s(bool,happ(s(fun(X426,bool),X427),s(X426,X428))),s(bool,happ(s(fun(X426,bool),X427),s(X426,X429)))),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X430),s(bool,happ(s(fun(X426,bool),X427),s(X426,'const.bool.COND_3'(s(bool,X425),s(X426,X428),s(X426,X429)))))))) )),
  inference(superposition,[],[f839,f479])).
fof(f13349,plain,(
  ( ! [X414,X416,X418,X413,X415,X417] : (~sP5('F_0','const.bool.COND_3'(s(bool,X413),s(bool,happ(s(fun(X414,bool),X415),s(X414,X416))),s(bool,happ(s(fun(X414,bool),X415),s(X414,X417)))),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X418),s(bool,happ(s(fun(X414,bool),X415),s(X414,'const.bool.COND_3'(s(bool,X413),s(X414,X416),s(X414,X417)))))))) )),
  inference(superposition,[],[f838,f479])).
fof(f13348,plain,(
  ( ! [X412,X408,X410,X407,X409,X411] : (sP5('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X412),s(bool,happ(s(fun(X408,bool),X409),s(X408,'const.bool.COND_3'(s(bool,X407),s(X408,X410),s(X408,X411)))))),'const.bool.COND_3'(s(bool,X407),s(bool,happ(s(fun(X408,bool),X409),s(X408,X410))),s(bool,happ(s(fun(X408,bool),X409),s(X408,X411)))))) )),
  inference(superposition,[],[f837,f479])).
fof(f13346,plain,(
  ( ! [X397,X399,X395,X396,X398,X400] : (~sP5('F_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X400),s(bool,happ(s(fun(X396,bool),X397),s(X396,'const.bool.COND_3'(s(bool,X395),s(X396,X398),s(X396,X399)))))),'const.bool.COND_3'(s(bool,X395),s(bool,happ(s(fun(X396,bool),X397),s(X396,X398))),s(bool,happ(s(fun(X396,bool),X397),s(X396,X399)))))) )),
  inference(superposition,[],[f836,f479])).
fof(f13343,plain,(
  ( ! [X377,X379,X380,X382,X378,X381] : (sP4('T_0','const.bool.COND_3'(s(bool,X377),s(bool,happ(s(fun(X378,bool),X379),s(X378,X380))),s(bool,happ(s(fun(X378,bool),X379),s(X378,X381)))),'const.bool.COND_3'(s(bool,'T_0'),s(bool,happ(s(fun(X378,bool),X379),s(X378,'const.bool.COND_3'(s(bool,X377),s(X378,X380),s(X378,X381))))),s(bool,X382)))) )),
  inference(superposition,[],[f825,f479])).
fof(f13337,plain,(
  ( ! [X341,X337,X339,X340,X342,X338] : (sP5('T_0','const.bool.COND_3'(s(bool,X337),s(bool,happ(s(fun(X338,bool),X339),s(X338,X340))),s(bool,happ(s(fun(X338,bool),X339),s(X338,X341)))),'const.bool.COND_3'(s(bool,'T_0'),s(bool,happ(s(fun(X338,bool),X339),s(X338,'const.bool.COND_3'(s(bool,X337),s(X338,X340),s(X338,X341))))),s(bool,X342)))) )),
  inference(superposition,[],[f778,f479])).
fof(f13335,plain,(
  ( ! [X325,X327,X329,X326,X328,X330] : (~sP5('F_0','const.bool.COND_3'(s(bool,X325),s(bool,happ(s(fun(X326,bool),X327),s(X326,X328))),s(bool,happ(s(fun(X326,bool),X327),s(X326,X329)))),'const.bool.COND_3'(s(bool,'T_0'),s(bool,happ(s(fun(X326,bool),X327),s(X326,'const.bool.COND_3'(s(bool,X325),s(X326,X328),s(X326,X329))))),s(bool,X330)))) )),
  inference(superposition,[],[f777,f479])).
fof(f13333,plain,(
  ( ! [X313,X315,X316,X318,X314,X317] : (sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,happ(s(fun(X314,bool),X315),s(X314,'const.bool.COND_3'(s(bool,X313),s(X314,X316),s(X314,X317))))),s(bool,X318)),'const.bool.COND_3'(s(bool,X313),s(bool,happ(s(fun(X314,bool),X315),s(X314,X316))),s(bool,happ(s(fun(X314,bool),X315),s(X314,X317)))))) )),
  inference(superposition,[],[f776,f479])).
fof(f13331,plain,(
  ( ! [X302,X304,X306,X301,X303,X305] : (~sP5('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,happ(s(fun(X302,bool),X303),s(X302,'const.bool.COND_3'(s(bool,X301),s(X302,X304),s(X302,X305))))),s(bool,X306)),'const.bool.COND_3'(s(bool,X301),s(bool,happ(s(fun(X302,bool),X303),s(X302,X304))),s(bool,happ(s(fun(X302,bool),X303),s(X302,X305)))))) )),
  inference(superposition,[],[f775,f479])).
fof(f13329,plain,(
  ( ! [X292,X294,X288,X290,X293,X289,X291] : (~p(s(bool,happ(s(fun(X289,bool),X290),s(X289,'const.bool.COND_3'(s(bool,X288),s(X289,X291),s(X289,X292)))))) | ~sP8('const.bool.COND_3'(s(bool,X288),s(bool,happ(s(fun(X289,bool),X290),s(X289,X291))),s(bool,happ(s(fun(X289,bool),X290),s(X289,X292)))),X293,X294) | sP9('const.bool.COND_3'(s(bool,X288),s(bool,happ(s(fun(X289,bool),X290),s(X289,X291))),s(bool,happ(s(fun(X289,bool),X290),s(X289,X292)))),X293,X294)) )),
  inference(superposition,[],[f577,f479])).
fof(f13328,plain,(
  ( ! [X282,X285,X287,X281,X283,X284,X286] : (~p(s(bool,happ(s(fun(X282,bool),X283),s(X282,'const.bool.COND_3'(s(bool,X281),s(X282,X284),s(X282,X285)))))) | ~sP10('const.bool.COND_3'(s(bool,X281),s(bool,happ(s(fun(X282,bool),X283),s(X282,X284))),s(bool,happ(s(fun(X282,bool),X283),s(X282,X285)))),X286,X287) | sP11('const.bool.COND_3'(s(bool,X281),s(bool,happ(s(fun(X282,bool),X283),s(X282,X284))),s(bool,happ(s(fun(X282,bool),X283),s(X282,X285)))),X286,X287)) )),
  inference(superposition,[],[f576,f479])).
fof(f13327,plain,(
  ( ! [X280,X277,X279,X275,X276,X278,X274] : (~p(s(bool,happ(s(fun(X275,bool),X276),s(X275,'const.bool.COND_3'(s(bool,X274),s(X275,X277),s(X275,X278)))))) | ~p(s(bool,X279)) | sP11('const.bool.COND_3'(s(bool,X274),s(bool,happ(s(fun(X275,bool),X276),s(X275,X277))),s(bool,happ(s(fun(X275,bool),X276),s(X275,X278)))),X279,X280)) )),
  inference(superposition,[],[f572,f479])).
fof(f13326,plain,(
  ( ! [X269,X271,X267,X273,X268,X270,X272] : (~p(s(bool,happ(s(fun(X268,bool),X269),s(X268,'const.bool.COND_3'(s(bool,X267),s(X268,X270),s(X268,X271)))))) | p(s(bool,X272)) | sP11('const.bool.COND_3'(s(bool,X267),s(bool,happ(s(fun(X268,bool),X269),s(X268,X270))),s(bool,happ(s(fun(X268,bool),X269),s(X268,X271)))),X273,X272)) )),
  inference(superposition,[],[f571,f479])).
fof(f13325,plain,(
  ( ! [X261,X263,X265,X260,X262,X264,X266] : (~p(s(bool,happ(s(fun(X261,bool),X262),s(X261,'const.bool.COND_3'(s(bool,X260),s(X261,X263),s(X261,X264)))))) | ~p(s(bool,X265)) | sP9('const.bool.COND_3'(s(bool,X260),s(bool,happ(s(fun(X261,bool),X262),s(X261,X263))),s(bool,happ(s(fun(X261,bool),X262),s(X261,X264)))),X265,X266)) )),
  inference(superposition,[],[f555,f479])).
fof(f13324,plain,(
  ( ! [X255,X253,X257,X259,X254,X256,X258] : (~p(s(bool,happ(s(fun(X254,bool),X255),s(X254,'const.bool.COND_3'(s(bool,X253),s(X254,X256),s(X254,X257)))))) | ~p(s(bool,X258)) | sP9('const.bool.COND_3'(s(bool,X253),s(bool,happ(s(fun(X254,bool),X255),s(X254,X256))),s(bool,happ(s(fun(X254,bool),X255),s(X254,X257)))),X259,X258)) )),
  inference(superposition,[],[f554,f479])).
fof(f13323,plain,(
  ( ! [X251,X249,X246,X252,X250,X248,X247] : (~p(s(bool,happ(s(fun(X247,bool),X248),s(X247,'const.bool.COND_3'(s(bool,X246),s(X247,X249),s(X247,X250)))))) | ~p(s(bool,X251)) | p(s(bool,X252)) | ~sP7(X252,X251,'const.bool.COND_3'(s(bool,X246),s(bool,happ(s(fun(X247,bool),X248),s(X247,X249))),s(bool,happ(s(fun(X247,bool),X248),s(X247,X250)))))) )),
  inference(superposition,[],[f541,f479])).
fof(f13315,plain,(
  ( ! [X191,X196,X194,X192,X195,X193] : (~p(s(bool,happ(s(fun(X192,bool),X193),s(X192,'const.bool.COND_3'(s(bool,X191),s(X192,X194),s(X192,X195)))))) | p(s(bool,X196)) | sP1(X196,'const.bool.COND_3'(s(bool,X191),s(bool,happ(s(fun(X192,bool),X193),s(X192,X194))),s(bool,happ(s(fun(X192,bool),X193),s(X192,X195)))))) )),
  inference(superposition,[],[f462,f479])).
fof(f13314,plain,(
  ( ! [X189,X187,X185,X190,X188,X186] : (~p(s(bool,happ(s(fun(X186,bool),X187),s(X186,'const.bool.COND_3'(s(bool,X185),s(X186,X188),s(X186,X189)))))) | p(s(bool,X190)) | sP1('const.bool.COND_3'(s(bool,X185),s(bool,happ(s(fun(X186,bool),X187),s(X186,X188))),s(bool,happ(s(fun(X186,bool),X187),s(X186,X189)))),X190)) )),
  inference(superposition,[],[f459,f479])).
fof(f13313,plain,(
  ( ! [X182,X180,X184,X183,X181,X179] : (~p(s(bool,happ(s(fun(X180,bool),X181),s(X180,'const.bool.COND_3'(s(bool,X179),s(X180,X182),s(X180,X183)))))) | sP0('const.bool.COND_3'(s(bool,X179),s(bool,happ(s(fun(X180,bool),X181),s(X180,X182))),s(bool,happ(s(fun(X180,bool),X181),s(X180,X183)))),X184)) )),
  inference(superposition,[],[f444,f479])).
fof(f13312,plain,(
  ( ! [X177,X175,X173,X178,X176,X174] : (~p(s(bool,happ(s(fun(X174,bool),X175),s(X174,'const.bool.COND_3'(s(bool,X173),s(X174,X176),s(X174,X177)))))) | sP0(X178,'const.bool.COND_3'(s(bool,X173),s(bool,happ(s(fun(X174,bool),X175),s(X174,X176))),s(bool,happ(s(fun(X174,bool),X175),s(X174,X177)))))) )),
  inference(superposition,[],[f443,f479])).
fof(f13298,plain,(
  ( ! [X70,X74,X72,X71,X73] : (~p(s(bool,happ(s(fun(X71,bool),X72),s(X71,'const.bool.COND_3'(s(bool,X70),s(X71,X73),s(X71,X74)))))) | p(s(bool,X70)) | p(s(bool,happ(s(fun(X71,bool),X72),s(X71,X74))))) )),
  inference(superposition,[],[f485,f479])).
fof(f13285,plain,(
  ( ! [X12,X10,X8,X11,X9] : (s(bool,'F_0') = s(bool,happ(s(fun(X9,bool),X10),s(X9,'const.bool.COND_3'(s(bool,X8),s(X9,X11),s(X9,X12))))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X8),s(bool,happ(s(fun(X9,bool),X10),s(X9,X11))),s(bool,happ(s(fun(X9,bool),X10),s(X9,X12)))))) )),
  inference(superposition,[],[f479,f343])).
fof(f13284,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',happ(s(fun(X4,'type.num.num'),X5),s(X4,X6))),s('type.num.num',happ(s(fun(X4,'type.num.num'),X5),s(X4,X7))))) = s('type.num.num',happ(s(fun(X4,'type.num.num'),X5),s(X4,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun(X4,'type.num.num'),X5),s(X4,X6))),s('type.num.num',happ(s(fun(X4,'type.num.num'),X5),s(X4,X7))))),s(X4,X7),s(X4,X6)))))) )),
  inference(superposition,[],[f479,f432])).
fof(f13283,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,X2))),s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,X3))))) = s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,X2))),s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,X3))))),s(X0,X2),s(X0,X3)))))) )),
  inference(superposition,[],[f479,f431])).
fof(f13282,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'const.bool.COND_3'(s(bool,X3),s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))),s(bool,'F_0'))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.bool.COND_3'(s(bool,X3),s(X0,X4),s(X0,X2))))) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) )),
  inference(superposition,[],[f479,f343])).
fof(f13263,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'const.bool.COND_3'(s(bool,X3),s(bool,'F_0'),s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.bool.COND_3'(s(bool,X3),s(X0,X2),s(X0,X4))))) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) )),
  inference(superposition,[],[f479,f343])).
fof(f12096,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP5('F_0','const.bool.COND_3'(s(bool,X0),s(bool,happ(s(fun(X1,bool),X2),s(X1,X3))),s(bool,happ(s(fun(X1,bool),X2),s(X1,X4)))),happ(s(fun(X1,bool),X2),s(X1,'const.bool.COND_3'(s(bool,X0),s(X1,X3),s(X1,X4)))))) )),
  inference(unit_resulting_resolution,[],[f592,f479,f524])).
fof(f12097,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP5('T_0','const.bool.COND_3'(s(bool,X0),s(bool,happ(s(fun(X1,bool),X2),s(X1,X3))),s(bool,happ(s(fun(X1,bool),X2),s(X1,X4)))),happ(s(fun(X1,bool),X2),s(X1,'const.bool.COND_3'(s(bool,X0),s(X1,X3),s(X1,X4)))))) )),
  inference(unit_resulting_resolution,[],[f335,f479,f522])).
fof(f12883,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP5('F_0',happ(s(fun(X0,bool),X1),s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X3),s(X0,X4)))),'const.bool.COND_3'(s(bool,X2),s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))),s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))))) )),
  inference(unit_resulting_resolution,[],[f592,f479,f524])).
fof(f12884,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP5('T_0',happ(s(fun(X0,bool),X1),s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X3),s(X0,X4)))),'const.bool.COND_3'(s(bool,X2),s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))),s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))))) )),
  inference(unit_resulting_resolution,[],[f335,f479,f522])).
fof(f11665,plain,(
  ( ! [X8,X7] : (sP1('const.bool.COND_3'(s(bool,'F_0'),s(bool,X8),s(bool,'T_0')),X7) | p(s(bool,X7))) )),
  inference(forward_demodulation,[],[f11664,f1595])).
fof(f11664,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X7)) | sP1('const.bool.COND_3'(s(bool,'F_0'),s(bool,X8),s(bool,'const.bool.COND_3'(s(bool,X6),s(bool,'T_0'),s(bool,'T_0')))),X7)) )),
  inference(subsumption_resolution,[],[f11661,f335])).
fof(f11661,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'T_0')) | p(s(bool,X7)) | sP1('const.bool.COND_3'(s(bool,'F_0'),s(bool,X8),s(bool,'const.bool.COND_3'(s(bool,X6),s(bool,'T_0'),s(bool,'T_0')))),X7)) )),
  inference(superposition,[],[f863,f1595])).
fof(f11657,plain,(
  ( ! [X0] : (sP1('const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f335,f592,f863])).
fof(f7360,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,X1) != s(X0,X2) | s(X0,X3) != s(X0,X4) | s(X0,'const.bool.COND_3'(s(bool,X5),s(X0,X3),s(X0,X1))) = s(X0,'const.bool.COND_3'(s(bool,X5),s(X0,X4),s(X0,X2)))) )),
  inference(equality_resolution,[],[f428])).
fof(f11643,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (s(bool,'F_0') != s(bool,X8) | s(bool,'T_0') = s(bool,X7) | s(X9,'const.bool.COND_3'(s(bool,X7),s(X9,X12),s(X9,X10))) = s(X9,'const.bool.COND_3'(s(bool,X8),s(X9,X13),s(X9,X11))) | s(X9,X10) != s(X9,X11)) )),
  inference(global_subsumption,[],[f334,f333,f335,f343,f355,f354,f413,f428,f427,f426,f429,f430,f437,f434,f444,f443,f593,f462,f459,f458,f457,f478,f479,f480,f481,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f819,f825,f828,f830,f840,f888,f898,f901,f917,f923,f933,f939,f946,f954,f960,f966,f971,f975,f980,f984,f994,f1000,f1007,f1012,f1024,f1032,f1047,f1050,f1062,f1068,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1239,f1275,f1339,f1332,f1373,f1385,f1402,f1422,f1433,f1437,f1461,f1501,f1512,f1516,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1662,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1981,f2015,f2030,f2073,f2082,f2119,f2118,f2128,f2157,f2170,f2252,f2222,f2262,f2270,f2617,f2622,f2841,f2844,f2751,f2544,f2504,f2465,f2432,f2400,f2371,f2351,f2317,f1837,f1818,f424,f423,f422,f421,f420,f419,f418,f417,f412,f411,f1854,f1861,f1852,f1859,f2052,f2050,f2043,f2041,f482,f433,f1542,f1492,f1486,f489,f488,f487,f485,f484,f1219,f1249,f1217,f1247,f1212,f1256,f1285,f1210,f1254,f1283,f2099,f2097,f2090,f2088,f2001,f1931,f1230,f483,f477,f468,f467,f432,f431,f2848,f2847,f2770,f2765,f2545,f2505,f2466,f2433,f2401,f2398,f2373,f2372,f2352,f2349,f2319,f2318,f2272,f2271,f2269,f2251,f1838,f1835,f1819,f1816,f1663,f1661,f1494,f1411,f1410,f1403,f1366,f1365,f1333,f1267,f1231,f1069,f1059,f1057,f1048,f1035,f1033,f1025,f1023,f1001,f991,f961,f942,f940,f930,f926,f924,f904,f902,f891,f889,f873,f871,f870,f869,f868,f867,f866,f865,f864,f863,f862,f861,f852,f839,f838,f837,f836,f829,f818,f787,f592,f414,f359,f358,f4838,f4843,f4844,f4853,f4855,f4857,f4858,f4859,f4860,f4861,f4862,f4863,f4864,f4644,f4646,f4648,f4650,f4870,f4655,f4658,f4660,f4661,f4663,f4669,f4977,f4986,f5268,f5276,f7384,f7394,f7492,f7497,f7506,f7666,f7671,f7680,f7351])).
fof(f7351,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (s(bool,'F_0') != s(bool,X8) | s(X9,X10) != s(X9,X11) | s(X9,X12) != s(X9,X13) | s(X9,'const.bool.COND_3'(s(bool,X7),s(X9,X12),s(X9,X10))) = s(X9,'const.bool.COND_3'(s(bool,X8),s(X9,X13),s(X9,X11))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f428,f343])).
fof(f7680,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X3),s('type.num.num',X4))))) = s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X2),s('type.num.num',X1))),s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X4),s('type.num.num',X3)))))) )),
  inference(forward_demodulation,[],[f7187,f414])).
fof(f7187,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X3),s('type.num.num',X4))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X5),s(bool,X0))),s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X2),s('type.num.num',X1))),s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X4),s('type.num.num',X3)))))) )),
  inference(unit_resulting_resolution,[],[f430,f430,f414,f428])).
fof(f7671,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X3),s('type.num.num',X4))))) = s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X2),s('type.num.num',X1))),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X4),s('type.num.num',X3)))))) )),
  inference(forward_demodulation,[],[f7192,f414])).
fof(f7192,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X3),s('type.num.num',X4))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X5),s(bool,X0))),s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X2),s('type.num.num',X1))),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X4),s('type.num.num',X3)))))) )),
  inference(unit_resulting_resolution,[],[f429,f430,f414,f428])).
fof(f7666,plain,(
  ( ! [X4,X2,X0,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num',X4))) = s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X2),s('type.num.num',X1))),s('type.num.num',X4)))) )),
  inference(forward_demodulation,[],[f7665,f414])).
fof(f7665,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X2),s('type.num.num',X1))),s('type.num.num',X4))) = s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X3),s('type.num.num',X4)))))) )),
  inference(forward_demodulation,[],[f7195,f414])).
fof(f7195,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X3),s('type.num.num',X4))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X5),s(bool,X0))),s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X2),s('type.num.num',X1))),s('type.num.num',X4)))) )),
  inference(unit_resulting_resolution,[],[f414,f430,f414,f428])).
fof(f7506,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X3),s('type.num.num',X4))))) = s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X2),s('type.num.num',X1))),s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X4),s('type.num.num',X3)))))) )),
  inference(forward_demodulation,[],[f7277,f414])).
fof(f7277,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X3),s('type.num.num',X4))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X5),s(bool,X0))),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X2),s('type.num.num',X1))),s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X4),s('type.num.num',X3)))))) )),
  inference(unit_resulting_resolution,[],[f430,f429,f414,f428])).
fof(f7497,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X3),s('type.num.num',X4))))) = s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X2),s('type.num.num',X1))),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X4),s('type.num.num',X3)))))) )),
  inference(forward_demodulation,[],[f7282,f414])).
fof(f7282,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X3),s('type.num.num',X4))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X5),s(bool,X0))),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X2),s('type.num.num',X1))),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X4),s('type.num.num',X3)))))) )),
  inference(unit_resulting_resolution,[],[f429,f429,f414,f428])).
fof(f7492,plain,(
  ( ! [X4,X2,X0,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num',X4))) = s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X2),s('type.num.num',X1))),s('type.num.num',X4)))) )),
  inference(forward_demodulation,[],[f7491,f414])).
fof(f7491,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X2),s('type.num.num',X1))),s('type.num.num',X4))) = s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X3),s('type.num.num',X4)))))) )),
  inference(forward_demodulation,[],[f7285,f414])).
fof(f7285,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X3),s('type.num.num',X4))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X5),s(bool,X0))),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X2),s('type.num.num',X1))),s('type.num.num',X4)))) )),
  inference(unit_resulting_resolution,[],[f414,f429,f414,f428])).
fof(f7394,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num',X1),s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X2),s('type.num.num',X3))))) = s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num',X1),s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X3),s('type.num.num',X2)))))) )),
  inference(forward_demodulation,[],[f7393,f414])).
fof(f7393,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num',X1),s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X2),s('type.num.num',X3))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X4),s(bool,X0))),s('type.num.num',X1),s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X3),s('type.num.num',X2)))))) )),
  inference(forward_demodulation,[],[f7334,f414])).
fof(f7334,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num',X1),s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X2),s('type.num.num',X3))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X4),s(bool,X0))),s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X5),s('type.num.num',X1))),s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X3),s('type.num.num',X2)))))) )),
  inference(unit_resulting_resolution,[],[f430,f414,f414,f428])).
fof(f7384,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num',X1),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X2),s('type.num.num',X3))))) = s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num',X1),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X3),s('type.num.num',X2)))))) )),
  inference(forward_demodulation,[],[f7383,f414])).
fof(f7383,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num',X1),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X2),s('type.num.num',X3))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X4),s(bool,X0))),s('type.num.num',X1),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X3),s('type.num.num',X2)))))) )),
  inference(forward_demodulation,[],[f7339,f414])).
fof(f7339,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num',X1),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X2),s('type.num.num',X3))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X4),s(bool,X0))),s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X5),s('type.num.num',X1))),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X3),s('type.num.num',X2)))))) )),
  inference(unit_resulting_resolution,[],[f429,f414,f414,f428])).
fof(f5276,plain,(
  ( ! [X8,X7] : (sP1(X7,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X8))) | p(s(bool,X7))) )),
  inference(forward_demodulation,[],[f5275,f1595])).
fof(f5275,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X7)) | sP1(X7,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,X6),s(bool,'T_0'),s(bool,'T_0'))),s(bool,X8)))) )),
  inference(subsumption_resolution,[],[f5272,f335])).
fof(f5272,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'T_0')) | p(s(bool,X7)) | sP1(X7,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,X6),s(bool,'T_0'),s(bool,'T_0'))),s(bool,X8)))) )),
  inference(superposition,[],[f799,f1595])).
fof(f5268,plain,(
  ( ! [X0] : (sP1('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f335,f592,f799])).
fof(f4986,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X1),s(bool,X2),s(bool,'T_0')))) )),
  inference(subsumption_resolution,[],[f4971,f592])).
fof(f4971,plain,(
  ( ! [X2,X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X1),s(bool,X2),s(bool,'T_0')))) )),
  inference(superposition,[],[f4855,f343])).
fof(f4977,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f4960,f592])).
fof(f4960,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,'T_0')))) | p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f4855,f343])).
fof(f4669,plain,(
  ( ! [X140,X141] : (sP4('T_0','const.bool.COND_3'(s(bool,X140),s(bool,'T_0'),s(bool,'T_0')),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X141),s(bool,'T_0')))) )),
  inference(superposition,[],[f930,f1595])).
fof(f4663,plain,(
  ( ! [X127,X126] : (sP5('T_0','const.bool.COND_3'(s(bool,X126),s(bool,'T_0'),s(bool,'T_0')),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X127),s(bool,'T_0')))) )),
  inference(superposition,[],[f839,f1595])).
fof(f4661,plain,(
  ( ! [X123,X122] : (~sP5('F_0','const.bool.COND_3'(s(bool,X122),s(bool,'T_0'),s(bool,'T_0')),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X123),s(bool,'T_0')))) )),
  inference(superposition,[],[f838,f1595])).
fof(f4660,plain,(
  ( ! [X121,X120] : (sP5('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X121),s(bool,'T_0')),'const.bool.COND_3'(s(bool,X120),s(bool,'T_0'),s(bool,'T_0')))) )),
  inference(superposition,[],[f837,f1595])).
fof(f4658,plain,(
  ( ! [X116,X117] : (~sP5('F_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X117),s(bool,'T_0')),'const.bool.COND_3'(s(bool,X116),s(bool,'T_0'),s(bool,'T_0')))) )),
  inference(superposition,[],[f836,f1595])).
fof(f4655,plain,(
  ( ! [X111,X110] : (sP4('T_0','const.bool.COND_3'(s(bool,X110),s(bool,'T_0'),s(bool,'T_0')),'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X111)))) )),
  inference(superposition,[],[f825,f1595])).
fof(f4870,plain,(
  ( ! [X109,X108] : (sP1('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X109)),X108) | p(s(bool,X108))) )),
  inference(forward_demodulation,[],[f4869,f1595])).
fof(f4869,plain,(
  ( ! [X109,X107,X108] : (p(s(bool,X108)) | sP1('const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,X107),s(bool,'T_0'),s(bool,'T_0'))),s(bool,X109)),X108)) )),
  inference(subsumption_resolution,[],[f4654,f335])).
fof(f4654,plain,(
  ( ! [X109,X107,X108] : (~p(s(bool,'T_0')) | p(s(bool,X108)) | sP1('const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,X107),s(bool,'T_0'),s(bool,'T_0'))),s(bool,X109)),X108)) )),
  inference(superposition,[],[f798,f1595])).
fof(f4650,plain,(
  ( ! [X97,X98] : (sP5('T_0','const.bool.COND_3'(s(bool,X97),s(bool,'T_0'),s(bool,'T_0')),'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X98)))) )),
  inference(superposition,[],[f778,f1595])).
fof(f4648,plain,(
  ( ! [X94,X93] : (~sP5('F_0','const.bool.COND_3'(s(bool,X93),s(bool,'T_0'),s(bool,'T_0')),'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X94)))) )),
  inference(superposition,[],[f777,f1595])).
fof(f4646,plain,(
  ( ! [X90,X89] : (sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X90)),'const.bool.COND_3'(s(bool,X89),s(bool,'T_0'),s(bool,'T_0')))) )),
  inference(superposition,[],[f776,f1595])).
fof(f4644,plain,(
  ( ! [X85,X86] : (~sP5('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X86)),'const.bool.COND_3'(s(bool,X85),s(bool,'T_0'),s(bool,'T_0')))) )),
  inference(superposition,[],[f775,f1595])).
fof(f4864,plain,(
  ( ! [X80,X81,X82] : (~sP8('const.bool.COND_3'(s(bool,X80),s(bool,'T_0'),s(bool,'T_0')),X81,X82) | sP9('const.bool.COND_3'(s(bool,X80),s(bool,'T_0'),s(bool,'T_0')),X81,X82)) )),
  inference(subsumption_resolution,[],[f4642,f335])).
fof(f4642,plain,(
  ( ! [X80,X81,X82] : (~p(s(bool,'T_0')) | ~sP8('const.bool.COND_3'(s(bool,X80),s(bool,'T_0'),s(bool,'T_0')),X81,X82) | sP9('const.bool.COND_3'(s(bool,X80),s(bool,'T_0'),s(bool,'T_0')),X81,X82)) )),
  inference(superposition,[],[f577,f1595])).
fof(f4863,plain,(
  ( ! [X78,X79,X77] : (~sP10('const.bool.COND_3'(s(bool,X77),s(bool,'T_0'),s(bool,'T_0')),X78,X79) | sP11('const.bool.COND_3'(s(bool,X77),s(bool,'T_0'),s(bool,'T_0')),X78,X79)) )),
  inference(subsumption_resolution,[],[f4641,f335])).
fof(f4641,plain,(
  ( ! [X78,X79,X77] : (~p(s(bool,'T_0')) | ~sP10('const.bool.COND_3'(s(bool,X77),s(bool,'T_0'),s(bool,'T_0')),X78,X79) | sP11('const.bool.COND_3'(s(bool,X77),s(bool,'T_0'),s(bool,'T_0')),X78,X79)) )),
  inference(superposition,[],[f576,f1595])).
fof(f4862,plain,(
  ( ! [X76,X74,X75] : (sP11('const.bool.COND_3'(s(bool,X74),s(bool,'T_0'),s(bool,'T_0')),X75,X76) | ~p(s(bool,X75))) )),
  inference(subsumption_resolution,[],[f4640,f335])).
fof(f4640,plain,(
  ( ! [X76,X74,X75] : (~p(s(bool,'T_0')) | ~p(s(bool,X75)) | sP11('const.bool.COND_3'(s(bool,X74),s(bool,'T_0'),s(bool,'T_0')),X75,X76)) )),
  inference(superposition,[],[f572,f1595])).
fof(f4861,plain,(
  ( ! [X72,X71,X73] : (sP11('const.bool.COND_3'(s(bool,X71),s(bool,'T_0'),s(bool,'T_0')),X73,X72) | p(s(bool,X72))) )),
  inference(subsumption_resolution,[],[f4639,f335])).
fof(f4639,plain,(
  ( ! [X72,X71,X73] : (~p(s(bool,'T_0')) | p(s(bool,X72)) | sP11('const.bool.COND_3'(s(bool,X71),s(bool,'T_0'),s(bool,'T_0')),X73,X72)) )),
  inference(superposition,[],[f571,f1595])).
fof(f4860,plain,(
  ( ! [X70,X68,X69] : (sP9('const.bool.COND_3'(s(bool,X68),s(bool,'T_0'),s(bool,'T_0')),X69,X70) | ~p(s(bool,X69))) )),
  inference(subsumption_resolution,[],[f4638,f335])).
fof(f4638,plain,(
  ( ! [X70,X68,X69] : (~p(s(bool,'T_0')) | ~p(s(bool,X69)) | sP9('const.bool.COND_3'(s(bool,X68),s(bool,'T_0'),s(bool,'T_0')),X69,X70)) )),
  inference(superposition,[],[f555,f1595])).
fof(f4859,plain,(
  ( ! [X66,X67,X65] : (sP9('const.bool.COND_3'(s(bool,X65),s(bool,'T_0'),s(bool,'T_0')),X67,X66) | ~p(s(bool,X66))) )),
  inference(subsumption_resolution,[],[f4637,f335])).
fof(f4637,plain,(
  ( ! [X66,X67,X65] : (~p(s(bool,'T_0')) | ~p(s(bool,X66)) | sP9('const.bool.COND_3'(s(bool,X65),s(bool,'T_0'),s(bool,'T_0')),X67,X66)) )),
  inference(superposition,[],[f554,f1595])).
fof(f4858,plain,(
  ( ! [X64,X62,X63] : (~p(s(bool,X63)) | p(s(bool,X64)) | ~sP7(X64,X63,'const.bool.COND_3'(s(bool,X62),s(bool,'T_0'),s(bool,'T_0')))) )),
  inference(subsumption_resolution,[],[f4636,f335])).
fof(f4636,plain,(
  ( ! [X64,X62,X63] : (~p(s(bool,'T_0')) | ~p(s(bool,X63)) | p(s(bool,X64)) | ~sP7(X64,X63,'const.bool.COND_3'(s(bool,X62),s(bool,'T_0'),s(bool,'T_0')))) )),
  inference(superposition,[],[f541,f1595])).
fof(f4857,plain,(
  ( ! [X61,X60] : (~p(s(bool,X60)) | p(s(bool,'const.bool.COND_3'(s(bool,X61),s(bool,'T_0'),s(bool,X60))))) )),
  inference(forward_demodulation,[],[f4856,f1595])).
fof(f4856,plain,(
  ( ! [X61,X59,X60] : (~p(s(bool,X60)) | p(s(bool,'const.bool.COND_3'(s(bool,X61),s(bool,'const.bool.COND_3'(s(bool,X59),s(bool,'T_0'),s(bool,'T_0'))),s(bool,X60))))) )),
  inference(subsumption_resolution,[],[f4635,f335])).
fof(f4635,plain,(
  ( ! [X61,X59,X60] : (~p(s(bool,'T_0')) | ~p(s(bool,X60)) | p(s(bool,'const.bool.COND_3'(s(bool,X61),s(bool,'const.bool.COND_3'(s(bool,X59),s(bool,'T_0'),s(bool,'T_0'))),s(bool,X60))))) )),
  inference(superposition,[],[f489,f1595])).
fof(f4855,plain,(
  ( ! [X57,X58] : (p(s(bool,'const.bool.COND_3'(s(bool,X57),s(bool,X58),s(bool,'T_0')))) | p(s(bool,X57))) )),
  inference(forward_demodulation,[],[f4854,f1595])).
fof(f4854,plain,(
  ( ! [X57,X58,X56] : (p(s(bool,'const.bool.COND_3'(s(bool,X57),s(bool,X58),s(bool,'const.bool.COND_3'(s(bool,X56),s(bool,'T_0'),s(bool,'T_0')))))) | p(s(bool,X57))) )),
  inference(subsumption_resolution,[],[f4634,f335])).
fof(f4634,plain,(
  ( ! [X57,X58,X56] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.COND_3'(s(bool,X57),s(bool,X58),s(bool,'const.bool.COND_3'(s(bool,X56),s(bool,'T_0'),s(bool,'T_0')))))) | p(s(bool,X57))) )),
  inference(superposition,[],[f488,f1595])).
fof(f4853,plain,(
  ( ! [X54,X55] : (~p(s(bool,X54)) | p(s(bool,'const.bool.COND_3'(s(bool,X54),s(bool,'T_0'),s(bool,X55))))) )),
  inference(forward_demodulation,[],[f4852,f1595])).
fof(f4852,plain,(
  ( ! [X54,X55,X53] : (~p(s(bool,X54)) | p(s(bool,'const.bool.COND_3'(s(bool,X54),s(bool,'const.bool.COND_3'(s(bool,X53),s(bool,'T_0'),s(bool,'T_0'))),s(bool,X55))))) )),
  inference(subsumption_resolution,[],[f4633,f335])).
fof(f4633,plain,(
  ( ! [X54,X55,X53] : (~p(s(bool,'T_0')) | ~p(s(bool,X54)) | p(s(bool,'const.bool.COND_3'(s(bool,X54),s(bool,'const.bool.COND_3'(s(bool,X53),s(bool,'T_0'),s(bool,'T_0'))),s(bool,X55))))) )),
  inference(superposition,[],[f487,f1595])).
fof(f4844,plain,(
  ( ! [X39,X40] : (sP1(X40,'const.bool.COND_3'(s(bool,X39),s(bool,'T_0'),s(bool,'T_0'))) | p(s(bool,X40))) )),
  inference(subsumption_resolution,[],[f4628,f335])).
fof(f4628,plain,(
  ( ! [X39,X40] : (~p(s(bool,'T_0')) | p(s(bool,X40)) | sP1(X40,'const.bool.COND_3'(s(bool,X39),s(bool,'T_0'),s(bool,'T_0')))) )),
  inference(superposition,[],[f462,f1595])).
fof(f4843,plain,(
  ( ! [X37,X38] : (sP1('const.bool.COND_3'(s(bool,X37),s(bool,'T_0'),s(bool,'T_0')),X38) | p(s(bool,X38))) )),
  inference(subsumption_resolution,[],[f4627,f335])).
fof(f4627,plain,(
  ( ! [X37,X38] : (~p(s(bool,'T_0')) | p(s(bool,X38)) | sP1('const.bool.COND_3'(s(bool,X37),s(bool,'T_0'),s(bool,'T_0')),X38)) )),
  inference(superposition,[],[f459,f1595])).
fof(f4838,plain,(
  ( ! [X6,X12,X10,X8,X7,X11,X9] : (s(X7,'const.bool.COND_3'(s(bool,'T_0'),s(X7,X11),s(X7,X12))) = s(X7,'const.bool.COND_3'(s(bool,X8),s(X7,X9),s(X7,X10))) | s(X7,X9) != s(X7,X11) | s(bool,X8) != s(bool,'const.bool.COND_3'(s(bool,X6),s(bool,'T_0'),s(bool,'T_0')))) )),
  inference(forward_demodulation,[],[f4837,f1595])).
fof(f4837,plain,(
  ( ! [X6,X12,X10,X8,X7,X11,X9] : (s(X7,'const.bool.COND_3'(s(bool,X8),s(X7,X9),s(X7,X10))) = s(X7,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,X6),s(bool,'T_0'),s(bool,'T_0'))),s(X7,X11),s(X7,X12))) | s(X7,X9) != s(X7,X11) | s(bool,X8) != s(bool,'const.bool.COND_3'(s(bool,X6),s(bool,'T_0'),s(bool,'T_0')))) )),
  inference(subsumption_resolution,[],[f4619,f335])).
fof(f4619,plain,(
  ( ! [X6,X12,X10,X8,X7,X11,X9] : (~p(s(bool,'T_0')) | s(X7,'const.bool.COND_3'(s(bool,X8),s(X7,X9),s(X7,X10))) = s(X7,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,X6),s(bool,'T_0'),s(bool,'T_0'))),s(X7,X11),s(X7,X12))) | s(X7,X9) != s(X7,X11) | s(bool,X8) != s(bool,'const.bool.COND_3'(s(bool,X6),s(bool,'T_0'),s(bool,'T_0')))) )),
  inference(superposition,[],[f426,f1595])).
fof(f358,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f228])).
fof(f228,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f90])).
fof(f90,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f89])).
fof(f89,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X9] : (s(bool,'F_0') = s(bool,X9) <=> ~p(s(bool,X9)))),
  file('thm.arithmetic.MIN_MAX_EQ','thm.bool.EQ_CLAUSES|split|3')).
fof(f359,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f228])).
fof(f787,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f413,f343])).
fof(f818,plain,(
  ( ! [X2,X3] : (~sP5('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X3)),X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f775,f343])).
fof(f829,plain,(
  ( ! [X2,X3] : (sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X3)),X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f776,f343])).
fof(f836,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,X1)),X1)) )),
  inference(unit_resulting_resolution,[],[f592,f414,f524])).
fof(f837,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,X1)),X1)) )),
  inference(unit_resulting_resolution,[],[f335,f414,f522])).
fof(f838,plain,(
  ( ! [X0,X1] : (~sP5('F_0',X0,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f592,f414,f524])).
fof(f839,plain,(
  ( ! [X0,X1] : (sP5('T_0',X0,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f335,f414,f522])).
fof(f852,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,X1))) | s(bool,'F_0') = s(bool,X1)) )),
  inference(superposition,[],[f414,f343])).
fof(f861,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,X20)) | sP0(X21,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X19),s(bool,X20)))) )),
  inference(superposition,[],[f443,f414])).
fof(f862,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X23)) | sP0('const.bool.COND_3'(s(bool,'F_0'),s(bool,X22),s(bool,X23)),X24)) )),
  inference(superposition,[],[f444,f414])).
fof(f863,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,X26)) | p(s(bool,X27)) | sP1('const.bool.COND_3'(s(bool,'F_0'),s(bool,X25),s(bool,X26)),X27)) )),
  inference(superposition,[],[f459,f414])).
fof(f864,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,X29)) | p(s(bool,X30)) | sP1(X30,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X28),s(bool,X29)))) )),
  inference(superposition,[],[f462,f414])).
fof(f865,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,X33)) | ~p(s(bool,X32)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X31),s(bool,X32)),X34,X33)) )),
  inference(superposition,[],[f554,f414])).
fof(f866,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X37)) | ~p(s(bool,X36)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X35),s(bool,X36)),X37,X38)) )),
  inference(superposition,[],[f555,f414])).
fof(f867,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X40)) | p(s(bool,X41)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X39),s(bool,X40)),X42,X41)) )),
  inference(superposition,[],[f571,f414])).
fof(f868,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X45)) | ~p(s(bool,X44)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X43),s(bool,X44)),X45,X46)) )),
  inference(superposition,[],[f572,f414])).
fof(f869,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X48)) | ~sP10('const.bool.COND_3'(s(bool,'F_0'),s(bool,X47),s(bool,X48)),X49,X50) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X47),s(bool,X48)),X49,X50)) )),
  inference(superposition,[],[f576,f414])).
fof(f870,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X52)) | ~sP8('const.bool.COND_3'(s(bool,'F_0'),s(bool,X51),s(bool,X52)),X53,X54) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X51),s(bool,X52)),X53,X54)) )),
  inference(superposition,[],[f577,f414])).
fof(f871,plain,(
  ( ! [X57,X56,X55] : (~sP5('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X56),s(bool,X57)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X55),s(bool,X56)))) )),
  inference(superposition,[],[f775,f414])).
fof(f873,plain,(
  ( ! [X61,X62,X63] : (sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X62),s(bool,X63)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X61),s(bool,X62)))) )),
  inference(superposition,[],[f776,f414])).
fof(f889,plain,(
  ( ! [X2,X3] : (~sP5('F_0',X2,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X3))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f777,f343])).
fof(f891,plain,(
  ( ! [X8,X7,X9] : (~sP5('F_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X7),s(bool,X8)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X8),s(bool,X9)))) )),
  inference(superposition,[],[f777,f414])).
fof(f902,plain,(
  ( ! [X2,X3] : (sP5('T_0',X2,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X3))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f778,f343])).
fof(f904,plain,(
  ( ! [X8,X7,X9] : (sP5('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X7),s(bool,X8)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X8),s(bool,X9)))) )),
  inference(superposition,[],[f778,f414])).
fof(f924,plain,(
  ( ! [X2,X3] : (~sP5('F_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,'F_0')),X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f836,f343])).
fof(f926,plain,(
  ( ! [X8,X7,X9] : (~sP5('F_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X9),s(bool,X8)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X7),s(bool,X8)))) )),
  inference(superposition,[],[f836,f414])).
fof(f930,plain,(
  ( ! [X0,X1] : (sP4('T_0',X0,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f837,f504])).
fof(f940,plain,(
  ( ! [X2,X3] : (sP5('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,'F_0')),X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f837,f343])).
fof(f942,plain,(
  ( ! [X8,X7,X9] : (sP5('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X9),s(bool,X8)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X7),s(bool,X8)))) )),
  inference(superposition,[],[f837,f414])).
fof(f961,plain,(
  ( ! [X2,X3] : (~sP5('F_0',X2,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f838,f343])).
fof(f991,plain,(
  ( ! [X0,X1] : (sP4('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,X1)),X1)) )),
  inference(unit_resulting_resolution,[],[f839,f504])).
fof(f1001,plain,(
  ( ! [X2,X3] : (sP5('T_0',X2,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f839,f343])).
fof(f1023,plain,(
  ( ! [X2,X3] : (sP4('T_0',X2,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X3))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f825,f343])).
fof(f1025,plain,(
  ( ! [X8,X7,X9] : (sP4('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X7),s(bool,X8)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X8),s(bool,X9)))) )),
  inference(superposition,[],[f825,f414])).
fof(f1033,plain,(
  ( ! [X2,X3] : (sP4('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X3)),X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f898,f343])).
fof(f1035,plain,(
  ( ! [X8,X7,X9] : (sP4('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X8),s(bool,X9)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X7),s(bool,X8)))) )),
  inference(superposition,[],[f898,f414])).
fof(f1057,plain,(
  ( ! [X2,X3] : (sP4('T_0',X2,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f930,f343])).
fof(f1059,plain,(
  ( ! [X8,X7,X9] : (sP4('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X7),s(bool,X8)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X9),s(bool,X8)))) )),
  inference(superposition,[],[f930,f414])).
fof(f1069,plain,(
  ( ! [X2,X3] : (sP4('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,'F_0')),X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f991,f343])).
fof(f1231,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f335,f861])).
fof(f1267,plain,(
  ( ! [X0,X1] : (sP0('const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f335,f862])).
fof(f1333,plain,(
  ( ! [X0] : (~sP3('const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f605,f836,f599,f335,f506])).
fof(f1365,plain,(
  ( ! [X0] : (~sP2('const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X0)),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f777,f599,f602,f592,f507])).
fof(f1366,plain,(
  ( ! [X0] : (~sP2('const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'F_0')),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f838,f599,f602,f592,f507])).
fof(f1403,plain,(
  ( ! [X0] : (~sP3('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f604,f838,f597,f592,f508])).
fof(f1410,plain,(
  ( ! [X0] : (~sP2('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X0)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f775,f598,f601,f592,f508])).
fof(f1411,plain,(
  ( ! [X0] : (~sP2('F_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'F_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f836,f598,f601,f592,f508])).
fof(f1494,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,'F_0')))) | p(s(bool,X0))) )),
  inference(condensation,[],[f1493])).
fof(f1493,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'const.bool.COND_3'(s(bool,X4),s(bool,X5),s(bool,'F_0')))) | p(s(bool,X4)) | p(s(bool,X3))) )),
  inference(subsumption_resolution,[],[f1478,f354])).
fof(f1478,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'const.bool.COND_3'(s(bool,X4),s(bool,X5),s(bool,'F_0')))) | p(s(bool,X4)) | p(s(bool,X3)) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f485,f343])).
fof(f1661,plain,(
  ( ! [X2,X0,X1] : (sP5('const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,X2)),X2)) )),
  inference(unit_resulting_resolution,[],[f414,f1582,f522])).
fof(f1663,plain,(
  ( ! [X2,X0,X1] : (sP5('const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),X1,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X1)))) )),
  inference(unit_resulting_resolution,[],[f414,f1582,f522])).
fof(f1816,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X1)))) | s(bool,'T_0') = s(bool,sK14(X0,bool))) )),
  inference(superposition,[],[f411,f343])).
fof(f1819,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,happ(s(fun(X4,bool),X6),s(X4,X7)))) | ~p(s(bool,happ(s(fun(X4,bool),X6),s(X4,sK14('const.bool.COND_3'(s(bool,'F_0'),s(fun(X4,bool),X5),s(fun(X4,bool),X6)),X4)))))) )),
  inference(forward_demodulation,[],[f1815,f414])).
fof(f1815,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,happ(s(fun(X4,bool),X6),s(X4,sK14('const.bool.COND_3'(s(bool,'F_0'),s(fun(X4,bool),X5),s(fun(X4,bool),X6)),X4))))) | p(s(bool,happ(s(fun(X4,bool),'const.bool.COND_3'(s(bool,'F_0'),s(fun(X4,bool),X5),s(fun(X4,bool),X6))),s(X4,X7))))) )),
  inference(superposition,[],[f411,f414])).
fof(f1835,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X1)))) | s(bool,'T_0') = s(bool,sK15(X0,bool))) )),
  inference(superposition,[],[f412,f343])).
fof(f1838,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,happ(s(fun(X4,bool),X6),s(X4,X7)))) | ~p(s(bool,happ(s(fun(X4,bool),X6),s(X4,sK15('const.bool.COND_3'(s(bool,'F_0'),s(fun(X4,bool),X5),s(fun(X4,bool),X6)),X4)))))) )),
  inference(forward_demodulation,[],[f1834,f414])).
fof(f1834,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,happ(s(fun(X4,bool),X6),s(X4,sK15('const.bool.COND_3'(s(bool,'F_0'),s(fun(X4,bool),X5),s(fun(X4,bool),X6)),X4))))) | p(s(bool,happ(s(fun(X4,bool),'const.bool.COND_3'(s(bool,'F_0'),s(fun(X4,bool),X5),s(fun(X4,bool),X6))),s(X4,X7))))) )),
  inference(superposition,[],[f412,f414])).
fof(f2251,plain,(
  ( ! [X0,X1] : (~sP3('const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,'T_0')),'F_0')) )),
  inference(forward_demodulation,[],[f2226,f1595])).
fof(f2226,plain,(
  ( ! [X0,X1] : (~sP3('const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f592,f838,f599,f1658,f508])).
fof(f2269,plain,(
  ( ! [X0,X1] : (~sP3('const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')),'const.bool.COND_3'(s(bool,X1),s(bool,'T_0'),s(bool,'T_0')),'F_0')) )),
  inference(forward_demodulation,[],[f2261,f1595])).
fof(f2261,plain,(
  ( ! [X0,X1] : (~sP3('const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'const.bool.COND_3'(s(bool,X1),s(bool,'T_0'),s(bool,'T_0')))),'const.bool.COND_3'(s(bool,X1),s(bool,'T_0'),s(bool,'T_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f592,f836,f599,f1659,f508])).
fof(f2271,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK22(X0,X1,X2,bool))) )),
  inference(superposition,[],[f478,f343])).
fof(f2272,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK22(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK22(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f478,f343])).
fof(f2318,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~p(s(bool,happ(s(fun(X5,bool),X7),s(X5,sK16('const.bool.COND_3'(s(bool,'F_0'),s(fun(X5,bool),X6),s(fun(X5,bool),X7)),X8,X5))))) | p(s(bool,happ(s(fun(X5,bool),X8),s(X5,X9)))) | ~p(s(bool,happ(s(fun(X5,bool),X8),s(X5,sK16('const.bool.COND_3'(s(bool,'F_0'),s(fun(X5,bool),X6),s(fun(X5,bool),X7)),X8,X5)))))) )),
  inference(superposition,[],[f417,f414])).
fof(f2319,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X1),s(bool,X2)))) | ~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,sK16(X0,X1,bool))))) | s(bool,'T_0') = s(bool,sK16(X0,X1,bool))) )),
  inference(superposition,[],[f417,f343])).
fof(f2349,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X2)))) | ~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,sK16(X0,X1,bool))))) | s(bool,'T_0') = s(bool,sK16(X0,X1,bool))) )),
  inference(superposition,[],[f418,f343])).
fof(f2352,plain,(
  ( ! [X6,X8,X7,X5,X9] : (p(s(bool,happ(s(fun(X5,bool),X7),s(X5,X9)))) | ~p(s(bool,happ(s(fun(X5,bool),X7),s(X5,sK16('const.bool.COND_3'(s(bool,'F_0'),s(fun(X5,bool),X6),s(fun(X5,bool),X7)),X8,X5))))) | ~p(s(bool,happ(s(fun(X5,bool),X8),s(X5,sK16('const.bool.COND_3'(s(bool,'F_0'),s(fun(X5,bool),X6),s(fun(X5,bool),X7)),X8,X5)))))) )),
  inference(forward_demodulation,[],[f2348,f414])).
fof(f2348,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~p(s(bool,happ(s(fun(X5,bool),X7),s(X5,sK16('const.bool.COND_3'(s(bool,'F_0'),s(fun(X5,bool),X6),s(fun(X5,bool),X7)),X8,X5))))) | p(s(bool,happ(s(fun(X5,bool),'const.bool.COND_3'(s(bool,'F_0'),s(fun(X5,bool),X6),s(fun(X5,bool),X7))),s(X5,X9)))) | ~p(s(bool,happ(s(fun(X5,bool),X8),s(X5,sK16('const.bool.COND_3'(s(bool,'F_0'),s(fun(X5,bool),X6),s(fun(X5,bool),X7)),X8,X5)))))) )),
  inference(superposition,[],[f418,f414])).
fof(f2372,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~p(s(bool,happ(s(fun(X5,bool),X7),s(X5,sK18('const.bool.COND_3'(s(bool,'F_0'),s(fun(X5,bool),X6),s(fun(X5,bool),X7)),X8,X5))))) | p(s(bool,happ(s(fun(X5,bool),X8),s(X5,X9)))) | ~p(s(bool,happ(s(fun(X5,bool),X8),s(X5,sK17('const.bool.COND_3'(s(bool,'F_0'),s(fun(X5,bool),X6),s(fun(X5,bool),X7)),X8,X5)))))) )),
  inference(superposition,[],[f419,f414])).
fof(f2373,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X1),s(bool,X2)))) | ~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,sK17(X0,X1,bool))))) | s(bool,'T_0') = s(bool,sK18(X0,X1,bool))) )),
  inference(superposition,[],[f419,f343])).
fof(f2398,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X2)))) | ~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,sK17(X0,X1,bool))))) | s(bool,'T_0') = s(bool,sK18(X0,X1,bool))) )),
  inference(superposition,[],[f420,f343])).
fof(f2401,plain,(
  ( ! [X6,X8,X7,X5,X9] : (p(s(bool,happ(s(fun(X5,bool),X7),s(X5,X9)))) | ~p(s(bool,happ(s(fun(X5,bool),X7),s(X5,sK18('const.bool.COND_3'(s(bool,'F_0'),s(fun(X5,bool),X6),s(fun(X5,bool),X7)),X8,X5))))) | ~p(s(bool,happ(s(fun(X5,bool),X8),s(X5,sK17('const.bool.COND_3'(s(bool,'F_0'),s(fun(X5,bool),X6),s(fun(X5,bool),X7)),X8,X5)))))) )),
  inference(forward_demodulation,[],[f2397,f414])).
fof(f2397,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~p(s(bool,happ(s(fun(X5,bool),X7),s(X5,sK18('const.bool.COND_3'(s(bool,'F_0'),s(fun(X5,bool),X6),s(fun(X5,bool),X7)),X8,X5))))) | p(s(bool,happ(s(fun(X5,bool),'const.bool.COND_3'(s(bool,'F_0'),s(fun(X5,bool),X6),s(fun(X5,bool),X7))),s(X5,X9)))) | ~p(s(bool,happ(s(fun(X5,bool),X8),s(X5,sK17('const.bool.COND_3'(s(bool,'F_0'),s(fun(X5,bool),X6),s(fun(X5,bool),X7)),X8,X5)))))) )),
  inference(superposition,[],[f420,f414])).
fof(f2433,plain,(
  ( ! [X6,X8,X7,X5,X9] : (p(s(bool,happ(s(fun(X5,bool),X7),s(X5,sK19(X9,'const.bool.COND_3'(s(bool,'F_0'),s(fun(X5,bool),X6),s(fun(X5,bool),X7)),X5))))) | ~p(s(bool,happ(s(fun(X5,bool),X7),s(X5,X8)))) | p(s(bool,happ(s(fun(X5,bool),X9),s(X5,sK20(X9,'const.bool.COND_3'(s(bool,'F_0'),s(fun(X5,bool),X6),s(fun(X5,bool),X7)),X5)))))) )),
  inference(forward_demodulation,[],[f2418,f414])).
fof(f2418,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~p(s(bool,happ(s(fun(X5,bool),X7),s(X5,X8)))) | p(s(bool,happ(s(fun(X5,bool),'const.bool.COND_3'(s(bool,'F_0'),s(fun(X5,bool),X6),s(fun(X5,bool),X7))),s(X5,sK19(X9,'const.bool.COND_3'(s(bool,'F_0'),s(fun(X5,bool),X6),s(fun(X5,bool),X7)),X5))))) | p(s(bool,happ(s(fun(X5,bool),X9),s(X5,sK20(X9,'const.bool.COND_3'(s(bool,'F_0'),s(fun(X5,bool),X6),s(fun(X5,bool),X7)),X5)))))) )),
  inference(superposition,[],[f421,f414])).
fof(f2466,plain,(
  ( ! [X6,X8,X7,X5,X9] : (p(s(bool,happ(s(fun(X5,bool),X7),s(X5,sK20('const.bool.COND_3'(s(bool,'F_0'),s(fun(X5,bool),X6),s(fun(X5,bool),X7)),X9,X5))))) | ~p(s(bool,happ(s(fun(X5,bool),X7),s(X5,X8)))) | p(s(bool,happ(s(fun(X5,bool),X9),s(X5,sK19('const.bool.COND_3'(s(bool,'F_0'),s(fun(X5,bool),X6),s(fun(X5,bool),X7)),X9,X5)))))) )),
  inference(forward_demodulation,[],[f2451,f414])).
fof(f2451,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~p(s(bool,happ(s(fun(X5,bool),X7),s(X5,X8)))) | p(s(bool,happ(s(fun(X5,bool),X9),s(X5,sK19('const.bool.COND_3'(s(bool,'F_0'),s(fun(X5,bool),X6),s(fun(X5,bool),X7)),X9,X5))))) | p(s(bool,happ(s(fun(X5,bool),'const.bool.COND_3'(s(bool,'F_0'),s(fun(X5,bool),X6),s(fun(X5,bool),X7))),s(X5,sK20('const.bool.COND_3'(s(bool,'F_0'),s(fun(X5,bool),X6),s(fun(X5,bool),X7)),X9,X5)))))) )),
  inference(superposition,[],[f422,f414])).
fof(f2505,plain,(
  ( ! [X6,X8,X7,X5,X9] : (p(s(bool,happ(s(fun(X5,bool),X7),s(X5,sK21(X9,'const.bool.COND_3'(s(bool,'F_0'),s(fun(X5,bool),X6),s(fun(X5,bool),X7)),X5))))) | ~p(s(bool,happ(s(fun(X5,bool),X7),s(X5,X8)))) | p(s(bool,happ(s(fun(X5,bool),X9),s(X5,sK21(X9,'const.bool.COND_3'(s(bool,'F_0'),s(fun(X5,bool),X6),s(fun(X5,bool),X7)),X5)))))) )),
  inference(forward_demodulation,[],[f2490,f414])).
fof(f2490,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~p(s(bool,happ(s(fun(X5,bool),X7),s(X5,X8)))) | p(s(bool,happ(s(fun(X5,bool),'const.bool.COND_3'(s(bool,'F_0'),s(fun(X5,bool),X6),s(fun(X5,bool),X7))),s(X5,sK21(X9,'const.bool.COND_3'(s(bool,'F_0'),s(fun(X5,bool),X6),s(fun(X5,bool),X7)),X5))))) | p(s(bool,happ(s(fun(X5,bool),X9),s(X5,sK21(X9,'const.bool.COND_3'(s(bool,'F_0'),s(fun(X5,bool),X6),s(fun(X5,bool),X7)),X5)))))) )),
  inference(superposition,[],[f423,f414])).
fof(f2545,plain,(
  ( ! [X6,X8,X7,X5,X9] : (p(s(bool,happ(s(fun(X5,bool),X7),s(X5,sK21('const.bool.COND_3'(s(bool,'F_0'),s(fun(X5,bool),X6),s(fun(X5,bool),X7)),X9,X5))))) | ~p(s(bool,happ(s(fun(X5,bool),X7),s(X5,X8)))) | p(s(bool,happ(s(fun(X5,bool),X9),s(X5,sK21('const.bool.COND_3'(s(bool,'F_0'),s(fun(X5,bool),X6),s(fun(X5,bool),X7)),X9,X5)))))) )),
  inference(forward_demodulation,[],[f2530,f414])).
fof(f2530,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~p(s(bool,happ(s(fun(X5,bool),X7),s(X5,X8)))) | p(s(bool,happ(s(fun(X5,bool),X9),s(X5,sK21('const.bool.COND_3'(s(bool,'F_0'),s(fun(X5,bool),X6),s(fun(X5,bool),X7)),X9,X5))))) | p(s(bool,happ(s(fun(X5,bool),'const.bool.COND_3'(s(bool,'F_0'),s(fun(X5,bool),X6),s(fun(X5,bool),X7))),s(X5,sK21('const.bool.COND_3'(s(bool,'F_0'),s(fun(X5,bool),X6),s(fun(X5,bool),X7)),X9,X5)))))) )),
  inference(superposition,[],[f424,f414])).
fof(f2765,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X2),s('type.num.num',X1))) = s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X0),s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X1),s('type.num.num',X2)))))) )),
  inference(forward_demodulation,[],[f2764,f414])).
fof(f2764,plain,(
  ( ! [X4,X2,X0,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X0),s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X1),s('type.num.num',X2))))) = s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X4),s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X2),s('type.num.num',X1)))))) )),
  inference(forward_demodulation,[],[f2738,f414])).
fof(f2738,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X0),s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X1),s('type.num.num',X2))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,'F_0'))),s('type.num.num',X4),s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X2),s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f592,f430,f414,f427])).
fof(f2770,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X2),s('type.num.num',X1))) = s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X0),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X2)))))) )),
  inference(forward_demodulation,[],[f2769,f414])).
fof(f2769,plain,(
  ( ! [X4,X2,X0,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X0),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X2))))) = s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X4),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X2),s('type.num.num',X1)))))) )),
  inference(forward_demodulation,[],[f2736,f414])).
fof(f2736,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X0),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X2))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,'F_0'))),s('type.num.num',X4),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X2),s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f592,f429,f414,f427])).
fof(f2847,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(X2,X3) != s(X2,X4) | s(X2,'const.bool.COND_3'(s(bool,'F_0'),s(X2,X6),s(X2,X4))) = s(X2,'const.bool.COND_3'(s(bool,X1),s(X2,X5),s(X2,X3))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f2747,f592])).
fof(f2747,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(X2,X3) != s(X2,X4) | p(s(bool,'F_0')) | s(X2,'const.bool.COND_3'(s(bool,'F_0'),s(X2,X6),s(X2,X4))) = s(X2,'const.bool.COND_3'(s(bool,X1),s(X2,X5),s(X2,X3))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f427,f343])).
fof(f2848,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (s(bool,'F_0') != s(bool,X8) | s(X9,X10) != s(X9,X11) | p(s(bool,X7)) | s(X9,'const.bool.COND_3'(s(bool,X7),s(X9,X13),s(X9,X11))) = s(X9,'const.bool.COND_3'(s(bool,X8),s(X9,X12),s(X9,X10)))) )),
  inference(subsumption_resolution,[],[f2748,f354])).
fof(f2748,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (s(bool,'F_0') != s(bool,X8) | s(X9,X10) != s(X9,X11) | p(s(bool,X7)) | s(X9,'const.bool.COND_3'(s(bool,X7),s(X9,X13),s(X9,X11))) = s(X9,'const.bool.COND_3'(s(bool,X8),s(X9,X12),s(X9,X10))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f427,f343])).
fof(f468,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f286])).
fof(f477,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(ennf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0,X1] : ~(p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) & p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X6,X7] : ~(p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num',X7)))) & p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num',X6)))))),
  file('thm.arithmetic.MIN_MAX_EQ','thm.arithmetic.LESS_ANTISYM')).
fof(f483,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X2))))) )),
  inference(cnf_transformation,[],[f200])).
fof(f200,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X2)))))),
  inference(flattening,[],[f199])).
fof(f199,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X2))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X2)))))),
  inference(ennf_transformation,[],[f155])).
fof(f155,plain,(
  ! [X0,X1,X2] : ((p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) & p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X2))))) => p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X2)))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X6,X7,X8] : ((p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num',X7)))) & p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num',X8))))) => p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num',X8)))))),
  file('thm.arithmetic.MIN_MAX_EQ','thm.arithmetic.LESS_TRANS')).
fof(f2088,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.num.num',X2)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num',X3)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X3))))) )),
  inference(superposition,[],[f483,f430])).
fof(f2090,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X9),s('type.num.num',X8))),s('type.num.num',X10)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X10),s('type.num.num',X11)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X8),s('type.num.num',X9))),s('type.num.num',X11))))) )),
  inference(superposition,[],[f483,f429])).
fof(f2097,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X1),s('type.num.num',X0)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X3)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num',X3))))) )),
  inference(superposition,[],[f483,f430])).
fof(f2099,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X10),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X9),s('type.num.num',X8)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X8),s('type.num.num',X9))),s('type.num.num',X11)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X10),s('type.num.num',X11))))) )),
  inference(superposition,[],[f483,f429])).
fof(f1283,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X2),s('type.num.num',X1))),s('type.num.num',X3)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X2))))))) )),
  inference(superposition,[],[f477,f429])).
fof(f1254,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X2),s('type.num.num',X1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num',X3))))) )),
  inference(superposition,[],[f468,f429])).
fof(f1210,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X2),s('type.num.num',X1))),s('type.num.num',X3)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X2))))))) )),
  inference(superposition,[],[f467,f429])).
fof(f1285,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X8),s('type.num.num',X7))),s('type.num.num',X9)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X7),s('type.num.num',X8))))))) )),
  inference(superposition,[],[f477,f430])).
fof(f1256,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X8),s('type.num.num',X7)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X7),s('type.num.num',X8))),s('type.num.num',X9))))) )),
  inference(superposition,[],[f468,f430])).
fof(f1212,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X8),s('type.num.num',X7))),s('type.num.num',X9)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X7),s('type.num.num',X8))))))) )),
  inference(superposition,[],[f467,f430])).
fof(f1247,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X2),s('type.num.num',X1))),s('type.num.num',X3)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X2))))))) )),
  inference(superposition,[],[f468,f429])).
fof(f1217,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X2),s('type.num.num',X1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num',X3))))) )),
  inference(superposition,[],[f467,f429])).
fof(f1249,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X8),s('type.num.num',X7))),s('type.num.num',X9)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X9),s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X7),s('type.num.num',X8))))))) )),
  inference(superposition,[],[f468,f430])).
fof(f1219,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X9),s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X8),s('type.num.num',X7)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X7),s('type.num.num',X8))),s('type.num.num',X9))))) )),
  inference(superposition,[],[f467,f430])).
fof(f484,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X2))))) )),
  inference(cnf_transformation,[],[f293])).
fof(f487,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X2))))) )),
  inference(cnf_transformation,[],[f293])).
fof(f488,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X2)))) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f293])).
fof(f489,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X2)) | p(s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X2))))) )),
  inference(cnf_transformation,[],[f293])).
fof(f1492,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.COND_3'(s(bool,X1),s(bool,X2),s(bool,X0)))) | p(s(bool,X1)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f1477,f592])).
fof(f1477,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.COND_3'(s(bool,X1),s(bool,X2),s(bool,X0)))) | p(s(bool,X1)) | p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f485,f343])).
fof(f1542,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X0)))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f1528,f592])).
fof(f1528,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X0)))) | p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1494,f343])).
fof(f433,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | s('type.num.num',X0) = s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f186])).
fof(f186,plain,(
  ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | s('type.num.num',X0) = s('type.num.num',X1))),
  inference(flattening,[],[f185])).
fof(f185,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) | s('type.num.num',X0) = s('type.num.num',X1))),
  inference(ennf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0,X1] : ((p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) => s('type.num.num',X0) = s('type.num.num',X1))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X7,X6] : ((p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',X6)))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num',X7))))) => s('type.num.num',X6) = s('type.num.num',X7))),
  file('thm.arithmetic.MIN_MAX_EQ','thm.arithmetic.LESS_EQUAL_ANTISYM')).
fof(f482,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X2))))) )),
  inference(cnf_transformation,[],[f198])).
fof(f198,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X2)))))),
  inference(flattening,[],[f197])).
fof(f197,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X2)))))),
  inference(ennf_transformation,[],[f154])).
fof(f154,plain,(
  ! [X0,X1,X2] : ((p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2))))) => p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X2)))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X6,X7,X8] : ((p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num',X7)))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',X8))))) => p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num',X8)))))),
  file('thm.arithmetic.MIN_MAX_EQ','thm.arithmetic.LESS_EQ_TRANS')).
fof(f2041,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.num.num',X2)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X3)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X3))))) )),
  inference(superposition,[],[f482,f430])).
fof(f2043,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X9),s('type.num.num',X8))),s('type.num.num',X10)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X10),s('type.num.num',X11)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X8),s('type.num.num',X9))),s('type.num.num',X11))))) )),
  inference(superposition,[],[f482,f429])).
fof(f2050,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X1),s('type.num.num',X0)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X3)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X3))))) )),
  inference(superposition,[],[f482,f430])).
fof(f2052,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X10),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X9),s('type.num.num',X8)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X8),s('type.num.num',X9))),s('type.num.num',X11)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X10),s('type.num.num',X11))))) )),
  inference(superposition,[],[f482,f429])).
fof(f1859,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X2),s('type.num.num',X1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num',X3)))) | s('type.num.num',X3) = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X2)))) )),
  inference(superposition,[],[f433,f429])).
fof(f1852,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X2),s('type.num.num',X1))),s('type.num.num',X3)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X2)))))) | s('type.num.num',X3) = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X2)))) )),
  inference(superposition,[],[f433,f429])).
fof(f1861,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X9),s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X8),s('type.num.num',X7)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X7),s('type.num.num',X8))),s('type.num.num',X9)))) | s('type.num.num',X9) = s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X7),s('type.num.num',X8)))) )),
  inference(superposition,[],[f433,f430])).
fof(f1854,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X8),s('type.num.num',X7))),s('type.num.num',X9)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X9),s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X7),s('type.num.num',X8)))))) | s('type.num.num',X9) = s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X7),s('type.num.num',X8)))) )),
  inference(superposition,[],[f433,f430])).
fof(f411,plain,(
  ( ! [X0,X5,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK14(X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))) )),
  inference(cnf_transformation,[],[f264])).
fof(f264,plain,(
  ! [X0,X1] : ((! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK14(X1,X0)))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK15(X1,X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK14,sK15])],[f263])).
fof(f263,plain,(
  ! [X0,X1] : ((! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(rectify,[],[f262])).
fof(f262,plain,(
  ! [X0,X1] : ((! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(nnf_transformation,[],[f182])).
fof(f182,plain,(
  ! [X0,X1] : (? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(ennf_transformation,[],[f114])).
fof(f114,plain,(
  ! [X0,X1] : (~! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(flattening,[],[f113])).
fof(f113,plain,(
  ! [X0] : ! [X1] : (~! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(rectify,[],[f49])).
fof(f49,axiom,(
  ! [X12] : ! [X16] : (~! [X13] : p(s(bool,happ(s(fun(X12,bool),X16),s(X12,X13)))) <=> ? [X13] : ~p(s(bool,happ(s(fun(X12,bool),X16),s(X12,X13)))))),
  file('thm.arithmetic.MIN_MAX_EQ','thm.bool.NOT_FORALL_THM')).
fof(f412,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK15(X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f264])).
fof(f417,plain,(
  ( ! [X2,X0,X7,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK16(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X7)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK16(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f269])).
fof(f269,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK16(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK16(X2,X1,X0)))))) | (! [X7] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X7)))) & ! [X6] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK17(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK18(X2,X1,X0))))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK16,sK17,sK18])],[f268])).
fof(f268,plain,(
  ! [X0,X1,X2] : ((? [X8] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X8)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X8))))) | (! [X7] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X7)))) & ! [X6] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))))) & (? [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(rectify,[],[f267])).
fof(f267,plain,(
  ! [X0,X1,X2] : ((? [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (? [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(flattening,[],[f266])).
fof(f266,plain,(
  ! [X0,X1,X2] : ((? [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & ((? [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0,X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f121])).
fof(f121,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f48])).
fof(f48,axiom,(
  ! [X12] : ! [X16,X17] : (! [X13] : (p(s(bool,happ(s(fun(X12,bool),X16),s(X12,X13)))) & p(s(bool,happ(s(fun(X12,bool),X17),s(X12,X13))))) <=> (! [X13] : p(s(bool,happ(s(fun(X12,bool),X16),s(X12,X13)))) & ! [X13] : p(s(bool,happ(s(fun(X12,bool),X17),s(X12,X13))))))),
  file('thm.arithmetic.MIN_MAX_EQ','thm.bool.FORALL_AND_THM')).
fof(f418,plain,(
  ( ! [X6,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK16(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK16(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f269])).
fof(f419,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK18(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK17(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f269])).
fof(f420,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK18(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK17(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f269])).
fof(f421,plain,(
  ( ! [X2,X0,X8,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X8)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK19(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK20(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f273])).
fof(f273,plain,(
  ! [X0,X1,X2] : ((! [X8] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X8)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X8))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK19(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK20(X2,X1,X0)))))) & ((! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK21(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK21(X2,X1,X0))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK19,sK20,sK21])],[f272])).
fof(f272,plain,(
  ! [X0,X1,X2] : ((! [X8] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X8)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X8))))) | ? [X7] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X7)))) | ? [X6] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) & ((! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(rectify,[],[f271])).
fof(f271,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) & ((! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(flattening,[],[f270])).
fof(f270,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & ((! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f123])).
fof(f123,plain,(
  ! [X0] : ! [X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f51])).
fof(f51,axiom,(
  ! [X12] : ! [X16,X17] : (? [X13] : (p(s(bool,happ(s(fun(X12,bool),X16),s(X12,X13)))) | p(s(bool,happ(s(fun(X12,bool),X17),s(X12,X13))))) <=> (? [X13] : p(s(bool,happ(s(fun(X12,bool),X16),s(X12,X13)))) | ? [X13] : p(s(bool,happ(s(fun(X12,bool),X17),s(X12,X13))))))),
  file('thm.arithmetic.MIN_MAX_EQ','thm.bool.EXISTS_OR_THM')).
fof(f422,plain,(
  ( ! [X2,X0,X8,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X8)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK19(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK20(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f273])).
fof(f423,plain,(
  ( ! [X2,X0,X5,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK21(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK21(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f273])).
fof(f424,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK21(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK21(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f273])).
fof(f1818,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK14('const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2)),X0)))))) )),
  inference(forward_demodulation,[],[f1814,f413])).
fof(f1814,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK14('const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2)),X0))))) | p(s(bool,happ(s(fun(X0,bool),'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2))),s(X0,X3))))) )),
  inference(superposition,[],[f411,f413])).
fof(f1837,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK15('const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2)),X0)))))) )),
  inference(forward_demodulation,[],[f1833,f413])).
fof(f1833,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK15('const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2)),X0))))) | p(s(bool,happ(s(fun(X0,bool),'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2))),s(X0,X3))))) )),
  inference(superposition,[],[f412,f413])).
fof(f2317,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK16('const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2)),X3,X0))))) | p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,sK16('const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2)),X3,X0)))))) )),
  inference(superposition,[],[f417,f413])).
fof(f2351,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK16('const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2)),X3,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,sK16('const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2)),X3,X0)))))) )),
  inference(forward_demodulation,[],[f2347,f413])).
fof(f2347,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK16('const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2)),X3,X0))))) | p(s(bool,happ(s(fun(X0,bool),'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2))),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,sK16('const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2)),X3,X0)))))) )),
  inference(superposition,[],[f418,f413])).
fof(f2371,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK18('const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2)),X3,X0))))) | p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,sK17('const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2)),X3,X0)))))) )),
  inference(superposition,[],[f419,f413])).
fof(f2400,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK18('const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2)),X3,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,sK17('const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2)),X3,X0)))))) )),
  inference(forward_demodulation,[],[f2396,f413])).
fof(f2396,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK18('const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2)),X3,X0))))) | p(s(bool,happ(s(fun(X0,bool),'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2))),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,sK17('const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2)),X3,X0)))))) )),
  inference(superposition,[],[f420,f413])).
fof(f2432,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK19(X4,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2)),X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X4),s(X0,sK20(X4,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2)),X0)))))) )),
  inference(forward_demodulation,[],[f2417,f413])).
fof(f2417,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2))),s(X0,sK19(X4,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2)),X0))))) | p(s(bool,happ(s(fun(X0,bool),X4),s(X0,sK20(X4,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2)),X0)))))) )),
  inference(superposition,[],[f421,f413])).
fof(f2465,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK20('const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2)),X4,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X4),s(X0,sK19('const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2)),X4,X0)))))) )),
  inference(forward_demodulation,[],[f2450,f413])).
fof(f2450,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X4),s(X0,sK19('const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2)),X4,X0))))) | p(s(bool,happ(s(fun(X0,bool),'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2))),s(X0,sK20('const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2)),X4,X0)))))) )),
  inference(superposition,[],[f422,f413])).
fof(f2504,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK21(X4,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2)),X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X4),s(X0,sK21(X4,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2)),X0)))))) )),
  inference(forward_demodulation,[],[f2489,f413])).
fof(f2489,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2))),s(X0,sK21(X4,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2)),X0))))) | p(s(bool,happ(s(fun(X0,bool),X4),s(X0,sK21(X4,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2)),X0)))))) )),
  inference(superposition,[],[f423,f413])).
fof(f2544,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK21('const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2)),X4,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X4),s(X0,sK21('const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2)),X4,X0)))))) )),
  inference(forward_demodulation,[],[f2529,f413])).
fof(f2529,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X4),s(X0,sK21('const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2)),X4,X0))))) | p(s(bool,happ(s(fun(X0,bool),'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2))),s(X0,sK21('const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2)),X4,X0)))))) )),
  inference(superposition,[],[f424,f413])).
fof(f2751,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,X1) != s(X0,X2) | p(s(bool,X3)) | s(X0,'const.bool.COND_3'(s(bool,X3),s(X0,X4),s(X0,X1))) = s(X0,'const.bool.COND_3'(s(bool,X3),s(X0,X5),s(X0,X2)))) )),
  inference(equality_resolution,[],[f427])).
fof(f2844,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (s(X9,X10) != s(X9,X11) | p(s(bool,X8)) | s(X9,'const.bool.COND_3'(s(bool,X7),s(X9,X12),s(X9,X10))) = s(X9,'const.bool.COND_3'(s(bool,X8),s(X9,X13),s(X9,X11))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f2744,f359])).
fof(f2744,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (s(bool,'F_0') != s(bool,X8) | s(X9,X10) != s(X9,X11) | p(s(bool,X8)) | s(X9,'const.bool.COND_3'(s(bool,X7),s(X9,X12),s(X9,X10))) = s(X9,'const.bool.COND_3'(s(bool,X8),s(X9,X13),s(X9,X11))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f427,f343])).
fof(f2622,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X0))) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X2)))) )),
  inference(forward_demodulation,[],[f2621,f413])).
fof(f2621,plain,(
  ( ! [X4,X2,X0,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X2))) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.num.num',X4)))) )),
  inference(forward_demodulation,[],[f2593,f414])).
fof(f2593,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X2))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,'T_0'))),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.num.num',X4)))) )),
  inference(unit_resulting_resolution,[],[f335,f429,f414,f426])).
fof(f2617,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X1),s('type.num.num',X0))) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X2)))) )),
  inference(forward_demodulation,[],[f2616,f413])).
fof(f2616,plain,(
  ( ! [X4,X2,X0,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X2))) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.num.num',X4)))) )),
  inference(forward_demodulation,[],[f2595,f414])).
fof(f2595,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X2))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,'T_0'))),s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.num.num',X4)))) )),
  inference(unit_resulting_resolution,[],[f335,f430,f414,f426])).
fof(f2270,plain,(
  ( ! [X0,X1] : (~sP3('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X1)),'const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),'F_0')) )),
  inference(forward_demodulation,[],[f2260,f1595])).
fof(f2260,plain,(
  ( ! [X0,X1] : (~sP3('const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0'))),s(bool,X1)),'const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f592,f775,f599,f1659,f508])).
fof(f2262,plain,(
  ( ! [X0] : (sP5('F_0','F_0','const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f592,f602,f599,f1659,f507])).
fof(f2222,plain,(
  ( ! [X0] : (sP5('F_0','const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f592,f601,f599,f1658,f508])).
fof(f2252,plain,(
  ( ! [X0,X1] : (~sP3('const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X1)),'F_0')) )),
  inference(forward_demodulation,[],[f2225,f1595])).
fof(f2225,plain,(
  ( ! [X0,X1] : (~sP3('const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),'const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0'))),s(bool,X1)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f592,f777,f599,f1658,f508])).
fof(f2170,plain,(
  ( ! [X4,X5,X3] : (s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X4),s('type.num.num',X5))))) = s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X5),s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X3),s('type.num.num',X4)))))) )),
  inference(superposition,[],[f481,f430])).
fof(f2157,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X1),s('type.num.num',X2))))) = s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.num.num',X2)))) )),
  inference(superposition,[],[f481,f430])).
fof(f2128,plain,(
  ( ! [X0] : (sP5('const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f592,f602,f598,f1657,f507])).
fof(f2118,plain,(
  ( ! [X0] : (sP5('const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),'T_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f335,f605,f597,f1656,f506])).
fof(f2119,plain,(
  ( ! [X0,X1] : (sP5('const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),'const.bool.COND_3'(s(bool,X1),s(bool,'T_0'),s(bool,'T_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f335,f605,f1647,f1656,f506])).
fof(f2082,plain,(
  ( ! [X0] : (~sP5('F_0','T_0','const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f1652,f503])).
fof(f2073,plain,(
  ( ! [X0] : (~sP5('F_0','const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f1648,f503])).
fof(f2030,plain,(
  ( ! [X0] : (sP5('T_0','const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f335,f605,f600,f1647,f506])).
fof(f2015,plain,(
  ( ! [X0] : (~sP5('T_0','F_0','const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f1639,f504])).
fof(f1981,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X0),s('type.num.num',X0))) = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X0),s('type.num.num',X0)))) )),
  inference(superposition,[],[f432,f431])).
fof(f1595,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f335,f1582,f437])).
fof(f1596,plain,(
  ( ! [X0,X1] : (s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0'))) = s(bool,'const.bool.COND_3'(s(bool,X1),s(bool,'T_0'),s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f1582,f1582,f437])).
fof(f1597,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.bool.COND_3'(s(bool,X1),s(bool,'T_0'),s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f1582,f443])).
fof(f1598,plain,(
  ( ! [X0,X1] : (sP0('const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f1582,f444])).
fof(f1599,plain,(
  ( ! [X0] : (~sP1('const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f335,f1582,f458])).
fof(f1602,plain,(
  ( ! [X0] : (~sP1('T_0','const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f335,f1582,f458])).
fof(f1603,plain,(
  ( ! [X0,X1] : (~sP1('const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),'const.bool.COND_3'(s(bool,X1),s(bool,'T_0'),s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f1582,f1582,f458])).
fof(f1604,plain,(
  ( ! [X0] : (sP1('const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f592,f1582,f459])).
fof(f1605,plain,(
  ( ! [X0] : (sP1('F_0','const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f592,f1582,f462])).
fof(f1606,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0'))),s(bool,'F_0'),s(bool,X1))))) )),
  inference(unit_resulting_resolution,[],[f592,f1582,f484])).
fof(f1608,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0'))),s(bool,'T_0'),s(bool,X1))))) )),
  inference(unit_resulting_resolution,[],[f335,f1582,f487])).
fof(f1611,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0'))),s(bool,'const.bool.COND_3'(s(bool,X1),s(bool,'T_0'),s(bool,'T_0'))),s(bool,X2))))) )),
  inference(unit_resulting_resolution,[],[f1582,f1582,f487])).
fof(f1613,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,X1),s(bool,'T_0'),s(bool,'T_0'))))))) )),
  inference(unit_resulting_resolution,[],[f335,f1582,f489])).
fof(f1615,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,'const.bool.COND_3'(s(bool,X1),s(bool,'T_0'),s(bool,'T_0'))),s(bool,'T_0'))))) )),
  inference(unit_resulting_resolution,[],[f335,f1582,f489])).
fof(f1616,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,'const.bool.COND_3'(s(bool,X1),s(bool,'T_0'),s(bool,'T_0'))),s(bool,'const.bool.COND_3'(s(bool,X2),s(bool,'T_0'),s(bool,'T_0'))))))) )),
  inference(unit_resulting_resolution,[],[f1582,f1582,f489])).
fof(f1619,plain,(
  ( ! [X0] : (~sP5('T_0','const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f335,f592,f1582,f505])).
fof(f1620,plain,(
  ( ! [X0,X1] : (~sP5('const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),'const.bool.COND_3'(s(bool,X1),s(bool,'T_0'),s(bool,'T_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1582,f592,f1582,f505])).
fof(f1634,plain,(
  ( ! [X0] : (sP5('T_0','T_0','const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f603,f597,f600,f1582,f506])).
fof(f1635,plain,(
  ( ! [X0] : (~sP4('const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f335,f592,f1582,f509])).
fof(f1639,plain,(
  ( ! [X0] : (~sP4('T_0','const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f335,f592,f1582,f509])).
fof(f1640,plain,(
  ( ! [X0,X1] : (~sP4('const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),'const.bool.COND_3'(s(bool,X1),s(bool,'T_0'),s(bool,'T_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1582,f592,f1582,f509])).
fof(f1647,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,'const.bool.COND_3'(s(bool,X2),s(bool,'T_0'),s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f1582,f510])).
fof(f1648,plain,(
  ( ! [X0] : (~sP3('const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f335,f592,f1582,f513])).
fof(f1652,plain,(
  ( ! [X0] : (~sP3('T_0','const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f335,f592,f1582,f513])).
fof(f1653,plain,(
  ( ! [X0,X1] : (~sP3('const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),'const.bool.COND_3'(s(bool,X1),s(bool,'T_0'),s(bool,'T_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1582,f592,f1582,f513])).
fof(f1656,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,'const.bool.COND_3'(s(bool,X2),s(bool,'T_0'),s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f1582,f514])).
fof(f1657,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,'const.bool.COND_3'(s(bool,X2),s(bool,'T_0'),s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f1582,f518])).
fof(f1658,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,'const.bool.COND_3'(s(bool,X1),s(bool,'T_0'),s(bool,'T_0')),X2)) )),
  inference(unit_resulting_resolution,[],[f1582,f519])).
fof(f1659,plain,(
  ( ! [X2,X0,X1] : (sP2('const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),X1,X2)) )),
  inference(unit_resulting_resolution,[],[f1582,f520])).
fof(f1660,plain,(
  ( ! [X2,X0,X1] : (sP5('const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2)),X1)) )),
  inference(unit_resulting_resolution,[],[f413,f1582,f522])).
fof(f1662,plain,(
  ( ! [X2,X0,X1] : (sP5('const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),X1,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2)))) )),
  inference(unit_resulting_resolution,[],[f413,f1582,f522])).
fof(f1664,plain,(
  ( ! [X0] : (~sP5('const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f615,f1582,f523])).
fof(f1665,plain,(
  ( ! [X0] : (~sP5('const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f615,f1582,f523])).
fof(f1666,plain,(
  ( ! [X0,X1] : (~sP7('const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f592,f1582,f526])).
fof(f1667,plain,(
  ( ! [X0,X1] : (~sP7('const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f592,f1582,f527])).
fof(f1674,plain,(
  ( ! [X0] : (~sP6('const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f335,f592,f1582,f532])).
fof(f1678,plain,(
  ( ! [X0] : (~sP6('T_0','const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f335,f592,f1582,f532])).
fof(f1679,plain,(
  ( ! [X0,X1] : (~sP6('const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),'const.bool.COND_3'(s(bool,X1),s(bool,'T_0'),s(bool,'T_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1582,f592,f1582,f532])).
fof(f1682,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,'const.bool.COND_3'(s(bool,X2),s(bool,'T_0'),s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f1582,f533])).
fof(f1683,plain,(
  ( ! [X0] : (sP7('T_0','const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f335,f335,f1582,f538])).
fof(f1688,plain,(
  ( ! [X0] : (sP7('const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),'T_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f335,f335,f1582,f538])).
fof(f1689,plain,(
  ( ! [X0,X1] : (sP7('const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),'const.bool.COND_3'(s(bool,X1),s(bool,'T_0'),s(bool,'T_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f1582,f335,f1582,f538])).
fof(f1692,plain,(
  ( ! [X0] : (sP7('T_0','T_0','const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f335,f335,f1582,f538])).
fof(f1693,plain,(
  ( ! [X0,X1] : (sP7('const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),'T_0','const.bool.COND_3'(s(bool,X1),s(bool,'T_0'),s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f1582,f335,f1582,f538])).
fof(f1694,plain,(
  ( ! [X0,X1] : (sP7('T_0','const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),'const.bool.COND_3'(s(bool,X1),s(bool,'T_0'),s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f335,f1582,f1582,f538])).
fof(f1695,plain,(
  ( ! [X2,X0,X1] : (sP7('const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),'const.bool.COND_3'(s(bool,X1),s(bool,'T_0'),s(bool,'T_0')),'const.bool.COND_3'(s(bool,X2),s(bool,'T_0'),s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f1582,f1582,f1582,f538])).
fof(f1697,plain,(
  ( ! [X0] : (~sP7('F_0','const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f335,f592,f1582,f541])).
fof(f1700,plain,(
  ( ! [X0] : (~sP7('F_0','T_0','const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f335,f592,f1582,f541])).
fof(f1701,plain,(
  ( ! [X0,X1] : (~sP7('F_0','const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),'const.bool.COND_3'(s(bool,X1),s(bool,'T_0'),s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f1582,f592,f1582,f541])).
fof(f1703,plain,(
  ( ! [X0,X1] : (~sP9('F_0',X0,'const.bool.COND_3'(s(bool,X1),s(bool,'T_0'),s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f592,f1582,f542])).
fof(f1705,plain,(
  ( ! [X0,X1] : (~sP9('F_0','const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f592,f1582,f543])).
fof(f1709,plain,(
  ( ! [X0] : (~sP8('const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f592,f592,f1582,f549])).
fof(f1712,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,'const.bool.COND_3'(s(bool,X2),s(bool,'T_0'),s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f1582,f550])).
fof(f1713,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,'const.bool.COND_3'(s(bool,X1),s(bool,'T_0'),s(bool,'T_0')),X2)) )),
  inference(unit_resulting_resolution,[],[f1582,f551])).
fof(f1716,plain,(
  ( ! [X0,X1] : (sP9('T_0',X0,'const.bool.COND_3'(s(bool,X1),s(bool,'T_0'),s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f335,f1582,f554])).
fof(f1717,plain,(
  ( ! [X2,X0,X1] : (sP9('const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),X1,'const.bool.COND_3'(s(bool,X2),s(bool,'T_0'),s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f1582,f1582,f554])).
fof(f1719,plain,(
  ( ! [X0,X1] : (sP9('T_0','const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f335,f1582,f555])).
fof(f1723,plain,(
  ( ! [X2,X0,X1] : (sP9('const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),'const.bool.COND_3'(s(bool,X1),s(bool,'T_0'),s(bool,'T_0')),X2)) )),
  inference(unit_resulting_resolution,[],[f1582,f1582,f555])).
fof(f1724,plain,(
  ( ! [X0] : (~sP9('const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f592,f592,f1582,f556])).
fof(f1725,plain,(
  ( ! [X0,X1] : (~sP11('F_0','const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f592,f1582,f560])).
fof(f1737,plain,(
  ( ! [X0] : (~sP10('const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f592,f335,f1582,f566])).
fof(f1741,plain,(
  ( ! [X0] : (~sP10('T_0','F_0','const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f335,f592,f1582,f566])).
fof(f1742,plain,(
  ( ! [X0,X1] : (~sP10('const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),'F_0','const.bool.COND_3'(s(bool,X1),s(bool,'T_0'),s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f1582,f592,f1582,f566])).
fof(f1745,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,'const.bool.COND_3'(s(bool,X1),s(bool,'T_0'),s(bool,'T_0')),X2)) )),
  inference(unit_resulting_resolution,[],[f1582,f568])).
fof(f1746,plain,(
  ( ! [X0] : (sP11('F_0','F_0','const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f592,f592,f1582,f570])).
fof(f1749,plain,(
  ( ! [X0,X1] : (sP11('T_0','const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f335,f1582,f572])).
fof(f1753,plain,(
  ( ! [X2,X0,X1] : (sP11('const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),'const.bool.COND_3'(s(bool,X1),s(bool,'T_0'),s(bool,'T_0')),X2)) )),
  inference(unit_resulting_resolution,[],[f1582,f1582,f572])).
fof(f1754,plain,(
  ( ! [X0] : (~sP11('const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f592,f335,f1582,f573])).
fof(f1756,plain,(
  ( ! [X0] : (~sP11('T_0','F_0','const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f592,f335,f1582,f573])).
fof(f1757,plain,(
  ( ! [X0,X1] : (~sP11('const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),'F_0','const.bool.COND_3'(s(bool,X1),s(bool,'T_0'),s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f592,f1582,f1582,f573])).
fof(f1759,plain,(
  ( ! [X0,X1] : (sP11('const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f612,f1582,f576])).
fof(f1760,plain,(
  ( ! [X0,X1] : (sP11('const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f613,f1582,f576])).
fof(f1761,plain,(
  ( ! [X0,X1] : (sP9('const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f609,f1582,f577])).
fof(f1762,plain,(
  ( ! [X0,X1] : (sP9('const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f610,f1582,f577])).
fof(f1763,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,X1),s(bool,'T_0'),s(bool,'T_0'))),s(bool,X2)))) )),
  inference(unit_resulting_resolution,[],[f1582,f796])).
fof(f1764,plain,(
  ( ! [X2,X0,X1] : (sP0('const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0'))),s(bool,X1)),X2)) )),
  inference(unit_resulting_resolution,[],[f1582,f797])).
fof(f1765,plain,(
  ( ! [X0,X1] : (sP1('const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0'))),s(bool,X1)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f592,f1582,f798])).
fof(f1766,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,'const.bool.COND_3'(s(bool,X2),s(bool,'T_0'),s(bool,'T_0')))))) )),
  inference(unit_resulting_resolution,[],[f1582,f861])).
fof(f1767,plain,(
  ( ! [X2,X0,X1] : (sP0('const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'const.bool.COND_3'(s(bool,X1),s(bool,'T_0'),s(bool,'T_0')))),X2)) )),
  inference(unit_resulting_resolution,[],[f1582,f862])).
fof(f1582,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0'))))) )),
  inference(unit_resulting_resolution,[],[f335,f335,f489])).
fof(f1569,plain,(
  ( ! [X0] : (sP1('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f335,f592,f798])).
fof(f1516,plain,(
  ( ! [X0,X1] : (~sP2('F_0','const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X0)),'F_0') | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1411,f343])).
fof(f1512,plain,(
  ( ! [X0] : (~sP2('F_0','const.bool.COND_3'(s(bool,X0),s(bool,X0),s(bool,X0)),'F_0') | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1411,f343])).
fof(f1501,plain,(
  ( ! [X0,X1] : (~sP2('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)),'F_0') | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1410,f343])).
fof(f1461,plain,(
  ( ! [X0,X1] : (~sP3('T_0','const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,'T_0')),'F_0') | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1403,f343])).
fof(f1437,plain,(
  ( ! [X0,X1] : (~sP2('const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X0)),'F_0','F_0') | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1366,f343])).
fof(f1433,plain,(
  ( ! [X0] : (~sP2('const.bool.COND_3'(s(bool,X0),s(bool,X0),s(bool,X0)),'F_0','F_0') | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1366,f343])).
fof(f1422,plain,(
  ( ! [X0,X1] : (~sP2('const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)),'F_0','F_0') | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1365,f343])).
fof(f1402,plain,(
  ( ! [X0] : (~sP3('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f604,f777,f597,f592,f508])).
fof(f1385,plain,(
  ( ! [X0,X1] : (~sP3('const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,'T_0')),'T_0','F_0') | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1333,f343])).
fof(f1373,plain,(
  sP5('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f603,f598,f601,f592,f507])).
fof(f1332,plain,(
  ( ! [X0] : (~sP3('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f605,f775,f599,f335,f506])).
fof(f1339,plain,(
  sP5('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f605,f597,f600,f335,f506])).
fof(f1275,plain,(
  ( ! [X2,X0,X1] : (sP0('const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,'T_0')),X2) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1267,f343])).
fof(f1239,plain,(
  ( ! [X2,X0,X1] : (sP0(X1,'const.bool.COND_3'(s(bool,X0),s(bool,X2),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1231,f343])).
fof(f1196,plain,(
  ( ! [X0,X1] : (sP0('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)),X1)) )),
  inference(unit_resulting_resolution,[],[f335,f797])).
fof(f1157,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X1)))) )),
  inference(unit_resulting_resolution,[],[f335,f796])).
fof(f1119,plain,(
  ~sP11('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f1114,f561])).
fof(f1114,plain,(
  ~sP10('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f335,f335,f592,f566])).
fof(f1111,plain,(
  sP11('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f614,f335,f592,f565])).
fof(f1087,plain,(
  ~sP9('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1082,f544])).
fof(f1082,plain,(
  ~sP8('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f335,f592,f592,f549])).
fof(f1079,plain,(
  sP9('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f611,f592,f592,f548])).
fof(f1068,plain,(
  ( ! [X0,X1] : (sP4('T_0','const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X0)),'F_0') | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f991,f343])).
fof(f1062,plain,(
  ( ! [X2,X0,X1] : (sP4('T_0','const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X2)),X2) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f991,f343])).
fof(f1050,plain,(
  ( ! [X2,X0,X1] : (sP4('T_0',X1,'const.bool.COND_3'(s(bool,X0),s(bool,X2),s(bool,X1))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f930,f343])).
fof(f1032,plain,(
  ( ! [X0,X1] : (sP4('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)),'F_0') | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f898,f343])).
fof(f1024,plain,(
  ( ! [X6,X4,X5] : (sP4('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X4),s(bool,X5)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X4),s(bool,X6)))) )),
  inference(superposition,[],[f825,f413])).
fof(f1012,plain,(
  ~sP7('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f1007,f525])).
fof(f1007,plain,(
  ~sP6('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f335,f592,f335,f532])).
fof(f1000,plain,(
  ( ! [X0,X1] : (sP5('T_0','F_0','const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f839,f343])).
fof(f994,plain,(
  ( ! [X2,X0,X1] : (sP5('T_0',X1,'const.bool.COND_3'(s(bool,X0),s(bool,X2),s(bool,X1))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f839,f343])).
fof(f984,plain,(
  sP7('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f606,f335,f335,f528])).
fof(f980,plain,(
  ~sP5('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f975,f502])).
fof(f975,plain,(
  ~sP2('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f592,f592,f592,f517])).
fof(f971,plain,(
  ~sP5('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f966,f503])).
fof(f966,plain,(
  ~sP3('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f335,f592,f335,f513])).
fof(f960,plain,(
  ( ! [X0,X1] : (~sP5('F_0','F_0','const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f838,f343])).
fof(f954,plain,(
  ( ! [X2,X0,X1] : (~sP5('F_0',X1,'const.bool.COND_3'(s(bool,X0),s(bool,X2),s(bool,X1))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f838,f343])).
fof(f946,plain,(
  ~sP4('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f335,f592,f335,f509])).
fof(f939,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X0)),'F_0') | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f837,f343])).
fof(f933,plain,(
  ( ! [X2,X0,X1] : (sP5('T_0','const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X2)),X2) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f837,f343])).
fof(f923,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X0)),'F_0') | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f836,f343])).
fof(f917,plain,(
  ( ! [X2,X0,X1] : (~sP5('F_0','const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X2)),X2) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f836,f343])).
fof(f901,plain,(
  ( ! [X0,X1] : (sP5('T_0','F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f778,f343])).
fof(f898,plain,(
  ( ! [X0,X1] : (sP4('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)),X0)) )),
  inference(unit_resulting_resolution,[],[f778,f504])).
fof(f888,plain,(
  ( ! [X0,X1] : (~sP5('F_0','F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f777,f343])).
fof(f840,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.bool.COND_3'(s(bool,X0),s(X1,X3),s(X1,X2))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f414,f343])).
fof(f830,plain,(
  ( ! [X6,X4,X5] : (sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X4),s(bool,X6)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X4),s(bool,X5)))) )),
  inference(superposition,[],[f776,f413])).
fof(f828,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)),'F_0') | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f776,f343])).
fof(f825,plain,(
  ( ! [X0,X1] : (sP4('T_0',X0,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)))) )),
  inference(unit_resulting_resolution,[],[f776,f504])).
fof(f819,plain,(
  ( ! [X6,X4,X5] : (~sP5('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X4),s(bool,X6)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X4),s(bool,X5)))) )),
  inference(superposition,[],[f775,f413])).
fof(f817,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)),'F_0') | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f775,f343])).
fof(f805,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X51)) | ~sP8('const.bool.COND_3'(s(bool,'T_0'),s(bool,X51),s(bool,X52)),X53,X54) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X51),s(bool,X52)),X53,X54)) )),
  inference(superposition,[],[f577,f413])).
fof(f804,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X47)) | ~sP10('const.bool.COND_3'(s(bool,'T_0'),s(bool,X47),s(bool,X48)),X49,X50) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X47),s(bool,X48)),X49,X50)) )),
  inference(superposition,[],[f576,f413])).
fof(f803,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X45)) | ~p(s(bool,X43)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X43),s(bool,X44)),X45,X46)) )),
  inference(superposition,[],[f572,f413])).
fof(f802,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X39)) | p(s(bool,X41)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X39),s(bool,X40)),X42,X41)) )),
  inference(superposition,[],[f571,f413])).
fof(f801,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X37)) | ~p(s(bool,X35)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X35),s(bool,X36)),X37,X38)) )),
  inference(superposition,[],[f555,f413])).
fof(f800,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,X33)) | ~p(s(bool,X31)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X31),s(bool,X32)),X34,X33)) )),
  inference(superposition,[],[f554,f413])).
fof(f799,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,X28)) | p(s(bool,X30)) | sP1(X30,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X28),s(bool,X29)))) )),
  inference(superposition,[],[f462,f413])).
fof(f798,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,X25)) | p(s(bool,X27)) | sP1('const.bool.COND_3'(s(bool,'T_0'),s(bool,X25),s(bool,X26)),X27)) )),
  inference(superposition,[],[f459,f413])).
fof(f797,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | sP0('const.bool.COND_3'(s(bool,'T_0'),s(bool,X22),s(bool,X23)),X24)) )),
  inference(superposition,[],[f444,f413])).
fof(f796,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,X19)) | sP0(X21,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X19),s(bool,X20)))) )),
  inference(superposition,[],[f443,f413])).
fof(f775,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)),X0)) )),
  inference(unit_resulting_resolution,[],[f592,f413,f524])).
fof(f776,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)),X0)) )),
  inference(unit_resulting_resolution,[],[f335,f413,f522])).
fof(f777,plain,(
  ( ! [X0,X1] : (~sP5('F_0',X0,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)))) )),
  inference(unit_resulting_resolution,[],[f592,f413,f524])).
fof(f778,plain,(
  ( ! [X0,X1] : (sP5('T_0',X0,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)))) )),
  inference(unit_resulting_resolution,[],[f335,f413,f522])).
fof(f764,plain,(
  ~sP5('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f335,f615,f523])).
fof(f766,plain,(
  ~sP5('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f335,f615,f523])).
fof(f755,plain,(
  sP5('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f592,f615,f521])).
fof(f756,plain,(
  sP5('F_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f592,f615,f521])).
fof(f677,plain,(
  ( ! [X0] : (sP11('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f335,f335,f572])).
fof(f671,plain,(
  ( ! [X0] : (sP11('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f592,f335,f571])).
fof(f668,plain,(
  ( ! [X0] : (~sP11('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f335,f592,f560])).
fof(f667,plain,(
  ( ! [X0] : (~sP11('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f592,f592,f559])).
fof(f663,plain,(
  ( ! [X0] : (sP9('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f335,f335,f555])).
fof(f656,plain,(
  ( ! [X0] : (sP9('T_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f335,f335,f554])).
fof(f654,plain,(
  ( ! [X0] : (~sP9('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f335,f592,f543])).
fof(f653,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f335,f592,f542])).
fof(f646,plain,(
  ( ! [X0] : (sP7('F_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f592,f592,f537])).
fof(f643,plain,(
  ( ! [X0] : (sP7('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f592,f592,f536])).
fof(f640,plain,(
  ( ! [X0] : (~sP7('T_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f335,f592,f527])).
fof(f637,plain,(
  ( ! [X0] : (~sP7('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f335,f592,f526])).
fof(f636,plain,(
  ~sP0('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f592,f592,f593])).
fof(f631,plain,(
  sP1('F_0','T_0')),
  inference(unit_resulting_resolution,[],[f592,f335,f462])).
fof(f625,plain,(
  sP1('T_0','F_0')),
  inference(unit_resulting_resolution,[],[f592,f335,f459])).
fof(f623,plain,(
  ~sP1('T_0','T_0')),
  inference(unit_resulting_resolution,[],[f335,f335,f458])).
fof(f622,plain,(
  ~sP1('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f592,f592,f457])).
fof(f615,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f592,f354])).
fof(f614,plain,(
  ( ! [X0,X1] : (sP10('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f592,f569])).
fof(f613,plain,(
  ( ! [X0,X1] : (sP10(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f335,f568])).
fof(f612,plain,(
  ( ! [X0,X1] : (sP10(X0,X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f592,f567])).
fof(f611,plain,(
  ( ! [X0,X1] : (sP8('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f592,f552])).
fof(f610,plain,(
  ( ! [X0,X1] : (sP8(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f335,f551])).
fof(f609,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f335,f550])).
fof(f608,plain,(
  ( ! [X0,X1] : (sP6('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f592,f535])).
fof(f607,plain,(
  ( ! [X0,X1] : (sP6(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f592,f534])).
fof(f606,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f335,f533])).
fof(f605,plain,(
  ( ! [X0,X1] : (sP2('T_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f335,f520])).
fof(f604,plain,(
  ( ! [X0,X1] : (sP2(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f335,f519])).
fof(f603,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f335,f518])).
fof(f602,plain,(
  ( ! [X0,X1] : (sP3('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f592,f516])).
fof(f601,plain,(
  ( ! [X0,X1] : (sP3(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f592,f515])).
fof(f600,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f335,f514])).
fof(f599,plain,(
  ( ! [X0,X1] : (sP4('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f592,f512])).
fof(f598,plain,(
  ( ! [X0,X1] : (sP4(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f592,f511])).
fof(f597,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f335,f510])).
fof(f596,plain,(
  ( ! [X0] : (sP0('T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f335,f444])).
fof(f595,plain,(
  ( ! [X0] : (sP0(X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f335,f443])).
fof(f570,plain,(
  ( ! [X2,X0,X1] : (sP11(X2,X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f332])).
fof(f332,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f331])).
fof(f331,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f221])).
fof(f221,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP11(X2,X0,X1))),
  inference(definition_folding,[],[f203,f220,f219])).
fof(f219,plain,(
  ! [X2,X0,X1] : (sP10(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f220,plain,(
  ! [X2,X0,X1] : (sP11(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f203,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f167])).
fof(f167,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f166])).
fof(f166,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X14,X15,X8] : ((p(s(bool,X8)) <=> (p(s(bool,X15)) => p(s(bool,X14)))) <=> ((p(s(bool,X8)) | p(s(bool,X15))) & (p(s(bool,X8)) | ~p(s(bool,X14))) & (~p(s(bool,X15)) | p(s(bool,X14)) | ~p(s(bool,X8)))))),
  file('thm.arithmetic.MIN_MAX_EQ','thm.sat.sat_27_prove_ltb')).
fof(f571,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f332])).
fof(f572,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f332])).
fof(f573,plain,(
  ( ! [X2,X0,X1] : (~sP11(X2,X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f332])).
fof(f566,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f330])).
fof(f330,plain,(
  ! [X0,X1,X2] : ((~sP10(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP10(X0,X1,X2)))),
  inference(rectify,[],[f329])).
fof(f329,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(flattening,[],[f328])).
fof(f328,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(nnf_transformation,[],[f219])).
fof(f567,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f330])).
fof(f568,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f330])).
fof(f569,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f330])).
fof(f559,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f327])).
fof(f327,plain,(
  ! [X0,X1,X2] : ((~sP11(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP10(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP10(X0,X1,X2) | sP11(X0,X1,X2)))),
  inference(rectify,[],[f326])).
fof(f326,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1) | sP11(X2,X0,X1)))),
  inference(flattening,[],[f325])).
fof(f325,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1)) | sP11(X2,X0,X1)))),
  inference(nnf_transformation,[],[f220])).
fof(f560,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f327])).
fof(f561,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f327])).
fof(f576,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP10(X0,X1,X2) | sP11(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f562])).
fof(f562,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f327])).
fof(f563,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f327])).
fof(f564,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f327])).
fof(f565,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f327])).
fof(f553,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f324])).
fof(f324,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f323])).
fof(f323,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f218])).
fof(f218,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP9(X2,X0,X1))),
  inference(definition_folding,[],[f165,f217,f216])).
fof(f216,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f217,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f165,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f164])).
fof(f164,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X14,X15,X8] : ((p(s(bool,X8)) <=> (p(s(bool,X15)) | p(s(bool,X14)))) <=> ((p(s(bool,X8)) | ~p(s(bool,X15))) & (p(s(bool,X8)) | ~p(s(bool,X14))) & (p(s(bool,X15)) | p(s(bool,X14)) | ~p(s(bool,X8)))))),
  file('thm.arithmetic.MIN_MAX_EQ','thm.sat.sat_25_prove_ltb')).
fof(f554,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f324])).
fof(f555,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f324])).
fof(f556,plain,(
  ( ! [X2,X0,X1] : (~sP9(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f324])).
fof(f549,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f322])).
fof(f322,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f321])).
fof(f321,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f320])).
fof(f320,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f216])).
fof(f550,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f322])).
fof(f551,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f322])).
fof(f552,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f322])).
fof(f542,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f319])).
fof(f319,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP8(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f318])).
fof(f318,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f317])).
fof(f317,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1)) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f217])).
fof(f543,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f319])).
fof(f544,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f319])).
fof(f577,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f545])).
fof(f545,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f319])).
fof(f546,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f319])).
fof(f547,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f319])).
fof(f548,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f319])).
fof(f536,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f316])).
fof(f537,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f316])).
fof(f538,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f316])).
fof(f532,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f314])).
fof(f314,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(flattening,[],[f313])).
fof(f313,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(nnf_transformation,[],[f213])).
fof(f533,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f314])).
fof(f534,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f314])).
fof(f535,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f314])).
fof(f525,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f312])).
fof(f312,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | (sP6(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP6(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f311])).
fof(f311,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP7(X2,X0,X1)))),
  inference(flattening,[],[f310])).
fof(f310,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP7(X2,X0,X1)))),
  inference(nnf_transformation,[],[f214])).
fof(f526,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f312])).
fof(f527,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f312])).
fof(f528,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f312])).
fof(f578,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(duplicate_literal_removal,[],[f531])).
fof(f531,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f312])).
fof(f521,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f309])).
fof(f309,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP5(X2,X1,X0)) & (~sP5(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f212])).
fof(f212,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP5(X2,X1,X0))),
  inference(definition_folding,[],[f161,f211,f210,f209,f208])).
fof(f208,plain,(
  ! [X0,X1,X2] : (sP2(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f209,plain,(
  ! [X1,X0,X2] : (sP3(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f210,plain,(
  ! [X2,X0,X1] : (sP4(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f211,plain,(
  ! [X2,X1,X0] : (sP5(X2,X1,X0) <=> (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f161,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f160])).
fof(f160,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X14,X15,X8] : ((p(s(bool,X8)) <=> s(bool,X14) = s(bool,X15)) <=> ((p(s(bool,X8)) | p(s(bool,X15)) | p(s(bool,X14))) & (p(s(bool,X8)) | ~p(s(bool,X14)) | ~p(s(bool,X15))) & (p(s(bool,X15)) | ~p(s(bool,X14)) | ~p(s(bool,X8))) & (p(s(bool,X14)) | ~p(s(bool,X15)) | ~p(s(bool,X8)))))),
  file('thm.arithmetic.MIN_MAX_EQ','thm.sat.sat_2|49|_prove_ltb')).
fof(f522,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f309])).
fof(f523,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f309])).
fof(f524,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f309])).
fof(f517,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f308])).
fof(f308,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(flattening,[],[f307])).
fof(f307,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(nnf_transformation,[],[f208])).
fof(f518,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f308])).
fof(f519,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f308])).
fof(f520,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f308])).
fof(f513,plain,(
  ( ! [X2,X0,X1] : (~sP3(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f306])).
fof(f306,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f305])).
fof(f305,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(flattening,[],[f304])).
fof(f304,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(nnf_transformation,[],[f209])).
fof(f514,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f306])).
fof(f515,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f306])).
fof(f516,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f306])).
fof(f509,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f303])).
fof(f303,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f302])).
fof(f302,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(flattening,[],[f301])).
fof(f301,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(nnf_transformation,[],[f210])).
fof(f510,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f303])).
fof(f511,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f303])).
fof(f512,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f303])).
fof(f502,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f300])).
fof(f300,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | (sP2(X2,X1,X0) & sP3(X1,X2,X0) & sP4(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP2(X2,X1,X0) | ~sP3(X1,X2,X0) | ~sP4(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f299])).
fof(f299,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP5(X2,X1,X0)))),
  inference(flattening,[],[f298])).
fof(f298,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP5(X2,X1,X0)))),
  inference(nnf_transformation,[],[f211])).
fof(f503,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP3(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f300])).
fof(f504,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP4(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f300])).
fof(f505,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f300])).
fof(f506,plain,(
  ( ! [X2,X0,X1] : (~sP3(X1,X2,X0) | ~p(s(bool,X2)) | ~sP4(X0,X2,X1) | sP5(X0,X1,X2) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f300])).
fof(f507,plain,(
  ( ! [X2,X0,X1] : (~sP3(X1,X2,X0) | p(s(bool,X1)) | ~sP4(X0,X2,X1) | sP5(X0,X1,X2) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f300])).
fof(f508,plain,(
  ( ! [X2,X0,X1] : (~sP3(X1,X2,X0) | p(s(bool,X0)) | ~sP4(X0,X2,X1) | sP5(X0,X1,X2) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f300])).
fof(f481,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X1),s('type.num.num',X2))))) = s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X2)))) )),
  inference(cnf_transformation,[],[f153])).
fof(f153,plain,(
  ! [X0,X1,X2] : s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X1),s('type.num.num',X2))))) = s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X2)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6,X7,X8] : s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X7),s('type.num.num',X8))))) = s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X6),s('type.num.num',X7))),s('type.num.num',X8)))),
  file('thm.arithmetic.MIN_MAX_EQ','thm.arithmetic.MIN_ASSOC')).
fof(f480,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),'const.bool.COND_3'(s(bool,X2),s(fun(X1,X0),X3),s(fun(X1,X0),X4))),s(X1,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,happ(s(fun(X1,X0),X3),s(X1,X5))),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5)))))) )),
  inference(cnf_transformation,[],[f152])).
fof(f152,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X1,X0),'const.bool.COND_3'(s(bool,X2),s(fun(X1,X0),X3),s(fun(X1,X0),X4))),s(X1,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,happ(s(fun(X1,X0),X3),s(X1,X5))),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5)))))),
  inference(flattening,[],[f151])).
fof(f151,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5] : s(X0,happ(s(fun(X1,X0),'const.bool.COND_3'(s(bool,X2),s(fun(X1,X0),X3),s(fun(X1,X0),X4))),s(X1,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,happ(s(fun(X1,X0),X3),s(X1,X5))),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5)))))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X23,X12] : ! [X22,X24,X25,X13] : s(X23,happ(s(fun(X12,X23),'const.bool.COND_3'(s(bool,X22),s(fun(X12,X23),X24),s(fun(X12,X23),X25))),s(X12,X13))) = s(X23,'const.bool.COND_3'(s(bool,X22),s(X23,happ(s(fun(X12,X23),X24),s(X12,X13))),s(X23,happ(s(fun(X12,X23),X25),s(X12,X13)))))),
  file('thm.arithmetic.MIN_MAX_EQ','thm.bool.COND_RATOR')).
fof(f479,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,'const.bool.COND_3'(s(bool,X3),s(X1,X4),s(X1,X5))))) = s(X0,'const.bool.COND_3'(s(bool,X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))) )),
  inference(cnf_transformation,[],[f150])).
fof(f150,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X2),s(X1,'const.bool.COND_3'(s(bool,X3),s(X1,X4),s(X1,X5))))) = s(X0,'const.bool.COND_3'(s(bool,X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))),
  inference(flattening,[],[f149])).
fof(f149,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X2),s(X1,'const.bool.COND_3'(s(bool,X3),s(X1,X4),s(X1,X5))))) = s(X0,'const.bool.COND_3'(s(bool,X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X23,X12] : ! [X24,X22,X13,X19] : s(X23,happ(s(fun(X12,X23),X24),s(X12,'const.bool.COND_3'(s(bool,X22),s(X12,X13),s(X12,X19))))) = s(X23,'const.bool.COND_3'(s(bool,X22),s(X23,happ(s(fun(X12,X23),X24),s(X12,X13))),s(X23,happ(s(fun(X12,X23),X24),s(X12,X19)))))),
  file('thm.arithmetic.MIN_MAX_EQ','thm.bool.COND_RAND')).
fof(f478,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK22(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK22(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f291])).
fof(f291,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK22(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK22(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK22])],[f196])).
fof(f196,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f148])).
fof(f148,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f147])).
fof(f147,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.arithmetic.MIN_MAX_EQ',aHL_EXT)).
fof(f457,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f284])).
fof(f284,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1)))),
  inference(rectify,[],[f283])).
fof(f283,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0)))),
  inference(flattening,[],[f282])).
fof(f282,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP1(X1,X0)))),
  inference(nnf_transformation,[],[f206])).
fof(f206,plain,(
  ! [X1,X0] : (sP1(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f458,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f284])).
fof(f459,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f284])).
fof(f462,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f284])).
fof(f593,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f442,f592])).
fof(f442,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f277])).
fof(f277,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f276])).
fof(f276,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f204])).
fof(f204,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f443,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f277])).
fof(f444,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f277])).
fof(f437,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f188])).
fof(f430,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X1),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X1),s('type.num.num',X0)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6,X7] : s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X6),s('type.num.num',X7))) = s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X7),s('type.num.num',X6)))),
  file('thm.arithmetic.MIN_MAX_EQ','thm.arithmetic.MIN_COMM')).
fof(f429,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f127])).
fof(f127,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X0)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6,X7] : s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X6),s('type.num.num',X7))) = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X7),s('type.num.num',X6)))),
  file('thm.arithmetic.MIN_MAX_EQ','thm.arithmetic.MAX_COMM')).
fof(f426,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,X2)) | s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))) | s(X0,X3) != s(X0,X4) | s(bool,X1) != s(bool,X2)) )),
  inference(cnf_transformation,[],[f184])).
fof(f184,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : (s(bool,X1) != s(bool,X2) | (p(s(bool,X2)) & s(X0,X3) != s(X0,X4)) | (~p(s(bool,X2)) & s(X0,X5) != s(X0,X6)) | s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))))),
  inference(flattening,[],[f183])).
fof(f183,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ((s(bool,X1) != s(bool,X2) | (p(s(bool,X2)) & s(X0,X3) != s(X0,X4)) | (~p(s(bool,X2)) & s(X0,X5) != s(X0,X6))) | s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))))),
  inference(ennf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ((s(bool,X1) = s(bool,X2) & (p(s(bool,X2)) => s(X0,X3) = s(X0,X4)) & (~p(s(bool,X2)) => s(X0,X5) = s(X0,X6))) => s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))))),
  inference(flattening,[],[f125])).
fof(f125,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6] : ((s(bool,X1) = s(bool,X2) & (p(s(bool,X2)) => s(X0,X3) = s(X0,X4)) & (~p(s(bool,X2)) => s(X0,X5) = s(X0,X6))) => s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X12] : ! [X16,X17,X13,X18,X19,X20] : ((s(bool,X16) = s(bool,X17) & (p(s(bool,X17)) => s(X12,X13) = s(X12,X18)) & (~p(s(bool,X17)) => s(X12,X19) = s(X12,X20))) => s(X12,'const.bool.COND_3'(s(bool,X16),s(X12,X13),s(X12,X19))) = s(X12,'const.bool.COND_3'(s(bool,X17),s(X12,X18),s(X12,X20))))),
  file('thm.arithmetic.MIN_MAX_EQ','thm.bool.COND_CONG')).
fof(f427,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,X1) != s(bool,X2) | s(X0,X5) != s(X0,X6) | p(s(bool,X2)) | s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6)))) )),
  inference(cnf_transformation,[],[f184])).
fof(f428,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,X1) != s(bool,X2) | s(X0,X5) != s(X0,X6) | s(X0,X3) != s(X0,X4) | s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6)))) )),
  inference(cnf_transformation,[],[f184])).
fof(f355,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f226])).
fof(f28416,plain,(
  ( ! [X28,X29] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X29),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f27319,f28278])).
fof(f27319,plain,(
  ( ! [X28,X29,X27] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X29),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X27),s(bool,X28))) | s(bool,'F_0') = s(bool,X28)) )),
  inference(superposition,[],[f852,f852])).
fof(f28412,plain,(
  ( ! [X103,X104] : (~p(s(bool,'const.bool.COND_3'(s(bool,X103),s(bool,'T_0'),s(bool,X104)))) | p(s(bool,X104)) | p(s(bool,X103))) ) | $spl98),
  inference(cnf_transformation,[],[f28412_D])).
fof(f28412_D,plain,(
  ( ! [X103,X104] : (~p(s(bool,'const.bool.COND_3'(s(bool,X103),s(bool,'T_0'),s(bool,X104)))) | p(s(bool,X104)) | p(s(bool,X103))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f27357,plain,(
  ( ! [X103,X102,X104] : (~p(s(bool,'const.bool.COND_3'(s(bool,X103),s(bool,'T_0'),s(bool,X104)))) | p(s(bool,X103)) | p(s(bool,X104)) | s(bool,'F_0') = s(bool,X102)) )),
  inference(superposition,[],[f485,f852])).
fof(f28398,plain,(
  ( ! [X232] : (~sP5('F_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,'T_0'),s(bool,X232)),X232)) ) | $spl96),
  inference(cnf_transformation,[],[f28398_D])).
fof(f28398_D,plain,(
  ( ! [X232] : (~sP5('F_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,'T_0'),s(bool,X232)),X232)) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f27391,plain,(
  ( ! [X231,X232] : (~sP5('F_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,'T_0'),s(bool,X232)),X232) | s(bool,'F_0') = s(bool,X231)) )),
  inference(superposition,[],[f836,f852])).
fof(f28397,plain,(
  ( ! [X238] : (sP5('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,'T_0'),s(bool,X238)),X238)) ) | $spl94),
  inference(cnf_transformation,[],[f28397_D])).
fof(f28397_D,plain,(
  ( ! [X238] : (sP5('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,'T_0'),s(bool,X238)),X238)) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f27393,plain,(
  ( ! [X237,X238] : (sP5('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,'T_0'),s(bool,X238)),X238) | s(bool,'F_0') = s(bool,X237)) )),
  inference(superposition,[],[f837,f852])).
fof(f28396,plain,(
  ( ! [X247] : (~sP5('F_0',X247,'const.bool.COND_3'(s(bool,'F_0'),s(bool,'T_0'),s(bool,X247)))) ) | $spl92),
  inference(cnf_transformation,[],[f28396_D])).
fof(f28396_D,plain,(
  ( ! [X247] : (~sP5('F_0',X247,'const.bool.COND_3'(s(bool,'F_0'),s(bool,'T_0'),s(bool,X247)))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f27396,plain,(
  ( ! [X246,X247] : (~sP5('F_0',X247,'const.bool.COND_3'(s(bool,'F_0'),s(bool,'T_0'),s(bool,X247))) | s(bool,'F_0') = s(bool,X246)) )),
  inference(superposition,[],[f838,f852])).
fof(f28395,plain,(
  ( ! [X253] : (sP5('T_0',X253,'const.bool.COND_3'(s(bool,'F_0'),s(bool,'T_0'),s(bool,X253)))) ) | $spl90),
  inference(cnf_transformation,[],[f28395_D])).
fof(f28395_D,plain,(
  ( ! [X253] : (sP5('T_0',X253,'const.bool.COND_3'(s(bool,'F_0'),s(bool,'T_0'),s(bool,X253)))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f27398,plain,(
  ( ! [X253,X252] : (sP5('T_0',X253,'const.bool.COND_3'(s(bool,'F_0'),s(bool,'T_0'),s(bool,X253))) | s(bool,'F_0') = s(bool,X252)) )),
  inference(superposition,[],[f839,f852])).
fof(f28390,plain,(
  ( ! [X306] : (sP4('T_0',X306,'const.bool.COND_3'(s(bool,'F_0'),s(bool,'T_0'),s(bool,X306)))) ) | $spl88),
  inference(cnf_transformation,[],[f28390_D])).
fof(f28390_D,plain,(
  ( ! [X306] : (sP4('T_0',X306,'const.bool.COND_3'(s(bool,'F_0'),s(bool,'T_0'),s(bool,X306)))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f27411,plain,(
  ( ! [X306,X305] : (sP4('T_0',X306,'const.bool.COND_3'(s(bool,'F_0'),s(bool,'T_0'),s(bool,X306))) | s(bool,'F_0') = s(bool,X305)) )),
  inference(superposition,[],[f930,f852])).
fof(f28389,plain,(
  ( ! [X309] : (sP4('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,'T_0'),s(bool,X309)),X309)) ) | $spl86),
  inference(cnf_transformation,[],[f28389_D])).
fof(f28389_D,plain,(
  ( ! [X309] : (sP4('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,'T_0'),s(bool,X309)),X309)) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f27412,plain,(
  ( ! [X308,X309] : (sP4('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,'T_0'),s(bool,X309)),X309) | s(bool,'F_0') = s(bool,X308)) )),
  inference(superposition,[],[f991,f852])).
fof(f28370,plain,(
  ( ! [X561,X560] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X560),s(bool,X561))) | s(bool,'F_0') = s(bool,X561)) ) | $spl68),
  inference(cnf_transformation,[],[f28370_D])).
fof(f28370_D,plain,(
  ( ! [X561,X560] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X560),s(bool,X561))) | s(bool,'F_0') = s(bool,X561)) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f28388,plain,(
  ( ! [X325,X324] : (sP0(X324,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X325),s(bool,'T_0')))) ) | $spl84),
  inference(cnf_transformation,[],[f28388_D])).
fof(f28388_D,plain,(
  ( ! [X325,X324] : (sP0(X324,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X325),s(bool,'T_0')))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f27417,plain,(
  ( ! [X325,X323,X324,X322] : (sP0(X324,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X325),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X322),s(bool,X323))) | s(bool,'F_0') = s(bool,X323)) )),
  inference(superposition,[],[f1239,f852])).
fof(f28387,plain,(
  ( ! [X335,X336] : (sP0('const.bool.COND_3'(s(bool,'T_0'),s(bool,X335),s(bool,'T_0')),X336)) ) | $spl82),
  inference(cnf_transformation,[],[f28387_D])).
fof(f28387_D,plain,(
  ( ! [X335,X336] : (sP0('const.bool.COND_3'(s(bool,'T_0'),s(bool,X335),s(bool,'T_0')),X336)) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f27420,plain,(
  ( ! [X333,X335,X334,X336] : (sP0('const.bool.COND_3'(s(bool,'T_0'),s(bool,X335),s(bool,'T_0')),X336) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X333),s(bool,X334))) | s(bool,'F_0') = s(bool,X334)) )),
  inference(superposition,[],[f1275,f852])).
fof(f28386,plain,(
  ~sP2('const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,'T_0')),'F_0','F_0') | $spl81),
  inference(cnf_transformation,[],[f28386_D])).
fof(f28386_D,plain,(
  ~sP2('const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,'T_0')),'F_0','F_0') <=> ~$spl81),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl81])])).
fof(f27424,plain,(
  ( ! [X346] : (~sP2('const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,'T_0')),'F_0','F_0') | s(bool,'F_0') = s(bool,X346)) )),
  inference(superposition,[],[f1365,f852])).
fof(f28384,plain,(
  ~sP2('const.bool.COND_3'(s(bool,'F_0'),s(bool,'T_0'),s(bool,'F_0')),'F_0','F_0') | $spl79),
  inference(cnf_transformation,[],[f28384_D])).
fof(f28384_D,plain,(
  ~sP2('const.bool.COND_3'(s(bool,'F_0'),s(bool,'T_0'),s(bool,'F_0')),'F_0','F_0') <=> ~$spl79),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl79])])).
fof(f27425,plain,(
  ( ! [X348] : (~sP2('const.bool.COND_3'(s(bool,'F_0'),s(bool,'T_0'),s(bool,'F_0')),'F_0','F_0') | s(bool,'F_0') = s(bool,X348)) )),
  inference(superposition,[],[f1366,f852])).
fof(f28382,plain,(
  ~sP2('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,'T_0')),'F_0') | $spl77),
  inference(cnf_transformation,[],[f28382_D])).
fof(f28382_D,plain,(
  ~sP2('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,'T_0')),'F_0') <=> ~$spl77),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl77])])).
fof(f27428,plain,(
  ( ! [X354] : (~sP2('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,'T_0')),'F_0') | s(bool,'F_0') = s(bool,X354)) )),
  inference(superposition,[],[f1410,f852])).
fof(f28380,plain,(
  ~sP2('F_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,'T_0'),s(bool,'F_0')),'F_0') | $spl75),
  inference(cnf_transformation,[],[f28380_D])).
fof(f28380_D,plain,(
  ~sP2('F_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,'T_0'),s(bool,'F_0')),'F_0') <=> ~$spl75),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl75])])).
fof(f27429,plain,(
  ( ! [X356] : (~sP2('F_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,'T_0'),s(bool,'F_0')),'F_0') | s(bool,'F_0') = s(bool,X356)) )),
  inference(superposition,[],[f1411,f852])).
fof(f28378,plain,(
  ( ! [X362] : (~p(s(bool,'const.bool.COND_3'(s(bool,X362),s(bool,'T_0'),s(bool,'F_0')))) | p(s(bool,X362))) ) | $spl72),
  inference(cnf_transformation,[],[f28378_D])).
fof(f28378_D,plain,(
  ( ! [X362] : (~p(s(bool,'const.bool.COND_3'(s(bool,X362),s(bool,'T_0'),s(bool,'F_0')))) | p(s(bool,X362))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f27431,plain,(
  ( ! [X362,X361] : (~p(s(bool,'const.bool.COND_3'(s(bool,X362),s(bool,'T_0'),s(bool,'F_0')))) | p(s(bool,X362)) | s(bool,'F_0') = s(bool,X361)) )),
  inference(superposition,[],[f1494,f852])).
fof(f28371,plain,(
  ( ! [X562] : (sP1('const.bool.COND_3'(s(bool,'T_0'),s(bool,X562),s(bool,'T_0')),'F_0')) ) | $spl70),
  inference(cnf_transformation,[],[f28371_D])).
fof(f28371_D,plain,(
  ( ! [X562] : (sP1('const.bool.COND_3'(s(bool,'T_0'),s(bool,X562),s(bool,'T_0')),'F_0')) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f27506,plain,(
  ( ! [X561,X560,X562] : (sP1('const.bool.COND_3'(s(bool,'T_0'),s(bool,X562),s(bool,'T_0')),'F_0') | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X560),s(bool,X561))) | s(bool,'F_0') = s(bool,X561)) )),
  inference(superposition,[],[f16717,f852])).
fof(f28368,plain,(
  ( ! [X575,X576] : (p(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X575),s(bool,X576)))) | s(bool,'F_0') = s(bool,X576)) ) | $spl64),
  inference(cnf_transformation,[],[f28368_D])).
fof(f28368_D,plain,(
  ( ! [X575,X576] : (p(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X575),s(bool,X576)))) | s(bool,'F_0') = s(bool,X576)) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f28369,plain,(
  ( ! [X577] : (sP1('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X577)))) ) | $spl66),
  inference(cnf_transformation,[],[f28369_D])).
fof(f28369_D,plain,(
  ( ! [X577] : (sP1('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X577)))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f27512,plain,(
  ( ! [X575,X576,X577] : (sP1('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X577))) | p(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X575),s(bool,X576)))) | s(bool,'F_0') = s(bool,X576)) ) | $spl14),
  inference(superposition,[],[f25970,f852])).
fof(f25970,plain,(
  ( ! [X15,X16] : (sP1('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X15),s(bool,X16))) | p(s(bool,X15))) ) | $spl14),
  inference(superposition,[],[f25838,f413])).
fof(f28367,plain,(
  ( ! [X625] : (s(bool,X625) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,'T_0'),s(bool,X625)))) ) | $spl62),
  inference(cnf_transformation,[],[f28367_D])).
fof(f28367_D,plain,(
  ( ! [X625] : (s(bool,X625) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,'T_0'),s(bool,X625)))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f27528,plain,(
  ( ! [X625,X624] : (s(bool,X625) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,'T_0'),s(bool,X625))) | s(bool,'F_0') = s(bool,X624)) )),
  inference(superposition,[],[f414,f852])).
fof(f27529,plain,(
  ( ! [X627,X628] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X628),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X627)) )),
  inference(superposition,[],[f414,f852])).
fof(f28365,plain,(
  ( ! [X631,X632] : (~p(s(bool,happ(s(fun(bool,bool),X631),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X632),s(bool,sK20(X632,X631,bool))))) | p(s(bool,happ(s(fun(bool,bool),X631),s(bool,sK19(X632,X631,bool)))))) ) | $spl58),
  inference(cnf_transformation,[],[f28365_D])).
fof(f28365_D,plain,(
  ( ! [X631,X632] : (~p(s(bool,happ(s(fun(bool,bool),X631),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X632),s(bool,sK20(X632,X631,bool))))) | p(s(bool,happ(s(fun(bool,bool),X631),s(bool,sK19(X632,X631,bool)))))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f27530,plain,(
  ( ! [X631,X632,X630] : (~p(s(bool,happ(s(fun(bool,bool),X631),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X631),s(bool,sK19(X632,X631,bool))))) | p(s(bool,happ(s(fun(bool,bool),X632),s(bool,sK20(X632,X631,bool))))) | s(bool,'F_0') = s(bool,X630)) )),
  inference(superposition,[],[f421,f852])).
fof(f28364,plain,(
  ( ! [X635,X636] : (~p(s(bool,happ(s(fun(bool,bool),X635),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X635),s(bool,sK20(X635,X636,bool))))) | p(s(bool,happ(s(fun(bool,bool),X636),s(bool,sK19(X635,X636,bool)))))) ) | $spl56),
  inference(cnf_transformation,[],[f28364_D])).
fof(f28364_D,plain,(
  ( ! [X635,X636] : (~p(s(bool,happ(s(fun(bool,bool),X635),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X635),s(bool,sK20(X635,X636,bool))))) | p(s(bool,happ(s(fun(bool,bool),X636),s(bool,sK19(X635,X636,bool)))))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f27531,plain,(
  ( ! [X635,X636,X634] : (~p(s(bool,happ(s(fun(bool,bool),X635),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X636),s(bool,sK19(X635,X636,bool))))) | p(s(bool,happ(s(fun(bool,bool),X635),s(bool,sK20(X635,X636,bool))))) | s(bool,'F_0') = s(bool,X634)) )),
  inference(superposition,[],[f422,f852])).
fof(f28363,plain,(
  ( ! [X639,X640] : (~p(s(bool,happ(s(fun(bool,bool),X639),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X640),s(bool,sK21(X640,X639,bool))))) | p(s(bool,happ(s(fun(bool,bool),X639),s(bool,sK21(X640,X639,bool)))))) ) | $spl54),
  inference(cnf_transformation,[],[f28363_D])).
fof(f28363_D,plain,(
  ( ! [X639,X640] : (~p(s(bool,happ(s(fun(bool,bool),X639),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X640),s(bool,sK21(X640,X639,bool))))) | p(s(bool,happ(s(fun(bool,bool),X639),s(bool,sK21(X640,X639,bool)))))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f27532,plain,(
  ( ! [X639,X640,X638] : (~p(s(bool,happ(s(fun(bool,bool),X639),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X639),s(bool,sK21(X640,X639,bool))))) | p(s(bool,happ(s(fun(bool,bool),X640),s(bool,sK21(X640,X639,bool))))) | s(bool,'F_0') = s(bool,X638)) )),
  inference(superposition,[],[f423,f852])).
fof(f28362,plain,(
  ( ! [X644,X643] : (~p(s(bool,happ(s(fun(bool,bool),X643),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X643),s(bool,sK21(X643,X644,bool))))) | p(s(bool,happ(s(fun(bool,bool),X644),s(bool,sK21(X643,X644,bool)))))) ) | $spl52),
  inference(cnf_transformation,[],[f28362_D])).
fof(f28362_D,plain,(
  ( ! [X644,X643] : (~p(s(bool,happ(s(fun(bool,bool),X643),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X643),s(bool,sK21(X643,X644,bool))))) | p(s(bool,happ(s(fun(bool,bool),X644),s(bool,sK21(X643,X644,bool)))))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f27533,plain,(
  ( ! [X644,X642,X643] : (~p(s(bool,happ(s(fun(bool,bool),X643),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X644),s(bool,sK21(X643,X644,bool))))) | p(s(bool,happ(s(fun(bool,bool),X643),s(bool,sK21(X643,X644,bool))))) | s(bool,'F_0') = s(bool,X642)) )),
  inference(superposition,[],[f424,f852])).
fof(f28258,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f334,f333,f335,f343,f355,f354,f413,f428,f427,f426,f429,f430,f437,f434,f444,f443,f593,f462,f459,f458,f457,f478,f479,f480,f481,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f819,f825,f828,f830,f840,f888,f898,f901,f917,f923,f933,f939,f946,f954,f960,f966,f971,f975,f980,f984,f994,f1000,f1007,f1012,f1024,f1032,f1047,f1050,f1062,f1068,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1239,f1275,f1339,f1332,f1373,f1385,f1402,f1422,f1433,f1437,f1461,f1501,f1512,f1516,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1662,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1981,f2015,f2030,f2073,f2082,f2119,f2118,f2128,f2157,f2170,f2252,f2222,f2262,f2270,f2617,f2622,f2841,f2844,f2751,f2544,f2504,f2465,f2432,f2400,f2371,f2351,f2317,f1837,f1818,f424,f423,f422,f421,f420,f419,f418,f417,f412,f411,f1854,f1861,f1852,f1859,f2052,f2050,f2043,f2041,f482,f433,f1542,f1492,f1486,f489,f488,f487,f485,f484,f1219,f1249,f1217,f1247,f1212,f1256,f1285,f1210,f1254,f1283,f2099,f2097,f2090,f2088,f2001,f1931,f1230,f483,f477,f468,f467,f432,f431,f2848,f2847,f2770,f2765,f2545,f2505,f2466,f2433,f2401,f2398,f2373,f2372,f2352,f2349,f2319,f2318,f2272,f2271,f2269,f2251,f1838,f1835,f1819,f1816,f1663,f1661,f1494,f1411,f1410,f1403,f1366,f1365,f1333,f1267,f1231,f1069,f1059,f1057,f1048,f1035,f1033,f1025,f1023,f1001,f991,f961,f942,f940,f930,f926,f924,f904,f902,f891,f889,f873,f871,f870,f869,f868,f867,f866,f865,f864,f863,f862,f861,f852,f839,f838,f837,f836,f829,f818,f787,f592,f414,f359,f358,f4838,f4843,f4844,f4853,f4855,f4857,f4858,f4859,f4860,f4861,f4862,f4863,f4864,f4644,f4646,f4648,f4650,f4870,f4655,f4658,f4660,f4661,f4663,f4669,f4977,f4986,f5268,f5276,f7384,f7394,f7492,f7497,f7506,f7666,f7671,f7680,f7351,f11643,f7360,f11657,f11665,f12884,f12883,f12097,f12096,f13263,f13282,f13283,f13284,f13285,f13298,f13312,f13313,f13314,f13315,f13323,f13324,f13325,f13326,f13327,f13328,f13329,f13331,f13333,f13335,f13337,f13343,f13346,f13348,f13349,f13351,f13356,f13358,f13361,f16694,f16695,f13454,f16696,f13456,f16697,f16698,f16699,f16700,f16701,f16717,f16728,f18264,f18263,f17257,f17256,f18761,f18782,f18783,f18784,f18785,f18786,f18802,f18816,f18817,f18818,f18819,f18828,f18829,f18830,f18831,f18832,f18833,f18834,f18836,f18838,f18840,f18842,f18848,f18851,f18853,f18854,f18856,f18861,f18863,f18866,f23086,f23087,f18961,f23088,f18963,f23089,f23090,f23091,f23092,f23093,f23129,f23139,f23156,f23157,f23158,f23159,f23173,f23183,f23184,f23185,f23186,f25796,f23209,f23235,f23237,f23239,f23241,f23247,f23250,f23252,f23253,f23255,f23261,f23263,f23266,f25806,f23476,f23477,f23508,f23513,f23514,f23515,f23516,f23517,f23518,f23519,f23520,f23521,f23522,f23523,f25824,f23848,f23849,f23850,f23851,f23852,f23853,f23854,f23855,f23856,f23857,f23858,f23859,f23860,f23861,f23862,f23863,f23864,f23865,f23866,f23867,f23868,f23869,f23870,f23871,f23872,f23873,f23874,f23875,f23876,f23877,f23878,f23879,f23880,f23881,f23882,f23883,f23884,f23885,f23886,f23887,f23888,f23889,f23890,f23891,f23892,f23893,f23894,f23895,f23896,f23897,f23898,f23899,f23900,f23901,f23902,f23903,f23904,f23905,f23906,f23907,f23908,f23909,f23910,f23911,f23912,f23913,f23914,f23915,f23916,f23917,f23918,f23919,f23920,f23921,f23922,f23923,f23924,f23925,f23926,f23927,f23928,f23929,f23930,f23931,f23932,f23933,f23934,f23935,f23936,f23937,f23938,f23939,f25830,f23947,f25837,f25840,f25846,f25849,f25852,f25855,f25858,f25861,f25864,f23973,f23975,f23977,f23979,f23985,f23988,f23990,f23991,f23993,f24001,f25872,f25873,f24118,f25875,f25876,f25877,f25878,f25879,f25880,f25883,f25884,f25885,f25886,f25887,f25888,f25889,f25890,f25891,f24144,f24146,f24148,f24150,f25892,f25893,f25894,f25895,f24156,f24159,f24161,f24162,f24164,f25896,f24170,f24172,f24175,f25897,f25898,f25900,f25901,f25902,f25903,f25904,f25905,f25906,f25907,f25908,f25909,f25910,f25911,f25912,f25913,f25914,f25915,f25916,f25917,f25918,f24416,f24421,f24422,f24423,f24424,f24425,f24426,f24427,f24428,f24429,f24430,f24431,f25920,f25926,f25928,f25933,f25935,f25937,f25939,f25941,f25943,f25945,f25998,f26013,f26015,f26018,f26019,f26009,f26010,f26011,f26012,f26024,f26039,f26041,f26044,f26045,f26035,f26036,f26037,f26038,f26053,f26068,f26070,f26073,f26075,f26090,f26105,f26107,f26110,f26111,f26101,f26102,f26103,f26104,f26118,f26133,f26135,f26138,f26139,f26129,f26130,f26131,f26132,f26141,f26160,f26178,f26193,f26195,f26198,f26199,f26189,f26190,f26191,f26192,f26203,f26233,f26260,f26275,f26277,f26280,f26282,f26290,f26320,f26347,f26362,f26364,f26367,f26368,f26358,f26359,f26360,f26361,f26372,f26391,f26408,f26427,f26442,f26458,f26459,f26460,f26461,f26709,f26711,f26714,f26717,f26718,f26719,f26720,f26721,f26723,f26731,f26733,f26735,f26736,f26737,f26738,f26739,f26740,f26741,f26742,f26521,f26523,f26525,f26527,f26751,f26752,f26755,f26537,f26540,f26542,f26543,f26545,f26757,f26759,f26761,f26763,f26764,f26765,f26767,f26768,f26555,f26557,f26560,f26772,f26777,f26783,f26803,f26820,f26840,f26865,f26866,f26867,f26868,f26869,f26921,f26923,f26925,f26927,f26928,f26948,f26949,f26950,f26951,f26952,f26953,f26954,f26955,f26957,f26959,f26960,f26961,f26962,f26963,f26966,f26967,f26968,f26970,f26971,f26972,f26973,f26974,f26975,f26976,f26977,f26978,f26979,f26984,f26985,f26987,f26989,f26991,f26993,f26995,f26996,f26997,f26998,f26999,f27000,f27001,f27002,f27003,f27006,f27008,f27009,f27011,f27013,f27014,f27015,f27016,f27017,f27018,f27019,f27020,f27021,f27023,f27026,f27030,f27033,f27043,f27106,f27107,f27109,f27119,f27132,f27133,f27134,f27135,f27136,f27137,f27138,f27139,f27140,f27141,f27155,f27156,f27157,f27158,f27159,f27160,f27161,f27162,f27166,f27167,f27168,f27169,f27199,f25807,f27298,f27559,f27562,f27567,f27568,f27569,f27570,f28094,f28095,f28097,f28099,f27582,f27583,f27585,f27587,f27589,f27590,f27592,f27594,f27597,f27599,f27602,f27606,f27607,f27608,f27609,f27610,f27611,f27613,f27614,f27619,f27620,f27635,f27636,f27639,f27640,f27641,f27642,f27643,f27645,f27665,f28144,f28149,f28151,f28169,f28171,f28173,f28175,f28178,f28181,f28184,f28187,f28190,f28192,f28194,f28196,f28198,f28206,f28210,f28214,f28235,f28237,f28239,f28241,f28249,f28253,f28257])).
fof(f28257,plain,(
  ( ! [X315,X314,X309] : (sP11('const.bool.COND_3'(s(bool,X309),s(bool,'T_0'),s(bool,'T_0')),X314,X315)) ) | $spl8),
  inference(forward_demodulation,[],[f28256,f27559])).
fof(f28256,plain,(
  ( ! [X313,X315,X310,X312,X314,X309] : (sP11('const.bool.COND_3'(s(bool,X309),s(bool,'T_0'),s(bool,happ(s(fun(X310,bool),X313),s(X310,X312)))),X314,X315)) ) | $spl8),
  inference(subsumption_resolution,[],[f28255,f25807])).
fof(f28255,plain,(
  ( ! [X313,X315,X310,X312,X314,X309,X311] : (sP11('const.bool.COND_3'(s(bool,X309),s(bool,'T_0'),s(bool,happ(s(fun(X310,bool),X313),s(X310,X312)))),X314,X315) | ~p(s(bool,happ(s(fun(X310,bool),'const.bool.COND_3'(s(bool,X309),s(fun(X310,bool),X311),s(fun(X310,bool),X313))),s(X310,X312))))) ) | $spl8),
  inference(subsumption_resolution,[],[f27967,f25807])).
fof(f27967,plain,(
  ( ! [X313,X315,X310,X312,X314,X309,X311] : (sP11('const.bool.COND_3'(s(bool,X309),s(bool,'T_0'),s(bool,happ(s(fun(X310,bool),X313),s(X310,X312)))),X314,X315) | ~p(s(bool,happ(s(fun(X310,bool),'const.bool.COND_3'(s(bool,X309),s(fun(X310,bool),X311),s(fun(X310,bool),X313))),s(X310,X312)))) | ~p(s(bool,X314))) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f18832])).
fof(f28253,plain,(
  ( ! [X300,X295,X301] : (sP9('const.bool.COND_3'(s(bool,X295),s(bool,'T_0'),s(bool,'T_0')),X300,X301)) ) | $spl8),
  inference(forward_demodulation,[],[f28252,f27559])).
fof(f28252,plain,(
  ( ! [X300,X296,X298,X295,X301,X299] : (sP9('const.bool.COND_3'(s(bool,X295),s(bool,'T_0'),s(bool,happ(s(fun(X296,bool),X299),s(X296,X298)))),X300,X301)) ) | $spl8),
  inference(subsumption_resolution,[],[f28251,f25807])).
fof(f28251,plain,(
  ( ! [X300,X296,X298,X295,X301,X297,X299] : (sP9('const.bool.COND_3'(s(bool,X295),s(bool,'T_0'),s(bool,happ(s(fun(X296,bool),X299),s(X296,X298)))),X300,X301) | ~p(s(bool,happ(s(fun(X296,bool),'const.bool.COND_3'(s(bool,X295),s(fun(X296,bool),X297),s(fun(X296,bool),X299))),s(X296,X298))))) ) | $spl8),
  inference(subsumption_resolution,[],[f27965,f25807])).
fof(f27965,plain,(
  ( ! [X300,X296,X298,X295,X301,X297,X299] : (sP9('const.bool.COND_3'(s(bool,X295),s(bool,'T_0'),s(bool,happ(s(fun(X296,bool),X299),s(X296,X298)))),X300,X301) | ~p(s(bool,happ(s(fun(X296,bool),'const.bool.COND_3'(s(bool,X295),s(fun(X296,bool),X297),s(fun(X296,bool),X299))),s(X296,X298)))) | ~p(s(bool,X300))) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f18830])).
fof(f28249,plain,(
  ( ! [X294,X288,X293] : (sP9('const.bool.COND_3'(s(bool,X288),s(bool,'T_0'),s(bool,'T_0')),X294,X293)) ) | $spl8),
  inference(forward_demodulation,[],[f28248,f27559])).
fof(f28248,plain,(
  ( ! [X292,X294,X288,X293,X289,X291] : (sP9('const.bool.COND_3'(s(bool,X288),s(bool,'T_0'),s(bool,happ(s(fun(X289,bool),X292),s(X289,X291)))),X294,X293)) ) | $spl8),
  inference(subsumption_resolution,[],[f28247,f25807])).
fof(f28247,plain,(
  ( ! [X292,X294,X288,X290,X293,X289,X291] : (sP9('const.bool.COND_3'(s(bool,X288),s(bool,'T_0'),s(bool,happ(s(fun(X289,bool),X292),s(X289,X291)))),X294,X293) | ~p(s(bool,happ(s(fun(X289,bool),'const.bool.COND_3'(s(bool,X288),s(fun(X289,bool),X290),s(fun(X289,bool),X292))),s(X289,X291))))) ) | $spl8),
  inference(subsumption_resolution,[],[f27964,f25807])).
fof(f27964,plain,(
  ( ! [X292,X294,X288,X290,X293,X289,X291] : (sP9('const.bool.COND_3'(s(bool,X288),s(bool,'T_0'),s(bool,happ(s(fun(X289,bool),X292),s(X289,X291)))),X294,X293) | ~p(s(bool,happ(s(fun(X289,bool),'const.bool.COND_3'(s(bool,X288),s(fun(X289,bool),X290),s(fun(X289,bool),X292))),s(X289,X291)))) | ~p(s(bool,X293))) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f18829])).
fof(f28241,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP5('T_0',happ(s(fun(X0,bool),'const.bool.COND_3'(s(bool,X1),s(fun(X0,bool),X2),s(fun(X0,bool),X3))),s(X0,X4)),'const.bool.COND_3'(s(bool,X1),s(bool,'T_0'),s(bool,'T_0')))) ) | $spl8),
  inference(forward_demodulation,[],[f27957,f27559])).
fof(f27957,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP5('T_0',happ(s(fun(X0,bool),'const.bool.COND_3'(s(bool,X1),s(fun(X0,bool),X2),s(fun(X0,bool),X3))),s(X0,X4)),'const.bool.COND_3'(s(bool,X1),s(bool,'T_0'),s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))))) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f18264])).
fof(f28239,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP5('F_0',happ(s(fun(X0,bool),'const.bool.COND_3'(s(bool,X1),s(fun(X0,bool),X2),s(fun(X0,bool),X3))),s(X0,X4)),'const.bool.COND_3'(s(bool,X1),s(bool,'T_0'),s(bool,'T_0')))) ) | $spl8),
  inference(forward_demodulation,[],[f27956,f27559])).
fof(f27956,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP5('F_0',happ(s(fun(X0,bool),'const.bool.COND_3'(s(bool,X1),s(fun(X0,bool),X2),s(fun(X0,bool),X3))),s(X0,X4)),'const.bool.COND_3'(s(bool,X1),s(bool,'T_0'),s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))))) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f18263])).
fof(f28237,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP5('T_0','const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),happ(s(fun(X1,bool),'const.bool.COND_3'(s(bool,X0),s(fun(X1,bool),X2),s(fun(X1,bool),X4))),s(X1,X3)))) ) | $spl8),
  inference(forward_demodulation,[],[f27955,f27559])).
fof(f27955,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP5('T_0','const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,happ(s(fun(X1,bool),X4),s(X1,X3)))),happ(s(fun(X1,bool),'const.bool.COND_3'(s(bool,X0),s(fun(X1,bool),X2),s(fun(X1,bool),X4))),s(X1,X3)))) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f17257])).
fof(f28235,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP5('F_0','const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),happ(s(fun(X1,bool),'const.bool.COND_3'(s(bool,X0),s(fun(X1,bool),X2),s(fun(X1,bool),X4))),s(X1,X3)))) ) | $spl8),
  inference(forward_demodulation,[],[f27954,f27559])).
fof(f27954,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP5('F_0','const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,happ(s(fun(X1,bool),X4),s(X1,X3)))),happ(s(fun(X1,bool),'const.bool.COND_3'(s(bool,X0),s(fun(X1,bool),X2),s(fun(X1,bool),X4))),s(X1,X3)))) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f17256])).
fof(f28214,plain,(
  ( ! [X280,X279,X274] : (sP11('const.bool.COND_3'(s(bool,X274),s(bool,'T_0'),s(bool,'T_0')),X279,X280)) ) | $spl8),
  inference(forward_demodulation,[],[f28213,f27559])).
fof(f28213,plain,(
  ( ! [X280,X279,X275,X276,X278,X274] : (sP11('const.bool.COND_3'(s(bool,X274),s(bool,'T_0'),s(bool,happ(s(fun(X275,bool),X276),s(X275,X278)))),X279,X280)) ) | $spl8),
  inference(subsumption_resolution,[],[f28212,f25807])).
fof(f28212,plain,(
  ( ! [X280,X277,X279,X275,X276,X278,X274] : (sP11('const.bool.COND_3'(s(bool,X274),s(bool,'T_0'),s(bool,happ(s(fun(X275,bool),X276),s(X275,X278)))),X279,X280) | ~p(s(bool,happ(s(fun(X275,bool),X276),s(X275,'const.bool.COND_3'(s(bool,X274),s(X275,X277),s(X275,X278))))))) ) | $spl8),
  inference(subsumption_resolution,[],[f27943,f25807])).
fof(f27943,plain,(
  ( ! [X280,X277,X279,X275,X276,X278,X274] : (sP11('const.bool.COND_3'(s(bool,X274),s(bool,'T_0'),s(bool,happ(s(fun(X275,bool),X276),s(X275,X278)))),X279,X280) | ~p(s(bool,happ(s(fun(X275,bool),X276),s(X275,'const.bool.COND_3'(s(bool,X274),s(X275,X277),s(X275,X278)))))) | ~p(s(bool,X279))) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f13327])).
fof(f28210,plain,(
  ( ! [X265,X260,X266] : (sP9('const.bool.COND_3'(s(bool,X260),s(bool,'T_0'),s(bool,'T_0')),X265,X266)) ) | $spl8),
  inference(forward_demodulation,[],[f28209,f27559])).
fof(f28209,plain,(
  ( ! [X261,X265,X260,X262,X264,X266] : (sP9('const.bool.COND_3'(s(bool,X260),s(bool,'T_0'),s(bool,happ(s(fun(X261,bool),X262),s(X261,X264)))),X265,X266)) ) | $spl8),
  inference(subsumption_resolution,[],[f28208,f25807])).
fof(f28208,plain,(
  ( ! [X261,X263,X265,X260,X262,X264,X266] : (sP9('const.bool.COND_3'(s(bool,X260),s(bool,'T_0'),s(bool,happ(s(fun(X261,bool),X262),s(X261,X264)))),X265,X266) | ~p(s(bool,happ(s(fun(X261,bool),X262),s(X261,'const.bool.COND_3'(s(bool,X260),s(X261,X263),s(X261,X264))))))) ) | $spl8),
  inference(subsumption_resolution,[],[f27941,f25807])).
fof(f27941,plain,(
  ( ! [X261,X263,X265,X260,X262,X264,X266] : (sP9('const.bool.COND_3'(s(bool,X260),s(bool,'T_0'),s(bool,happ(s(fun(X261,bool),X262),s(X261,X264)))),X265,X266) | ~p(s(bool,happ(s(fun(X261,bool),X262),s(X261,'const.bool.COND_3'(s(bool,X260),s(X261,X263),s(X261,X264)))))) | ~p(s(bool,X265))) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f13325])).
fof(f28206,plain,(
  ( ! [X253,X259,X258] : (sP9('const.bool.COND_3'(s(bool,X253),s(bool,'T_0'),s(bool,'T_0')),X259,X258)) ) | $spl8),
  inference(forward_demodulation,[],[f28205,f27559])).
fof(f28205,plain,(
  ( ! [X255,X253,X257,X259,X254,X258] : (sP9('const.bool.COND_3'(s(bool,X253),s(bool,'T_0'),s(bool,happ(s(fun(X254,bool),X255),s(X254,X257)))),X259,X258)) ) | $spl8),
  inference(subsumption_resolution,[],[f28204,f25807])).
fof(f28204,plain,(
  ( ! [X255,X253,X257,X259,X254,X256,X258] : (sP9('const.bool.COND_3'(s(bool,X253),s(bool,'T_0'),s(bool,happ(s(fun(X254,bool),X255),s(X254,X257)))),X259,X258) | ~p(s(bool,happ(s(fun(X254,bool),X255),s(X254,'const.bool.COND_3'(s(bool,X253),s(X254,X256),s(X254,X257))))))) ) | $spl8),
  inference(subsumption_resolution,[],[f27940,f25807])).
fof(f27940,plain,(
  ( ! [X255,X253,X257,X259,X254,X256,X258] : (sP9('const.bool.COND_3'(s(bool,X253),s(bool,'T_0'),s(bool,happ(s(fun(X254,bool),X255),s(X254,X257)))),X259,X258) | ~p(s(bool,happ(s(fun(X254,bool),X255),s(X254,'const.bool.COND_3'(s(bool,X253),s(X254,X256),s(X254,X257)))))) | ~p(s(bool,X258))) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f13324])).
fof(f28198,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP5('T_0',happ(s(fun(X0,bool),X1),s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X3),s(X0,X4)))),'const.bool.COND_3'(s(bool,X2),s(bool,'T_0'),s(bool,'T_0')))) ) | $spl8),
  inference(forward_demodulation,[],[f27933,f27559])).
fof(f27933,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP5('T_0',happ(s(fun(X0,bool),X1),s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X3),s(X0,X4)))),'const.bool.COND_3'(s(bool,X2),s(bool,'T_0'),s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))))) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f12884])).
fof(f28196,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP5('F_0',happ(s(fun(X0,bool),X1),s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X3),s(X0,X4)))),'const.bool.COND_3'(s(bool,X2),s(bool,'T_0'),s(bool,'T_0')))) ) | $spl8),
  inference(forward_demodulation,[],[f27932,f27559])).
fof(f27932,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP5('F_0',happ(s(fun(X0,bool),X1),s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X3),s(X0,X4)))),'const.bool.COND_3'(s(bool,X2),s(bool,'T_0'),s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))))) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f12883])).
fof(f28194,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP5('T_0','const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),happ(s(fun(X1,bool),X2),s(X1,'const.bool.COND_3'(s(bool,X0),s(X1,X3),s(X1,X4)))))) ) | $spl8),
  inference(forward_demodulation,[],[f27931,f27559])).
fof(f27931,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP5('T_0','const.bool.COND_3'(s(bool,X0),s(bool,happ(s(fun(X1,bool),X2),s(X1,X3))),s(bool,'T_0')),happ(s(fun(X1,bool),X2),s(X1,'const.bool.COND_3'(s(bool,X0),s(X1,X3),s(X1,X4)))))) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f12097])).
fof(f28192,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP5('F_0','const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),happ(s(fun(X1,bool),X2),s(X1,'const.bool.COND_3'(s(bool,X0),s(X1,X3),s(X1,X4)))))) ) | $spl8),
  inference(forward_demodulation,[],[f27930,f27559])).
fof(f27930,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP5('F_0','const.bool.COND_3'(s(bool,X0),s(bool,happ(s(fun(X1,bool),X2),s(X1,X3))),s(bool,'T_0')),happ(s(fun(X1,bool),X2),s(X1,'const.bool.COND_3'(s(bool,X0),s(X1,X3),s(X1,X4)))))) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f12096])).
fof(f28190,plain,(
  ( ! [X2,X3,X1] : (s('type.num.num',X3) = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X2)))) ) | $spl8),
  inference(subsumption_resolution,[],[f28189,f25807])).
fof(f28189,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X2),s('type.num.num',X1))),s('type.num.num',X3)))) | s('type.num.num',X3) = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X2)))) ) | $spl8),
  inference(subsumption_resolution,[],[f27909,f335])).
fof(f27909,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X2),s('type.num.num',X1))),s('type.num.num',X3)))) | s('type.num.num',X3) = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X2)))) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f1852])).
fof(f28187,plain,(
  ( ! [X2,X3,X1] : (s('type.num.num',X3) = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X2)))) ) | $spl8),
  inference(subsumption_resolution,[],[f28186,f25807])).
fof(f28186,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num',X3)))) | s('type.num.num',X3) = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X2)))) ) | $spl8),
  inference(subsumption_resolution,[],[f27908,f335])).
fof(f27908,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num',X3)))) | s('type.num.num',X3) = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X2)))) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f1859])).
fof(f28184,plain,(
  ( ! [X8,X7,X9] : (s('type.num.num',X9) = s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X7),s('type.num.num',X8)))) ) | $spl8),
  inference(subsumption_resolution,[],[f28183,f25807])).
fof(f28183,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X8),s('type.num.num',X7))),s('type.num.num',X9)))) | s('type.num.num',X9) = s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X7),s('type.num.num',X8)))) ) | $spl8),
  inference(subsumption_resolution,[],[f27907,f335])).
fof(f27907,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X8),s('type.num.num',X7))),s('type.num.num',X9)))) | s('type.num.num',X9) = s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X7),s('type.num.num',X8)))) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f1854])).
fof(f28181,plain,(
  ( ! [X8,X7,X9] : (s('type.num.num',X9) = s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X7),s('type.num.num',X8)))) ) | $spl8),
  inference(subsumption_resolution,[],[f28180,f25807])).
fof(f28180,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X7),s('type.num.num',X8))),s('type.num.num',X9)))) | s('type.num.num',X9) = s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X7),s('type.num.num',X8)))) ) | $spl8),
  inference(subsumption_resolution,[],[f27906,f335])).
fof(f27906,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X7),s('type.num.num',X8))),s('type.num.num',X9)))) | s('type.num.num',X9) = s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X7),s('type.num.num',X8)))) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f1861])).
fof(f28178,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) = s('type.num.num',X1)) ) | $spl8),
  inference(subsumption_resolution,[],[f28177,f25807])).
fof(f28177,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | s('type.num.num',X0) = s('type.num.num',X1)) ) | $spl8),
  inference(subsumption_resolution,[],[f27894,f335])).
fof(f27894,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | s('type.num.num',X0) = s('type.num.num',X1)) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f433])).
fof(f28175,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,X2))) = s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,X2))),s('type.num.num',happ(s(fun(X0,'type.num.num'),X3),s(X0,X2)))))) ) | $spl8),
  inference(forward_demodulation,[],[f27843,f413])).
fof(f27843,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.num.num',happ(s(fun(X0,'type.num.num'),'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,'type.num.num'),X1),s(fun(X0,'type.num.num'),X3))),s(X0,X2))) = s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,X2))),s('type.num.num',happ(s(fun(X0,'type.num.num'),X3),s(X0,X2)))))) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f18783])).
fof(f28173,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.num.num',happ(s(fun(X4,'type.num.num'),X7),s(X4,X6))) = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',happ(s(fun(X4,'type.num.num'),X5),s(X4,X6))),s('type.num.num',happ(s(fun(X4,'type.num.num'),X7),s(X4,X6)))))) ) | $spl8),
  inference(forward_demodulation,[],[f27842,f413])).
fof(f27842,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.num.num',happ(s(fun(X4,'type.num.num'),'const.bool.COND_3'(s(bool,'T_0'),s(fun(X4,'type.num.num'),X7),s(fun(X4,'type.num.num'),X5))),s(X4,X6))) = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',happ(s(fun(X4,'type.num.num'),X5),s(X4,X6))),s('type.num.num',happ(s(fun(X4,'type.num.num'),X7),s(X4,X6)))))) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f18784])).
fof(f28171,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,X2))) = s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,X2))),s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,X3)))))) ) | $spl8),
  inference(forward_demodulation,[],[f27841,f413])).
fof(f27841,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X2),s(X0,X3))))) = s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,X2))),s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,X3)))))) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f13283])).
fof(f28169,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.num.num',happ(s(fun(X4,'type.num.num'),X5),s(X4,X7))) = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',happ(s(fun(X4,'type.num.num'),X5),s(X4,X6))),s('type.num.num',happ(s(fun(X4,'type.num.num'),X5),s(X4,X7)))))) ) | $spl8),
  inference(forward_demodulation,[],[f27840,f413])).
fof(f27840,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.num.num',happ(s(fun(X4,'type.num.num'),X5),s(X4,'const.bool.COND_3'(s(bool,'T_0'),s(X4,X7),s(X4,X6))))) = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',happ(s(fun(X4,'type.num.num'),X5),s(X4,X6))),s('type.num.num',happ(s(fun(X4,'type.num.num'),X5),s(X4,X7)))))) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f13284])).
fof(f28151,plain,(
  ( ! [X0,X1] : (s('type.num.num',X1) = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X0),s('type.num.num',X1)))) ) | $spl8),
  inference(forward_demodulation,[],[f27807,f413])).
fof(f27807,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',X1),s('type.num.num',X0)))) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f432])).
fof(f28149,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X0),s('type.num.num',X1)))) ) | $spl8),
  inference(forward_demodulation,[],[f27806,f413])).
fof(f27806,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',X0),s('type.num.num',X1)))) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f431])).
fof(f28144,plain,(
  ( ! [X10] : (s(bool,'F_0') = s(bool,X10)) ) | $spl8),
  inference(subsumption_resolution,[],[f27749,f27559])).
fof(f27749,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X10),s(bool,X11))) | s(bool,'F_0') = s(bool,X10)) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f26709])).
fof(f27665,plain,(
  ( ! [X0] : (sP1('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,'T_0')))) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f23129])).
fof(f27645,plain,(
  ( ! [X0] : (sP1('const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,'T_0')),'F_0')) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f11657])).
fof(f27643,plain,(
  ( ! [X140,X141] : (sP4('T_0','const.bool.COND_3'(s(bool,X140),s(bool,'T_0'),s(bool,'T_0')),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X141),s(bool,'T_0')))) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f4669])).
fof(f27642,plain,(
  ( ! [X127,X126] : (sP5('T_0','const.bool.COND_3'(s(bool,X126),s(bool,'T_0'),s(bool,'T_0')),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X127),s(bool,'T_0')))) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f4663])).
fof(f27641,plain,(
  ( ! [X123,X122] : (~sP5('F_0','const.bool.COND_3'(s(bool,X122),s(bool,'T_0'),s(bool,'T_0')),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X123),s(bool,'T_0')))) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f4661])).
fof(f27640,plain,(
  ( ! [X121,X120] : (sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X121),s(bool,'T_0')),'const.bool.COND_3'(s(bool,X120),s(bool,'T_0'),s(bool,'T_0')))) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f4660])).
fof(f27639,plain,(
  ( ! [X116,X117] : (~sP5('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X117),s(bool,'T_0')),'const.bool.COND_3'(s(bool,X116),s(bool,'T_0'),s(bool,'T_0')))) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f4658])).
fof(f27636,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X2),s('type.num.num',X1))) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',X0),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X2)))))) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f2770])).
fof(f27635,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X2),s('type.num.num',X1))) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',X0),s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X1),s('type.num.num',X2)))))) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f2765])).
fof(f27620,plain,(
  ( ! [X0,X1] : (~sP3('const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,'T_0')),'const.bool.COND_3'(s(bool,X1),s(bool,'T_0'),s(bool,'T_0')),'F_0')) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f2269])).
fof(f27619,plain,(
  ( ! [X0,X1] : (~sP3('const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,'T_0')),'F_0')) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f2251])).
fof(f27614,plain,(
  ( ! [X2,X0,X1] : (sP5('const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),X1,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X2),s(bool,X1)))) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f1663])).
fof(f27613,plain,(
  ( ! [X2,X0,X1] : (sP5('const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2)),X2)) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f1661])).
fof(f27611,plain,(
  ( ! [X0] : (~sP2('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,'T_0')),'F_0')) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f1411])).
fof(f27610,plain,(
  ( ! [X0] : (~sP2('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)),'F_0')) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f1410])).
fof(f27609,plain,(
  ( ! [X0] : (~sP3('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,'T_0')),'F_0')) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f1403])).
fof(f27608,plain,(
  ( ! [X0] : (~sP2('const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,'T_0')),'F_0','F_0')) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f1366])).
fof(f27607,plain,(
  ( ! [X0] : (~sP2('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)),'F_0','F_0')) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f1365])).
fof(f27606,plain,(
  ( ! [X0] : (~sP3('const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,'T_0')),'T_0','F_0')) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f1333])).
fof(f27602,plain,(
  ( ! [X8,X7,X9] : (sP4('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X7),s(bool,X8)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X9),s(bool,X8)))) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f1059])).
fof(f27599,plain,(
  ( ! [X8,X7,X9] : (sP4('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X8),s(bool,X9)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X7),s(bool,X8)))) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f1035])).
fof(f27597,plain,(
  ( ! [X8,X7,X9] : (sP4('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X7),s(bool,X8)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X8),s(bool,X9)))) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f1025])).
fof(f27594,plain,(
  ( ! [X0,X1] : (sP4('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)),X1)) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f991])).
fof(f27592,plain,(
  ( ! [X8,X7,X9] : (sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X9),s(bool,X8)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X7),s(bool,X8)))) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f942])).
fof(f27590,plain,(
  ( ! [X0,X1] : (sP4('T_0',X0,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X0)))) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f930])).
fof(f27589,plain,(
  ( ! [X8,X7,X9] : (~sP5('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X9),s(bool,X8)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X7),s(bool,X8)))) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f926])).
fof(f27587,plain,(
  ( ! [X8,X7,X9] : (sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X7),s(bool,X8)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X8),s(bool,X9)))) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f904])).
fof(f27585,plain,(
  ( ! [X8,X7,X9] : (~sP5('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X7),s(bool,X8)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X8),s(bool,X9)))) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f891])).
fof(f27583,plain,(
  ( ! [X61,X62,X63] : (sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X62),s(bool,X63)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X61),s(bool,X62)))) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f873])).
fof(f27582,plain,(
  ( ! [X57,X56,X55] : (~sP5('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X56),s(bool,X57)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X55),s(bool,X56)))) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f871])).
fof(f28099,plain,(
  ( ! [X45,X43,X46,X44] : (sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X43),s(bool,X44)),X45,X46)) ) | $spl8),
  inference(subsumption_resolution,[],[f28098,f25807])).
fof(f28098,plain,(
  ( ! [X45,X43,X46,X44] : (sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X43),s(bool,X44)),X45,X46) | ~p(s(bool,X45))) ) | $spl8),
  inference(subsumption_resolution,[],[f27579,f25807])).
fof(f27579,plain,(
  ( ! [X45,X43,X46,X44] : (sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X43),s(bool,X44)),X45,X46) | ~p(s(bool,X45)) | ~p(s(bool,X44))) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f868])).
fof(f28097,plain,(
  ( ! [X33,X31,X34,X32] : (sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X31),s(bool,X32)),X34,X33)) ) | $spl8),
  inference(subsumption_resolution,[],[f28096,f25807])).
fof(f28096,plain,(
  ( ! [X33,X31,X34,X32] : (sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X31),s(bool,X32)),X34,X33) | ~p(s(bool,X33))) ) | $spl8),
  inference(subsumption_resolution,[],[f27576,f25807])).
fof(f27576,plain,(
  ( ! [X33,X31,X34,X32] : (sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X31),s(bool,X32)),X34,X33) | ~p(s(bool,X33)) | ~p(s(bool,X32))) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f865])).
fof(f28095,plain,(
  ( ! [X24,X23,X22] : (sP0('const.bool.COND_3'(s(bool,'T_0'),s(bool,X22),s(bool,X23)),X24)) ) | $spl8),
  inference(subsumption_resolution,[],[f27573,f25807])).
fof(f27573,plain,(
  ( ! [X24,X23,X22] : (sP0('const.bool.COND_3'(s(bool,'T_0'),s(bool,X22),s(bool,X23)),X24) | ~p(s(bool,X23))) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f862])).
fof(f28094,plain,(
  ( ! [X21,X19,X20] : (sP0(X21,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X19),s(bool,X20)))) ) | $spl8),
  inference(subsumption_resolution,[],[f27572,f25807])).
fof(f27572,plain,(
  ( ! [X21,X19,X20] : (sP0(X21,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X19),s(bool,X20))) | ~p(s(bool,X20))) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f861])).
fof(f27570,plain,(
  ( ! [X0,X1] : (sP5('T_0',X0,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X0)))) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f839])).
fof(f27569,plain,(
  ( ! [X0,X1] : (~sP5('F_0',X0,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X0)))) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f838])).
fof(f27568,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)),X1)) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f837])).
fof(f27567,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)),X1)) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f836])).
fof(f27562,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f414])).
fof(f27559,plain,(
  ( ! [X16] : (s(bool,'T_0') = s(bool,X16)) ) | $spl8),
  inference(global_subsumption,[],[f334,f333,f335,f343,f355,f354,f413,f428,f427,f426,f429,f430,f437,f434,f444,f443,f593,f462,f459,f458,f457,f478,f479,f480,f481,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f819,f825,f828,f830,f840,f888,f898,f901,f917,f923,f933,f939,f946,f954,f960,f966,f971,f975,f980,f984,f994,f1000,f1007,f1012,f1024,f1032,f1047,f1050,f1062,f1068,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1239,f1275,f1339,f1332,f1373,f1385,f1402,f1422,f1433,f1437,f1461,f1501,f1512,f1516,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1662,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1981,f2015,f2030,f2073,f2082,f2119,f2118,f2128,f2157,f2170,f2252,f2222,f2262,f2270,f2617,f2622,f2841,f2844,f2751,f2544,f2504,f2465,f2432,f2400,f2371,f2351,f2317,f1837,f1818,f424,f423,f422,f421,f420,f419,f418,f417,f412,f411,f1854,f1861,f1852,f1859,f2052,f2050,f2043,f2041,f482,f433,f1542,f1492,f1486,f489,f488,f487,f485,f484,f1219,f1249,f1217,f1247,f1212,f1256,f1285,f1210,f1254,f1283,f2099,f2097,f2090,f2088,f2001,f1931,f1230,f483,f477,f468,f467,f432,f431,f2848,f2847,f2770,f2765,f2545,f2505,f2466,f2433,f2401,f2398,f2373,f2372,f2352,f2349,f2319,f2318,f2272,f2271,f2269,f2251,f1838,f1835,f1819,f1816,f1663,f1661,f1494,f1411,f1410,f1403,f1366,f1365,f1333,f1267,f1231,f1069,f1059,f1057,f1048,f1035,f1033,f1025,f1023,f1001,f991,f961,f942,f940,f930,f926,f924,f904,f902,f891,f889,f873,f871,f870,f869,f868,f867,f866,f865,f864,f863,f862,f861,f852,f839,f838,f837,f836,f829,f818,f787,f592,f414,f359,f358,f4838,f4843,f4844,f4853,f4855,f4857,f4858,f4859,f4860,f4861,f4862,f4863,f4864,f4644,f4646,f4648,f4650,f4870,f4655,f4658,f4660,f4661,f4663,f4669,f4977,f4986,f5268,f5276,f7384,f7394,f7492,f7497,f7506,f7666,f7671,f7680,f7351,f11643,f7360,f11657,f11665,f12884,f12883,f12097,f12096,f13263,f13282,f13283,f13284,f13285,f13298,f13312,f13313,f13314,f13315,f13323,f13324,f13325,f13326,f13327,f13328,f13329,f13331,f13333,f13335,f13337,f13343,f13346,f13348,f13349,f13351,f13356,f13358,f13361,f16694,f16695,f13454,f16696,f13456,f16697,f16698,f16699,f16700,f16701,f16717,f16728,f18264,f18263,f17257,f17256,f18761,f18782,f18783,f18784,f18785,f18786,f18802,f18816,f18817,f18818,f18819,f18828,f18829,f18830,f18831,f18832,f18833,f18834,f18836,f18838,f18840,f18842,f18848,f18851,f18853,f18854,f18856,f18861,f18863,f18866,f23086,f23087,f18961,f23088,f18963,f23089,f23090,f23091,f23092,f23093,f23129,f23139,f23156,f23157,f23158,f23159,f23173,f23183,f23184,f23185,f23186,f25796,f23209,f23235,f23237,f23239,f23241,f23247,f23250,f23252,f23253,f23255,f23261,f23263,f23266,f25806,f23476,f23477,f23508,f23513,f23514,f23515,f23516,f23517,f23518,f23519,f23520,f23521,f23522,f23523,f25824,f23848,f23849,f23850,f23851,f23852,f23853,f23854,f23855,f23856,f23857,f23858,f23859,f23860,f23861,f23862,f23863,f23864,f23865,f23866,f23867,f23868,f23869,f23870,f23871,f23872,f23873,f23874,f23875,f23876,f23877,f23878,f23879,f23880,f23881,f23882,f23883,f23884,f23885,f23886,f23887,f23888,f23889,f23890,f23891,f23892,f23893,f23894,f23895,f23896,f23897,f23898,f23899,f23900,f23901,f23902,f23903,f23904,f23905,f23906,f23907,f23908,f23909,f23910,f23911,f23912,f23913,f23914,f23915,f23916,f23917,f23918,f23919,f23920,f23921,f23922,f23923,f23924,f23925,f23926,f23927,f23928,f23929,f23930,f23931,f23932,f23933,f23934,f23935,f23936,f23937,f23938,f23939,f25830,f23947,f25837,f25840,f25846,f25849,f25852,f25855,f25858,f25861,f25864,f23973,f23975,f23977,f23979,f23985,f23988,f23990,f23991,f23993,f24001,f25872,f25873,f24118,f25875,f25876,f25877,f25878,f25879,f25880,f25883,f25884,f25885,f25886,f25887,f25888,f25889,f25890,f25891,f24144,f24146,f24148,f24150,f25892,f25893,f25894,f25895,f24156,f24159,f24161,f24162,f24164,f25896,f24170,f24172,f24175,f25897,f25898,f25900,f25901,f25902,f25903,f25904,f25905,f25906,f25907,f25908,f25909,f25910,f25911,f25912,f25913,f25914,f25915,f25916,f25917,f25918,f24416,f24421,f24422,f24423,f24424,f24425,f24426,f24427,f24428,f24429,f24430,f24431,f25920,f25926,f25928,f25933,f25935,f25937,f25939,f25941,f25943,f25945,f25998,f26013,f26015,f26018,f26019,f26009,f26010,f26011,f26012,f26024,f26039,f26041,f26044,f26045,f26035,f26036,f26037,f26038,f26053,f26068,f26070,f26073,f26075,f26090,f26105,f26107,f26110,f26111,f26101,f26102,f26103,f26104,f26118,f26133,f26135,f26138,f26139,f26129,f26130,f26131,f26132,f26141,f26160,f26178,f26193,f26195,f26198,f26199,f26189,f26190,f26191,f26192,f26203,f26233,f26260,f26275,f26277,f26280,f26282,f26290,f26320,f26347,f26362,f26364,f26367,f26368,f26358,f26359,f26360,f26361,f26372,f26391,f26408,f26427,f26442,f26458,f26459,f26460,f26461,f26709,f26711,f26714,f26717,f26718,f26719,f26720,f26721,f26723,f26731,f26733,f26735,f26736,f26737,f26738,f26739,f26740,f26741,f26742,f26521,f26523,f26525,f26527,f26751,f26752,f26755,f26537,f26540,f26542,f26543,f26545,f26757,f26759,f26761,f26763,f26764,f26765,f26767,f26768,f26555,f26557,f26560,f26772,f26777,f26783,f26803,f26820,f26840,f26865,f26866,f26867,f26868,f26869,f26921,f26923,f26925,f26927,f26928,f26948,f26949,f26950,f26951,f26952,f26953,f26954,f26955,f26957,f26959,f26960,f26961,f26962,f26963,f26966,f26967,f26968,f26970,f26971,f26972,f26973,f26974,f26975,f26976,f26977,f26978,f26979,f26984,f26985,f26987,f26989,f26991,f26993,f26995,f26996,f26997,f26998,f26999,f27000,f27001,f27002,f27003,f27006,f27008,f27009,f27011,f27013,f27014,f27015,f27016,f27017,f27018,f27019,f27020,f27021,f27023,f27026,f27030,f27033,f27043,f27106,f27107,f27109,f27119,f27132,f27133,f27134,f27135,f27136,f27137,f27138,f27139,f27140,f27141,f27155,f27156,f27157,f27158,f27159,f27160,f27161,f27162,f27166,f27167,f27168,f27169,f27199,f25807,f27298])).
fof(f28254,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f334,f333,f335,f343,f355,f354,f413,f428,f427,f426,f429,f430,f437,f434,f444,f443,f593,f462,f459,f458,f457,f478,f479,f480,f481,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f819,f825,f828,f830,f840,f888,f898,f901,f917,f923,f933,f939,f946,f954,f960,f966,f971,f975,f980,f984,f994,f1000,f1007,f1012,f1024,f1032,f1047,f1050,f1062,f1068,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1239,f1275,f1339,f1332,f1373,f1385,f1402,f1422,f1433,f1437,f1461,f1501,f1512,f1516,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1662,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1981,f2015,f2030,f2073,f2082,f2119,f2118,f2128,f2157,f2170,f2252,f2222,f2262,f2270,f2617,f2622,f2841,f2844,f2751,f2544,f2504,f2465,f2432,f2400,f2371,f2351,f2317,f1837,f1818,f424,f423,f422,f421,f420,f419,f418,f417,f412,f411,f1854,f1861,f1852,f1859,f2052,f2050,f2043,f2041,f482,f433,f1542,f1492,f1486,f489,f488,f487,f485,f484,f1219,f1249,f1217,f1247,f1212,f1256,f1285,f1210,f1254,f1283,f2099,f2097,f2090,f2088,f2001,f1931,f1230,f483,f477,f468,f467,f432,f431,f2848,f2847,f2770,f2765,f2545,f2505,f2466,f2433,f2401,f2398,f2373,f2372,f2352,f2349,f2319,f2318,f2272,f2271,f2269,f2251,f1838,f1835,f1819,f1816,f1663,f1661,f1494,f1411,f1410,f1403,f1366,f1365,f1333,f1267,f1231,f1069,f1059,f1057,f1048,f1035,f1033,f1025,f1023,f1001,f991,f961,f942,f940,f930,f926,f924,f904,f902,f891,f889,f873,f871,f870,f869,f868,f867,f866,f865,f864,f863,f862,f861,f852,f839,f838,f837,f836,f829,f818,f787,f592,f414,f359,f358,f4838,f4843,f4844,f4853,f4855,f4857,f4858,f4859,f4860,f4861,f4862,f4863,f4864,f4644,f4646,f4648,f4650,f4870,f4655,f4658,f4660,f4661,f4663,f4669,f4977,f4986,f5268,f5276,f7384,f7394,f7492,f7497,f7506,f7666,f7671,f7680,f7351,f11643,f7360,f11657,f11665,f12884,f12883,f12097,f12096,f13263,f13282,f13283,f13284,f13285,f13298,f13312,f13313,f13314,f13315,f13323,f13324,f13325,f13326,f13327,f13328,f13329,f13331,f13333,f13335,f13337,f13343,f13346,f13348,f13349,f13351,f13356,f13358,f13361,f16694,f16695,f13454,f16696,f13456,f16697,f16698,f16699,f16700,f16701,f16717,f16728,f18264,f18263,f17257,f17256,f18761,f18782,f18783,f18784,f18785,f18786,f18802,f18816,f18817,f18818,f18819,f18828,f18829,f18830,f18831,f18832,f18833,f18834,f18836,f18838,f18840,f18842,f18848,f18851,f18853,f18854,f18856,f18861,f18863,f18866,f23086,f23087,f18961,f23088,f18963,f23089,f23090,f23091,f23092,f23093,f23129,f23139,f23156,f23157,f23158,f23159,f23173,f23183,f23184,f23185,f23186,f25796,f23209,f23235,f23237,f23239,f23241,f23247,f23250,f23252,f23253,f23255,f23261,f23263,f23266,f25806,f23476,f23477,f23508,f23513,f23514,f23515,f23516,f23517,f23518,f23519,f23520,f23521,f23522,f23523,f25824,f23848,f23849,f23850,f23851,f23852,f23853,f23854,f23855,f23856,f23857,f23858,f23859,f23860,f23861,f23862,f23863,f23864,f23865,f23866,f23867,f23868,f23869,f23870,f23871,f23872,f23873,f23874,f23875,f23876,f23877,f23878,f23879,f23880,f23881,f23882,f23883,f23884,f23885,f23886,f23887,f23888,f23889,f23890,f23891,f23892,f23893,f23894,f23895,f23896,f23897,f23898,f23899,f23900,f23901,f23902,f23903,f23904,f23905,f23906,f23907,f23908,f23909,f23910,f23911,f23912,f23913,f23914,f23915,f23916,f23917,f23918,f23919,f23920,f23921,f23922,f23923,f23924,f23925,f23926,f23927,f23928,f23929,f23930,f23931,f23932,f23933,f23934,f23935,f23936,f23937,f23938,f23939,f25830,f23947,f25837,f25840,f25846,f25849,f25852,f25855,f25858,f25861,f25864,f23973,f23975,f23977,f23979,f23985,f23988,f23990,f23991,f23993,f24001,f25872,f25873,f24118,f25875,f25876,f25877,f25878,f25879,f25880,f25883,f25884,f25885,f25886,f25887,f25888,f25889,f25890,f25891,f24144,f24146,f24148,f24150,f25892,f25893,f25894,f25895,f24156,f24159,f24161,f24162,f24164,f25896,f24170,f24172,f24175,f25897,f25898,f25900,f25901,f25902,f25903,f25904,f25905,f25906,f25907,f25908,f25909,f25910,f25911,f25912,f25913,f25914,f25915,f25916,f25917,f25918,f24416,f24421,f24422,f24423,f24424,f24425,f24426,f24427,f24428,f24429,f24430,f24431,f25920,f25926,f25928,f25933,f25935,f25937,f25939,f25941,f25943,f25945,f25998,f26013,f26015,f26018,f26019,f26009,f26010,f26011,f26012,f26024,f26039,f26041,f26044,f26045,f26035,f26036,f26037,f26038,f26053,f26068,f26070,f26073,f26075,f26090,f26105,f26107,f26110,f26111,f26101,f26102,f26103,f26104,f26118,f26133,f26135,f26138,f26139,f26129,f26130,f26131,f26132,f26141,f26160,f26178,f26193,f26195,f26198,f26199,f26189,f26190,f26191,f26192,f26203,f26233,f26260,f26275,f26277,f26280,f26282,f26290,f26320,f26347,f26362,f26364,f26367,f26368,f26358,f26359,f26360,f26361,f26372,f26391,f26408,f26427,f26442,f26458,f26459,f26460,f26461,f26709,f26711,f26714,f26717,f26718,f26719,f26720,f26721,f26723,f26731,f26733,f26735,f26736,f26737,f26738,f26739,f26740,f26741,f26742,f26521,f26523,f26525,f26527,f26751,f26752,f26755,f26537,f26540,f26542,f26543,f26545,f26757,f26759,f26761,f26763,f26764,f26765,f26767,f26768,f26555,f26557,f26560,f26772,f26777,f26783,f26803,f26820,f26840,f26865,f26866,f26867,f26868,f26869,f26921,f26923,f26925,f26927,f26928,f26948,f26949,f26950,f26951,f26952,f26953,f26954,f26955,f26957,f26959,f26960,f26961,f26962,f26963,f26966,f26967,f26968,f26970,f26971,f26972,f26973,f26974,f26975,f26976,f26977,f26978,f26979,f26984,f26985,f26987,f26989,f26991,f26993,f26995,f26996,f26997,f26998,f26999,f27000,f27001,f27002,f27003,f27006,f27008,f27009,f27011,f27013,f27014,f27015,f27016,f27017,f27018,f27019,f27020,f27021,f27023,f27026,f27030,f27033,f27043,f27106,f27107,f27109,f27119,f27132,f27133,f27134,f27135,f27136,f27137,f27138,f27139,f27140,f27141,f27155,f27156,f27157,f27158,f27159,f27160,f27161,f27162,f27166,f27167,f27168,f27169,f27199,f25807,f27298,f27559,f27562,f27567,f27568,f27569,f27570,f28094,f28095,f28097,f28099,f27582,f27583,f27585,f27587,f27589,f27590,f27592,f27594,f27597,f27599,f27602,f27606,f27607,f27608,f27609,f27610,f27611,f27613,f27614,f27619,f27620,f27635,f27636,f27639,f27640,f27641,f27642,f27643,f27645,f27665,f28144,f28149,f28151,f28169,f28171,f28173,f28175,f28178,f28181,f28184,f28187,f28190,f28192,f28194,f28196,f28198,f28206,f28210,f28214,f28235,f28237,f28239,f28241,f28249,f28253])).
fof(f28250,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f334,f333,f335,f343,f355,f354,f413,f428,f427,f426,f429,f430,f437,f434,f444,f443,f593,f462,f459,f458,f457,f478,f479,f480,f481,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f819,f825,f828,f830,f840,f888,f898,f901,f917,f923,f933,f939,f946,f954,f960,f966,f971,f975,f980,f984,f994,f1000,f1007,f1012,f1024,f1032,f1047,f1050,f1062,f1068,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1239,f1275,f1339,f1332,f1373,f1385,f1402,f1422,f1433,f1437,f1461,f1501,f1512,f1516,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1662,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1981,f2015,f2030,f2073,f2082,f2119,f2118,f2128,f2157,f2170,f2252,f2222,f2262,f2270,f2617,f2622,f2841,f2844,f2751,f2544,f2504,f2465,f2432,f2400,f2371,f2351,f2317,f1837,f1818,f424,f423,f422,f421,f420,f419,f418,f417,f412,f411,f1854,f1861,f1852,f1859,f2052,f2050,f2043,f2041,f482,f433,f1542,f1492,f1486,f489,f488,f487,f485,f484,f1219,f1249,f1217,f1247,f1212,f1256,f1285,f1210,f1254,f1283,f2099,f2097,f2090,f2088,f2001,f1931,f1230,f483,f477,f468,f467,f432,f431,f2848,f2847,f2770,f2765,f2545,f2505,f2466,f2433,f2401,f2398,f2373,f2372,f2352,f2349,f2319,f2318,f2272,f2271,f2269,f2251,f1838,f1835,f1819,f1816,f1663,f1661,f1494,f1411,f1410,f1403,f1366,f1365,f1333,f1267,f1231,f1069,f1059,f1057,f1048,f1035,f1033,f1025,f1023,f1001,f991,f961,f942,f940,f930,f926,f924,f904,f902,f891,f889,f873,f871,f870,f869,f868,f867,f866,f865,f864,f863,f862,f861,f852,f839,f838,f837,f836,f829,f818,f787,f592,f414,f359,f358,f4838,f4843,f4844,f4853,f4855,f4857,f4858,f4859,f4860,f4861,f4862,f4863,f4864,f4644,f4646,f4648,f4650,f4870,f4655,f4658,f4660,f4661,f4663,f4669,f4977,f4986,f5268,f5276,f7384,f7394,f7492,f7497,f7506,f7666,f7671,f7680,f7351,f11643,f7360,f11657,f11665,f12884,f12883,f12097,f12096,f13263,f13282,f13283,f13284,f13285,f13298,f13312,f13313,f13314,f13315,f13323,f13324,f13325,f13326,f13327,f13328,f13329,f13331,f13333,f13335,f13337,f13343,f13346,f13348,f13349,f13351,f13356,f13358,f13361,f16694,f16695,f13454,f16696,f13456,f16697,f16698,f16699,f16700,f16701,f16717,f16728,f18264,f18263,f17257,f17256,f18761,f18782,f18783,f18784,f18785,f18786,f18802,f18816,f18817,f18818,f18819,f18828,f18829,f18830,f18831,f18832,f18833,f18834,f18836,f18838,f18840,f18842,f18848,f18851,f18853,f18854,f18856,f18861,f18863,f18866,f23086,f23087,f18961,f23088,f18963,f23089,f23090,f23091,f23092,f23093,f23129,f23139,f23156,f23157,f23158,f23159,f23173,f23183,f23184,f23185,f23186,f25796,f23209,f23235,f23237,f23239,f23241,f23247,f23250,f23252,f23253,f23255,f23261,f23263,f23266,f25806,f23476,f23477,f23508,f23513,f23514,f23515,f23516,f23517,f23518,f23519,f23520,f23521,f23522,f23523,f25824,f23848,f23849,f23850,f23851,f23852,f23853,f23854,f23855,f23856,f23857,f23858,f23859,f23860,f23861,f23862,f23863,f23864,f23865,f23866,f23867,f23868,f23869,f23870,f23871,f23872,f23873,f23874,f23875,f23876,f23877,f23878,f23879,f23880,f23881,f23882,f23883,f23884,f23885,f23886,f23887,f23888,f23889,f23890,f23891,f23892,f23893,f23894,f23895,f23896,f23897,f23898,f23899,f23900,f23901,f23902,f23903,f23904,f23905,f23906,f23907,f23908,f23909,f23910,f23911,f23912,f23913,f23914,f23915,f23916,f23917,f23918,f23919,f23920,f23921,f23922,f23923,f23924,f23925,f23926,f23927,f23928,f23929,f23930,f23931,f23932,f23933,f23934,f23935,f23936,f23937,f23938,f23939,f25830,f23947,f25837,f25840,f25846,f25849,f25852,f25855,f25858,f25861,f25864,f23973,f23975,f23977,f23979,f23985,f23988,f23990,f23991,f23993,f24001,f25872,f25873,f24118,f25875,f25876,f25877,f25878,f25879,f25880,f25883,f25884,f25885,f25886,f25887,f25888,f25889,f25890,f25891,f24144,f24146,f24148,f24150,f25892,f25893,f25894,f25895,f24156,f24159,f24161,f24162,f24164,f25896,f24170,f24172,f24175,f25897,f25898,f25900,f25901,f25902,f25903,f25904,f25905,f25906,f25907,f25908,f25909,f25910,f25911,f25912,f25913,f25914,f25915,f25916,f25917,f25918,f24416,f24421,f24422,f24423,f24424,f24425,f24426,f24427,f24428,f24429,f24430,f24431,f25920,f25926,f25928,f25933,f25935,f25937,f25939,f25941,f25943,f25945,f25998,f26013,f26015,f26018,f26019,f26009,f26010,f26011,f26012,f26024,f26039,f26041,f26044,f26045,f26035,f26036,f26037,f26038,f26053,f26068,f26070,f26073,f26075,f26090,f26105,f26107,f26110,f26111,f26101,f26102,f26103,f26104,f26118,f26133,f26135,f26138,f26139,f26129,f26130,f26131,f26132,f26141,f26160,f26178,f26193,f26195,f26198,f26199,f26189,f26190,f26191,f26192,f26203,f26233,f26260,f26275,f26277,f26280,f26282,f26290,f26320,f26347,f26362,f26364,f26367,f26368,f26358,f26359,f26360,f26361,f26372,f26391,f26408,f26427,f26442,f26458,f26459,f26460,f26461,f26709,f26711,f26714,f26717,f26718,f26719,f26720,f26721,f26723,f26731,f26733,f26735,f26736,f26737,f26738,f26739,f26740,f26741,f26742,f26521,f26523,f26525,f26527,f26751,f26752,f26755,f26537,f26540,f26542,f26543,f26545,f26757,f26759,f26761,f26763,f26764,f26765,f26767,f26768,f26555,f26557,f26560,f26772,f26777,f26783,f26803,f26820,f26840,f26865,f26866,f26867,f26868,f26869,f26921,f26923,f26925,f26927,f26928,f26948,f26949,f26950,f26951,f26952,f26953,f26954,f26955,f26957,f26959,f26960,f26961,f26962,f26963,f26966,f26967,f26968,f26970,f26971,f26972,f26973,f26974,f26975,f26976,f26977,f26978,f26979,f26984,f26985,f26987,f26989,f26991,f26993,f26995,f26996,f26997,f26998,f26999,f27000,f27001,f27002,f27003,f27006,f27008,f27009,f27011,f27013,f27014,f27015,f27016,f27017,f27018,f27019,f27020,f27021,f27023,f27026,f27030,f27033,f27043,f27106,f27107,f27109,f27119,f27132,f27133,f27134,f27135,f27136,f27137,f27138,f27139,f27140,f27141,f27155,f27156,f27157,f27158,f27159,f27160,f27161,f27162,f27166,f27167,f27168,f27169,f27199,f25807,f27298,f27559,f27562,f27567,f27568,f27569,f27570,f28094,f28095,f28097,f28099,f27582,f27583,f27585,f27587,f27589,f27590,f27592,f27594,f27597,f27599,f27602,f27606,f27607,f27608,f27609,f27610,f27611,f27613,f27614,f27619,f27620,f27635,f27636,f27639,f27640,f27641,f27642,f27643,f27645,f27665,f28144,f28149,f28151,f28169,f28171,f28173,f28175,f28178,f28181,f28184,f28187,f28190,f28192,f28194,f28196,f28198,f28206,f28210,f28214,f28235,f28237,f28239,f28241,f28249])).
fof(f28242,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f334,f333,f335,f343,f355,f354,f413,f428,f427,f426,f429,f430,f437,f434,f444,f443,f593,f462,f459,f458,f457,f478,f479,f480,f481,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f819,f825,f828,f830,f840,f888,f898,f901,f917,f923,f933,f939,f946,f954,f960,f966,f971,f975,f980,f984,f994,f1000,f1007,f1012,f1024,f1032,f1047,f1050,f1062,f1068,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1239,f1275,f1339,f1332,f1373,f1385,f1402,f1422,f1433,f1437,f1461,f1501,f1512,f1516,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1662,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1981,f2015,f2030,f2073,f2082,f2119,f2118,f2128,f2157,f2170,f2252,f2222,f2262,f2270,f2617,f2622,f2841,f2844,f2751,f2544,f2504,f2465,f2432,f2400,f2371,f2351,f2317,f1837,f1818,f424,f423,f422,f421,f420,f419,f418,f417,f412,f411,f1854,f1861,f1852,f1859,f2052,f2050,f2043,f2041,f482,f433,f1542,f1492,f1486,f489,f488,f487,f485,f484,f1219,f1249,f1217,f1247,f1212,f1256,f1285,f1210,f1254,f1283,f2099,f2097,f2090,f2088,f2001,f1931,f1230,f483,f477,f468,f467,f432,f431,f2848,f2847,f2770,f2765,f2545,f2505,f2466,f2433,f2401,f2398,f2373,f2372,f2352,f2349,f2319,f2318,f2272,f2271,f2269,f2251,f1838,f1835,f1819,f1816,f1663,f1661,f1494,f1411,f1410,f1403,f1366,f1365,f1333,f1267,f1231,f1069,f1059,f1057,f1048,f1035,f1033,f1025,f1023,f1001,f991,f961,f942,f940,f930,f926,f924,f904,f902,f891,f889,f873,f871,f870,f869,f868,f867,f866,f865,f864,f863,f862,f861,f852,f839,f838,f837,f836,f829,f818,f787,f592,f414,f359,f358,f4838,f4843,f4844,f4853,f4855,f4857,f4858,f4859,f4860,f4861,f4862,f4863,f4864,f4644,f4646,f4648,f4650,f4870,f4655,f4658,f4660,f4661,f4663,f4669,f4977,f4986,f5268,f5276,f7384,f7394,f7492,f7497,f7506,f7666,f7671,f7680,f7351,f11643,f7360,f11657,f11665,f12884,f12883,f12097,f12096,f13263,f13282,f13283,f13284,f13285,f13298,f13312,f13313,f13314,f13315,f13323,f13324,f13325,f13326,f13327,f13328,f13329,f13331,f13333,f13335,f13337,f13343,f13346,f13348,f13349,f13351,f13356,f13358,f13361,f16694,f16695,f13454,f16696,f13456,f16697,f16698,f16699,f16700,f16701,f16717,f16728,f18264,f18263,f17257,f17256,f18761,f18782,f18783,f18784,f18785,f18786,f18802,f18816,f18817,f18818,f18819,f18828,f18829,f18830,f18831,f18832,f18833,f18834,f18836,f18838,f18840,f18842,f18848,f18851,f18853,f18854,f18856,f18861,f18863,f18866,f23086,f23087,f18961,f23088,f18963,f23089,f23090,f23091,f23092,f23093,f23129,f23139,f23156,f23157,f23158,f23159,f23173,f23183,f23184,f23185,f23186,f25796,f23209,f23235,f23237,f23239,f23241,f23247,f23250,f23252,f23253,f23255,f23261,f23263,f23266,f25806,f23476,f23477,f23508,f23513,f23514,f23515,f23516,f23517,f23518,f23519,f23520,f23521,f23522,f23523,f25824,f23848,f23849,f23850,f23851,f23852,f23853,f23854,f23855,f23856,f23857,f23858,f23859,f23860,f23861,f23862,f23863,f23864,f23865,f23866,f23867,f23868,f23869,f23870,f23871,f23872,f23873,f23874,f23875,f23876,f23877,f23878,f23879,f23880,f23881,f23882,f23883,f23884,f23885,f23886,f23887,f23888,f23889,f23890,f23891,f23892,f23893,f23894,f23895,f23896,f23897,f23898,f23899,f23900,f23901,f23902,f23903,f23904,f23905,f23906,f23907,f23908,f23909,f23910,f23911,f23912,f23913,f23914,f23915,f23916,f23917,f23918,f23919,f23920,f23921,f23922,f23923,f23924,f23925,f23926,f23927,f23928,f23929,f23930,f23931,f23932,f23933,f23934,f23935,f23936,f23937,f23938,f23939,f25830,f23947,f25837,f25840,f25846,f25849,f25852,f25855,f25858,f25861,f25864,f23973,f23975,f23977,f23979,f23985,f23988,f23990,f23991,f23993,f24001,f25872,f25873,f24118,f25875,f25876,f25877,f25878,f25879,f25880,f25883,f25884,f25885,f25886,f25887,f25888,f25889,f25890,f25891,f24144,f24146,f24148,f24150,f25892,f25893,f25894,f25895,f24156,f24159,f24161,f24162,f24164,f25896,f24170,f24172,f24175,f25897,f25898,f25900,f25901,f25902,f25903,f25904,f25905,f25906,f25907,f25908,f25909,f25910,f25911,f25912,f25913,f25914,f25915,f25916,f25917,f25918,f24416,f24421,f24422,f24423,f24424,f24425,f24426,f24427,f24428,f24429,f24430,f24431,f25920,f25926,f25928,f25933,f25935,f25937,f25939,f25941,f25943,f25945,f25998,f26013,f26015,f26018,f26019,f26009,f26010,f26011,f26012,f26024,f26039,f26041,f26044,f26045,f26035,f26036,f26037,f26038,f26053,f26068,f26070,f26073,f26075,f26090,f26105,f26107,f26110,f26111,f26101,f26102,f26103,f26104,f26118,f26133,f26135,f26138,f26139,f26129,f26130,f26131,f26132,f26141,f26160,f26178,f26193,f26195,f26198,f26199,f26189,f26190,f26191,f26192,f26203,f26233,f26260,f26275,f26277,f26280,f26282,f26290,f26320,f26347,f26362,f26364,f26367,f26368,f26358,f26359,f26360,f26361,f26372,f26391,f26408,f26427,f26442,f26458,f26459,f26460,f26461,f26709,f26711,f26714,f26717,f26718,f26719,f26720,f26721,f26723,f26731,f26733,f26735,f26736,f26737,f26738,f26739,f26740,f26741,f26742,f26521,f26523,f26525,f26527,f26751,f26752,f26755,f26537,f26540,f26542,f26543,f26545,f26757,f26759,f26761,f26763,f26764,f26765,f26767,f26768,f26555,f26557,f26560,f26772,f26777,f26783,f26803,f26820,f26840,f26865,f26866,f26867,f26868,f26869,f26921,f26923,f26925,f26927,f26928,f26948,f26949,f26950,f26951,f26952,f26953,f26954,f26955,f26957,f26959,f26960,f26961,f26962,f26963,f26966,f26967,f26968,f26970,f26971,f26972,f26973,f26974,f26975,f26976,f26977,f26978,f26979,f26984,f26985,f26987,f26989,f26991,f26993,f26995,f26996,f26997,f26998,f26999,f27000,f27001,f27002,f27003,f27006,f27008,f27009,f27011,f27013,f27014,f27015,f27016,f27017,f27018,f27019,f27020,f27021,f27023,f27026,f27030,f27033,f27043,f27106,f27107,f27109,f27119,f27132,f27133,f27134,f27135,f27136,f27137,f27138,f27139,f27140,f27141,f27155,f27156,f27157,f27158,f27159,f27160,f27161,f27162,f27166,f27167,f27168,f27169,f27199,f25807,f27298,f27559,f27562,f27567,f27568,f27569,f27570,f28094,f28095,f28097,f28099,f27582,f27583,f27585,f27587,f27589,f27590,f27592,f27594,f27597,f27599,f27602,f27606,f27607,f27608,f27609,f27610,f27611,f27613,f27614,f27619,f27620,f27635,f27636,f27639,f27640,f27641,f27642,f27643,f27645,f27665,f28144,f28149,f28151,f28169,f28171,f28173,f28175,f28178,f28181,f28184,f28187,f28190,f28192,f28194,f28196,f28198,f28206,f28210,f28214,f28235,f28237,f28239,f28241])).
fof(f28240,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f334,f333,f335,f343,f355,f354,f413,f428,f427,f426,f429,f430,f437,f434,f444,f443,f593,f462,f459,f458,f457,f478,f479,f480,f481,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f819,f825,f828,f830,f840,f888,f898,f901,f917,f923,f933,f939,f946,f954,f960,f966,f971,f975,f980,f984,f994,f1000,f1007,f1012,f1024,f1032,f1047,f1050,f1062,f1068,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1239,f1275,f1339,f1332,f1373,f1385,f1402,f1422,f1433,f1437,f1461,f1501,f1512,f1516,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1662,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1981,f2015,f2030,f2073,f2082,f2119,f2118,f2128,f2157,f2170,f2252,f2222,f2262,f2270,f2617,f2622,f2841,f2844,f2751,f2544,f2504,f2465,f2432,f2400,f2371,f2351,f2317,f1837,f1818,f424,f423,f422,f421,f420,f419,f418,f417,f412,f411,f1854,f1861,f1852,f1859,f2052,f2050,f2043,f2041,f482,f433,f1542,f1492,f1486,f489,f488,f487,f485,f484,f1219,f1249,f1217,f1247,f1212,f1256,f1285,f1210,f1254,f1283,f2099,f2097,f2090,f2088,f2001,f1931,f1230,f483,f477,f468,f467,f432,f431,f2848,f2847,f2770,f2765,f2545,f2505,f2466,f2433,f2401,f2398,f2373,f2372,f2352,f2349,f2319,f2318,f2272,f2271,f2269,f2251,f1838,f1835,f1819,f1816,f1663,f1661,f1494,f1411,f1410,f1403,f1366,f1365,f1333,f1267,f1231,f1069,f1059,f1057,f1048,f1035,f1033,f1025,f1023,f1001,f991,f961,f942,f940,f930,f926,f924,f904,f902,f891,f889,f873,f871,f870,f869,f868,f867,f866,f865,f864,f863,f862,f861,f852,f839,f838,f837,f836,f829,f818,f787,f592,f414,f359,f358,f4838,f4843,f4844,f4853,f4855,f4857,f4858,f4859,f4860,f4861,f4862,f4863,f4864,f4644,f4646,f4648,f4650,f4870,f4655,f4658,f4660,f4661,f4663,f4669,f4977,f4986,f5268,f5276,f7384,f7394,f7492,f7497,f7506,f7666,f7671,f7680,f7351,f11643,f7360,f11657,f11665,f12884,f12883,f12097,f12096,f13263,f13282,f13283,f13284,f13285,f13298,f13312,f13313,f13314,f13315,f13323,f13324,f13325,f13326,f13327,f13328,f13329,f13331,f13333,f13335,f13337,f13343,f13346,f13348,f13349,f13351,f13356,f13358,f13361,f16694,f16695,f13454,f16696,f13456,f16697,f16698,f16699,f16700,f16701,f16717,f16728,f18264,f18263,f17257,f17256,f18761,f18782,f18783,f18784,f18785,f18786,f18802,f18816,f18817,f18818,f18819,f18828,f18829,f18830,f18831,f18832,f18833,f18834,f18836,f18838,f18840,f18842,f18848,f18851,f18853,f18854,f18856,f18861,f18863,f18866,f23086,f23087,f18961,f23088,f18963,f23089,f23090,f23091,f23092,f23093,f23129,f23139,f23156,f23157,f23158,f23159,f23173,f23183,f23184,f23185,f23186,f25796,f23209,f23235,f23237,f23239,f23241,f23247,f23250,f23252,f23253,f23255,f23261,f23263,f23266,f25806,f23476,f23477,f23508,f23513,f23514,f23515,f23516,f23517,f23518,f23519,f23520,f23521,f23522,f23523,f25824,f23848,f23849,f23850,f23851,f23852,f23853,f23854,f23855,f23856,f23857,f23858,f23859,f23860,f23861,f23862,f23863,f23864,f23865,f23866,f23867,f23868,f23869,f23870,f23871,f23872,f23873,f23874,f23875,f23876,f23877,f23878,f23879,f23880,f23881,f23882,f23883,f23884,f23885,f23886,f23887,f23888,f23889,f23890,f23891,f23892,f23893,f23894,f23895,f23896,f23897,f23898,f23899,f23900,f23901,f23902,f23903,f23904,f23905,f23906,f23907,f23908,f23909,f23910,f23911,f23912,f23913,f23914,f23915,f23916,f23917,f23918,f23919,f23920,f23921,f23922,f23923,f23924,f23925,f23926,f23927,f23928,f23929,f23930,f23931,f23932,f23933,f23934,f23935,f23936,f23937,f23938,f23939,f25830,f23947,f25837,f25840,f25846,f25849,f25852,f25855,f25858,f25861,f25864,f23973,f23975,f23977,f23979,f23985,f23988,f23990,f23991,f23993,f24001,f25872,f25873,f24118,f25875,f25876,f25877,f25878,f25879,f25880,f25883,f25884,f25885,f25886,f25887,f25888,f25889,f25890,f25891,f24144,f24146,f24148,f24150,f25892,f25893,f25894,f25895,f24156,f24159,f24161,f24162,f24164,f25896,f24170,f24172,f24175,f25897,f25898,f25900,f25901,f25902,f25903,f25904,f25905,f25906,f25907,f25908,f25909,f25910,f25911,f25912,f25913,f25914,f25915,f25916,f25917,f25918,f24416,f24421,f24422,f24423,f24424,f24425,f24426,f24427,f24428,f24429,f24430,f24431,f25920,f25926,f25928,f25933,f25935,f25937,f25939,f25941,f25943,f25945,f25998,f26013,f26015,f26018,f26019,f26009,f26010,f26011,f26012,f26024,f26039,f26041,f26044,f26045,f26035,f26036,f26037,f26038,f26053,f26068,f26070,f26073,f26075,f26090,f26105,f26107,f26110,f26111,f26101,f26102,f26103,f26104,f26118,f26133,f26135,f26138,f26139,f26129,f26130,f26131,f26132,f26141,f26160,f26178,f26193,f26195,f26198,f26199,f26189,f26190,f26191,f26192,f26203,f26233,f26260,f26275,f26277,f26280,f26282,f26290,f26320,f26347,f26362,f26364,f26367,f26368,f26358,f26359,f26360,f26361,f26372,f26391,f26408,f26427,f26442,f26458,f26459,f26460,f26461,f26709,f26711,f26714,f26717,f26718,f26719,f26720,f26721,f26723,f26731,f26733,f26735,f26736,f26737,f26738,f26739,f26740,f26741,f26742,f26521,f26523,f26525,f26527,f26751,f26752,f26755,f26537,f26540,f26542,f26543,f26545,f26757,f26759,f26761,f26763,f26764,f26765,f26767,f26768,f26555,f26557,f26560,f26772,f26777,f26783,f26803,f26820,f26840,f26865,f26866,f26867,f26868,f26869,f26921,f26923,f26925,f26927,f26928,f26948,f26949,f26950,f26951,f26952,f26953,f26954,f26955,f26957,f26959,f26960,f26961,f26962,f26963,f26966,f26967,f26968,f26970,f26971,f26972,f26973,f26974,f26975,f26976,f26977,f26978,f26979,f26984,f26985,f26987,f26989,f26991,f26993,f26995,f26996,f26997,f26998,f26999,f27000,f27001,f27002,f27003,f27006,f27008,f27009,f27011,f27013,f27014,f27015,f27016,f27017,f27018,f27019,f27020,f27021,f27023,f27026,f27030,f27033,f27043,f27106,f27107,f27109,f27119,f27132,f27133,f27134,f27135,f27136,f27137,f27138,f27139,f27140,f27141,f27155,f27156,f27157,f27158,f27159,f27160,f27161,f27162,f27166,f27167,f27168,f27169,f27199,f25807,f27298,f27559,f27562,f27567,f27568,f27569,f27570,f28094,f28095,f28097,f28099,f27582,f27583,f27585,f27587,f27589,f27590,f27592,f27594,f27597,f27599,f27602,f27606,f27607,f27608,f27609,f27610,f27611,f27613,f27614,f27619,f27620,f27635,f27636,f27639,f27640,f27641,f27642,f27643,f27645,f27665,f28144,f28149,f28151,f28169,f28171,f28173,f28175,f28178,f28181,f28184,f28187,f28190,f28192,f28194,f28196,f28198,f28206,f28210,f28214,f28235,f28237,f28239])).
fof(f28238,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f334,f333,f335,f343,f355,f354,f413,f428,f427,f426,f429,f430,f437,f434,f444,f443,f593,f462,f459,f458,f457,f478,f479,f480,f481,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f819,f825,f828,f830,f840,f888,f898,f901,f917,f923,f933,f939,f946,f954,f960,f966,f971,f975,f980,f984,f994,f1000,f1007,f1012,f1024,f1032,f1047,f1050,f1062,f1068,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1239,f1275,f1339,f1332,f1373,f1385,f1402,f1422,f1433,f1437,f1461,f1501,f1512,f1516,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1662,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1981,f2015,f2030,f2073,f2082,f2119,f2118,f2128,f2157,f2170,f2252,f2222,f2262,f2270,f2617,f2622,f2841,f2844,f2751,f2544,f2504,f2465,f2432,f2400,f2371,f2351,f2317,f1837,f1818,f424,f423,f422,f421,f420,f419,f418,f417,f412,f411,f1854,f1861,f1852,f1859,f2052,f2050,f2043,f2041,f482,f433,f1542,f1492,f1486,f489,f488,f487,f485,f484,f1219,f1249,f1217,f1247,f1212,f1256,f1285,f1210,f1254,f1283,f2099,f2097,f2090,f2088,f2001,f1931,f1230,f483,f477,f468,f467,f432,f431,f2848,f2847,f2770,f2765,f2545,f2505,f2466,f2433,f2401,f2398,f2373,f2372,f2352,f2349,f2319,f2318,f2272,f2271,f2269,f2251,f1838,f1835,f1819,f1816,f1663,f1661,f1494,f1411,f1410,f1403,f1366,f1365,f1333,f1267,f1231,f1069,f1059,f1057,f1048,f1035,f1033,f1025,f1023,f1001,f991,f961,f942,f940,f930,f926,f924,f904,f902,f891,f889,f873,f871,f870,f869,f868,f867,f866,f865,f864,f863,f862,f861,f852,f839,f838,f837,f836,f829,f818,f787,f592,f414,f359,f358,f4838,f4843,f4844,f4853,f4855,f4857,f4858,f4859,f4860,f4861,f4862,f4863,f4864,f4644,f4646,f4648,f4650,f4870,f4655,f4658,f4660,f4661,f4663,f4669,f4977,f4986,f5268,f5276,f7384,f7394,f7492,f7497,f7506,f7666,f7671,f7680,f7351,f11643,f7360,f11657,f11665,f12884,f12883,f12097,f12096,f13263,f13282,f13283,f13284,f13285,f13298,f13312,f13313,f13314,f13315,f13323,f13324,f13325,f13326,f13327,f13328,f13329,f13331,f13333,f13335,f13337,f13343,f13346,f13348,f13349,f13351,f13356,f13358,f13361,f16694,f16695,f13454,f16696,f13456,f16697,f16698,f16699,f16700,f16701,f16717,f16728,f18264,f18263,f17257,f17256,f18761,f18782,f18783,f18784,f18785,f18786,f18802,f18816,f18817,f18818,f18819,f18828,f18829,f18830,f18831,f18832,f18833,f18834,f18836,f18838,f18840,f18842,f18848,f18851,f18853,f18854,f18856,f18861,f18863,f18866,f23086,f23087,f18961,f23088,f18963,f23089,f23090,f23091,f23092,f23093,f23129,f23139,f23156,f23157,f23158,f23159,f23173,f23183,f23184,f23185,f23186,f25796,f23209,f23235,f23237,f23239,f23241,f23247,f23250,f23252,f23253,f23255,f23261,f23263,f23266,f25806,f23476,f23477,f23508,f23513,f23514,f23515,f23516,f23517,f23518,f23519,f23520,f23521,f23522,f23523,f25824,f23848,f23849,f23850,f23851,f23852,f23853,f23854,f23855,f23856,f23857,f23858,f23859,f23860,f23861,f23862,f23863,f23864,f23865,f23866,f23867,f23868,f23869,f23870,f23871,f23872,f23873,f23874,f23875,f23876,f23877,f23878,f23879,f23880,f23881,f23882,f23883,f23884,f23885,f23886,f23887,f23888,f23889,f23890,f23891,f23892,f23893,f23894,f23895,f23896,f23897,f23898,f23899,f23900,f23901,f23902,f23903,f23904,f23905,f23906,f23907,f23908,f23909,f23910,f23911,f23912,f23913,f23914,f23915,f23916,f23917,f23918,f23919,f23920,f23921,f23922,f23923,f23924,f23925,f23926,f23927,f23928,f23929,f23930,f23931,f23932,f23933,f23934,f23935,f23936,f23937,f23938,f23939,f25830,f23947,f25837,f25840,f25846,f25849,f25852,f25855,f25858,f25861,f25864,f23973,f23975,f23977,f23979,f23985,f23988,f23990,f23991,f23993,f24001,f25872,f25873,f24118,f25875,f25876,f25877,f25878,f25879,f25880,f25883,f25884,f25885,f25886,f25887,f25888,f25889,f25890,f25891,f24144,f24146,f24148,f24150,f25892,f25893,f25894,f25895,f24156,f24159,f24161,f24162,f24164,f25896,f24170,f24172,f24175,f25897,f25898,f25900,f25901,f25902,f25903,f25904,f25905,f25906,f25907,f25908,f25909,f25910,f25911,f25912,f25913,f25914,f25915,f25916,f25917,f25918,f24416,f24421,f24422,f24423,f24424,f24425,f24426,f24427,f24428,f24429,f24430,f24431,f25920,f25926,f25928,f25933,f25935,f25937,f25939,f25941,f25943,f25945,f25998,f26013,f26015,f26018,f26019,f26009,f26010,f26011,f26012,f26024,f26039,f26041,f26044,f26045,f26035,f26036,f26037,f26038,f26053,f26068,f26070,f26073,f26075,f26090,f26105,f26107,f26110,f26111,f26101,f26102,f26103,f26104,f26118,f26133,f26135,f26138,f26139,f26129,f26130,f26131,f26132,f26141,f26160,f26178,f26193,f26195,f26198,f26199,f26189,f26190,f26191,f26192,f26203,f26233,f26260,f26275,f26277,f26280,f26282,f26290,f26320,f26347,f26362,f26364,f26367,f26368,f26358,f26359,f26360,f26361,f26372,f26391,f26408,f26427,f26442,f26458,f26459,f26460,f26461,f26709,f26711,f26714,f26717,f26718,f26719,f26720,f26721,f26723,f26731,f26733,f26735,f26736,f26737,f26738,f26739,f26740,f26741,f26742,f26521,f26523,f26525,f26527,f26751,f26752,f26755,f26537,f26540,f26542,f26543,f26545,f26757,f26759,f26761,f26763,f26764,f26765,f26767,f26768,f26555,f26557,f26560,f26772,f26777,f26783,f26803,f26820,f26840,f26865,f26866,f26867,f26868,f26869,f26921,f26923,f26925,f26927,f26928,f26948,f26949,f26950,f26951,f26952,f26953,f26954,f26955,f26957,f26959,f26960,f26961,f26962,f26963,f26966,f26967,f26968,f26970,f26971,f26972,f26973,f26974,f26975,f26976,f26977,f26978,f26979,f26984,f26985,f26987,f26989,f26991,f26993,f26995,f26996,f26997,f26998,f26999,f27000,f27001,f27002,f27003,f27006,f27008,f27009,f27011,f27013,f27014,f27015,f27016,f27017,f27018,f27019,f27020,f27021,f27023,f27026,f27030,f27033,f27043,f27106,f27107,f27109,f27119,f27132,f27133,f27134,f27135,f27136,f27137,f27138,f27139,f27140,f27141,f27155,f27156,f27157,f27158,f27159,f27160,f27161,f27162,f27166,f27167,f27168,f27169,f27199,f25807,f27298,f27559,f27562,f27567,f27568,f27569,f27570,f28094,f28095,f28097,f28099,f27582,f27583,f27585,f27587,f27589,f27590,f27592,f27594,f27597,f27599,f27602,f27606,f27607,f27608,f27609,f27610,f27611,f27613,f27614,f27619,f27620,f27635,f27636,f27639,f27640,f27641,f27642,f27643,f27645,f27665,f28144,f28149,f28151,f28169,f28171,f28173,f28175,f28178,f28181,f28184,f28187,f28190,f28192,f28194,f28196,f28198,f28206,f28210,f28214,f28235,f28237])).
fof(f28236,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f334,f333,f335,f343,f355,f354,f413,f428,f427,f426,f429,f430,f437,f434,f444,f443,f593,f462,f459,f458,f457,f478,f479,f480,f481,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f819,f825,f828,f830,f840,f888,f898,f901,f917,f923,f933,f939,f946,f954,f960,f966,f971,f975,f980,f984,f994,f1000,f1007,f1012,f1024,f1032,f1047,f1050,f1062,f1068,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1239,f1275,f1339,f1332,f1373,f1385,f1402,f1422,f1433,f1437,f1461,f1501,f1512,f1516,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1662,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1981,f2015,f2030,f2073,f2082,f2119,f2118,f2128,f2157,f2170,f2252,f2222,f2262,f2270,f2617,f2622,f2841,f2844,f2751,f2544,f2504,f2465,f2432,f2400,f2371,f2351,f2317,f1837,f1818,f424,f423,f422,f421,f420,f419,f418,f417,f412,f411,f1854,f1861,f1852,f1859,f2052,f2050,f2043,f2041,f482,f433,f1542,f1492,f1486,f489,f488,f487,f485,f484,f1219,f1249,f1217,f1247,f1212,f1256,f1285,f1210,f1254,f1283,f2099,f2097,f2090,f2088,f2001,f1931,f1230,f483,f477,f468,f467,f432,f431,f2848,f2847,f2770,f2765,f2545,f2505,f2466,f2433,f2401,f2398,f2373,f2372,f2352,f2349,f2319,f2318,f2272,f2271,f2269,f2251,f1838,f1835,f1819,f1816,f1663,f1661,f1494,f1411,f1410,f1403,f1366,f1365,f1333,f1267,f1231,f1069,f1059,f1057,f1048,f1035,f1033,f1025,f1023,f1001,f991,f961,f942,f940,f930,f926,f924,f904,f902,f891,f889,f873,f871,f870,f869,f868,f867,f866,f865,f864,f863,f862,f861,f852,f839,f838,f837,f836,f829,f818,f787,f592,f414,f359,f358,f4838,f4843,f4844,f4853,f4855,f4857,f4858,f4859,f4860,f4861,f4862,f4863,f4864,f4644,f4646,f4648,f4650,f4870,f4655,f4658,f4660,f4661,f4663,f4669,f4977,f4986,f5268,f5276,f7384,f7394,f7492,f7497,f7506,f7666,f7671,f7680,f7351,f11643,f7360,f11657,f11665,f12884,f12883,f12097,f12096,f13263,f13282,f13283,f13284,f13285,f13298,f13312,f13313,f13314,f13315,f13323,f13324,f13325,f13326,f13327,f13328,f13329,f13331,f13333,f13335,f13337,f13343,f13346,f13348,f13349,f13351,f13356,f13358,f13361,f16694,f16695,f13454,f16696,f13456,f16697,f16698,f16699,f16700,f16701,f16717,f16728,f18264,f18263,f17257,f17256,f18761,f18782,f18783,f18784,f18785,f18786,f18802,f18816,f18817,f18818,f18819,f18828,f18829,f18830,f18831,f18832,f18833,f18834,f18836,f18838,f18840,f18842,f18848,f18851,f18853,f18854,f18856,f18861,f18863,f18866,f23086,f23087,f18961,f23088,f18963,f23089,f23090,f23091,f23092,f23093,f23129,f23139,f23156,f23157,f23158,f23159,f23173,f23183,f23184,f23185,f23186,f25796,f23209,f23235,f23237,f23239,f23241,f23247,f23250,f23252,f23253,f23255,f23261,f23263,f23266,f25806,f23476,f23477,f23508,f23513,f23514,f23515,f23516,f23517,f23518,f23519,f23520,f23521,f23522,f23523,f25824,f23848,f23849,f23850,f23851,f23852,f23853,f23854,f23855,f23856,f23857,f23858,f23859,f23860,f23861,f23862,f23863,f23864,f23865,f23866,f23867,f23868,f23869,f23870,f23871,f23872,f23873,f23874,f23875,f23876,f23877,f23878,f23879,f23880,f23881,f23882,f23883,f23884,f23885,f23886,f23887,f23888,f23889,f23890,f23891,f23892,f23893,f23894,f23895,f23896,f23897,f23898,f23899,f23900,f23901,f23902,f23903,f23904,f23905,f23906,f23907,f23908,f23909,f23910,f23911,f23912,f23913,f23914,f23915,f23916,f23917,f23918,f23919,f23920,f23921,f23922,f23923,f23924,f23925,f23926,f23927,f23928,f23929,f23930,f23931,f23932,f23933,f23934,f23935,f23936,f23937,f23938,f23939,f25830,f23947,f25837,f25840,f25846,f25849,f25852,f25855,f25858,f25861,f25864,f23973,f23975,f23977,f23979,f23985,f23988,f23990,f23991,f23993,f24001,f25872,f25873,f24118,f25875,f25876,f25877,f25878,f25879,f25880,f25883,f25884,f25885,f25886,f25887,f25888,f25889,f25890,f25891,f24144,f24146,f24148,f24150,f25892,f25893,f25894,f25895,f24156,f24159,f24161,f24162,f24164,f25896,f24170,f24172,f24175,f25897,f25898,f25900,f25901,f25902,f25903,f25904,f25905,f25906,f25907,f25908,f25909,f25910,f25911,f25912,f25913,f25914,f25915,f25916,f25917,f25918,f24416,f24421,f24422,f24423,f24424,f24425,f24426,f24427,f24428,f24429,f24430,f24431,f25920,f25926,f25928,f25933,f25935,f25937,f25939,f25941,f25943,f25945,f25998,f26013,f26015,f26018,f26019,f26009,f26010,f26011,f26012,f26024,f26039,f26041,f26044,f26045,f26035,f26036,f26037,f26038,f26053,f26068,f26070,f26073,f26075,f26090,f26105,f26107,f26110,f26111,f26101,f26102,f26103,f26104,f26118,f26133,f26135,f26138,f26139,f26129,f26130,f26131,f26132,f26141,f26160,f26178,f26193,f26195,f26198,f26199,f26189,f26190,f26191,f26192,f26203,f26233,f26260,f26275,f26277,f26280,f26282,f26290,f26320,f26347,f26362,f26364,f26367,f26368,f26358,f26359,f26360,f26361,f26372,f26391,f26408,f26427,f26442,f26458,f26459,f26460,f26461,f26709,f26711,f26714,f26717,f26718,f26719,f26720,f26721,f26723,f26731,f26733,f26735,f26736,f26737,f26738,f26739,f26740,f26741,f26742,f26521,f26523,f26525,f26527,f26751,f26752,f26755,f26537,f26540,f26542,f26543,f26545,f26757,f26759,f26761,f26763,f26764,f26765,f26767,f26768,f26555,f26557,f26560,f26772,f26777,f26783,f26803,f26820,f26840,f26865,f26866,f26867,f26868,f26869,f26921,f26923,f26925,f26927,f26928,f26948,f26949,f26950,f26951,f26952,f26953,f26954,f26955,f26957,f26959,f26960,f26961,f26962,f26963,f26966,f26967,f26968,f26970,f26971,f26972,f26973,f26974,f26975,f26976,f26977,f26978,f26979,f26984,f26985,f26987,f26989,f26991,f26993,f26995,f26996,f26997,f26998,f26999,f27000,f27001,f27002,f27003,f27006,f27008,f27009,f27011,f27013,f27014,f27015,f27016,f27017,f27018,f27019,f27020,f27021,f27023,f27026,f27030,f27033,f27043,f27106,f27107,f27109,f27119,f27132,f27133,f27134,f27135,f27136,f27137,f27138,f27139,f27140,f27141,f27155,f27156,f27157,f27158,f27159,f27160,f27161,f27162,f27166,f27167,f27168,f27169,f27199,f25807,f27298,f27559,f27562,f27567,f27568,f27569,f27570,f28094,f28095,f28097,f28099,f27582,f27583,f27585,f27587,f27589,f27590,f27592,f27594,f27597,f27599,f27602,f27606,f27607,f27608,f27609,f27610,f27611,f27613,f27614,f27619,f27620,f27635,f27636,f27639,f27640,f27641,f27642,f27643,f27645,f27665,f28144,f28149,f28151,f28169,f28171,f28173,f28175,f28178,f28181,f28184,f28187,f28190,f28192,f28194,f28196,f28198,f28206,f28210,f28214,f28235])).
fof(f28215,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f334,f333,f335,f343,f355,f354,f413,f428,f427,f426,f429,f430,f437,f434,f444,f443,f593,f462,f459,f458,f457,f478,f479,f480,f481,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f819,f825,f828,f830,f840,f888,f898,f901,f917,f923,f933,f939,f946,f954,f960,f966,f971,f975,f980,f984,f994,f1000,f1007,f1012,f1024,f1032,f1047,f1050,f1062,f1068,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1239,f1275,f1339,f1332,f1373,f1385,f1402,f1422,f1433,f1437,f1461,f1501,f1512,f1516,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1662,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1981,f2015,f2030,f2073,f2082,f2119,f2118,f2128,f2157,f2170,f2252,f2222,f2262,f2270,f2617,f2622,f2841,f2844,f2751,f2544,f2504,f2465,f2432,f2400,f2371,f2351,f2317,f1837,f1818,f424,f423,f422,f421,f420,f419,f418,f417,f412,f411,f1854,f1861,f1852,f1859,f2052,f2050,f2043,f2041,f482,f433,f1542,f1492,f1486,f489,f488,f487,f485,f484,f1219,f1249,f1217,f1247,f1212,f1256,f1285,f1210,f1254,f1283,f2099,f2097,f2090,f2088,f2001,f1931,f1230,f483,f477,f468,f467,f432,f431,f2848,f2847,f2770,f2765,f2545,f2505,f2466,f2433,f2401,f2398,f2373,f2372,f2352,f2349,f2319,f2318,f2272,f2271,f2269,f2251,f1838,f1835,f1819,f1816,f1663,f1661,f1494,f1411,f1410,f1403,f1366,f1365,f1333,f1267,f1231,f1069,f1059,f1057,f1048,f1035,f1033,f1025,f1023,f1001,f991,f961,f942,f940,f930,f926,f924,f904,f902,f891,f889,f873,f871,f870,f869,f868,f867,f866,f865,f864,f863,f862,f861,f852,f839,f838,f837,f836,f829,f818,f787,f592,f414,f359,f358,f4838,f4843,f4844,f4853,f4855,f4857,f4858,f4859,f4860,f4861,f4862,f4863,f4864,f4644,f4646,f4648,f4650,f4870,f4655,f4658,f4660,f4661,f4663,f4669,f4977,f4986,f5268,f5276,f7384,f7394,f7492,f7497,f7506,f7666,f7671,f7680,f7351,f11643,f7360,f11657,f11665,f12884,f12883,f12097,f12096,f13263,f13282,f13283,f13284,f13285,f13298,f13312,f13313,f13314,f13315,f13323,f13324,f13325,f13326,f13327,f13328,f13329,f13331,f13333,f13335,f13337,f13343,f13346,f13348,f13349,f13351,f13356,f13358,f13361,f16694,f16695,f13454,f16696,f13456,f16697,f16698,f16699,f16700,f16701,f16717,f16728,f18264,f18263,f17257,f17256,f18761,f18782,f18783,f18784,f18785,f18786,f18802,f18816,f18817,f18818,f18819,f18828,f18829,f18830,f18831,f18832,f18833,f18834,f18836,f18838,f18840,f18842,f18848,f18851,f18853,f18854,f18856,f18861,f18863,f18866,f23086,f23087,f18961,f23088,f18963,f23089,f23090,f23091,f23092,f23093,f23129,f23139,f23156,f23157,f23158,f23159,f23173,f23183,f23184,f23185,f23186,f25796,f23209,f23235,f23237,f23239,f23241,f23247,f23250,f23252,f23253,f23255,f23261,f23263,f23266,f25806,f23476,f23477,f23508,f23513,f23514,f23515,f23516,f23517,f23518,f23519,f23520,f23521,f23522,f23523,f25824,f23848,f23849,f23850,f23851,f23852,f23853,f23854,f23855,f23856,f23857,f23858,f23859,f23860,f23861,f23862,f23863,f23864,f23865,f23866,f23867,f23868,f23869,f23870,f23871,f23872,f23873,f23874,f23875,f23876,f23877,f23878,f23879,f23880,f23881,f23882,f23883,f23884,f23885,f23886,f23887,f23888,f23889,f23890,f23891,f23892,f23893,f23894,f23895,f23896,f23897,f23898,f23899,f23900,f23901,f23902,f23903,f23904,f23905,f23906,f23907,f23908,f23909,f23910,f23911,f23912,f23913,f23914,f23915,f23916,f23917,f23918,f23919,f23920,f23921,f23922,f23923,f23924,f23925,f23926,f23927,f23928,f23929,f23930,f23931,f23932,f23933,f23934,f23935,f23936,f23937,f23938,f23939,f25830,f23947,f25837,f25840,f25846,f25849,f25852,f25855,f25858,f25861,f25864,f23973,f23975,f23977,f23979,f23985,f23988,f23990,f23991,f23993,f24001,f25872,f25873,f24118,f25875,f25876,f25877,f25878,f25879,f25880,f25883,f25884,f25885,f25886,f25887,f25888,f25889,f25890,f25891,f24144,f24146,f24148,f24150,f25892,f25893,f25894,f25895,f24156,f24159,f24161,f24162,f24164,f25896,f24170,f24172,f24175,f25897,f25898,f25900,f25901,f25902,f25903,f25904,f25905,f25906,f25907,f25908,f25909,f25910,f25911,f25912,f25913,f25914,f25915,f25916,f25917,f25918,f24416,f24421,f24422,f24423,f24424,f24425,f24426,f24427,f24428,f24429,f24430,f24431,f25920,f25926,f25928,f25933,f25935,f25937,f25939,f25941,f25943,f25945,f25998,f26013,f26015,f26018,f26019,f26009,f26010,f26011,f26012,f26024,f26039,f26041,f26044,f26045,f26035,f26036,f26037,f26038,f26053,f26068,f26070,f26073,f26075,f26090,f26105,f26107,f26110,f26111,f26101,f26102,f26103,f26104,f26118,f26133,f26135,f26138,f26139,f26129,f26130,f26131,f26132,f26141,f26160,f26178,f26193,f26195,f26198,f26199,f26189,f26190,f26191,f26192,f26203,f26233,f26260,f26275,f26277,f26280,f26282,f26290,f26320,f26347,f26362,f26364,f26367,f26368,f26358,f26359,f26360,f26361,f26372,f26391,f26408,f26427,f26442,f26458,f26459,f26460,f26461,f26709,f26711,f26714,f26717,f26718,f26719,f26720,f26721,f26723,f26731,f26733,f26735,f26736,f26737,f26738,f26739,f26740,f26741,f26742,f26521,f26523,f26525,f26527,f26751,f26752,f26755,f26537,f26540,f26542,f26543,f26545,f26757,f26759,f26761,f26763,f26764,f26765,f26767,f26768,f26555,f26557,f26560,f26772,f26777,f26783,f26803,f26820,f26840,f26865,f26866,f26867,f26868,f26869,f26921,f26923,f26925,f26927,f26928,f26948,f26949,f26950,f26951,f26952,f26953,f26954,f26955,f26957,f26959,f26960,f26961,f26962,f26963,f26966,f26967,f26968,f26970,f26971,f26972,f26973,f26974,f26975,f26976,f26977,f26978,f26979,f26984,f26985,f26987,f26989,f26991,f26993,f26995,f26996,f26997,f26998,f26999,f27000,f27001,f27002,f27003,f27006,f27008,f27009,f27011,f27013,f27014,f27015,f27016,f27017,f27018,f27019,f27020,f27021,f27023,f27026,f27030,f27033,f27043,f27106,f27107,f27109,f27119,f27132,f27133,f27134,f27135,f27136,f27137,f27138,f27139,f27140,f27141,f27155,f27156,f27157,f27158,f27159,f27160,f27161,f27162,f27166,f27167,f27168,f27169,f27199,f25807,f27298,f27559,f27562,f27567,f27568,f27569,f27570,f28094,f28095,f28097,f28099,f27582,f27583,f27585,f27587,f27589,f27590,f27592,f27594,f27597,f27599,f27602,f27606,f27607,f27608,f27609,f27610,f27611,f27613,f27614,f27619,f27620,f27635,f27636,f27639,f27640,f27641,f27642,f27643,f27645,f27665,f28144,f28149,f28151,f28169,f28171,f28173,f28175,f28178,f28181,f28184,f28187,f28190,f28192,f28194,f28196,f28198,f28206,f28210,f28214])).
fof(f28211,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f334,f333,f335,f343,f355,f354,f413,f428,f427,f426,f429,f430,f437,f434,f444,f443,f593,f462,f459,f458,f457,f478,f479,f480,f481,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f819,f825,f828,f830,f840,f888,f898,f901,f917,f923,f933,f939,f946,f954,f960,f966,f971,f975,f980,f984,f994,f1000,f1007,f1012,f1024,f1032,f1047,f1050,f1062,f1068,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1239,f1275,f1339,f1332,f1373,f1385,f1402,f1422,f1433,f1437,f1461,f1501,f1512,f1516,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1662,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1981,f2015,f2030,f2073,f2082,f2119,f2118,f2128,f2157,f2170,f2252,f2222,f2262,f2270,f2617,f2622,f2841,f2844,f2751,f2544,f2504,f2465,f2432,f2400,f2371,f2351,f2317,f1837,f1818,f424,f423,f422,f421,f420,f419,f418,f417,f412,f411,f1854,f1861,f1852,f1859,f2052,f2050,f2043,f2041,f482,f433,f1542,f1492,f1486,f489,f488,f487,f485,f484,f1219,f1249,f1217,f1247,f1212,f1256,f1285,f1210,f1254,f1283,f2099,f2097,f2090,f2088,f2001,f1931,f1230,f483,f477,f468,f467,f432,f431,f2848,f2847,f2770,f2765,f2545,f2505,f2466,f2433,f2401,f2398,f2373,f2372,f2352,f2349,f2319,f2318,f2272,f2271,f2269,f2251,f1838,f1835,f1819,f1816,f1663,f1661,f1494,f1411,f1410,f1403,f1366,f1365,f1333,f1267,f1231,f1069,f1059,f1057,f1048,f1035,f1033,f1025,f1023,f1001,f991,f961,f942,f940,f930,f926,f924,f904,f902,f891,f889,f873,f871,f870,f869,f868,f867,f866,f865,f864,f863,f862,f861,f852,f839,f838,f837,f836,f829,f818,f787,f592,f414,f359,f358,f4838,f4843,f4844,f4853,f4855,f4857,f4858,f4859,f4860,f4861,f4862,f4863,f4864,f4644,f4646,f4648,f4650,f4870,f4655,f4658,f4660,f4661,f4663,f4669,f4977,f4986,f5268,f5276,f7384,f7394,f7492,f7497,f7506,f7666,f7671,f7680,f7351,f11643,f7360,f11657,f11665,f12884,f12883,f12097,f12096,f13263,f13282,f13283,f13284,f13285,f13298,f13312,f13313,f13314,f13315,f13323,f13324,f13325,f13326,f13327,f13328,f13329,f13331,f13333,f13335,f13337,f13343,f13346,f13348,f13349,f13351,f13356,f13358,f13361,f16694,f16695,f13454,f16696,f13456,f16697,f16698,f16699,f16700,f16701,f16717,f16728,f18264,f18263,f17257,f17256,f18761,f18782,f18783,f18784,f18785,f18786,f18802,f18816,f18817,f18818,f18819,f18828,f18829,f18830,f18831,f18832,f18833,f18834,f18836,f18838,f18840,f18842,f18848,f18851,f18853,f18854,f18856,f18861,f18863,f18866,f23086,f23087,f18961,f23088,f18963,f23089,f23090,f23091,f23092,f23093,f23129,f23139,f23156,f23157,f23158,f23159,f23173,f23183,f23184,f23185,f23186,f25796,f23209,f23235,f23237,f23239,f23241,f23247,f23250,f23252,f23253,f23255,f23261,f23263,f23266,f25806,f23476,f23477,f23508,f23513,f23514,f23515,f23516,f23517,f23518,f23519,f23520,f23521,f23522,f23523,f25824,f23848,f23849,f23850,f23851,f23852,f23853,f23854,f23855,f23856,f23857,f23858,f23859,f23860,f23861,f23862,f23863,f23864,f23865,f23866,f23867,f23868,f23869,f23870,f23871,f23872,f23873,f23874,f23875,f23876,f23877,f23878,f23879,f23880,f23881,f23882,f23883,f23884,f23885,f23886,f23887,f23888,f23889,f23890,f23891,f23892,f23893,f23894,f23895,f23896,f23897,f23898,f23899,f23900,f23901,f23902,f23903,f23904,f23905,f23906,f23907,f23908,f23909,f23910,f23911,f23912,f23913,f23914,f23915,f23916,f23917,f23918,f23919,f23920,f23921,f23922,f23923,f23924,f23925,f23926,f23927,f23928,f23929,f23930,f23931,f23932,f23933,f23934,f23935,f23936,f23937,f23938,f23939,f25830,f23947,f25837,f25840,f25846,f25849,f25852,f25855,f25858,f25861,f25864,f23973,f23975,f23977,f23979,f23985,f23988,f23990,f23991,f23993,f24001,f25872,f25873,f24118,f25875,f25876,f25877,f25878,f25879,f25880,f25883,f25884,f25885,f25886,f25887,f25888,f25889,f25890,f25891,f24144,f24146,f24148,f24150,f25892,f25893,f25894,f25895,f24156,f24159,f24161,f24162,f24164,f25896,f24170,f24172,f24175,f25897,f25898,f25900,f25901,f25902,f25903,f25904,f25905,f25906,f25907,f25908,f25909,f25910,f25911,f25912,f25913,f25914,f25915,f25916,f25917,f25918,f24416,f24421,f24422,f24423,f24424,f24425,f24426,f24427,f24428,f24429,f24430,f24431,f25920,f25926,f25928,f25933,f25935,f25937,f25939,f25941,f25943,f25945,f25998,f26013,f26015,f26018,f26019,f26009,f26010,f26011,f26012,f26024,f26039,f26041,f26044,f26045,f26035,f26036,f26037,f26038,f26053,f26068,f26070,f26073,f26075,f26090,f26105,f26107,f26110,f26111,f26101,f26102,f26103,f26104,f26118,f26133,f26135,f26138,f26139,f26129,f26130,f26131,f26132,f26141,f26160,f26178,f26193,f26195,f26198,f26199,f26189,f26190,f26191,f26192,f26203,f26233,f26260,f26275,f26277,f26280,f26282,f26290,f26320,f26347,f26362,f26364,f26367,f26368,f26358,f26359,f26360,f26361,f26372,f26391,f26408,f26427,f26442,f26458,f26459,f26460,f26461,f26709,f26711,f26714,f26717,f26718,f26719,f26720,f26721,f26723,f26731,f26733,f26735,f26736,f26737,f26738,f26739,f26740,f26741,f26742,f26521,f26523,f26525,f26527,f26751,f26752,f26755,f26537,f26540,f26542,f26543,f26545,f26757,f26759,f26761,f26763,f26764,f26765,f26767,f26768,f26555,f26557,f26560,f26772,f26777,f26783,f26803,f26820,f26840,f26865,f26866,f26867,f26868,f26869,f26921,f26923,f26925,f26927,f26928,f26948,f26949,f26950,f26951,f26952,f26953,f26954,f26955,f26957,f26959,f26960,f26961,f26962,f26963,f26966,f26967,f26968,f26970,f26971,f26972,f26973,f26974,f26975,f26976,f26977,f26978,f26979,f26984,f26985,f26987,f26989,f26991,f26993,f26995,f26996,f26997,f26998,f26999,f27000,f27001,f27002,f27003,f27006,f27008,f27009,f27011,f27013,f27014,f27015,f27016,f27017,f27018,f27019,f27020,f27021,f27023,f27026,f27030,f27033,f27043,f27106,f27107,f27109,f27119,f27132,f27133,f27134,f27135,f27136,f27137,f27138,f27139,f27140,f27141,f27155,f27156,f27157,f27158,f27159,f27160,f27161,f27162,f27166,f27167,f27168,f27169,f27199,f25807,f27298,f27559,f27562,f27567,f27568,f27569,f27570,f28094,f28095,f28097,f28099,f27582,f27583,f27585,f27587,f27589,f27590,f27592,f27594,f27597,f27599,f27602,f27606,f27607,f27608,f27609,f27610,f27611,f27613,f27614,f27619,f27620,f27635,f27636,f27639,f27640,f27641,f27642,f27643,f27645,f27665,f28144,f28149,f28151,f28169,f28171,f28173,f28175,f28178,f28181,f28184,f28187,f28190,f28192,f28194,f28196,f28198,f28206,f28210])).
fof(f28207,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f334,f333,f335,f343,f355,f354,f413,f428,f427,f426,f429,f430,f437,f434,f444,f443,f593,f462,f459,f458,f457,f478,f479,f480,f481,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f819,f825,f828,f830,f840,f888,f898,f901,f917,f923,f933,f939,f946,f954,f960,f966,f971,f975,f980,f984,f994,f1000,f1007,f1012,f1024,f1032,f1047,f1050,f1062,f1068,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1239,f1275,f1339,f1332,f1373,f1385,f1402,f1422,f1433,f1437,f1461,f1501,f1512,f1516,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1662,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1981,f2015,f2030,f2073,f2082,f2119,f2118,f2128,f2157,f2170,f2252,f2222,f2262,f2270,f2617,f2622,f2841,f2844,f2751,f2544,f2504,f2465,f2432,f2400,f2371,f2351,f2317,f1837,f1818,f424,f423,f422,f421,f420,f419,f418,f417,f412,f411,f1854,f1861,f1852,f1859,f2052,f2050,f2043,f2041,f482,f433,f1542,f1492,f1486,f489,f488,f487,f485,f484,f1219,f1249,f1217,f1247,f1212,f1256,f1285,f1210,f1254,f1283,f2099,f2097,f2090,f2088,f2001,f1931,f1230,f483,f477,f468,f467,f432,f431,f2848,f2847,f2770,f2765,f2545,f2505,f2466,f2433,f2401,f2398,f2373,f2372,f2352,f2349,f2319,f2318,f2272,f2271,f2269,f2251,f1838,f1835,f1819,f1816,f1663,f1661,f1494,f1411,f1410,f1403,f1366,f1365,f1333,f1267,f1231,f1069,f1059,f1057,f1048,f1035,f1033,f1025,f1023,f1001,f991,f961,f942,f940,f930,f926,f924,f904,f902,f891,f889,f873,f871,f870,f869,f868,f867,f866,f865,f864,f863,f862,f861,f852,f839,f838,f837,f836,f829,f818,f787,f592,f414,f359,f358,f4838,f4843,f4844,f4853,f4855,f4857,f4858,f4859,f4860,f4861,f4862,f4863,f4864,f4644,f4646,f4648,f4650,f4870,f4655,f4658,f4660,f4661,f4663,f4669,f4977,f4986,f5268,f5276,f7384,f7394,f7492,f7497,f7506,f7666,f7671,f7680,f7351,f11643,f7360,f11657,f11665,f12884,f12883,f12097,f12096,f13263,f13282,f13283,f13284,f13285,f13298,f13312,f13313,f13314,f13315,f13323,f13324,f13325,f13326,f13327,f13328,f13329,f13331,f13333,f13335,f13337,f13343,f13346,f13348,f13349,f13351,f13356,f13358,f13361,f16694,f16695,f13454,f16696,f13456,f16697,f16698,f16699,f16700,f16701,f16717,f16728,f18264,f18263,f17257,f17256,f18761,f18782,f18783,f18784,f18785,f18786,f18802,f18816,f18817,f18818,f18819,f18828,f18829,f18830,f18831,f18832,f18833,f18834,f18836,f18838,f18840,f18842,f18848,f18851,f18853,f18854,f18856,f18861,f18863,f18866,f23086,f23087,f18961,f23088,f18963,f23089,f23090,f23091,f23092,f23093,f23129,f23139,f23156,f23157,f23158,f23159,f23173,f23183,f23184,f23185,f23186,f25796,f23209,f23235,f23237,f23239,f23241,f23247,f23250,f23252,f23253,f23255,f23261,f23263,f23266,f25806,f23476,f23477,f23508,f23513,f23514,f23515,f23516,f23517,f23518,f23519,f23520,f23521,f23522,f23523,f25824,f23848,f23849,f23850,f23851,f23852,f23853,f23854,f23855,f23856,f23857,f23858,f23859,f23860,f23861,f23862,f23863,f23864,f23865,f23866,f23867,f23868,f23869,f23870,f23871,f23872,f23873,f23874,f23875,f23876,f23877,f23878,f23879,f23880,f23881,f23882,f23883,f23884,f23885,f23886,f23887,f23888,f23889,f23890,f23891,f23892,f23893,f23894,f23895,f23896,f23897,f23898,f23899,f23900,f23901,f23902,f23903,f23904,f23905,f23906,f23907,f23908,f23909,f23910,f23911,f23912,f23913,f23914,f23915,f23916,f23917,f23918,f23919,f23920,f23921,f23922,f23923,f23924,f23925,f23926,f23927,f23928,f23929,f23930,f23931,f23932,f23933,f23934,f23935,f23936,f23937,f23938,f23939,f25830,f23947,f25837,f25840,f25846,f25849,f25852,f25855,f25858,f25861,f25864,f23973,f23975,f23977,f23979,f23985,f23988,f23990,f23991,f23993,f24001,f25872,f25873,f24118,f25875,f25876,f25877,f25878,f25879,f25880,f25883,f25884,f25885,f25886,f25887,f25888,f25889,f25890,f25891,f24144,f24146,f24148,f24150,f25892,f25893,f25894,f25895,f24156,f24159,f24161,f24162,f24164,f25896,f24170,f24172,f24175,f25897,f25898,f25900,f25901,f25902,f25903,f25904,f25905,f25906,f25907,f25908,f25909,f25910,f25911,f25912,f25913,f25914,f25915,f25916,f25917,f25918,f24416,f24421,f24422,f24423,f24424,f24425,f24426,f24427,f24428,f24429,f24430,f24431,f25920,f25926,f25928,f25933,f25935,f25937,f25939,f25941,f25943,f25945,f25998,f26013,f26015,f26018,f26019,f26009,f26010,f26011,f26012,f26024,f26039,f26041,f26044,f26045,f26035,f26036,f26037,f26038,f26053,f26068,f26070,f26073,f26075,f26090,f26105,f26107,f26110,f26111,f26101,f26102,f26103,f26104,f26118,f26133,f26135,f26138,f26139,f26129,f26130,f26131,f26132,f26141,f26160,f26178,f26193,f26195,f26198,f26199,f26189,f26190,f26191,f26192,f26203,f26233,f26260,f26275,f26277,f26280,f26282,f26290,f26320,f26347,f26362,f26364,f26367,f26368,f26358,f26359,f26360,f26361,f26372,f26391,f26408,f26427,f26442,f26458,f26459,f26460,f26461,f26709,f26711,f26714,f26717,f26718,f26719,f26720,f26721,f26723,f26731,f26733,f26735,f26736,f26737,f26738,f26739,f26740,f26741,f26742,f26521,f26523,f26525,f26527,f26751,f26752,f26755,f26537,f26540,f26542,f26543,f26545,f26757,f26759,f26761,f26763,f26764,f26765,f26767,f26768,f26555,f26557,f26560,f26772,f26777,f26783,f26803,f26820,f26840,f26865,f26866,f26867,f26868,f26869,f26921,f26923,f26925,f26927,f26928,f26948,f26949,f26950,f26951,f26952,f26953,f26954,f26955,f26957,f26959,f26960,f26961,f26962,f26963,f26966,f26967,f26968,f26970,f26971,f26972,f26973,f26974,f26975,f26976,f26977,f26978,f26979,f26984,f26985,f26987,f26989,f26991,f26993,f26995,f26996,f26997,f26998,f26999,f27000,f27001,f27002,f27003,f27006,f27008,f27009,f27011,f27013,f27014,f27015,f27016,f27017,f27018,f27019,f27020,f27021,f27023,f27026,f27030,f27033,f27043,f27106,f27107,f27109,f27119,f27132,f27133,f27134,f27135,f27136,f27137,f27138,f27139,f27140,f27141,f27155,f27156,f27157,f27158,f27159,f27160,f27161,f27162,f27166,f27167,f27168,f27169,f27199,f25807,f27298,f27559,f27562,f27567,f27568,f27569,f27570,f28094,f28095,f28097,f28099,f27582,f27583,f27585,f27587,f27589,f27590,f27592,f27594,f27597,f27599,f27602,f27606,f27607,f27608,f27609,f27610,f27611,f27613,f27614,f27619,f27620,f27635,f27636,f27639,f27640,f27641,f27642,f27643,f27645,f27665,f28144,f28149,f28151,f28169,f28171,f28173,f28175,f28178,f28181,f28184,f28187,f28190,f28192,f28194,f28196,f28198,f28206])).
fof(f28199,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f334,f333,f335,f343,f355,f354,f413,f428,f427,f426,f429,f430,f437,f434,f444,f443,f593,f462,f459,f458,f457,f478,f479,f480,f481,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f819,f825,f828,f830,f840,f888,f898,f901,f917,f923,f933,f939,f946,f954,f960,f966,f971,f975,f980,f984,f994,f1000,f1007,f1012,f1024,f1032,f1047,f1050,f1062,f1068,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1239,f1275,f1339,f1332,f1373,f1385,f1402,f1422,f1433,f1437,f1461,f1501,f1512,f1516,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1662,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1981,f2015,f2030,f2073,f2082,f2119,f2118,f2128,f2157,f2170,f2252,f2222,f2262,f2270,f2617,f2622,f2841,f2844,f2751,f2544,f2504,f2465,f2432,f2400,f2371,f2351,f2317,f1837,f1818,f424,f423,f422,f421,f420,f419,f418,f417,f412,f411,f1854,f1861,f1852,f1859,f2052,f2050,f2043,f2041,f482,f433,f1542,f1492,f1486,f489,f488,f487,f485,f484,f1219,f1249,f1217,f1247,f1212,f1256,f1285,f1210,f1254,f1283,f2099,f2097,f2090,f2088,f2001,f1931,f1230,f483,f477,f468,f467,f432,f431,f2848,f2847,f2770,f2765,f2545,f2505,f2466,f2433,f2401,f2398,f2373,f2372,f2352,f2349,f2319,f2318,f2272,f2271,f2269,f2251,f1838,f1835,f1819,f1816,f1663,f1661,f1494,f1411,f1410,f1403,f1366,f1365,f1333,f1267,f1231,f1069,f1059,f1057,f1048,f1035,f1033,f1025,f1023,f1001,f991,f961,f942,f940,f930,f926,f924,f904,f902,f891,f889,f873,f871,f870,f869,f868,f867,f866,f865,f864,f863,f862,f861,f852,f839,f838,f837,f836,f829,f818,f787,f592,f414,f359,f358,f4838,f4843,f4844,f4853,f4855,f4857,f4858,f4859,f4860,f4861,f4862,f4863,f4864,f4644,f4646,f4648,f4650,f4870,f4655,f4658,f4660,f4661,f4663,f4669,f4977,f4986,f5268,f5276,f7384,f7394,f7492,f7497,f7506,f7666,f7671,f7680,f7351,f11643,f7360,f11657,f11665,f12884,f12883,f12097,f12096,f13263,f13282,f13283,f13284,f13285,f13298,f13312,f13313,f13314,f13315,f13323,f13324,f13325,f13326,f13327,f13328,f13329,f13331,f13333,f13335,f13337,f13343,f13346,f13348,f13349,f13351,f13356,f13358,f13361,f16694,f16695,f13454,f16696,f13456,f16697,f16698,f16699,f16700,f16701,f16717,f16728,f18264,f18263,f17257,f17256,f18761,f18782,f18783,f18784,f18785,f18786,f18802,f18816,f18817,f18818,f18819,f18828,f18829,f18830,f18831,f18832,f18833,f18834,f18836,f18838,f18840,f18842,f18848,f18851,f18853,f18854,f18856,f18861,f18863,f18866,f23086,f23087,f18961,f23088,f18963,f23089,f23090,f23091,f23092,f23093,f23129,f23139,f23156,f23157,f23158,f23159,f23173,f23183,f23184,f23185,f23186,f25796,f23209,f23235,f23237,f23239,f23241,f23247,f23250,f23252,f23253,f23255,f23261,f23263,f23266,f25806,f23476,f23477,f23508,f23513,f23514,f23515,f23516,f23517,f23518,f23519,f23520,f23521,f23522,f23523,f25824,f23848,f23849,f23850,f23851,f23852,f23853,f23854,f23855,f23856,f23857,f23858,f23859,f23860,f23861,f23862,f23863,f23864,f23865,f23866,f23867,f23868,f23869,f23870,f23871,f23872,f23873,f23874,f23875,f23876,f23877,f23878,f23879,f23880,f23881,f23882,f23883,f23884,f23885,f23886,f23887,f23888,f23889,f23890,f23891,f23892,f23893,f23894,f23895,f23896,f23897,f23898,f23899,f23900,f23901,f23902,f23903,f23904,f23905,f23906,f23907,f23908,f23909,f23910,f23911,f23912,f23913,f23914,f23915,f23916,f23917,f23918,f23919,f23920,f23921,f23922,f23923,f23924,f23925,f23926,f23927,f23928,f23929,f23930,f23931,f23932,f23933,f23934,f23935,f23936,f23937,f23938,f23939,f25830,f23947,f25837,f25840,f25846,f25849,f25852,f25855,f25858,f25861,f25864,f23973,f23975,f23977,f23979,f23985,f23988,f23990,f23991,f23993,f24001,f25872,f25873,f24118,f25875,f25876,f25877,f25878,f25879,f25880,f25883,f25884,f25885,f25886,f25887,f25888,f25889,f25890,f25891,f24144,f24146,f24148,f24150,f25892,f25893,f25894,f25895,f24156,f24159,f24161,f24162,f24164,f25896,f24170,f24172,f24175,f25897,f25898,f25900,f25901,f25902,f25903,f25904,f25905,f25906,f25907,f25908,f25909,f25910,f25911,f25912,f25913,f25914,f25915,f25916,f25917,f25918,f24416,f24421,f24422,f24423,f24424,f24425,f24426,f24427,f24428,f24429,f24430,f24431,f25920,f25926,f25928,f25933,f25935,f25937,f25939,f25941,f25943,f25945,f25998,f26013,f26015,f26018,f26019,f26009,f26010,f26011,f26012,f26024,f26039,f26041,f26044,f26045,f26035,f26036,f26037,f26038,f26053,f26068,f26070,f26073,f26075,f26090,f26105,f26107,f26110,f26111,f26101,f26102,f26103,f26104,f26118,f26133,f26135,f26138,f26139,f26129,f26130,f26131,f26132,f26141,f26160,f26178,f26193,f26195,f26198,f26199,f26189,f26190,f26191,f26192,f26203,f26233,f26260,f26275,f26277,f26280,f26282,f26290,f26320,f26347,f26362,f26364,f26367,f26368,f26358,f26359,f26360,f26361,f26372,f26391,f26408,f26427,f26442,f26458,f26459,f26460,f26461,f26709,f26711,f26714,f26717,f26718,f26719,f26720,f26721,f26723,f26731,f26733,f26735,f26736,f26737,f26738,f26739,f26740,f26741,f26742,f26521,f26523,f26525,f26527,f26751,f26752,f26755,f26537,f26540,f26542,f26543,f26545,f26757,f26759,f26761,f26763,f26764,f26765,f26767,f26768,f26555,f26557,f26560,f26772,f26777,f26783,f26803,f26820,f26840,f26865,f26866,f26867,f26868,f26869,f26921,f26923,f26925,f26927,f26928,f26948,f26949,f26950,f26951,f26952,f26953,f26954,f26955,f26957,f26959,f26960,f26961,f26962,f26963,f26966,f26967,f26968,f26970,f26971,f26972,f26973,f26974,f26975,f26976,f26977,f26978,f26979,f26984,f26985,f26987,f26989,f26991,f26993,f26995,f26996,f26997,f26998,f26999,f27000,f27001,f27002,f27003,f27006,f27008,f27009,f27011,f27013,f27014,f27015,f27016,f27017,f27018,f27019,f27020,f27021,f27023,f27026,f27030,f27033,f27043,f27106,f27107,f27109,f27119,f27132,f27133,f27134,f27135,f27136,f27137,f27138,f27139,f27140,f27141,f27155,f27156,f27157,f27158,f27159,f27160,f27161,f27162,f27166,f27167,f27168,f27169,f27199,f25807,f27298,f27559,f27562,f27567,f27568,f27569,f27570,f28094,f28095,f28097,f28099,f27582,f27583,f27585,f27587,f27589,f27590,f27592,f27594,f27597,f27599,f27602,f27606,f27607,f27608,f27609,f27610,f27611,f27613,f27614,f27619,f27620,f27635,f27636,f27639,f27640,f27641,f27642,f27643,f27645,f27665,f28144,f28149,f28151,f28169,f28171,f28173,f28175,f28178,f28181,f28184,f28187,f28190,f28192,f28194,f28196,f28198])).
fof(f28197,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f334,f333,f335,f343,f355,f354,f413,f428,f427,f426,f429,f430,f437,f434,f444,f443,f593,f462,f459,f458,f457,f478,f479,f480,f481,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f819,f825,f828,f830,f840,f888,f898,f901,f917,f923,f933,f939,f946,f954,f960,f966,f971,f975,f980,f984,f994,f1000,f1007,f1012,f1024,f1032,f1047,f1050,f1062,f1068,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1239,f1275,f1339,f1332,f1373,f1385,f1402,f1422,f1433,f1437,f1461,f1501,f1512,f1516,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1662,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1981,f2015,f2030,f2073,f2082,f2119,f2118,f2128,f2157,f2170,f2252,f2222,f2262,f2270,f2617,f2622,f2841,f2844,f2751,f2544,f2504,f2465,f2432,f2400,f2371,f2351,f2317,f1837,f1818,f424,f423,f422,f421,f420,f419,f418,f417,f412,f411,f1854,f1861,f1852,f1859,f2052,f2050,f2043,f2041,f482,f433,f1542,f1492,f1486,f489,f488,f487,f485,f484,f1219,f1249,f1217,f1247,f1212,f1256,f1285,f1210,f1254,f1283,f2099,f2097,f2090,f2088,f2001,f1931,f1230,f483,f477,f468,f467,f432,f431,f2848,f2847,f2770,f2765,f2545,f2505,f2466,f2433,f2401,f2398,f2373,f2372,f2352,f2349,f2319,f2318,f2272,f2271,f2269,f2251,f1838,f1835,f1819,f1816,f1663,f1661,f1494,f1411,f1410,f1403,f1366,f1365,f1333,f1267,f1231,f1069,f1059,f1057,f1048,f1035,f1033,f1025,f1023,f1001,f991,f961,f942,f940,f930,f926,f924,f904,f902,f891,f889,f873,f871,f870,f869,f868,f867,f866,f865,f864,f863,f862,f861,f852,f839,f838,f837,f836,f829,f818,f787,f592,f414,f359,f358,f4838,f4843,f4844,f4853,f4855,f4857,f4858,f4859,f4860,f4861,f4862,f4863,f4864,f4644,f4646,f4648,f4650,f4870,f4655,f4658,f4660,f4661,f4663,f4669,f4977,f4986,f5268,f5276,f7384,f7394,f7492,f7497,f7506,f7666,f7671,f7680,f7351,f11643,f7360,f11657,f11665,f12884,f12883,f12097,f12096,f13263,f13282,f13283,f13284,f13285,f13298,f13312,f13313,f13314,f13315,f13323,f13324,f13325,f13326,f13327,f13328,f13329,f13331,f13333,f13335,f13337,f13343,f13346,f13348,f13349,f13351,f13356,f13358,f13361,f16694,f16695,f13454,f16696,f13456,f16697,f16698,f16699,f16700,f16701,f16717,f16728,f18264,f18263,f17257,f17256,f18761,f18782,f18783,f18784,f18785,f18786,f18802,f18816,f18817,f18818,f18819,f18828,f18829,f18830,f18831,f18832,f18833,f18834,f18836,f18838,f18840,f18842,f18848,f18851,f18853,f18854,f18856,f18861,f18863,f18866,f23086,f23087,f18961,f23088,f18963,f23089,f23090,f23091,f23092,f23093,f23129,f23139,f23156,f23157,f23158,f23159,f23173,f23183,f23184,f23185,f23186,f25796,f23209,f23235,f23237,f23239,f23241,f23247,f23250,f23252,f23253,f23255,f23261,f23263,f23266,f25806,f23476,f23477,f23508,f23513,f23514,f23515,f23516,f23517,f23518,f23519,f23520,f23521,f23522,f23523,f25824,f23848,f23849,f23850,f23851,f23852,f23853,f23854,f23855,f23856,f23857,f23858,f23859,f23860,f23861,f23862,f23863,f23864,f23865,f23866,f23867,f23868,f23869,f23870,f23871,f23872,f23873,f23874,f23875,f23876,f23877,f23878,f23879,f23880,f23881,f23882,f23883,f23884,f23885,f23886,f23887,f23888,f23889,f23890,f23891,f23892,f23893,f23894,f23895,f23896,f23897,f23898,f23899,f23900,f23901,f23902,f23903,f23904,f23905,f23906,f23907,f23908,f23909,f23910,f23911,f23912,f23913,f23914,f23915,f23916,f23917,f23918,f23919,f23920,f23921,f23922,f23923,f23924,f23925,f23926,f23927,f23928,f23929,f23930,f23931,f23932,f23933,f23934,f23935,f23936,f23937,f23938,f23939,f25830,f23947,f25837,f25840,f25846,f25849,f25852,f25855,f25858,f25861,f25864,f23973,f23975,f23977,f23979,f23985,f23988,f23990,f23991,f23993,f24001,f25872,f25873,f24118,f25875,f25876,f25877,f25878,f25879,f25880,f25883,f25884,f25885,f25886,f25887,f25888,f25889,f25890,f25891,f24144,f24146,f24148,f24150,f25892,f25893,f25894,f25895,f24156,f24159,f24161,f24162,f24164,f25896,f24170,f24172,f24175,f25897,f25898,f25900,f25901,f25902,f25903,f25904,f25905,f25906,f25907,f25908,f25909,f25910,f25911,f25912,f25913,f25914,f25915,f25916,f25917,f25918,f24416,f24421,f24422,f24423,f24424,f24425,f24426,f24427,f24428,f24429,f24430,f24431,f25920,f25926,f25928,f25933,f25935,f25937,f25939,f25941,f25943,f25945,f25998,f26013,f26015,f26018,f26019,f26009,f26010,f26011,f26012,f26024,f26039,f26041,f26044,f26045,f26035,f26036,f26037,f26038,f26053,f26068,f26070,f26073,f26075,f26090,f26105,f26107,f26110,f26111,f26101,f26102,f26103,f26104,f26118,f26133,f26135,f26138,f26139,f26129,f26130,f26131,f26132,f26141,f26160,f26178,f26193,f26195,f26198,f26199,f26189,f26190,f26191,f26192,f26203,f26233,f26260,f26275,f26277,f26280,f26282,f26290,f26320,f26347,f26362,f26364,f26367,f26368,f26358,f26359,f26360,f26361,f26372,f26391,f26408,f26427,f26442,f26458,f26459,f26460,f26461,f26709,f26711,f26714,f26717,f26718,f26719,f26720,f26721,f26723,f26731,f26733,f26735,f26736,f26737,f26738,f26739,f26740,f26741,f26742,f26521,f26523,f26525,f26527,f26751,f26752,f26755,f26537,f26540,f26542,f26543,f26545,f26757,f26759,f26761,f26763,f26764,f26765,f26767,f26768,f26555,f26557,f26560,f26772,f26777,f26783,f26803,f26820,f26840,f26865,f26866,f26867,f26868,f26869,f26921,f26923,f26925,f26927,f26928,f26948,f26949,f26950,f26951,f26952,f26953,f26954,f26955,f26957,f26959,f26960,f26961,f26962,f26963,f26966,f26967,f26968,f26970,f26971,f26972,f26973,f26974,f26975,f26976,f26977,f26978,f26979,f26984,f26985,f26987,f26989,f26991,f26993,f26995,f26996,f26997,f26998,f26999,f27000,f27001,f27002,f27003,f27006,f27008,f27009,f27011,f27013,f27014,f27015,f27016,f27017,f27018,f27019,f27020,f27021,f27023,f27026,f27030,f27033,f27043,f27106,f27107,f27109,f27119,f27132,f27133,f27134,f27135,f27136,f27137,f27138,f27139,f27140,f27141,f27155,f27156,f27157,f27158,f27159,f27160,f27161,f27162,f27166,f27167,f27168,f27169,f27199,f25807,f27298,f27559,f27562,f27567,f27568,f27569,f27570,f28094,f28095,f28097,f28099,f27582,f27583,f27585,f27587,f27589,f27590,f27592,f27594,f27597,f27599,f27602,f27606,f27607,f27608,f27609,f27610,f27611,f27613,f27614,f27619,f27620,f27635,f27636,f27639,f27640,f27641,f27642,f27643,f27645,f27665,f28144,f28149,f28151,f28169,f28171,f28173,f28175,f28178,f28181,f28184,f28187,f28190,f28192,f28194,f28196])).
fof(f28195,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f334,f333,f335,f343,f355,f354,f413,f428,f427,f426,f429,f430,f437,f434,f444,f443,f593,f462,f459,f458,f457,f478,f479,f480,f481,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f819,f825,f828,f830,f840,f888,f898,f901,f917,f923,f933,f939,f946,f954,f960,f966,f971,f975,f980,f984,f994,f1000,f1007,f1012,f1024,f1032,f1047,f1050,f1062,f1068,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1239,f1275,f1339,f1332,f1373,f1385,f1402,f1422,f1433,f1437,f1461,f1501,f1512,f1516,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1662,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1981,f2015,f2030,f2073,f2082,f2119,f2118,f2128,f2157,f2170,f2252,f2222,f2262,f2270,f2617,f2622,f2841,f2844,f2751,f2544,f2504,f2465,f2432,f2400,f2371,f2351,f2317,f1837,f1818,f424,f423,f422,f421,f420,f419,f418,f417,f412,f411,f1854,f1861,f1852,f1859,f2052,f2050,f2043,f2041,f482,f433,f1542,f1492,f1486,f489,f488,f487,f485,f484,f1219,f1249,f1217,f1247,f1212,f1256,f1285,f1210,f1254,f1283,f2099,f2097,f2090,f2088,f2001,f1931,f1230,f483,f477,f468,f467,f432,f431,f2848,f2847,f2770,f2765,f2545,f2505,f2466,f2433,f2401,f2398,f2373,f2372,f2352,f2349,f2319,f2318,f2272,f2271,f2269,f2251,f1838,f1835,f1819,f1816,f1663,f1661,f1494,f1411,f1410,f1403,f1366,f1365,f1333,f1267,f1231,f1069,f1059,f1057,f1048,f1035,f1033,f1025,f1023,f1001,f991,f961,f942,f940,f930,f926,f924,f904,f902,f891,f889,f873,f871,f870,f869,f868,f867,f866,f865,f864,f863,f862,f861,f852,f839,f838,f837,f836,f829,f818,f787,f592,f414,f359,f358,f4838,f4843,f4844,f4853,f4855,f4857,f4858,f4859,f4860,f4861,f4862,f4863,f4864,f4644,f4646,f4648,f4650,f4870,f4655,f4658,f4660,f4661,f4663,f4669,f4977,f4986,f5268,f5276,f7384,f7394,f7492,f7497,f7506,f7666,f7671,f7680,f7351,f11643,f7360,f11657,f11665,f12884,f12883,f12097,f12096,f13263,f13282,f13283,f13284,f13285,f13298,f13312,f13313,f13314,f13315,f13323,f13324,f13325,f13326,f13327,f13328,f13329,f13331,f13333,f13335,f13337,f13343,f13346,f13348,f13349,f13351,f13356,f13358,f13361,f16694,f16695,f13454,f16696,f13456,f16697,f16698,f16699,f16700,f16701,f16717,f16728,f18264,f18263,f17257,f17256,f18761,f18782,f18783,f18784,f18785,f18786,f18802,f18816,f18817,f18818,f18819,f18828,f18829,f18830,f18831,f18832,f18833,f18834,f18836,f18838,f18840,f18842,f18848,f18851,f18853,f18854,f18856,f18861,f18863,f18866,f23086,f23087,f18961,f23088,f18963,f23089,f23090,f23091,f23092,f23093,f23129,f23139,f23156,f23157,f23158,f23159,f23173,f23183,f23184,f23185,f23186,f25796,f23209,f23235,f23237,f23239,f23241,f23247,f23250,f23252,f23253,f23255,f23261,f23263,f23266,f25806,f23476,f23477,f23508,f23513,f23514,f23515,f23516,f23517,f23518,f23519,f23520,f23521,f23522,f23523,f25824,f23848,f23849,f23850,f23851,f23852,f23853,f23854,f23855,f23856,f23857,f23858,f23859,f23860,f23861,f23862,f23863,f23864,f23865,f23866,f23867,f23868,f23869,f23870,f23871,f23872,f23873,f23874,f23875,f23876,f23877,f23878,f23879,f23880,f23881,f23882,f23883,f23884,f23885,f23886,f23887,f23888,f23889,f23890,f23891,f23892,f23893,f23894,f23895,f23896,f23897,f23898,f23899,f23900,f23901,f23902,f23903,f23904,f23905,f23906,f23907,f23908,f23909,f23910,f23911,f23912,f23913,f23914,f23915,f23916,f23917,f23918,f23919,f23920,f23921,f23922,f23923,f23924,f23925,f23926,f23927,f23928,f23929,f23930,f23931,f23932,f23933,f23934,f23935,f23936,f23937,f23938,f23939,f25830,f23947,f25837,f25840,f25846,f25849,f25852,f25855,f25858,f25861,f25864,f23973,f23975,f23977,f23979,f23985,f23988,f23990,f23991,f23993,f24001,f25872,f25873,f24118,f25875,f25876,f25877,f25878,f25879,f25880,f25883,f25884,f25885,f25886,f25887,f25888,f25889,f25890,f25891,f24144,f24146,f24148,f24150,f25892,f25893,f25894,f25895,f24156,f24159,f24161,f24162,f24164,f25896,f24170,f24172,f24175,f25897,f25898,f25900,f25901,f25902,f25903,f25904,f25905,f25906,f25907,f25908,f25909,f25910,f25911,f25912,f25913,f25914,f25915,f25916,f25917,f25918,f24416,f24421,f24422,f24423,f24424,f24425,f24426,f24427,f24428,f24429,f24430,f24431,f25920,f25926,f25928,f25933,f25935,f25937,f25939,f25941,f25943,f25945,f25998,f26013,f26015,f26018,f26019,f26009,f26010,f26011,f26012,f26024,f26039,f26041,f26044,f26045,f26035,f26036,f26037,f26038,f26053,f26068,f26070,f26073,f26075,f26090,f26105,f26107,f26110,f26111,f26101,f26102,f26103,f26104,f26118,f26133,f26135,f26138,f26139,f26129,f26130,f26131,f26132,f26141,f26160,f26178,f26193,f26195,f26198,f26199,f26189,f26190,f26191,f26192,f26203,f26233,f26260,f26275,f26277,f26280,f26282,f26290,f26320,f26347,f26362,f26364,f26367,f26368,f26358,f26359,f26360,f26361,f26372,f26391,f26408,f26427,f26442,f26458,f26459,f26460,f26461,f26709,f26711,f26714,f26717,f26718,f26719,f26720,f26721,f26723,f26731,f26733,f26735,f26736,f26737,f26738,f26739,f26740,f26741,f26742,f26521,f26523,f26525,f26527,f26751,f26752,f26755,f26537,f26540,f26542,f26543,f26545,f26757,f26759,f26761,f26763,f26764,f26765,f26767,f26768,f26555,f26557,f26560,f26772,f26777,f26783,f26803,f26820,f26840,f26865,f26866,f26867,f26868,f26869,f26921,f26923,f26925,f26927,f26928,f26948,f26949,f26950,f26951,f26952,f26953,f26954,f26955,f26957,f26959,f26960,f26961,f26962,f26963,f26966,f26967,f26968,f26970,f26971,f26972,f26973,f26974,f26975,f26976,f26977,f26978,f26979,f26984,f26985,f26987,f26989,f26991,f26993,f26995,f26996,f26997,f26998,f26999,f27000,f27001,f27002,f27003,f27006,f27008,f27009,f27011,f27013,f27014,f27015,f27016,f27017,f27018,f27019,f27020,f27021,f27023,f27026,f27030,f27033,f27043,f27106,f27107,f27109,f27119,f27132,f27133,f27134,f27135,f27136,f27137,f27138,f27139,f27140,f27141,f27155,f27156,f27157,f27158,f27159,f27160,f27161,f27162,f27166,f27167,f27168,f27169,f27199,f25807,f27298,f27559,f27562,f27567,f27568,f27569,f27570,f28094,f28095,f28097,f28099,f27582,f27583,f27585,f27587,f27589,f27590,f27592,f27594,f27597,f27599,f27602,f27606,f27607,f27608,f27609,f27610,f27611,f27613,f27614,f27619,f27620,f27635,f27636,f27639,f27640,f27641,f27642,f27643,f27645,f27665,f28144,f28149,f28151,f28169,f28171,f28173,f28175,f28178,f28181,f28184,f28187,f28190,f28192,f28194])).
fof(f28193,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f334,f333,f335,f343,f355,f354,f413,f428,f427,f426,f429,f430,f437,f434,f444,f443,f593,f462,f459,f458,f457,f478,f479,f480,f481,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f819,f825,f828,f830,f840,f888,f898,f901,f917,f923,f933,f939,f946,f954,f960,f966,f971,f975,f980,f984,f994,f1000,f1007,f1012,f1024,f1032,f1047,f1050,f1062,f1068,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1239,f1275,f1339,f1332,f1373,f1385,f1402,f1422,f1433,f1437,f1461,f1501,f1512,f1516,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1662,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1981,f2015,f2030,f2073,f2082,f2119,f2118,f2128,f2157,f2170,f2252,f2222,f2262,f2270,f2617,f2622,f2841,f2844,f2751,f2544,f2504,f2465,f2432,f2400,f2371,f2351,f2317,f1837,f1818,f424,f423,f422,f421,f420,f419,f418,f417,f412,f411,f1854,f1861,f1852,f1859,f2052,f2050,f2043,f2041,f482,f433,f1542,f1492,f1486,f489,f488,f487,f485,f484,f1219,f1249,f1217,f1247,f1212,f1256,f1285,f1210,f1254,f1283,f2099,f2097,f2090,f2088,f2001,f1931,f1230,f483,f477,f468,f467,f432,f431,f2848,f2847,f2770,f2765,f2545,f2505,f2466,f2433,f2401,f2398,f2373,f2372,f2352,f2349,f2319,f2318,f2272,f2271,f2269,f2251,f1838,f1835,f1819,f1816,f1663,f1661,f1494,f1411,f1410,f1403,f1366,f1365,f1333,f1267,f1231,f1069,f1059,f1057,f1048,f1035,f1033,f1025,f1023,f1001,f991,f961,f942,f940,f930,f926,f924,f904,f902,f891,f889,f873,f871,f870,f869,f868,f867,f866,f865,f864,f863,f862,f861,f852,f839,f838,f837,f836,f829,f818,f787,f592,f414,f359,f358,f4838,f4843,f4844,f4853,f4855,f4857,f4858,f4859,f4860,f4861,f4862,f4863,f4864,f4644,f4646,f4648,f4650,f4870,f4655,f4658,f4660,f4661,f4663,f4669,f4977,f4986,f5268,f5276,f7384,f7394,f7492,f7497,f7506,f7666,f7671,f7680,f7351,f11643,f7360,f11657,f11665,f12884,f12883,f12097,f12096,f13263,f13282,f13283,f13284,f13285,f13298,f13312,f13313,f13314,f13315,f13323,f13324,f13325,f13326,f13327,f13328,f13329,f13331,f13333,f13335,f13337,f13343,f13346,f13348,f13349,f13351,f13356,f13358,f13361,f16694,f16695,f13454,f16696,f13456,f16697,f16698,f16699,f16700,f16701,f16717,f16728,f18264,f18263,f17257,f17256,f18761,f18782,f18783,f18784,f18785,f18786,f18802,f18816,f18817,f18818,f18819,f18828,f18829,f18830,f18831,f18832,f18833,f18834,f18836,f18838,f18840,f18842,f18848,f18851,f18853,f18854,f18856,f18861,f18863,f18866,f23086,f23087,f18961,f23088,f18963,f23089,f23090,f23091,f23092,f23093,f23129,f23139,f23156,f23157,f23158,f23159,f23173,f23183,f23184,f23185,f23186,f25796,f23209,f23235,f23237,f23239,f23241,f23247,f23250,f23252,f23253,f23255,f23261,f23263,f23266,f25806,f23476,f23477,f23508,f23513,f23514,f23515,f23516,f23517,f23518,f23519,f23520,f23521,f23522,f23523,f25824,f23848,f23849,f23850,f23851,f23852,f23853,f23854,f23855,f23856,f23857,f23858,f23859,f23860,f23861,f23862,f23863,f23864,f23865,f23866,f23867,f23868,f23869,f23870,f23871,f23872,f23873,f23874,f23875,f23876,f23877,f23878,f23879,f23880,f23881,f23882,f23883,f23884,f23885,f23886,f23887,f23888,f23889,f23890,f23891,f23892,f23893,f23894,f23895,f23896,f23897,f23898,f23899,f23900,f23901,f23902,f23903,f23904,f23905,f23906,f23907,f23908,f23909,f23910,f23911,f23912,f23913,f23914,f23915,f23916,f23917,f23918,f23919,f23920,f23921,f23922,f23923,f23924,f23925,f23926,f23927,f23928,f23929,f23930,f23931,f23932,f23933,f23934,f23935,f23936,f23937,f23938,f23939,f25830,f23947,f25837,f25840,f25846,f25849,f25852,f25855,f25858,f25861,f25864,f23973,f23975,f23977,f23979,f23985,f23988,f23990,f23991,f23993,f24001,f25872,f25873,f24118,f25875,f25876,f25877,f25878,f25879,f25880,f25883,f25884,f25885,f25886,f25887,f25888,f25889,f25890,f25891,f24144,f24146,f24148,f24150,f25892,f25893,f25894,f25895,f24156,f24159,f24161,f24162,f24164,f25896,f24170,f24172,f24175,f25897,f25898,f25900,f25901,f25902,f25903,f25904,f25905,f25906,f25907,f25908,f25909,f25910,f25911,f25912,f25913,f25914,f25915,f25916,f25917,f25918,f24416,f24421,f24422,f24423,f24424,f24425,f24426,f24427,f24428,f24429,f24430,f24431,f25920,f25926,f25928,f25933,f25935,f25937,f25939,f25941,f25943,f25945,f25998,f26013,f26015,f26018,f26019,f26009,f26010,f26011,f26012,f26024,f26039,f26041,f26044,f26045,f26035,f26036,f26037,f26038,f26053,f26068,f26070,f26073,f26075,f26090,f26105,f26107,f26110,f26111,f26101,f26102,f26103,f26104,f26118,f26133,f26135,f26138,f26139,f26129,f26130,f26131,f26132,f26141,f26160,f26178,f26193,f26195,f26198,f26199,f26189,f26190,f26191,f26192,f26203,f26233,f26260,f26275,f26277,f26280,f26282,f26290,f26320,f26347,f26362,f26364,f26367,f26368,f26358,f26359,f26360,f26361,f26372,f26391,f26408,f26427,f26442,f26458,f26459,f26460,f26461,f26709,f26711,f26714,f26717,f26718,f26719,f26720,f26721,f26723,f26731,f26733,f26735,f26736,f26737,f26738,f26739,f26740,f26741,f26742,f26521,f26523,f26525,f26527,f26751,f26752,f26755,f26537,f26540,f26542,f26543,f26545,f26757,f26759,f26761,f26763,f26764,f26765,f26767,f26768,f26555,f26557,f26560,f26772,f26777,f26783,f26803,f26820,f26840,f26865,f26866,f26867,f26868,f26869,f26921,f26923,f26925,f26927,f26928,f26948,f26949,f26950,f26951,f26952,f26953,f26954,f26955,f26957,f26959,f26960,f26961,f26962,f26963,f26966,f26967,f26968,f26970,f26971,f26972,f26973,f26974,f26975,f26976,f26977,f26978,f26979,f26984,f26985,f26987,f26989,f26991,f26993,f26995,f26996,f26997,f26998,f26999,f27000,f27001,f27002,f27003,f27006,f27008,f27009,f27011,f27013,f27014,f27015,f27016,f27017,f27018,f27019,f27020,f27021,f27023,f27026,f27030,f27033,f27043,f27106,f27107,f27109,f27119,f27132,f27133,f27134,f27135,f27136,f27137,f27138,f27139,f27140,f27141,f27155,f27156,f27157,f27158,f27159,f27160,f27161,f27162,f27166,f27167,f27168,f27169,f27199,f25807,f27298,f27559,f27562,f27567,f27568,f27569,f27570,f28094,f28095,f28097,f28099,f27582,f27583,f27585,f27587,f27589,f27590,f27592,f27594,f27597,f27599,f27602,f27606,f27607,f27608,f27609,f27610,f27611,f27613,f27614,f27619,f27620,f27635,f27636,f27639,f27640,f27641,f27642,f27643,f27645,f27665,f28144,f28149,f28151,f28169,f28171,f28173,f28175,f28178,f28181,f28184,f28187,f28190,f28192])).
fof(f28191,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f334,f333,f335,f343,f355,f354,f413,f428,f427,f426,f429,f430,f437,f434,f444,f443,f593,f462,f459,f458,f457,f478,f479,f480,f481,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f819,f825,f828,f830,f840,f888,f898,f901,f917,f923,f933,f939,f946,f954,f960,f966,f971,f975,f980,f984,f994,f1000,f1007,f1012,f1024,f1032,f1047,f1050,f1062,f1068,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1239,f1275,f1339,f1332,f1373,f1385,f1402,f1422,f1433,f1437,f1461,f1501,f1512,f1516,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1662,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1981,f2015,f2030,f2073,f2082,f2119,f2118,f2128,f2157,f2170,f2252,f2222,f2262,f2270,f2617,f2622,f2841,f2844,f2751,f2544,f2504,f2465,f2432,f2400,f2371,f2351,f2317,f1837,f1818,f424,f423,f422,f421,f420,f419,f418,f417,f412,f411,f1854,f1861,f1852,f1859,f2052,f2050,f2043,f2041,f482,f433,f1542,f1492,f1486,f489,f488,f487,f485,f484,f1219,f1249,f1217,f1247,f1212,f1256,f1285,f1210,f1254,f1283,f2099,f2097,f2090,f2088,f2001,f1931,f1230,f483,f477,f468,f467,f432,f431,f2848,f2847,f2770,f2765,f2545,f2505,f2466,f2433,f2401,f2398,f2373,f2372,f2352,f2349,f2319,f2318,f2272,f2271,f2269,f2251,f1838,f1835,f1819,f1816,f1663,f1661,f1494,f1411,f1410,f1403,f1366,f1365,f1333,f1267,f1231,f1069,f1059,f1057,f1048,f1035,f1033,f1025,f1023,f1001,f991,f961,f942,f940,f930,f926,f924,f904,f902,f891,f889,f873,f871,f870,f869,f868,f867,f866,f865,f864,f863,f862,f861,f852,f839,f838,f837,f836,f829,f818,f787,f592,f414,f359,f358,f4838,f4843,f4844,f4853,f4855,f4857,f4858,f4859,f4860,f4861,f4862,f4863,f4864,f4644,f4646,f4648,f4650,f4870,f4655,f4658,f4660,f4661,f4663,f4669,f4977,f4986,f5268,f5276,f7384,f7394,f7492,f7497,f7506,f7666,f7671,f7680,f7351,f11643,f7360,f11657,f11665,f12884,f12883,f12097,f12096,f13263,f13282,f13283,f13284,f13285,f13298,f13312,f13313,f13314,f13315,f13323,f13324,f13325,f13326,f13327,f13328,f13329,f13331,f13333,f13335,f13337,f13343,f13346,f13348,f13349,f13351,f13356,f13358,f13361,f16694,f16695,f13454,f16696,f13456,f16697,f16698,f16699,f16700,f16701,f16717,f16728,f18264,f18263,f17257,f17256,f18761,f18782,f18783,f18784,f18785,f18786,f18802,f18816,f18817,f18818,f18819,f18828,f18829,f18830,f18831,f18832,f18833,f18834,f18836,f18838,f18840,f18842,f18848,f18851,f18853,f18854,f18856,f18861,f18863,f18866,f23086,f23087,f18961,f23088,f18963,f23089,f23090,f23091,f23092,f23093,f23129,f23139,f23156,f23157,f23158,f23159,f23173,f23183,f23184,f23185,f23186,f25796,f23209,f23235,f23237,f23239,f23241,f23247,f23250,f23252,f23253,f23255,f23261,f23263,f23266,f25806,f23476,f23477,f23508,f23513,f23514,f23515,f23516,f23517,f23518,f23519,f23520,f23521,f23522,f23523,f25824,f23848,f23849,f23850,f23851,f23852,f23853,f23854,f23855,f23856,f23857,f23858,f23859,f23860,f23861,f23862,f23863,f23864,f23865,f23866,f23867,f23868,f23869,f23870,f23871,f23872,f23873,f23874,f23875,f23876,f23877,f23878,f23879,f23880,f23881,f23882,f23883,f23884,f23885,f23886,f23887,f23888,f23889,f23890,f23891,f23892,f23893,f23894,f23895,f23896,f23897,f23898,f23899,f23900,f23901,f23902,f23903,f23904,f23905,f23906,f23907,f23908,f23909,f23910,f23911,f23912,f23913,f23914,f23915,f23916,f23917,f23918,f23919,f23920,f23921,f23922,f23923,f23924,f23925,f23926,f23927,f23928,f23929,f23930,f23931,f23932,f23933,f23934,f23935,f23936,f23937,f23938,f23939,f25830,f23947,f25837,f25840,f25846,f25849,f25852,f25855,f25858,f25861,f25864,f23973,f23975,f23977,f23979,f23985,f23988,f23990,f23991,f23993,f24001,f25872,f25873,f24118,f25875,f25876,f25877,f25878,f25879,f25880,f25883,f25884,f25885,f25886,f25887,f25888,f25889,f25890,f25891,f24144,f24146,f24148,f24150,f25892,f25893,f25894,f25895,f24156,f24159,f24161,f24162,f24164,f25896,f24170,f24172,f24175,f25897,f25898,f25900,f25901,f25902,f25903,f25904,f25905,f25906,f25907,f25908,f25909,f25910,f25911,f25912,f25913,f25914,f25915,f25916,f25917,f25918,f24416,f24421,f24422,f24423,f24424,f24425,f24426,f24427,f24428,f24429,f24430,f24431,f25920,f25926,f25928,f25933,f25935,f25937,f25939,f25941,f25943,f25945,f25998,f26013,f26015,f26018,f26019,f26009,f26010,f26011,f26012,f26024,f26039,f26041,f26044,f26045,f26035,f26036,f26037,f26038,f26053,f26068,f26070,f26073,f26075,f26090,f26105,f26107,f26110,f26111,f26101,f26102,f26103,f26104,f26118,f26133,f26135,f26138,f26139,f26129,f26130,f26131,f26132,f26141,f26160,f26178,f26193,f26195,f26198,f26199,f26189,f26190,f26191,f26192,f26203,f26233,f26260,f26275,f26277,f26280,f26282,f26290,f26320,f26347,f26362,f26364,f26367,f26368,f26358,f26359,f26360,f26361,f26372,f26391,f26408,f26427,f26442,f26458,f26459,f26460,f26461,f26709,f26711,f26714,f26717,f26718,f26719,f26720,f26721,f26723,f26731,f26733,f26735,f26736,f26737,f26738,f26739,f26740,f26741,f26742,f26521,f26523,f26525,f26527,f26751,f26752,f26755,f26537,f26540,f26542,f26543,f26545,f26757,f26759,f26761,f26763,f26764,f26765,f26767,f26768,f26555,f26557,f26560,f26772,f26777,f26783,f26803,f26820,f26840,f26865,f26866,f26867,f26868,f26869,f26921,f26923,f26925,f26927,f26928,f26948,f26949,f26950,f26951,f26952,f26953,f26954,f26955,f26957,f26959,f26960,f26961,f26962,f26963,f26966,f26967,f26968,f26970,f26971,f26972,f26973,f26974,f26975,f26976,f26977,f26978,f26979,f26984,f26985,f26987,f26989,f26991,f26993,f26995,f26996,f26997,f26998,f26999,f27000,f27001,f27002,f27003,f27006,f27008,f27009,f27011,f27013,f27014,f27015,f27016,f27017,f27018,f27019,f27020,f27021,f27023,f27026,f27030,f27033,f27043,f27106,f27107,f27109,f27119,f27132,f27133,f27134,f27135,f27136,f27137,f27138,f27139,f27140,f27141,f27155,f27156,f27157,f27158,f27159,f27160,f27161,f27162,f27166,f27167,f27168,f27169,f27199,f25807,f27298,f27559,f27562,f27567,f27568,f27569,f27570,f28094,f28095,f28097,f28099,f27582,f27583,f27585,f27587,f27589,f27590,f27592,f27594,f27597,f27599,f27602,f27606,f27607,f27608,f27609,f27610,f27611,f27613,f27614,f27619,f27620,f27635,f27636,f27639,f27640,f27641,f27642,f27643,f27645,f27665,f28144,f28149,f28151,f28169,f28171,f28173,f28175,f28178,f28181,f28184,f28187,f28190])).
fof(f28188,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f334,f333,f335,f343,f355,f354,f413,f428,f427,f426,f429,f430,f437,f434,f444,f443,f593,f462,f459,f458,f457,f478,f479,f480,f481,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f819,f825,f828,f830,f840,f888,f898,f901,f917,f923,f933,f939,f946,f954,f960,f966,f971,f975,f980,f984,f994,f1000,f1007,f1012,f1024,f1032,f1047,f1050,f1062,f1068,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1239,f1275,f1339,f1332,f1373,f1385,f1402,f1422,f1433,f1437,f1461,f1501,f1512,f1516,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1662,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1981,f2015,f2030,f2073,f2082,f2119,f2118,f2128,f2157,f2170,f2252,f2222,f2262,f2270,f2617,f2622,f2841,f2844,f2751,f2544,f2504,f2465,f2432,f2400,f2371,f2351,f2317,f1837,f1818,f424,f423,f422,f421,f420,f419,f418,f417,f412,f411,f1854,f1861,f1852,f1859,f2052,f2050,f2043,f2041,f482,f433,f1542,f1492,f1486,f489,f488,f487,f485,f484,f1219,f1249,f1217,f1247,f1212,f1256,f1285,f1210,f1254,f1283,f2099,f2097,f2090,f2088,f2001,f1931,f1230,f483,f477,f468,f467,f432,f431,f2848,f2847,f2770,f2765,f2545,f2505,f2466,f2433,f2401,f2398,f2373,f2372,f2352,f2349,f2319,f2318,f2272,f2271,f2269,f2251,f1838,f1835,f1819,f1816,f1663,f1661,f1494,f1411,f1410,f1403,f1366,f1365,f1333,f1267,f1231,f1069,f1059,f1057,f1048,f1035,f1033,f1025,f1023,f1001,f991,f961,f942,f940,f930,f926,f924,f904,f902,f891,f889,f873,f871,f870,f869,f868,f867,f866,f865,f864,f863,f862,f861,f852,f839,f838,f837,f836,f829,f818,f787,f592,f414,f359,f358,f4838,f4843,f4844,f4853,f4855,f4857,f4858,f4859,f4860,f4861,f4862,f4863,f4864,f4644,f4646,f4648,f4650,f4870,f4655,f4658,f4660,f4661,f4663,f4669,f4977,f4986,f5268,f5276,f7384,f7394,f7492,f7497,f7506,f7666,f7671,f7680,f7351,f11643,f7360,f11657,f11665,f12884,f12883,f12097,f12096,f13263,f13282,f13283,f13284,f13285,f13298,f13312,f13313,f13314,f13315,f13323,f13324,f13325,f13326,f13327,f13328,f13329,f13331,f13333,f13335,f13337,f13343,f13346,f13348,f13349,f13351,f13356,f13358,f13361,f16694,f16695,f13454,f16696,f13456,f16697,f16698,f16699,f16700,f16701,f16717,f16728,f18264,f18263,f17257,f17256,f18761,f18782,f18783,f18784,f18785,f18786,f18802,f18816,f18817,f18818,f18819,f18828,f18829,f18830,f18831,f18832,f18833,f18834,f18836,f18838,f18840,f18842,f18848,f18851,f18853,f18854,f18856,f18861,f18863,f18866,f23086,f23087,f18961,f23088,f18963,f23089,f23090,f23091,f23092,f23093,f23129,f23139,f23156,f23157,f23158,f23159,f23173,f23183,f23184,f23185,f23186,f25796,f23209,f23235,f23237,f23239,f23241,f23247,f23250,f23252,f23253,f23255,f23261,f23263,f23266,f25806,f23476,f23477,f23508,f23513,f23514,f23515,f23516,f23517,f23518,f23519,f23520,f23521,f23522,f23523,f25824,f23848,f23849,f23850,f23851,f23852,f23853,f23854,f23855,f23856,f23857,f23858,f23859,f23860,f23861,f23862,f23863,f23864,f23865,f23866,f23867,f23868,f23869,f23870,f23871,f23872,f23873,f23874,f23875,f23876,f23877,f23878,f23879,f23880,f23881,f23882,f23883,f23884,f23885,f23886,f23887,f23888,f23889,f23890,f23891,f23892,f23893,f23894,f23895,f23896,f23897,f23898,f23899,f23900,f23901,f23902,f23903,f23904,f23905,f23906,f23907,f23908,f23909,f23910,f23911,f23912,f23913,f23914,f23915,f23916,f23917,f23918,f23919,f23920,f23921,f23922,f23923,f23924,f23925,f23926,f23927,f23928,f23929,f23930,f23931,f23932,f23933,f23934,f23935,f23936,f23937,f23938,f23939,f25830,f23947,f25837,f25840,f25846,f25849,f25852,f25855,f25858,f25861,f25864,f23973,f23975,f23977,f23979,f23985,f23988,f23990,f23991,f23993,f24001,f25872,f25873,f24118,f25875,f25876,f25877,f25878,f25879,f25880,f25883,f25884,f25885,f25886,f25887,f25888,f25889,f25890,f25891,f24144,f24146,f24148,f24150,f25892,f25893,f25894,f25895,f24156,f24159,f24161,f24162,f24164,f25896,f24170,f24172,f24175,f25897,f25898,f25900,f25901,f25902,f25903,f25904,f25905,f25906,f25907,f25908,f25909,f25910,f25911,f25912,f25913,f25914,f25915,f25916,f25917,f25918,f24416,f24421,f24422,f24423,f24424,f24425,f24426,f24427,f24428,f24429,f24430,f24431,f25920,f25926,f25928,f25933,f25935,f25937,f25939,f25941,f25943,f25945,f25998,f26013,f26015,f26018,f26019,f26009,f26010,f26011,f26012,f26024,f26039,f26041,f26044,f26045,f26035,f26036,f26037,f26038,f26053,f26068,f26070,f26073,f26075,f26090,f26105,f26107,f26110,f26111,f26101,f26102,f26103,f26104,f26118,f26133,f26135,f26138,f26139,f26129,f26130,f26131,f26132,f26141,f26160,f26178,f26193,f26195,f26198,f26199,f26189,f26190,f26191,f26192,f26203,f26233,f26260,f26275,f26277,f26280,f26282,f26290,f26320,f26347,f26362,f26364,f26367,f26368,f26358,f26359,f26360,f26361,f26372,f26391,f26408,f26427,f26442,f26458,f26459,f26460,f26461,f26709,f26711,f26714,f26717,f26718,f26719,f26720,f26721,f26723,f26731,f26733,f26735,f26736,f26737,f26738,f26739,f26740,f26741,f26742,f26521,f26523,f26525,f26527,f26751,f26752,f26755,f26537,f26540,f26542,f26543,f26545,f26757,f26759,f26761,f26763,f26764,f26765,f26767,f26768,f26555,f26557,f26560,f26772,f26777,f26783,f26803,f26820,f26840,f26865,f26866,f26867,f26868,f26869,f26921,f26923,f26925,f26927,f26928,f26948,f26949,f26950,f26951,f26952,f26953,f26954,f26955,f26957,f26959,f26960,f26961,f26962,f26963,f26966,f26967,f26968,f26970,f26971,f26972,f26973,f26974,f26975,f26976,f26977,f26978,f26979,f26984,f26985,f26987,f26989,f26991,f26993,f26995,f26996,f26997,f26998,f26999,f27000,f27001,f27002,f27003,f27006,f27008,f27009,f27011,f27013,f27014,f27015,f27016,f27017,f27018,f27019,f27020,f27021,f27023,f27026,f27030,f27033,f27043,f27106,f27107,f27109,f27119,f27132,f27133,f27134,f27135,f27136,f27137,f27138,f27139,f27140,f27141,f27155,f27156,f27157,f27158,f27159,f27160,f27161,f27162,f27166,f27167,f27168,f27169,f27199,f25807,f27298,f27559,f27562,f27567,f27568,f27569,f27570,f28094,f28095,f28097,f28099,f27582,f27583,f27585,f27587,f27589,f27590,f27592,f27594,f27597,f27599,f27602,f27606,f27607,f27608,f27609,f27610,f27611,f27613,f27614,f27619,f27620,f27635,f27636,f27639,f27640,f27641,f27642,f27643,f27645,f27665,f28144,f28149,f28151,f28169,f28171,f28173,f28175,f28178,f28181,f28184,f28187])).
fof(f28185,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f334,f333,f335,f343,f355,f354,f413,f428,f427,f426,f429,f430,f437,f434,f444,f443,f593,f462,f459,f458,f457,f478,f479,f480,f481,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f819,f825,f828,f830,f840,f888,f898,f901,f917,f923,f933,f939,f946,f954,f960,f966,f971,f975,f980,f984,f994,f1000,f1007,f1012,f1024,f1032,f1047,f1050,f1062,f1068,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1239,f1275,f1339,f1332,f1373,f1385,f1402,f1422,f1433,f1437,f1461,f1501,f1512,f1516,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1662,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1981,f2015,f2030,f2073,f2082,f2119,f2118,f2128,f2157,f2170,f2252,f2222,f2262,f2270,f2617,f2622,f2841,f2844,f2751,f2544,f2504,f2465,f2432,f2400,f2371,f2351,f2317,f1837,f1818,f424,f423,f422,f421,f420,f419,f418,f417,f412,f411,f1854,f1861,f1852,f1859,f2052,f2050,f2043,f2041,f482,f433,f1542,f1492,f1486,f489,f488,f487,f485,f484,f1219,f1249,f1217,f1247,f1212,f1256,f1285,f1210,f1254,f1283,f2099,f2097,f2090,f2088,f2001,f1931,f1230,f483,f477,f468,f467,f432,f431,f2848,f2847,f2770,f2765,f2545,f2505,f2466,f2433,f2401,f2398,f2373,f2372,f2352,f2349,f2319,f2318,f2272,f2271,f2269,f2251,f1838,f1835,f1819,f1816,f1663,f1661,f1494,f1411,f1410,f1403,f1366,f1365,f1333,f1267,f1231,f1069,f1059,f1057,f1048,f1035,f1033,f1025,f1023,f1001,f991,f961,f942,f940,f930,f926,f924,f904,f902,f891,f889,f873,f871,f870,f869,f868,f867,f866,f865,f864,f863,f862,f861,f852,f839,f838,f837,f836,f829,f818,f787,f592,f414,f359,f358,f4838,f4843,f4844,f4853,f4855,f4857,f4858,f4859,f4860,f4861,f4862,f4863,f4864,f4644,f4646,f4648,f4650,f4870,f4655,f4658,f4660,f4661,f4663,f4669,f4977,f4986,f5268,f5276,f7384,f7394,f7492,f7497,f7506,f7666,f7671,f7680,f7351,f11643,f7360,f11657,f11665,f12884,f12883,f12097,f12096,f13263,f13282,f13283,f13284,f13285,f13298,f13312,f13313,f13314,f13315,f13323,f13324,f13325,f13326,f13327,f13328,f13329,f13331,f13333,f13335,f13337,f13343,f13346,f13348,f13349,f13351,f13356,f13358,f13361,f16694,f16695,f13454,f16696,f13456,f16697,f16698,f16699,f16700,f16701,f16717,f16728,f18264,f18263,f17257,f17256,f18761,f18782,f18783,f18784,f18785,f18786,f18802,f18816,f18817,f18818,f18819,f18828,f18829,f18830,f18831,f18832,f18833,f18834,f18836,f18838,f18840,f18842,f18848,f18851,f18853,f18854,f18856,f18861,f18863,f18866,f23086,f23087,f18961,f23088,f18963,f23089,f23090,f23091,f23092,f23093,f23129,f23139,f23156,f23157,f23158,f23159,f23173,f23183,f23184,f23185,f23186,f25796,f23209,f23235,f23237,f23239,f23241,f23247,f23250,f23252,f23253,f23255,f23261,f23263,f23266,f25806,f23476,f23477,f23508,f23513,f23514,f23515,f23516,f23517,f23518,f23519,f23520,f23521,f23522,f23523,f25824,f23848,f23849,f23850,f23851,f23852,f23853,f23854,f23855,f23856,f23857,f23858,f23859,f23860,f23861,f23862,f23863,f23864,f23865,f23866,f23867,f23868,f23869,f23870,f23871,f23872,f23873,f23874,f23875,f23876,f23877,f23878,f23879,f23880,f23881,f23882,f23883,f23884,f23885,f23886,f23887,f23888,f23889,f23890,f23891,f23892,f23893,f23894,f23895,f23896,f23897,f23898,f23899,f23900,f23901,f23902,f23903,f23904,f23905,f23906,f23907,f23908,f23909,f23910,f23911,f23912,f23913,f23914,f23915,f23916,f23917,f23918,f23919,f23920,f23921,f23922,f23923,f23924,f23925,f23926,f23927,f23928,f23929,f23930,f23931,f23932,f23933,f23934,f23935,f23936,f23937,f23938,f23939,f25830,f23947,f25837,f25840,f25846,f25849,f25852,f25855,f25858,f25861,f25864,f23973,f23975,f23977,f23979,f23985,f23988,f23990,f23991,f23993,f24001,f25872,f25873,f24118,f25875,f25876,f25877,f25878,f25879,f25880,f25883,f25884,f25885,f25886,f25887,f25888,f25889,f25890,f25891,f24144,f24146,f24148,f24150,f25892,f25893,f25894,f25895,f24156,f24159,f24161,f24162,f24164,f25896,f24170,f24172,f24175,f25897,f25898,f25900,f25901,f25902,f25903,f25904,f25905,f25906,f25907,f25908,f25909,f25910,f25911,f25912,f25913,f25914,f25915,f25916,f25917,f25918,f24416,f24421,f24422,f24423,f24424,f24425,f24426,f24427,f24428,f24429,f24430,f24431,f25920,f25926,f25928,f25933,f25935,f25937,f25939,f25941,f25943,f25945,f25998,f26013,f26015,f26018,f26019,f26009,f26010,f26011,f26012,f26024,f26039,f26041,f26044,f26045,f26035,f26036,f26037,f26038,f26053,f26068,f26070,f26073,f26075,f26090,f26105,f26107,f26110,f26111,f26101,f26102,f26103,f26104,f26118,f26133,f26135,f26138,f26139,f26129,f26130,f26131,f26132,f26141,f26160,f26178,f26193,f26195,f26198,f26199,f26189,f26190,f26191,f26192,f26203,f26233,f26260,f26275,f26277,f26280,f26282,f26290,f26320,f26347,f26362,f26364,f26367,f26368,f26358,f26359,f26360,f26361,f26372,f26391,f26408,f26427,f26442,f26458,f26459,f26460,f26461,f26709,f26711,f26714,f26717,f26718,f26719,f26720,f26721,f26723,f26731,f26733,f26735,f26736,f26737,f26738,f26739,f26740,f26741,f26742,f26521,f26523,f26525,f26527,f26751,f26752,f26755,f26537,f26540,f26542,f26543,f26545,f26757,f26759,f26761,f26763,f26764,f26765,f26767,f26768,f26555,f26557,f26560,f26772,f26777,f26783,f26803,f26820,f26840,f26865,f26866,f26867,f26868,f26869,f26921,f26923,f26925,f26927,f26928,f26948,f26949,f26950,f26951,f26952,f26953,f26954,f26955,f26957,f26959,f26960,f26961,f26962,f26963,f26966,f26967,f26968,f26970,f26971,f26972,f26973,f26974,f26975,f26976,f26977,f26978,f26979,f26984,f26985,f26987,f26989,f26991,f26993,f26995,f26996,f26997,f26998,f26999,f27000,f27001,f27002,f27003,f27006,f27008,f27009,f27011,f27013,f27014,f27015,f27016,f27017,f27018,f27019,f27020,f27021,f27023,f27026,f27030,f27033,f27043,f27106,f27107,f27109,f27119,f27132,f27133,f27134,f27135,f27136,f27137,f27138,f27139,f27140,f27141,f27155,f27156,f27157,f27158,f27159,f27160,f27161,f27162,f27166,f27167,f27168,f27169,f27199,f25807,f27298,f27559,f27562,f27567,f27568,f27569,f27570,f28094,f28095,f28097,f28099,f27582,f27583,f27585,f27587,f27589,f27590,f27592,f27594,f27597,f27599,f27602,f27606,f27607,f27608,f27609,f27610,f27611,f27613,f27614,f27619,f27620,f27635,f27636,f27639,f27640,f27641,f27642,f27643,f27645,f27665,f28144,f28149,f28151,f28169,f28171,f28173,f28175,f28178,f28181,f28184])).
fof(f28182,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f334,f333,f335,f343,f355,f354,f413,f428,f427,f426,f429,f430,f437,f434,f444,f443,f593,f462,f459,f458,f457,f478,f479,f480,f481,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f819,f825,f828,f830,f840,f888,f898,f901,f917,f923,f933,f939,f946,f954,f960,f966,f971,f975,f980,f984,f994,f1000,f1007,f1012,f1024,f1032,f1047,f1050,f1062,f1068,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1239,f1275,f1339,f1332,f1373,f1385,f1402,f1422,f1433,f1437,f1461,f1501,f1512,f1516,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1662,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1981,f2015,f2030,f2073,f2082,f2119,f2118,f2128,f2157,f2170,f2252,f2222,f2262,f2270,f2617,f2622,f2841,f2844,f2751,f2544,f2504,f2465,f2432,f2400,f2371,f2351,f2317,f1837,f1818,f424,f423,f422,f421,f420,f419,f418,f417,f412,f411,f1854,f1861,f1852,f1859,f2052,f2050,f2043,f2041,f482,f433,f1542,f1492,f1486,f489,f488,f487,f485,f484,f1219,f1249,f1217,f1247,f1212,f1256,f1285,f1210,f1254,f1283,f2099,f2097,f2090,f2088,f2001,f1931,f1230,f483,f477,f468,f467,f432,f431,f2848,f2847,f2770,f2765,f2545,f2505,f2466,f2433,f2401,f2398,f2373,f2372,f2352,f2349,f2319,f2318,f2272,f2271,f2269,f2251,f1838,f1835,f1819,f1816,f1663,f1661,f1494,f1411,f1410,f1403,f1366,f1365,f1333,f1267,f1231,f1069,f1059,f1057,f1048,f1035,f1033,f1025,f1023,f1001,f991,f961,f942,f940,f930,f926,f924,f904,f902,f891,f889,f873,f871,f870,f869,f868,f867,f866,f865,f864,f863,f862,f861,f852,f839,f838,f837,f836,f829,f818,f787,f592,f414,f359,f358,f4838,f4843,f4844,f4853,f4855,f4857,f4858,f4859,f4860,f4861,f4862,f4863,f4864,f4644,f4646,f4648,f4650,f4870,f4655,f4658,f4660,f4661,f4663,f4669,f4977,f4986,f5268,f5276,f7384,f7394,f7492,f7497,f7506,f7666,f7671,f7680,f7351,f11643,f7360,f11657,f11665,f12884,f12883,f12097,f12096,f13263,f13282,f13283,f13284,f13285,f13298,f13312,f13313,f13314,f13315,f13323,f13324,f13325,f13326,f13327,f13328,f13329,f13331,f13333,f13335,f13337,f13343,f13346,f13348,f13349,f13351,f13356,f13358,f13361,f16694,f16695,f13454,f16696,f13456,f16697,f16698,f16699,f16700,f16701,f16717,f16728,f18264,f18263,f17257,f17256,f18761,f18782,f18783,f18784,f18785,f18786,f18802,f18816,f18817,f18818,f18819,f18828,f18829,f18830,f18831,f18832,f18833,f18834,f18836,f18838,f18840,f18842,f18848,f18851,f18853,f18854,f18856,f18861,f18863,f18866,f23086,f23087,f18961,f23088,f18963,f23089,f23090,f23091,f23092,f23093,f23129,f23139,f23156,f23157,f23158,f23159,f23173,f23183,f23184,f23185,f23186,f25796,f23209,f23235,f23237,f23239,f23241,f23247,f23250,f23252,f23253,f23255,f23261,f23263,f23266,f25806,f23476,f23477,f23508,f23513,f23514,f23515,f23516,f23517,f23518,f23519,f23520,f23521,f23522,f23523,f25824,f23848,f23849,f23850,f23851,f23852,f23853,f23854,f23855,f23856,f23857,f23858,f23859,f23860,f23861,f23862,f23863,f23864,f23865,f23866,f23867,f23868,f23869,f23870,f23871,f23872,f23873,f23874,f23875,f23876,f23877,f23878,f23879,f23880,f23881,f23882,f23883,f23884,f23885,f23886,f23887,f23888,f23889,f23890,f23891,f23892,f23893,f23894,f23895,f23896,f23897,f23898,f23899,f23900,f23901,f23902,f23903,f23904,f23905,f23906,f23907,f23908,f23909,f23910,f23911,f23912,f23913,f23914,f23915,f23916,f23917,f23918,f23919,f23920,f23921,f23922,f23923,f23924,f23925,f23926,f23927,f23928,f23929,f23930,f23931,f23932,f23933,f23934,f23935,f23936,f23937,f23938,f23939,f25830,f23947,f25837,f25840,f25846,f25849,f25852,f25855,f25858,f25861,f25864,f23973,f23975,f23977,f23979,f23985,f23988,f23990,f23991,f23993,f24001,f25872,f25873,f24118,f25875,f25876,f25877,f25878,f25879,f25880,f25883,f25884,f25885,f25886,f25887,f25888,f25889,f25890,f25891,f24144,f24146,f24148,f24150,f25892,f25893,f25894,f25895,f24156,f24159,f24161,f24162,f24164,f25896,f24170,f24172,f24175,f25897,f25898,f25900,f25901,f25902,f25903,f25904,f25905,f25906,f25907,f25908,f25909,f25910,f25911,f25912,f25913,f25914,f25915,f25916,f25917,f25918,f24416,f24421,f24422,f24423,f24424,f24425,f24426,f24427,f24428,f24429,f24430,f24431,f25920,f25926,f25928,f25933,f25935,f25937,f25939,f25941,f25943,f25945,f25998,f26013,f26015,f26018,f26019,f26009,f26010,f26011,f26012,f26024,f26039,f26041,f26044,f26045,f26035,f26036,f26037,f26038,f26053,f26068,f26070,f26073,f26075,f26090,f26105,f26107,f26110,f26111,f26101,f26102,f26103,f26104,f26118,f26133,f26135,f26138,f26139,f26129,f26130,f26131,f26132,f26141,f26160,f26178,f26193,f26195,f26198,f26199,f26189,f26190,f26191,f26192,f26203,f26233,f26260,f26275,f26277,f26280,f26282,f26290,f26320,f26347,f26362,f26364,f26367,f26368,f26358,f26359,f26360,f26361,f26372,f26391,f26408,f26427,f26442,f26458,f26459,f26460,f26461,f26709,f26711,f26714,f26717,f26718,f26719,f26720,f26721,f26723,f26731,f26733,f26735,f26736,f26737,f26738,f26739,f26740,f26741,f26742,f26521,f26523,f26525,f26527,f26751,f26752,f26755,f26537,f26540,f26542,f26543,f26545,f26757,f26759,f26761,f26763,f26764,f26765,f26767,f26768,f26555,f26557,f26560,f26772,f26777,f26783,f26803,f26820,f26840,f26865,f26866,f26867,f26868,f26869,f26921,f26923,f26925,f26927,f26928,f26948,f26949,f26950,f26951,f26952,f26953,f26954,f26955,f26957,f26959,f26960,f26961,f26962,f26963,f26966,f26967,f26968,f26970,f26971,f26972,f26973,f26974,f26975,f26976,f26977,f26978,f26979,f26984,f26985,f26987,f26989,f26991,f26993,f26995,f26996,f26997,f26998,f26999,f27000,f27001,f27002,f27003,f27006,f27008,f27009,f27011,f27013,f27014,f27015,f27016,f27017,f27018,f27019,f27020,f27021,f27023,f27026,f27030,f27033,f27043,f27106,f27107,f27109,f27119,f27132,f27133,f27134,f27135,f27136,f27137,f27138,f27139,f27140,f27141,f27155,f27156,f27157,f27158,f27159,f27160,f27161,f27162,f27166,f27167,f27168,f27169,f27199,f25807,f27298,f27559,f27562,f27567,f27568,f27569,f27570,f28094,f28095,f28097,f28099,f27582,f27583,f27585,f27587,f27589,f27590,f27592,f27594,f27597,f27599,f27602,f27606,f27607,f27608,f27609,f27610,f27611,f27613,f27614,f27619,f27620,f27635,f27636,f27639,f27640,f27641,f27642,f27643,f27645,f27665,f28144,f28149,f28151,f28169,f28171,f28173,f28175,f28178,f28181])).
fof(f28179,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f334,f333,f335,f343,f355,f354,f413,f428,f427,f426,f429,f430,f437,f434,f444,f443,f593,f462,f459,f458,f457,f478,f479,f480,f481,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f819,f825,f828,f830,f840,f888,f898,f901,f917,f923,f933,f939,f946,f954,f960,f966,f971,f975,f980,f984,f994,f1000,f1007,f1012,f1024,f1032,f1047,f1050,f1062,f1068,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1239,f1275,f1339,f1332,f1373,f1385,f1402,f1422,f1433,f1437,f1461,f1501,f1512,f1516,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1662,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1981,f2015,f2030,f2073,f2082,f2119,f2118,f2128,f2157,f2170,f2252,f2222,f2262,f2270,f2617,f2622,f2841,f2844,f2751,f2544,f2504,f2465,f2432,f2400,f2371,f2351,f2317,f1837,f1818,f424,f423,f422,f421,f420,f419,f418,f417,f412,f411,f1854,f1861,f1852,f1859,f2052,f2050,f2043,f2041,f482,f433,f1542,f1492,f1486,f489,f488,f487,f485,f484,f1219,f1249,f1217,f1247,f1212,f1256,f1285,f1210,f1254,f1283,f2099,f2097,f2090,f2088,f2001,f1931,f1230,f483,f477,f468,f467,f432,f431,f2848,f2847,f2770,f2765,f2545,f2505,f2466,f2433,f2401,f2398,f2373,f2372,f2352,f2349,f2319,f2318,f2272,f2271,f2269,f2251,f1838,f1835,f1819,f1816,f1663,f1661,f1494,f1411,f1410,f1403,f1366,f1365,f1333,f1267,f1231,f1069,f1059,f1057,f1048,f1035,f1033,f1025,f1023,f1001,f991,f961,f942,f940,f930,f926,f924,f904,f902,f891,f889,f873,f871,f870,f869,f868,f867,f866,f865,f864,f863,f862,f861,f852,f839,f838,f837,f836,f829,f818,f787,f592,f414,f359,f358,f4838,f4843,f4844,f4853,f4855,f4857,f4858,f4859,f4860,f4861,f4862,f4863,f4864,f4644,f4646,f4648,f4650,f4870,f4655,f4658,f4660,f4661,f4663,f4669,f4977,f4986,f5268,f5276,f7384,f7394,f7492,f7497,f7506,f7666,f7671,f7680,f7351,f11643,f7360,f11657,f11665,f12884,f12883,f12097,f12096,f13263,f13282,f13283,f13284,f13285,f13298,f13312,f13313,f13314,f13315,f13323,f13324,f13325,f13326,f13327,f13328,f13329,f13331,f13333,f13335,f13337,f13343,f13346,f13348,f13349,f13351,f13356,f13358,f13361,f16694,f16695,f13454,f16696,f13456,f16697,f16698,f16699,f16700,f16701,f16717,f16728,f18264,f18263,f17257,f17256,f18761,f18782,f18783,f18784,f18785,f18786,f18802,f18816,f18817,f18818,f18819,f18828,f18829,f18830,f18831,f18832,f18833,f18834,f18836,f18838,f18840,f18842,f18848,f18851,f18853,f18854,f18856,f18861,f18863,f18866,f23086,f23087,f18961,f23088,f18963,f23089,f23090,f23091,f23092,f23093,f23129,f23139,f23156,f23157,f23158,f23159,f23173,f23183,f23184,f23185,f23186,f25796,f23209,f23235,f23237,f23239,f23241,f23247,f23250,f23252,f23253,f23255,f23261,f23263,f23266,f25806,f23476,f23477,f23508,f23513,f23514,f23515,f23516,f23517,f23518,f23519,f23520,f23521,f23522,f23523,f25824,f23848,f23849,f23850,f23851,f23852,f23853,f23854,f23855,f23856,f23857,f23858,f23859,f23860,f23861,f23862,f23863,f23864,f23865,f23866,f23867,f23868,f23869,f23870,f23871,f23872,f23873,f23874,f23875,f23876,f23877,f23878,f23879,f23880,f23881,f23882,f23883,f23884,f23885,f23886,f23887,f23888,f23889,f23890,f23891,f23892,f23893,f23894,f23895,f23896,f23897,f23898,f23899,f23900,f23901,f23902,f23903,f23904,f23905,f23906,f23907,f23908,f23909,f23910,f23911,f23912,f23913,f23914,f23915,f23916,f23917,f23918,f23919,f23920,f23921,f23922,f23923,f23924,f23925,f23926,f23927,f23928,f23929,f23930,f23931,f23932,f23933,f23934,f23935,f23936,f23937,f23938,f23939,f25830,f23947,f25837,f25840,f25846,f25849,f25852,f25855,f25858,f25861,f25864,f23973,f23975,f23977,f23979,f23985,f23988,f23990,f23991,f23993,f24001,f25872,f25873,f24118,f25875,f25876,f25877,f25878,f25879,f25880,f25883,f25884,f25885,f25886,f25887,f25888,f25889,f25890,f25891,f24144,f24146,f24148,f24150,f25892,f25893,f25894,f25895,f24156,f24159,f24161,f24162,f24164,f25896,f24170,f24172,f24175,f25897,f25898,f25900,f25901,f25902,f25903,f25904,f25905,f25906,f25907,f25908,f25909,f25910,f25911,f25912,f25913,f25914,f25915,f25916,f25917,f25918,f24416,f24421,f24422,f24423,f24424,f24425,f24426,f24427,f24428,f24429,f24430,f24431,f25920,f25926,f25928,f25933,f25935,f25937,f25939,f25941,f25943,f25945,f25998,f26013,f26015,f26018,f26019,f26009,f26010,f26011,f26012,f26024,f26039,f26041,f26044,f26045,f26035,f26036,f26037,f26038,f26053,f26068,f26070,f26073,f26075,f26090,f26105,f26107,f26110,f26111,f26101,f26102,f26103,f26104,f26118,f26133,f26135,f26138,f26139,f26129,f26130,f26131,f26132,f26141,f26160,f26178,f26193,f26195,f26198,f26199,f26189,f26190,f26191,f26192,f26203,f26233,f26260,f26275,f26277,f26280,f26282,f26290,f26320,f26347,f26362,f26364,f26367,f26368,f26358,f26359,f26360,f26361,f26372,f26391,f26408,f26427,f26442,f26458,f26459,f26460,f26461,f26709,f26711,f26714,f26717,f26718,f26719,f26720,f26721,f26723,f26731,f26733,f26735,f26736,f26737,f26738,f26739,f26740,f26741,f26742,f26521,f26523,f26525,f26527,f26751,f26752,f26755,f26537,f26540,f26542,f26543,f26545,f26757,f26759,f26761,f26763,f26764,f26765,f26767,f26768,f26555,f26557,f26560,f26772,f26777,f26783,f26803,f26820,f26840,f26865,f26866,f26867,f26868,f26869,f26921,f26923,f26925,f26927,f26928,f26948,f26949,f26950,f26951,f26952,f26953,f26954,f26955,f26957,f26959,f26960,f26961,f26962,f26963,f26966,f26967,f26968,f26970,f26971,f26972,f26973,f26974,f26975,f26976,f26977,f26978,f26979,f26984,f26985,f26987,f26989,f26991,f26993,f26995,f26996,f26997,f26998,f26999,f27000,f27001,f27002,f27003,f27006,f27008,f27009,f27011,f27013,f27014,f27015,f27016,f27017,f27018,f27019,f27020,f27021,f27023,f27026,f27030,f27033,f27043,f27106,f27107,f27109,f27119,f27132,f27133,f27134,f27135,f27136,f27137,f27138,f27139,f27140,f27141,f27155,f27156,f27157,f27158,f27159,f27160,f27161,f27162,f27166,f27167,f27168,f27169,f27199,f25807,f27298,f27559,f27562,f27567,f27568,f27569,f27570,f28094,f28095,f28097,f28099,f27582,f27583,f27585,f27587,f27589,f27590,f27592,f27594,f27597,f27599,f27602,f27606,f27607,f27608,f27609,f27610,f27611,f27613,f27614,f27619,f27620,f27635,f27636,f27639,f27640,f27641,f27642,f27643,f27645,f27665,f28144,f28149,f28151,f28169,f28171,f28173,f28175,f28178])).
fof(f28176,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f334,f333,f335,f343,f355,f354,f413,f428,f427,f426,f429,f430,f437,f434,f444,f443,f593,f462,f459,f458,f457,f478,f479,f480,f481,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f819,f825,f828,f830,f840,f888,f898,f901,f917,f923,f933,f939,f946,f954,f960,f966,f971,f975,f980,f984,f994,f1000,f1007,f1012,f1024,f1032,f1047,f1050,f1062,f1068,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1239,f1275,f1339,f1332,f1373,f1385,f1402,f1422,f1433,f1437,f1461,f1501,f1512,f1516,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1662,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1981,f2015,f2030,f2073,f2082,f2119,f2118,f2128,f2157,f2170,f2252,f2222,f2262,f2270,f2617,f2622,f2841,f2844,f2751,f2544,f2504,f2465,f2432,f2400,f2371,f2351,f2317,f1837,f1818,f424,f423,f422,f421,f420,f419,f418,f417,f412,f411,f1854,f1861,f1852,f1859,f2052,f2050,f2043,f2041,f482,f433,f1542,f1492,f1486,f489,f488,f487,f485,f484,f1219,f1249,f1217,f1247,f1212,f1256,f1285,f1210,f1254,f1283,f2099,f2097,f2090,f2088,f2001,f1931,f1230,f483,f477,f468,f467,f432,f431,f2848,f2847,f2770,f2765,f2545,f2505,f2466,f2433,f2401,f2398,f2373,f2372,f2352,f2349,f2319,f2318,f2272,f2271,f2269,f2251,f1838,f1835,f1819,f1816,f1663,f1661,f1494,f1411,f1410,f1403,f1366,f1365,f1333,f1267,f1231,f1069,f1059,f1057,f1048,f1035,f1033,f1025,f1023,f1001,f991,f961,f942,f940,f930,f926,f924,f904,f902,f891,f889,f873,f871,f870,f869,f868,f867,f866,f865,f864,f863,f862,f861,f852,f839,f838,f837,f836,f829,f818,f787,f592,f414,f359,f358,f4838,f4843,f4844,f4853,f4855,f4857,f4858,f4859,f4860,f4861,f4862,f4863,f4864,f4644,f4646,f4648,f4650,f4870,f4655,f4658,f4660,f4661,f4663,f4669,f4977,f4986,f5268,f5276,f7384,f7394,f7492,f7497,f7506,f7666,f7671,f7680,f7351,f11643,f7360,f11657,f11665,f12884,f12883,f12097,f12096,f13263,f13282,f13283,f13284,f13285,f13298,f13312,f13313,f13314,f13315,f13323,f13324,f13325,f13326,f13327,f13328,f13329,f13331,f13333,f13335,f13337,f13343,f13346,f13348,f13349,f13351,f13356,f13358,f13361,f16694,f16695,f13454,f16696,f13456,f16697,f16698,f16699,f16700,f16701,f16717,f16728,f18264,f18263,f17257,f17256,f18761,f18782,f18783,f18784,f18785,f18786,f18802,f18816,f18817,f18818,f18819,f18828,f18829,f18830,f18831,f18832,f18833,f18834,f18836,f18838,f18840,f18842,f18848,f18851,f18853,f18854,f18856,f18861,f18863,f18866,f23086,f23087,f18961,f23088,f18963,f23089,f23090,f23091,f23092,f23093,f23129,f23139,f23156,f23157,f23158,f23159,f23173,f23183,f23184,f23185,f23186,f25796,f23209,f23235,f23237,f23239,f23241,f23247,f23250,f23252,f23253,f23255,f23261,f23263,f23266,f25806,f23476,f23477,f23508,f23513,f23514,f23515,f23516,f23517,f23518,f23519,f23520,f23521,f23522,f23523,f25824,f23848,f23849,f23850,f23851,f23852,f23853,f23854,f23855,f23856,f23857,f23858,f23859,f23860,f23861,f23862,f23863,f23864,f23865,f23866,f23867,f23868,f23869,f23870,f23871,f23872,f23873,f23874,f23875,f23876,f23877,f23878,f23879,f23880,f23881,f23882,f23883,f23884,f23885,f23886,f23887,f23888,f23889,f23890,f23891,f23892,f23893,f23894,f23895,f23896,f23897,f23898,f23899,f23900,f23901,f23902,f23903,f23904,f23905,f23906,f23907,f23908,f23909,f23910,f23911,f23912,f23913,f23914,f23915,f23916,f23917,f23918,f23919,f23920,f23921,f23922,f23923,f23924,f23925,f23926,f23927,f23928,f23929,f23930,f23931,f23932,f23933,f23934,f23935,f23936,f23937,f23938,f23939,f25830,f23947,f25837,f25840,f25846,f25849,f25852,f25855,f25858,f25861,f25864,f23973,f23975,f23977,f23979,f23985,f23988,f23990,f23991,f23993,f24001,f25872,f25873,f24118,f25875,f25876,f25877,f25878,f25879,f25880,f25883,f25884,f25885,f25886,f25887,f25888,f25889,f25890,f25891,f24144,f24146,f24148,f24150,f25892,f25893,f25894,f25895,f24156,f24159,f24161,f24162,f24164,f25896,f24170,f24172,f24175,f25897,f25898,f25900,f25901,f25902,f25903,f25904,f25905,f25906,f25907,f25908,f25909,f25910,f25911,f25912,f25913,f25914,f25915,f25916,f25917,f25918,f24416,f24421,f24422,f24423,f24424,f24425,f24426,f24427,f24428,f24429,f24430,f24431,f25920,f25926,f25928,f25933,f25935,f25937,f25939,f25941,f25943,f25945,f25998,f26013,f26015,f26018,f26019,f26009,f26010,f26011,f26012,f26024,f26039,f26041,f26044,f26045,f26035,f26036,f26037,f26038,f26053,f26068,f26070,f26073,f26075,f26090,f26105,f26107,f26110,f26111,f26101,f26102,f26103,f26104,f26118,f26133,f26135,f26138,f26139,f26129,f26130,f26131,f26132,f26141,f26160,f26178,f26193,f26195,f26198,f26199,f26189,f26190,f26191,f26192,f26203,f26233,f26260,f26275,f26277,f26280,f26282,f26290,f26320,f26347,f26362,f26364,f26367,f26368,f26358,f26359,f26360,f26361,f26372,f26391,f26408,f26427,f26442,f26458,f26459,f26460,f26461,f26709,f26711,f26714,f26717,f26718,f26719,f26720,f26721,f26723,f26731,f26733,f26735,f26736,f26737,f26738,f26739,f26740,f26741,f26742,f26521,f26523,f26525,f26527,f26751,f26752,f26755,f26537,f26540,f26542,f26543,f26545,f26757,f26759,f26761,f26763,f26764,f26765,f26767,f26768,f26555,f26557,f26560,f26772,f26777,f26783,f26803,f26820,f26840,f26865,f26866,f26867,f26868,f26869,f26921,f26923,f26925,f26927,f26928,f26948,f26949,f26950,f26951,f26952,f26953,f26954,f26955,f26957,f26959,f26960,f26961,f26962,f26963,f26966,f26967,f26968,f26970,f26971,f26972,f26973,f26974,f26975,f26976,f26977,f26978,f26979,f26984,f26985,f26987,f26989,f26991,f26993,f26995,f26996,f26997,f26998,f26999,f27000,f27001,f27002,f27003,f27006,f27008,f27009,f27011,f27013,f27014,f27015,f27016,f27017,f27018,f27019,f27020,f27021,f27023,f27026,f27030,f27033,f27043,f27106,f27107,f27109,f27119,f27132,f27133,f27134,f27135,f27136,f27137,f27138,f27139,f27140,f27141,f27155,f27156,f27157,f27158,f27159,f27160,f27161,f27162,f27166,f27167,f27168,f27169,f27199,f25807,f27298,f27559,f27562,f27567,f27568,f27569,f27570,f28094,f28095,f28097,f28099,f27582,f27583,f27585,f27587,f27589,f27590,f27592,f27594,f27597,f27599,f27602,f27606,f27607,f27608,f27609,f27610,f27611,f27613,f27614,f27619,f27620,f27635,f27636,f27639,f27640,f27641,f27642,f27643,f27645,f27665,f28144,f28149,f28151,f28169,f28171,f28173,f28175])).
fof(f28174,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f334,f333,f335,f343,f355,f354,f413,f428,f427,f426,f429,f430,f437,f434,f444,f443,f593,f462,f459,f458,f457,f478,f479,f480,f481,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f819,f825,f828,f830,f840,f888,f898,f901,f917,f923,f933,f939,f946,f954,f960,f966,f971,f975,f980,f984,f994,f1000,f1007,f1012,f1024,f1032,f1047,f1050,f1062,f1068,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1239,f1275,f1339,f1332,f1373,f1385,f1402,f1422,f1433,f1437,f1461,f1501,f1512,f1516,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1662,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1981,f2015,f2030,f2073,f2082,f2119,f2118,f2128,f2157,f2170,f2252,f2222,f2262,f2270,f2617,f2622,f2841,f2844,f2751,f2544,f2504,f2465,f2432,f2400,f2371,f2351,f2317,f1837,f1818,f424,f423,f422,f421,f420,f419,f418,f417,f412,f411,f1854,f1861,f1852,f1859,f2052,f2050,f2043,f2041,f482,f433,f1542,f1492,f1486,f489,f488,f487,f485,f484,f1219,f1249,f1217,f1247,f1212,f1256,f1285,f1210,f1254,f1283,f2099,f2097,f2090,f2088,f2001,f1931,f1230,f483,f477,f468,f467,f432,f431,f2848,f2847,f2770,f2765,f2545,f2505,f2466,f2433,f2401,f2398,f2373,f2372,f2352,f2349,f2319,f2318,f2272,f2271,f2269,f2251,f1838,f1835,f1819,f1816,f1663,f1661,f1494,f1411,f1410,f1403,f1366,f1365,f1333,f1267,f1231,f1069,f1059,f1057,f1048,f1035,f1033,f1025,f1023,f1001,f991,f961,f942,f940,f930,f926,f924,f904,f902,f891,f889,f873,f871,f870,f869,f868,f867,f866,f865,f864,f863,f862,f861,f852,f839,f838,f837,f836,f829,f818,f787,f592,f414,f359,f358,f4838,f4843,f4844,f4853,f4855,f4857,f4858,f4859,f4860,f4861,f4862,f4863,f4864,f4644,f4646,f4648,f4650,f4870,f4655,f4658,f4660,f4661,f4663,f4669,f4977,f4986,f5268,f5276,f7384,f7394,f7492,f7497,f7506,f7666,f7671,f7680,f7351,f11643,f7360,f11657,f11665,f12884,f12883,f12097,f12096,f13263,f13282,f13283,f13284,f13285,f13298,f13312,f13313,f13314,f13315,f13323,f13324,f13325,f13326,f13327,f13328,f13329,f13331,f13333,f13335,f13337,f13343,f13346,f13348,f13349,f13351,f13356,f13358,f13361,f16694,f16695,f13454,f16696,f13456,f16697,f16698,f16699,f16700,f16701,f16717,f16728,f18264,f18263,f17257,f17256,f18761,f18782,f18783,f18784,f18785,f18786,f18802,f18816,f18817,f18818,f18819,f18828,f18829,f18830,f18831,f18832,f18833,f18834,f18836,f18838,f18840,f18842,f18848,f18851,f18853,f18854,f18856,f18861,f18863,f18866,f23086,f23087,f18961,f23088,f18963,f23089,f23090,f23091,f23092,f23093,f23129,f23139,f23156,f23157,f23158,f23159,f23173,f23183,f23184,f23185,f23186,f25796,f23209,f23235,f23237,f23239,f23241,f23247,f23250,f23252,f23253,f23255,f23261,f23263,f23266,f25806,f23476,f23477,f23508,f23513,f23514,f23515,f23516,f23517,f23518,f23519,f23520,f23521,f23522,f23523,f25824,f23848,f23849,f23850,f23851,f23852,f23853,f23854,f23855,f23856,f23857,f23858,f23859,f23860,f23861,f23862,f23863,f23864,f23865,f23866,f23867,f23868,f23869,f23870,f23871,f23872,f23873,f23874,f23875,f23876,f23877,f23878,f23879,f23880,f23881,f23882,f23883,f23884,f23885,f23886,f23887,f23888,f23889,f23890,f23891,f23892,f23893,f23894,f23895,f23896,f23897,f23898,f23899,f23900,f23901,f23902,f23903,f23904,f23905,f23906,f23907,f23908,f23909,f23910,f23911,f23912,f23913,f23914,f23915,f23916,f23917,f23918,f23919,f23920,f23921,f23922,f23923,f23924,f23925,f23926,f23927,f23928,f23929,f23930,f23931,f23932,f23933,f23934,f23935,f23936,f23937,f23938,f23939,f25830,f23947,f25837,f25840,f25846,f25849,f25852,f25855,f25858,f25861,f25864,f23973,f23975,f23977,f23979,f23985,f23988,f23990,f23991,f23993,f24001,f25872,f25873,f24118,f25875,f25876,f25877,f25878,f25879,f25880,f25883,f25884,f25885,f25886,f25887,f25888,f25889,f25890,f25891,f24144,f24146,f24148,f24150,f25892,f25893,f25894,f25895,f24156,f24159,f24161,f24162,f24164,f25896,f24170,f24172,f24175,f25897,f25898,f25900,f25901,f25902,f25903,f25904,f25905,f25906,f25907,f25908,f25909,f25910,f25911,f25912,f25913,f25914,f25915,f25916,f25917,f25918,f24416,f24421,f24422,f24423,f24424,f24425,f24426,f24427,f24428,f24429,f24430,f24431,f25920,f25926,f25928,f25933,f25935,f25937,f25939,f25941,f25943,f25945,f25998,f26013,f26015,f26018,f26019,f26009,f26010,f26011,f26012,f26024,f26039,f26041,f26044,f26045,f26035,f26036,f26037,f26038,f26053,f26068,f26070,f26073,f26075,f26090,f26105,f26107,f26110,f26111,f26101,f26102,f26103,f26104,f26118,f26133,f26135,f26138,f26139,f26129,f26130,f26131,f26132,f26141,f26160,f26178,f26193,f26195,f26198,f26199,f26189,f26190,f26191,f26192,f26203,f26233,f26260,f26275,f26277,f26280,f26282,f26290,f26320,f26347,f26362,f26364,f26367,f26368,f26358,f26359,f26360,f26361,f26372,f26391,f26408,f26427,f26442,f26458,f26459,f26460,f26461,f26709,f26711,f26714,f26717,f26718,f26719,f26720,f26721,f26723,f26731,f26733,f26735,f26736,f26737,f26738,f26739,f26740,f26741,f26742,f26521,f26523,f26525,f26527,f26751,f26752,f26755,f26537,f26540,f26542,f26543,f26545,f26757,f26759,f26761,f26763,f26764,f26765,f26767,f26768,f26555,f26557,f26560,f26772,f26777,f26783,f26803,f26820,f26840,f26865,f26866,f26867,f26868,f26869,f26921,f26923,f26925,f26927,f26928,f26948,f26949,f26950,f26951,f26952,f26953,f26954,f26955,f26957,f26959,f26960,f26961,f26962,f26963,f26966,f26967,f26968,f26970,f26971,f26972,f26973,f26974,f26975,f26976,f26977,f26978,f26979,f26984,f26985,f26987,f26989,f26991,f26993,f26995,f26996,f26997,f26998,f26999,f27000,f27001,f27002,f27003,f27006,f27008,f27009,f27011,f27013,f27014,f27015,f27016,f27017,f27018,f27019,f27020,f27021,f27023,f27026,f27030,f27033,f27043,f27106,f27107,f27109,f27119,f27132,f27133,f27134,f27135,f27136,f27137,f27138,f27139,f27140,f27141,f27155,f27156,f27157,f27158,f27159,f27160,f27161,f27162,f27166,f27167,f27168,f27169,f27199,f25807,f27298,f27559,f27562,f27567,f27568,f27569,f27570,f28094,f28095,f28097,f28099,f27582,f27583,f27585,f27587,f27589,f27590,f27592,f27594,f27597,f27599,f27602,f27606,f27607,f27608,f27609,f27610,f27611,f27613,f27614,f27619,f27620,f27635,f27636,f27639,f27640,f27641,f27642,f27643,f27645,f27665,f28144,f28149,f28151,f28169,f28171,f28173])).
fof(f28172,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f334,f333,f335,f343,f355,f354,f413,f428,f427,f426,f429,f430,f437,f434,f444,f443,f593,f462,f459,f458,f457,f478,f479,f480,f481,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f819,f825,f828,f830,f840,f888,f898,f901,f917,f923,f933,f939,f946,f954,f960,f966,f971,f975,f980,f984,f994,f1000,f1007,f1012,f1024,f1032,f1047,f1050,f1062,f1068,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1239,f1275,f1339,f1332,f1373,f1385,f1402,f1422,f1433,f1437,f1461,f1501,f1512,f1516,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1662,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1981,f2015,f2030,f2073,f2082,f2119,f2118,f2128,f2157,f2170,f2252,f2222,f2262,f2270,f2617,f2622,f2841,f2844,f2751,f2544,f2504,f2465,f2432,f2400,f2371,f2351,f2317,f1837,f1818,f424,f423,f422,f421,f420,f419,f418,f417,f412,f411,f1854,f1861,f1852,f1859,f2052,f2050,f2043,f2041,f482,f433,f1542,f1492,f1486,f489,f488,f487,f485,f484,f1219,f1249,f1217,f1247,f1212,f1256,f1285,f1210,f1254,f1283,f2099,f2097,f2090,f2088,f2001,f1931,f1230,f483,f477,f468,f467,f432,f431,f2848,f2847,f2770,f2765,f2545,f2505,f2466,f2433,f2401,f2398,f2373,f2372,f2352,f2349,f2319,f2318,f2272,f2271,f2269,f2251,f1838,f1835,f1819,f1816,f1663,f1661,f1494,f1411,f1410,f1403,f1366,f1365,f1333,f1267,f1231,f1069,f1059,f1057,f1048,f1035,f1033,f1025,f1023,f1001,f991,f961,f942,f940,f930,f926,f924,f904,f902,f891,f889,f873,f871,f870,f869,f868,f867,f866,f865,f864,f863,f862,f861,f852,f839,f838,f837,f836,f829,f818,f787,f592,f414,f359,f358,f4838,f4843,f4844,f4853,f4855,f4857,f4858,f4859,f4860,f4861,f4862,f4863,f4864,f4644,f4646,f4648,f4650,f4870,f4655,f4658,f4660,f4661,f4663,f4669,f4977,f4986,f5268,f5276,f7384,f7394,f7492,f7497,f7506,f7666,f7671,f7680,f7351,f11643,f7360,f11657,f11665,f12884,f12883,f12097,f12096,f13263,f13282,f13283,f13284,f13285,f13298,f13312,f13313,f13314,f13315,f13323,f13324,f13325,f13326,f13327,f13328,f13329,f13331,f13333,f13335,f13337,f13343,f13346,f13348,f13349,f13351,f13356,f13358,f13361,f16694,f16695,f13454,f16696,f13456,f16697,f16698,f16699,f16700,f16701,f16717,f16728,f18264,f18263,f17257,f17256,f18761,f18782,f18783,f18784,f18785,f18786,f18802,f18816,f18817,f18818,f18819,f18828,f18829,f18830,f18831,f18832,f18833,f18834,f18836,f18838,f18840,f18842,f18848,f18851,f18853,f18854,f18856,f18861,f18863,f18866,f23086,f23087,f18961,f23088,f18963,f23089,f23090,f23091,f23092,f23093,f23129,f23139,f23156,f23157,f23158,f23159,f23173,f23183,f23184,f23185,f23186,f25796,f23209,f23235,f23237,f23239,f23241,f23247,f23250,f23252,f23253,f23255,f23261,f23263,f23266,f25806,f23476,f23477,f23508,f23513,f23514,f23515,f23516,f23517,f23518,f23519,f23520,f23521,f23522,f23523,f25824,f23848,f23849,f23850,f23851,f23852,f23853,f23854,f23855,f23856,f23857,f23858,f23859,f23860,f23861,f23862,f23863,f23864,f23865,f23866,f23867,f23868,f23869,f23870,f23871,f23872,f23873,f23874,f23875,f23876,f23877,f23878,f23879,f23880,f23881,f23882,f23883,f23884,f23885,f23886,f23887,f23888,f23889,f23890,f23891,f23892,f23893,f23894,f23895,f23896,f23897,f23898,f23899,f23900,f23901,f23902,f23903,f23904,f23905,f23906,f23907,f23908,f23909,f23910,f23911,f23912,f23913,f23914,f23915,f23916,f23917,f23918,f23919,f23920,f23921,f23922,f23923,f23924,f23925,f23926,f23927,f23928,f23929,f23930,f23931,f23932,f23933,f23934,f23935,f23936,f23937,f23938,f23939,f25830,f23947,f25837,f25840,f25846,f25849,f25852,f25855,f25858,f25861,f25864,f23973,f23975,f23977,f23979,f23985,f23988,f23990,f23991,f23993,f24001,f25872,f25873,f24118,f25875,f25876,f25877,f25878,f25879,f25880,f25883,f25884,f25885,f25886,f25887,f25888,f25889,f25890,f25891,f24144,f24146,f24148,f24150,f25892,f25893,f25894,f25895,f24156,f24159,f24161,f24162,f24164,f25896,f24170,f24172,f24175,f25897,f25898,f25900,f25901,f25902,f25903,f25904,f25905,f25906,f25907,f25908,f25909,f25910,f25911,f25912,f25913,f25914,f25915,f25916,f25917,f25918,f24416,f24421,f24422,f24423,f24424,f24425,f24426,f24427,f24428,f24429,f24430,f24431,f25920,f25926,f25928,f25933,f25935,f25937,f25939,f25941,f25943,f25945,f25998,f26013,f26015,f26018,f26019,f26009,f26010,f26011,f26012,f26024,f26039,f26041,f26044,f26045,f26035,f26036,f26037,f26038,f26053,f26068,f26070,f26073,f26075,f26090,f26105,f26107,f26110,f26111,f26101,f26102,f26103,f26104,f26118,f26133,f26135,f26138,f26139,f26129,f26130,f26131,f26132,f26141,f26160,f26178,f26193,f26195,f26198,f26199,f26189,f26190,f26191,f26192,f26203,f26233,f26260,f26275,f26277,f26280,f26282,f26290,f26320,f26347,f26362,f26364,f26367,f26368,f26358,f26359,f26360,f26361,f26372,f26391,f26408,f26427,f26442,f26458,f26459,f26460,f26461,f26709,f26711,f26714,f26717,f26718,f26719,f26720,f26721,f26723,f26731,f26733,f26735,f26736,f26737,f26738,f26739,f26740,f26741,f26742,f26521,f26523,f26525,f26527,f26751,f26752,f26755,f26537,f26540,f26542,f26543,f26545,f26757,f26759,f26761,f26763,f26764,f26765,f26767,f26768,f26555,f26557,f26560,f26772,f26777,f26783,f26803,f26820,f26840,f26865,f26866,f26867,f26868,f26869,f26921,f26923,f26925,f26927,f26928,f26948,f26949,f26950,f26951,f26952,f26953,f26954,f26955,f26957,f26959,f26960,f26961,f26962,f26963,f26966,f26967,f26968,f26970,f26971,f26972,f26973,f26974,f26975,f26976,f26977,f26978,f26979,f26984,f26985,f26987,f26989,f26991,f26993,f26995,f26996,f26997,f26998,f26999,f27000,f27001,f27002,f27003,f27006,f27008,f27009,f27011,f27013,f27014,f27015,f27016,f27017,f27018,f27019,f27020,f27021,f27023,f27026,f27030,f27033,f27043,f27106,f27107,f27109,f27119,f27132,f27133,f27134,f27135,f27136,f27137,f27138,f27139,f27140,f27141,f27155,f27156,f27157,f27158,f27159,f27160,f27161,f27162,f27166,f27167,f27168,f27169,f27199,f25807,f27298,f27559,f27562,f27567,f27568,f27569,f27570,f28094,f28095,f28097,f28099,f27582,f27583,f27585,f27587,f27589,f27590,f27592,f27594,f27597,f27599,f27602,f27606,f27607,f27608,f27609,f27610,f27611,f27613,f27614,f27619,f27620,f27635,f27636,f27639,f27640,f27641,f27642,f27643,f27645,f27665,f28144,f28149,f28151,f28169,f28171])).
fof(f28170,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f334,f333,f335,f343,f355,f354,f413,f428,f427,f426,f429,f430,f437,f434,f444,f443,f593,f462,f459,f458,f457,f478,f479,f480,f481,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f819,f825,f828,f830,f840,f888,f898,f901,f917,f923,f933,f939,f946,f954,f960,f966,f971,f975,f980,f984,f994,f1000,f1007,f1012,f1024,f1032,f1047,f1050,f1062,f1068,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1239,f1275,f1339,f1332,f1373,f1385,f1402,f1422,f1433,f1437,f1461,f1501,f1512,f1516,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1662,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1981,f2015,f2030,f2073,f2082,f2119,f2118,f2128,f2157,f2170,f2252,f2222,f2262,f2270,f2617,f2622,f2841,f2844,f2751,f2544,f2504,f2465,f2432,f2400,f2371,f2351,f2317,f1837,f1818,f424,f423,f422,f421,f420,f419,f418,f417,f412,f411,f1854,f1861,f1852,f1859,f2052,f2050,f2043,f2041,f482,f433,f1542,f1492,f1486,f489,f488,f487,f485,f484,f1219,f1249,f1217,f1247,f1212,f1256,f1285,f1210,f1254,f1283,f2099,f2097,f2090,f2088,f2001,f1931,f1230,f483,f477,f468,f467,f432,f431,f2848,f2847,f2770,f2765,f2545,f2505,f2466,f2433,f2401,f2398,f2373,f2372,f2352,f2349,f2319,f2318,f2272,f2271,f2269,f2251,f1838,f1835,f1819,f1816,f1663,f1661,f1494,f1411,f1410,f1403,f1366,f1365,f1333,f1267,f1231,f1069,f1059,f1057,f1048,f1035,f1033,f1025,f1023,f1001,f991,f961,f942,f940,f930,f926,f924,f904,f902,f891,f889,f873,f871,f870,f869,f868,f867,f866,f865,f864,f863,f862,f861,f852,f839,f838,f837,f836,f829,f818,f787,f592,f414,f359,f358,f4838,f4843,f4844,f4853,f4855,f4857,f4858,f4859,f4860,f4861,f4862,f4863,f4864,f4644,f4646,f4648,f4650,f4870,f4655,f4658,f4660,f4661,f4663,f4669,f4977,f4986,f5268,f5276,f7384,f7394,f7492,f7497,f7506,f7666,f7671,f7680,f7351,f11643,f7360,f11657,f11665,f12884,f12883,f12097,f12096,f13263,f13282,f13283,f13284,f13285,f13298,f13312,f13313,f13314,f13315,f13323,f13324,f13325,f13326,f13327,f13328,f13329,f13331,f13333,f13335,f13337,f13343,f13346,f13348,f13349,f13351,f13356,f13358,f13361,f16694,f16695,f13454,f16696,f13456,f16697,f16698,f16699,f16700,f16701,f16717,f16728,f18264,f18263,f17257,f17256,f18761,f18782,f18783,f18784,f18785,f18786,f18802,f18816,f18817,f18818,f18819,f18828,f18829,f18830,f18831,f18832,f18833,f18834,f18836,f18838,f18840,f18842,f18848,f18851,f18853,f18854,f18856,f18861,f18863,f18866,f23086,f23087,f18961,f23088,f18963,f23089,f23090,f23091,f23092,f23093,f23129,f23139,f23156,f23157,f23158,f23159,f23173,f23183,f23184,f23185,f23186,f25796,f23209,f23235,f23237,f23239,f23241,f23247,f23250,f23252,f23253,f23255,f23261,f23263,f23266,f25806,f23476,f23477,f23508,f23513,f23514,f23515,f23516,f23517,f23518,f23519,f23520,f23521,f23522,f23523,f25824,f23848,f23849,f23850,f23851,f23852,f23853,f23854,f23855,f23856,f23857,f23858,f23859,f23860,f23861,f23862,f23863,f23864,f23865,f23866,f23867,f23868,f23869,f23870,f23871,f23872,f23873,f23874,f23875,f23876,f23877,f23878,f23879,f23880,f23881,f23882,f23883,f23884,f23885,f23886,f23887,f23888,f23889,f23890,f23891,f23892,f23893,f23894,f23895,f23896,f23897,f23898,f23899,f23900,f23901,f23902,f23903,f23904,f23905,f23906,f23907,f23908,f23909,f23910,f23911,f23912,f23913,f23914,f23915,f23916,f23917,f23918,f23919,f23920,f23921,f23922,f23923,f23924,f23925,f23926,f23927,f23928,f23929,f23930,f23931,f23932,f23933,f23934,f23935,f23936,f23937,f23938,f23939,f25830,f23947,f25837,f25840,f25846,f25849,f25852,f25855,f25858,f25861,f25864,f23973,f23975,f23977,f23979,f23985,f23988,f23990,f23991,f23993,f24001,f25872,f25873,f24118,f25875,f25876,f25877,f25878,f25879,f25880,f25883,f25884,f25885,f25886,f25887,f25888,f25889,f25890,f25891,f24144,f24146,f24148,f24150,f25892,f25893,f25894,f25895,f24156,f24159,f24161,f24162,f24164,f25896,f24170,f24172,f24175,f25897,f25898,f25900,f25901,f25902,f25903,f25904,f25905,f25906,f25907,f25908,f25909,f25910,f25911,f25912,f25913,f25914,f25915,f25916,f25917,f25918,f24416,f24421,f24422,f24423,f24424,f24425,f24426,f24427,f24428,f24429,f24430,f24431,f25920,f25926,f25928,f25933,f25935,f25937,f25939,f25941,f25943,f25945,f25998,f26013,f26015,f26018,f26019,f26009,f26010,f26011,f26012,f26024,f26039,f26041,f26044,f26045,f26035,f26036,f26037,f26038,f26053,f26068,f26070,f26073,f26075,f26090,f26105,f26107,f26110,f26111,f26101,f26102,f26103,f26104,f26118,f26133,f26135,f26138,f26139,f26129,f26130,f26131,f26132,f26141,f26160,f26178,f26193,f26195,f26198,f26199,f26189,f26190,f26191,f26192,f26203,f26233,f26260,f26275,f26277,f26280,f26282,f26290,f26320,f26347,f26362,f26364,f26367,f26368,f26358,f26359,f26360,f26361,f26372,f26391,f26408,f26427,f26442,f26458,f26459,f26460,f26461,f26709,f26711,f26714,f26717,f26718,f26719,f26720,f26721,f26723,f26731,f26733,f26735,f26736,f26737,f26738,f26739,f26740,f26741,f26742,f26521,f26523,f26525,f26527,f26751,f26752,f26755,f26537,f26540,f26542,f26543,f26545,f26757,f26759,f26761,f26763,f26764,f26765,f26767,f26768,f26555,f26557,f26560,f26772,f26777,f26783,f26803,f26820,f26840,f26865,f26866,f26867,f26868,f26869,f26921,f26923,f26925,f26927,f26928,f26948,f26949,f26950,f26951,f26952,f26953,f26954,f26955,f26957,f26959,f26960,f26961,f26962,f26963,f26966,f26967,f26968,f26970,f26971,f26972,f26973,f26974,f26975,f26976,f26977,f26978,f26979,f26984,f26985,f26987,f26989,f26991,f26993,f26995,f26996,f26997,f26998,f26999,f27000,f27001,f27002,f27003,f27006,f27008,f27009,f27011,f27013,f27014,f27015,f27016,f27017,f27018,f27019,f27020,f27021,f27023,f27026,f27030,f27033,f27043,f27106,f27107,f27109,f27119,f27132,f27133,f27134,f27135,f27136,f27137,f27138,f27139,f27140,f27141,f27155,f27156,f27157,f27158,f27159,f27160,f27161,f27162,f27166,f27167,f27168,f27169,f27199,f25807,f27298,f27559,f27562,f27567,f27568,f27569,f27570,f28094,f28095,f28097,f28099,f27582,f27583,f27585,f27587,f27589,f27590,f27592,f27594,f27597,f27599,f27602,f27606,f27607,f27608,f27609,f27610,f27611,f27613,f27614,f27619,f27620,f27635,f27636,f27639,f27640,f27641,f27642,f27643,f27645,f27665,f28144,f28149,f28151,f28169])).
fof(f28168,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f28167,f25807])).
fof(f28167,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num',X3))))) ) | $spl8),
  inference(subsumption_resolution,[],[f27837,f335])).
fof(f27837,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num',X3))))) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f1254])).
fof(f28166,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f28165,f25807])).
fof(f28165,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X7),s('type.num.num',X8))),s('type.num.num',X9))))) ) | $spl8),
  inference(subsumption_resolution,[],[f27835,f335])).
fof(f27835,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X7),s('type.num.num',X8))),s('type.num.num',X9))))) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f1256])).
fof(f28164,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f28163,f25807])).
fof(f28163,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X2))))))) ) | $spl8),
  inference(subsumption_resolution,[],[f27832,f335])).
fof(f27832,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X2))))))) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f1247])).
fof(f28162,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f28161,f25807])).
fof(f28161,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X2))))))) ) | $spl8),
  inference(subsumption_resolution,[],[f27831,f335])).
fof(f27831,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X2))))))) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f1283])).
fof(f28160,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f28159,f25807])).
fof(f28159,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X9),s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X7),s('type.num.num',X8))))))) ) | $spl8),
  inference(subsumption_resolution,[],[f27829,f335])).
fof(f27829,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X9),s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X7),s('type.num.num',X8))))))) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f1249])).
fof(f28158,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f28157,f25807])).
fof(f28157,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X7),s('type.num.num',X8))))))) ) | $spl8),
  inference(subsumption_resolution,[],[f27828,f335])).
fof(f27828,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X7),s('type.num.num',X8))))))) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f1285])).
fof(f28156,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f28155,f25807])).
fof(f28155,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))))) ) | $spl8),
  inference(subsumption_resolution,[],[f27810,f335])).
fof(f27810,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))))) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f477])).
fof(f28154,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f28153,f25807])).
fof(f28153,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) ) | $spl8),
  inference(subsumption_resolution,[],[f27809,f335])).
fof(f27809,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f468])).
fof(f28152,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f334,f333,f335,f343,f355,f354,f413,f428,f427,f426,f429,f430,f437,f434,f444,f443,f593,f462,f459,f458,f457,f478,f479,f480,f481,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f819,f825,f828,f830,f840,f888,f898,f901,f917,f923,f933,f939,f946,f954,f960,f966,f971,f975,f980,f984,f994,f1000,f1007,f1012,f1024,f1032,f1047,f1050,f1062,f1068,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1239,f1275,f1339,f1332,f1373,f1385,f1402,f1422,f1433,f1437,f1461,f1501,f1512,f1516,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1662,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1981,f2015,f2030,f2073,f2082,f2119,f2118,f2128,f2157,f2170,f2252,f2222,f2262,f2270,f2617,f2622,f2841,f2844,f2751,f2544,f2504,f2465,f2432,f2400,f2371,f2351,f2317,f1837,f1818,f424,f423,f422,f421,f420,f419,f418,f417,f412,f411,f1854,f1861,f1852,f1859,f2052,f2050,f2043,f2041,f482,f433,f1542,f1492,f1486,f489,f488,f487,f485,f484,f1219,f1249,f1217,f1247,f1212,f1256,f1285,f1210,f1254,f1283,f2099,f2097,f2090,f2088,f2001,f1931,f1230,f483,f477,f468,f467,f432,f431,f2848,f2847,f2770,f2765,f2545,f2505,f2466,f2433,f2401,f2398,f2373,f2372,f2352,f2349,f2319,f2318,f2272,f2271,f2269,f2251,f1838,f1835,f1819,f1816,f1663,f1661,f1494,f1411,f1410,f1403,f1366,f1365,f1333,f1267,f1231,f1069,f1059,f1057,f1048,f1035,f1033,f1025,f1023,f1001,f991,f961,f942,f940,f930,f926,f924,f904,f902,f891,f889,f873,f871,f870,f869,f868,f867,f866,f865,f864,f863,f862,f861,f852,f839,f838,f837,f836,f829,f818,f787,f592,f414,f359,f358,f4838,f4843,f4844,f4853,f4855,f4857,f4858,f4859,f4860,f4861,f4862,f4863,f4864,f4644,f4646,f4648,f4650,f4870,f4655,f4658,f4660,f4661,f4663,f4669,f4977,f4986,f5268,f5276,f7384,f7394,f7492,f7497,f7506,f7666,f7671,f7680,f7351,f11643,f7360,f11657,f11665,f12884,f12883,f12097,f12096,f13263,f13282,f13283,f13284,f13285,f13298,f13312,f13313,f13314,f13315,f13323,f13324,f13325,f13326,f13327,f13328,f13329,f13331,f13333,f13335,f13337,f13343,f13346,f13348,f13349,f13351,f13356,f13358,f13361,f16694,f16695,f13454,f16696,f13456,f16697,f16698,f16699,f16700,f16701,f16717,f16728,f18264,f18263,f17257,f17256,f18761,f18782,f18783,f18784,f18785,f18786,f18802,f18816,f18817,f18818,f18819,f18828,f18829,f18830,f18831,f18832,f18833,f18834,f18836,f18838,f18840,f18842,f18848,f18851,f18853,f18854,f18856,f18861,f18863,f18866,f23086,f23087,f18961,f23088,f18963,f23089,f23090,f23091,f23092,f23093,f23129,f23139,f23156,f23157,f23158,f23159,f23173,f23183,f23184,f23185,f23186,f25796,f23209,f23235,f23237,f23239,f23241,f23247,f23250,f23252,f23253,f23255,f23261,f23263,f23266,f25806,f23476,f23477,f23508,f23513,f23514,f23515,f23516,f23517,f23518,f23519,f23520,f23521,f23522,f23523,f25824,f23848,f23849,f23850,f23851,f23852,f23853,f23854,f23855,f23856,f23857,f23858,f23859,f23860,f23861,f23862,f23863,f23864,f23865,f23866,f23867,f23868,f23869,f23870,f23871,f23872,f23873,f23874,f23875,f23876,f23877,f23878,f23879,f23880,f23881,f23882,f23883,f23884,f23885,f23886,f23887,f23888,f23889,f23890,f23891,f23892,f23893,f23894,f23895,f23896,f23897,f23898,f23899,f23900,f23901,f23902,f23903,f23904,f23905,f23906,f23907,f23908,f23909,f23910,f23911,f23912,f23913,f23914,f23915,f23916,f23917,f23918,f23919,f23920,f23921,f23922,f23923,f23924,f23925,f23926,f23927,f23928,f23929,f23930,f23931,f23932,f23933,f23934,f23935,f23936,f23937,f23938,f23939,f25830,f23947,f25837,f25840,f25846,f25849,f25852,f25855,f25858,f25861,f25864,f23973,f23975,f23977,f23979,f23985,f23988,f23990,f23991,f23993,f24001,f25872,f25873,f24118,f25875,f25876,f25877,f25878,f25879,f25880,f25883,f25884,f25885,f25886,f25887,f25888,f25889,f25890,f25891,f24144,f24146,f24148,f24150,f25892,f25893,f25894,f25895,f24156,f24159,f24161,f24162,f24164,f25896,f24170,f24172,f24175,f25897,f25898,f25900,f25901,f25902,f25903,f25904,f25905,f25906,f25907,f25908,f25909,f25910,f25911,f25912,f25913,f25914,f25915,f25916,f25917,f25918,f24416,f24421,f24422,f24423,f24424,f24425,f24426,f24427,f24428,f24429,f24430,f24431,f25920,f25926,f25928,f25933,f25935,f25937,f25939,f25941,f25943,f25945,f25998,f26013,f26015,f26018,f26019,f26009,f26010,f26011,f26012,f26024,f26039,f26041,f26044,f26045,f26035,f26036,f26037,f26038,f26053,f26068,f26070,f26073,f26075,f26090,f26105,f26107,f26110,f26111,f26101,f26102,f26103,f26104,f26118,f26133,f26135,f26138,f26139,f26129,f26130,f26131,f26132,f26141,f26160,f26178,f26193,f26195,f26198,f26199,f26189,f26190,f26191,f26192,f26203,f26233,f26260,f26275,f26277,f26280,f26282,f26290,f26320,f26347,f26362,f26364,f26367,f26368,f26358,f26359,f26360,f26361,f26372,f26391,f26408,f26427,f26442,f26458,f26459,f26460,f26461,f26709,f26711,f26714,f26717,f26718,f26719,f26720,f26721,f26723,f26731,f26733,f26735,f26736,f26737,f26738,f26739,f26740,f26741,f26742,f26521,f26523,f26525,f26527,f26751,f26752,f26755,f26537,f26540,f26542,f26543,f26545,f26757,f26759,f26761,f26763,f26764,f26765,f26767,f26768,f26555,f26557,f26560,f26772,f26777,f26783,f26803,f26820,f26840,f26865,f26866,f26867,f26868,f26869,f26921,f26923,f26925,f26927,f26928,f26948,f26949,f26950,f26951,f26952,f26953,f26954,f26955,f26957,f26959,f26960,f26961,f26962,f26963,f26966,f26967,f26968,f26970,f26971,f26972,f26973,f26974,f26975,f26976,f26977,f26978,f26979,f26984,f26985,f26987,f26989,f26991,f26993,f26995,f26996,f26997,f26998,f26999,f27000,f27001,f27002,f27003,f27006,f27008,f27009,f27011,f27013,f27014,f27015,f27016,f27017,f27018,f27019,f27020,f27021,f27023,f27026,f27030,f27033,f27043,f27106,f27107,f27109,f27119,f27132,f27133,f27134,f27135,f27136,f27137,f27138,f27139,f27140,f27141,f27155,f27156,f27157,f27158,f27159,f27160,f27161,f27162,f27166,f27167,f27168,f27169,f27199,f25807,f27298,f27559,f27562,f27567,f27568,f27569,f27570,f28094,f28095,f28097,f28099,f27582,f27583,f27585,f27587,f27589,f27590,f27592,f27594,f27597,f27599,f27602,f27606,f27607,f27608,f27609,f27610,f27611,f27613,f27614,f27619,f27620,f27635,f27636,f27639,f27640,f27641,f27642,f27643,f27645,f27665,f28144,f28149,f28151])).
fof(f28150,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f334,f333,f335,f343,f355,f354,f413,f428,f427,f426,f429,f430,f437,f434,f444,f443,f593,f462,f459,f458,f457,f478,f479,f480,f481,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f819,f825,f828,f830,f840,f888,f898,f901,f917,f923,f933,f939,f946,f954,f960,f966,f971,f975,f980,f984,f994,f1000,f1007,f1012,f1024,f1032,f1047,f1050,f1062,f1068,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1239,f1275,f1339,f1332,f1373,f1385,f1402,f1422,f1433,f1437,f1461,f1501,f1512,f1516,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1662,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1981,f2015,f2030,f2073,f2082,f2119,f2118,f2128,f2157,f2170,f2252,f2222,f2262,f2270,f2617,f2622,f2841,f2844,f2751,f2544,f2504,f2465,f2432,f2400,f2371,f2351,f2317,f1837,f1818,f424,f423,f422,f421,f420,f419,f418,f417,f412,f411,f1854,f1861,f1852,f1859,f2052,f2050,f2043,f2041,f482,f433,f1542,f1492,f1486,f489,f488,f487,f485,f484,f1219,f1249,f1217,f1247,f1212,f1256,f1285,f1210,f1254,f1283,f2099,f2097,f2090,f2088,f2001,f1931,f1230,f483,f477,f468,f467,f432,f431,f2848,f2847,f2770,f2765,f2545,f2505,f2466,f2433,f2401,f2398,f2373,f2372,f2352,f2349,f2319,f2318,f2272,f2271,f2269,f2251,f1838,f1835,f1819,f1816,f1663,f1661,f1494,f1411,f1410,f1403,f1366,f1365,f1333,f1267,f1231,f1069,f1059,f1057,f1048,f1035,f1033,f1025,f1023,f1001,f991,f961,f942,f940,f930,f926,f924,f904,f902,f891,f889,f873,f871,f870,f869,f868,f867,f866,f865,f864,f863,f862,f861,f852,f839,f838,f837,f836,f829,f818,f787,f592,f414,f359,f358,f4838,f4843,f4844,f4853,f4855,f4857,f4858,f4859,f4860,f4861,f4862,f4863,f4864,f4644,f4646,f4648,f4650,f4870,f4655,f4658,f4660,f4661,f4663,f4669,f4977,f4986,f5268,f5276,f7384,f7394,f7492,f7497,f7506,f7666,f7671,f7680,f7351,f11643,f7360,f11657,f11665,f12884,f12883,f12097,f12096,f13263,f13282,f13283,f13284,f13285,f13298,f13312,f13313,f13314,f13315,f13323,f13324,f13325,f13326,f13327,f13328,f13329,f13331,f13333,f13335,f13337,f13343,f13346,f13348,f13349,f13351,f13356,f13358,f13361,f16694,f16695,f13454,f16696,f13456,f16697,f16698,f16699,f16700,f16701,f16717,f16728,f18264,f18263,f17257,f17256,f18761,f18782,f18783,f18784,f18785,f18786,f18802,f18816,f18817,f18818,f18819,f18828,f18829,f18830,f18831,f18832,f18833,f18834,f18836,f18838,f18840,f18842,f18848,f18851,f18853,f18854,f18856,f18861,f18863,f18866,f23086,f23087,f18961,f23088,f18963,f23089,f23090,f23091,f23092,f23093,f23129,f23139,f23156,f23157,f23158,f23159,f23173,f23183,f23184,f23185,f23186,f25796,f23209,f23235,f23237,f23239,f23241,f23247,f23250,f23252,f23253,f23255,f23261,f23263,f23266,f25806,f23476,f23477,f23508,f23513,f23514,f23515,f23516,f23517,f23518,f23519,f23520,f23521,f23522,f23523,f25824,f23848,f23849,f23850,f23851,f23852,f23853,f23854,f23855,f23856,f23857,f23858,f23859,f23860,f23861,f23862,f23863,f23864,f23865,f23866,f23867,f23868,f23869,f23870,f23871,f23872,f23873,f23874,f23875,f23876,f23877,f23878,f23879,f23880,f23881,f23882,f23883,f23884,f23885,f23886,f23887,f23888,f23889,f23890,f23891,f23892,f23893,f23894,f23895,f23896,f23897,f23898,f23899,f23900,f23901,f23902,f23903,f23904,f23905,f23906,f23907,f23908,f23909,f23910,f23911,f23912,f23913,f23914,f23915,f23916,f23917,f23918,f23919,f23920,f23921,f23922,f23923,f23924,f23925,f23926,f23927,f23928,f23929,f23930,f23931,f23932,f23933,f23934,f23935,f23936,f23937,f23938,f23939,f25830,f23947,f25837,f25840,f25846,f25849,f25852,f25855,f25858,f25861,f25864,f23973,f23975,f23977,f23979,f23985,f23988,f23990,f23991,f23993,f24001,f25872,f25873,f24118,f25875,f25876,f25877,f25878,f25879,f25880,f25883,f25884,f25885,f25886,f25887,f25888,f25889,f25890,f25891,f24144,f24146,f24148,f24150,f25892,f25893,f25894,f25895,f24156,f24159,f24161,f24162,f24164,f25896,f24170,f24172,f24175,f25897,f25898,f25900,f25901,f25902,f25903,f25904,f25905,f25906,f25907,f25908,f25909,f25910,f25911,f25912,f25913,f25914,f25915,f25916,f25917,f25918,f24416,f24421,f24422,f24423,f24424,f24425,f24426,f24427,f24428,f24429,f24430,f24431,f25920,f25926,f25928,f25933,f25935,f25937,f25939,f25941,f25943,f25945,f25998,f26013,f26015,f26018,f26019,f26009,f26010,f26011,f26012,f26024,f26039,f26041,f26044,f26045,f26035,f26036,f26037,f26038,f26053,f26068,f26070,f26073,f26075,f26090,f26105,f26107,f26110,f26111,f26101,f26102,f26103,f26104,f26118,f26133,f26135,f26138,f26139,f26129,f26130,f26131,f26132,f26141,f26160,f26178,f26193,f26195,f26198,f26199,f26189,f26190,f26191,f26192,f26203,f26233,f26260,f26275,f26277,f26280,f26282,f26290,f26320,f26347,f26362,f26364,f26367,f26368,f26358,f26359,f26360,f26361,f26372,f26391,f26408,f26427,f26442,f26458,f26459,f26460,f26461,f26709,f26711,f26714,f26717,f26718,f26719,f26720,f26721,f26723,f26731,f26733,f26735,f26736,f26737,f26738,f26739,f26740,f26741,f26742,f26521,f26523,f26525,f26527,f26751,f26752,f26755,f26537,f26540,f26542,f26543,f26545,f26757,f26759,f26761,f26763,f26764,f26765,f26767,f26768,f26555,f26557,f26560,f26772,f26777,f26783,f26803,f26820,f26840,f26865,f26866,f26867,f26868,f26869,f26921,f26923,f26925,f26927,f26928,f26948,f26949,f26950,f26951,f26952,f26953,f26954,f26955,f26957,f26959,f26960,f26961,f26962,f26963,f26966,f26967,f26968,f26970,f26971,f26972,f26973,f26974,f26975,f26976,f26977,f26978,f26979,f26984,f26985,f26987,f26989,f26991,f26993,f26995,f26996,f26997,f26998,f26999,f27000,f27001,f27002,f27003,f27006,f27008,f27009,f27011,f27013,f27014,f27015,f27016,f27017,f27018,f27019,f27020,f27021,f27023,f27026,f27030,f27033,f27043,f27106,f27107,f27109,f27119,f27132,f27133,f27134,f27135,f27136,f27137,f27138,f27139,f27140,f27141,f27155,f27156,f27157,f27158,f27159,f27160,f27161,f27162,f27166,f27167,f27168,f27169,f27199,f25807,f27298,f27559,f27562,f27567,f27568,f27569,f27570,f28094,f28095,f28097,f28099,f27582,f27583,f27585,f27587,f27589,f27590,f27592,f27594,f27597,f27599,f27602,f27606,f27607,f27608,f27609,f27610,f27611,f27613,f27614,f27619,f27620,f27635,f27636,f27639,f27640,f27641,f27642,f27643,f27645,f27665,f28144,f28149])).
fof(f28148,plain,(
  $false | ($spl8 | $spl14)),
  inference(global_subsumption,[],[f27803,f334,f333,f335,f343,f355,f354,f413,f428,f427,f426,f429,f430,f437,f434,f444,f443,f593,f462,f459,f458,f457,f478,f479,f480,f481,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f819,f825,f828,f830,f840,f888,f898,f901,f917,f923,f933,f939,f946,f954,f960,f966,f971,f975,f980,f984,f994,f1000,f1007,f1012,f1024,f1032,f1047,f1050,f1062,f1068,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1239,f1275,f1339,f1332,f1373,f1385,f1402,f1422,f1433,f1437,f1461,f1501,f1512,f1516,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1662,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1981,f2015,f2030,f2073,f2082,f2119,f2118,f2128,f2157,f2170,f2252,f2222,f2262,f2270,f2617,f2622,f2841,f2844,f2751,f2544,f2504,f2465,f2432,f2400,f2371,f2351,f2317,f1837,f1818,f424,f423,f422,f421,f420,f419,f418,f417,f412,f411,f1854,f1861,f1852,f1859,f2052,f2050,f2043,f2041,f482,f433,f1542,f1492,f1486,f489,f488,f487,f485,f484,f1219,f1249,f1217,f1247,f1212,f1256,f1285,f1210,f1254,f1283,f2099,f2097,f2090,f2088,f2001,f1931,f1230,f483,f477,f468,f467,f432,f431,f2848,f2847,f2770,f2765,f2545,f2505,f2466,f2433,f2401,f2398,f2373,f2372,f2352,f2349,f2319,f2318,f2272,f2271,f2269,f2251,f1838,f1835,f1819,f1816,f1663,f1661,f1494,f1411,f1410,f1403,f1366,f1365,f1333,f1267,f1231,f1069,f1059,f1057,f1048,f1035,f1033,f1025,f1023,f1001,f991,f961,f942,f940,f930,f926,f924,f904,f902,f891,f889,f873,f871,f870,f869,f868,f867,f866,f865,f864,f863,f862,f861,f852,f839,f838,f837,f836,f829,f818,f787,f592,f414,f359,f358,f4838,f4843,f4844,f4853,f4855,f4857,f4858,f4859,f4860,f4861,f4862,f4863,f4864,f4644,f4646,f4648,f4650,f4870,f4655,f4658,f4660,f4661,f4663,f4669,f4977,f4986,f5268,f5276,f7384,f7394,f7492,f7497,f7506,f7666,f7671,f7680,f7351,f11643,f7360,f11657,f11665,f12884,f12883,f12097,f12096,f13263,f13282,f13283,f13284,f13285,f13298,f13312,f13313,f13314,f13315,f13323,f13324,f13325,f13326,f13327,f13328,f13329,f13331,f13333,f13335,f13337,f13343,f13346,f13348,f13349,f13351,f13356,f13358,f13361,f16694,f16695,f13454,f16696,f13456,f16697,f16698,f16699,f16700,f16701,f16717,f16728,f18264,f18263,f17257,f17256,f18761,f18782,f18783,f18784,f18785,f18786,f18802,f18816,f18817,f18818,f18819,f18828,f18829,f18830,f18831,f18832,f18833,f18834,f18836,f18838,f18840,f18842,f18848,f18851,f18853,f18854,f18856,f18861,f18863,f18866,f23086,f23087,f18961,f23088,f18963,f23089,f23090,f23091,f23092,f23093,f23129,f23139,f23156,f23157,f23158,f23159,f23173,f23183,f23184,f23185,f23186,f25796,f23209,f23235,f23237,f23239,f23241,f23247,f23250,f23252,f23253,f23255,f23261,f23263,f23266,f25806,f23476,f23477,f23508,f23513,f23514,f23515,f23516,f23517,f23518,f23519,f23520,f23521,f23522,f23523,f25824,f23848,f23849,f23850,f23851,f23852,f23853,f23854,f23855,f23856,f23857,f23858,f23859,f23860,f23861,f23862,f23863,f23864,f23865,f23866,f23867,f23868,f23869,f23870,f23871,f23872,f23873,f23874,f23875,f23876,f23877,f23878,f23879,f23880,f23881,f23882,f23883,f23884,f23885,f23886,f23887,f23888,f23889,f23890,f23891,f23892,f23893,f23894,f23895,f23896,f23897,f23898,f23899,f23900,f23901,f23902,f23903,f23904,f23905,f23906,f23907,f23908,f23909,f23910,f23911,f23912,f23913,f23914,f23915,f23916,f23917,f23918,f23919,f23920,f23921,f23922,f23923,f23924,f23925,f23926,f23927,f23928,f23929,f23930,f23931,f23932,f23933,f23934,f23935,f23936,f23937,f23938,f23939,f25830,f23947,f25837,f25840,f25846,f25849,f25852,f25855,f25858,f25861,f25864,f23973,f23975,f23977,f23979,f23985,f23988,f23990,f23991,f23993,f24001,f25872,f25873,f24118,f25875,f25876,f25877,f25878,f25879,f25880,f25883,f25884,f25885,f25886,f25887,f25888,f25889,f25890,f25891,f24144,f24146,f24148,f24150,f25892,f25893,f25894,f25895,f24156,f24159,f24161,f24162,f24164,f25896,f24170,f24172,f24175,f25897,f25898,f25900,f25901,f25902,f25903,f25904,f25905,f25906,f25907,f25908,f25909,f25910,f25911,f25912,f25913,f25914,f25915,f25916,f25917,f25918,f24416,f24421,f24422,f24423,f24424,f24425,f24426,f24427,f24428,f24429,f24430,f24431,f25920,f25926,f25928,f25933,f25935,f25937,f25939,f25941,f25943,f25945,f25998,f26013,f26015,f26018,f26019,f26009,f26010,f26011,f26012,f26024,f26039,f26041,f26044,f26045,f26035,f26036,f26037,f26038,f26053,f26068,f26070,f26073,f26075,f26090,f26105,f26107,f26110,f26111,f26101,f26102,f26103,f26104,f26118,f26133,f26135,f26138,f26139,f26129,f26130,f26131,f26132,f26141,f26160,f26178,f26193,f26195,f26198,f26199,f26189,f26190,f26191,f26192,f26203,f26233,f26260,f26275,f26277,f26280,f26282,f26290,f26320,f26347,f26362,f26364,f26367,f26368,f26358,f26359,f26360,f26361,f26372,f26391,f26408,f26427,f26442,f26458,f26459,f26460,f26461,f26709,f26711,f26714,f26717,f26718,f26719,f26720,f26721,f26723,f26731,f26733,f26735,f26736,f26737,f26738,f26739,f26740,f26741,f26742,f26521,f26523,f26525,f26527,f26751,f26752,f26755,f26537,f26540,f26542,f26543,f26545,f26757,f26759,f26761,f26763,f26764,f26765,f26767,f26768,f26555,f26557,f26560,f26772,f26777,f26783,f26803,f26820,f26840,f26865,f26866,f26867,f26868,f26869,f26921,f26923,f26925,f26927,f26928,f26948,f26949,f26950,f26951,f26952,f26953,f26954,f26955,f26957,f26959,f26960,f26961,f26962,f26963,f26966,f26967,f26968,f26970,f26971,f26972,f26973,f26974,f26975,f26976,f26977,f26978,f26979,f26984,f26985,f26987,f26989,f26991,f26993,f26995,f26996,f26997,f26998,f26999,f27000,f27001,f27002,f27003,f27006,f27008,f27009,f27011,f27013,f27014,f27015,f27016,f27017,f27018,f27019,f27020,f27021,f27023,f27026,f27030,f27033,f27043,f27106,f27107,f27109,f27119,f27132,f27133,f27134,f27135,f27136,f27137,f27138,f27139,f27140,f27141,f27155,f27156,f27157,f27158,f27159,f27160,f27161,f27162,f27166,f27167,f27168,f27169,f27199,f25807,f27298,f27559,f27562,f27567,f27568,f27569,f27570,f28094,f28095,f28097,f28099,f27582,f27583,f27585,f27587,f27589,f27590,f27592,f27594,f27597,f27599,f27602,f27606,f27607,f27608,f27609,f27610,f27611,f27613,f27614,f27619,f27620,f27635,f27636,f27639,f27640,f27641,f27642,f27643,f27645,f27665,f28144])).
fof(f27803,plain,(
  ( ! [X0] : (sP1('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)))) ) | ($spl8 | $spl14)),
  inference(backward_demodulation,[],[f27559,f27232])).
fof(f27232,plain,(
  ( ! [X0] : (sP1('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X0)))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f592,f25970])).
fof(f28145,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f334,f333,f335,f343,f355,f354,f413,f428,f427,f426,f429,f430,f437,f434,f444,f443,f593,f462,f459,f458,f457,f478,f479,f480,f481,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f819,f825,f828,f830,f840,f888,f898,f901,f917,f923,f933,f939,f946,f954,f960,f966,f971,f975,f980,f984,f994,f1000,f1007,f1012,f1024,f1032,f1047,f1050,f1062,f1068,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1239,f1275,f1339,f1332,f1373,f1385,f1402,f1422,f1433,f1437,f1461,f1501,f1512,f1516,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1662,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1981,f2015,f2030,f2073,f2082,f2119,f2118,f2128,f2157,f2170,f2252,f2222,f2262,f2270,f2617,f2622,f2841,f2844,f2751,f2544,f2504,f2465,f2432,f2400,f2371,f2351,f2317,f1837,f1818,f424,f423,f422,f421,f420,f419,f418,f417,f412,f411,f1854,f1861,f1852,f1859,f2052,f2050,f2043,f2041,f482,f433,f1542,f1492,f1486,f489,f488,f487,f485,f484,f1219,f1249,f1217,f1247,f1212,f1256,f1285,f1210,f1254,f1283,f2099,f2097,f2090,f2088,f2001,f1931,f1230,f483,f477,f468,f467,f432,f431,f2848,f2847,f2770,f2765,f2545,f2505,f2466,f2433,f2401,f2398,f2373,f2372,f2352,f2349,f2319,f2318,f2272,f2271,f2269,f2251,f1838,f1835,f1819,f1816,f1663,f1661,f1494,f1411,f1410,f1403,f1366,f1365,f1333,f1267,f1231,f1069,f1059,f1057,f1048,f1035,f1033,f1025,f1023,f1001,f991,f961,f942,f940,f930,f926,f924,f904,f902,f891,f889,f873,f871,f870,f869,f868,f867,f866,f865,f864,f863,f862,f861,f852,f839,f838,f837,f836,f829,f818,f787,f592,f414,f359,f358,f4838,f4843,f4844,f4853,f4855,f4857,f4858,f4859,f4860,f4861,f4862,f4863,f4864,f4644,f4646,f4648,f4650,f4870,f4655,f4658,f4660,f4661,f4663,f4669,f4977,f4986,f5268,f5276,f7384,f7394,f7492,f7497,f7506,f7666,f7671,f7680,f7351,f11643,f7360,f11657,f11665,f12884,f12883,f12097,f12096,f13263,f13282,f13283,f13284,f13285,f13298,f13312,f13313,f13314,f13315,f13323,f13324,f13325,f13326,f13327,f13328,f13329,f13331,f13333,f13335,f13337,f13343,f13346,f13348,f13349,f13351,f13356,f13358,f13361,f16694,f16695,f13454,f16696,f13456,f16697,f16698,f16699,f16700,f16701,f16717,f16728,f18264,f18263,f17257,f17256,f18761,f18782,f18783,f18784,f18785,f18786,f18802,f18816,f18817,f18818,f18819,f18828,f18829,f18830,f18831,f18832,f18833,f18834,f18836,f18838,f18840,f18842,f18848,f18851,f18853,f18854,f18856,f18861,f18863,f18866,f23086,f23087,f18961,f23088,f18963,f23089,f23090,f23091,f23092,f23093,f23129,f23139,f23156,f23157,f23158,f23159,f23173,f23183,f23184,f23185,f23186,f25796,f23209,f23235,f23237,f23239,f23241,f23247,f23250,f23252,f23253,f23255,f23261,f23263,f23266,f25806,f23476,f23477,f23508,f23513,f23514,f23515,f23516,f23517,f23518,f23519,f23520,f23521,f23522,f23523,f25824,f23848,f23849,f23850,f23851,f23852,f23853,f23854,f23855,f23856,f23857,f23858,f23859,f23860,f23861,f23862,f23863,f23864,f23865,f23866,f23867,f23868,f23869,f23870,f23871,f23872,f23873,f23874,f23875,f23876,f23877,f23878,f23879,f23880,f23881,f23882,f23883,f23884,f23885,f23886,f23887,f23888,f23889,f23890,f23891,f23892,f23893,f23894,f23895,f23896,f23897,f23898,f23899,f23900,f23901,f23902,f23903,f23904,f23905,f23906,f23907,f23908,f23909,f23910,f23911,f23912,f23913,f23914,f23915,f23916,f23917,f23918,f23919,f23920,f23921,f23922,f23923,f23924,f23925,f23926,f23927,f23928,f23929,f23930,f23931,f23932,f23933,f23934,f23935,f23936,f23937,f23938,f23939,f25830,f23947,f25837,f25840,f25846,f25849,f25852,f25855,f25858,f25861,f25864,f23973,f23975,f23977,f23979,f23985,f23988,f23990,f23991,f23993,f24001,f25872,f25873,f24118,f25875,f25876,f25877,f25878,f25879,f25880,f25883,f25884,f25885,f25886,f25887,f25888,f25889,f25890,f25891,f24144,f24146,f24148,f24150,f25892,f25893,f25894,f25895,f24156,f24159,f24161,f24162,f24164,f25896,f24170,f24172,f24175,f25897,f25898,f25900,f25901,f25902,f25903,f25904,f25905,f25906,f25907,f25908,f25909,f25910,f25911,f25912,f25913,f25914,f25915,f25916,f25917,f25918,f24416,f24421,f24422,f24423,f24424,f24425,f24426,f24427,f24428,f24429,f24430,f24431,f25920,f25926,f25928,f25933,f25935,f25937,f25939,f25941,f25943,f25945,f25998,f26013,f26015,f26018,f26019,f26009,f26010,f26011,f26012,f26024,f26039,f26041,f26044,f26045,f26035,f26036,f26037,f26038,f26053,f26068,f26070,f26073,f26075,f26090,f26105,f26107,f26110,f26111,f26101,f26102,f26103,f26104,f26118,f26133,f26135,f26138,f26139,f26129,f26130,f26131,f26132,f26141,f26160,f26178,f26193,f26195,f26198,f26199,f26189,f26190,f26191,f26192,f26203,f26233,f26260,f26275,f26277,f26280,f26282,f26290,f26320,f26347,f26362,f26364,f26367,f26368,f26358,f26359,f26360,f26361,f26372,f26391,f26408,f26427,f26442,f26458,f26459,f26460,f26461,f26709,f26711,f26714,f26717,f26718,f26719,f26720,f26721,f26723,f26731,f26733,f26735,f26736,f26737,f26738,f26739,f26740,f26741,f26742,f26521,f26523,f26525,f26527,f26751,f26752,f26755,f26537,f26540,f26542,f26543,f26545,f26757,f26759,f26761,f26763,f26764,f26765,f26767,f26768,f26555,f26557,f26560,f26772,f26777,f26783,f26803,f26820,f26840,f26865,f26866,f26867,f26868,f26869,f26921,f26923,f26925,f26927,f26928,f26948,f26949,f26950,f26951,f26952,f26953,f26954,f26955,f26957,f26959,f26960,f26961,f26962,f26963,f26966,f26967,f26968,f26970,f26971,f26972,f26973,f26974,f26975,f26976,f26977,f26978,f26979,f26984,f26985,f26987,f26989,f26991,f26993,f26995,f26996,f26997,f26998,f26999,f27000,f27001,f27002,f27003,f27006,f27008,f27009,f27011,f27013,f27014,f27015,f27016,f27017,f27018,f27019,f27020,f27021,f27023,f27026,f27030,f27033,f27043,f27106,f27107,f27109,f27119,f27132,f27133,f27134,f27135,f27136,f27137,f27138,f27139,f27140,f27141,f27155,f27156,f27157,f27158,f27159,f27160,f27161,f27162,f27166,f27167,f27168,f27169,f27199,f25807,f27298,f27559,f27562,f27567,f27568,f27569,f27570,f28094,f28095,f28097,f28099,f27582,f27583,f27585,f27587,f27589,f27590,f27592,f27594,f27597,f27599,f27602,f27606,f27607,f27608,f27609,f27610,f27611,f27613,f27614,f27619,f27620,f27635,f27636,f27639,f27640,f27641,f27642,f27643,f27645,f27665,f28144])).
fof(f28093,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f27563,f335])).
fof(f27563,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(backward_demodulation,[],[f27559,f592])).
fof(f28092,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f28091,f25807])).
fof(f28091,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(subsumption_resolution,[],[f27560,f27559])).
fof(f27560,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl8),
  inference(backward_demodulation,[],[f27559,f358])).
fof(f27277,plain,(
  ( ! [X8,X7] : (sP1('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X7),s(bool,X8))) | s(bool,'T_0') = s(bool,X7)) ) | $spl48),
  inference(cnf_transformation,[],[f27277_D])).
fof(f27277_D,plain,(
  ( ! [X8,X7] : (sP1('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X7),s(bool,X8))) | s(bool,'T_0') = s(bool,X7)) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f27278,plain,(
  ( ! [X10,X11,X9] : (sP1('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X10),s(bool,X11))) | p(s(bool,X9)) | s(bool,'T_0') = s(bool,X10)) ) | $spl14),
  inference(subsumption_resolution,[],[f27243,f354])).
fof(f27243,plain,(
  ( ! [X10,X11,X9] : (sP1('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X10),s(bool,X11))) | p(s(bool,X9)) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X9)) ) | $spl14),
  inference(superposition,[],[f25970,f701])).
fof(f27276,plain,(
  ( ! [X6,X8,X7] : (sP1('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X7),s(bool,X8))) | p(s(bool,X6)) | s(bool,'T_0') = s(bool,X7)) ) | $spl14),
  inference(subsumption_resolution,[],[f27242,f354])).
fof(f27242,plain,(
  ( ! [X6,X8,X7] : (sP1('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X7),s(bool,X8))) | p(s(bool,X6)) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,X7)) ) | $spl14),
  inference(superposition,[],[f25970,f701])).
fof(f26365,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X1)),'T_0',X3)) ) | $spl46),
  inference(cnf_transformation,[],[f26365_D])).
fof(f26365_D,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X1)),'T_0',X3)) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f26367,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X1)),'T_0',X3)) )),
  inference(condensation,[],[f26366])).
fof(f26366,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X30)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X31),s(bool,X30)),'T_0',X32) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f26355,f747])).
fof(f26355,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | ~p(s(bool,X30)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X31),s(bool,X30)),'T_0',X32) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f868,f701])).
fof(f26364,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X1)),'T_0',X3)) )),
  inference(condensation,[],[f26363])).
fof(f26363,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X25)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X26),s(bool,X25)),'T_0',X27) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f26354,f747])).
fof(f26354,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X26),s(bool,X25)),'T_0',X27) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f868,f701])).
fof(f26278,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,'T_0')),X3,X1)) ) | $spl44),
  inference(cnf_transformation,[],[f26278_D])).
fof(f26278_D,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,'T_0')),X3,X1)) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f26280,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,'T_0')),X3,X1)) )),
  inference(condensation,[],[f26279])).
fof(f26279,plain,(
  ( ! [X30,X33,X31,X29,X32] : (p(s(bool,X30)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X31),s(bool,'T_0')),X32,X30) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f26268,f747])).
fof(f26268,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | p(s(bool,X30)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X31),s(bool,'T_0')),X32,X30) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f867,f701])).
fof(f26277,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,'T_0')),X3,X1)) )),
  inference(condensation,[],[f26276])).
fof(f26276,plain,(
  ( ! [X28,X26,X24,X27,X25] : (p(s(bool,X25)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X26),s(bool,'T_0')),X27,X25) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f26267,f747])).
fof(f26267,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | p(s(bool,X25)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X26),s(bool,'T_0')),X27,X25) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f867,f701])).
fof(f26196,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X1)),'T_0',X3)) ) | $spl42),
  inference(cnf_transformation,[],[f26196_D])).
fof(f26196_D,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X1)),'T_0',X3)) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f26198,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X1)),'T_0',X3)) )),
  inference(condensation,[],[f26197])).
fof(f26197,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X30)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X31),s(bool,X30)),'T_0',X32) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f26186,f747])).
fof(f26186,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | ~p(s(bool,X30)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X31),s(bool,X30)),'T_0',X32) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f866,f701])).
fof(f26195,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X1)),'T_0',X3)) )),
  inference(condensation,[],[f26194])).
fof(f26194,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X25)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X26),s(bool,X25)),'T_0',X27) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f26185,f747])).
fof(f26185,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X26),s(bool,X25)),'T_0',X27) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f866,f701])).
fof(f26136,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X1)),X3,'T_0')) ) | $spl40),
  inference(cnf_transformation,[],[f26136_D])).
fof(f26136_D,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X1)),X3,'T_0')) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f26138,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X1)),X3,'T_0')) )),
  inference(condensation,[],[f26137])).
fof(f26137,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X30)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X31),s(bool,X30)),X32,'T_0') | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f26126,f747])).
fof(f26126,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | ~p(s(bool,X30)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X31),s(bool,X30)),X32,'T_0') | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f865,f701])).
fof(f26135,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X1)),X3,'T_0')) )),
  inference(condensation,[],[f26134])).
fof(f26134,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X25)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X26),s(bool,X25)),X27,'T_0') | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f26125,f747])).
fof(f26125,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X26),s(bool,X25)),X27,'T_0') | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f865,f701])).
fof(f26108,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2)),'T_0',X3)) ) | $spl38),
  inference(cnf_transformation,[],[f26108_D])).
fof(f26108_D,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2)),'T_0',X3)) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f26110,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2)),'T_0',X3)) )),
  inference(condensation,[],[f26109])).
fof(f26109,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X30)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X30),s(bool,X31)),'T_0',X32) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f26098,f747])).
fof(f26098,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | ~p(s(bool,X30)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X30),s(bool,X31)),'T_0',X32) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f803,f701])).
fof(f26107,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2)),'T_0',X3)) )),
  inference(condensation,[],[f26106])).
fof(f26106,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X25)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X25),s(bool,X26)),'T_0',X27) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f26097,f747])).
fof(f26097,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X25),s(bool,X26)),'T_0',X27) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f803,f701])).
fof(f26071,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X2)),X3,X1)) ) | $spl36),
  inference(cnf_transformation,[],[f26071_D])).
fof(f26071_D,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X2)),X3,X1)) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f26073,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X2)),X3,X1)) )),
  inference(condensation,[],[f26072])).
fof(f26072,plain,(
  ( ! [X30,X33,X31,X29,X32] : (p(s(bool,X30)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X31)),X32,X30) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f26061,f747])).
fof(f26061,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | p(s(bool,X30)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X31)),X32,X30) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f802,f701])).
fof(f26070,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X2)),X3,X1)) )),
  inference(condensation,[],[f26069])).
fof(f26069,plain,(
  ( ! [X28,X26,X24,X27,X25] : (p(s(bool,X25)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X26)),X27,X25) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f26060,f747])).
fof(f26060,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | p(s(bool,X25)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X26)),X27,X25) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f802,f701])).
fof(f26042,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2)),'T_0',X3)) ) | $spl34),
  inference(cnf_transformation,[],[f26042_D])).
fof(f26042_D,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2)),'T_0',X3)) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f26044,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2)),'T_0',X3)) )),
  inference(condensation,[],[f26043])).
fof(f26043,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X30)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X30),s(bool,X31)),'T_0',X32) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f26032,f747])).
fof(f26032,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | ~p(s(bool,X30)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X30),s(bool,X31)),'T_0',X32) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f801,f701])).
fof(f26041,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2)),'T_0',X3)) )),
  inference(condensation,[],[f26040])).
fof(f26040,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X25)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X25),s(bool,X26)),'T_0',X27) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f26031,f747])).
fof(f26031,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X25),s(bool,X26)),'T_0',X27) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f801,f701])).
fof(f26016,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2)),X3,'T_0')) ) | $spl32),
  inference(cnf_transformation,[],[f26016_D])).
fof(f26016_D,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2)),X3,'T_0')) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f26018,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2)),X3,'T_0')) )),
  inference(condensation,[],[f26017])).
fof(f26017,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X30)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X30),s(bool,X31)),X32,'T_0') | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f26006,f747])).
fof(f26006,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | ~p(s(bool,X30)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X30),s(bool,X31)),X32,'T_0') | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f800,f701])).
fof(f26015,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2)),X3,'T_0')) )),
  inference(condensation,[],[f26014])).
fof(f26014,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X25)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X25),s(bool,X26)),X27,'T_0') | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f26005,f747])).
fof(f26005,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X25),s(bool,X26)),X27,'T_0') | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f800,f701])).
fof(f25865,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~sP8('T_0',X1,X2)) ) | $spl30),
  inference(cnf_transformation,[],[f25865_D])).
fof(f25865_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~sP8('T_0',X1,X2)) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f25945,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('T_0',X1,X2) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f25944])).
fof(f25944,plain,(
  ( ! [X478,X480,X477,X479] : (~sP8('T_0',X478,X479) | sP9('T_0',X478,X479) | s(bool,X477) = s(bool,X480) | s(bool,'T_0') = s(bool,X480)) )),
  inference(subsumption_resolution,[],[f24556,f747])).
fof(f24556,plain,(
  ( ! [X478,X480,X477,X479] : (~p(s(bool,X477)) | ~sP8('T_0',X478,X479) | sP9('T_0',X478,X479) | s(bool,X477) = s(bool,X480) | s(bool,'T_0') = s(bool,X480)) )),
  inference(superposition,[],[f577,f701])).
fof(f25862,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~sP10('T_0',X1,X2)) ) | $spl28),
  inference(cnf_transformation,[],[f25862_D])).
fof(f25862_D,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~sP10('T_0',X1,X2)) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f25943,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP10('T_0',X1,X2) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f25942])).
fof(f25942,plain,(
  ( ! [X476,X474,X475,X473] : (~sP10('T_0',X474,X475) | sP11('T_0',X474,X475) | s(bool,X473) = s(bool,X476) | s(bool,'T_0') = s(bool,X476)) )),
  inference(subsumption_resolution,[],[f24555,f747])).
fof(f24555,plain,(
  ( ! [X476,X474,X475,X473] : (~p(s(bool,X473)) | ~sP10('T_0',X474,X475) | sP11('T_0',X474,X475) | s(bool,X473) = s(bool,X476) | s(bool,'T_0') = s(bool,X476)) )),
  inference(superposition,[],[f576,f701])).
fof(f25859,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~p(s(bool,X1))) ) | $spl26),
  inference(cnf_transformation,[],[f25859_D])).
fof(f25859_D,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~p(s(bool,X1))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f25941,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f25940])).
fof(f25940,plain,(
  ( ! [X472,X469,X471,X470] : (~p(s(bool,X470)) | sP11('T_0',X470,X471) | s(bool,X469) = s(bool,X472) | s(bool,'T_0') = s(bool,X472)) )),
  inference(subsumption_resolution,[],[f24554,f747])).
fof(f24554,plain,(
  ( ! [X472,X469,X471,X470] : (~p(s(bool,X469)) | ~p(s(bool,X470)) | sP11('T_0',X470,X471) | s(bool,X469) = s(bool,X472) | s(bool,'T_0') = s(bool,X472)) )),
  inference(superposition,[],[f572,f701])).
fof(f25856,plain,(
  ( ! [X2,X1] : (sP11('T_0',X2,X1) | p(s(bool,X1))) ) | $spl24),
  inference(cnf_transformation,[],[f25856_D])).
fof(f25856_D,plain,(
  ( ! [X2,X1] : (sP11('T_0',X2,X1) | p(s(bool,X1))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f25939,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f25938])).
fof(f25938,plain,(
  ( ! [X466,X465,X467,X468] : (p(s(bool,X466)) | sP11('T_0',X467,X466) | s(bool,X465) = s(bool,X468) | s(bool,'T_0') = s(bool,X468)) )),
  inference(subsumption_resolution,[],[f24553,f747])).
fof(f24553,plain,(
  ( ! [X466,X465,X467,X468] : (~p(s(bool,X465)) | p(s(bool,X466)) | sP11('T_0',X467,X466) | s(bool,X465) = s(bool,X468) | s(bool,'T_0') = s(bool,X468)) )),
  inference(superposition,[],[f571,f701])).
fof(f25853,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~p(s(bool,X1))) ) | $spl22),
  inference(cnf_transformation,[],[f25853_D])).
fof(f25853_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~p(s(bool,X1))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f25937,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f25936])).
fof(f25936,plain,(
  ( ! [X461,X463,X462,X464] : (~p(s(bool,X462)) | sP9('T_0',X462,X463) | s(bool,X461) = s(bool,X464) | s(bool,'T_0') = s(bool,X464)) )),
  inference(subsumption_resolution,[],[f24552,f747])).
fof(f24552,plain,(
  ( ! [X461,X463,X462,X464] : (~p(s(bool,X461)) | ~p(s(bool,X462)) | sP9('T_0',X462,X463) | s(bool,X461) = s(bool,X464) | s(bool,'T_0') = s(bool,X464)) )),
  inference(superposition,[],[f555,f701])).
fof(f25850,plain,(
  ( ! [X2,X1] : (sP9('T_0',X2,X1) | ~p(s(bool,X1))) ) | $spl20),
  inference(cnf_transformation,[],[f25850_D])).
fof(f25850_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X2,X1) | ~p(s(bool,X1))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f25935,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f25934])).
fof(f25934,plain,(
  ( ! [X457,X459,X460,X458] : (~p(s(bool,X458)) | sP9('T_0',X459,X458) | s(bool,X457) = s(bool,X460) | s(bool,'T_0') = s(bool,X460)) )),
  inference(subsumption_resolution,[],[f24551,f747])).
fof(f24551,plain,(
  ( ! [X457,X459,X460,X458] : (~p(s(bool,X457)) | ~p(s(bool,X458)) | sP9('T_0',X459,X458) | s(bool,X457) = s(bool,X460) | s(bool,'T_0') = s(bool,X460)) )),
  inference(superposition,[],[f554,f701])).
fof(f25847,plain,(
  ( ! [X2,X1] : (~sP7(X2,X1,'T_0') | p(s(bool,X2)) | ~p(s(bool,X1))) ) | $spl18),
  inference(cnf_transformation,[],[f25847_D])).
fof(f25847_D,plain,(
  ( ! [X2,X1] : (~sP7(X2,X1,'T_0') | p(s(bool,X2)) | ~p(s(bool,X1))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f25933,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,X1,'T_0')) )),
  inference(condensation,[],[f25932])).
fof(f25932,plain,(
  ( ! [X453,X455,X454,X456] : (~p(s(bool,X454)) | p(s(bool,X455)) | ~sP7(X455,X454,'T_0') | s(bool,X453) = s(bool,X456) | s(bool,'T_0') = s(bool,X456)) )),
  inference(subsumption_resolution,[],[f24550,f747])).
fof(f24550,plain,(
  ( ! [X453,X455,X454,X456] : (~p(s(bool,X453)) | ~p(s(bool,X454)) | p(s(bool,X455)) | ~sP7(X455,X454,'T_0') | s(bool,X453) = s(bool,X456) | s(bool,'T_0') = s(bool,X456)) )),
  inference(superposition,[],[f541,f701])).
fof(f25841,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) | $spl16),
  inference(cnf_transformation,[],[f25841_D])).
fof(f25841_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f25928,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f25927])).
fof(f25927,plain,(
  ( ! [X408,X410,X409] : (p(s(bool,X409)) | sP1(X409,'T_0') | s(bool,X408) = s(bool,X410) | s(bool,'T_0') = s(bool,X410)) )),
  inference(subsumption_resolution,[],[f24540,f747])).
fof(f24540,plain,(
  ( ! [X408,X410,X409] : (~p(s(bool,X408)) | p(s(bool,X409)) | sP1(X409,'T_0') | s(bool,X408) = s(bool,X410) | s(bool,'T_0') = s(bool,X410)) )),
  inference(superposition,[],[f462,f701])).
fof(f25838,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) | $spl14),
  inference(cnf_transformation,[],[f25838_D])).
fof(f25838_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f25926,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f25925])).
fof(f25925,plain,(
  ( ! [X405,X407,X406] : (p(s(bool,X406)) | sP1('T_0',X406) | s(bool,X405) = s(bool,X407) | s(bool,'T_0') = s(bool,X407)) )),
  inference(subsumption_resolution,[],[f24539,f747])).
fof(f24539,plain,(
  ( ! [X405,X407,X406] : (~p(s(bool,X405)) | p(s(bool,X406)) | sP1('T_0',X406) | s(bool,X405) = s(bool,X407) | s(bool,'T_0') = s(bool,X407)) )),
  inference(superposition,[],[f459,f701])).
fof(f25831,plain,(
  ( ! [X6,X4,X2,X5,X3,X1] : (s(X1,'const.bool.COND_3'(s(bool,'T_0'),s(X1,X2),s(X1,X3))) = s(X1,'const.bool.COND_3'(s(bool,X4),s(X1,X5),s(X1,X6))) | s(bool,'T_0') != s(bool,X4) | s(X1,X2) != s(X1,X5)) ) | $spl12),
  inference(cnf_transformation,[],[f25831_D])).
fof(f25831_D,plain,(
  ( ! [X6,X4,X2,X5,X3,X1] : (s(X1,'const.bool.COND_3'(s(bool,'T_0'),s(X1,X2),s(X1,X3))) = s(X1,'const.bool.COND_3'(s(bool,X4),s(X1,X5),s(X1,X6))) | s(bool,'T_0') != s(bool,X4) | s(X1,X2) != s(X1,X5)) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f25920,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(X1,'const.bool.COND_3'(s(bool,'T_0'),s(X1,X2),s(X1,X3))) = s(X1,'const.bool.COND_3'(s(bool,X4),s(X1,X5),s(X1,X6))) | s(X1,X2) != s(X1,X5) | s(bool,'T_0') != s(bool,X4)) )),
  inference(condensation,[],[f25919])).
fof(f25919,plain,(
  ( ! [X350,X356,X352,X354,X351,X357,X353,X355] : (s(X351,'const.bool.COND_3'(s(bool,'T_0'),s(X351,X355),s(X351,X356))) = s(X351,'const.bool.COND_3'(s(bool,X352),s(X351,X353),s(X351,X354))) | s(X351,X353) != s(X351,X355) | s(bool,'T_0') != s(bool,X352) | s(bool,X350) = s(bool,X357) | s(bool,'T_0') = s(bool,X357)) )),
  inference(subsumption_resolution,[],[f24529,f747])).
fof(f24529,plain,(
  ( ! [X350,X356,X352,X354,X351,X357,X353,X355] : (~p(s(bool,X350)) | s(X351,'const.bool.COND_3'(s(bool,'T_0'),s(X351,X355),s(X351,X356))) = s(X351,'const.bool.COND_3'(s(bool,X352),s(X351,X353),s(X351,X354))) | s(X351,X353) != s(X351,X355) | s(bool,'T_0') != s(bool,X352) | s(bool,X350) = s(bool,X357) | s(bool,'T_0') = s(bool,X357)) )),
  inference(superposition,[],[f426,f701])).
fof(f25864,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('T_0',X1,X2) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f25863])).
fof(f25863,plain,(
  ( ! [X478,X480,X477,X479] : (~sP8('T_0',X478,X479) | sP9('T_0',X478,X479) | s(bool,X477) = s(bool,X480) | s(bool,'T_0') = s(bool,X480)) )),
  inference(subsumption_resolution,[],[f23971,f747])).
fof(f23971,plain,(
  ( ! [X478,X480,X477,X479] : (~p(s(bool,X477)) | ~sP8('T_0',X478,X479) | sP9('T_0',X478,X479) | s(bool,X477) = s(bool,X480) | s(bool,'T_0') = s(bool,X480)) )),
  inference(superposition,[],[f577,f701])).
fof(f25861,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP10('T_0',X1,X2) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f25860])).
fof(f25860,plain,(
  ( ! [X476,X474,X475,X473] : (~sP10('T_0',X474,X475) | sP11('T_0',X474,X475) | s(bool,X473) = s(bool,X476) | s(bool,'T_0') = s(bool,X476)) )),
  inference(subsumption_resolution,[],[f23970,f747])).
fof(f23970,plain,(
  ( ! [X476,X474,X475,X473] : (~p(s(bool,X473)) | ~sP10('T_0',X474,X475) | sP11('T_0',X474,X475) | s(bool,X473) = s(bool,X476) | s(bool,'T_0') = s(bool,X476)) )),
  inference(superposition,[],[f576,f701])).
fof(f25858,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f25857])).
fof(f25857,plain,(
  ( ! [X472,X469,X471,X470] : (~p(s(bool,X470)) | sP11('T_0',X470,X471) | s(bool,X469) = s(bool,X472) | s(bool,'T_0') = s(bool,X472)) )),
  inference(subsumption_resolution,[],[f23969,f747])).
fof(f23969,plain,(
  ( ! [X472,X469,X471,X470] : (~p(s(bool,X469)) | ~p(s(bool,X470)) | sP11('T_0',X470,X471) | s(bool,X469) = s(bool,X472) | s(bool,'T_0') = s(bool,X472)) )),
  inference(superposition,[],[f572,f701])).
fof(f25855,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f25854])).
fof(f25854,plain,(
  ( ! [X466,X465,X467,X468] : (p(s(bool,X466)) | sP11('T_0',X467,X466) | s(bool,X465) = s(bool,X468) | s(bool,'T_0') = s(bool,X468)) )),
  inference(subsumption_resolution,[],[f23968,f747])).
fof(f23968,plain,(
  ( ! [X466,X465,X467,X468] : (~p(s(bool,X465)) | p(s(bool,X466)) | sP11('T_0',X467,X466) | s(bool,X465) = s(bool,X468) | s(bool,'T_0') = s(bool,X468)) )),
  inference(superposition,[],[f571,f701])).
fof(f25852,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f25851])).
fof(f25851,plain,(
  ( ! [X461,X463,X462,X464] : (~p(s(bool,X462)) | sP9('T_0',X462,X463) | s(bool,X461) = s(bool,X464) | s(bool,'T_0') = s(bool,X464)) )),
  inference(subsumption_resolution,[],[f23967,f747])).
fof(f23967,plain,(
  ( ! [X461,X463,X462,X464] : (~p(s(bool,X461)) | ~p(s(bool,X462)) | sP9('T_0',X462,X463) | s(bool,X461) = s(bool,X464) | s(bool,'T_0') = s(bool,X464)) )),
  inference(superposition,[],[f555,f701])).
fof(f25849,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f25848])).
fof(f25848,plain,(
  ( ! [X457,X459,X460,X458] : (~p(s(bool,X458)) | sP9('T_0',X459,X458) | s(bool,X457) = s(bool,X460) | s(bool,'T_0') = s(bool,X460)) )),
  inference(subsumption_resolution,[],[f23966,f747])).
fof(f23966,plain,(
  ( ! [X457,X459,X460,X458] : (~p(s(bool,X457)) | ~p(s(bool,X458)) | sP9('T_0',X459,X458) | s(bool,X457) = s(bool,X460) | s(bool,'T_0') = s(bool,X460)) )),
  inference(superposition,[],[f554,f701])).
fof(f25846,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,X1,'T_0')) )),
  inference(condensation,[],[f25845])).
fof(f25845,plain,(
  ( ! [X453,X455,X454,X456] : (~p(s(bool,X454)) | p(s(bool,X455)) | ~sP7(X455,X454,'T_0') | s(bool,X453) = s(bool,X456) | s(bool,'T_0') = s(bool,X456)) )),
  inference(subsumption_resolution,[],[f23965,f747])).
fof(f23965,plain,(
  ( ! [X453,X455,X454,X456] : (~p(s(bool,X453)) | ~p(s(bool,X454)) | p(s(bool,X455)) | ~sP7(X455,X454,'T_0') | s(bool,X453) = s(bool,X456) | s(bool,'T_0') = s(bool,X456)) )),
  inference(superposition,[],[f541,f701])).
fof(f25840,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f25839])).
fof(f25839,plain,(
  ( ! [X408,X410,X409] : (p(s(bool,X409)) | sP1(X409,'T_0') | s(bool,X408) = s(bool,X410) | s(bool,'T_0') = s(bool,X410)) )),
  inference(subsumption_resolution,[],[f23955,f747])).
fof(f23955,plain,(
  ( ! [X408,X410,X409] : (~p(s(bool,X408)) | p(s(bool,X409)) | sP1(X409,'T_0') | s(bool,X408) = s(bool,X410) | s(bool,'T_0') = s(bool,X410)) )),
  inference(superposition,[],[f462,f701])).
fof(f25837,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f25836])).
fof(f25836,plain,(
  ( ! [X405,X407,X406] : (p(s(bool,X406)) | sP1('T_0',X406) | s(bool,X405) = s(bool,X407) | s(bool,'T_0') = s(bool,X407)) )),
  inference(subsumption_resolution,[],[f23954,f747])).
fof(f23954,plain,(
  ( ! [X405,X407,X406] : (~p(s(bool,X405)) | p(s(bool,X406)) | sP1('T_0',X406) | s(bool,X405) = s(bool,X407) | s(bool,'T_0') = s(bool,X407)) )),
  inference(superposition,[],[f459,f701])).
fof(f25830,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(X1,'const.bool.COND_3'(s(bool,'T_0'),s(X1,X2),s(X1,X3))) = s(X1,'const.bool.COND_3'(s(bool,X4),s(X1,X5),s(X1,X6))) | s(X1,X2) != s(X1,X5) | s(bool,'T_0') != s(bool,X4)) )),
  inference(condensation,[],[f25829])).
fof(f25829,plain,(
  ( ! [X350,X356,X352,X354,X351,X357,X353,X355] : (s(X351,'const.bool.COND_3'(s(bool,'T_0'),s(X351,X355),s(X351,X356))) = s(X351,'const.bool.COND_3'(s(bool,X352),s(X351,X353),s(X351,X354))) | s(X351,X353) != s(X351,X355) | s(bool,'T_0') != s(bool,X352) | s(bool,X350) = s(bool,X357) | s(bool,'T_0') = s(bool,X357)) )),
  inference(subsumption_resolution,[],[f23944,f747])).
fof(f23944,plain,(
  ( ! [X350,X356,X352,X354,X351,X357,X353,X355] : (~p(s(bool,X350)) | s(X351,'const.bool.COND_3'(s(bool,'T_0'),s(X351,X355),s(X351,X356))) = s(X351,'const.bool.COND_3'(s(bool,X352),s(X351,X353),s(X351,X354))) | s(X351,X353) != s(X351,X355) | s(bool,'T_0') != s(bool,X352) | s(bool,X350) = s(bool,X357) | s(bool,'T_0') = s(bool,X357)) )),
  inference(superposition,[],[f426,f701])).
fof(f25808,plain,(
  ( ! [X294,X293] : (~p(s(bool,'const.bool.COND_3'(s(bool,X293),s(bool,X294),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X293)) ) | $spl10),
  inference(cnf_transformation,[],[f25808_D])).
fof(f25808_D,plain,(
  ( ! [X294,X293] : (~p(s(bool,'const.bool.COND_3'(s(bool,X293),s(bool,X294),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X293)) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f25824,plain,(
  ( ! [X292,X294,X293] : (~p(s(bool,'const.bool.COND_3'(s(bool,X293),s(bool,X294),s(bool,'F_0')))) | p(s(bool,X292)) | s(bool,'T_0') = s(bool,X293)) )),
  inference(subsumption_resolution,[],[f23606,f354])).
fof(f23606,plain,(
  ( ! [X292,X294,X293] : (~p(s(bool,'const.bool.COND_3'(s(bool,X293),s(bool,X294),s(bool,'F_0')))) | p(s(bool,X292)) | s(bool,'T_0') = s(bool,X292) | s(bool,'T_0') = s(bool,X293)) )),
  inference(superposition,[],[f1494,f701])).
fof(f25806,plain,(
  ( ! [X292,X294,X293] : (~p(s(bool,'const.bool.COND_3'(s(bool,X293),s(bool,X294),s(bool,'F_0')))) | p(s(bool,X292)) | s(bool,'T_0') = s(bool,X293)) )),
  inference(subsumption_resolution,[],[f23283,f354])).
fof(f23283,plain,(
  ( ! [X292,X294,X293] : (~p(s(bool,'const.bool.COND_3'(s(bool,X293),s(bool,X294),s(bool,'F_0')))) | p(s(bool,X292)) | s(bool,'T_0') = s(bool,X293) | s(bool,'T_0') = s(bool,X292)) )),
  inference(superposition,[],[f1494,f701])).
fof(f5081,plain,(
  $false | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f594,f1981])).
fof(f594,plain,(
  s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',sK12),s('type.num.num',sK12))) != s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f589,f590])).
fof(f5137,plain,(
  $false | ($spl1 | $spl2)),
  inference(trivial_inequality_removal,[],[f5105])).
fof(f5105,plain,(
  s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',sK12),s('type.num.num',sK12))) != s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | ($spl1 | $spl2)),
  inference(superposition,[],[f594,f1981])).
fof(f3140,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f334,f333,f335,f592,f343,f355,f354,f359,f358,f412,f411,f413,f414,f420,f419,f418,f417,f424,f423,f422,f421,f428,f427,f426,f429,f430,f431,f432,f433,f437,f434,f444,f443,f593,f462,f459,f458,f457,f468,f467,f477,f478,f479,f480,f481,f482,f483,f489,f488,f487,f485,f484,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f787,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f818,f819,f825,f828,f829,f830,f839,f838,f837,f836,f840,f852,f861,f862,f863,f864,f865,f866,f867,f868,f869,f870,f871,f873,f888,f889,f891,f898,f901,f902,f904,f917,f923,f924,f926,f930,f933,f939,f940,f942,f946,f954,f960,f961,f966,f971,f975,f980,f984,f991,f994,f1000,f1001,f1007,f1012,f1023,f1024,f1025,f1032,f1033,f1035,f1047,f1048,f1050,f1057,f1059,f1062,f1068,f1069,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1210,f1212,f1217,f1219,f1230,f1231,f1239,f1247,f1249,f1254,f1256,f1267,f1275,f1283,f1285,f1339,f1333,f1332,f1373,f1366,f1365,f1385,f1411,f1410,f1403,f1402,f1422,f1433,f1437,f1461,f1486,f1492,f1494,f1501,f1512,f1516,f1542,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1663,f1662,f1661,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1818,f1819,f1816,f1837,f1838,f1835,f1852,f1854,f1859,f1861,f1931,f2001,f1981,f2015,f2030,f2041,f2043,f2050,f2052,f2073,f2082,f2088,f2090,f2097,f2099,f2119,f2118,f2128,f2157,f2170,f2251,f2252,f2222,f2262,f2269,f2270,f2271,f2272,f2317,f2318,f2319,f2351,f2352,f2349,f2371,f2372,f2373,f2400,f2401,f2398,f2432,f2433,f2465,f2466,f2504,f2505,f2544,f2545,f2617,f2622,f2765,f2770,f2841,f2844,f2847,f2848,f2751,f2842,f2996,f2855,f2860,f2861,f2862,f2863,f2865,f2866,f2867,f2868,f2869,f2870,f2871,f2872,f3012,f3014,f2875,f2876,f2878,f2880,f2882,f2883,f2885,f2887,f2890,f2892,f2893,f2895,f2897,f2898,f2899,f2900,f2901,f2902,f2903,f2904,f3037,f2906,f2907,f2912,f2913,f2928,f2929,f3067,f3069,f3071,f3109,f3111,f3115,f3127,f3131,f3135,f3139])).
fof(f3139,plain,(
  ( ! [X8,X7,X9] : (s('type.num.num',X9) = s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X7),s('type.num.num',X8)))) ) | $spl4),
  inference(subsumption_resolution,[],[f3138,f335])).
fof(f3138,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,'T_0')) | s('type.num.num',X9) = s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X7),s('type.num.num',X8)))) ) | $spl4),
  inference(forward_demodulation,[],[f3137,f2842])).
fof(f3137,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X8),s('type.num.num',X7))),s('type.num.num',X9)))) | s('type.num.num',X9) = s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X7),s('type.num.num',X8)))) ) | $spl4),
  inference(subsumption_resolution,[],[f2970,f335])).
fof(f2970,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X8),s('type.num.num',X7))),s('type.num.num',X9)))) | s('type.num.num',X9) = s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X7),s('type.num.num',X8)))) ) | $spl4),
  inference(backward_demodulation,[],[f2842,f1854])).
fof(f3135,plain,(
  ( ! [X8,X7,X9] : (s('type.num.num',X9) = s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X7),s('type.num.num',X8)))) ) | $spl4),
  inference(subsumption_resolution,[],[f3134,f335])).
fof(f3134,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,'T_0')) | s('type.num.num',X9) = s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X7),s('type.num.num',X8)))) ) | $spl4),
  inference(forward_demodulation,[],[f3133,f2842])).
fof(f3133,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X7),s('type.num.num',X8))),s('type.num.num',X9)))) | s('type.num.num',X9) = s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X7),s('type.num.num',X8)))) ) | $spl4),
  inference(subsumption_resolution,[],[f2969,f335])).
fof(f2969,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X7),s('type.num.num',X8))),s('type.num.num',X9)))) | s('type.num.num',X9) = s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X7),s('type.num.num',X8)))) ) | $spl4),
  inference(backward_demodulation,[],[f2842,f1861])).
fof(f3131,plain,(
  ( ! [X2,X3,X1] : (s('type.num.num',X3) = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X2)))) ) | $spl4),
  inference(subsumption_resolution,[],[f3130,f335])).
fof(f3130,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,'T_0')) | s('type.num.num',X3) = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X2)))) ) | $spl4),
  inference(forward_demodulation,[],[f3129,f2842])).
fof(f3129,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X2),s('type.num.num',X1))),s('type.num.num',X3)))) | s('type.num.num',X3) = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X2)))) ) | $spl4),
  inference(subsumption_resolution,[],[f2968,f335])).
fof(f2968,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X2),s('type.num.num',X1))),s('type.num.num',X3)))) | s('type.num.num',X3) = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X2)))) ) | $spl4),
  inference(backward_demodulation,[],[f2842,f1852])).
fof(f3127,plain,(
  ( ! [X2,X3,X1] : (s('type.num.num',X3) = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X2)))) ) | $spl4),
  inference(subsumption_resolution,[],[f3126,f335])).
fof(f3126,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,'T_0')) | s('type.num.num',X3) = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X2)))) ) | $spl4),
  inference(forward_demodulation,[],[f3125,f2842])).
fof(f3125,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num',X3)))) | s('type.num.num',X3) = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X2)))) ) | $spl4),
  inference(subsumption_resolution,[],[f2967,f335])).
fof(f2967,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num',X3)))) | s('type.num.num',X3) = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X2)))) ) | $spl4),
  inference(backward_demodulation,[],[f2842,f1859])).
fof(f3115,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) = s('type.num.num',X1)) ) | $spl4),
  inference(subsumption_resolution,[],[f3114,f335])).
fof(f3114,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s('type.num.num',X0) = s('type.num.num',X1)) ) | $spl4),
  inference(forward_demodulation,[],[f3113,f2842])).
fof(f3113,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | s('type.num.num',X0) = s('type.num.num',X1)) ) | $spl4),
  inference(subsumption_resolution,[],[f2961,f335])).
fof(f2961,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | s('type.num.num',X0) = s('type.num.num',X1)) ) | $spl4),
  inference(backward_demodulation,[],[f2842,f433])).
fof(f3111,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(subsumption_resolution,[],[f2994,f335])).
fof(f2994,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'T_0'))) ) | $spl4),
  inference(condensation,[],[f2954])).
fof(f2954,plain,(
  ( ! [X2,X0] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | p(s(bool,X2))) ) | $spl4),
  inference(backward_demodulation,[],[f2842,f485])).
fof(f3109,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(subsumption_resolution,[],[f2995,f335])).
fof(f2995,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,X0))) ) | $spl4),
  inference(condensation,[],[f2953])).
fof(f2953,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X1))) ) | $spl4),
  inference(backward_demodulation,[],[f2842,f484])).
fof(f3071,plain,(
  ( ! [X0,X1] : (s('type.num.num',X1) = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X0),s('type.num.num',X1)))) ) | $spl4),
  inference(forward_demodulation,[],[f2933,f413])).
fof(f2933,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',X1),s('type.num.num',X0)))) ) | $spl4),
  inference(backward_demodulation,[],[f2842,f432])).
fof(f3069,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X0),s('type.num.num',X1)))) ) | $spl4),
  inference(forward_demodulation,[],[f2932,f413])).
fof(f2932,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',X0),s('type.num.num',X1)))) ) | $spl4),
  inference(backward_demodulation,[],[f2842,f431])).
fof(f3067,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (s(X9,X10) != s(X9,X11) | p(s(bool,X7)) | s(X9,'const.bool.COND_3'(s(bool,X7),s(X9,X13),s(X9,X11))) = s(X9,'const.bool.COND_3'(s(bool,X8),s(X9,X12),s(X9,X10)))) ) | $spl4),
  inference(subsumption_resolution,[],[f2931,f2844])).
fof(f2931,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (s(bool,'T_0') != s(bool,X8) | s(X9,X10) != s(X9,X11) | p(s(bool,X7)) | s(X9,'const.bool.COND_3'(s(bool,X7),s(X9,X13),s(X9,X11))) = s(X9,'const.bool.COND_3'(s(bool,X8),s(X9,X12),s(X9,X10)))) ) | $spl4),
  inference(backward_demodulation,[],[f2842,f2848])).
fof(f2929,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X2),s('type.num.num',X1))) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',X0),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X2)))))) ) | $spl4),
  inference(backward_demodulation,[],[f2842,f2770])).
fof(f2928,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X2),s('type.num.num',X1))) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',X0),s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X1),s('type.num.num',X2)))))) ) | $spl4),
  inference(backward_demodulation,[],[f2842,f2765])).
fof(f2913,plain,(
  ( ! [X0,X1] : (~sP3('const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,'T_0')),'const.bool.COND_3'(s(bool,X1),s(bool,'T_0'),s(bool,'T_0')),'F_0')) ) | $spl4),
  inference(backward_demodulation,[],[f2842,f2269])).
fof(f2912,plain,(
  ( ! [X0,X1] : (~sP3('const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,'T_0')),'F_0')) ) | $spl4),
  inference(backward_demodulation,[],[f2842,f2251])).
fof(f2907,plain,(
  ( ! [X2,X0,X1] : (sP5('const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),X1,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X2),s(bool,X1)))) ) | $spl4),
  inference(backward_demodulation,[],[f2842,f1663])).
fof(f2906,plain,(
  ( ! [X2,X0,X1] : (sP5('const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,'T_0')),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2)),X2)) ) | $spl4),
  inference(backward_demodulation,[],[f2842,f1661])).
fof(f3037,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(subsumption_resolution,[],[f3036,f335])).
fof(f3036,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,X0))) ) | $spl4),
  inference(forward_demodulation,[],[f2905,f2842])).
fof(f2905,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,'T_0')))) | p(s(bool,X0))) ) | $spl4),
  inference(backward_demodulation,[],[f2842,f1494])).
fof(f2904,plain,(
  ( ! [X0] : (~sP2('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,'T_0')),'F_0')) ) | $spl4),
  inference(backward_demodulation,[],[f2842,f1411])).
fof(f2903,plain,(
  ( ! [X0] : (~sP2('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)),'F_0')) ) | $spl4),
  inference(backward_demodulation,[],[f2842,f1410])).
fof(f2902,plain,(
  ( ! [X0] : (~sP3('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,'T_0')),'F_0')) ) | $spl4),
  inference(backward_demodulation,[],[f2842,f1403])).
fof(f2901,plain,(
  ( ! [X0] : (~sP2('const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,'T_0')),'F_0','F_0')) ) | $spl4),
  inference(backward_demodulation,[],[f2842,f1366])).
fof(f2900,plain,(
  ( ! [X0] : (~sP2('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)),'F_0','F_0')) ) | $spl4),
  inference(backward_demodulation,[],[f2842,f1365])).
fof(f2899,plain,(
  ( ! [X0] : (~sP3('const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,'T_0')),'T_0','F_0')) ) | $spl4),
  inference(backward_demodulation,[],[f2842,f1333])).
fof(f2898,plain,(
  ( ! [X0,X1] : (sP0('const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,'T_0')),X1)) ) | $spl4),
  inference(backward_demodulation,[],[f2842,f1267])).
fof(f2897,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,'T_0')))) ) | $spl4),
  inference(backward_demodulation,[],[f2842,f1231])).
fof(f2895,plain,(
  ( ! [X8,X7,X9] : (sP4('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X7),s(bool,X8)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X9),s(bool,X8)))) ) | $spl4),
  inference(backward_demodulation,[],[f2842,f1059])).
fof(f2893,plain,(
  ( ! [X12,X10,X13,X11] : (~sP7(X13,X12,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X10),s(bool,X11))) | ~p(s(bool,X11)) | ~p(s(bool,X12)) | p(s(bool,X13))) ) | $spl4),
  inference(backward_demodulation,[],[f2842,f1048])).
fof(f2892,plain,(
  ( ! [X8,X7,X9] : (sP4('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X8),s(bool,X9)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X7),s(bool,X8)))) ) | $spl4),
  inference(backward_demodulation,[],[f2842,f1035])).
fof(f2890,plain,(
  ( ! [X8,X7,X9] : (sP4('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X7),s(bool,X8)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X8),s(bool,X9)))) ) | $spl4),
  inference(backward_demodulation,[],[f2842,f1025])).
fof(f2887,plain,(
  ( ! [X0,X1] : (sP4('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)),X1)) ) | $spl4),
  inference(backward_demodulation,[],[f2842,f991])).
fof(f2885,plain,(
  ( ! [X8,X7,X9] : (sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X9),s(bool,X8)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X7),s(bool,X8)))) ) | $spl4),
  inference(backward_demodulation,[],[f2842,f942])).
fof(f2883,plain,(
  ( ! [X0,X1] : (sP4('T_0',X0,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X0)))) ) | $spl4),
  inference(backward_demodulation,[],[f2842,f930])).
fof(f2882,plain,(
  ( ! [X8,X7,X9] : (~sP5('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X9),s(bool,X8)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X7),s(bool,X8)))) ) | $spl4),
  inference(backward_demodulation,[],[f2842,f926])).
fof(f2880,plain,(
  ( ! [X8,X7,X9] : (sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X7),s(bool,X8)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X8),s(bool,X9)))) ) | $spl4),
  inference(backward_demodulation,[],[f2842,f904])).
fof(f2878,plain,(
  ( ! [X8,X7,X9] : (~sP5('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X7),s(bool,X8)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X8),s(bool,X9)))) ) | $spl4),
  inference(backward_demodulation,[],[f2842,f891])).
fof(f2876,plain,(
  ( ! [X61,X62,X63] : (sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X62),s(bool,X63)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X61),s(bool,X62)))) ) | $spl4),
  inference(backward_demodulation,[],[f2842,f873])).
fof(f2875,plain,(
  ( ! [X57,X56,X55] : (~sP5('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X56),s(bool,X57)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X55),s(bool,X56)))) ) | $spl4),
  inference(backward_demodulation,[],[f2842,f871])).
fof(f3014,plain,(
  ( ! [X54,X52,X53,X51] : (sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X51),s(bool,X52)),X53,X54) | ~sP8('const.bool.COND_3'(s(bool,'T_0'),s(bool,X51),s(bool,X52)),X53,X54) | ~p(s(bool,X52))) ) | $spl4),
  inference(forward_demodulation,[],[f2874,f2842])).
fof(f2874,plain,(
  ( ! [X54,X52,X53,X51] : (~sP8('const.bool.COND_3'(s(bool,'T_0'),s(bool,X51),s(bool,X52)),X53,X54) | ~p(s(bool,X52)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X51),s(bool,X52)),X53,X54)) ) | $spl4),
  inference(backward_demodulation,[],[f2842,f870])).
fof(f3012,plain,(
  ( ! [X47,X50,X48,X49] : (sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X47),s(bool,X48)),X49,X50) | ~sP10('const.bool.COND_3'(s(bool,'T_0'),s(bool,X47),s(bool,X48)),X49,X50) | ~p(s(bool,X48))) ) | $spl4),
  inference(forward_demodulation,[],[f2873,f2842])).
fof(f2873,plain,(
  ( ! [X47,X50,X48,X49] : (~sP10('const.bool.COND_3'(s(bool,'T_0'),s(bool,X47),s(bool,X48)),X49,X50) | ~p(s(bool,X48)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X47),s(bool,X48)),X49,X50)) ) | $spl4),
  inference(backward_demodulation,[],[f2842,f869])).
fof(f2872,plain,(
  ( ! [X45,X43,X46,X44] : (sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X43),s(bool,X44)),X45,X46) | ~p(s(bool,X44)) | ~p(s(bool,X45))) ) | $spl4),
  inference(backward_demodulation,[],[f2842,f868])).
fof(f2871,plain,(
  ( ! [X39,X41,X42,X40] : (sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X39),s(bool,X40)),X42,X41) | ~p(s(bool,X40)) | p(s(bool,X41))) ) | $spl4),
  inference(backward_demodulation,[],[f2842,f867])).
fof(f2870,plain,(
  ( ! [X37,X35,X38,X36] : (sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X35),s(bool,X36)),X37,X38) | ~p(s(bool,X36)) | ~p(s(bool,X37))) ) | $spl4),
  inference(backward_demodulation,[],[f2842,f866])).
fof(f2869,plain,(
  ( ! [X33,X31,X34,X32] : (sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X31),s(bool,X32)),X34,X33) | ~p(s(bool,X32)) | ~p(s(bool,X33))) ) | $spl4),
  inference(backward_demodulation,[],[f2842,f865])).
fof(f2868,plain,(
  ( ! [X30,X28,X29] : (sP1(X30,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X28),s(bool,X29))) | ~p(s(bool,X29)) | p(s(bool,X30))) ) | $spl4),
  inference(backward_demodulation,[],[f2842,f864])).
fof(f2867,plain,(
  ( ! [X26,X27,X25] : (sP1('const.bool.COND_3'(s(bool,'T_0'),s(bool,X25),s(bool,X26)),X27) | ~p(s(bool,X26)) | p(s(bool,X27))) ) | $spl4),
  inference(backward_demodulation,[],[f2842,f863])).
fof(f2866,plain,(
  ( ! [X24,X23,X22] : (sP0('const.bool.COND_3'(s(bool,'T_0'),s(bool,X22),s(bool,X23)),X24) | ~p(s(bool,X23))) ) | $spl4),
  inference(backward_demodulation,[],[f2842,f862])).
fof(f2865,plain,(
  ( ! [X21,X19,X20] : (sP0(X21,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X19),s(bool,X20))) | ~p(s(bool,X20))) ) | $spl4),
  inference(backward_demodulation,[],[f2842,f861])).
fof(f2863,plain,(
  ( ! [X0,X1] : (sP5('T_0',X0,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X0)))) ) | $spl4),
  inference(backward_demodulation,[],[f2842,f839])).
fof(f2862,plain,(
  ( ! [X0,X1] : (~sP5('F_0',X0,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X0)))) ) | $spl4),
  inference(backward_demodulation,[],[f2842,f838])).
fof(f2861,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)),X1)) ) | $spl4),
  inference(backward_demodulation,[],[f2842,f837])).
fof(f2860,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)),X1)) ) | $spl4),
  inference(backward_demodulation,[],[f2842,f836])).
fof(f2855,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) ) | $spl4),
  inference(backward_demodulation,[],[f2842,f414])).
fof(f2996,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl4),
  inference(subsumption_resolution,[],[f2853,f2842])).
fof(f2853,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl4),
  inference(backward_demodulation,[],[f2842,f358])).
fof(f3136,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f334,f333,f335,f592,f343,f355,f354,f359,f358,f412,f411,f413,f414,f420,f419,f418,f417,f424,f423,f422,f421,f428,f427,f426,f429,f430,f431,f432,f433,f437,f434,f444,f443,f593,f462,f459,f458,f457,f468,f467,f477,f478,f479,f480,f481,f482,f483,f489,f488,f487,f485,f484,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f787,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f818,f819,f825,f828,f829,f830,f839,f838,f837,f836,f840,f852,f861,f862,f863,f864,f865,f866,f867,f868,f869,f870,f871,f873,f888,f889,f891,f898,f901,f902,f904,f917,f923,f924,f926,f930,f933,f939,f940,f942,f946,f954,f960,f961,f966,f971,f975,f980,f984,f991,f994,f1000,f1001,f1007,f1012,f1023,f1024,f1025,f1032,f1033,f1035,f1047,f1048,f1050,f1057,f1059,f1062,f1068,f1069,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1210,f1212,f1217,f1219,f1230,f1231,f1239,f1247,f1249,f1254,f1256,f1267,f1275,f1283,f1285,f1339,f1333,f1332,f1373,f1366,f1365,f1385,f1411,f1410,f1403,f1402,f1422,f1433,f1437,f1461,f1486,f1492,f1494,f1501,f1512,f1516,f1542,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1663,f1662,f1661,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1818,f1819,f1816,f1837,f1838,f1835,f1852,f1854,f1859,f1861,f1931,f2001,f1981,f2015,f2030,f2041,f2043,f2050,f2052,f2073,f2082,f2088,f2090,f2097,f2099,f2119,f2118,f2128,f2157,f2170,f2251,f2252,f2222,f2262,f2269,f2270,f2271,f2272,f2317,f2318,f2319,f2351,f2352,f2349,f2371,f2372,f2373,f2400,f2401,f2398,f2432,f2433,f2465,f2466,f2504,f2505,f2544,f2545,f2617,f2622,f2765,f2770,f2841,f2844,f2847,f2848,f2751,f2842,f2996,f2855,f2860,f2861,f2862,f2863,f2865,f2866,f2867,f2868,f2869,f2870,f2871,f2872,f3012,f3014,f2875,f2876,f2878,f2880,f2882,f2883,f2885,f2887,f2890,f2892,f2893,f2895,f2897,f2898,f2899,f2900,f2901,f2902,f2903,f2904,f3037,f2906,f2907,f2912,f2913,f2928,f2929,f3067,f3069,f3071,f3109,f3111,f3115,f3127,f3131,f3135])).
fof(f3132,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f334,f333,f335,f592,f343,f355,f354,f359,f358,f412,f411,f413,f414,f420,f419,f418,f417,f424,f423,f422,f421,f428,f427,f426,f429,f430,f431,f432,f433,f437,f434,f444,f443,f593,f462,f459,f458,f457,f468,f467,f477,f478,f479,f480,f481,f482,f483,f489,f488,f487,f485,f484,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f787,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f818,f819,f825,f828,f829,f830,f839,f838,f837,f836,f840,f852,f861,f862,f863,f864,f865,f866,f867,f868,f869,f870,f871,f873,f888,f889,f891,f898,f901,f902,f904,f917,f923,f924,f926,f930,f933,f939,f940,f942,f946,f954,f960,f961,f966,f971,f975,f980,f984,f991,f994,f1000,f1001,f1007,f1012,f1023,f1024,f1025,f1032,f1033,f1035,f1047,f1048,f1050,f1057,f1059,f1062,f1068,f1069,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1210,f1212,f1217,f1219,f1230,f1231,f1239,f1247,f1249,f1254,f1256,f1267,f1275,f1283,f1285,f1339,f1333,f1332,f1373,f1366,f1365,f1385,f1411,f1410,f1403,f1402,f1422,f1433,f1437,f1461,f1486,f1492,f1494,f1501,f1512,f1516,f1542,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1663,f1662,f1661,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1818,f1819,f1816,f1837,f1838,f1835,f1852,f1854,f1859,f1861,f1931,f2001,f1981,f2015,f2030,f2041,f2043,f2050,f2052,f2073,f2082,f2088,f2090,f2097,f2099,f2119,f2118,f2128,f2157,f2170,f2251,f2252,f2222,f2262,f2269,f2270,f2271,f2272,f2317,f2318,f2319,f2351,f2352,f2349,f2371,f2372,f2373,f2400,f2401,f2398,f2432,f2433,f2465,f2466,f2504,f2505,f2544,f2545,f2617,f2622,f2765,f2770,f2841,f2844,f2847,f2848,f2751,f2842,f2996,f2855,f2860,f2861,f2862,f2863,f2865,f2866,f2867,f2868,f2869,f2870,f2871,f2872,f3012,f3014,f2875,f2876,f2878,f2880,f2882,f2883,f2885,f2887,f2890,f2892,f2893,f2895,f2897,f2898,f2899,f2900,f2901,f2902,f2903,f2904,f3037,f2906,f2907,f2912,f2913,f2928,f2929,f3067,f3069,f3071,f3109,f3111,f3115,f3127,f3131])).
fof(f3128,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f334,f333,f335,f592,f343,f355,f354,f359,f358,f412,f411,f413,f414,f420,f419,f418,f417,f424,f423,f422,f421,f428,f427,f426,f429,f430,f431,f432,f433,f437,f434,f444,f443,f593,f462,f459,f458,f457,f468,f467,f477,f478,f479,f480,f481,f482,f483,f489,f488,f487,f485,f484,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f787,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f818,f819,f825,f828,f829,f830,f839,f838,f837,f836,f840,f852,f861,f862,f863,f864,f865,f866,f867,f868,f869,f870,f871,f873,f888,f889,f891,f898,f901,f902,f904,f917,f923,f924,f926,f930,f933,f939,f940,f942,f946,f954,f960,f961,f966,f971,f975,f980,f984,f991,f994,f1000,f1001,f1007,f1012,f1023,f1024,f1025,f1032,f1033,f1035,f1047,f1048,f1050,f1057,f1059,f1062,f1068,f1069,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1210,f1212,f1217,f1219,f1230,f1231,f1239,f1247,f1249,f1254,f1256,f1267,f1275,f1283,f1285,f1339,f1333,f1332,f1373,f1366,f1365,f1385,f1411,f1410,f1403,f1402,f1422,f1433,f1437,f1461,f1486,f1492,f1494,f1501,f1512,f1516,f1542,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1663,f1662,f1661,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1818,f1819,f1816,f1837,f1838,f1835,f1852,f1854,f1859,f1861,f1931,f2001,f1981,f2015,f2030,f2041,f2043,f2050,f2052,f2073,f2082,f2088,f2090,f2097,f2099,f2119,f2118,f2128,f2157,f2170,f2251,f2252,f2222,f2262,f2269,f2270,f2271,f2272,f2317,f2318,f2319,f2351,f2352,f2349,f2371,f2372,f2373,f2400,f2401,f2398,f2432,f2433,f2465,f2466,f2504,f2505,f2544,f2545,f2617,f2622,f2765,f2770,f2841,f2844,f2847,f2848,f2751,f2842,f2996,f2855,f2860,f2861,f2862,f2863,f2865,f2866,f2867,f2868,f2869,f2870,f2871,f2872,f3012,f3014,f2875,f2876,f2878,f2880,f2882,f2883,f2885,f2887,f2890,f2892,f2893,f2895,f2897,f2898,f2899,f2900,f2901,f2902,f2903,f2904,f3037,f2906,f2907,f2912,f2913,f2928,f2929,f3067,f3069,f3071,f3109,f3111,f3115,f3127])).
fof(f3116,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f334,f333,f335,f592,f343,f355,f354,f359,f358,f412,f411,f413,f414,f420,f419,f418,f417,f424,f423,f422,f421,f428,f427,f426,f429,f430,f431,f432,f433,f437,f434,f444,f443,f593,f462,f459,f458,f457,f468,f467,f477,f478,f479,f480,f481,f482,f483,f489,f488,f487,f485,f484,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f787,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f818,f819,f825,f828,f829,f830,f839,f838,f837,f836,f840,f852,f861,f862,f863,f864,f865,f866,f867,f868,f869,f870,f871,f873,f888,f889,f891,f898,f901,f902,f904,f917,f923,f924,f926,f930,f933,f939,f940,f942,f946,f954,f960,f961,f966,f971,f975,f980,f984,f991,f994,f1000,f1001,f1007,f1012,f1023,f1024,f1025,f1032,f1033,f1035,f1047,f1048,f1050,f1057,f1059,f1062,f1068,f1069,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1210,f1212,f1217,f1219,f1230,f1231,f1239,f1247,f1249,f1254,f1256,f1267,f1275,f1283,f1285,f1339,f1333,f1332,f1373,f1366,f1365,f1385,f1411,f1410,f1403,f1402,f1422,f1433,f1437,f1461,f1486,f1492,f1494,f1501,f1512,f1516,f1542,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1663,f1662,f1661,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1818,f1819,f1816,f1837,f1838,f1835,f1852,f1854,f1859,f1861,f1931,f2001,f1981,f2015,f2030,f2041,f2043,f2050,f2052,f2073,f2082,f2088,f2090,f2097,f2099,f2119,f2118,f2128,f2157,f2170,f2251,f2252,f2222,f2262,f2269,f2270,f2271,f2272,f2317,f2318,f2319,f2351,f2352,f2349,f2371,f2372,f2373,f2400,f2401,f2398,f2432,f2433,f2465,f2466,f2504,f2505,f2544,f2545,f2617,f2622,f2765,f2770,f2841,f2844,f2847,f2848,f2751,f2842,f2996,f2855,f2860,f2861,f2862,f2863,f2865,f2866,f2867,f2868,f2869,f2870,f2871,f2872,f3012,f3014,f2875,f2876,f2878,f2880,f2882,f2883,f2885,f2887,f2890,f2892,f2893,f2895,f2897,f2898,f2899,f2900,f2901,f2902,f2903,f2904,f3037,f2906,f2907,f2912,f2913,f2928,f2929,f3067,f3069,f3071,f3109,f3111,f3115])).
fof(f3112,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f334,f333,f335,f592,f343,f355,f354,f359,f358,f412,f411,f413,f414,f420,f419,f418,f417,f424,f423,f422,f421,f428,f427,f426,f429,f430,f431,f432,f433,f437,f434,f444,f443,f593,f462,f459,f458,f457,f468,f467,f477,f478,f479,f480,f481,f482,f483,f489,f488,f487,f485,f484,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f787,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f818,f819,f825,f828,f829,f830,f839,f838,f837,f836,f840,f852,f861,f862,f863,f864,f865,f866,f867,f868,f869,f870,f871,f873,f888,f889,f891,f898,f901,f902,f904,f917,f923,f924,f926,f930,f933,f939,f940,f942,f946,f954,f960,f961,f966,f971,f975,f980,f984,f991,f994,f1000,f1001,f1007,f1012,f1023,f1024,f1025,f1032,f1033,f1035,f1047,f1048,f1050,f1057,f1059,f1062,f1068,f1069,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1210,f1212,f1217,f1219,f1230,f1231,f1239,f1247,f1249,f1254,f1256,f1267,f1275,f1283,f1285,f1339,f1333,f1332,f1373,f1366,f1365,f1385,f1411,f1410,f1403,f1402,f1422,f1433,f1437,f1461,f1486,f1492,f1494,f1501,f1512,f1516,f1542,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1663,f1662,f1661,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1818,f1819,f1816,f1837,f1838,f1835,f1852,f1854,f1859,f1861,f1931,f2001,f1981,f2015,f2030,f2041,f2043,f2050,f2052,f2073,f2082,f2088,f2090,f2097,f2099,f2119,f2118,f2128,f2157,f2170,f2251,f2252,f2222,f2262,f2269,f2270,f2271,f2272,f2317,f2318,f2319,f2351,f2352,f2349,f2371,f2372,f2373,f2400,f2401,f2398,f2432,f2433,f2465,f2466,f2504,f2505,f2544,f2545,f2617,f2622,f2765,f2770,f2841,f2844,f2847,f2848,f2751,f2842,f2996,f2855,f2860,f2861,f2862,f2863,f2865,f2866,f2867,f2868,f2869,f2870,f2871,f2872,f3012,f3014,f2875,f2876,f2878,f2880,f2882,f2883,f2885,f2887,f2890,f2892,f2893,f2895,f2897,f2898,f2899,f2900,f2901,f2902,f2903,f2904,f3037,f2906,f2907,f2912,f2913,f2928,f2929,f3067,f3069,f3071,f3109,f3111])).
fof(f3110,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f334,f333,f335,f592,f343,f355,f354,f359,f358,f412,f411,f413,f414,f420,f419,f418,f417,f424,f423,f422,f421,f428,f427,f426,f429,f430,f431,f432,f433,f437,f434,f444,f443,f593,f462,f459,f458,f457,f468,f467,f477,f478,f479,f480,f481,f482,f483,f489,f488,f487,f485,f484,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f787,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f818,f819,f825,f828,f829,f830,f839,f838,f837,f836,f840,f852,f861,f862,f863,f864,f865,f866,f867,f868,f869,f870,f871,f873,f888,f889,f891,f898,f901,f902,f904,f917,f923,f924,f926,f930,f933,f939,f940,f942,f946,f954,f960,f961,f966,f971,f975,f980,f984,f991,f994,f1000,f1001,f1007,f1012,f1023,f1024,f1025,f1032,f1033,f1035,f1047,f1048,f1050,f1057,f1059,f1062,f1068,f1069,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1210,f1212,f1217,f1219,f1230,f1231,f1239,f1247,f1249,f1254,f1256,f1267,f1275,f1283,f1285,f1339,f1333,f1332,f1373,f1366,f1365,f1385,f1411,f1410,f1403,f1402,f1422,f1433,f1437,f1461,f1486,f1492,f1494,f1501,f1512,f1516,f1542,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1663,f1662,f1661,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1818,f1819,f1816,f1837,f1838,f1835,f1852,f1854,f1859,f1861,f1931,f2001,f1981,f2015,f2030,f2041,f2043,f2050,f2052,f2073,f2082,f2088,f2090,f2097,f2099,f2119,f2118,f2128,f2157,f2170,f2251,f2252,f2222,f2262,f2269,f2270,f2271,f2272,f2317,f2318,f2319,f2351,f2352,f2349,f2371,f2372,f2373,f2400,f2401,f2398,f2432,f2433,f2465,f2466,f2504,f2505,f2544,f2545,f2617,f2622,f2765,f2770,f2841,f2844,f2847,f2848,f2751,f2842,f2996,f2855,f2860,f2861,f2862,f2863,f2865,f2866,f2867,f2868,f2869,f2870,f2871,f2872,f3012,f3014,f2875,f2876,f2878,f2880,f2882,f2883,f2885,f2887,f2890,f2892,f2893,f2895,f2897,f2898,f2899,f2900,f2901,f2902,f2903,f2904,f3037,f2906,f2907,f2912,f2913,f2928,f2929,f3067,f3069,f3071,f3109])).
fof(f3108,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f3107,f335])).
fof(f3107,plain,(
  ~p(s(bool,'T_0')) | $spl4),
  inference(forward_demodulation,[],[f3106,f2842])).
fof(f3106,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X9),s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X7),s('type.num.num',X8))))))) ) | $spl4),
  inference(subsumption_resolution,[],[f2951,f335])).
fof(f2951,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X9),s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X7),s('type.num.num',X8))))))) ) | $spl4),
  inference(backward_demodulation,[],[f2842,f1249])).
fof(f3105,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f3104,f335])).
fof(f3104,plain,(
  ~p(s(bool,'T_0')) | $spl4),
  inference(forward_demodulation,[],[f3103,f2842])).
fof(f3103,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X2))))))) ) | $spl4),
  inference(subsumption_resolution,[],[f2949,f335])).
fof(f2949,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X2))))))) ) | $spl4),
  inference(backward_demodulation,[],[f2842,f1247])).
fof(f3102,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f3101,f335])).
fof(f3101,plain,(
  ~p(s(bool,'T_0')) | $spl4),
  inference(forward_demodulation,[],[f3100,f2842])).
fof(f3100,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X7),s('type.num.num',X8))),s('type.num.num',X9))))) ) | $spl4),
  inference(subsumption_resolution,[],[f2947,f335])).
fof(f2947,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X7),s('type.num.num',X8))),s('type.num.num',X9))))) ) | $spl4),
  inference(backward_demodulation,[],[f2842,f1256])).
fof(f3099,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f3098,f335])).
fof(f3098,plain,(
  ~p(s(bool,'T_0')) | $spl4),
  inference(forward_demodulation,[],[f3097,f2842])).
fof(f3097,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X8),s('type.num.num',X7))),s('type.num.num',X9))))) ) | $spl4),
  inference(subsumption_resolution,[],[f2946,f335])).
fof(f2946,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X8),s('type.num.num',X7))),s('type.num.num',X9))))) ) | $spl4),
  inference(backward_demodulation,[],[f2842,f1285])).
fof(f3096,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f3095,f335])).
fof(f3095,plain,(
  ~p(s(bool,'T_0')) | $spl4),
  inference(forward_demodulation,[],[f3094,f2842])).
fof(f3094,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num',X3))))) ) | $spl4),
  inference(subsumption_resolution,[],[f2944,f335])).
fof(f2944,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num',X3))))) ) | $spl4),
  inference(backward_demodulation,[],[f2842,f1254])).
fof(f3093,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f3092,f335])).
fof(f3092,plain,(
  ~p(s(bool,'T_0')) | $spl4),
  inference(forward_demodulation,[],[f3091,f2842])).
fof(f3091,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X2),s('type.num.num',X1))),s('type.num.num',X3))))) ) | $spl4),
  inference(subsumption_resolution,[],[f2943,f335])).
fof(f2943,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X2),s('type.num.num',X1))),s('type.num.num',X3))))) ) | $spl4),
  inference(backward_demodulation,[],[f2842,f1283])).
fof(f3078,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f3077,f335])).
fof(f3077,plain,(
  ~p(s(bool,'T_0')) | $spl4),
  inference(forward_demodulation,[],[f3076,f2842])).
fof(f3076,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))))) ) | $spl4),
  inference(subsumption_resolution,[],[f2936,f335])).
fof(f2936,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))))) ) | $spl4),
  inference(backward_demodulation,[],[f2842,f477])).
fof(f3075,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f3074,f335])).
fof(f3074,plain,(
  ~p(s(bool,'T_0')) | $spl4),
  inference(forward_demodulation,[],[f3073,f2842])).
fof(f3073,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) ) | $spl4),
  inference(subsumption_resolution,[],[f2935,f335])).
fof(f2935,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) ) | $spl4),
  inference(backward_demodulation,[],[f2842,f468])).
fof(f3072,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f334,f333,f335,f592,f343,f355,f354,f359,f358,f412,f411,f413,f414,f420,f419,f418,f417,f424,f423,f422,f421,f428,f427,f426,f429,f430,f431,f432,f433,f437,f434,f444,f443,f593,f462,f459,f458,f457,f468,f467,f477,f478,f479,f480,f481,f482,f483,f489,f488,f487,f485,f484,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f787,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f818,f819,f825,f828,f829,f830,f839,f838,f837,f836,f840,f852,f861,f862,f863,f864,f865,f866,f867,f868,f869,f870,f871,f873,f888,f889,f891,f898,f901,f902,f904,f917,f923,f924,f926,f930,f933,f939,f940,f942,f946,f954,f960,f961,f966,f971,f975,f980,f984,f991,f994,f1000,f1001,f1007,f1012,f1023,f1024,f1025,f1032,f1033,f1035,f1047,f1048,f1050,f1057,f1059,f1062,f1068,f1069,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1210,f1212,f1217,f1219,f1230,f1231,f1239,f1247,f1249,f1254,f1256,f1267,f1275,f1283,f1285,f1339,f1333,f1332,f1373,f1366,f1365,f1385,f1411,f1410,f1403,f1402,f1422,f1433,f1437,f1461,f1486,f1492,f1494,f1501,f1512,f1516,f1542,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1663,f1662,f1661,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1818,f1819,f1816,f1837,f1838,f1835,f1852,f1854,f1859,f1861,f1931,f2001,f1981,f2015,f2030,f2041,f2043,f2050,f2052,f2073,f2082,f2088,f2090,f2097,f2099,f2119,f2118,f2128,f2157,f2170,f2251,f2252,f2222,f2262,f2269,f2270,f2271,f2272,f2317,f2318,f2319,f2351,f2352,f2349,f2371,f2372,f2373,f2400,f2401,f2398,f2432,f2433,f2465,f2466,f2504,f2505,f2544,f2545,f2617,f2622,f2765,f2770,f2841,f2844,f2847,f2848,f2751,f2842,f2996,f2855,f2860,f2861,f2862,f2863,f2865,f2866,f2867,f2868,f2869,f2870,f2871,f2872,f3012,f3014,f2875,f2876,f2878,f2880,f2882,f2883,f2885,f2887,f2890,f2892,f2893,f2895,f2897,f2898,f2899,f2900,f2901,f2902,f2903,f2904,f3037,f2906,f2907,f2912,f2913,f2928,f2929,f3067,f3069,f3071])).
fof(f3070,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f334,f333,f335,f592,f343,f355,f354,f359,f358,f412,f411,f413,f414,f420,f419,f418,f417,f424,f423,f422,f421,f428,f427,f426,f429,f430,f431,f432,f433,f437,f434,f444,f443,f593,f462,f459,f458,f457,f468,f467,f477,f478,f479,f480,f481,f482,f483,f489,f488,f487,f485,f484,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f787,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f818,f819,f825,f828,f829,f830,f839,f838,f837,f836,f840,f852,f861,f862,f863,f864,f865,f866,f867,f868,f869,f870,f871,f873,f888,f889,f891,f898,f901,f902,f904,f917,f923,f924,f926,f930,f933,f939,f940,f942,f946,f954,f960,f961,f966,f971,f975,f980,f984,f991,f994,f1000,f1001,f1007,f1012,f1023,f1024,f1025,f1032,f1033,f1035,f1047,f1048,f1050,f1057,f1059,f1062,f1068,f1069,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1210,f1212,f1217,f1219,f1230,f1231,f1239,f1247,f1249,f1254,f1256,f1267,f1275,f1283,f1285,f1339,f1333,f1332,f1373,f1366,f1365,f1385,f1411,f1410,f1403,f1402,f1422,f1433,f1437,f1461,f1486,f1492,f1494,f1501,f1512,f1516,f1542,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1663,f1662,f1661,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1818,f1819,f1816,f1837,f1838,f1835,f1852,f1854,f1859,f1861,f1931,f2001,f1981,f2015,f2030,f2041,f2043,f2050,f2052,f2073,f2082,f2088,f2090,f2097,f2099,f2119,f2118,f2128,f2157,f2170,f2251,f2252,f2222,f2262,f2269,f2270,f2271,f2272,f2317,f2318,f2319,f2351,f2352,f2349,f2371,f2372,f2373,f2400,f2401,f2398,f2432,f2433,f2465,f2466,f2504,f2505,f2544,f2545,f2617,f2622,f2765,f2770,f2841,f2844,f2847,f2848,f2751,f2842,f2996,f2855,f2860,f2861,f2862,f2863,f2865,f2866,f2867,f2868,f2869,f2870,f2871,f2872,f3012,f3014,f2875,f2876,f2878,f2880,f2882,f2883,f2885,f2887,f2890,f2892,f2893,f2895,f2897,f2898,f2899,f2900,f2901,f2902,f2903,f2904,f3037,f2906,f2907,f2912,f2913,f2928,f2929,f3067,f3069])).
fof(f3068,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f334,f333,f335,f592,f343,f355,f354,f359,f358,f412,f411,f413,f414,f420,f419,f418,f417,f424,f423,f422,f421,f428,f427,f426,f429,f430,f431,f432,f433,f437,f434,f444,f443,f593,f462,f459,f458,f457,f468,f467,f477,f478,f479,f480,f481,f482,f483,f489,f488,f487,f485,f484,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f787,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f818,f819,f825,f828,f829,f830,f839,f838,f837,f836,f840,f852,f861,f862,f863,f864,f865,f866,f867,f868,f869,f870,f871,f873,f888,f889,f891,f898,f901,f902,f904,f917,f923,f924,f926,f930,f933,f939,f940,f942,f946,f954,f960,f961,f966,f971,f975,f980,f984,f991,f994,f1000,f1001,f1007,f1012,f1023,f1024,f1025,f1032,f1033,f1035,f1047,f1048,f1050,f1057,f1059,f1062,f1068,f1069,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1210,f1212,f1217,f1219,f1230,f1231,f1239,f1247,f1249,f1254,f1256,f1267,f1275,f1283,f1285,f1339,f1333,f1332,f1373,f1366,f1365,f1385,f1411,f1410,f1403,f1402,f1422,f1433,f1437,f1461,f1486,f1492,f1494,f1501,f1512,f1516,f1542,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1663,f1662,f1661,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1818,f1819,f1816,f1837,f1838,f1835,f1852,f1854,f1859,f1861,f1931,f2001,f1981,f2015,f2030,f2041,f2043,f2050,f2052,f2073,f2082,f2088,f2090,f2097,f2099,f2119,f2118,f2128,f2157,f2170,f2251,f2252,f2222,f2262,f2269,f2270,f2271,f2272,f2317,f2318,f2319,f2351,f2352,f2349,f2371,f2372,f2373,f2400,f2401,f2398,f2432,f2433,f2465,f2466,f2504,f2505,f2544,f2545,f2617,f2622,f2765,f2770,f2841,f2844,f2847,f2848,f2751,f2842,f2996,f2855,f2860,f2861,f2862,f2863,f2865,f2866,f2867,f2868,f2869,f2870,f2871,f2872,f3012,f3014,f2875,f2876,f2878,f2880,f2882,f2883,f2885,f2887,f2890,f2892,f2893,f2895,f2897,f2898,f2899,f2900,f2901,f2902,f2903,f2904,f3037,f2906,f2907,f2912,f2913,f2928,f2929,f3067])).
fof(f3066,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f334,f333,f335,f592,f343,f355,f354,f359,f358,f412,f411,f413,f414,f420,f419,f418,f417,f424,f423,f422,f421,f428,f427,f426,f429,f430,f431,f432,f433,f437,f434,f444,f443,f593,f462,f459,f458,f457,f468,f467,f477,f478,f479,f480,f481,f482,f483,f489,f488,f487,f485,f484,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f787,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f818,f819,f825,f828,f829,f830,f839,f838,f837,f836,f840,f852,f861,f862,f863,f864,f865,f866,f867,f868,f869,f870,f871,f873,f888,f889,f891,f898,f901,f902,f904,f917,f923,f924,f926,f930,f933,f939,f940,f942,f946,f954,f960,f961,f966,f971,f975,f980,f984,f991,f994,f1000,f1001,f1007,f1012,f1023,f1024,f1025,f1032,f1033,f1035,f1047,f1048,f1050,f1057,f1059,f1062,f1068,f1069,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1210,f1212,f1217,f1219,f1230,f1231,f1239,f1247,f1249,f1254,f1256,f1267,f1275,f1283,f1285,f1339,f1333,f1332,f1373,f1366,f1365,f1385,f1411,f1410,f1403,f1402,f1422,f1433,f1437,f1461,f1486,f1492,f1494,f1501,f1512,f1516,f1542,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1663,f1662,f1661,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1818,f1819,f1816,f1837,f1838,f1835,f1852,f1854,f1859,f1861,f1931,f2001,f1981,f2015,f2030,f2041,f2043,f2050,f2052,f2073,f2082,f2088,f2090,f2097,f2099,f2119,f2118,f2128,f2157,f2170,f2251,f2252,f2222,f2262,f2269,f2270,f2271,f2272,f2317,f2318,f2319,f2351,f2352,f2349,f2371,f2372,f2373,f2400,f2401,f2398,f2432,f2433,f2465,f2466,f2504,f2505,f2544,f2545,f2617,f2622,f2765,f2770,f2841,f2844,f2847,f2848,f2751,f2842,f2996,f2855,f2860,f2861,f2862,f2863,f2865,f2866,f2867,f2868,f2869,f2870,f2871,f2872,f3012,f3014,f2875,f2876,f2878,f2880,f2882,f2883,f2885,f2887,f2890,f2892,f2893,f2895,f2897,f2898,f2899,f2900,f2901,f2902,f2903,f2904,f3037,f2906,f2907,f2912,f2913,f2928,f2929])).
fof(f3065,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f334,f333,f335,f592,f343,f355,f354,f359,f358,f412,f411,f413,f414,f420,f419,f418,f417,f424,f423,f422,f421,f428,f427,f426,f429,f430,f431,f432,f433,f437,f434,f444,f443,f593,f462,f459,f458,f457,f468,f467,f477,f478,f479,f480,f481,f482,f483,f489,f488,f487,f485,f484,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f787,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f818,f819,f825,f828,f829,f830,f839,f838,f837,f836,f840,f852,f861,f862,f863,f864,f865,f866,f867,f868,f869,f870,f871,f873,f888,f889,f891,f898,f901,f902,f904,f917,f923,f924,f926,f930,f933,f939,f940,f942,f946,f954,f960,f961,f966,f971,f975,f980,f984,f991,f994,f1000,f1001,f1007,f1012,f1023,f1024,f1025,f1032,f1033,f1035,f1047,f1048,f1050,f1057,f1059,f1062,f1068,f1069,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1210,f1212,f1217,f1219,f1230,f1231,f1239,f1247,f1249,f1254,f1256,f1267,f1275,f1283,f1285,f1339,f1333,f1332,f1373,f1366,f1365,f1385,f1411,f1410,f1403,f1402,f1422,f1433,f1437,f1461,f1486,f1492,f1494,f1501,f1512,f1516,f1542,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1663,f1662,f1661,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1818,f1819,f1816,f1837,f1838,f1835,f1852,f1854,f1859,f1861,f1931,f2001,f1981,f2015,f2030,f2041,f2043,f2050,f2052,f2073,f2082,f2088,f2090,f2097,f2099,f2119,f2118,f2128,f2157,f2170,f2251,f2252,f2222,f2262,f2269,f2270,f2271,f2272,f2317,f2318,f2319,f2351,f2352,f2349,f2371,f2372,f2373,f2400,f2401,f2398,f2432,f2433,f2465,f2466,f2504,f2505,f2544,f2545,f2617,f2622,f2765,f2770,f2841,f2844,f2847,f2848,f2751,f2842,f2996,f2855,f2860,f2861,f2862,f2863,f2865,f2866,f2867,f2868,f2869,f2870,f2871,f2872,f3012,f3014,f2875,f2876,f2878,f2880,f2882,f2883,f2885,f2887,f2890,f2892,f2893,f2895,f2897,f2898,f2899,f2900,f2901,f2902,f2903,f2904,f3037,f2906,f2907,f2912,f2913,f2928])).
fof(f3048,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f334,f333,f335,f592,f343,f355,f354,f359,f358,f412,f411,f413,f414,f420,f419,f418,f417,f424,f423,f422,f421,f428,f427,f426,f429,f430,f431,f432,f433,f437,f434,f444,f443,f593,f462,f459,f458,f457,f468,f467,f477,f478,f479,f480,f481,f482,f483,f489,f488,f487,f485,f484,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f787,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f818,f819,f825,f828,f829,f830,f839,f838,f837,f836,f840,f852,f861,f862,f863,f864,f865,f866,f867,f868,f869,f870,f871,f873,f888,f889,f891,f898,f901,f902,f904,f917,f923,f924,f926,f930,f933,f939,f940,f942,f946,f954,f960,f961,f966,f971,f975,f980,f984,f991,f994,f1000,f1001,f1007,f1012,f1023,f1024,f1025,f1032,f1033,f1035,f1047,f1048,f1050,f1057,f1059,f1062,f1068,f1069,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1210,f1212,f1217,f1219,f1230,f1231,f1239,f1247,f1249,f1254,f1256,f1267,f1275,f1283,f1285,f1339,f1333,f1332,f1373,f1366,f1365,f1385,f1411,f1410,f1403,f1402,f1422,f1433,f1437,f1461,f1486,f1492,f1494,f1501,f1512,f1516,f1542,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1663,f1662,f1661,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1818,f1819,f1816,f1837,f1838,f1835,f1852,f1854,f1859,f1861,f1931,f2001,f1981,f2015,f2030,f2041,f2043,f2050,f2052,f2073,f2082,f2088,f2090,f2097,f2099,f2119,f2118,f2128,f2157,f2170,f2251,f2252,f2222,f2262,f2269,f2270,f2271,f2272,f2317,f2318,f2319,f2351,f2352,f2349,f2371,f2372,f2373,f2400,f2401,f2398,f2432,f2433,f2465,f2466,f2504,f2505,f2544,f2545,f2617,f2622,f2765,f2770,f2841,f2844,f2847,f2848,f2751,f2842,f2996,f2855,f2860,f2861,f2862,f2863,f2865,f2866,f2867,f2868,f2869,f2870,f2871,f2872,f3012,f3014,f2875,f2876,f2878,f2880,f2882,f2883,f2885,f2887,f2890,f2892,f2893,f2895,f2897,f2898,f2899,f2900,f2901,f2902,f2903,f2904,f3037,f2906,f2907,f2912,f2913])).
fof(f3047,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f334,f333,f335,f592,f343,f355,f354,f359,f358,f412,f411,f413,f414,f420,f419,f418,f417,f424,f423,f422,f421,f428,f427,f426,f429,f430,f431,f432,f433,f437,f434,f444,f443,f593,f462,f459,f458,f457,f468,f467,f477,f478,f479,f480,f481,f482,f483,f489,f488,f487,f485,f484,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f787,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f818,f819,f825,f828,f829,f830,f839,f838,f837,f836,f840,f852,f861,f862,f863,f864,f865,f866,f867,f868,f869,f870,f871,f873,f888,f889,f891,f898,f901,f902,f904,f917,f923,f924,f926,f930,f933,f939,f940,f942,f946,f954,f960,f961,f966,f971,f975,f980,f984,f991,f994,f1000,f1001,f1007,f1012,f1023,f1024,f1025,f1032,f1033,f1035,f1047,f1048,f1050,f1057,f1059,f1062,f1068,f1069,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1210,f1212,f1217,f1219,f1230,f1231,f1239,f1247,f1249,f1254,f1256,f1267,f1275,f1283,f1285,f1339,f1333,f1332,f1373,f1366,f1365,f1385,f1411,f1410,f1403,f1402,f1422,f1433,f1437,f1461,f1486,f1492,f1494,f1501,f1512,f1516,f1542,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1663,f1662,f1661,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1818,f1819,f1816,f1837,f1838,f1835,f1852,f1854,f1859,f1861,f1931,f2001,f1981,f2015,f2030,f2041,f2043,f2050,f2052,f2073,f2082,f2088,f2090,f2097,f2099,f2119,f2118,f2128,f2157,f2170,f2251,f2252,f2222,f2262,f2269,f2270,f2271,f2272,f2317,f2318,f2319,f2351,f2352,f2349,f2371,f2372,f2373,f2400,f2401,f2398,f2432,f2433,f2465,f2466,f2504,f2505,f2544,f2545,f2617,f2622,f2765,f2770,f2841,f2844,f2847,f2848,f2751,f2842,f2996,f2855,f2860,f2861,f2862,f2863,f2865,f2866,f2867,f2868,f2869,f2870,f2871,f2872,f3012,f3014,f2875,f2876,f2878,f2880,f2882,f2883,f2885,f2887,f2890,f2892,f2893,f2895,f2897,f2898,f2899,f2900,f2901,f2902,f2903,f2904,f3037,f2906,f2907,f2912])).
fof(f3040,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f334,f333,f335,f592,f343,f355,f354,f359,f358,f412,f411,f413,f414,f420,f419,f418,f417,f424,f423,f422,f421,f428,f427,f426,f429,f430,f431,f432,f433,f437,f434,f444,f443,f593,f462,f459,f458,f457,f468,f467,f477,f478,f479,f480,f481,f482,f483,f489,f488,f487,f485,f484,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f787,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f818,f819,f825,f828,f829,f830,f839,f838,f837,f836,f840,f852,f861,f862,f863,f864,f865,f866,f867,f868,f869,f870,f871,f873,f888,f889,f891,f898,f901,f902,f904,f917,f923,f924,f926,f930,f933,f939,f940,f942,f946,f954,f960,f961,f966,f971,f975,f980,f984,f991,f994,f1000,f1001,f1007,f1012,f1023,f1024,f1025,f1032,f1033,f1035,f1047,f1048,f1050,f1057,f1059,f1062,f1068,f1069,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1210,f1212,f1217,f1219,f1230,f1231,f1239,f1247,f1249,f1254,f1256,f1267,f1275,f1283,f1285,f1339,f1333,f1332,f1373,f1366,f1365,f1385,f1411,f1410,f1403,f1402,f1422,f1433,f1437,f1461,f1486,f1492,f1494,f1501,f1512,f1516,f1542,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1663,f1662,f1661,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1818,f1819,f1816,f1837,f1838,f1835,f1852,f1854,f1859,f1861,f1931,f2001,f1981,f2015,f2030,f2041,f2043,f2050,f2052,f2073,f2082,f2088,f2090,f2097,f2099,f2119,f2118,f2128,f2157,f2170,f2251,f2252,f2222,f2262,f2269,f2270,f2271,f2272,f2317,f2318,f2319,f2351,f2352,f2349,f2371,f2372,f2373,f2400,f2401,f2398,f2432,f2433,f2465,f2466,f2504,f2505,f2544,f2545,f2617,f2622,f2765,f2770,f2841,f2844,f2847,f2848,f2751,f2842,f2996,f2855,f2860,f2861,f2862,f2863,f2865,f2866,f2867,f2868,f2869,f2870,f2871,f2872,f3012,f3014,f2875,f2876,f2878,f2880,f2882,f2883,f2885,f2887,f2890,f2892,f2893,f2895,f2897,f2898,f2899,f2900,f2901,f2902,f2903,f2904,f3037,f2906,f2907])).
fof(f3039,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f334,f333,f335,f592,f343,f355,f354,f359,f358,f412,f411,f413,f414,f420,f419,f418,f417,f424,f423,f422,f421,f428,f427,f426,f429,f430,f431,f432,f433,f437,f434,f444,f443,f593,f462,f459,f458,f457,f468,f467,f477,f478,f479,f480,f481,f482,f483,f489,f488,f487,f485,f484,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f787,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f818,f819,f825,f828,f829,f830,f839,f838,f837,f836,f840,f852,f861,f862,f863,f864,f865,f866,f867,f868,f869,f870,f871,f873,f888,f889,f891,f898,f901,f902,f904,f917,f923,f924,f926,f930,f933,f939,f940,f942,f946,f954,f960,f961,f966,f971,f975,f980,f984,f991,f994,f1000,f1001,f1007,f1012,f1023,f1024,f1025,f1032,f1033,f1035,f1047,f1048,f1050,f1057,f1059,f1062,f1068,f1069,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1210,f1212,f1217,f1219,f1230,f1231,f1239,f1247,f1249,f1254,f1256,f1267,f1275,f1283,f1285,f1339,f1333,f1332,f1373,f1366,f1365,f1385,f1411,f1410,f1403,f1402,f1422,f1433,f1437,f1461,f1486,f1492,f1494,f1501,f1512,f1516,f1542,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1663,f1662,f1661,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1818,f1819,f1816,f1837,f1838,f1835,f1852,f1854,f1859,f1861,f1931,f2001,f1981,f2015,f2030,f2041,f2043,f2050,f2052,f2073,f2082,f2088,f2090,f2097,f2099,f2119,f2118,f2128,f2157,f2170,f2251,f2252,f2222,f2262,f2269,f2270,f2271,f2272,f2317,f2318,f2319,f2351,f2352,f2349,f2371,f2372,f2373,f2400,f2401,f2398,f2432,f2433,f2465,f2466,f2504,f2505,f2544,f2545,f2617,f2622,f2765,f2770,f2841,f2844,f2847,f2848,f2751,f2842,f2996,f2855,f2860,f2861,f2862,f2863,f2865,f2866,f2867,f2868,f2869,f2870,f2871,f2872,f3012,f3014,f2875,f2876,f2878,f2880,f2882,f2883,f2885,f2887,f2890,f2892,f2893,f2895,f2897,f2898,f2899,f2900,f2901,f2902,f2903,f2904,f3037,f2906])).
fof(f3038,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f334,f333,f335,f592,f343,f355,f354,f359,f358,f412,f411,f413,f414,f420,f419,f418,f417,f424,f423,f422,f421,f428,f427,f426,f429,f430,f431,f432,f433,f437,f434,f444,f443,f593,f462,f459,f458,f457,f468,f467,f477,f478,f479,f480,f481,f482,f483,f489,f488,f487,f485,f484,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f787,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f818,f819,f825,f828,f829,f830,f839,f838,f837,f836,f840,f852,f861,f862,f863,f864,f865,f866,f867,f868,f869,f870,f871,f873,f888,f889,f891,f898,f901,f902,f904,f917,f923,f924,f926,f930,f933,f939,f940,f942,f946,f954,f960,f961,f966,f971,f975,f980,f984,f991,f994,f1000,f1001,f1007,f1012,f1023,f1024,f1025,f1032,f1033,f1035,f1047,f1048,f1050,f1057,f1059,f1062,f1068,f1069,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1210,f1212,f1217,f1219,f1230,f1231,f1239,f1247,f1249,f1254,f1256,f1267,f1275,f1283,f1285,f1339,f1333,f1332,f1373,f1366,f1365,f1385,f1411,f1410,f1403,f1402,f1422,f1433,f1437,f1461,f1486,f1492,f1494,f1501,f1512,f1516,f1542,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1663,f1662,f1661,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1818,f1819,f1816,f1837,f1838,f1835,f1852,f1854,f1859,f1861,f1931,f2001,f1981,f2015,f2030,f2041,f2043,f2050,f2052,f2073,f2082,f2088,f2090,f2097,f2099,f2119,f2118,f2128,f2157,f2170,f2251,f2252,f2222,f2262,f2269,f2270,f2271,f2272,f2317,f2318,f2319,f2351,f2352,f2349,f2371,f2372,f2373,f2400,f2401,f2398,f2432,f2433,f2465,f2466,f2504,f2505,f2544,f2545,f2617,f2622,f2765,f2770,f2841,f2844,f2847,f2848,f2751,f2842,f2996,f2855,f2860,f2861,f2862,f2863,f2865,f2866,f2867,f2868,f2869,f2870,f2871,f2872,f3012,f3014,f2875,f2876,f2878,f2880,f2882,f2883,f2885,f2887,f2890,f2892,f2893,f2895,f2897,f2898,f2899,f2900,f2901,f2902,f2903,f2904,f3037])).
fof(f3035,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f334,f333,f335,f592,f343,f355,f354,f359,f358,f412,f411,f413,f414,f420,f419,f418,f417,f424,f423,f422,f421,f428,f427,f426,f429,f430,f431,f432,f433,f437,f434,f444,f443,f593,f462,f459,f458,f457,f468,f467,f477,f478,f479,f480,f481,f482,f483,f489,f488,f487,f485,f484,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f787,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f818,f819,f825,f828,f829,f830,f839,f838,f837,f836,f840,f852,f861,f862,f863,f864,f865,f866,f867,f868,f869,f870,f871,f873,f888,f889,f891,f898,f901,f902,f904,f917,f923,f924,f926,f930,f933,f939,f940,f942,f946,f954,f960,f961,f966,f971,f975,f980,f984,f991,f994,f1000,f1001,f1007,f1012,f1023,f1024,f1025,f1032,f1033,f1035,f1047,f1048,f1050,f1057,f1059,f1062,f1068,f1069,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1210,f1212,f1217,f1219,f1230,f1231,f1239,f1247,f1249,f1254,f1256,f1267,f1275,f1283,f1285,f1339,f1333,f1332,f1373,f1366,f1365,f1385,f1411,f1410,f1403,f1402,f1422,f1433,f1437,f1461,f1486,f1492,f1494,f1501,f1512,f1516,f1542,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1663,f1662,f1661,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1818,f1819,f1816,f1837,f1838,f1835,f1852,f1854,f1859,f1861,f1931,f2001,f1981,f2015,f2030,f2041,f2043,f2050,f2052,f2073,f2082,f2088,f2090,f2097,f2099,f2119,f2118,f2128,f2157,f2170,f2251,f2252,f2222,f2262,f2269,f2270,f2271,f2272,f2317,f2318,f2319,f2351,f2352,f2349,f2371,f2372,f2373,f2400,f2401,f2398,f2432,f2433,f2465,f2466,f2504,f2505,f2544,f2545,f2617,f2622,f2765,f2770,f2841,f2844,f2847,f2848,f2751,f2842,f2996,f2855,f2860,f2861,f2862,f2863,f2865,f2866,f2867,f2868,f2869,f2870,f2871,f2872,f3012,f3014,f2875,f2876,f2878,f2880,f2882,f2883,f2885,f2887,f2890,f2892,f2893,f2895,f2897,f2898,f2899,f2900,f2901,f2902,f2903,f2904])).
fof(f3034,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f334,f333,f335,f592,f343,f355,f354,f359,f358,f412,f411,f413,f414,f420,f419,f418,f417,f424,f423,f422,f421,f428,f427,f426,f429,f430,f431,f432,f433,f437,f434,f444,f443,f593,f462,f459,f458,f457,f468,f467,f477,f478,f479,f480,f481,f482,f483,f489,f488,f487,f485,f484,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f787,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f818,f819,f825,f828,f829,f830,f839,f838,f837,f836,f840,f852,f861,f862,f863,f864,f865,f866,f867,f868,f869,f870,f871,f873,f888,f889,f891,f898,f901,f902,f904,f917,f923,f924,f926,f930,f933,f939,f940,f942,f946,f954,f960,f961,f966,f971,f975,f980,f984,f991,f994,f1000,f1001,f1007,f1012,f1023,f1024,f1025,f1032,f1033,f1035,f1047,f1048,f1050,f1057,f1059,f1062,f1068,f1069,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1210,f1212,f1217,f1219,f1230,f1231,f1239,f1247,f1249,f1254,f1256,f1267,f1275,f1283,f1285,f1339,f1333,f1332,f1373,f1366,f1365,f1385,f1411,f1410,f1403,f1402,f1422,f1433,f1437,f1461,f1486,f1492,f1494,f1501,f1512,f1516,f1542,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1663,f1662,f1661,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1818,f1819,f1816,f1837,f1838,f1835,f1852,f1854,f1859,f1861,f1931,f2001,f1981,f2015,f2030,f2041,f2043,f2050,f2052,f2073,f2082,f2088,f2090,f2097,f2099,f2119,f2118,f2128,f2157,f2170,f2251,f2252,f2222,f2262,f2269,f2270,f2271,f2272,f2317,f2318,f2319,f2351,f2352,f2349,f2371,f2372,f2373,f2400,f2401,f2398,f2432,f2433,f2465,f2466,f2504,f2505,f2544,f2545,f2617,f2622,f2765,f2770,f2841,f2844,f2847,f2848,f2751,f2842,f2996,f2855,f2860,f2861,f2862,f2863,f2865,f2866,f2867,f2868,f2869,f2870,f2871,f2872,f3012,f3014,f2875,f2876,f2878,f2880,f2882,f2883,f2885,f2887,f2890,f2892,f2893,f2895,f2897,f2898,f2899,f2900,f2901,f2902,f2903])).
fof(f3033,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f334,f333,f335,f592,f343,f355,f354,f359,f358,f412,f411,f413,f414,f420,f419,f418,f417,f424,f423,f422,f421,f428,f427,f426,f429,f430,f431,f432,f433,f437,f434,f444,f443,f593,f462,f459,f458,f457,f468,f467,f477,f478,f479,f480,f481,f482,f483,f489,f488,f487,f485,f484,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f787,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f818,f819,f825,f828,f829,f830,f839,f838,f837,f836,f840,f852,f861,f862,f863,f864,f865,f866,f867,f868,f869,f870,f871,f873,f888,f889,f891,f898,f901,f902,f904,f917,f923,f924,f926,f930,f933,f939,f940,f942,f946,f954,f960,f961,f966,f971,f975,f980,f984,f991,f994,f1000,f1001,f1007,f1012,f1023,f1024,f1025,f1032,f1033,f1035,f1047,f1048,f1050,f1057,f1059,f1062,f1068,f1069,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1210,f1212,f1217,f1219,f1230,f1231,f1239,f1247,f1249,f1254,f1256,f1267,f1275,f1283,f1285,f1339,f1333,f1332,f1373,f1366,f1365,f1385,f1411,f1410,f1403,f1402,f1422,f1433,f1437,f1461,f1486,f1492,f1494,f1501,f1512,f1516,f1542,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1663,f1662,f1661,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1818,f1819,f1816,f1837,f1838,f1835,f1852,f1854,f1859,f1861,f1931,f2001,f1981,f2015,f2030,f2041,f2043,f2050,f2052,f2073,f2082,f2088,f2090,f2097,f2099,f2119,f2118,f2128,f2157,f2170,f2251,f2252,f2222,f2262,f2269,f2270,f2271,f2272,f2317,f2318,f2319,f2351,f2352,f2349,f2371,f2372,f2373,f2400,f2401,f2398,f2432,f2433,f2465,f2466,f2504,f2505,f2544,f2545,f2617,f2622,f2765,f2770,f2841,f2844,f2847,f2848,f2751,f2842,f2996,f2855,f2860,f2861,f2862,f2863,f2865,f2866,f2867,f2868,f2869,f2870,f2871,f2872,f3012,f3014,f2875,f2876,f2878,f2880,f2882,f2883,f2885,f2887,f2890,f2892,f2893,f2895,f2897,f2898,f2899,f2900,f2901,f2902])).
fof(f3032,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f334,f333,f335,f592,f343,f355,f354,f359,f358,f412,f411,f413,f414,f420,f419,f418,f417,f424,f423,f422,f421,f428,f427,f426,f429,f430,f431,f432,f433,f437,f434,f444,f443,f593,f462,f459,f458,f457,f468,f467,f477,f478,f479,f480,f481,f482,f483,f489,f488,f487,f485,f484,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f787,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f818,f819,f825,f828,f829,f830,f839,f838,f837,f836,f840,f852,f861,f862,f863,f864,f865,f866,f867,f868,f869,f870,f871,f873,f888,f889,f891,f898,f901,f902,f904,f917,f923,f924,f926,f930,f933,f939,f940,f942,f946,f954,f960,f961,f966,f971,f975,f980,f984,f991,f994,f1000,f1001,f1007,f1012,f1023,f1024,f1025,f1032,f1033,f1035,f1047,f1048,f1050,f1057,f1059,f1062,f1068,f1069,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1210,f1212,f1217,f1219,f1230,f1231,f1239,f1247,f1249,f1254,f1256,f1267,f1275,f1283,f1285,f1339,f1333,f1332,f1373,f1366,f1365,f1385,f1411,f1410,f1403,f1402,f1422,f1433,f1437,f1461,f1486,f1492,f1494,f1501,f1512,f1516,f1542,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1663,f1662,f1661,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1818,f1819,f1816,f1837,f1838,f1835,f1852,f1854,f1859,f1861,f1931,f2001,f1981,f2015,f2030,f2041,f2043,f2050,f2052,f2073,f2082,f2088,f2090,f2097,f2099,f2119,f2118,f2128,f2157,f2170,f2251,f2252,f2222,f2262,f2269,f2270,f2271,f2272,f2317,f2318,f2319,f2351,f2352,f2349,f2371,f2372,f2373,f2400,f2401,f2398,f2432,f2433,f2465,f2466,f2504,f2505,f2544,f2545,f2617,f2622,f2765,f2770,f2841,f2844,f2847,f2848,f2751,f2842,f2996,f2855,f2860,f2861,f2862,f2863,f2865,f2866,f2867,f2868,f2869,f2870,f2871,f2872,f3012,f3014,f2875,f2876,f2878,f2880,f2882,f2883,f2885,f2887,f2890,f2892,f2893,f2895,f2897,f2898,f2899,f2900,f2901])).
fof(f3031,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f334,f333,f335,f592,f343,f355,f354,f359,f358,f412,f411,f413,f414,f420,f419,f418,f417,f424,f423,f422,f421,f428,f427,f426,f429,f430,f431,f432,f433,f437,f434,f444,f443,f593,f462,f459,f458,f457,f468,f467,f477,f478,f479,f480,f481,f482,f483,f489,f488,f487,f485,f484,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f787,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f818,f819,f825,f828,f829,f830,f839,f838,f837,f836,f840,f852,f861,f862,f863,f864,f865,f866,f867,f868,f869,f870,f871,f873,f888,f889,f891,f898,f901,f902,f904,f917,f923,f924,f926,f930,f933,f939,f940,f942,f946,f954,f960,f961,f966,f971,f975,f980,f984,f991,f994,f1000,f1001,f1007,f1012,f1023,f1024,f1025,f1032,f1033,f1035,f1047,f1048,f1050,f1057,f1059,f1062,f1068,f1069,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1210,f1212,f1217,f1219,f1230,f1231,f1239,f1247,f1249,f1254,f1256,f1267,f1275,f1283,f1285,f1339,f1333,f1332,f1373,f1366,f1365,f1385,f1411,f1410,f1403,f1402,f1422,f1433,f1437,f1461,f1486,f1492,f1494,f1501,f1512,f1516,f1542,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1663,f1662,f1661,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1818,f1819,f1816,f1837,f1838,f1835,f1852,f1854,f1859,f1861,f1931,f2001,f1981,f2015,f2030,f2041,f2043,f2050,f2052,f2073,f2082,f2088,f2090,f2097,f2099,f2119,f2118,f2128,f2157,f2170,f2251,f2252,f2222,f2262,f2269,f2270,f2271,f2272,f2317,f2318,f2319,f2351,f2352,f2349,f2371,f2372,f2373,f2400,f2401,f2398,f2432,f2433,f2465,f2466,f2504,f2505,f2544,f2545,f2617,f2622,f2765,f2770,f2841,f2844,f2847,f2848,f2751,f2842,f2996,f2855,f2860,f2861,f2862,f2863,f2865,f2866,f2867,f2868,f2869,f2870,f2871,f2872,f3012,f3014,f2875,f2876,f2878,f2880,f2882,f2883,f2885,f2887,f2890,f2892,f2893,f2895,f2897,f2898,f2899,f2900])).
fof(f3030,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f334,f333,f335,f592,f343,f355,f354,f359,f358,f412,f411,f413,f414,f420,f419,f418,f417,f424,f423,f422,f421,f428,f427,f426,f429,f430,f431,f432,f433,f437,f434,f444,f443,f593,f462,f459,f458,f457,f468,f467,f477,f478,f479,f480,f481,f482,f483,f489,f488,f487,f485,f484,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f787,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f818,f819,f825,f828,f829,f830,f839,f838,f837,f836,f840,f852,f861,f862,f863,f864,f865,f866,f867,f868,f869,f870,f871,f873,f888,f889,f891,f898,f901,f902,f904,f917,f923,f924,f926,f930,f933,f939,f940,f942,f946,f954,f960,f961,f966,f971,f975,f980,f984,f991,f994,f1000,f1001,f1007,f1012,f1023,f1024,f1025,f1032,f1033,f1035,f1047,f1048,f1050,f1057,f1059,f1062,f1068,f1069,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1210,f1212,f1217,f1219,f1230,f1231,f1239,f1247,f1249,f1254,f1256,f1267,f1275,f1283,f1285,f1339,f1333,f1332,f1373,f1366,f1365,f1385,f1411,f1410,f1403,f1402,f1422,f1433,f1437,f1461,f1486,f1492,f1494,f1501,f1512,f1516,f1542,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1663,f1662,f1661,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1818,f1819,f1816,f1837,f1838,f1835,f1852,f1854,f1859,f1861,f1931,f2001,f1981,f2015,f2030,f2041,f2043,f2050,f2052,f2073,f2082,f2088,f2090,f2097,f2099,f2119,f2118,f2128,f2157,f2170,f2251,f2252,f2222,f2262,f2269,f2270,f2271,f2272,f2317,f2318,f2319,f2351,f2352,f2349,f2371,f2372,f2373,f2400,f2401,f2398,f2432,f2433,f2465,f2466,f2504,f2505,f2544,f2545,f2617,f2622,f2765,f2770,f2841,f2844,f2847,f2848,f2751,f2842,f2996,f2855,f2860,f2861,f2862,f2863,f2865,f2866,f2867,f2868,f2869,f2870,f2871,f2872,f3012,f3014,f2875,f2876,f2878,f2880,f2882,f2883,f2885,f2887,f2890,f2892,f2893,f2895,f2897,f2898,f2899])).
fof(f3029,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f334,f333,f335,f592,f343,f355,f354,f359,f358,f412,f411,f413,f414,f420,f419,f418,f417,f424,f423,f422,f421,f428,f427,f426,f429,f430,f431,f432,f433,f437,f434,f444,f443,f593,f462,f459,f458,f457,f468,f467,f477,f478,f479,f480,f481,f482,f483,f489,f488,f487,f485,f484,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f787,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f818,f819,f825,f828,f829,f830,f839,f838,f837,f836,f840,f852,f861,f862,f863,f864,f865,f866,f867,f868,f869,f870,f871,f873,f888,f889,f891,f898,f901,f902,f904,f917,f923,f924,f926,f930,f933,f939,f940,f942,f946,f954,f960,f961,f966,f971,f975,f980,f984,f991,f994,f1000,f1001,f1007,f1012,f1023,f1024,f1025,f1032,f1033,f1035,f1047,f1048,f1050,f1057,f1059,f1062,f1068,f1069,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1210,f1212,f1217,f1219,f1230,f1231,f1239,f1247,f1249,f1254,f1256,f1267,f1275,f1283,f1285,f1339,f1333,f1332,f1373,f1366,f1365,f1385,f1411,f1410,f1403,f1402,f1422,f1433,f1437,f1461,f1486,f1492,f1494,f1501,f1512,f1516,f1542,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1663,f1662,f1661,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1818,f1819,f1816,f1837,f1838,f1835,f1852,f1854,f1859,f1861,f1931,f2001,f1981,f2015,f2030,f2041,f2043,f2050,f2052,f2073,f2082,f2088,f2090,f2097,f2099,f2119,f2118,f2128,f2157,f2170,f2251,f2252,f2222,f2262,f2269,f2270,f2271,f2272,f2317,f2318,f2319,f2351,f2352,f2349,f2371,f2372,f2373,f2400,f2401,f2398,f2432,f2433,f2465,f2466,f2504,f2505,f2544,f2545,f2617,f2622,f2765,f2770,f2841,f2844,f2847,f2848,f2751,f2842,f2996,f2855,f2860,f2861,f2862,f2863,f2865,f2866,f2867,f2868,f2869,f2870,f2871,f2872,f3012,f3014,f2875,f2876,f2878,f2880,f2882,f2883,f2885,f2887,f2890,f2892,f2893,f2895,f2897,f2898])).
fof(f3028,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f334,f333,f335,f592,f343,f355,f354,f359,f358,f412,f411,f413,f414,f420,f419,f418,f417,f424,f423,f422,f421,f428,f427,f426,f429,f430,f431,f432,f433,f437,f434,f444,f443,f593,f462,f459,f458,f457,f468,f467,f477,f478,f479,f480,f481,f482,f483,f489,f488,f487,f485,f484,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f787,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f818,f819,f825,f828,f829,f830,f839,f838,f837,f836,f840,f852,f861,f862,f863,f864,f865,f866,f867,f868,f869,f870,f871,f873,f888,f889,f891,f898,f901,f902,f904,f917,f923,f924,f926,f930,f933,f939,f940,f942,f946,f954,f960,f961,f966,f971,f975,f980,f984,f991,f994,f1000,f1001,f1007,f1012,f1023,f1024,f1025,f1032,f1033,f1035,f1047,f1048,f1050,f1057,f1059,f1062,f1068,f1069,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1210,f1212,f1217,f1219,f1230,f1231,f1239,f1247,f1249,f1254,f1256,f1267,f1275,f1283,f1285,f1339,f1333,f1332,f1373,f1366,f1365,f1385,f1411,f1410,f1403,f1402,f1422,f1433,f1437,f1461,f1486,f1492,f1494,f1501,f1512,f1516,f1542,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1663,f1662,f1661,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1818,f1819,f1816,f1837,f1838,f1835,f1852,f1854,f1859,f1861,f1931,f2001,f1981,f2015,f2030,f2041,f2043,f2050,f2052,f2073,f2082,f2088,f2090,f2097,f2099,f2119,f2118,f2128,f2157,f2170,f2251,f2252,f2222,f2262,f2269,f2270,f2271,f2272,f2317,f2318,f2319,f2351,f2352,f2349,f2371,f2372,f2373,f2400,f2401,f2398,f2432,f2433,f2465,f2466,f2504,f2505,f2544,f2545,f2617,f2622,f2765,f2770,f2841,f2844,f2847,f2848,f2751,f2842,f2996,f2855,f2860,f2861,f2862,f2863,f2865,f2866,f2867,f2868,f2869,f2870,f2871,f2872,f3012,f3014,f2875,f2876,f2878,f2880,f2882,f2883,f2885,f2887,f2890,f2892,f2893,f2895,f2897])).
fof(f3027,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f334,f333,f335,f592,f343,f355,f354,f359,f358,f412,f411,f413,f414,f420,f419,f418,f417,f424,f423,f422,f421,f428,f427,f426,f429,f430,f431,f432,f433,f437,f434,f444,f443,f593,f462,f459,f458,f457,f468,f467,f477,f478,f479,f480,f481,f482,f483,f489,f488,f487,f485,f484,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f787,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f818,f819,f825,f828,f829,f830,f839,f838,f837,f836,f840,f852,f861,f862,f863,f864,f865,f866,f867,f868,f869,f870,f871,f873,f888,f889,f891,f898,f901,f902,f904,f917,f923,f924,f926,f930,f933,f939,f940,f942,f946,f954,f960,f961,f966,f971,f975,f980,f984,f991,f994,f1000,f1001,f1007,f1012,f1023,f1024,f1025,f1032,f1033,f1035,f1047,f1048,f1050,f1057,f1059,f1062,f1068,f1069,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1210,f1212,f1217,f1219,f1230,f1231,f1239,f1247,f1249,f1254,f1256,f1267,f1275,f1283,f1285,f1339,f1333,f1332,f1373,f1366,f1365,f1385,f1411,f1410,f1403,f1402,f1422,f1433,f1437,f1461,f1486,f1492,f1494,f1501,f1512,f1516,f1542,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1663,f1662,f1661,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1818,f1819,f1816,f1837,f1838,f1835,f1852,f1854,f1859,f1861,f1931,f2001,f1981,f2015,f2030,f2041,f2043,f2050,f2052,f2073,f2082,f2088,f2090,f2097,f2099,f2119,f2118,f2128,f2157,f2170,f2251,f2252,f2222,f2262,f2269,f2270,f2271,f2272,f2317,f2318,f2319,f2351,f2352,f2349,f2371,f2372,f2373,f2400,f2401,f2398,f2432,f2433,f2465,f2466,f2504,f2505,f2544,f2545,f2617,f2622,f2765,f2770,f2841,f2844,f2847,f2848,f2751,f2842,f2996,f2855,f2860,f2861,f2862,f2863,f2865,f2866,f2867,f2868,f2869,f2870,f2871,f2872,f3012,f3014,f2875,f2876,f2878,f2880,f2882,f2883,f2885,f2887,f2890,f2892,f2893,f2895])).
fof(f3026,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f334,f333,f335,f592,f343,f355,f354,f359,f358,f412,f411,f413,f414,f420,f419,f418,f417,f424,f423,f422,f421,f428,f427,f426,f429,f430,f431,f432,f433,f437,f434,f444,f443,f593,f462,f459,f458,f457,f468,f467,f477,f478,f479,f480,f481,f482,f483,f489,f488,f487,f485,f484,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f787,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f818,f819,f825,f828,f829,f830,f839,f838,f837,f836,f840,f852,f861,f862,f863,f864,f865,f866,f867,f868,f869,f870,f871,f873,f888,f889,f891,f898,f901,f902,f904,f917,f923,f924,f926,f930,f933,f939,f940,f942,f946,f954,f960,f961,f966,f971,f975,f980,f984,f991,f994,f1000,f1001,f1007,f1012,f1023,f1024,f1025,f1032,f1033,f1035,f1047,f1048,f1050,f1057,f1059,f1062,f1068,f1069,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1210,f1212,f1217,f1219,f1230,f1231,f1239,f1247,f1249,f1254,f1256,f1267,f1275,f1283,f1285,f1339,f1333,f1332,f1373,f1366,f1365,f1385,f1411,f1410,f1403,f1402,f1422,f1433,f1437,f1461,f1486,f1492,f1494,f1501,f1512,f1516,f1542,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1663,f1662,f1661,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1818,f1819,f1816,f1837,f1838,f1835,f1852,f1854,f1859,f1861,f1931,f2001,f1981,f2015,f2030,f2041,f2043,f2050,f2052,f2073,f2082,f2088,f2090,f2097,f2099,f2119,f2118,f2128,f2157,f2170,f2251,f2252,f2222,f2262,f2269,f2270,f2271,f2272,f2317,f2318,f2319,f2351,f2352,f2349,f2371,f2372,f2373,f2400,f2401,f2398,f2432,f2433,f2465,f2466,f2504,f2505,f2544,f2545,f2617,f2622,f2765,f2770,f2841,f2844,f2847,f2848,f2751,f2842,f2996,f2855,f2860,f2861,f2862,f2863,f2865,f2866,f2867,f2868,f2869,f2870,f2871,f2872,f3012,f3014,f2875,f2876,f2878,f2880,f2882,f2883,f2885,f2887,f2890,f2892,f2893])).
fof(f3025,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f334,f333,f335,f592,f343,f355,f354,f359,f358,f412,f411,f413,f414,f420,f419,f418,f417,f424,f423,f422,f421,f428,f427,f426,f429,f430,f431,f432,f433,f437,f434,f444,f443,f593,f462,f459,f458,f457,f468,f467,f477,f478,f479,f480,f481,f482,f483,f489,f488,f487,f485,f484,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f787,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f818,f819,f825,f828,f829,f830,f839,f838,f837,f836,f840,f852,f861,f862,f863,f864,f865,f866,f867,f868,f869,f870,f871,f873,f888,f889,f891,f898,f901,f902,f904,f917,f923,f924,f926,f930,f933,f939,f940,f942,f946,f954,f960,f961,f966,f971,f975,f980,f984,f991,f994,f1000,f1001,f1007,f1012,f1023,f1024,f1025,f1032,f1033,f1035,f1047,f1048,f1050,f1057,f1059,f1062,f1068,f1069,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1210,f1212,f1217,f1219,f1230,f1231,f1239,f1247,f1249,f1254,f1256,f1267,f1275,f1283,f1285,f1339,f1333,f1332,f1373,f1366,f1365,f1385,f1411,f1410,f1403,f1402,f1422,f1433,f1437,f1461,f1486,f1492,f1494,f1501,f1512,f1516,f1542,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1663,f1662,f1661,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1818,f1819,f1816,f1837,f1838,f1835,f1852,f1854,f1859,f1861,f1931,f2001,f1981,f2015,f2030,f2041,f2043,f2050,f2052,f2073,f2082,f2088,f2090,f2097,f2099,f2119,f2118,f2128,f2157,f2170,f2251,f2252,f2222,f2262,f2269,f2270,f2271,f2272,f2317,f2318,f2319,f2351,f2352,f2349,f2371,f2372,f2373,f2400,f2401,f2398,f2432,f2433,f2465,f2466,f2504,f2505,f2544,f2545,f2617,f2622,f2765,f2770,f2841,f2844,f2847,f2848,f2751,f2842,f2996,f2855,f2860,f2861,f2862,f2863,f2865,f2866,f2867,f2868,f2869,f2870,f2871,f2872,f3012,f3014,f2875,f2876,f2878,f2880,f2882,f2883,f2885,f2887,f2890,f2892])).
fof(f3024,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f334,f333,f335,f592,f343,f355,f354,f359,f358,f412,f411,f413,f414,f420,f419,f418,f417,f424,f423,f422,f421,f428,f427,f426,f429,f430,f431,f432,f433,f437,f434,f444,f443,f593,f462,f459,f458,f457,f468,f467,f477,f478,f479,f480,f481,f482,f483,f489,f488,f487,f485,f484,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f787,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f818,f819,f825,f828,f829,f830,f839,f838,f837,f836,f840,f852,f861,f862,f863,f864,f865,f866,f867,f868,f869,f870,f871,f873,f888,f889,f891,f898,f901,f902,f904,f917,f923,f924,f926,f930,f933,f939,f940,f942,f946,f954,f960,f961,f966,f971,f975,f980,f984,f991,f994,f1000,f1001,f1007,f1012,f1023,f1024,f1025,f1032,f1033,f1035,f1047,f1048,f1050,f1057,f1059,f1062,f1068,f1069,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1210,f1212,f1217,f1219,f1230,f1231,f1239,f1247,f1249,f1254,f1256,f1267,f1275,f1283,f1285,f1339,f1333,f1332,f1373,f1366,f1365,f1385,f1411,f1410,f1403,f1402,f1422,f1433,f1437,f1461,f1486,f1492,f1494,f1501,f1512,f1516,f1542,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1663,f1662,f1661,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1818,f1819,f1816,f1837,f1838,f1835,f1852,f1854,f1859,f1861,f1931,f2001,f1981,f2015,f2030,f2041,f2043,f2050,f2052,f2073,f2082,f2088,f2090,f2097,f2099,f2119,f2118,f2128,f2157,f2170,f2251,f2252,f2222,f2262,f2269,f2270,f2271,f2272,f2317,f2318,f2319,f2351,f2352,f2349,f2371,f2372,f2373,f2400,f2401,f2398,f2432,f2433,f2465,f2466,f2504,f2505,f2544,f2545,f2617,f2622,f2765,f2770,f2841,f2844,f2847,f2848,f2751,f2842,f2996,f2855,f2860,f2861,f2862,f2863,f2865,f2866,f2867,f2868,f2869,f2870,f2871,f2872,f3012,f3014,f2875,f2876,f2878,f2880,f2882,f2883,f2885,f2887,f2890])).
fof(f3023,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f334,f333,f335,f592,f343,f355,f354,f359,f358,f412,f411,f413,f414,f420,f419,f418,f417,f424,f423,f422,f421,f428,f427,f426,f429,f430,f431,f432,f433,f437,f434,f444,f443,f593,f462,f459,f458,f457,f468,f467,f477,f478,f479,f480,f481,f482,f483,f489,f488,f487,f485,f484,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f787,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f818,f819,f825,f828,f829,f830,f839,f838,f837,f836,f840,f852,f861,f862,f863,f864,f865,f866,f867,f868,f869,f870,f871,f873,f888,f889,f891,f898,f901,f902,f904,f917,f923,f924,f926,f930,f933,f939,f940,f942,f946,f954,f960,f961,f966,f971,f975,f980,f984,f991,f994,f1000,f1001,f1007,f1012,f1023,f1024,f1025,f1032,f1033,f1035,f1047,f1048,f1050,f1057,f1059,f1062,f1068,f1069,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1210,f1212,f1217,f1219,f1230,f1231,f1239,f1247,f1249,f1254,f1256,f1267,f1275,f1283,f1285,f1339,f1333,f1332,f1373,f1366,f1365,f1385,f1411,f1410,f1403,f1402,f1422,f1433,f1437,f1461,f1486,f1492,f1494,f1501,f1512,f1516,f1542,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1663,f1662,f1661,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1818,f1819,f1816,f1837,f1838,f1835,f1852,f1854,f1859,f1861,f1931,f2001,f1981,f2015,f2030,f2041,f2043,f2050,f2052,f2073,f2082,f2088,f2090,f2097,f2099,f2119,f2118,f2128,f2157,f2170,f2251,f2252,f2222,f2262,f2269,f2270,f2271,f2272,f2317,f2318,f2319,f2351,f2352,f2349,f2371,f2372,f2373,f2400,f2401,f2398,f2432,f2433,f2465,f2466,f2504,f2505,f2544,f2545,f2617,f2622,f2765,f2770,f2841,f2844,f2847,f2848,f2751,f2842,f2996,f2855,f2860,f2861,f2862,f2863,f2865,f2866,f2867,f2868,f2869,f2870,f2871,f2872,f3012,f3014,f2875,f2876,f2878,f2880,f2882,f2883,f2885,f2887])).
fof(f3022,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f334,f333,f335,f592,f343,f355,f354,f359,f358,f412,f411,f413,f414,f420,f419,f418,f417,f424,f423,f422,f421,f428,f427,f426,f429,f430,f431,f432,f433,f437,f434,f444,f443,f593,f462,f459,f458,f457,f468,f467,f477,f478,f479,f480,f481,f482,f483,f489,f488,f487,f485,f484,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f787,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f818,f819,f825,f828,f829,f830,f839,f838,f837,f836,f840,f852,f861,f862,f863,f864,f865,f866,f867,f868,f869,f870,f871,f873,f888,f889,f891,f898,f901,f902,f904,f917,f923,f924,f926,f930,f933,f939,f940,f942,f946,f954,f960,f961,f966,f971,f975,f980,f984,f991,f994,f1000,f1001,f1007,f1012,f1023,f1024,f1025,f1032,f1033,f1035,f1047,f1048,f1050,f1057,f1059,f1062,f1068,f1069,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1210,f1212,f1217,f1219,f1230,f1231,f1239,f1247,f1249,f1254,f1256,f1267,f1275,f1283,f1285,f1339,f1333,f1332,f1373,f1366,f1365,f1385,f1411,f1410,f1403,f1402,f1422,f1433,f1437,f1461,f1486,f1492,f1494,f1501,f1512,f1516,f1542,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1663,f1662,f1661,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1818,f1819,f1816,f1837,f1838,f1835,f1852,f1854,f1859,f1861,f1931,f2001,f1981,f2015,f2030,f2041,f2043,f2050,f2052,f2073,f2082,f2088,f2090,f2097,f2099,f2119,f2118,f2128,f2157,f2170,f2251,f2252,f2222,f2262,f2269,f2270,f2271,f2272,f2317,f2318,f2319,f2351,f2352,f2349,f2371,f2372,f2373,f2400,f2401,f2398,f2432,f2433,f2465,f2466,f2504,f2505,f2544,f2545,f2617,f2622,f2765,f2770,f2841,f2844,f2847,f2848,f2751,f2842,f2996,f2855,f2860,f2861,f2862,f2863,f2865,f2866,f2867,f2868,f2869,f2870,f2871,f2872,f3012,f3014,f2875,f2876,f2878,f2880,f2882,f2883,f2885])).
fof(f3021,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f334,f333,f335,f592,f343,f355,f354,f359,f358,f412,f411,f413,f414,f420,f419,f418,f417,f424,f423,f422,f421,f428,f427,f426,f429,f430,f431,f432,f433,f437,f434,f444,f443,f593,f462,f459,f458,f457,f468,f467,f477,f478,f479,f480,f481,f482,f483,f489,f488,f487,f485,f484,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f787,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f818,f819,f825,f828,f829,f830,f839,f838,f837,f836,f840,f852,f861,f862,f863,f864,f865,f866,f867,f868,f869,f870,f871,f873,f888,f889,f891,f898,f901,f902,f904,f917,f923,f924,f926,f930,f933,f939,f940,f942,f946,f954,f960,f961,f966,f971,f975,f980,f984,f991,f994,f1000,f1001,f1007,f1012,f1023,f1024,f1025,f1032,f1033,f1035,f1047,f1048,f1050,f1057,f1059,f1062,f1068,f1069,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1210,f1212,f1217,f1219,f1230,f1231,f1239,f1247,f1249,f1254,f1256,f1267,f1275,f1283,f1285,f1339,f1333,f1332,f1373,f1366,f1365,f1385,f1411,f1410,f1403,f1402,f1422,f1433,f1437,f1461,f1486,f1492,f1494,f1501,f1512,f1516,f1542,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1663,f1662,f1661,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1818,f1819,f1816,f1837,f1838,f1835,f1852,f1854,f1859,f1861,f1931,f2001,f1981,f2015,f2030,f2041,f2043,f2050,f2052,f2073,f2082,f2088,f2090,f2097,f2099,f2119,f2118,f2128,f2157,f2170,f2251,f2252,f2222,f2262,f2269,f2270,f2271,f2272,f2317,f2318,f2319,f2351,f2352,f2349,f2371,f2372,f2373,f2400,f2401,f2398,f2432,f2433,f2465,f2466,f2504,f2505,f2544,f2545,f2617,f2622,f2765,f2770,f2841,f2844,f2847,f2848,f2751,f2842,f2996,f2855,f2860,f2861,f2862,f2863,f2865,f2866,f2867,f2868,f2869,f2870,f2871,f2872,f3012,f3014,f2875,f2876,f2878,f2880,f2882,f2883])).
fof(f3020,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f334,f333,f335,f592,f343,f355,f354,f359,f358,f412,f411,f413,f414,f420,f419,f418,f417,f424,f423,f422,f421,f428,f427,f426,f429,f430,f431,f432,f433,f437,f434,f444,f443,f593,f462,f459,f458,f457,f468,f467,f477,f478,f479,f480,f481,f482,f483,f489,f488,f487,f485,f484,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f787,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f818,f819,f825,f828,f829,f830,f839,f838,f837,f836,f840,f852,f861,f862,f863,f864,f865,f866,f867,f868,f869,f870,f871,f873,f888,f889,f891,f898,f901,f902,f904,f917,f923,f924,f926,f930,f933,f939,f940,f942,f946,f954,f960,f961,f966,f971,f975,f980,f984,f991,f994,f1000,f1001,f1007,f1012,f1023,f1024,f1025,f1032,f1033,f1035,f1047,f1048,f1050,f1057,f1059,f1062,f1068,f1069,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1210,f1212,f1217,f1219,f1230,f1231,f1239,f1247,f1249,f1254,f1256,f1267,f1275,f1283,f1285,f1339,f1333,f1332,f1373,f1366,f1365,f1385,f1411,f1410,f1403,f1402,f1422,f1433,f1437,f1461,f1486,f1492,f1494,f1501,f1512,f1516,f1542,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1663,f1662,f1661,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1818,f1819,f1816,f1837,f1838,f1835,f1852,f1854,f1859,f1861,f1931,f2001,f1981,f2015,f2030,f2041,f2043,f2050,f2052,f2073,f2082,f2088,f2090,f2097,f2099,f2119,f2118,f2128,f2157,f2170,f2251,f2252,f2222,f2262,f2269,f2270,f2271,f2272,f2317,f2318,f2319,f2351,f2352,f2349,f2371,f2372,f2373,f2400,f2401,f2398,f2432,f2433,f2465,f2466,f2504,f2505,f2544,f2545,f2617,f2622,f2765,f2770,f2841,f2844,f2847,f2848,f2751,f2842,f2996,f2855,f2860,f2861,f2862,f2863,f2865,f2866,f2867,f2868,f2869,f2870,f2871,f2872,f3012,f3014,f2875,f2876,f2878,f2880,f2882])).
fof(f3019,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f334,f333,f335,f592,f343,f355,f354,f359,f358,f412,f411,f413,f414,f420,f419,f418,f417,f424,f423,f422,f421,f428,f427,f426,f429,f430,f431,f432,f433,f437,f434,f444,f443,f593,f462,f459,f458,f457,f468,f467,f477,f478,f479,f480,f481,f482,f483,f489,f488,f487,f485,f484,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f787,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f818,f819,f825,f828,f829,f830,f839,f838,f837,f836,f840,f852,f861,f862,f863,f864,f865,f866,f867,f868,f869,f870,f871,f873,f888,f889,f891,f898,f901,f902,f904,f917,f923,f924,f926,f930,f933,f939,f940,f942,f946,f954,f960,f961,f966,f971,f975,f980,f984,f991,f994,f1000,f1001,f1007,f1012,f1023,f1024,f1025,f1032,f1033,f1035,f1047,f1048,f1050,f1057,f1059,f1062,f1068,f1069,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1210,f1212,f1217,f1219,f1230,f1231,f1239,f1247,f1249,f1254,f1256,f1267,f1275,f1283,f1285,f1339,f1333,f1332,f1373,f1366,f1365,f1385,f1411,f1410,f1403,f1402,f1422,f1433,f1437,f1461,f1486,f1492,f1494,f1501,f1512,f1516,f1542,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1663,f1662,f1661,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1818,f1819,f1816,f1837,f1838,f1835,f1852,f1854,f1859,f1861,f1931,f2001,f1981,f2015,f2030,f2041,f2043,f2050,f2052,f2073,f2082,f2088,f2090,f2097,f2099,f2119,f2118,f2128,f2157,f2170,f2251,f2252,f2222,f2262,f2269,f2270,f2271,f2272,f2317,f2318,f2319,f2351,f2352,f2349,f2371,f2372,f2373,f2400,f2401,f2398,f2432,f2433,f2465,f2466,f2504,f2505,f2544,f2545,f2617,f2622,f2765,f2770,f2841,f2844,f2847,f2848,f2751,f2842,f2996,f2855,f2860,f2861,f2862,f2863,f2865,f2866,f2867,f2868,f2869,f2870,f2871,f2872,f3012,f3014,f2875,f2876,f2878,f2880])).
fof(f3018,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f334,f333,f335,f592,f343,f355,f354,f359,f358,f412,f411,f413,f414,f420,f419,f418,f417,f424,f423,f422,f421,f428,f427,f426,f429,f430,f431,f432,f433,f437,f434,f444,f443,f593,f462,f459,f458,f457,f468,f467,f477,f478,f479,f480,f481,f482,f483,f489,f488,f487,f485,f484,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f787,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f818,f819,f825,f828,f829,f830,f839,f838,f837,f836,f840,f852,f861,f862,f863,f864,f865,f866,f867,f868,f869,f870,f871,f873,f888,f889,f891,f898,f901,f902,f904,f917,f923,f924,f926,f930,f933,f939,f940,f942,f946,f954,f960,f961,f966,f971,f975,f980,f984,f991,f994,f1000,f1001,f1007,f1012,f1023,f1024,f1025,f1032,f1033,f1035,f1047,f1048,f1050,f1057,f1059,f1062,f1068,f1069,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1210,f1212,f1217,f1219,f1230,f1231,f1239,f1247,f1249,f1254,f1256,f1267,f1275,f1283,f1285,f1339,f1333,f1332,f1373,f1366,f1365,f1385,f1411,f1410,f1403,f1402,f1422,f1433,f1437,f1461,f1486,f1492,f1494,f1501,f1512,f1516,f1542,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1663,f1662,f1661,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1818,f1819,f1816,f1837,f1838,f1835,f1852,f1854,f1859,f1861,f1931,f2001,f1981,f2015,f2030,f2041,f2043,f2050,f2052,f2073,f2082,f2088,f2090,f2097,f2099,f2119,f2118,f2128,f2157,f2170,f2251,f2252,f2222,f2262,f2269,f2270,f2271,f2272,f2317,f2318,f2319,f2351,f2352,f2349,f2371,f2372,f2373,f2400,f2401,f2398,f2432,f2433,f2465,f2466,f2504,f2505,f2544,f2545,f2617,f2622,f2765,f2770,f2841,f2844,f2847,f2848,f2751,f2842,f2996,f2855,f2860,f2861,f2862,f2863,f2865,f2866,f2867,f2868,f2869,f2870,f2871,f2872,f3012,f3014,f2875,f2876,f2878])).
fof(f3017,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f334,f333,f335,f592,f343,f355,f354,f359,f358,f412,f411,f413,f414,f420,f419,f418,f417,f424,f423,f422,f421,f428,f427,f426,f429,f430,f431,f432,f433,f437,f434,f444,f443,f593,f462,f459,f458,f457,f468,f467,f477,f478,f479,f480,f481,f482,f483,f489,f488,f487,f485,f484,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f787,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f818,f819,f825,f828,f829,f830,f839,f838,f837,f836,f840,f852,f861,f862,f863,f864,f865,f866,f867,f868,f869,f870,f871,f873,f888,f889,f891,f898,f901,f902,f904,f917,f923,f924,f926,f930,f933,f939,f940,f942,f946,f954,f960,f961,f966,f971,f975,f980,f984,f991,f994,f1000,f1001,f1007,f1012,f1023,f1024,f1025,f1032,f1033,f1035,f1047,f1048,f1050,f1057,f1059,f1062,f1068,f1069,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1210,f1212,f1217,f1219,f1230,f1231,f1239,f1247,f1249,f1254,f1256,f1267,f1275,f1283,f1285,f1339,f1333,f1332,f1373,f1366,f1365,f1385,f1411,f1410,f1403,f1402,f1422,f1433,f1437,f1461,f1486,f1492,f1494,f1501,f1512,f1516,f1542,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1663,f1662,f1661,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1818,f1819,f1816,f1837,f1838,f1835,f1852,f1854,f1859,f1861,f1931,f2001,f1981,f2015,f2030,f2041,f2043,f2050,f2052,f2073,f2082,f2088,f2090,f2097,f2099,f2119,f2118,f2128,f2157,f2170,f2251,f2252,f2222,f2262,f2269,f2270,f2271,f2272,f2317,f2318,f2319,f2351,f2352,f2349,f2371,f2372,f2373,f2400,f2401,f2398,f2432,f2433,f2465,f2466,f2504,f2505,f2544,f2545,f2617,f2622,f2765,f2770,f2841,f2844,f2847,f2848,f2751,f2842,f2996,f2855,f2860,f2861,f2862,f2863,f2865,f2866,f2867,f2868,f2869,f2870,f2871,f2872,f3012,f3014,f2875,f2876])).
fof(f3016,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f334,f333,f335,f592,f343,f355,f354,f359,f358,f412,f411,f413,f414,f420,f419,f418,f417,f424,f423,f422,f421,f428,f427,f426,f429,f430,f431,f432,f433,f437,f434,f444,f443,f593,f462,f459,f458,f457,f468,f467,f477,f478,f479,f480,f481,f482,f483,f489,f488,f487,f485,f484,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f787,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f818,f819,f825,f828,f829,f830,f839,f838,f837,f836,f840,f852,f861,f862,f863,f864,f865,f866,f867,f868,f869,f870,f871,f873,f888,f889,f891,f898,f901,f902,f904,f917,f923,f924,f926,f930,f933,f939,f940,f942,f946,f954,f960,f961,f966,f971,f975,f980,f984,f991,f994,f1000,f1001,f1007,f1012,f1023,f1024,f1025,f1032,f1033,f1035,f1047,f1048,f1050,f1057,f1059,f1062,f1068,f1069,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1210,f1212,f1217,f1219,f1230,f1231,f1239,f1247,f1249,f1254,f1256,f1267,f1275,f1283,f1285,f1339,f1333,f1332,f1373,f1366,f1365,f1385,f1411,f1410,f1403,f1402,f1422,f1433,f1437,f1461,f1486,f1492,f1494,f1501,f1512,f1516,f1542,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1663,f1662,f1661,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1818,f1819,f1816,f1837,f1838,f1835,f1852,f1854,f1859,f1861,f1931,f2001,f1981,f2015,f2030,f2041,f2043,f2050,f2052,f2073,f2082,f2088,f2090,f2097,f2099,f2119,f2118,f2128,f2157,f2170,f2251,f2252,f2222,f2262,f2269,f2270,f2271,f2272,f2317,f2318,f2319,f2351,f2352,f2349,f2371,f2372,f2373,f2400,f2401,f2398,f2432,f2433,f2465,f2466,f2504,f2505,f2544,f2545,f2617,f2622,f2765,f2770,f2841,f2844,f2847,f2848,f2751,f2842,f2996,f2855,f2860,f2861,f2862,f2863,f2865,f2866,f2867,f2868,f2869,f2870,f2871,f2872,f3012,f3014,f2875])).
fof(f3015,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f334,f333,f335,f592,f343,f355,f354,f359,f358,f412,f411,f413,f414,f420,f419,f418,f417,f424,f423,f422,f421,f428,f427,f426,f429,f430,f431,f432,f433,f437,f434,f444,f443,f593,f462,f459,f458,f457,f468,f467,f477,f478,f479,f480,f481,f482,f483,f489,f488,f487,f485,f484,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f787,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f818,f819,f825,f828,f829,f830,f839,f838,f837,f836,f840,f852,f861,f862,f863,f864,f865,f866,f867,f868,f869,f870,f871,f873,f888,f889,f891,f898,f901,f902,f904,f917,f923,f924,f926,f930,f933,f939,f940,f942,f946,f954,f960,f961,f966,f971,f975,f980,f984,f991,f994,f1000,f1001,f1007,f1012,f1023,f1024,f1025,f1032,f1033,f1035,f1047,f1048,f1050,f1057,f1059,f1062,f1068,f1069,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1210,f1212,f1217,f1219,f1230,f1231,f1239,f1247,f1249,f1254,f1256,f1267,f1275,f1283,f1285,f1339,f1333,f1332,f1373,f1366,f1365,f1385,f1411,f1410,f1403,f1402,f1422,f1433,f1437,f1461,f1486,f1492,f1494,f1501,f1512,f1516,f1542,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1663,f1662,f1661,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1818,f1819,f1816,f1837,f1838,f1835,f1852,f1854,f1859,f1861,f1931,f2001,f1981,f2015,f2030,f2041,f2043,f2050,f2052,f2073,f2082,f2088,f2090,f2097,f2099,f2119,f2118,f2128,f2157,f2170,f2251,f2252,f2222,f2262,f2269,f2270,f2271,f2272,f2317,f2318,f2319,f2351,f2352,f2349,f2371,f2372,f2373,f2400,f2401,f2398,f2432,f2433,f2465,f2466,f2504,f2505,f2544,f2545,f2617,f2622,f2765,f2770,f2841,f2844,f2847,f2848,f2751,f2842,f2996,f2855,f2860,f2861,f2862,f2863,f2865,f2866,f2867,f2868,f2869,f2870,f2871,f2872,f3012,f3014])).
fof(f3013,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f334,f333,f335,f592,f343,f355,f354,f359,f358,f412,f411,f413,f414,f420,f419,f418,f417,f424,f423,f422,f421,f428,f427,f426,f429,f430,f431,f432,f433,f437,f434,f444,f443,f593,f462,f459,f458,f457,f468,f467,f477,f478,f479,f480,f481,f482,f483,f489,f488,f487,f485,f484,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f787,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f818,f819,f825,f828,f829,f830,f839,f838,f837,f836,f840,f852,f861,f862,f863,f864,f865,f866,f867,f868,f869,f870,f871,f873,f888,f889,f891,f898,f901,f902,f904,f917,f923,f924,f926,f930,f933,f939,f940,f942,f946,f954,f960,f961,f966,f971,f975,f980,f984,f991,f994,f1000,f1001,f1007,f1012,f1023,f1024,f1025,f1032,f1033,f1035,f1047,f1048,f1050,f1057,f1059,f1062,f1068,f1069,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1210,f1212,f1217,f1219,f1230,f1231,f1239,f1247,f1249,f1254,f1256,f1267,f1275,f1283,f1285,f1339,f1333,f1332,f1373,f1366,f1365,f1385,f1411,f1410,f1403,f1402,f1422,f1433,f1437,f1461,f1486,f1492,f1494,f1501,f1512,f1516,f1542,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1663,f1662,f1661,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1818,f1819,f1816,f1837,f1838,f1835,f1852,f1854,f1859,f1861,f1931,f2001,f1981,f2015,f2030,f2041,f2043,f2050,f2052,f2073,f2082,f2088,f2090,f2097,f2099,f2119,f2118,f2128,f2157,f2170,f2251,f2252,f2222,f2262,f2269,f2270,f2271,f2272,f2317,f2318,f2319,f2351,f2352,f2349,f2371,f2372,f2373,f2400,f2401,f2398,f2432,f2433,f2465,f2466,f2504,f2505,f2544,f2545,f2617,f2622,f2765,f2770,f2841,f2844,f2847,f2848,f2751,f2842,f2996,f2855,f2860,f2861,f2862,f2863,f2865,f2866,f2867,f2868,f2869,f2870,f2871,f2872,f3012])).
fof(f3011,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f334,f333,f335,f592,f343,f355,f354,f359,f358,f412,f411,f413,f414,f420,f419,f418,f417,f424,f423,f422,f421,f428,f427,f426,f429,f430,f431,f432,f433,f437,f434,f444,f443,f593,f462,f459,f458,f457,f468,f467,f477,f478,f479,f480,f481,f482,f483,f489,f488,f487,f485,f484,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f787,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f818,f819,f825,f828,f829,f830,f839,f838,f837,f836,f840,f852,f861,f862,f863,f864,f865,f866,f867,f868,f869,f870,f871,f873,f888,f889,f891,f898,f901,f902,f904,f917,f923,f924,f926,f930,f933,f939,f940,f942,f946,f954,f960,f961,f966,f971,f975,f980,f984,f991,f994,f1000,f1001,f1007,f1012,f1023,f1024,f1025,f1032,f1033,f1035,f1047,f1048,f1050,f1057,f1059,f1062,f1068,f1069,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1210,f1212,f1217,f1219,f1230,f1231,f1239,f1247,f1249,f1254,f1256,f1267,f1275,f1283,f1285,f1339,f1333,f1332,f1373,f1366,f1365,f1385,f1411,f1410,f1403,f1402,f1422,f1433,f1437,f1461,f1486,f1492,f1494,f1501,f1512,f1516,f1542,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1663,f1662,f1661,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1818,f1819,f1816,f1837,f1838,f1835,f1852,f1854,f1859,f1861,f1931,f2001,f1981,f2015,f2030,f2041,f2043,f2050,f2052,f2073,f2082,f2088,f2090,f2097,f2099,f2119,f2118,f2128,f2157,f2170,f2251,f2252,f2222,f2262,f2269,f2270,f2271,f2272,f2317,f2318,f2319,f2351,f2352,f2349,f2371,f2372,f2373,f2400,f2401,f2398,f2432,f2433,f2465,f2466,f2504,f2505,f2544,f2545,f2617,f2622,f2765,f2770,f2841,f2844,f2847,f2848,f2751,f2842,f2996,f2855,f2860,f2861,f2862,f2863,f2865,f2866,f2867,f2868,f2869,f2870,f2871,f2872])).
fof(f3010,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f334,f333,f335,f592,f343,f355,f354,f359,f358,f412,f411,f413,f414,f420,f419,f418,f417,f424,f423,f422,f421,f428,f427,f426,f429,f430,f431,f432,f433,f437,f434,f444,f443,f593,f462,f459,f458,f457,f468,f467,f477,f478,f479,f480,f481,f482,f483,f489,f488,f487,f485,f484,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f787,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f818,f819,f825,f828,f829,f830,f839,f838,f837,f836,f840,f852,f861,f862,f863,f864,f865,f866,f867,f868,f869,f870,f871,f873,f888,f889,f891,f898,f901,f902,f904,f917,f923,f924,f926,f930,f933,f939,f940,f942,f946,f954,f960,f961,f966,f971,f975,f980,f984,f991,f994,f1000,f1001,f1007,f1012,f1023,f1024,f1025,f1032,f1033,f1035,f1047,f1048,f1050,f1057,f1059,f1062,f1068,f1069,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1210,f1212,f1217,f1219,f1230,f1231,f1239,f1247,f1249,f1254,f1256,f1267,f1275,f1283,f1285,f1339,f1333,f1332,f1373,f1366,f1365,f1385,f1411,f1410,f1403,f1402,f1422,f1433,f1437,f1461,f1486,f1492,f1494,f1501,f1512,f1516,f1542,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1663,f1662,f1661,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1818,f1819,f1816,f1837,f1838,f1835,f1852,f1854,f1859,f1861,f1931,f2001,f1981,f2015,f2030,f2041,f2043,f2050,f2052,f2073,f2082,f2088,f2090,f2097,f2099,f2119,f2118,f2128,f2157,f2170,f2251,f2252,f2222,f2262,f2269,f2270,f2271,f2272,f2317,f2318,f2319,f2351,f2352,f2349,f2371,f2372,f2373,f2400,f2401,f2398,f2432,f2433,f2465,f2466,f2504,f2505,f2544,f2545,f2617,f2622,f2765,f2770,f2841,f2844,f2847,f2848,f2751,f2842,f2996,f2855,f2860,f2861,f2862,f2863,f2865,f2866,f2867,f2868,f2869,f2870,f2871])).
fof(f3009,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f334,f333,f335,f592,f343,f355,f354,f359,f358,f412,f411,f413,f414,f420,f419,f418,f417,f424,f423,f422,f421,f428,f427,f426,f429,f430,f431,f432,f433,f437,f434,f444,f443,f593,f462,f459,f458,f457,f468,f467,f477,f478,f479,f480,f481,f482,f483,f489,f488,f487,f485,f484,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f787,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f818,f819,f825,f828,f829,f830,f839,f838,f837,f836,f840,f852,f861,f862,f863,f864,f865,f866,f867,f868,f869,f870,f871,f873,f888,f889,f891,f898,f901,f902,f904,f917,f923,f924,f926,f930,f933,f939,f940,f942,f946,f954,f960,f961,f966,f971,f975,f980,f984,f991,f994,f1000,f1001,f1007,f1012,f1023,f1024,f1025,f1032,f1033,f1035,f1047,f1048,f1050,f1057,f1059,f1062,f1068,f1069,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1210,f1212,f1217,f1219,f1230,f1231,f1239,f1247,f1249,f1254,f1256,f1267,f1275,f1283,f1285,f1339,f1333,f1332,f1373,f1366,f1365,f1385,f1411,f1410,f1403,f1402,f1422,f1433,f1437,f1461,f1486,f1492,f1494,f1501,f1512,f1516,f1542,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1663,f1662,f1661,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1818,f1819,f1816,f1837,f1838,f1835,f1852,f1854,f1859,f1861,f1931,f2001,f1981,f2015,f2030,f2041,f2043,f2050,f2052,f2073,f2082,f2088,f2090,f2097,f2099,f2119,f2118,f2128,f2157,f2170,f2251,f2252,f2222,f2262,f2269,f2270,f2271,f2272,f2317,f2318,f2319,f2351,f2352,f2349,f2371,f2372,f2373,f2400,f2401,f2398,f2432,f2433,f2465,f2466,f2504,f2505,f2544,f2545,f2617,f2622,f2765,f2770,f2841,f2844,f2847,f2848,f2751,f2842,f2996,f2855,f2860,f2861,f2862,f2863,f2865,f2866,f2867,f2868,f2869,f2870])).
fof(f3008,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f334,f333,f335,f592,f343,f355,f354,f359,f358,f412,f411,f413,f414,f420,f419,f418,f417,f424,f423,f422,f421,f428,f427,f426,f429,f430,f431,f432,f433,f437,f434,f444,f443,f593,f462,f459,f458,f457,f468,f467,f477,f478,f479,f480,f481,f482,f483,f489,f488,f487,f485,f484,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f787,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f818,f819,f825,f828,f829,f830,f839,f838,f837,f836,f840,f852,f861,f862,f863,f864,f865,f866,f867,f868,f869,f870,f871,f873,f888,f889,f891,f898,f901,f902,f904,f917,f923,f924,f926,f930,f933,f939,f940,f942,f946,f954,f960,f961,f966,f971,f975,f980,f984,f991,f994,f1000,f1001,f1007,f1012,f1023,f1024,f1025,f1032,f1033,f1035,f1047,f1048,f1050,f1057,f1059,f1062,f1068,f1069,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1210,f1212,f1217,f1219,f1230,f1231,f1239,f1247,f1249,f1254,f1256,f1267,f1275,f1283,f1285,f1339,f1333,f1332,f1373,f1366,f1365,f1385,f1411,f1410,f1403,f1402,f1422,f1433,f1437,f1461,f1486,f1492,f1494,f1501,f1512,f1516,f1542,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1663,f1662,f1661,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1818,f1819,f1816,f1837,f1838,f1835,f1852,f1854,f1859,f1861,f1931,f2001,f1981,f2015,f2030,f2041,f2043,f2050,f2052,f2073,f2082,f2088,f2090,f2097,f2099,f2119,f2118,f2128,f2157,f2170,f2251,f2252,f2222,f2262,f2269,f2270,f2271,f2272,f2317,f2318,f2319,f2351,f2352,f2349,f2371,f2372,f2373,f2400,f2401,f2398,f2432,f2433,f2465,f2466,f2504,f2505,f2544,f2545,f2617,f2622,f2765,f2770,f2841,f2844,f2847,f2848,f2751,f2842,f2996,f2855,f2860,f2861,f2862,f2863,f2865,f2866,f2867,f2868,f2869])).
fof(f3007,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f334,f333,f335,f592,f343,f355,f354,f359,f358,f412,f411,f413,f414,f420,f419,f418,f417,f424,f423,f422,f421,f428,f427,f426,f429,f430,f431,f432,f433,f437,f434,f444,f443,f593,f462,f459,f458,f457,f468,f467,f477,f478,f479,f480,f481,f482,f483,f489,f488,f487,f485,f484,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f787,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f818,f819,f825,f828,f829,f830,f839,f838,f837,f836,f840,f852,f861,f862,f863,f864,f865,f866,f867,f868,f869,f870,f871,f873,f888,f889,f891,f898,f901,f902,f904,f917,f923,f924,f926,f930,f933,f939,f940,f942,f946,f954,f960,f961,f966,f971,f975,f980,f984,f991,f994,f1000,f1001,f1007,f1012,f1023,f1024,f1025,f1032,f1033,f1035,f1047,f1048,f1050,f1057,f1059,f1062,f1068,f1069,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1210,f1212,f1217,f1219,f1230,f1231,f1239,f1247,f1249,f1254,f1256,f1267,f1275,f1283,f1285,f1339,f1333,f1332,f1373,f1366,f1365,f1385,f1411,f1410,f1403,f1402,f1422,f1433,f1437,f1461,f1486,f1492,f1494,f1501,f1512,f1516,f1542,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1663,f1662,f1661,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1818,f1819,f1816,f1837,f1838,f1835,f1852,f1854,f1859,f1861,f1931,f2001,f1981,f2015,f2030,f2041,f2043,f2050,f2052,f2073,f2082,f2088,f2090,f2097,f2099,f2119,f2118,f2128,f2157,f2170,f2251,f2252,f2222,f2262,f2269,f2270,f2271,f2272,f2317,f2318,f2319,f2351,f2352,f2349,f2371,f2372,f2373,f2400,f2401,f2398,f2432,f2433,f2465,f2466,f2504,f2505,f2544,f2545,f2617,f2622,f2765,f2770,f2841,f2844,f2847,f2848,f2751,f2842,f2996,f2855,f2860,f2861,f2862,f2863,f2865,f2866,f2867,f2868])).
fof(f3006,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f334,f333,f335,f592,f343,f355,f354,f359,f358,f412,f411,f413,f414,f420,f419,f418,f417,f424,f423,f422,f421,f428,f427,f426,f429,f430,f431,f432,f433,f437,f434,f444,f443,f593,f462,f459,f458,f457,f468,f467,f477,f478,f479,f480,f481,f482,f483,f489,f488,f487,f485,f484,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f787,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f818,f819,f825,f828,f829,f830,f839,f838,f837,f836,f840,f852,f861,f862,f863,f864,f865,f866,f867,f868,f869,f870,f871,f873,f888,f889,f891,f898,f901,f902,f904,f917,f923,f924,f926,f930,f933,f939,f940,f942,f946,f954,f960,f961,f966,f971,f975,f980,f984,f991,f994,f1000,f1001,f1007,f1012,f1023,f1024,f1025,f1032,f1033,f1035,f1047,f1048,f1050,f1057,f1059,f1062,f1068,f1069,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1210,f1212,f1217,f1219,f1230,f1231,f1239,f1247,f1249,f1254,f1256,f1267,f1275,f1283,f1285,f1339,f1333,f1332,f1373,f1366,f1365,f1385,f1411,f1410,f1403,f1402,f1422,f1433,f1437,f1461,f1486,f1492,f1494,f1501,f1512,f1516,f1542,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1663,f1662,f1661,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1818,f1819,f1816,f1837,f1838,f1835,f1852,f1854,f1859,f1861,f1931,f2001,f1981,f2015,f2030,f2041,f2043,f2050,f2052,f2073,f2082,f2088,f2090,f2097,f2099,f2119,f2118,f2128,f2157,f2170,f2251,f2252,f2222,f2262,f2269,f2270,f2271,f2272,f2317,f2318,f2319,f2351,f2352,f2349,f2371,f2372,f2373,f2400,f2401,f2398,f2432,f2433,f2465,f2466,f2504,f2505,f2544,f2545,f2617,f2622,f2765,f2770,f2841,f2844,f2847,f2848,f2751,f2842,f2996,f2855,f2860,f2861,f2862,f2863,f2865,f2866,f2867])).
fof(f3005,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f334,f333,f335,f592,f343,f355,f354,f359,f358,f412,f411,f413,f414,f420,f419,f418,f417,f424,f423,f422,f421,f428,f427,f426,f429,f430,f431,f432,f433,f437,f434,f444,f443,f593,f462,f459,f458,f457,f468,f467,f477,f478,f479,f480,f481,f482,f483,f489,f488,f487,f485,f484,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f787,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f818,f819,f825,f828,f829,f830,f839,f838,f837,f836,f840,f852,f861,f862,f863,f864,f865,f866,f867,f868,f869,f870,f871,f873,f888,f889,f891,f898,f901,f902,f904,f917,f923,f924,f926,f930,f933,f939,f940,f942,f946,f954,f960,f961,f966,f971,f975,f980,f984,f991,f994,f1000,f1001,f1007,f1012,f1023,f1024,f1025,f1032,f1033,f1035,f1047,f1048,f1050,f1057,f1059,f1062,f1068,f1069,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1210,f1212,f1217,f1219,f1230,f1231,f1239,f1247,f1249,f1254,f1256,f1267,f1275,f1283,f1285,f1339,f1333,f1332,f1373,f1366,f1365,f1385,f1411,f1410,f1403,f1402,f1422,f1433,f1437,f1461,f1486,f1492,f1494,f1501,f1512,f1516,f1542,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1663,f1662,f1661,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1818,f1819,f1816,f1837,f1838,f1835,f1852,f1854,f1859,f1861,f1931,f2001,f1981,f2015,f2030,f2041,f2043,f2050,f2052,f2073,f2082,f2088,f2090,f2097,f2099,f2119,f2118,f2128,f2157,f2170,f2251,f2252,f2222,f2262,f2269,f2270,f2271,f2272,f2317,f2318,f2319,f2351,f2352,f2349,f2371,f2372,f2373,f2400,f2401,f2398,f2432,f2433,f2465,f2466,f2504,f2505,f2544,f2545,f2617,f2622,f2765,f2770,f2841,f2844,f2847,f2848,f2751,f2842,f2996,f2855,f2860,f2861,f2862,f2863,f2865,f2866])).
fof(f3004,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f334,f333,f335,f592,f343,f355,f354,f359,f358,f412,f411,f413,f414,f420,f419,f418,f417,f424,f423,f422,f421,f428,f427,f426,f429,f430,f431,f432,f433,f437,f434,f444,f443,f593,f462,f459,f458,f457,f468,f467,f477,f478,f479,f480,f481,f482,f483,f489,f488,f487,f485,f484,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f787,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f818,f819,f825,f828,f829,f830,f839,f838,f837,f836,f840,f852,f861,f862,f863,f864,f865,f866,f867,f868,f869,f870,f871,f873,f888,f889,f891,f898,f901,f902,f904,f917,f923,f924,f926,f930,f933,f939,f940,f942,f946,f954,f960,f961,f966,f971,f975,f980,f984,f991,f994,f1000,f1001,f1007,f1012,f1023,f1024,f1025,f1032,f1033,f1035,f1047,f1048,f1050,f1057,f1059,f1062,f1068,f1069,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1210,f1212,f1217,f1219,f1230,f1231,f1239,f1247,f1249,f1254,f1256,f1267,f1275,f1283,f1285,f1339,f1333,f1332,f1373,f1366,f1365,f1385,f1411,f1410,f1403,f1402,f1422,f1433,f1437,f1461,f1486,f1492,f1494,f1501,f1512,f1516,f1542,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1663,f1662,f1661,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1818,f1819,f1816,f1837,f1838,f1835,f1852,f1854,f1859,f1861,f1931,f2001,f1981,f2015,f2030,f2041,f2043,f2050,f2052,f2073,f2082,f2088,f2090,f2097,f2099,f2119,f2118,f2128,f2157,f2170,f2251,f2252,f2222,f2262,f2269,f2270,f2271,f2272,f2317,f2318,f2319,f2351,f2352,f2349,f2371,f2372,f2373,f2400,f2401,f2398,f2432,f2433,f2465,f2466,f2504,f2505,f2544,f2545,f2617,f2622,f2765,f2770,f2841,f2844,f2847,f2848,f2751,f2842,f2996,f2855,f2860,f2861,f2862,f2863,f2865])).
fof(f3003,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f334,f333,f335,f592,f343,f355,f354,f359,f358,f412,f411,f413,f414,f420,f419,f418,f417,f424,f423,f422,f421,f428,f427,f426,f429,f430,f431,f432,f433,f437,f434,f444,f443,f593,f462,f459,f458,f457,f468,f467,f477,f478,f479,f480,f481,f482,f483,f489,f488,f487,f485,f484,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f787,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f818,f819,f825,f828,f829,f830,f839,f838,f837,f836,f840,f852,f861,f862,f863,f864,f865,f866,f867,f868,f869,f870,f871,f873,f888,f889,f891,f898,f901,f902,f904,f917,f923,f924,f926,f930,f933,f939,f940,f942,f946,f954,f960,f961,f966,f971,f975,f980,f984,f991,f994,f1000,f1001,f1007,f1012,f1023,f1024,f1025,f1032,f1033,f1035,f1047,f1048,f1050,f1057,f1059,f1062,f1068,f1069,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1210,f1212,f1217,f1219,f1230,f1231,f1239,f1247,f1249,f1254,f1256,f1267,f1275,f1283,f1285,f1339,f1333,f1332,f1373,f1366,f1365,f1385,f1411,f1410,f1403,f1402,f1422,f1433,f1437,f1461,f1486,f1492,f1494,f1501,f1512,f1516,f1542,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1663,f1662,f1661,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1818,f1819,f1816,f1837,f1838,f1835,f1852,f1854,f1859,f1861,f1931,f2001,f1981,f2015,f2030,f2041,f2043,f2050,f2052,f2073,f2082,f2088,f2090,f2097,f2099,f2119,f2118,f2128,f2157,f2170,f2251,f2252,f2222,f2262,f2269,f2270,f2271,f2272,f2317,f2318,f2319,f2351,f2352,f2349,f2371,f2372,f2373,f2400,f2401,f2398,f2432,f2433,f2465,f2466,f2504,f2505,f2544,f2545,f2617,f2622,f2765,f2770,f2841,f2844,f2847,f2848,f2751,f2842,f2996,f2855,f2860,f2861,f2862,f2863])).
fof(f3002,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f334,f333,f335,f592,f343,f355,f354,f359,f358,f412,f411,f413,f414,f420,f419,f418,f417,f424,f423,f422,f421,f428,f427,f426,f429,f430,f431,f432,f433,f437,f434,f444,f443,f593,f462,f459,f458,f457,f468,f467,f477,f478,f479,f480,f481,f482,f483,f489,f488,f487,f485,f484,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f787,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f818,f819,f825,f828,f829,f830,f839,f838,f837,f836,f840,f852,f861,f862,f863,f864,f865,f866,f867,f868,f869,f870,f871,f873,f888,f889,f891,f898,f901,f902,f904,f917,f923,f924,f926,f930,f933,f939,f940,f942,f946,f954,f960,f961,f966,f971,f975,f980,f984,f991,f994,f1000,f1001,f1007,f1012,f1023,f1024,f1025,f1032,f1033,f1035,f1047,f1048,f1050,f1057,f1059,f1062,f1068,f1069,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1210,f1212,f1217,f1219,f1230,f1231,f1239,f1247,f1249,f1254,f1256,f1267,f1275,f1283,f1285,f1339,f1333,f1332,f1373,f1366,f1365,f1385,f1411,f1410,f1403,f1402,f1422,f1433,f1437,f1461,f1486,f1492,f1494,f1501,f1512,f1516,f1542,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1663,f1662,f1661,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1818,f1819,f1816,f1837,f1838,f1835,f1852,f1854,f1859,f1861,f1931,f2001,f1981,f2015,f2030,f2041,f2043,f2050,f2052,f2073,f2082,f2088,f2090,f2097,f2099,f2119,f2118,f2128,f2157,f2170,f2251,f2252,f2222,f2262,f2269,f2270,f2271,f2272,f2317,f2318,f2319,f2351,f2352,f2349,f2371,f2372,f2373,f2400,f2401,f2398,f2432,f2433,f2465,f2466,f2504,f2505,f2544,f2545,f2617,f2622,f2765,f2770,f2841,f2844,f2847,f2848,f2751,f2842,f2996,f2855,f2860,f2861,f2862])).
fof(f3001,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f334,f333,f335,f592,f343,f355,f354,f359,f358,f412,f411,f413,f414,f420,f419,f418,f417,f424,f423,f422,f421,f428,f427,f426,f429,f430,f431,f432,f433,f437,f434,f444,f443,f593,f462,f459,f458,f457,f468,f467,f477,f478,f479,f480,f481,f482,f483,f489,f488,f487,f485,f484,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f787,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f818,f819,f825,f828,f829,f830,f839,f838,f837,f836,f840,f852,f861,f862,f863,f864,f865,f866,f867,f868,f869,f870,f871,f873,f888,f889,f891,f898,f901,f902,f904,f917,f923,f924,f926,f930,f933,f939,f940,f942,f946,f954,f960,f961,f966,f971,f975,f980,f984,f991,f994,f1000,f1001,f1007,f1012,f1023,f1024,f1025,f1032,f1033,f1035,f1047,f1048,f1050,f1057,f1059,f1062,f1068,f1069,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1210,f1212,f1217,f1219,f1230,f1231,f1239,f1247,f1249,f1254,f1256,f1267,f1275,f1283,f1285,f1339,f1333,f1332,f1373,f1366,f1365,f1385,f1411,f1410,f1403,f1402,f1422,f1433,f1437,f1461,f1486,f1492,f1494,f1501,f1512,f1516,f1542,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1663,f1662,f1661,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1818,f1819,f1816,f1837,f1838,f1835,f1852,f1854,f1859,f1861,f1931,f2001,f1981,f2015,f2030,f2041,f2043,f2050,f2052,f2073,f2082,f2088,f2090,f2097,f2099,f2119,f2118,f2128,f2157,f2170,f2251,f2252,f2222,f2262,f2269,f2270,f2271,f2272,f2317,f2318,f2319,f2351,f2352,f2349,f2371,f2372,f2373,f2400,f2401,f2398,f2432,f2433,f2465,f2466,f2504,f2505,f2544,f2545,f2617,f2622,f2765,f2770,f2841,f2844,f2847,f2848,f2751,f2842,f2996,f2855,f2860,f2861])).
fof(f3000,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f334,f333,f335,f592,f343,f355,f354,f359,f358,f412,f411,f413,f414,f420,f419,f418,f417,f424,f423,f422,f421,f428,f427,f426,f429,f430,f431,f432,f433,f437,f434,f444,f443,f593,f462,f459,f458,f457,f468,f467,f477,f478,f479,f480,f481,f482,f483,f489,f488,f487,f485,f484,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f787,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f818,f819,f825,f828,f829,f830,f839,f838,f837,f836,f840,f852,f861,f862,f863,f864,f865,f866,f867,f868,f869,f870,f871,f873,f888,f889,f891,f898,f901,f902,f904,f917,f923,f924,f926,f930,f933,f939,f940,f942,f946,f954,f960,f961,f966,f971,f975,f980,f984,f991,f994,f1000,f1001,f1007,f1012,f1023,f1024,f1025,f1032,f1033,f1035,f1047,f1048,f1050,f1057,f1059,f1062,f1068,f1069,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1210,f1212,f1217,f1219,f1230,f1231,f1239,f1247,f1249,f1254,f1256,f1267,f1275,f1283,f1285,f1339,f1333,f1332,f1373,f1366,f1365,f1385,f1411,f1410,f1403,f1402,f1422,f1433,f1437,f1461,f1486,f1492,f1494,f1501,f1512,f1516,f1542,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1663,f1662,f1661,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1818,f1819,f1816,f1837,f1838,f1835,f1852,f1854,f1859,f1861,f1931,f2001,f1981,f2015,f2030,f2041,f2043,f2050,f2052,f2073,f2082,f2088,f2090,f2097,f2099,f2119,f2118,f2128,f2157,f2170,f2251,f2252,f2222,f2262,f2269,f2270,f2271,f2272,f2317,f2318,f2319,f2351,f2352,f2349,f2371,f2372,f2373,f2400,f2401,f2398,f2432,f2433,f2465,f2466,f2504,f2505,f2544,f2545,f2617,f2622,f2765,f2770,f2841,f2844,f2847,f2848,f2751,f2842,f2996,f2855,f2860])).
fof(f2999,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f2856,f335])).
fof(f2856,plain,(
  ~p(s(bool,'T_0')) | $spl4),
  inference(backward_demodulation,[],[f2842,f592])).
fof(f2998,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f334,f333,f335,f592,f343,f355,f354,f359,f358,f412,f411,f413,f414,f420,f419,f418,f417,f424,f423,f422,f421,f428,f427,f426,f429,f430,f431,f432,f433,f437,f434,f444,f443,f593,f462,f459,f458,f457,f468,f467,f477,f478,f479,f480,f481,f482,f483,f489,f488,f487,f485,f484,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f787,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f818,f819,f825,f828,f829,f830,f839,f838,f837,f836,f840,f852,f861,f862,f863,f864,f865,f866,f867,f868,f869,f870,f871,f873,f888,f889,f891,f898,f901,f902,f904,f917,f923,f924,f926,f930,f933,f939,f940,f942,f946,f954,f960,f961,f966,f971,f975,f980,f984,f991,f994,f1000,f1001,f1007,f1012,f1023,f1024,f1025,f1032,f1033,f1035,f1047,f1048,f1050,f1057,f1059,f1062,f1068,f1069,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1210,f1212,f1217,f1219,f1230,f1231,f1239,f1247,f1249,f1254,f1256,f1267,f1275,f1283,f1285,f1339,f1333,f1332,f1373,f1366,f1365,f1385,f1411,f1410,f1403,f1402,f1422,f1433,f1437,f1461,f1486,f1492,f1494,f1501,f1512,f1516,f1542,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1663,f1662,f1661,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1818,f1819,f1816,f1837,f1838,f1835,f1852,f1854,f1859,f1861,f1931,f2001,f1981,f2015,f2030,f2041,f2043,f2050,f2052,f2073,f2082,f2088,f2090,f2097,f2099,f2119,f2118,f2128,f2157,f2170,f2251,f2252,f2222,f2262,f2269,f2270,f2271,f2272,f2317,f2318,f2319,f2351,f2352,f2349,f2371,f2372,f2373,f2400,f2401,f2398,f2432,f2433,f2465,f2466,f2504,f2505,f2544,f2545,f2617,f2622,f2765,f2770,f2841,f2844,f2847,f2848,f2751,f2842,f2996,f2855])).
fof(f2997,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f334,f333,f335,f592,f343,f355,f354,f359,f358,f412,f411,f413,f414,f420,f419,f418,f417,f424,f423,f422,f421,f428,f427,f426,f429,f430,f431,f432,f433,f437,f434,f444,f443,f593,f462,f459,f458,f457,f468,f467,f477,f478,f479,f480,f481,f482,f483,f489,f488,f487,f485,f484,f508,f507,f506,f505,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f520,f519,f518,f517,f524,f523,f522,f521,f578,f528,f527,f526,f525,f535,f534,f533,f532,f541,f538,f537,f536,f548,f547,f546,f577,f544,f543,f542,f552,f551,f550,f549,f556,f555,f554,f553,f565,f564,f563,f576,f561,f560,f559,f569,f568,f567,f566,f573,f572,f571,f570,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f622,f623,f625,f631,f636,f637,f640,f643,f646,f653,f654,f656,f663,f667,f668,f671,f677,f701,f747,f756,f755,f766,f764,f778,f777,f776,f775,f787,f796,f797,f798,f799,f800,f801,f802,f803,f804,f805,f817,f818,f819,f825,f828,f829,f830,f839,f838,f837,f836,f840,f852,f861,f862,f863,f864,f865,f866,f867,f868,f869,f870,f871,f873,f888,f889,f891,f898,f901,f902,f904,f917,f923,f924,f926,f930,f933,f939,f940,f942,f946,f954,f960,f961,f966,f971,f975,f980,f984,f991,f994,f1000,f1001,f1007,f1012,f1023,f1024,f1025,f1032,f1033,f1035,f1047,f1048,f1050,f1057,f1059,f1062,f1068,f1069,f1079,f1082,f1087,f1111,f1114,f1119,f1157,f1196,f1210,f1212,f1217,f1219,f1230,f1231,f1239,f1247,f1249,f1254,f1256,f1267,f1275,f1283,f1285,f1339,f1333,f1332,f1373,f1366,f1365,f1385,f1411,f1410,f1403,f1402,f1422,f1433,f1437,f1461,f1486,f1492,f1494,f1501,f1512,f1516,f1542,f1569,f1582,f1767,f1766,f1765,f1764,f1763,f1762,f1761,f1760,f1759,f1757,f1756,f1754,f1753,f1749,f1746,f1745,f1742,f1741,f1737,f1725,f1724,f1723,f1719,f1717,f1716,f1713,f1712,f1709,f1705,f1703,f1701,f1700,f1697,f1695,f1694,f1693,f1692,f1689,f1688,f1683,f1682,f1679,f1678,f1674,f1667,f1666,f1665,f1664,f1663,f1662,f1661,f1660,f1659,f1658,f1657,f1656,f1653,f1652,f1648,f1647,f1640,f1639,f1635,f1634,f1620,f1619,f1616,f1615,f1613,f1611,f1608,f1606,f1605,f1604,f1603,f1602,f1599,f1598,f1597,f1596,f1595,f1818,f1819,f1816,f1837,f1838,f1835,f1852,f1854,f1859,f1861,f1931,f2001,f1981,f2015,f2030,f2041,f2043,f2050,f2052,f2073,f2082,f2088,f2090,f2097,f2099,f2119,f2118,f2128,f2157,f2170,f2251,f2252,f2222,f2262,f2269,f2270,f2271,f2272,f2317,f2318,f2319,f2351,f2352,f2349,f2371,f2372,f2373,f2400,f2401,f2398,f2432,f2433,f2465,f2466,f2504,f2505,f2544,f2545,f2617,f2622,f2765,f2770,f2841,f2844,f2847,f2848,f2751,f2842,f2996])).
fof(f2843,plain,(
  ( ! [X6,X4,X2,X3,X1] : (p(s(bool,X1)) | s(X2,X3) = s(X2,'const.bool.COND_3'(s(bool,X1),s(X2,X6),s(X2,X4))) | s(X2,X3) != s(X2,X4)) ) | $spl6),
  inference(cnf_transformation,[],[f2843_D])).
fof(f2843_D,plain,(
  ( ! [X6,X4,X2,X3,X1] : (p(s(bool,X1)) | s(X2,X3) = s(X2,'const.bool.COND_3'(s(bool,X1),s(X2,X6),s(X2,X4))) | s(X2,X3) != s(X2,X4)) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f2841,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (s(X2,X3) = s(X2,'const.bool.COND_3'(s(bool,X1),s(X2,X6),s(X2,X4))) | s(X2,X3) != s(X2,X4) | p(s(bool,X1)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(forward_demodulation,[],[f2840,f414])).
fof(f2840,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X2,X3) != s(X2,X4) | p(s(bool,X1)) | s(X2,'const.bool.COND_3'(s(bool,'F_0'),s(X2,X5),s(X2,X3))) = s(X2,'const.bool.COND_3'(s(bool,X1),s(X2,X6),s(X2,X4))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f2743,f747])).
fof(f2743,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(X2,X3) != s(X2,X4) | p(s(bool,X1)) | s(X2,'const.bool.COND_3'(s(bool,'F_0'),s(X2,X5),s(X2,X3))) = s(X2,'const.bool.COND_3'(s(bool,X1),s(X2,X6),s(X2,X4))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f427,f343])).
fof(f588,plain,(
  s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',sK12),s('type.num.num',sK13))) = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',sK12),s('type.num.num',sK13))) | $spl0),
  inference(cnf_transformation,[],[f588_D])).
fof(f588_D,plain,(
  s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',sK12),s('type.num.num',sK13))) = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',sK12),s('type.num.num',sK13))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f590,plain,(
  s('type.num.num',sK12) = s('type.num.num',sK13) | $spl2),
  inference(cnf_transformation,[],[f590_D])).
fof(f590_D,plain,(
  s('type.num.num',sK12) = s('type.num.num',sK13) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f333,plain,(
  s('type.num.num',sK12) = s('type.num.num',sK13) | s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',sK12),s('type.num.num',sK13))) = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',sK12),s('type.num.num',sK13)))),
  inference(cnf_transformation,[],[f223])).
fof(f223,plain,(
  (s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',sK12),s('type.num.num',sK13))) = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',sK12),s('type.num.num',sK13))) | s('type.num.num',sK12) = s('type.num.num',sK13)) & (s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',sK12),s('type.num.num',sK13))) != s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',sK12),s('type.num.num',sK13))) | s('type.num.num',sK12) != s('type.num.num',sK13))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13])],[f222])).
fof(f222,plain,(
  ? [X0,X1] : ((s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X0),s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) & (s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X0),s('type.num.num',X1))) != s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X0),s('type.num.num',X1))) | s('type.num.num',X0) != s('type.num.num',X1)))),
  inference(nnf_transformation,[],[f168])).
fof(f168,plain,(
  ? [X0,X1] : (s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X0),s('type.num.num',X1))) <~> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0,X1] : (s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X0),s('type.num.num',X1))) <=> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X6),s('type.num.num',X7))) = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X6),s('type.num.num',X7))) <=> s('type.num.num',X6) = s('type.num.num',X7))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',X6),s('type.num.num',X7))) = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X6),s('type.num.num',X7))) <=> s('type.num.num',X6) = s('type.num.num',X7))),
  file('thm.arithmetic.MIN_MAX_EQ','thm.arithmetic.MIN_MAX_EQ')).
fof(f589,plain,(
  s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',sK12),s('type.num.num',sK13))) != s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',sK12),s('type.num.num',sK13))) | $spl1),
  inference(cnf_transformation,[],[f589_D])).
fof(f589_D,plain,(
  s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',sK12),s('type.num.num',sK13))) != s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',sK12),s('type.num.num',sK13))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f591,plain,(
  s('type.num.num',sK12) != s('type.num.num',sK13) | $spl3),
  inference(cnf_transformation,[],[f591_D])).
fof(f591_D,plain,(
  s('type.num.num',sK12) != s('type.num.num',sK13) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f334,plain,(
  s('type.num.num',sK12) != s('type.num.num',sK13) | s('type.num.num','const.arithmetic.MIN_2'(s('type.num.num',sK12),s('type.num.num',sK13))) != s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',sK12),s('type.num.num',sK13)))),
  inference(cnf_transformation,[],[f223])).
% SZS output end Proof for thm.arithmetic
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 486
% Passive clauses: 1999
% Generated clauses: 60802
% Final active clauses: 374
% Final passive clauses: 869
% Input formulas: 69
% Initial clauses: 243
% 
% Duplicate literals: 503
% Trivial inequalities: 20
% Fw subsumption resolutions: 598
% Fw demodulations: 31528
% Bw demodulations: 658
% Condensations: 1160
% Global subsumptions: 86
% 
% Simple tautologies: 133
% Equational tautologies: 69
% Forward subsumptions: 16059
% Fw demodulations to eq. taut.: 8459
% Bw demodulations to eq. taut.: 11
% 
% Unit resulting resolution: 15924
% Factoring: 1
% Forward superposition: 4933
% Backward superposition: 4220
% Self superposition: 78
% Equality factoring: 26
% Equality resolution: 3
% 
% Split clauses: 85
% Split components: 170
% SAT solver clauses: 2360
% SAT solver unit clauses: 490
% SAT solver binary clauses: 623
% 
% Sat splits: 85
% Sat splitting refutations: 103
% 
% Memory used [KB]: 18293
% Time elapsed: 1.774 s
% ------------------------------
% ------------------------------
% Success in time 3.07 s
