Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.numRing
% (92382)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 168
% Passive clauses: 4205
% Generated clauses: 22008
% Final active clauses: 160
% Final passive clauses: 3895
% Input formulas: 118
% Initial clauses: 123
% 
% Duplicate literals: 583
% Trivial inequalities: 405
% Fw subsumption resolutions: 1070
% Fw demodulations: 1335
% Bw demodulations: 4
% 
% Simple tautologies: 337
% Equational tautologies: 22
% Forward subsumptions: 13651
% Fw demodulations to eq. taut.: 12
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 31
% Forward superposition: 10551
% Backward superposition: 7654
% Self superposition: 217
% Equality factoring: 54
% Equality resolution: 31
% 
% Split clauses: 255
% Split components: 510
% SAT solver clauses: 512
% SAT solver unit clauses: 5
% SAT solver binary clauses: 100
% 
% Sat splits: 255
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 512
% TWLsolver calls for satisfiability: 67
% 
% Memory used [KB]: 21492
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.numRing
% (92394)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 127
% Passive clauses: 1335
% Generated clauses: 15761
% Final active clauses: 115
% Final passive clauses: 1064
% Input formulas: 118
% Initial clauses: 123
% 
% Duplicate literals: 579
% Trivial inequalities: 198
% Fw subsumption resolutions: 555
% Fw demodulations: 1018
% Bw demodulations: 121
% Condensations: 889
% 
% Simple tautologies: 175
% Equational tautologies: 25
% Forward subsumptions: 8855
% Fw demodulations to eq. taut.: 4
% Bw demodulations to eq. taut.: 26
% 
% Binary resolution: 43
% Forward superposition: 5957
% Backward superposition: 4599
% Self superposition: 211
% Equality factoring: 74
% Equality resolution: 19
% 
% Split clauses: 105
% Split components: 210
% SAT solver clauses: 2073
% SAT solver unit clauses: 235
% SAT solver binary clauses: 50
% 
% Sat splits: 105
% Sat splitting refutations: 30
% 
% TWLsolver clauses: 2073
% TWLsolver calls for satisfiability: 262
% 
% Memory used [KB]: 22515
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.numRing
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.numRing
% SZS output start Proof for thm.numRing
fof(f120229,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f120147])).
fof(f120147,plain,(
  s('type.num.num',sK1) != s('type.num.num',sK1)),
  inference(superposition,[],[f307,f119222])).
fof(f119222,plain,(
  ( ! [X2,X3] : (s('type.num.num',X3) = s('type.num.num','const.numRing.num_interp_sp_2'(s('type.quote.varmap'('type.num.num'),X2),s('type.canonical.spolynom'('type.num.num'),'const.canonical.SPconst_1'(s('type.num.num',X3)))))) )),
  inference(superposition,[],[f24398,f349])).
fof(f349,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X3) = s(X0,'const.canonical.interp_sp_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),'const.canonical.SPconst_1'(s(X0,X3)))))) )),
  inference(cnf_transformation,[],[f160])).
fof(f160,plain,(
  ! [X0,X1,X2,X3] : s(X0,X3) = s(X0,'const.canonical.interp_sp_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),'const.canonical.SPconst_1'(s(X0,X3)))))),
  inference(flattening,[],[f159])).
fof(f159,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,X3) = s(X0,'const.canonical.interp_sp_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),'const.canonical.SPconst_1'(s(X0,X3)))))),
  inference(rectify,[],[f53])).
fof(f53,axiom,(
  ! [X8] : ! [X13,X6,X7] : s(X8,X7) = s(X8,'const.canonical.interp_sp_3'(s('type.semi_ring.semi_ring'(X8),X13),s('type.quote.varmap'(X8),X6),s('type.canonical.spolynom'(X8),'const.canonical.SPconst_1'(s(X8,X7)))))),
  file('thm.numRing.num_ring_thms|split|2','thm.canonical.interp_sp_def|split|0')).
fof(f24398,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.numRing.num_interp_sp_2'(s('type.quote.varmap'('type.num.num'),X0),s('type.canonical.spolynom'('type.num.num'),X1))) = s('type.num.num','const.canonical.interp_sp_3'(s('type.semi_ring.semi_ring'('type.num.num'),'const.semi_ring.recordtype|46|semi_ring_4'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.+_0'),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.*_0'))),s('type.quote.varmap'('type.num.num'),X0),s('type.canonical.spolynom'('type.num.num'),X1)))) )),
  inference(forward_demodulation,[],[f24364,f415])).
fof(f415,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.numRing.num_interp_sp_2'(s('type.quote.varmap'('type.num.num'),X0),s('type.canonical.spolynom'('type.num.num'),X1))) = s('type.num.num',happ(s(fun('type.canonical.spolynom'('type.num.num'),'type.num.num'),happ(s(fun('type.quote.varmap'('type.num.num'),fun('type.canonical.spolynom'('type.num.num'),'type.num.num')),'const.numRing.num_interp_sp_0'),s('type.quote.varmap'('type.num.num'),X0))),s('type.canonical.spolynom'('type.num.num'),X1)))) )),
  inference(cnf_transformation,[],[f281])).
fof(f281,plain,(
  ! [X0,X1] : s('type.num.num','const.numRing.num_interp_sp_2'(s('type.quote.varmap'('type.num.num'),X0),s('type.canonical.spolynom'('type.num.num'),X1))) = s('type.num.num',happ(s(fun('type.canonical.spolynom'('type.num.num'),'type.num.num'),happ(s(fun('type.quote.varmap'('type.num.num'),fun('type.canonical.spolynom'('type.num.num'),'type.num.num')),'const.numRing.num_interp_sp_0'),s('type.quote.varmap'('type.num.num'),X0))),s('type.canonical.spolynom'('type.num.num'),X1)))),
  inference(rectify,[],[f83])).
fof(f83,axiom,(
  ! [X45,X46] : s('type.num.num','const.numRing.num_interp_sp_2'(s('type.quote.varmap'('type.num.num'),X45),s('type.canonical.spolynom'('type.num.num'),X46))) = s('type.num.num',happ(s(fun('type.canonical.spolynom'('type.num.num'),'type.num.num'),happ(s(fun('type.quote.varmap'('type.num.num'),fun('type.canonical.spolynom'('type.num.num'),'type.num.num')),'const.numRing.num_interp_sp_0'),s('type.quote.varmap'('type.num.num'),X45))),s('type.canonical.spolynom'('type.num.num'),X46)))),
  file('thm.numRing.num_ring_thms|split|2','ARITY_0_2_const.numRing.num_interp_sp')).
fof(f24364,plain,(
  ( ! [X0,X1] : (s('type.num.num',happ(s(fun('type.canonical.spolynom'('type.num.num'),'type.num.num'),happ(s(fun('type.quote.varmap'('type.num.num'),fun('type.canonical.spolynom'('type.num.num'),'type.num.num')),'const.numRing.num_interp_sp_0'),s('type.quote.varmap'('type.num.num'),X0))),s('type.canonical.spolynom'('type.num.num'),X1))) = s('type.num.num','const.canonical.interp_sp_3'(s('type.semi_ring.semi_ring'('type.num.num'),'const.semi_ring.recordtype|46|semi_ring_4'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.+_0'),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.*_0'))),s('type.quote.varmap'('type.num.num'),X0),s('type.canonical.spolynom'('type.num.num'),X1)))) )),
  inference(superposition,[],[f363,f314])).
fof(f314,plain,(
  s(fun('type.quote.varmap'('type.num.num'),fun('type.canonical.spolynom'('type.num.num'),'type.num.num')),'const.numRing.num_interp_sp_0') = s(fun('type.quote.varmap'('type.num.num'),fun('type.canonical.spolynom'('type.num.num'),'type.num.num')),'const.canonical.interp_sp_1'(s('type.semi_ring.semi_ring'('type.num.num'),'const.semi_ring.recordtype|46|semi_ring_4'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.+_0'),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.*_0')))))),
  inference(cnf_transformation,[],[f59])).
fof(f59,axiom,(
  s(fun('type.quote.varmap'('type.num.num'),fun('type.canonical.spolynom'('type.num.num'),'type.num.num')),'const.numRing.num_interp_sp_0') = s(fun('type.quote.varmap'('type.num.num'),fun('type.canonical.spolynom'('type.num.num'),'type.num.num')),'const.canonical.interp_sp_1'(s('type.semi_ring.semi_ring'('type.num.num'),'const.semi_ring.recordtype|46|semi_ring_4'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.+_0'),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.*_0')))))),
  file('thm.numRing.num_ring_thms|split|2','thm.numRing.num_interp_sp_def')).
fof(f363,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.canonical.interp_sp_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),X3))) = s(X0,happ(s(fun('type.canonical.spolynom'(X0),X0),happ(s(fun('type.quote.varmap'(X0),fun('type.canonical.spolynom'(X0),X0)),'const.canonical.interp_sp_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.canonical.spolynom'(X0),X3)))) )),
  inference(cnf_transformation,[],[f188])).
fof(f188,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.canonical.interp_sp_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),X3))) = s(X0,happ(s(fun('type.canonical.spolynom'(X0),X0),happ(s(fun('type.quote.varmap'(X0),fun('type.canonical.spolynom'(X0),X0)),'const.canonical.interp_sp_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.canonical.spolynom'(X0),X3)))),
  inference(flattening,[],[f187])).
fof(f187,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.canonical.interp_sp_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),X3))) = s(X0,happ(s(fun('type.canonical.spolynom'(X0),X0),happ(s(fun('type.quote.varmap'(X0),fun('type.canonical.spolynom'(X0),X0)),'const.canonical.interp_sp_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.canonical.spolynom'(X0),X3)))),
  inference(rectify,[],[f84])).
fof(f84,axiom,(
  ! [X8] : ! [X45,X46,X47] : s(X8,'const.canonical.interp_sp_3'(s('type.semi_ring.semi_ring'(X8),X45),s('type.quote.varmap'(X8),X46),s('type.canonical.spolynom'(X8),X47))) = s(X8,happ(s(fun('type.canonical.spolynom'(X8),X8),happ(s(fun('type.quote.varmap'(X8),fun('type.canonical.spolynom'(X8),X8)),'const.canonical.interp_sp_1'(s('type.semi_ring.semi_ring'(X8),X45))),s('type.quote.varmap'(X8),X46))),s('type.canonical.spolynom'(X8),X47)))),
  file('thm.numRing.num_ring_thms|split|2','ARITY_1_3_const.canonical.interp_sp')).
fof(f307,plain,(
  s('type.num.num',sK1) != s('type.num.num','const.numRing.num_interp_sp_2'(s('type.quote.varmap'('type.num.num'),sK0),s('type.canonical.spolynom'('type.num.num'),'const.canonical.SPconst_1'(s('type.num.num',sK1)))))),
  inference(cnf_transformation,[],[f301])).
fof(f301,plain,(
  s('type.num.num',sK1) != s('type.num.num','const.numRing.num_interp_sp_2'(s('type.quote.varmap'('type.num.num'),sK0),s('type.canonical.spolynom'('type.num.num'),'const.canonical.SPconst_1'(s('type.num.num',sK1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f298])).
fof(f298,plain,(
  ? [X0,X1] : s('type.num.num',X1) != s('type.num.num','const.numRing.num_interp_sp_2'(s('type.quote.varmap'('type.num.num'),X0),s('type.canonical.spolynom'('type.num.num'),'const.canonical.SPconst_1'(s('type.num.num',X1)))))),
  inference(ennf_transformation,[],[f120])).
fof(f120,plain,(
  ~! [X0,X1] : s('type.num.num',X1) = s('type.num.num','const.numRing.num_interp_sp_2'(s('type.quote.varmap'('type.num.num'),X0),s('type.canonical.spolynom'('type.num.num'),'const.canonical.SPconst_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : s('type.num.num',X7) = s('type.num.num','const.numRing.num_interp_sp_2'(s('type.quote.varmap'('type.num.num'),X6),s('type.canonical.spolynom'('type.num.num'),'const.canonical.SPconst_1'(s('type.num.num',X7)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : s('type.num.num',X7) = s('type.num.num','const.numRing.num_interp_sp_2'(s('type.quote.varmap'('type.num.num'),X6),s('type.canonical.spolynom'('type.num.num'),'const.canonical.SPconst_1'(s('type.num.num',X7)))))),
  file('thm.numRing.num_ring_thms|split|2','thm.numRing.num_ring_thms|split|2')).
% SZS output end Proof for thm.numRing
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 535
% Passive clauses: 9526
% Generated clauses: 123226
% Final active clauses: 218
% Final passive clauses: 1989
% Input formulas: 118
% Initial clauses: 123
% Discarded non-redundant clauses: 2933
% 
% Duplicate literals: 3466
% Trivial inequalities: 1175
% Fw subsumption resolutions: 3142
% Fw demodulations: 13958
% Bw demodulations: 1815
% Condensations: 4751
% Global subsumptions: 209
% 
% Simple tautologies: 599
% Equational tautologies: 264
% Forward subsumptions: 84956
% Fw demodulations to eq. taut.: 23
% Bw demodulations to eq. taut.: 116
% 
% Unit resulting resolution: 2729
% Forward superposition: 45615
% Backward superposition: 40961
% Self superposition: 719
% Equality factoring: 216
% Equality resolution: 48
% 
% Split clauses: 340
% Split components: 685
% SAT solver clauses: 10653
% SAT solver unit clauses: 1199
% SAT solver binary clauses: 2738
% 
% Sat splits: 340
% Sat splitting refutations: 236
% 
% Memory used [KB]: 157353
% Time elapsed: 6.687 s
% ------------------------------
% ------------------------------
% Success in time 8.013 s
