Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.ringNorm
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.ringNorm
% SZS output start Proof for thm.ringNorm
fof(f4650,plain,(
  $false),
  inference(equality_resolution,[],[f4621])).
fof(f4621,plain,(
  ( ! [X0] : (s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(sK0),sK1),s('type.canonical.spolynom'(sK0),'const.canonical.SPconst_1'(s(sK0,sK2))))) != s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(sK0),X0),s('type.canonical.spolynom'(sK0),'const.canonical.SPconst_1'(s(sK0,sK2)))))) )),
  inference(superposition,[],[f255,f4614])).
fof(f4614,plain,(
  ( ! [X2,X0,X1] : (s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.spolynom'(X0),X2)))) )),
  inference(forward_demodulation,[],[f282,f68])).
fof(f68,plain,(
  ( ! [X2,X0,X1] : (s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.r_spolynom_normalize_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.spolynom'(X0),X2)))) )),
  inference(cnf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0,X1,X2] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.r_spolynom_normalize_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.spolynom'(X0),X2)))),
  inference(flattening,[],[f34])).
fof(f34,plain,(
  ! [X0] : ! [X1,X2] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.r_spolynom_normalize_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.spolynom'(X0),X2)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X6] : ! [X14,X15] : s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X6),X14),s('type.canonical.spolynom'(X6),X15))) = s('type.canonical.canonical_sum'(X6),happ(s(fun('type.canonical.spolynom'(X6),'type.canonical.canonical_sum'(X6)),'const.ringNorm.r_spolynom_normalize_1'(s('type.ring.ring'(X6),X14))),s('type.canonical.spolynom'(X6),X15)))),
  file('thm.ringNorm.spolynom_normalize_def|split|1','ARITY_1_2_const.ringNorm.r_spolynom_normalize')).
fof(f282,plain,(
  ( ! [X2,X0,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.r_spolynom_normalize_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.spolynom'(X0),X2)))) )),
  inference(superposition,[],[f67,f63])).
fof(f63,plain,(
  ( ! [X0,X1] : (s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.r_spolynom_normalize_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.canonical.spolynom_normalize_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f25])).
fof(f25,plain,(
  ! [X0,X1] : s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.r_spolynom_normalize_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.canonical.spolynom_normalize_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f24])).
fof(f24,plain,(
  ! [X0] : ! [X1] : s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.r_spolynom_normalize_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.canonical.spolynom_normalize_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6] : ! [X7] : s(fun('type.canonical.spolynom'(X6),'type.canonical.canonical_sum'(X6)),'const.ringNorm.r_spolynom_normalize_1'(s('type.ring.ring'(X6),X7))) = s(fun('type.canonical.spolynom'(X6),'type.canonical.canonical_sum'(X6)),'const.canonical.spolynom_normalize_1'(s('type.semi_ring.semi_ring'(X6),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X6),X7)))))),
  file('thm.ringNorm.spolynom_normalize_def|split|1','thm.ringNorm.r_spolynom_normalize_def')).
fof(f67,plain,(
  ( ! [X2,X0,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.canonical.spolynom_normalize_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.spolynom'(X0),X2)))) )),
  inference(cnf_transformation,[],[f33])).
fof(f33,plain,(
  ! [X0,X1,X2] : s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.canonical.spolynom_normalize_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.spolynom'(X0),X2)))),
  inference(flattening,[],[f32])).
fof(f32,plain,(
  ! [X0] : ! [X1,X2] : s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.canonical.spolynom_normalize_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.spolynom'(X0),X2)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X6] : ! [X14,X15] : s('type.canonical.canonical_sum'(X6),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X6),X14),s('type.canonical.spolynom'(X6),X15))) = s('type.canonical.canonical_sum'(X6),happ(s(fun('type.canonical.spolynom'(X6),'type.canonical.canonical_sum'(X6)),'const.canonical.spolynom_normalize_1'(s('type.semi_ring.semi_ring'(X6),X14))),s('type.canonical.spolynom'(X6),X15)))),
  file('thm.ringNorm.spolynom_normalize_def|split|1','ARITY_1_2_const.canonical.spolynom_normalize')).
fof(f255,plain,(
  ( ! [X0] : (s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(sK0),sK1),s('type.canonical.spolynom'(sK0),'const.canonical.SPconst_1'(s(sK0,sK2))))) != s('type.canonical.canonical_sum'(sK0),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(sK0),X0),s('type.canonical.spolynom'(sK0),'const.canonical.SPconst_1'(s(sK0,sK2)))))) )),
  inference(superposition,[],[f59,f66])).
fof(f66,plain,(
  ( ! [X2,X0,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_monom_3'(s(X0,X2),s('type.list.list'('type.quote.index'),'const.list.NIL_0'),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0'))) = s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.spolynom'(X0),'const.canonical.SPconst_1'(s(X0,X2)))))) )),
  inference(cnf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0,X1,X2] : s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_monom_3'(s(X0,X2),s('type.list.list'('type.quote.index'),'const.list.NIL_0'),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0'))) = s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.spolynom'(X0),'const.canonical.SPconst_1'(s(X0,X2)))))),
  inference(flattening,[],[f30])).
fof(f30,plain,(
  ! [X0] : ! [X1,X2] : s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_monom_3'(s(X0,X2),s('type.list.list'('type.quote.index'),'const.list.NIL_0'),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0'))) = s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.spolynom'(X0),'const.canonical.SPconst_1'(s(X0,X2)))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6] : ! [X11,X8] : s('type.canonical.canonical_sum'(X6),'const.canonical.Cons_monom_3'(s(X6,X8),s('type.list.list'('type.quote.index'),'const.list.NIL_0'),s('type.canonical.canonical_sum'(X6),'const.canonical.Nil_monom_0'))) = s('type.canonical.canonical_sum'(X6),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X6),X11),s('type.canonical.spolynom'(X6),'const.canonical.SPconst_1'(s(X6,X8)))))),
  file('thm.ringNorm.spolynom_normalize_def|split|1','thm.canonical.spolynom_normalize_def|split|1')).
fof(f59,plain,(
  s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(sK0),sK1),s('type.canonical.spolynom'(sK0),'const.canonical.SPconst_1'(s(sK0,sK2))))) != s('type.canonical.canonical_sum'(sK0),'const.canonical.Cons_monom_3'(s(sK0,sK2),s('type.list.list'('type.quote.index'),'const.list.NIL_0'),s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0')))),
  inference(cnf_transformation,[],[f56])).
fof(f56,plain,(
  s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(sK0),sK1),s('type.canonical.spolynom'(sK0),'const.canonical.SPconst_1'(s(sK0,sK2))))) != s('type.canonical.canonical_sum'(sK0),'const.canonical.Cons_monom_3'(s(sK0,sK2),s('type.list.list'('type.quote.index'),'const.list.NIL_0'),s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0')))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f54])).
fof(f54,plain,(
  ? [X0,X1,X2] : s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_monom_3'(s(X0,X2),s('type.list.list'('type.quote.index'),'const.list.NIL_0'),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0'))) != s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X0),X1),s('type.canonical.spolynom'(X0),'const.canonical.SPconst_1'(s(X0,X2)))))),
  inference(ennf_transformation,[],[f22])).
fof(f22,plain,(
  ~! [X0,X1,X2] : s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_monom_3'(s(X0,X2),s('type.list.list'('type.quote.index'),'const.list.NIL_0'),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0'))) = s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X0),X1),s('type.canonical.spolynom'(X0),'const.canonical.SPconst_1'(s(X0,X2)))))),
  inference(flattening,[],[f21])).
fof(f21,plain,(
  ~! [X0] : ! [X1,X2] : s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_monom_3'(s(X0,X2),s('type.list.list'('type.quote.index'),'const.list.NIL_0'),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0'))) = s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X0),X1),s('type.canonical.spolynom'(X0),'const.canonical.SPconst_1'(s(X0,X2)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : s('type.canonical.canonical_sum'(X6),'const.canonical.Cons_monom_3'(s(X6,X8),s('type.list.list'('type.quote.index'),'const.list.NIL_0'),s('type.canonical.canonical_sum'(X6),'const.canonical.Nil_monom_0'))) = s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X6),X7),s('type.canonical.spolynom'(X6),'const.canonical.SPconst_1'(s(X6,X8)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : s('type.canonical.canonical_sum'(X6),'const.canonical.Cons_monom_3'(s(X6,X8),s('type.list.list'('type.quote.index'),'const.list.NIL_0'),s('type.canonical.canonical_sum'(X6),'const.canonical.Nil_monom_0'))) = s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X6),X7),s('type.canonical.spolynom'(X6),'const.canonical.SPconst_1'(s(X6,X8)))))),
  file('thm.ringNorm.spolynom_normalize_def|split|1','thm.ringNorm.spolynom_normalize_def|split|1')).
% SZS output end Proof for thm.ringNorm
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 45
% Passive clauses: 605
% Generated clauses: 4417
% Final active clauses: 43
% Final passive clauses: 526
% Input formulas: 19
% Initial clauses: 20
% 
% Duplicate literals: 439
% Trivial inequalities: 60
% Fw subsumption resolutions: 434
% Fw demodulations: 9
% 
% Simple tautologies: 138
% Equational tautologies: 89
% Forward subsumptions: 1982
% 
% Binary resolution: 16
% Forward superposition: 1526
% Backward superposition: 1340
% Self superposition: 296
% Equality factoring: 246
% Equality resolution: 7
% 
% Split clauses: 187
% Split components: 376
% SAT solver clauses: 693
% SAT solver unit clauses: 228
% SAT solver binary clauses: 234
% 
% Sat splits: 187
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 691
% TWLsolver calls for satisfiability: 27
% 
% Memory used [KB]: 2430
% Time elapsed: 0.205 s
% ------------------------------
% ------------------------------
% Success in time 0.207 s
