Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.sat
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.sat
% SZS output start Proof for thm.sat
fof(f839,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f838,f532,f484,f470,f124,f188,f138,f205,f588,f523,f544,f527,f494,f253,f128,f198,f200,f505,f500,f127,f194,f389,f240,f251,f249,f370,f368,f125,f190,f186,f134,f202,f204,f130,f199,f763,f602,f502,f754,f503,f131,f748,f632,f596,f594,f501,f642,f371,f619,f637,f197,f630,f727,f644,f629,f597,f643,f196,f638,f372,f129,f201,f195,f193,f191,f189,f187])).
fof(f187,plain,(
  s(bool,sK7) != s(bool,'const.bool.COND_3'(s(bool,sK6),s(bool,sK5),s(bool,sK4))) | $spl1),
  inference(cnf_transformation,[],[f187_D])).
fof(f187_D,plain,(
  s(bool,sK7) != s(bool,'const.bool.COND_3'(s(bool,sK6),s(bool,sK5),s(bool,sK4))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f189,plain,(
  ~sP0(sK4,sK6,sK7) | $spl3),
  inference(cnf_transformation,[],[f189_D])).
fof(f189_D,plain,(
  ~sP0(sK4,sK6,sK7) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f191,plain,(
  ~sP1(sK6,sK5,sK7) | $spl5),
  inference(cnf_transformation,[],[f191_D])).
fof(f191_D,plain,(
  ~sP1(sK6,sK5,sK7) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f193,plain,(
  ~sP2(sK4,sK5,sK7) | $spl7),
  inference(cnf_transformation,[],[f193_D])).
fof(f193_D,plain,(
  ~sP2(sK4,sK5,sK7) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f195,plain,(
  ~sP3(sK7,sK5,sK6) | $spl9),
  inference(cnf_transformation,[],[f195_D])).
fof(f195_D,plain,(
  ~sP3(sK7,sK5,sK6) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f201,plain,(
  ~p(s(bool,sK6)) | $spl15),
  inference(cnf_transformation,[],[f201_D])).
fof(f201_D,plain,(
  ~p(s(bool,sK6)) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f129,plain,(
  ~p(s(bool,sK6)) | ~sP3(sK7,sK5,sK6) | ~sP2(sK4,sK5,sK7) | ~sP1(sK6,sK5,sK7) | ~sP0(sK4,sK6,sK7) | s(bool,sK7) != s(bool,'const.bool.COND_3'(s(bool,sK6),s(bool,sK5),s(bool,sK4)))),
  inference(cnf_transformation,[],[f79])).
fof(f79,plain,(
  (s(bool,sK7) = s(bool,'const.bool.COND_3'(s(bool,sK6),s(bool,sK5),s(bool,sK4))) | (sP0(sK4,sK6,sK7) & sP1(sK6,sK5,sK7) & sP2(sK4,sK5,sK7) & sP3(sK7,sK5,sK6) & (p(s(bool,sK6)) | p(s(bool,sK4)) | ~p(s(bool,sK7))))) & (s(bool,sK7) != s(bool,'const.bool.COND_3'(s(bool,sK6),s(bool,sK5),s(bool,sK4))) | ~sP0(sK4,sK6,sK7) | ~sP1(sK6,sK5,sK7) | ~sP2(sK4,sK5,sK7) | ~sP3(sK7,sK5,sK6) | (~p(s(bool,sK6)) & ~p(s(bool,sK4)) & p(s(bool,sK7))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4,sK5,sK6,sK7])],[f78])).
fof(f78,plain,(
  ? [X0,X1,X2,X3] : ((s(bool,X3) = s(bool,'const.bool.COND_3'(s(bool,X2),s(bool,X1),s(bool,X0))) | (sP0(X0,X2,X3) & sP1(X2,X1,X3) & sP2(X0,X1,X3) & sP3(X3,X1,X2) & (p(s(bool,X2)) | p(s(bool,X0)) | ~p(s(bool,X3))))) & (s(bool,X3) != s(bool,'const.bool.COND_3'(s(bool,X2),s(bool,X1),s(bool,X0))) | ~sP0(X0,X2,X3) | ~sP1(X2,X1,X3) | ~sP2(X0,X1,X3) | ~sP3(X3,X1,X2) | (~p(s(bool,X2)) & ~p(s(bool,X0)) & p(s(bool,X3)))))),
  inference(flattening,[],[f77])).
fof(f77,plain,(
  ? [X0,X1,X2,X3] : ((s(bool,X3) = s(bool,'const.bool.COND_3'(s(bool,X2),s(bool,X1),s(bool,X0))) | (sP0(X0,X2,X3) & sP1(X2,X1,X3) & sP2(X0,X1,X3) & sP3(X3,X1,X2) & (p(s(bool,X2)) | p(s(bool,X0)) | ~p(s(bool,X3))))) & (s(bool,X3) != s(bool,'const.bool.COND_3'(s(bool,X2),s(bool,X1),s(bool,X0))) | (~sP0(X0,X2,X3) | ~sP1(X2,X1,X3) | ~sP2(X0,X1,X3) | ~sP3(X3,X1,X2) | (~p(s(bool,X2)) & ~p(s(bool,X0)) & p(s(bool,X3))))))),
  inference(nnf_transformation,[],[f64])).
fof(f64,plain,(
  ? [X0,X1,X2,X3] : (s(bool,X3) = s(bool,'const.bool.COND_3'(s(bool,X2),s(bool,X1),s(bool,X0))) <~> (sP0(X0,X2,X3) & sP1(X2,X1,X3) & sP2(X0,X1,X3) & sP3(X3,X1,X2) & (p(s(bool,X2)) | p(s(bool,X0)) | ~p(s(bool,X3)))))),
  inference(definition_folding,[],[f58,f63,f62,f61,f60])).
fof(f60,plain,(
  ! [X0,X2,X3] : (sP0(X0,X2,X3) <=> (p(s(bool,X3)) | p(s(bool,X2)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f61,plain,(
  ! [X2,X1,X3] : (sP1(X2,X1,X3) <=> (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f62,plain,(
  ! [X0,X1,X3] : (sP2(X0,X1,X3) <=> (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f63,plain,(
  ! [X3,X1,X2] : (sP3(X3,X1,X2) <=> (~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X3))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f58,plain,(
  ? [X0,X1,X2,X3] : (s(bool,X3) = s(bool,'const.bool.COND_3'(s(bool,X2),s(bool,X1),s(bool,X0))) <~> ((p(s(bool,X3)) | p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X3))) & (p(s(bool,X2)) | p(s(bool,X0)) | ~p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f29])).
fof(f29,plain,(
  ~! [X0,X1,X2,X3] : (s(bool,X3) = s(bool,'const.bool.COND_3'(s(bool,X2),s(bool,X1),s(bool,X0))) <=> ((p(s(bool,X3)) | p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X3))) & (p(s(bool,X2)) | p(s(bool,X0)) | ~p(s(bool,X3)))))),
  inference(flattening,[],[f28])).
fof(f28,plain,(
  ~! [X0,X1,X2,X3] : (s(bool,X3) = s(bool,'const.bool.COND_3'(s(bool,X2),s(bool,X1),s(bool,X0))) <=> ((p(s(bool,X3)) | p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X3))) & (p(s(bool,X2)) | p(s(bool,X0)) | ~p(s(bool,X3)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8,X9] : (s(bool,X9) = s(bool,'const.bool.COND_3'(s(bool,X8),s(bool,X7),s(bool,X6))) <=> ((p(s(bool,X9)) | p(s(bool,X8)) | ~p(s(bool,X6))) & (p(s(bool,X9)) | ~p(s(bool,X7)) | ~p(s(bool,X8))) & (p(s(bool,X9)) | ~p(s(bool,X7)) | ~p(s(bool,X6))) & (~p(s(bool,X8)) | p(s(bool,X7)) | ~p(s(bool,X9))) & (p(s(bool,X8)) | p(s(bool,X6)) | ~p(s(bool,X9)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8,X9] : (s(bool,X9) = s(bool,'const.bool.COND_3'(s(bool,X8),s(bool,X7),s(bool,X6))) <=> ((p(s(bool,X9)) | p(s(bool,X8)) | ~p(s(bool,X6))) & (p(s(bool,X9)) | ~p(s(bool,X7)) | ~p(s(bool,X8))) & (p(s(bool,X9)) | ~p(s(bool,X7)) | ~p(s(bool,X6))) & (~p(s(bool,X8)) | p(s(bool,X7)) | ~p(s(bool,X9))) & (p(s(bool,X8)) | p(s(bool,X6)) | ~p(s(bool,X9)))))),
  file('thm.sat.sat_3|49|_prove_ltb','thm.sat.sat_3|49|_prove_ltb')).
fof(f372,plain,(
  ~p(s(bool,sK5)) | $spl37),
  inference(cnf_transformation,[],[f372_D])).
fof(f372_D,plain,(
  ~p(s(bool,sK5)) <=> ~$spl37),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl37])])).
fof(f638,plain,(
  ~p(s(bool,sK5)) | $spl9),
  inference(resolution,[],[f195,f110])).
fof(f110,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f67])).
fof(f67,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f66])).
fof(f66,plain,(
  ! [X3,X1,X2] : ((~sP3(X3,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X3))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X3))) | sP3(X3,X1,X2)))),
  inference(flattening,[],[f65])).
fof(f65,plain,(
  ! [X3,X1,X2] : ((~sP3(X3,X1,X2) | (~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X3)))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X3))) | sP3(X3,X1,X2)))),
  inference(nnf_transformation,[],[f63])).
fof(f196,plain,(
  ~p(s(bool,sK7)) | $spl11),
  inference(cnf_transformation,[],[f196_D])).
fof(f196_D,plain,(
  ~p(s(bool,sK7)) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f643,plain,(
  ~p(s(bool,sK7)) | $spl7),
  inference(resolution,[],[f193,f113])).
fof(f113,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f70])).
fof(f70,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(rectify,[],[f69])).
fof(f69,plain,(
  ! [X0,X1,X3] : ((~sP2(X0,X1,X3) | p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X3)) & p(s(bool,X1)) & p(s(bool,X0))) | sP2(X0,X1,X3)))),
  inference(flattening,[],[f68])).
fof(f68,plain,(
  ! [X0,X1,X3] : ((~sP2(X0,X1,X3) | (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X3)) & p(s(bool,X1)) & p(s(bool,X0))) | sP2(X0,X1,X3)))),
  inference(nnf_transformation,[],[f62])).
fof(f597,plain,(
  ~p(s(bool,sK7)) | $spl5),
  inference(resolution,[],[f191,f117])).
fof(f117,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0,X1,X2] : ((~sP1(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1,X2)))),
  inference(rectify,[],[f72])).
fof(f72,plain,(
  ! [X2,X1,X3] : ((~sP1(X2,X1,X3) | p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2))) & ((~p(s(bool,X3)) & p(s(bool,X1)) & p(s(bool,X2))) | sP1(X2,X1,X3)))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ! [X2,X1,X3] : ((~sP1(X2,X1,X3) | (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))) & ((~p(s(bool,X3)) & p(s(bool,X1)) & p(s(bool,X2))) | sP1(X2,X1,X3)))),
  inference(nnf_transformation,[],[f61])).
fof(f629,plain,(
  ~p(s(bool,sK7)) | $spl3),
  inference(resolution,[],[f189,f121])).
fof(f121,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0,X1,X2] : ((~sP0(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP0(X0,X1,X2)))),
  inference(rectify,[],[f75])).
fof(f75,plain,(
  ! [X0,X2,X3] : ((~sP0(X0,X2,X3) | p(s(bool,X3)) | p(s(bool,X2)) | ~p(s(bool,X0))) & ((~p(s(bool,X3)) & ~p(s(bool,X2)) & p(s(bool,X0))) | sP0(X0,X2,X3)))),
  inference(flattening,[],[f74])).
fof(f74,plain,(
  ! [X0,X2,X3] : ((~sP0(X0,X2,X3) | (p(s(bool,X3)) | p(s(bool,X2)) | ~p(s(bool,X0)))) & ((~p(s(bool,X3)) & ~p(s(bool,X2)) & p(s(bool,X0))) | sP0(X0,X2,X3)))),
  inference(nnf_transformation,[],[f60])).
fof(f644,plain,(
  $false | ($spl7 | $spl13)),
  inference(subsumption_resolution,[],[f641,f199])).
fof(f641,plain,(
  p(s(bool,sK4)) | $spl7),
  inference(resolution,[],[f193,f115])).
fof(f115,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f70])).
fof(f727,plain,(
  $false | ($spl0 | $spl25 | $spl32 | $spl41)),
  inference(subsumption_resolution,[],[f726,f500])).
fof(f726,plain,(
  s(bool,'T_0') = s(bool,sK7) | ($spl0 | $spl25 | $spl32)),
  inference(forward_demodulation,[],[f675,f253])).
fof(f675,plain,(
  s(bool,sK4) = s(bool,sK7) | ($spl0 | $spl25)),
  inference(forward_demodulation,[],[f674,f182])).
fof(f182,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0,X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f54])).
fof(f54,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X11] : ! [X12,X13] : s(X11,X13) = s(X11,'const.bool.COND_3'(s(bool,'F_0'),s(X11,X12),s(X11,X13)))),
  file('thm.sat.sat_3|49|_prove_ltb','thm.bool.COND_CLAUSES|split|1')).
fof(f674,plain,(
  s(bool,sK7) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,sK5),s(bool,sK4))) | ($spl0 | $spl25)),
  inference(subsumption_resolution,[],[f661,f244])).
fof(f244,plain,(
  s(bool,'T_0') != s(bool,sK6) | $spl25),
  inference(cnf_transformation,[],[f244_D])).
fof(f244_D,plain,(
  s(bool,'T_0') != s(bool,sK6) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f661,plain,(
  s(bool,sK7) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,sK5),s(bool,sK4))) | s(bool,'T_0') = s(bool,sK6) | $spl0),
  inference(superposition,[],[f186,f139])).
fof(f139,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.sat.sat_3|49|_prove_ltb',aHL_BOOL_CASES)).
fof(f630,plain,(
  $false | ($spl3 | $spl13)),
  inference(subsumption_resolution,[],[f627,f199])).
fof(f627,plain,(
  p(s(bool,sK4)) | $spl3),
  inference(resolution,[],[f189,f123])).
fof(f123,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f76])).
fof(f197,plain,(
  p(s(bool,sK7)) | $spl10),
  inference(cnf_transformation,[],[f197_D])).
fof(f197_D,plain,(
  p(s(bool,sK7)) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f637,plain,(
  p(s(bool,sK7)) | $spl9),
  inference(resolution,[],[f195,f111])).
fof(f111,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f67])).
fof(f619,plain,(
  $false | ($spl11 | $spl16 | $spl40)),
  inference(subsumption_resolution,[],[f607,f202])).
fof(f607,plain,(
  ~p(s(bool,'T_0')) | ($spl11 | $spl40)),
  inference(backward_demodulation,[],[f501,f196])).
fof(f371,plain,(
  p(s(bool,sK5)) | $spl36),
  inference(cnf_transformation,[],[f371_D])).
fof(f371_D,plain,(
  p(s(bool,sK5)) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f642,plain,(
  p(s(bool,sK5)) | $spl7),
  inference(resolution,[],[f193,f114])).
fof(f114,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f70])).
fof(f501,plain,(
  s(bool,'T_0') = s(bool,sK7) | $spl40),
  inference(cnf_transformation,[],[f501_D])).
fof(f501_D,plain,(
  s(bool,'T_0') = s(bool,sK7) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f594,plain,(
  s(bool,'T_0') = s(bool,sK7) | ($spl0 | $spl24 | $spl28)),
  inference(forward_demodulation,[],[f593,f181])).
fof(f181,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f52])).
fof(f52,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X11] : ! [X12,X13] : s(X11,X12) = s(X11,'const.bool.COND_3'(s(bool,'T_0'),s(X11,X12),s(X11,X13)))),
  file('thm.sat.sat_3|49|_prove_ltb','thm.bool.COND_CLAUSES|split|0')).
fof(f593,plain,(
  s(bool,sK7) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,sK4))) | ($spl0 | $spl24 | $spl28)),
  inference(forward_demodulation,[],[f592,f245])).
fof(f245,plain,(
  s(bool,'T_0') = s(bool,sK6) | $spl24),
  inference(cnf_transformation,[],[f245_D])).
fof(f245_D,plain,(
  s(bool,'T_0') = s(bool,sK6) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f592,plain,(
  s(bool,sK7) = s(bool,'const.bool.COND_3'(s(bool,sK6),s(bool,'T_0'),s(bool,sK4))) | ($spl0 | $spl28)),
  inference(forward_demodulation,[],[f186,f249])).
fof(f596,plain,(
  p(s(bool,sK5)) | $spl5),
  inference(resolution,[],[f191,f118])).
fof(f118,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f73])).
fof(f632,plain,(
  $false | ($spl3 | $spl16 | $spl24)),
  inference(subsumption_resolution,[],[f631,f202])).
fof(f631,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl24)),
  inference(forward_demodulation,[],[f628,f245])).
fof(f628,plain,(
  ~p(s(bool,sK6)) | $spl3),
  inference(resolution,[],[f189,f122])).
fof(f122,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f76])).
fof(f748,plain,(
  $false | ($spl5 | $spl15 | $spl19)),
  inference(subsumption_resolution,[],[f747,f205])).
fof(f747,plain,(
  p(s(bool,'F_0')) | ($spl5 | $spl15)),
  inference(forward_demodulation,[],[f744,f656])).
fof(f656,plain,(
  s(bool,'F_0') = s(bool,sK6) | $spl15),
  inference(resolution,[],[f201,f146])).
fof(f146,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f37])).
fof(f37,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f36])).
fof(f36,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X10] : (s(bool,'F_0') = s(bool,X10) <=> ~p(s(bool,X10)))),
  file('thm.sat.sat_3|49|_prove_ltb','thm.bool.EQ_CLAUSES|split|2')).
fof(f744,plain,(
  p(s(bool,sK6)) | $spl5),
  inference(resolution,[],[f191,f119])).
fof(f119,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f73])).
fof(f131,plain,(
  p(s(bool,sK7)) | ~sP3(sK7,sK5,sK6) | ~sP2(sK4,sK5,sK7) | ~sP1(sK6,sK5,sK7) | ~sP0(sK4,sK6,sK7) | s(bool,sK7) != s(bool,'const.bool.COND_3'(s(bool,sK6),s(bool,sK5),s(bool,sK4)))),
  inference(cnf_transformation,[],[f79])).
fof(f503,plain,(
  s(bool,'F_0') != s(bool,sK4) | $spl43),
  inference(cnf_transformation,[],[f503_D])).
fof(f503_D,plain,(
  s(bool,'F_0') != s(bool,sK4) <=> ~$spl43),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl43])])).
fof(f754,plain,(
  s(bool,'F_0') != s(bool,sK4) | ~p(s(bool,sK7)) | ($spl0 | $spl15)),
  inference(forward_demodulation,[],[f753,f182])).
fof(f753,plain,(
  s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,sK5),s(bool,sK4))) | ~p(s(bool,sK7)) | ($spl0 | $spl15)),
  inference(forward_demodulation,[],[f667,f656])).
fof(f667,plain,(
  ~p(s(bool,sK7)) | s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,sK6),s(bool,sK5),s(bool,sK4))) | $spl0),
  inference(superposition,[],[f145,f186])).
fof(f145,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f84])).
fof(f502,plain,(
  s(bool,'F_0') = s(bool,sK4) | $spl42),
  inference(cnf_transformation,[],[f502_D])).
fof(f502_D,plain,(
  s(bool,'F_0') = s(bool,sK4) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f602,plain,(
  s(bool,'F_0') = s(bool,sK4) | $spl13),
  inference(resolution,[],[f199,f146])).
fof(f763,plain,(
  $false | ($spl9 | $spl15 | $spl19)),
  inference(subsumption_resolution,[],[f762,f205])).
fof(f762,plain,(
  p(s(bool,'F_0')) | ($spl9 | $spl15)),
  inference(forward_demodulation,[],[f761,f656])).
fof(f761,plain,(
  p(s(bool,sK6)) | $spl9),
  inference(resolution,[],[f195,f109])).
fof(f109,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f67])).
fof(f199,plain,(
  ~p(s(bool,sK4)) | $spl13),
  inference(cnf_transformation,[],[f199_D])).
fof(f199_D,plain,(
  ~p(s(bool,sK4)) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f130,plain,(
  ~p(s(bool,sK4)) | ~sP3(sK7,sK5,sK6) | ~sP2(sK4,sK5,sK7) | ~sP1(sK6,sK5,sK7) | ~sP0(sK4,sK6,sK7) | s(bool,sK7) != s(bool,'const.bool.COND_3'(s(bool,sK6),s(bool,sK5),s(bool,sK4)))),
  inference(cnf_transformation,[],[f79])).
fof(f204,plain,(
  p(s(bool,'F_0')) | $spl18),
  inference(cnf_transformation,[],[f204_D])).
fof(f204_D,plain,(
  p(s(bool,'F_0')) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f202,plain,(
  p(s(bool,'T_0')) | $spl16),
  inference(cnf_transformation,[],[f202_D])).
fof(f202_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f134,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f80])).
fof(f80,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f30])).
fof(f30,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f21])).
fof(f21,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.sat.sat_3|49|_prove_ltb','thm.bool.NOT_CLAUSES|split|2')).
fof(f186,plain,(
  s(bool,sK7) = s(bool,'const.bool.COND_3'(s(bool,sK6),s(bool,sK5),s(bool,sK4))) | $spl0),
  inference(cnf_transformation,[],[f186_D])).
fof(f186_D,plain,(
  s(bool,sK7) = s(bool,'const.bool.COND_3'(s(bool,sK6),s(bool,sK5),s(bool,sK4))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f190,plain,(
  sP1(sK6,sK5,sK7) | $spl4),
  inference(cnf_transformation,[],[f190_D])).
fof(f190_D,plain,(
  sP1(sK6,sK5,sK7) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f125,plain,(
  sP1(sK6,sK5,sK7) | s(bool,sK7) = s(bool,'const.bool.COND_3'(s(bool,sK6),s(bool,sK5),s(bool,sK4)))),
  inference(cnf_transformation,[],[f79])).
fof(f368,plain,(
  ~p(s(bool,sK5)) | ($spl4 | $spl11 | $spl16 | $spl19 | $spl24)),
  inference(subsumption_resolution,[],[f367,f202])).
fof(f367,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,sK5)) | ($spl4 | $spl11 | $spl19 | $spl24)),
  inference(forward_demodulation,[],[f366,f245])).
fof(f366,plain,(
  ~p(s(bool,sK5)) | ~p(s(bool,sK6)) | ($spl4 | $spl11 | $spl19)),
  inference(subsumption_resolution,[],[f365,f205])).
fof(f365,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,sK5)) | ~p(s(bool,sK6)) | ($spl4 | $spl11)),
  inference(forward_demodulation,[],[f329,f211])).
fof(f211,plain,(
  s(bool,'F_0') = s(bool,sK7) | $spl11),
  inference(resolution,[],[f146,f196])).
fof(f329,plain,(
  ~p(s(bool,sK5)) | p(s(bool,sK7)) | ~p(s(bool,sK6)) | $spl4),
  inference(resolution,[],[f116,f190])).
fof(f116,plain,(
  ( ! [X2,X0,X1] : (~sP1(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f73])).
fof(f370,plain,(
  $false | ($spl24 | $spl31)),
  inference(subsumption_resolution,[],[f369,f181])).
fof(f369,plain,(
  s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,sK4))) | ($spl24 | $spl31)),
  inference(forward_demodulation,[],[f251,f245])).
fof(f249,plain,(
  s(bool,'T_0') = s(bool,sK5) | $spl28),
  inference(cnf_transformation,[],[f249_D])).
fof(f249_D,plain,(
  s(bool,'T_0') = s(bool,sK5) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f251,plain,(
  s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,sK6),s(bool,'F_0'),s(bool,sK4))) | $spl31),
  inference(cnf_transformation,[],[f251_D])).
fof(f251_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,sK6),s(bool,'F_0'),s(bool,sK4))) <=> ~$spl31),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl31])])).
fof(f240,plain,(
  s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,sK6),s(bool,'F_0'),s(bool,sK4))) | s(bool,'T_0') = s(bool,sK5) | ($spl1 | $spl11)),
  inference(superposition,[],[f238,f139])).
fof(f238,plain,(
  s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,sK6),s(bool,sK5),s(bool,sK4))) | ($spl1 | $spl11)),
  inference(forward_demodulation,[],[f187,f211])).
fof(f389,plain,(
  $false | ($spl16 | $spl28 | $spl37)),
  inference(subsumption_resolution,[],[f379,f202])).
fof(f379,plain,(
  ~p(s(bool,'T_0')) | ($spl28 | $spl37)),
  inference(backward_demodulation,[],[f249,f372])).
fof(f194,plain,(
  sP3(sK7,sK5,sK6) | $spl8),
  inference(cnf_transformation,[],[f194_D])).
fof(f194_D,plain,(
  sP3(sK7,sK5,sK6) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f127,plain,(
  sP3(sK7,sK5,sK6) | s(bool,sK7) = s(bool,'const.bool.COND_3'(s(bool,sK6),s(bool,sK5),s(bool,sK4)))),
  inference(cnf_transformation,[],[f79])).
fof(f500,plain,(
  s(bool,'T_0') != s(bool,sK7) | $spl41),
  inference(cnf_transformation,[],[f500_D])).
fof(f500_D,plain,(
  s(bool,'T_0') != s(bool,sK7) <=> ~$spl41),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl41])])).
fof(f505,plain,(
  s(bool,'T_0') != s(bool,sK7) | ($spl1 | $spl15 | $spl32)),
  inference(backward_demodulation,[],[f253,f488])).
fof(f488,plain,(
  s(bool,sK4) != s(bool,sK7) | ($spl1 | $spl15)),
  inference(forward_demodulation,[],[f487,f182])).
fof(f487,plain,(
  s(bool,sK7) != s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,sK5),s(bool,sK4))) | ($spl1 | $spl15)),
  inference(forward_demodulation,[],[f187,f400])).
fof(f400,plain,(
  s(bool,'F_0') = s(bool,sK6) | $spl15),
  inference(resolution,[],[f201,f146])).
fof(f200,plain,(
  p(s(bool,sK6)) | $spl14),
  inference(cnf_transformation,[],[f200_D])).
fof(f200_D,plain,(
  p(s(bool,sK6)) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f198,plain,(
  p(s(bool,sK4)) | $spl12),
  inference(cnf_transformation,[],[f198_D])).
fof(f198_D,plain,(
  p(s(bool,sK4)) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f128,plain,(
  ~p(s(bool,sK7)) | p(s(bool,sK4)) | p(s(bool,sK6)) | s(bool,sK7) = s(bool,'const.bool.COND_3'(s(bool,sK6),s(bool,sK5),s(bool,sK4)))),
  inference(cnf_transformation,[],[f79])).
fof(f253,plain,(
  s(bool,'T_0') = s(bool,sK4) | $spl32),
  inference(cnf_transformation,[],[f253_D])).
fof(f253_D,plain,(
  s(bool,'T_0') = s(bool,sK4) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f494,plain,(
  s(bool,'T_0') = s(bool,sK4) | $spl12),
  inference(resolution,[],[f198,f142])).
fof(f142,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f34])).
fof(f34,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X10] : (s(bool,'T_0') = s(bool,X10) <=> p(s(bool,X10)))),
  file('thm.sat.sat_3|49|_prove_ltb','thm.bool.EQ_CLAUSES|split|0')).
fof(f527,plain,(
  ~p(s(bool,sK6)) | ($spl8 | $spl10 | $spl37)),
  inference(subsumption_resolution,[],[f526,f197])).
fof(f526,plain,(
  ~p(s(bool,sK6)) | ~p(s(bool,sK7)) | ($spl8 | $spl37)),
  inference(subsumption_resolution,[],[f291,f372])).
fof(f291,plain,(
  p(s(bool,sK5)) | ~p(s(bool,sK6)) | ~p(s(bool,sK7)) | $spl8),
  inference(resolution,[],[f108,f194])).
fof(f108,plain,(
  ( ! [X2,X0,X1] : (~sP3(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f67])).
fof(f544,plain,(
  $false | ($spl29 | $spl36)),
  inference(subsumption_resolution,[],[f542,f248])).
fof(f248,plain,(
  s(bool,'T_0') != s(bool,sK5) | $spl29),
  inference(cnf_transformation,[],[f248_D])).
fof(f248_D,plain,(
  s(bool,'T_0') != s(bool,sK5) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f542,plain,(
  s(bool,'T_0') = s(bool,sK5) | $spl36),
  inference(resolution,[],[f371,f142])).
fof(f523,plain,(
  s(bool,'T_0') = s(bool,sK7) | $spl10),
  inference(trivial_inequality_removal,[],[f522])).
fof(f522,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,sK7) | $spl10),
  inference(superposition,[],[f489,f139])).
fof(f489,plain,(
  s(bool,'F_0') != s(bool,sK7) | $spl10),
  inference(resolution,[],[f197,f145])).
fof(f588,plain,(
  s(bool,'T_0') != s(bool,sK7) | ($spl1 | $spl24 | $spl28)),
  inference(forward_demodulation,[],[f587,f181])).
fof(f587,plain,(
  s(bool,sK7) != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,sK4))) | ($spl1 | $spl24 | $spl28)),
  inference(forward_demodulation,[],[f586,f245])).
fof(f586,plain,(
  s(bool,sK7) != s(bool,'const.bool.COND_3'(s(bool,sK6),s(bool,'T_0'),s(bool,sK4))) | ($spl1 | $spl28)),
  inference(forward_demodulation,[],[f187,f249])).
fof(f205,plain,(
  ~p(s(bool,'F_0')) | $spl19),
  inference(cnf_transformation,[],[f205_D])).
fof(f205_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f138,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f32])).
fof(f32,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.sat.sat_3|49|_prove_ltb',aHL_FALSITY)).
fof(f188,plain,(
  sP0(sK4,sK6,sK7) | $spl2),
  inference(cnf_transformation,[],[f188_D])).
fof(f188_D,plain,(
  sP0(sK4,sK6,sK7) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f124,plain,(
  sP0(sK4,sK6,sK7) | s(bool,sK7) = s(bool,'const.bool.COND_3'(s(bool,sK6),s(bool,sK5),s(bool,sK4)))),
  inference(cnf_transformation,[],[f79])).
fof(f470,plain,(
  $false | ($spl1 | $spl11 | $spl13 | $spl15)),
  inference(subsumption_resolution,[],[f469,f182])).
fof(f469,plain,(
  s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,sK5),s(bool,'F_0'))) | ($spl1 | $spl11 | $spl13 | $spl15)),
  inference(forward_demodulation,[],[f468,f400])).
fof(f468,plain,(
  s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,sK6),s(bool,sK5),s(bool,'F_0'))) | ($spl1 | $spl11 | $spl13)),
  inference(forward_demodulation,[],[f238,f375])).
fof(f375,plain,(
  s(bool,'F_0') = s(bool,sK4) | $spl13),
  inference(resolution,[],[f199,f146])).
fof(f484,plain,(
  ~p(s(bool,sK4)) | ($spl2 | $spl11 | $spl15 | $spl19)),
  inference(subsumption_resolution,[],[f483,f205])).
fof(f483,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,sK4)) | ($spl2 | $spl11 | $spl15 | $spl19)),
  inference(forward_demodulation,[],[f482,f211])).
fof(f482,plain,(
  p(s(bool,sK7)) | ~p(s(bool,sK4)) | ($spl2 | $spl15 | $spl19)),
  inference(subsumption_resolution,[],[f481,f205])).
fof(f481,plain,(
  p(s(bool,'F_0')) | p(s(bool,sK7)) | ~p(s(bool,sK4)) | ($spl2 | $spl15)),
  inference(forward_demodulation,[],[f334,f400])).
fof(f334,plain,(
  p(s(bool,sK6)) | p(s(bool,sK7)) | ~p(s(bool,sK4)) | $spl2),
  inference(resolution,[],[f120,f188])).
fof(f120,plain,(
  ( ! [X2,X0,X1] : (~sP0(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f76])).
fof(f532,plain,(
  $false | ($spl14 | $spl25)),
  inference(subsumption_resolution,[],[f530,f244])).
fof(f530,plain,(
  s(bool,'T_0') = s(bool,sK6) | $spl14),
  inference(resolution,[],[f200,f142])).
fof(f838,plain,(
  $false | ($spl0 | $spl24 | $spl29 | $spl40)),
  inference(subsumption_resolution,[],[f837,f248])).
fof(f837,plain,(
  s(bool,'T_0') = s(bool,sK5) | ($spl0 | $spl24 | $spl40)),
  inference(forward_demodulation,[],[f836,f501])).
fof(f836,plain,(
  s(bool,sK5) = s(bool,sK7) | ($spl0 | $spl24)),
  inference(forward_demodulation,[],[f835,f181])).
fof(f835,plain,(
  s(bool,sK7) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,sK5),s(bool,sK4))) | ($spl0 | $spl24)),
  inference(forward_demodulation,[],[f186,f245])).
% SZS output end Proof for thm.sat
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 100
% Passive clauses: 189
% Generated clauses: 840
% Final active clauses: 35
% Final passive clauses: 25
% Input formulas: 26
% Initial clauses: 76
% 
% Duplicate literals: 3
% Trivial inequalities: 19
% Fw subsumption resolutions: 64
% Fw demodulations: 143
% Bw demodulations: 21
% 
% Simple tautologies: 40
% Equational tautologies: 3
% Forward subsumptions: 273
% 
% Binary resolution: 76
% Forward superposition: 93
% Backward superposition: 168
% Self superposition: 6
% Equality factoring: 2
% Equality resolution: 1
% 
% Split clauses: 51
% Split components: 122
% SAT solver clauses: 139
% SAT solver unit clauses: 6
% SAT solver binary clauses: 56
% SAT solver learnt clauses: 11
% SAT solver learnt literals: 22
% 
% Sat splits: 51
% Sat splitting refutations: 23
% 
% TWLsolver clauses: 128
% TWLsolver calls for satisfiability: 62
% 
% Memory used [KB]: 895
% Time elapsed: 0.023 s
% ------------------------------
% ------------------------------
% Success in time 0.025 s
