Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.option
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.option
% SZS output start Proof for thm.option
fof(f171,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f170,f67,f106,f104,f155,f68,f107,f105])).
fof(f105,plain,(
  s('type.option.option'('type.one.one'),'const.option.OPTION_GUARD_1'(s(bool,sK0))) != s('type.option.option'('type.one.one'),'const.option.SOME_1'(s('type.one.one','const.one.one_0'))) | $spl1),
  inference(cnf_transformation,[],[f105_D])).
fof(f105_D,plain,(
  s('type.option.option'('type.one.one'),'const.option.OPTION_GUARD_1'(s(bool,sK0))) != s('type.option.option'('type.one.one'),'const.option.SOME_1'(s('type.one.one','const.one.one_0'))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f107,plain,(
  ~p(s(bool,sK0)) | $spl3),
  inference(cnf_transformation,[],[f107_D])).
fof(f107_D,plain,(
  ~p(s(bool,sK0)) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f68,plain,(
  ~p(s(bool,sK0)) | s('type.option.option'('type.one.one'),'const.option.OPTION_GUARD_1'(s(bool,sK0))) != s('type.option.option'('type.one.one'),'const.option.SOME_1'(s('type.one.one','const.one.one_0')))),
  inference(cnf_transformation,[],[f55])).
fof(f55,plain,(
  (s('type.option.option'('type.one.one'),'const.option.OPTION_GUARD_1'(s(bool,sK0))) = s('type.option.option'('type.one.one'),'const.option.SOME_1'(s('type.one.one','const.one.one_0'))) | p(s(bool,sK0))) & (s('type.option.option'('type.one.one'),'const.option.OPTION_GUARD_1'(s(bool,sK0))) != s('type.option.option'('type.one.one'),'const.option.SOME_1'(s('type.one.one','const.one.one_0'))) | ~p(s(bool,sK0)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f54])).
fof(f54,plain,(
  ? [X0] : ((s('type.option.option'('type.one.one'),'const.option.SOME_1'(s('type.one.one','const.one.one_0'))) = s('type.option.option'('type.one.one'),'const.option.OPTION_GUARD_1'(s(bool,X0))) | p(s(bool,X0))) & (s('type.option.option'('type.one.one'),'const.option.SOME_1'(s('type.one.one','const.one.one_0'))) != s('type.option.option'('type.one.one'),'const.option.OPTION_GUARD_1'(s(bool,X0))) | ~p(s(bool,X0))))),
  inference(nnf_transformation,[],[f49])).
fof(f49,plain,(
  ? [X0] : (s('type.option.option'('type.one.one'),'const.option.SOME_1'(s('type.one.one','const.one.one_0'))) = s('type.option.option'('type.one.one'),'const.option.OPTION_GUARD_1'(s(bool,X0))) <~> p(s(bool,X0)))),
  inference(ennf_transformation,[],[f24])).
fof(f24,plain,(
  ~! [X0] : (s('type.option.option'('type.one.one'),'const.option.SOME_1'(s('type.one.one','const.one.one_0'))) = s('type.option.option'('type.one.one'),'const.option.OPTION_GUARD_1'(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : (s('type.option.option'('type.one.one'),'const.option.SOME_1'(s('type.one.one','const.one.one_0'))) = s('type.option.option'('type.one.one'),'const.option.OPTION_GUARD_1'(s(bool,X6))) <=> p(s(bool,X6)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : (s('type.option.option'('type.one.one'),'const.option.SOME_1'(s('type.one.one','const.one.one_0'))) = s('type.option.option'('type.one.one'),'const.option.OPTION_GUARD_1'(s(bool,X6))) <=> p(s(bool,X6)))),
  file('thm.option.OPTION_GUARD_EQ_THM|split|0','thm.option.OPTION_GUARD_EQ_THM|split|0')).
fof(f155,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f72,f154])).
fof(f154,plain,(
  s('type.option.option'('type.one.one'),'const.option.OPTION_GUARD_1'(s(bool,'T_0'))) != s('type.option.option'('type.one.one'),'const.option.SOME_1'(s('type.one.one','const.one.one_0'))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f105,f114])).
fof(f114,plain,(
  s(bool,'T_0') = s(bool,sK0) | $spl2),
  inference(resolution,[],[f82,f106])).
fof(f82,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f30])).
fof(f30,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X7] : (s(bool,'T_0') = s(bool,X7) <=> p(s(bool,X7)))),
  file('thm.option.OPTION_GUARD_EQ_THM|split|0','thm.bool.EQ_CLAUSES|split|0')).
fof(f72,plain,(
  s('type.option.option'('type.one.one'),'const.option.OPTION_GUARD_1'(s(bool,'T_0'))) = s('type.option.option'('type.one.one'),'const.option.SOME_1'(s('type.one.one','const.one.one_0')))),
  inference(cnf_transformation,[],[f22])).
fof(f22,axiom,(
  s('type.option.option'('type.one.one'),'const.option.OPTION_GUARD_1'(s(bool,'T_0'))) = s('type.option.option'('type.one.one'),'const.option.SOME_1'(s('type.one.one','const.one.one_0')))),
  file('thm.option.OPTION_GUARD_EQ_THM|split|0','thm.option.OPTION_GUARD_def|split|0')).
fof(f104,plain,(
  s('type.option.option'('type.one.one'),'const.option.OPTION_GUARD_1'(s(bool,sK0))) = s('type.option.option'('type.one.one'),'const.option.SOME_1'(s('type.one.one','const.one.one_0'))) | $spl0),
  inference(cnf_transformation,[],[f104_D])).
fof(f104_D,plain,(
  s('type.option.option'('type.one.one'),'const.option.OPTION_GUARD_1'(s(bool,sK0))) = s('type.option.option'('type.one.one'),'const.option.SOME_1'(s('type.one.one','const.one.one_0'))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f106,plain,(
  p(s(bool,sK0)) | $spl2),
  inference(cnf_transformation,[],[f106_D])).
fof(f106_D,plain,(
  p(s(bool,sK0)) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f67,plain,(
  p(s(bool,sK0)) | s('type.option.option'('type.one.one'),'const.option.OPTION_GUARD_1'(s(bool,sK0))) = s('type.option.option'('type.one.one'),'const.option.SOME_1'(s('type.one.one','const.one.one_0')))),
  inference(cnf_transformation,[],[f55])).
fof(f170,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f169,f93])).
fof(f93,plain,(
  ( ! [X0,X1] : (s('type.option.option'(X0),'const.option.NONE_0') != s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1)))) )),
  inference(cnf_transformation,[],[f41])).
fof(f41,plain,(
  ! [X0,X1] : s('type.option.option'(X0),'const.option.NONE_0') != s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1)))),
  inference(flattening,[],[f40])).
fof(f40,plain,(
  ! [X0] : ! [X1] : ~s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X10] : ! [X11] : ~s('type.option.option'(X10),'const.option.NONE_0') = s('type.option.option'(X10),'const.option.SOME_1'(s(X10,X11)))),
  file('thm.option.OPTION_GUARD_EQ_THM|split|0','thm.option.NOT_NONE_SOME')).
fof(f169,plain,(
  s('type.option.option'('type.one.one'),'const.option.NONE_0') = s('type.option.option'('type.one.one'),'const.option.SOME_1'(s('type.one.one','const.one.one_0'))) | ($spl0 | $spl3)),
  inference(forward_demodulation,[],[f168,f71])).
fof(f71,plain,(
  s('type.option.option'('type.one.one'),'const.option.NONE_0') = s('type.option.option'('type.one.one'),'const.option.OPTION_GUARD_1'(s(bool,'F_0')))),
  inference(cnf_transformation,[],[f23])).
fof(f23,axiom,(
  s('type.option.option'('type.one.one'),'const.option.NONE_0') = s('type.option.option'('type.one.one'),'const.option.OPTION_GUARD_1'(s(bool,'F_0')))),
  file('thm.option.OPTION_GUARD_EQ_THM|split|0','thm.option.OPTION_GUARD_def|split|1')).
fof(f168,plain,(
  s('type.option.option'('type.one.one'),'const.option.OPTION_GUARD_1'(s(bool,'F_0'))) = s('type.option.option'('type.one.one'),'const.option.SOME_1'(s('type.one.one','const.one.one_0'))) | ($spl0 | $spl3)),
  inference(forward_demodulation,[],[f104,f156])).
fof(f156,plain,(
  s(bool,'F_0') = s(bool,sK0) | $spl3),
  inference(resolution,[],[f107,f86])).
fof(f86,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f33])).
fof(f33,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f32])).
fof(f32,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X7] : (s(bool,'F_0') = s(bool,X7) <=> ~p(s(bool,X7)))),
  file('thm.option.OPTION_GUARD_EQ_THM|split|0','thm.bool.EQ_CLAUSES|split|3')).
% SZS output end Proof for thm.option
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 15
% Passive clauses: 37
% Generated clauses: 101
% Final active clauses: 11
% Final passive clauses: 11
% Input formulas: 22
% Initial clauses: 36
% 
% Duplicate literals: 1
% Trivial inequalities: 4
% Fw subsumption resolutions: 3
% Fw demodulations: 4
% Bw demodulations: 2
% 
% Simple tautologies: 10
% Equational tautologies: 4
% Forward subsumptions: 36
% 
% Binary resolution: 8
% Forward superposition: 6
% Backward superposition: 25
% Self superposition: 2
% Equality factoring: 2
% Equality resolution: 1
% 
% Split clauses: 7
% Split components: 14
% SAT solver clauses: 15
% SAT solver unit clauses: 5
% SAT solver binary clauses: 9
% SAT solver learnt clauses: 2
% SAT solver learnt literals: 1
% 
% Sat splits: 7
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 13
% TWLsolver calls for satisfiability: 4
% 
% Memory used [KB]: 639
% Time elapsed: 0.014 s
% ------------------------------
% ------------------------------
% Success in time 0.015 s
