Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.integer
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.integer
% SZS output start Proof for thm.integer
fof(f6899,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f6667,f624,f426,f431,f361,f429,f435,f279,f430,f4731,f4732,f3369,f6805,f3457])).
fof(f3457,plain,(
  $false | ($spl1 | $spl2 | $spl24)),
  inference(subsumption_resolution,[],[f3388,f431])).
fof(f3388,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl24)),
  inference(backward_demodulation,[],[f3369,f430])).
fof(f6805,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f6776,f4732])).
fof(f6776,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f376,f976])).
fof(f976,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(trivial_inequality_removal,[],[f975])).
fof(f975,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(superposition,[],[f267,f314])).
fof(f314,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.integer.INT_LT_REDUCE|split|2',aHL_BOOL_CASES)).
fof(f267,plain,(
  s(bool,'F_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(cnf_transformation,[],[f103])).
fof(f103,plain,(
  s(bool,'F_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(flattening,[],[f6])).
fof(f6,negated_conjecture,(
  ~s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.integer.INT_LT_REDUCE|split|2','thm.integer.INT_LT_REDUCE|split|2')).
fof(f376,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f174])).
fof(f174,plain,(
  ! [X0,X1] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X7,X8] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num',X8))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X8)))))),
  file('thm.integer.INT_LT_REDUCE|split|2','thm.integer.INT_LT_CALCULATE|split|0')).
fof(f3369,plain,(
  ( ! [X15] : (s(bool,'T_0') = s(bool,X15)) ) | $spl24),
  inference(cnf_transformation,[],[f3369_D])).
fof(f3369_D,plain,(
  ( ! [X15] : (s(bool,'T_0') = s(bool,X15)) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f4732,plain,(
  ( ! [X26] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',X26),s('type.num.num','const.num.|48|_0')))) ) | $spl30),
  inference(cnf_transformation,[],[f4732_D])).
fof(f4732_D,plain,(
  ( ! [X26] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',X26),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f4731,plain,(
  ( ! [X26,X25] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',X26),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X25)) ) | ($spl2 | $spl8 | $spl13)),
  inference(forward_demodulation,[],[f4711,f285])).
fof(f285,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f112])).
fof(f112,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f86])).
fof(f86,axiom,(
  ! [X7] : s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X7)))),
  file('thm.integer.INT_LT_REDUCE|split|2','thm.numeral.numeral_distrib|split|26')).
fof(f4711,plain,(
  ( ! [X26,X27,X25] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X26),s('type.num.num','const.num.|48|_0'))) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X27))) | s(bool,'T_0') = s(bool,X25)) ) | ($spl2 | $spl8 | $spl13)),
  inference(duplicate_literal_removal,[],[f4662])).
fof(f4662,plain,(
  ( ! [X26,X27,X25] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X26),s('type.num.num','const.num.|48|_0'))) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X27))) | s(bool,'T_0') = s(bool,X25) | s(bool,'T_0') = s(bool,X25)) ) | ($spl2 | $spl8 | $spl13)),
  inference(superposition,[],[f954,f3221])).
fof(f3221,plain,(
  ( ! [X14,X15] : (s(bool,X15) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X14),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X15)) ) | $spl13),
  inference(subsumption_resolution,[],[f2945,f624])).
fof(f2945,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X15) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X14),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f595,f291])).
fof(f291,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0] : s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f80])).
fof(f80,axiom,(
  ! [X7] : s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num','const.num.|48|_0')))),
  file('thm.integer.INT_LT_REDUCE|split|2','thm.numeral.numeral_distrib|split|20')).
fof(f595,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f314,f314])).
fof(f954,plain,(
  ( ! [X6,X7] : (s(bool,X7) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X6))) | s(bool,'T_0') = s(bool,X7)) ) | ($spl2 | $spl8)),
  inference(subsumption_resolution,[],[f947,f431])).
fof(f947,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,X7) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X6))) | s(bool,'T_0') = s(bool,X7)) ) | $spl8),
  inference(superposition,[],[f601,f285])).
fof(f601,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl8),
  inference(superposition,[],[f435,f314])).
fof(f430,plain,(
  ~p(s(bool,'F_0')) | $spl1),
  inference(cnf_transformation,[],[f430_D])).
fof(f430_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f279,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f106])).
fof(f106,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.integer.INT_LT_REDUCE|split|2',aHL_FALSITY)).
fof(f435,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl8),
  inference(cnf_transformation,[],[f435_D])).
fof(f435_D,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f429,plain,(
  p(s(bool,'F_0')) | $spl0),
  inference(cnf_transformation,[],[f429_D])).
fof(f429_D,plain,(
  p(s(bool,'F_0')) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f361,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | ~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f245])).
fof(f245,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X0)) & (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f244])).
fof(f244,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X0)) & (s(bool,'F_0') != s(bool,X0) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f207])).
fof(f207,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> s(bool,'F_0') = s(bool,X0))),
  inference(ennf_transformation,[],[f160])).
fof(f160,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> s(bool,'F_0') = s(bool,X0))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X11] : ((p(s(bool,X11)) => p(s(bool,'F_0'))) <=> s(bool,'F_0') = s(bool,X11))),
  file('thm.integer.INT_LT_REDUCE|split|2','thm.bool.IMP_F_EQ_F')).
fof(f431,plain,(
  p(s(bool,'T_0')) | $spl2),
  inference(cnf_transformation,[],[f431_D])).
fof(f431_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f426,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f364])).
fof(f364,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f247])).
fof(f247,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f164])).
fof(f164,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f163])).
fof(f163,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X10] : ! [X6] : (s(X10,X6) = s(X10,X6) <=> p(s(bool,'T_0')))),
  file('thm.integer.INT_LT_REDUCE|split|2','thm.bool.REFL_CLAUSE')).
fof(f624,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl13),
  inference(cnf_transformation,[],[f624_D])).
fof(f624_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f6667,plain,(
  ( ! [X31] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X31)) ) | ($spl2 | $spl8)),
  inference(forward_demodulation,[],[f6666,f443])).
fof(f443,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))),
  inference(resolution,[],[f316,f270])).
fof(f270,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(cnf_transformation,[],[f95])).
fof(f95,axiom,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0'))))),
  file('thm.integer.INT_LT_REDUCE|split|2','thm.numeral.numeral_distrib|split|35')).
fof(f316,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f217])).
fof(f217,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f141])).
fof(f141,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f48])).
fof(f48,axiom,(
  ! [X11] : (s(bool,'T_0') = s(bool,X11) <=> p(s(bool,X11)))),
  file('thm.integer.INT_LT_REDUCE|split|2','thm.bool.EQ_CLAUSES|split|0')).
fof(f6666,plain,(
  ( ! [X31] : (s(bool,'F_0') != s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X31)) ) | ($spl2 | $spl8)),
  inference(forward_demodulation,[],[f6633,f4991])).
fof(f4991,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num',X2)))) )),
  inference(resolution,[],[f4980,f320])).
fof(f320,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f219])).
fof(f219,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f144])).
fof(f144,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f143])).
fof(f143,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f50])).
fof(f50,axiom,(
  ! [X11] : (s(bool,'F_0') = s(bool,X11) <=> ~p(s(bool,X11)))),
  file('thm.integer.INT_LT_REDUCE|split|2','thm.bool.EQ_CLAUSES|split|2')).
fof(f4980,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X0))))) )),
  inference(forward_demodulation,[],[f4948,f371])).
fof(f371,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f169])).
fof(f169,plain,(
  ! [X0,X1] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X8,X7] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X8),s('type.num.num',X7))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.num.num',X7)))),
  file('thm.integer.INT_LT_REDUCE|split|2','thm.arithmetic.LESS_EQ')).
fof(f4948,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X0))))) )),
  inference(resolution,[],[f4903,f399])).
fof(f399,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f257])).
fof(f257,plain,(
  ! [X0,X1] : ((p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) & (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(nnf_transformation,[],[f193])).
fof(f193,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(flattening,[],[f192])).
fof(f192,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X8,X7] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X8),s('type.num.num',X7)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',X8)))))),
  file('thm.integer.INT_LT_REDUCE|split|2','thm.arithmetic.NOT_LESS')).
fof(f4903,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))) )),
  inference(equality_resolution,[],[f4831])).
fof(f4831,plain,(
  ( ! [X2,X3] : (s('type.num.num',X3) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num',X3))))) )),
  inference(superposition,[],[f388,f371])).
fof(f388,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X0) != s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f252])).
fof(f252,plain,(
  ! [X0,X1] : ((s('type.num.num',X0) != s('type.num.num',X1) | (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))) & (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | s('type.num.num',X0) = s('type.num.num',X1)))),
  inference(flattening,[],[f251])).
fof(f251,plain,(
  ! [X0,X1] : ((s('type.num.num',X0) != s('type.num.num',X1) | (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))) & ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) | s('type.num.num',X0) = s('type.num.num',X1)))),
  inference(nnf_transformation,[],[f186])).
fof(f186,plain,(
  ! [X0,X1] : (s('type.num.num',X0) = s('type.num.num',X1) <=> (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X8,X7] : (s('type.num.num',X7) = s('type.num.num',X8) <=> (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X8),s('type.num.num',X7)))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',X8))))))),
  file('thm.integer.INT_LT_REDUCE|split|2','thm.arithmetic.EQ_LESS_EQ')).
fof(f6633,plain,(
  ( ! [X31,X32] : (s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0'))) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X32),s('type.num.num',X32))) | s(bool,'T_0') = s(bool,X31)) ) | ($spl2 | $spl8)),
  inference(duplicate_literal_removal,[],[f6563])).
fof(f6563,plain,(
  ( ! [X31,X32] : (s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0'))) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X32),s('type.num.num',X32))) | s(bool,'T_0') = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) ) | ($spl2 | $spl8)),
  inference(superposition,[],[f956,f4990])).
fof(f4990,plain,(
  ( ! [X0,X1] : (s(bool,X1) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X0))) | s(bool,'T_0') = s(bool,X1)) ) | $spl2),
  inference(resolution,[],[f4980,f3048])).
fof(f3048,plain,(
  ( ! [X4,X5] : (p(s(bool,X4)) | s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) ) | $spl2),
  inference(superposition,[],[f431,f595])).
fof(f956,plain,(
  ( ! [X17] : (s(bool,X17) != s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X17)) ) | ($spl2 | $spl8)),
  inference(subsumption_resolution,[],[f953,f431])).
fof(f953,plain,(
  ( ! [X17] : (~p(s(bool,'T_0')) | s(bool,X17) != s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X17)) ) | $spl8),
  inference(superposition,[],[f601,f443])).
% SZS output end Proof for thm.integer
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 140
% Passive clauses: 937
% Generated clauses: 6756
% Final active clauses: 97
% Final passive clauses: 593
% Input formulas: 101
% Initial clauses: 159
% 
% Duplicate literals: 177
% Trivial inequalities: 130
% Fw subsumption resolutions: 294
% Fw demodulations: 1241
% Bw demodulations: 34
% 
% Simple tautologies: 116
% Equational tautologies: 23
% Forward subsumptions: 3721
% Fw demodulations to eq. taut.: 5
% 
% Binary resolution: 57
% Forward superposition: 2598
% Backward superposition: 1717
% Self superposition: 63
% Equality factoring: 34
% Equality resolution: 26
% 
% Split clauses: 164
% Split components: 332
% SAT solver clauses: 223
% SAT solver unit clauses: 13
% SAT solver binary clauses: 108
% SAT solver learnt clauses: 1
% 
% Sat splits: 164
% Sat splitting refutations: 6
% 
% TWLsolver clauses: 84
% TWLsolver calls for satisfiability: 30
% 
% Memory used [KB]: 3198
% Time elapsed: 0.119 s
% ------------------------------
% ------------------------------
% Success in time 0.124 s
