Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.real
% (93776)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 462
% Passive clauses: 3378
% Generated clauses: 21996
% Final active clauses: 402
% Final passive clauses: 2436
% Input formulas: 42
% Initial clauses: 132
% 
% Duplicate literals: 854
% Trivial inequalities: 163
% Fw subsumption resolutions: 2581
% Fw demodulations: 2308
% Bw demodulations: 53
% 
% Simple tautologies: 547
% Equational tautologies: 54
% Forward subsumptions: 11206
% Fw demodulations to eq. taut.: 19
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 7830
% Factoring: 23
% Forward superposition: 3941
% Backward superposition: 3660
% Self superposition: 135
% Equality factoring: 74
% Equality resolution: 21
% 
% Split clauses: 713
% Split components: 1454
% SAT solver clauses: 1175
% SAT solver unit clauses: 29
% SAT solver binary clauses: 154
% 
% Sat splits: 713
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 1158
% TWLsolver calls for satisfiability: 293
% 
% Memory used [KB]: 9338
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.real
% (93806)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 191
% Passive clauses: 650
% Generated clauses: 9332
% Final active clauses: 182
% Final passive clauses: 450
% Input formulas: 42
% Initial clauses: 132
% 
% Duplicate literals: 585
% Trivial inequalities: 72
% Fw subsumption resolutions: 876
% Fw demodulations: 683
% Bw demodulations: 14
% Condensations: 836
% Global subsumptions: 14
% 
% Simple tautologies: 242
% Equational tautologies: 49
% Forward subsumptions: 4183
% Fw demodulations to eq. taut.: 21
% 
% Binary resolution: 603
% Factoring: 10
% Forward superposition: 2053
% Backward superposition: 2531
% Self superposition: 157
% Equality factoring: 100
% Equality resolution: 11
% 
% Split clauses: 55
% Split components: 119
% SAT solver clauses: 1170
% SAT solver unit clauses: 145
% SAT solver binary clauses: 63
% SAT solver learnt clauses: 28
% SAT solver learnt literals: 10
% 
% Sat splits: 55
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 1114
% TWLsolver calls for satisfiability: 1043
% 
% Memory used [KB]: 8059
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.real
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.real
% SZS output start Proof for thm.real
fof(f5975,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f5974])).
fof(f5974,plain,(
  s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK8),s('type.realax.real',sK9))) != s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK8),s('type.realax.real',sK9)))),
  inference(forward_demodulation,[],[f5973,f1450])).
fof(f1450,plain,(
  ( ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK7),s('type.realax.real',X0))),s('type.realax.real',sK7)))) )),
  inference(superposition,[],[f1404,f250])).
fof(f250,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X0)))) )),
  inference(cnf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0,X1] : s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X0)))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X6,X7] : s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X6),s('type.realax.real',X7))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X7),s('type.realax.real',X6)))),
  file('thm.real.REAL_DIV_INNER_CANCEL','thm.realax.REAL_MUL_SYM')).
fof(f1404,plain,(
  ( ! [X6] : (s('type.realax.real',X6) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X6),s('type.realax.real',sK7))),s('type.realax.real',sK7)))) )),
  inference(forward_demodulation,[],[f1384,f1224])).
fof(f1224,plain,(
  ( ! [X1] : (s('type.realax.real',X1) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK7),s('type.realax.real',sK7))),s('type.realax.real',X1)))) )),
  inference(backward_demodulation,[],[f1222,f1058])).
fof(f1058,plain,(
  ( ! [X1] : (s('type.realax.real',X1) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X1)))) )),
  inference(superposition,[],[f208,f250])).
fof(f208,plain,(
  ( ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) )),
  inference(cnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0] : s('type.realax.real',X0) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X6] : s('type.realax.real',X6) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X6),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),
  file('thm.real.REAL_DIV_INNER_CANCEL','thm.real.REAL_MUL_RID')).
fof(f1222,plain,(
  s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK7),s('type.realax.real',sK7))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(forward_demodulation,[],[f1196,f1078])).
fof(f1078,plain,(
  ( ! [X2,X3] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X2),s('type.realax.real',X3))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))),s('type.realax.real',X2)))) )),
  inference(superposition,[],[f251,f250])).
fof(f251,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1)))))) )),
  inference(cnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0,X1] : s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1)))))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X6,X7] : s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X6),s('type.realax.real',X7))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X6),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X7)))))),
  file('thm.real.REAL_DIV_INNER_CANCEL','thm.real.real_div')).
fof(f1196,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK7))),s('type.realax.real',sK7)))),
  inference(unit_resulting_resolution,[],[f199,f216])).
fof(f216,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))),s('type.realax.real',X0))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f110])).
fof(f110,plain,(
  ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))),s('type.realax.real',X0))))),
  inference(ennf_transformation,[],[f56])).
fof(f56,plain,(
  ! [X0] : (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))),s('type.realax.real',X0))))),
  inference(flattening,[],[f55])).
fof(f55,plain,(
  ! [X0] : (~s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))),s('type.realax.real',X0))))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X6] : (~s('type.realax.real',X6) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X6))),s('type.realax.real',X6))))),
  file('thm.real.REAL_DIV_INNER_CANCEL','thm.real.REAL_MUL_LINV')).
fof(f199,plain,(
  s('type.realax.real',sK7) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))),
  inference(cnf_transformation,[],[f147])).
fof(f147,plain,(
  s('type.realax.real',sK7) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) & s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK8),s('type.realax.real',sK9))) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK8),s('type.realax.real',sK7))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK7),s('type.realax.real',sK9)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7,sK8,sK9])],[f104])).
fof(f104,plain,(
  ? [X0,X1,X2] : (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) & s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X1),s('type.realax.real',X2))) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X1),s('type.realax.real',X0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))))))),
  inference(ennf_transformation,[],[f45])).
fof(f45,plain,(
  ~! [X0,X1,X2] : (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X1),s('type.realax.real',X2))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X1),s('type.realax.real',X0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))))))),
  inference(flattening,[],[f44])).
fof(f44,plain,(
  ~! [X0,X1,X2] : (~s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X1),s('type.realax.real',X2))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X1),s('type.realax.real',X0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X2))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : (~s('type.realax.real',X6) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X7),s('type.realax.real',X8))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X7),s('type.realax.real',X6))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X6),s('type.realax.real',X8))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : (~s('type.realax.real',X6) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X7),s('type.realax.real',X8))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X7),s('type.realax.real',X6))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X6),s('type.realax.real',X8))))))),
  file('thm.real.REAL_DIV_INNER_CANCEL','thm.real.REAL_DIV_INNER_CANCEL')).
fof(f1384,plain,(
  ( ! [X6] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X6),s('type.realax.real',sK7))),s('type.realax.real',sK7))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK7),s('type.realax.real',sK7))),s('type.realax.real',X6)))) )),
  inference(superposition,[],[f1279,f1224])).
fof(f1279,plain,(
  ( ! [X4,X3] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X3),s('type.realax.real',X4))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X3),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X4),s('type.realax.real',sK7))),s('type.realax.real',sK7)))))) )),
  inference(forward_demodulation,[],[f1267,f1127])).
fof(f1127,plain,(
  ( ! [X12,X13,X11] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X11),s('type.realax.real',X12))),s('type.realax.real',X13))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X11),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X12),s('type.realax.real',X13)))))) )),
  inference(forward_demodulation,[],[f1114,f251])).
fof(f1114,plain,(
  ( ! [X12,X13,X11] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X11),s('type.realax.real',X12))),s('type.realax.real',X13))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X11),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X12),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X13)))))))) )),
  inference(superposition,[],[f276,f251])).
fof(f276,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real',X2)))) )),
  inference(cnf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0,X1,X2] : s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real',X2)))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X6,X7,X8] : s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X6),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X7),s('type.realax.real',X8))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X6),s('type.realax.real',X7))),s('type.realax.real',X8)))),
  file('thm.real.REAL_DIV_INNER_CANCEL','thm.realax.REAL_MUL_ASSOC')).
fof(f1267,plain,(
  ( ! [X4,X3] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X3),s('type.realax.real',X4))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X3),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X4),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK7),s('type.realax.real',sK7)))))))) )),
  inference(superposition,[],[f1226,f276])).
fof(f1226,plain,(
  ( ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK7),s('type.realax.real',sK7)))))) )),
  inference(backward_demodulation,[],[f1222,f208])).
fof(f5973,plain,(
  s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK8),s('type.realax.real',sK9))) != s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK7),s('type.realax.real',sK8))),s('type.realax.real',sK7))),s('type.realax.real',sK9)))),
  inference(forward_demodulation,[],[f5972,f1127])).
fof(f5972,plain,(
  s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK8),s('type.realax.real',sK9))) != s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK7),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK8),s('type.realax.real',sK7))))),s('type.realax.real',sK9)))),
  inference(forward_demodulation,[],[f5885,f250])).
fof(f5885,plain,(
  s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK8),s('type.realax.real',sK9))) != s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK8),s('type.realax.real',sK7))),s('type.realax.real',sK7))),s('type.realax.real',sK9)))),
  inference(superposition,[],[f200,f1127])).
fof(f200,plain,(
  s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK8),s('type.realax.real',sK9))) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK8),s('type.realax.real',sK7))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK7),s('type.realax.real',sK9)))))),
  inference(cnf_transformation,[],[f147])).
% SZS output end Proof for thm.real
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 173
% Passive clauses: 560
% Generated clauses: 5831
% Final active clauses: 137
% Final passive clauses: 215
% Input formulas: 42
% Initial clauses: 132
% 
% Duplicate literals: 54
% Trivial inequalities: 12
% Fw subsumption resolutions: 83
% Fw demodulations: 1495
% Bw demodulations: 116
% Condensations: 93
% Global subsumptions: 1
% 
% Simple tautologies: 60
% Equational tautologies: 43
% Forward subsumptions: 3323
% Fw demodulations to eq. taut.: 20
% Bw demodulations to eq. taut.: 4
% 
% Unit resulting resolution: 321
% Factoring: 1
% Forward superposition: 1559
% Backward superposition: 1708
% Self superposition: 115
% Equality factoring: 26
% Equality resolution: 2
% 
% Split clauses: 31
% Split components: 62
% SAT solver clauses: 630
% SAT solver unit clauses: 119
% SAT solver binary clauses: 240
% 
% Sat splits: 31
% Sat splitting refutations: 5
% 
% Memory used [KB]: 6268
% Time elapsed: 0.228 s
% ------------------------------
% ------------------------------
% Success in time 1.545 s
