Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.real
% (74284)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 141
% Passive clauses: 3366
% Generated clauses: 24370
% Final active clauses: 138
% Final passive clauses: 3004
% Input formulas: 22
% Initial clauses: 39
% 
% Duplicate literals: 991
% Trivial inequalities: 306
% Fw subsumption resolutions: 3520
% Fw demodulations: 3360
% Bw demodulations: 2
% 
% Simple tautologies: 367
% Equational tautologies: 40
% Forward subsumptions: 11314
% 
% Binary resolution: 887
% Factoring: 5
% Forward superposition: 7759
% Backward superposition: 7086
% Self superposition: 183
% Equality factoring: 94
% Equality resolution: 19
% 
% Split clauses: 605
% Split components: 1225
% SAT solver clauses: 1181
% SAT solver unit clauses: 26
% SAT solver binary clauses: 242
% 
% Sat splits: 605
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 1176
% TWLsolver calls for satisfiability: 199
% 
% Memory used [KB]: 12409
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.real
% (74297)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 126
% Passive clauses: 689
% Generated clauses: 16640
% Final active clauses: 103
% Final passive clauses: 497
% Input formulas: 22
% Initial clauses: 39
% 
% Duplicate literals: 1117
% Trivial inequalities: 288
% Fw subsumption resolutions: 1325
% Fw demodulations: 957
% Bw demodulations: 87
% Condensations: 2068
% Global subsumptions: 5
% 
% Simple tautologies: 300
% Equational tautologies: 38
% Forward subsumptions: 7818
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 138
% Factoring: 1
% Forward superposition: 4382
% Backward superposition: 5247
% Self superposition: 148
% Equality factoring: 78
% Equality resolution: 29
% 
% Split clauses: 103
% Split components: 211
% SAT solver clauses: 1246
% SAT solver unit clauses: 47
% SAT solver binary clauses: 34
% SAT solver learnt clauses: 10
% 
% Sat splits: 103
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 1226
% TWLsolver calls for satisfiability: 217
% 
% Memory used [KB]: 10490
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.real
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.real
% SZS output start Proof for thm.real
fof(f21698,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f2996,f2989,f2966,f2948,f2928,f2919,f2916,f2900,f2895,f3048,f2889,f3007,f2886,f2881,f2871,f3015,f3075,f2866,f3047,f2843,f3094,f2842,f3095,f2837,f3096,f2836,f3097,f2333,f3134,f3163,f3164,f3172,f3173,f3174,f3175,f2700,f3181,f3318,f3320,f3328,f3330,f3341,f3346,f3363,f3367,f3369,f3370,f3574,f3576,f3568,f3578,f3594,f3539,f3602,f3606,f3608,f3610,f3655,f3772,f3773,f5651,f5683,f7181,f7189,f7195,f7196,f7372,f7383,f7407,f7408,f7530,f7535,f7531,f7542,f7544,f7545,f3135,f7553,f7555,f7556,f7558,f7559,f7561,f7562,f7565,f7566,f7567,f8066,f8072,f8074,f8073,f8067,f8133,f8134,f8235,f8252,f8810,f8820,f9513,f9555,f9800,f9801,f9961,f9983,f9985,f9995,f9977,f10089,f10126,f10857,f10867,f11088,f11111,f11205,f11206,f11976,f11990,f13740,f13741,f3049,f13960,f13961,f14173,f14175,f14177,f14181,f14186,f14192,f7382,f14193,f14496,f14498,f14499,f14513,f14519,f14522,f14530,f14536,f14538,f14548,f14570,f14572,f14575,f20899,f3165,f20900,f21061,f3008,f21063,f21120,f21601,f21603,f21604,f21608,f21624,f21626,f21634,f21637,f21639,f21649,f21653,f21655,f21658,f21665,f21670,f21673,f21675,f21693,f21695])).
fof(f21695,plain,(
  $false | ($spl10 | $spl12)),
  inference(global_subsumption,[],[f70,f71,f76,f80,f79,f102,f107,f104,f108,f110,f131,f162,f207,f325,f326,f376,f558,f856,f575,f598,f611,f891,f651,f917,f923,f713,f715,f769,f792,f780,f996,f958,f1072,f1074,f1093,f1168,f1222,f1224,f1226,f1227,f1663,f1664,f1297,f1739,f1783,f2996,f2989,f2966,f3031,f2025,f2948,f2928,f2919,f2916,f2900,f2895,f3048,f2889,f2886,f3059,f2122,f2881,f2871,f3075,f2866,f2843,f2842,f2837,f2836,f3107,f2283,f2285,f3133,f2333,f3137,f3143,f3163,f3172,f3174,f3179,f3180,f2700,f3200,f3202,f456,f1665,f938,f1795,f1740,f1223,f714,f263,f249,f111,f455,f454,f3771,f3772,f5038,f5039,f5044,f6426,f6508,f6509,f6512,f6513,f6514,f6529,f6558,f6565,f6841,f7181,f7195,f7372,f7407,f8066,f8072,f8074,f8125,f8126,f8133,f8249,f8251,f8235,f8264,f8292,f8294,f8314,f8338,f8339,f8340,f8810,f9513,f9961,f9983,f9985,f10000,f10008,f10089,f11088,f11205,f11362,f11401,f11403,f11594,f11597,f11649,f11651,f11987,f11989,f11976,f12029,f12037,f12377,f13602,f12425,f13740,f690,f12040,f12038,f12036,f12010,f11996,f11978,f11973,f8132,f8127,f6839,f6838,f3141,f258,f77,f1294,f1165,f766,f687,f648,f13840,f13824,f13808,f13776,f13775,f13624,f13603,f13601,f13591,f13550,f13532,f13348,f13313,f13215,f13203,f13089,f13084,f12868,f12854,f12849,f12847,f12846,f12844,f12843,f12841,f12840,f12837,f12834,f12829,f12674,f12669,f12493,f12492,f12491,f12490,f12427,f12424,f12423,f12375,f12374,f12373,f12372,f12371,f12327,f12324,f12323,f12322,f12321,f12041,f12011,f12009,f12008,f11998,f11995,f11780,f11779,f11778,f11777,f11652,f11650,f11581,f11555,f11404,f11402,f11363,f11361,f11209,f11200,f11199,f11194,f11192,f9563,f9562,f9546,f9476,f9475,f9474,f9378,f9377,f9375,f9107,f9082,f8926,f8925,f8923,f8844,f8842,f8841,f8727,f8726,f8724,f8571,f8570,f8569,f8568,f8341,f8135,f8068,f8062,f8060,f7411,f7406,f7097,f7090,f6947,f6946,f6945,f6893,f6890,f6850,f6848,f6845,f6693,f6692,f6691,f6586,f6582,f6581,f6580,f6555,f6523,f6522,f6521,f6520,f6519,f6518,f6517,f6516,f6507,f6427,f6425,f6124,f5816,f5324,f5323,f5322,f5307,f5300,f5279,f5275,f5274,f5273,f5272,f5271,f5270,f5269,f5268,f5266,f5265,f5262,f5261,f5260,f5194,f5193,f5091,f5090,f5047,f4996,f4966,f4382,f4379,f4324,f4321,f4245,f4173,f4143,f4099,f3219,f3201,f3170,f3148,f3095,f3094,f3076,f2835,f2831,f2829,f2588,f2585,f2503,f2502,f2500,f2499,f2498,f2495,f2494,f2492,f2489,f2429,f2426,f2366,f2363,f2282,f2279,f2219,f2216,f2156,f2155,f2125,f2121,f2091,f2061,f2024,f1994,f1993,f1952,f1951,f1801,f1799,f1094,f1050,f957,f937,f934,f921,f612,f576,f542,f525,f524,f523,f522,f493,f453,f417,f264,f219,f187,f118,f103,f84,f83,f73,f12035,f12034,f12033,f11842,f11841,f8097,f8096,f6748,f6747,f6589,f5327,f74,f12802,f12642,f5147,f2480,f2417,f2270,f2207,f1157,f681,f642,f6770,f5146,f2354,f469,f468,f11961,f10086,f11962,f12275,f12274,f11553,f11552,f11188,f11187,f11119,f11118,f11010,f11009,f10928,f10927,f10674,f10673,f10120,f10095,f10093,f10090,f10074,f10073,f10018,f10017,f9332,f9314,f9190,f9189,f9150,f9149,f9081,f9080,f8058,f8057,f7988,f7987,f7199,f7180,f7177,f7163,f7162,f7108,f7107,f6583,f6579,f6568,f6567,f6564,f6563,f6554,f6549,f6547,f6546,f6545,f6544,f6543,f6515,f6058,f5321,f5310,f5309,f5306,f5305,f5299,f5294,f5292,f5291,f5290,f5289,f5288,f5267,f13719,f8832,f8681,f8680,f8679,f8080,f7247,f7246,f7176,f9851,f9850,f7389,f7387,f9473,f8898,f8897,f8896,f8895,f8678,f6944,f6690,f5192,f4995,f4172,f2023,f1944,f915,f450,f416,f16317,f16368,f16369,f20133,f20134,f20135,f20136,f20527,f20723,f20686,f20687,f20688,f20850,f20854,f20855,f20892,f20893,f20898,f20899,f20903,f20907,f20913,f20914,f20932,f20918,f20919,f20920,f20922,f20923,f21054,f21053,f21025,f21026,f21061,f21064,f21065,f21066,f21067,f21068,f21076,f21078,f21081,f21114,f21115,f21116,f21117,f21118,f21124,f21602,f21128,f21607,f21623,f21625,f21633,f21638,f21648,f21652,f21664,f21669,f21672,f21674,f21692,f21694])).
fof(f21694,plain,(
  ( ! [X17,X16] : (s('type.realax.real','const.real.abs_1'(s('type.realax.real',X16))) = s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real',X17),s('type.realax.real','const.real.abs_1'(s('type.realax.real',X16))))) | s('type.realax.real',X16) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X16)))) ) | ($spl10 | $spl12)),
  inference(forward_demodulation,[],[f21045,f21124])).
fof(f21045,plain,(
  ( ! [X17,X16] : (s('type.realax.real','const.real.abs_1'(s('type.realax.real',X16))) = s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real',X17),s('type.realax.real','const.real.abs_1'(s('type.realax.real',X16))))) | s('type.realax.real',X16) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X16)))) ) | $spl12),
  inference(superposition,[],[f103,f20932])).
fof(f21692,plain,(
  ( ! [X14,X15] : (s('type.realax.real',X15) = s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real','const.real.abs_1'(s('type.realax.real',X14))),s('type.realax.real',X15))) | s('type.realax.real',X14) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X14)))) ) | ($spl10 | $spl12)),
  inference(forward_demodulation,[],[f21044,f21124])).
fof(f21044,plain,(
  ( ! [X14,X15] : (s('type.realax.real',X15) = s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real','const.real.abs_1'(s('type.realax.real',X14))),s('type.realax.real',X15))) | s('type.realax.real',X14) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X14)))) ) | $spl12),
  inference(superposition,[],[f103,f20932])).
fof(f21674,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl10 | $spl12)),
  inference(subsumption_resolution,[],[f21592,f79])).
fof(f21592,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | ($spl10 | $spl12)),
  inference(condensation,[],[f21591])).
fof(f21591,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X0))) ) | ($spl10 | $spl12)),
  inference(condensation,[],[f21563])).
fof(f21563,plain,(
  ( ! [X68,X66,X69] : (s(bool,'T_0') = s(bool,X69) | p(s(bool,X68)) | p(s(bool,X69)) | s(bool,'T_0') = s(bool,X66)) ) | ($spl10 | $spl12)),
  inference(backward_demodulation,[],[f21124,f6547])).
fof(f21672,plain,(
  ( ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0)))) ) | ($spl10 | $spl12)),
  inference(duplicate_literal_removal,[],[f21671])).
fof(f21671,plain,(
  ( ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0))) | s('type.realax.real',X0) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0)))) ) | ($spl10 | $spl12)),
  inference(forward_demodulation,[],[f21547,f102])).
fof(f21547,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0))) = s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real',X0),s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0))))) | s('type.realax.real',X0) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0)))) ) | ($spl10 | $spl12)),
  inference(backward_demodulation,[],[f21124,f12035])).
fof(f21669,plain,(
  ( ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0)))) ) | ($spl10 | $spl12)),
  inference(forward_demodulation,[],[f21533,f102])).
fof(f21533,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0))) = s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real',X0),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))) ) | ($spl10 | $spl12)),
  inference(backward_demodulation,[],[f21124,f74])).
fof(f21664,plain,(
  ( ! [X9] : (s('type.realax.real','const.real.abs_1'(s('type.realax.real',X9))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X9)))) ) | ($spl10 | $spl12)),
  inference(subsumption_resolution,[],[f21440,f21124])).
fof(f21440,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X9))) | s('type.realax.real','const.real.abs_1'(s('type.realax.real',X9))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X9)))) ) | ($spl10 | $spl12)),
  inference(backward_demodulation,[],[f21124,f11995])).
fof(f21652,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | ($spl10 | $spl12)),
  inference(condensation,[],[f21651])).
fof(f21651,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X1)) ) | ($spl10 | $spl12)),
  inference(condensation,[],[f21650])).
fof(f21650,plain,(
  ( ! [X70,X69,X67] : (s(bool,X69) = s(bool,X70) | s(bool,'F_0') = s(bool,X70) | s(bool,X67) = s(bool,X69)) ) | ($spl10 | $spl12)),
  inference(subsumption_resolution,[],[f21366,f21124])).
fof(f21366,plain,(
  ( ! [X70,X68,X69,X67] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X68),s(bool,X67))) | s(bool,X69) = s(bool,X70) | s(bool,'F_0') = s(bool,X70) | s(bool,X67) = s(bool,X69)) ) | ($spl10 | $spl12)),
  inference(backward_demodulation,[],[f21124,f9562])).
fof(f21648,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | ($spl10 | $spl12)),
  inference(condensation,[],[f21647])).
fof(f21647,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X1)) ) | ($spl10 | $spl12)),
  inference(duplicate_literal_removal,[],[f21646])).
fof(f21646,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X1)) ) | ($spl10 | $spl12)),
  inference(condensation,[],[f21645])).
fof(f21645,plain,(
  ( ! [X4,X0,X3,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X3) = s(bool,X4) | s(bool,'F_0') = s(bool,X4) | s(bool,X0) = s(bool,X3)) ) | ($spl10 | $spl12)),
  inference(subsumption_resolution,[],[f21337,f21124])).
fof(f21337,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X2))) | s(bool,'F_0') = s(bool,X1) | s(bool,X3) = s(bool,X4) | s(bool,'F_0') = s(bool,X4) | s(bool,X0) = s(bool,X3)) ) | ($spl10 | $spl12)),
  inference(backward_demodulation,[],[f21124,f8841])).
fof(f21638,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | ($spl10 | $spl12)),
  inference(subsumption_resolution,[],[f21321,f21124])).
fof(f21321,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,X1))) | s(bool,'F_0') = s(bool,X0)) ) | ($spl10 | $spl12)),
  inference(backward_demodulation,[],[f21124,f8080])).
fof(f21633,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | ($spl10 | $spl12)),
  inference(condensation,[],[f21632])).
fof(f21632,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X1)) ) | ($spl10 | $spl12)),
  inference(duplicate_literal_removal,[],[f21631])).
fof(f21631,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X1)) ) | ($spl10 | $spl12)),
  inference(condensation,[],[f21630])).
fof(f21630,plain,(
  ( ! [X52,X56,X55,X53] : (s(bool,'F_0') = s(bool,X53) | s(bool,X55) = s(bool,X56) | s(bool,'F_0') = s(bool,X56) | s(bool,X52) = s(bool,X55)) ) | ($spl10 | $spl12)),
  inference(subsumption_resolution,[],[f21319,f21124])).
fof(f21319,plain,(
  ( ! [X54,X52,X56,X55,X53] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,X53),s(bool,X52),s(bool,X54))) | s(bool,'F_0') = s(bool,X53) | s(bool,X55) = s(bool,X56) | s(bool,'F_0') = s(bool,X56) | s(bool,X52) = s(bool,X55)) ) | ($spl10 | $spl12)),
  inference(backward_demodulation,[],[f21124,f8068])).
fof(f21625,plain,(
  ( ! [X7] : (s(bool,'F_0') = s(bool,X7)) ) | ($spl10 | $spl12)),
  inference(subsumption_resolution,[],[f21298,f21124])).
fof(f21298,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X6),s(bool,X7))) | s(bool,'F_0') = s(bool,X7)) ) | ($spl10 | $spl12)),
  inference(backward_demodulation,[],[f21124,f7097])).
fof(f21623,plain,(
  ( ! [X3] : (s('type.realax.real',X3) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X3)))) ) | ($spl10 | $spl12)),
  inference(duplicate_literal_removal,[],[f21622])).
fof(f21622,plain,(
  ( ! [X3] : (s('type.realax.real',X3) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X3))) | s('type.realax.real',X3) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X3)))) ) | ($spl10 | $spl12)),
  inference(forward_demodulation,[],[f21287,f102])).
fof(f21287,plain,(
  ( ! [X3] : (s('type.realax.real','const.real.abs_1'(s('type.realax.real',X3))) = s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real',X3),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X3))))) | s('type.realax.real',X3) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X3)))) ) | ($spl10 | $spl12)),
  inference(backward_demodulation,[],[f21124,f6770])).
fof(f21607,plain,(
  ( ! [X12] : (s(bool,'F_0') = s(bool,X12)) ) | ($spl10 | $spl12)),
  inference(subsumption_resolution,[],[f21152,f21124])).
fof(f21152,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X12),s(bool,X13))) | s(bool,'F_0') = s(bool,X12)) ) | ($spl10 | $spl12)),
  inference(backward_demodulation,[],[f21124,f1944])).
fof(f21128,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) ) | ($spl10 | $spl12)),
  inference(backward_demodulation,[],[f21124,f103])).
fof(f21602,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl10 | $spl12)),
  inference(subsumption_resolution,[],[f21126,f21124])).
fof(f21126,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | ($spl10 | $spl12)),
  inference(backward_demodulation,[],[f21124,f83])).
fof(f21124,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl10 | $spl12)),
  inference(condensation,[],[f21123])).
fof(f21123,plain,(
  ( ! [X10,X11] : (s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11)) ) | ($spl10 | $spl12)),
  inference(subsumption_resolution,[],[f21122,f21076])).
fof(f21122,plain,(
  ( ! [X10,X11] : (s(bool,X10) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11)) ) | ($spl10 | $spl12)),
  inference(forward_demodulation,[],[f21091,f21065])).
fof(f21091,plain,(
  ( ! [X10,X11] : (s(bool,X10) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11)) ) | ($spl10 | $spl12)),
  inference(backward_demodulation,[],[f21065,f8340])).
fof(f21118,plain,(
  s('type.realax.real',sK0) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK0))) | ($spl10 | $spl12)),
  inference(forward_demodulation,[],[f21088,f21065])).
fof(f21088,plain,(
  s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK0))) | ($spl10 | $spl12)),
  inference(backward_demodulation,[],[f21065,f8314])).
fof(f21117,plain,(
  ( ! [X70,X72,X71] : (s(bool,'F_0') != s(bool,X71) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X70),s(bool,X71),s(bool,X72))) | s(bool,'F_0') = s(bool,X70)) ) | ($spl10 | $spl12)),
  inference(forward_demodulation,[],[f21087,f21067])).
fof(f21087,plain,(
  ( ! [X70,X72,X71] : (s(bool,X71) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X70),s(bool,X71),s(bool,X72))) | s(bool,'F_0') = s(bool,X70)) ) | ($spl10 | $spl12)),
  inference(backward_demodulation,[],[f21065,f6693])).
fof(f21116,plain,(
  ( ! [X68,X69,X67] : (s(bool,'F_0') != s(bool,X69) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X67),s(bool,X68),s(bool,X69))) | s(bool,'T_0') = s(bool,X67)) ) | ($spl10 | $spl12)),
  inference(forward_demodulation,[],[f21086,f21067])).
fof(f21086,plain,(
  ( ! [X68,X69,X67] : (s(bool,X69) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X67),s(bool,X68),s(bool,X69))) | s(bool,'T_0') = s(bool,X67)) ) | ($spl10 | $spl12)),
  inference(backward_demodulation,[],[f21065,f6692])).
fof(f21115,plain,(
  ( ! [X66,X65] : (s(bool,'F_0') != s(bool,X66) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X65),s(bool,X66)))) ) | ($spl10 | $spl12)),
  inference(forward_demodulation,[],[f21085,f21067])).
fof(f21085,plain,(
  ( ! [X66,X65] : (s(bool,X66) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X65),s(bool,X66)))) ) | ($spl10 | $spl12)),
  inference(backward_demodulation,[],[f21065,f6691])).
fof(f21114,plain,(
  ( ! [X64,X63] : (s(bool,'F_0') != s(bool,X63) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X63),s(bool,X64)))) ) | ($spl10 | $spl12)),
  inference(forward_demodulation,[],[f21084,f21067])).
fof(f21084,plain,(
  ( ! [X64,X63] : (s(bool,X63) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X63),s(bool,X64)))) ) | ($spl10 | $spl12)),
  inference(backward_demodulation,[],[f21065,f6690])).
fof(f21081,plain,(
  ( ! [X36] : (s(bool,X36) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,'F_0') = s(bool,X36)) ) | ($spl10 | $spl12)),
  inference(backward_demodulation,[],[f21065,f4099])).
fof(f21078,plain,(
  ( ! [X76,X77] : (s(bool,X76) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,X76) = s(bool,X77) | s(bool,'F_0') = s(bool,X77)) ) | ($spl10 | $spl12)),
  inference(backward_demodulation,[],[f21065,f3076])).
fof(f21076,plain,(
  ( ! [X64,X65] : (s(bool,X64) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X65) | s(bool,X64) = s(bool,X65)) ) | ($spl10 | $spl12)),
  inference(backward_demodulation,[],[f21065,f1795])).
fof(f21068,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl10 | $spl12)),
  inference(backward_demodulation,[],[f21065,f249])).
fof(f21067,plain,(
  s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | ($spl10 | $spl12)),
  inference(backward_demodulation,[],[f21065,f118])).
fof(f21066,plain,(
  s(bool,'T_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | ($spl10 | $spl12)),
  inference(backward_demodulation,[],[f21065,f111])).
fof(f21065,plain,(
  s('type.realax.real',sK0) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))) | ($spl10 | $spl12)),
  inference(global_subsumption,[],[f70,f71,f76,f80,f79,f102,f107,f104,f108,f110,f131,f162,f207,f325,f326,f376,f558,f856,f575,f598,f611,f891,f651,f917,f923,f713,f715,f769,f792,f780,f996,f958,f1072,f1074,f1093,f1168,f1222,f1224,f1226,f1227,f1663,f1664,f1297,f1739,f1783,f2996,f2989,f2966,f3031,f2025,f2948,f2928,f2919,f2916,f2900,f2895,f3048,f2889,f2886,f3059,f2122,f2881,f2871,f3075,f2866,f2843,f2842,f2837,f2836,f3107,f2283,f2285,f3133,f2333,f3137,f3143,f3163,f3172,f3174,f3179,f3180,f2700,f3200,f3202,f456,f1665,f938,f1795,f1740,f1223,f714,f263,f249,f111,f455,f454,f3771,f3772,f5038,f5039,f5044,f6426,f6508,f6509,f6512,f6513,f6514,f6529,f6558,f6565,f6841,f7181,f7195,f7372,f7407,f8066,f8072,f8074,f8125,f8126,f8133,f8249,f8251,f8235,f8264,f8292,f8294,f8314,f8338,f8339,f8340,f8810,f9513,f9961,f9983,f9985,f10000,f10008,f10089,f11088,f11205,f11362,f11401,f11403,f11594,f11597,f11649,f11651,f11987,f11989,f11976,f12029,f12037,f12377,f13602,f12425,f13740,f690,f12040,f12038,f12036,f12010,f11996,f11978,f11973,f8132,f8127,f6839,f6838,f3141,f258,f77,f1294,f1165,f766,f687,f648,f13840,f13824,f13808,f13776,f13775,f13624,f13603,f13601,f13591,f13550,f13532,f13348,f13313,f13215,f13203,f13089,f13084,f12868,f12854,f12849,f12847,f12846,f12844,f12843,f12841,f12840,f12837,f12834,f12829,f12674,f12669,f12493,f12492,f12491,f12490,f12427,f12424,f12423,f12375,f12374,f12373,f12372,f12371,f12327,f12324,f12323,f12322,f12321,f12041,f12011,f12009,f12008,f11998,f11995,f11780,f11779,f11778,f11777,f11652,f11650,f11581,f11555,f11404,f11402,f11363,f11361,f11209,f11200,f11199,f11194,f11192,f9563,f9562,f9546,f9476,f9475,f9474,f9378,f9377,f9375,f9107,f9082,f8926,f8925,f8923,f8844,f8842,f8841,f8727,f8726,f8724,f8571,f8570,f8569,f8568,f8341,f8135,f8068,f8062,f8060,f7411,f7406,f7097,f7090,f6947,f6946,f6945,f6893,f6890,f6850,f6848,f6845,f6693,f6692,f6691,f6586,f6582,f6581,f6580,f6555,f6523,f6522,f6521,f6520,f6519,f6518,f6517,f6516,f6507,f6427,f6425,f6124,f5816,f5324,f5323,f5322,f5307,f5300,f5279,f5275,f5274,f5273,f5272,f5271,f5270,f5269,f5268,f5266,f5265,f5262,f5261,f5260,f5194,f5193,f5091,f5090,f5047,f4996,f4966,f4382,f4379,f4324,f4321,f4245,f4173,f4143,f4099,f3219,f3201,f3170,f3148,f3095,f3094,f3076,f2835,f2831,f2829,f2588,f2585,f2503,f2502,f2500,f2499,f2498,f2495,f2494,f2492,f2489,f2429,f2426,f2366,f2363,f2282,f2279,f2219,f2216,f2156,f2155,f2125,f2121,f2091,f2061,f2024,f1994,f1993,f1952,f1951,f1801,f1799,f1094,f1050,f957,f937,f934,f921,f612,f576,f542,f525,f524,f523,f522,f493,f453,f417,f264,f219,f187,f118,f103,f84,f83,f73,f12035,f12034,f12033,f11842,f11841,f8097,f8096,f6748,f6747,f6589,f5327,f74,f12802,f12642,f5147,f2480,f2417,f2270,f2207,f1157,f681,f642,f6770,f5146,f2354,f469,f468,f11961,f10086,f11962,f12275,f12274,f11553,f11552,f11188,f11187,f11119,f11118,f11010,f11009,f10928,f10927,f10674,f10673,f10120,f10095,f10093,f10090,f10074,f10073,f10018,f10017,f9332,f9314,f9190,f9189,f9150,f9149,f9081,f9080,f8058,f8057,f7988,f7987,f7199,f7180,f7177,f7163,f7162,f7108,f7107,f6583,f6579,f6568,f6567,f6564,f6563,f6554,f6549,f6547,f6546,f6545,f6544,f6543,f6515,f6058,f5321,f5310,f5309,f5306,f5305,f5299,f5294,f5292,f5291,f5290,f5289,f5288,f5267,f13719,f8832,f8681,f8680,f8679,f8080,f7247,f7246,f7176,f9851,f9850,f7389,f7387,f9473,f8898,f8897,f8896,f8895,f8678,f6944,f6690,f5192,f4995,f4172,f2023,f1944,f915,f450,f416,f16317,f16368,f16369,f20133,f20134,f20135,f20136,f20527,f20723,f20686,f20687,f20688,f20850,f20854,f20855,f20892,f20893,f20898,f20899,f20903,f20907,f20913,f20914,f20932,f20918,f20919,f20920,f20922,f20923,f21054,f21053,f21025,f21026,f21061,f21064])).
fof(f21064,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,X1) | s('type.realax.real',sK0) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0)))) ) | ($spl10 | $spl12)),
  inference(subsumption_resolution,[],[f21031,f4099])).
fof(f21031,plain,(
  ( ! [X1] : (s(bool,X1) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))) | s(bool,'F_0') = s(bool,X1) | s('type.realax.real',sK0) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0)))) ) | $spl12),
  inference(superposition,[],[f8341,f20932])).
fof(f21026,plain,(
  ( ! [X12,X13,X11] : (s('type.realax.real','const.real.abs_1'(s('type.realax.real',X13))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',X13))))) | s('type.realax.real','const.bool.COND_3'(s(bool,X11),s('type.realax.real',X12),s('type.realax.real',X13))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.bool.COND_3'(s(bool,X11),s('type.realax.real',X12),s('type.realax.real',X13))))) | s(bool,'T_0') = s(bool,X11)) ) | $spl12),
  inference(superposition,[],[f20932,f207])).
fof(f21025,plain,(
  ( ! [X10,X8,X9] : (s('type.realax.real','const.real.abs_1'(s('type.realax.real',X9))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',X9))))) | s('type.realax.real','const.bool.COND_3'(s(bool,X8),s('type.realax.real',X9),s('type.realax.real',X10))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.bool.COND_3'(s(bool,X8),s('type.realax.real',X9),s('type.realax.real',X10))))) | s(bool,'F_0') = s(bool,X8)) ) | $spl12),
  inference(superposition,[],[f20932,f3095])).
fof(f21053,plain,(
  ( ! [X1] : (s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))))))) | s('type.realax.real',X1) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X1)))) ) | $spl12),
  inference(duplicate_literal_removal,[],[f21019])).
fof(f21019,plain,(
  ( ! [X1] : (s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))))))) | s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))))))) | s('type.realax.real',X1) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X1)))) ) | $spl12),
  inference(superposition,[],[f20932,f20892])).
fof(f21054,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0))))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0))))))) | s('type.realax.real',X0) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0)))) ) | $spl12),
  inference(duplicate_literal_removal,[],[f21018])).
fof(f21018,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0))))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0))))))) | s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0))))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0))))))) | s('type.realax.real',X0) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0)))) ) | $spl12),
  inference(superposition,[],[f20932,f20932])).
fof(f20923,plain,(
  ( ! [X6] : (s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X6))))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X6))))) | s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X6))))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X6))))))) | s('type.realax.real',X6) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X6)))) ) | $spl12),
  inference(superposition,[],[f11987,f20892])).
fof(f20922,plain,(
  ( ! [X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X5)))))) | s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X5))))))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X5))))))))) | s('type.realax.real',X5) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X5)))) ) | $spl12),
  inference(superposition,[],[f6838,f20892])).
fof(f20920,plain,(
  ( ! [X3] : (s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X3))))) | s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X3))))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X3))))))) | s('type.realax.real',X3) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X3)))) ) | $spl12),
  inference(superposition,[],[f5279,f20892])).
fof(f20919,plain,(
  ( ! [X2] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2)))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))))))) | s('type.realax.real',X2) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X2)))) ) | $spl12),
  inference(superposition,[],[f258,f20892])).
fof(f20918,plain,(
  ( ! [X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1)))))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1)))))))) | s('type.realax.real',X1) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X1)))) ) | $spl12),
  inference(superposition,[],[f77,f20892])).
fof(f20932,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0))))) | s('type.realax.real',X0) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0)))) ) | $spl12),
  inference(duplicate_literal_removal,[],[f20916])).
fof(f20916,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0))))) | s('type.realax.real',X0) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0))) | s('type.realax.real',X0) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0)))) ) | $spl12),
  inference(superposition,[],[f20892,f11987])).
fof(f20914,plain,(
  ( ! [X12,X10,X11] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X12))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X12))))) | s('type.realax.real','const.bool.COND_3'(s(bool,X10),s('type.realax.real',X11),s('type.realax.real',X12))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.bool.COND_3'(s(bool,X10),s('type.realax.real',X11),s('type.realax.real',X12))))) | s(bool,'T_0') = s(bool,X10)) ) | $spl12),
  inference(superposition,[],[f20892,f207])).
fof(f20913,plain,(
  ( ! [X8,X7,X9] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X8))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X8))))) | s('type.realax.real','const.bool.COND_3'(s(bool,X7),s('type.realax.real',X8),s('type.realax.real',X9))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.bool.COND_3'(s(bool,X7),s('type.realax.real',X8),s('type.realax.real',X9))))) | s(bool,'F_0') = s(bool,X7)) ) | $spl12),
  inference(superposition,[],[f20892,f3095])).
fof(f20907,plain,(
  ( ! [X0] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))))) | s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))))) | s('type.realax.real',X0) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0)))) ) | $spl12),
  inference(superposition,[],[f20892,f20892])).
fof(f20903,plain,(
  ( ! [X41,X42] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X41))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X41))))) | s(bool,X42) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X41))) | s(bool,'F_0') = s(bool,X42)) ) | $spl12),
  inference(subsumption_resolution,[],[f20877,f73])).
fof(f20877,plain,(
  ( ! [X41,X42] : (p(s(bool,'F_0')) | s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X41))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X41))))) | s(bool,X42) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X41))) | s(bool,'F_0') = s(bool,X42)) ) | $spl12),
  inference(superposition,[],[f6838,f3059])).
fof(f20898,plain,(
  ( ! [X30,X31,X29] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X29))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X29))))) | s(bool,X30) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X29))) | s(bool,X30) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) ) | $spl12),
  inference(subsumption_resolution,[],[f20873,f73])).
fof(f20873,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,'F_0')) | s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X29))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X29))))) | s(bool,X30) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X29))) | s(bool,X30) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) ) | $spl12),
  inference(superposition,[],[f6838,f1664])).
fof(f20893,plain,(
  ( ! [X4] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X4))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X4))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X4)))) ) | $spl12),
  inference(subsumption_resolution,[],[f20862,f73])).
fof(f20862,plain,(
  ( ! [X4] : (p(s(bool,'F_0')) | s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X4))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X4))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X4)))) ) | $spl12),
  inference(superposition,[],[f6838,f76])).
fof(f20892,plain,(
  ( ! [X2] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))) | s('type.realax.real',X2) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X2)))) ) | $spl12),
  inference(subsumption_resolution,[],[f20860,f73])).
fof(f20860,plain,(
  ( ! [X2] : (p(s(bool,'F_0')) | s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))) | s('type.realax.real',X2) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X2)))) ) | $spl12),
  inference(superposition,[],[f6838,f5279])).
fof(f20855,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X11)))) | s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.bool.COND_3'(s(bool,X9),s('type.realax.real',X10),s('type.realax.real',X11))))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.bool.COND_3'(s(bool,X9),s('type.realax.real',X10),s('type.realax.real',X11))))))) | s(bool,'T_0') = s(bool,X9)) ) | $spl12),
  inference(superposition,[],[f6838,f207])).
fof(f20854,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X7)))) | s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real',X7),s('type.realax.real',X8))))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real',X7),s('type.realax.real',X8))))))) | s(bool,'F_0') = s(bool,X6)) ) | $spl12),
  inference(superposition,[],[f6838,f3095])).
fof(f20850,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0)))))) | s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))))) | s('type.realax.real',X0) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0)))) ) | $spl12),
  inference(superposition,[],[f6838,f11987])).
fof(f20688,plain,(
  ( ! [X80,X78,X79] : (s(bool,X80) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X78),s(bool,X79),s(bool,X80))) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f20527,f207])).
fof(f20687,plain,(
  ( ! [X76,X77,X75] : (s(bool,X76) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X75),s(bool,X76),s(bool,X77))) | s(bool,'F_0') = s(bool,X75)) ) | $spl12),
  inference(superposition,[],[f20527,f3095])).
fof(f20686,plain,(
  ( ! [X74,X73] : (s(bool,X74) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X73),s(bool,X74)))) )),
  inference(superposition,[],[f20527,f103])).
fof(f20723,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1))) | s(bool,X0) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0)))))))) )),
  inference(condensation,[],[f20722])).
fof(f20722,plain,(
  ( ! [X14,X15,X13] : (s(bool,X14) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))))) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X14),s(bool,X15)))) )),
  inference(duplicate_literal_removal,[],[f20662])).
fof(f20662,plain,(
  ( ! [X14,X15,X13] : (s(bool,X14) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))))) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X14),s(bool,X15))) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f20527,f326])).
fof(f20527,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(condensation,[],[f20526])).
fof(f20526,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(condensation,[],[f20525])).
fof(f20525,plain,(
  ( ! [X88,X87,X86] : (s(bool,X87) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))))) | s(bool,'T_0') = s(bool,X88) | s(bool,X86) = s(bool,X88) | s(bool,'T_0') = s(bool,X87)) )),
  inference(subsumption_resolution,[],[f20200,f1663])).
fof(f20200,plain,(
  ( ! [X88,X87,X86] : (s(bool,X87) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))))) | s(bool,'T_0') = s(bool,X88) | s(bool,X86) = s(bool,X88) | s(bool,'T_0') = s(bool,X87) | s(bool,'F_0') = s(bool,X86)) )),
  inference(superposition,[],[f1783,f8341])).
fof(f20136,plain,(
  ( ! [X80,X78,X79] : (s(bool,X80) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X78),s(bool,X79),s(bool,X80))) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f8341,f207])).
fof(f20135,plain,(
  ( ! [X76,X77,X75] : (s(bool,X76) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X75),s(bool,X76),s(bool,X77))) | s(bool,'F_0') = s(bool,X75)) ) | $spl12),
  inference(superposition,[],[f8341,f3095])).
fof(f20134,plain,(
  ( ! [X74,X73] : (s(bool,X74) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X73),s(bool,X74)))) )),
  inference(superposition,[],[f8341,f103])).
fof(f20133,plain,(
  ( ! [X72,X71] : (s(bool,X71) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X71),s(bool,X72)))) )),
  inference(superposition,[],[f8341,f102])).
fof(f16369,plain,(
  ( ! [X118,X116,X119,X117] : (s(bool,'T_0') != s(bool,X118) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X119),s(bool,'const.bool.COND_3'(s(bool,X116),s(bool,X117),s(bool,X118))))) | s(bool,'T_0') = s(bool,X116)) )),
  inference(superposition,[],[f13348,f207])).
fof(f16368,plain,(
  ( ! [X114,X112,X115,X113] : (s(bool,'T_0') != s(bool,X113) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X115),s(bool,'const.bool.COND_3'(s(bool,X112),s(bool,X113),s(bool,X114))))) | s(bool,'F_0') = s(bool,X112)) ) | $spl12),
  inference(superposition,[],[f13348,f3095])).
fof(f16317,plain,(
  ( ! [X14,X17,X15,X13,X16] : (s(bool,X13) != s(bool,X14) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X15),s(bool,X14))) | s(bool,X13) = s(bool,X16) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X16),s(bool,X17)))) )),
  inference(superposition,[],[f13348,f326])).
fof(f416,plain,(
  ( ! [X19,X17,X18] : (s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X17),s(bool,X18))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X17),s(bool,X19)))) )),
  inference(forward_demodulation,[],[f339,f118])).
fof(f339,plain,(
  ( ! [X19,X17,X18] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X17),s(bool,X19))) | s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X17),s(bool,X18))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0)))))) )),
  inference(superposition,[],[f263,f102])).
fof(f450,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X6),s(bool,X7))) | s(bool,X6) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0)))))) )),
  inference(subsumption_resolution,[],[f393,f71])).
fof(f393,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X6),s(bool,X7))) | s(bool,X6) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0)))))) )),
  inference(superposition,[],[f83,f263])).
fof(f915,plain,(
  ( ! [X12,X13,X11] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X11),s(bool,X12))) | s(bool,X11) = s(bool,X13) | s(bool,'T_0') = s(bool,X13)) )),
  inference(subsumption_resolution,[],[f694,f71])).
fof(f694,plain,(
  ( ! [X12,X13,X11] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X11),s(bool,X12))) | s(bool,X11) = s(bool,X13) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f83,f326])).
fof(f1944,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X12),s(bool,X13))) | s(bool,'F_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f1865,f71])).
fof(f1865,plain,(
  ( ! [X12,X13] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X12),s(bool,X13))) | s(bool,'F_0') = s(bool,X12)) )),
  inference(superposition,[],[f83,f187])).
fof(f2023,plain,(
  ( ! [X118,X116,X117,X115] : (s(bool,X115) = s(bool,X117) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X115),s(bool,X116))) | s(bool,X117) = s(bool,X118) | s(bool,'T_0') = s(bool,X118)) )),
  inference(superposition,[],[f1664,f102])).
fof(f4172,plain,(
  ( ! [X90,X91,X89] : (s(bool,X89) = s(bool,X91) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X89),s(bool,X90))) | s(bool,'F_0') = s(bool,X91)) )),
  inference(superposition,[],[f3059,f102])).
fof(f4995,plain,(
  ( ! [X90,X91,X89] : (s(bool,X89) != s(bool,X91) | s(bool,'T_0') = s(bool,X91) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X89),s(bool,X90)))) )),
  inference(superposition,[],[f3107,f102])).
fof(f5192,plain,(
  ( ! [X146,X149,X147,X145,X148] : (s(X147,X148) = s(X147,'const.bool.COND_3'(s(bool,X145),s(X147,X148),s(X147,X149))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X145),s(bool,X146)))) ) | $spl12),
  inference(superposition,[],[f3095,f102])).
fof(f6690,plain,(
  ( ! [X64,X63] : (s(bool,X63) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X63),s(bool,X64)))) ) | $spl10),
  inference(superposition,[],[f4099,f102])).
fof(f6944,plain,(
  ( ! [X109,X107,X108] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X109),s(bool,X107))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X107),s(bool,X108)))) )),
  inference(superposition,[],[f219,f102])).
fof(f8678,plain,(
  ( ! [X109,X110,X108] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X108),s(bool,'T_0'),s(bool,X110))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X108),s(bool,X109)))) ) | $spl12),
  inference(superposition,[],[f8080,f102])).
fof(f8895,plain,(
  ( ! [X109,X110,X108] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X108),s(bool,X110))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X108),s(bool,X109)))) )),
  inference(superposition,[],[f1944,f102])).
fof(f8896,plain,(
  ( ! [X111,X112,X113] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X112),s(bool,X113))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X111),s(bool,X112)))) )),
  inference(superposition,[],[f1944,f103])).
fof(f8897,plain,(
  ( ! [X116,X114,X117,X115] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X116),s(bool,X117))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X114),s(bool,X115),s(bool,X116))) | s(bool,'T_0') = s(bool,X114)) )),
  inference(superposition,[],[f1944,f207])).
fof(f8898,plain,(
  ( ! [X121,X118,X120,X119] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X119),s(bool,X121))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X118),s(bool,X119),s(bool,X120))) | s(bool,'F_0') = s(bool,X118)) ) | $spl12),
  inference(superposition,[],[f1944,f3095])).
fof(f9473,plain,(
  ( ! [X109,X110,X108] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X110),s(bool,X108))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X108),s(bool,X109)))) )),
  inference(superposition,[],[f7097,f102])).
fof(f7387,plain,(
  ( ! [X70,X68,X69,X67] : (p(s(bool,'const.bool.COND_3'(s(bool,X67),s(bool,X67),s(bool,X68)))) | s(bool,X69) = s(bool,X70) | s(bool,'F_0') = s(bool,X70) | s(bool,X67) = s(bool,X69)) ) | $spl12),
  inference(subsumption_resolution,[],[f7229,f110])).
fof(f7229,plain,(
  ( ! [X70,X68,X69,X67] : (p(s(bool,'const.bool.COND_3'(s(bool,X67),s(bool,X67),s(bool,X68)))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X69) = s(bool,X70) | s(bool,'F_0') = s(bool,X70) | s(bool,X67) = s(bool,X69)) ) | $spl12),
  inference(superposition,[],[f7176,f1664])).
fof(f7389,plain,(
  ( ! [X78,X79,X77] : (p(s(bool,'const.bool.COND_3'(s(bool,X77),s(bool,X77),s(bool,X78)))) | s(bool,'T_0') = s(bool,X79) | s(bool,X77) = s(bool,X79)) ) | $spl12),
  inference(forward_demodulation,[],[f7388,f102])).
fof(f7388,plain,(
  ( ! [X80,X78,X79,X77] : (p(s(bool,'const.bool.COND_3'(s(bool,X77),s(bool,X77),s(bool,X78)))) | s(bool,X77) = s(bool,X79) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X79),s(bool,X80)))) ) | $spl12),
  inference(subsumption_resolution,[],[f7232,f110])).
fof(f7232,plain,(
  ( ! [X80,X78,X79,X77] : (p(s(bool,'const.bool.COND_3'(s(bool,X77),s(bool,X77),s(bool,X78)))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X77) = s(bool,X79) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X79),s(bool,X80)))) ) | $spl12),
  inference(superposition,[],[f7176,f326])).
fof(f9850,plain,(
  ( ! [X158,X156,X159,X157,X160] : (p(s(bool,'const.bool.COND_3'(s(bool,X158),s(bool,X158),s(bool,X159)))) | s(bool,'T_0') = s(bool,X160) | s(bool,X160) = s(bool,'const.bool.COND_3'(s(bool,X156),s(bool,X157),s(bool,X158))) | s(bool,'T_0') = s(bool,X156)) ) | $spl12),
  inference(superposition,[],[f7389,f207])).
fof(f9851,plain,(
  ( ! [X165,X163,X161,X164,X162] : (p(s(bool,'const.bool.COND_3'(s(bool,X162),s(bool,X162),s(bool,X164)))) | s(bool,'T_0') = s(bool,X165) | s(bool,X165) = s(bool,'const.bool.COND_3'(s(bool,X161),s(bool,X162),s(bool,X163))) | s(bool,'F_0') = s(bool,X161)) ) | $spl12),
  inference(superposition,[],[f7389,f3095])).
fof(f7176,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,X1)))) | s(bool,'F_0') = s(bool,X0)) ) | $spl12),
  inference(condensation,[],[f7175])).
fof(f7175,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,X1)))) | s(bool,'F_0') = s(bool,X2)) ) | $spl12),
  inference(duplicate_literal_removal,[],[f7174])).
fof(f7174,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,X1)))) | s(bool,'F_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X2)) ) | $spl12),
  inference(condensation,[],[f7173])).
fof(f7173,plain,(
  ( ! [X90,X88,X87,X91,X89] : (p(s(bool,'const.bool.COND_3'(s(bool,X88),s(bool,'T_0'),s(bool,X89)))) | s(bool,'F_0') = s(bool,X88) | s(bool,X90) = s(bool,X91) | s(bool,'F_0') = s(bool,X91) | s(bool,X87) = s(bool,X90)) ) | $spl12),
  inference(trivial_inequality_removal,[],[f7145])).
fof(f7145,plain,(
  ( ! [X90,X88,X87,X91,X89] : (s(bool,X87) != s(bool,X87) | p(s(bool,'const.bool.COND_3'(s(bool,X88),s(bool,'T_0'),s(bool,X89)))) | s(bool,'F_0') = s(bool,X88) | s(bool,X90) = s(bool,X91) | s(bool,'F_0') = s(bool,X91) | s(bool,X87) = s(bool,X90)) ) | $spl12),
  inference(superposition,[],[f5288,f1664])).
fof(f7246,plain,(
  ( ! [X116,X114,X115,X113] : (p(s(bool,'const.bool.COND_3'(s(bool,X115),s(bool,'T_0'),s(bool,X116)))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X113),s(bool,X114),s(bool,X115))) | s(bool,'T_0') = s(bool,X113)) ) | $spl12),
  inference(superposition,[],[f7176,f207])).
fof(f7247,plain,(
  ( ! [X118,X120,X119,X117] : (p(s(bool,'const.bool.COND_3'(s(bool,X118),s(bool,'T_0'),s(bool,X120)))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X117),s(bool,X118),s(bool,X119))) | s(bool,'F_0') = s(bool,X117)) ) | $spl12),
  inference(superposition,[],[f7176,f3095])).
fof(f8080,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,X1))) | s(bool,'F_0') = s(bool,X0)) ) | $spl12),
  inference(condensation,[],[f8079])).
fof(f8079,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,X1))) | s(bool,'F_0') = s(bool,X2)) ) | $spl12),
  inference(duplicate_literal_removal,[],[f8078])).
fof(f8078,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,X1))) | s(bool,'F_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X2)) ) | $spl12),
  inference(condensation,[],[f8077])).
fof(f8077,plain,(
  ( ! [X90,X88,X87,X91,X89] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X88),s(bool,'T_0'),s(bool,X89))) | s(bool,'F_0') = s(bool,X88) | s(bool,X90) = s(bool,X91) | s(bool,'F_0') = s(bool,X91) | s(bool,X87) = s(bool,X90)) ) | $spl12),
  inference(subsumption_resolution,[],[f8040,f3148])).
fof(f8040,plain,(
  ( ! [X90,X88,X87,X91,X89] : (~p(s(bool,X87)) | s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X88),s(bool,'T_0'),s(bool,X89))) | s(bool,'F_0') = s(bool,X88) | s(bool,X90) = s(bool,X91) | s(bool,'F_0') = s(bool,X91) | s(bool,X87) = s(bool,X90)) ) | $spl12),
  inference(superposition,[],[f5290,f1664])).
fof(f8679,plain,(
  ( ! [X111,X112,X113] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X112),s(bool,'T_0'),s(bool,X113))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X111),s(bool,X112)))) ) | $spl12),
  inference(superposition,[],[f8080,f103])).
fof(f8680,plain,(
  ( ! [X116,X114,X117,X115] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X116),s(bool,'T_0'),s(bool,X117))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X114),s(bool,X115),s(bool,X116))) | s(bool,'T_0') = s(bool,X114)) ) | $spl12),
  inference(superposition,[],[f8080,f207])).
fof(f8681,plain,(
  ( ! [X121,X118,X120,X119] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X119),s(bool,'T_0'),s(bool,X121))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X118),s(bool,X119),s(bool,X120))) | s(bool,'F_0') = s(bool,X118)) ) | $spl12),
  inference(superposition,[],[f8080,f3095])).
fof(f8832,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,'T_0'),s(bool,X1))),s(bool,X2))) | s(bool,'F_0') = s(bool,X0)) ) | $spl12),
  inference(condensation,[],[f8831])).
fof(f8831,plain,(
  ( ! [X30,X33,X31,X32] : (s(bool,'F_0') = s(bool,X30) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,X30),s(bool,'T_0'),s(bool,X31))),s(bool,X33))) | s(bool,'T_0') = s(bool,X32)) ) | $spl12),
  inference(subsumption_resolution,[],[f8752,f76])).
fof(f8752,plain,(
  ( ! [X30,X33,X31,X32] : (s(bool,'F_0') != s(bool,X32) | s(bool,'F_0') = s(bool,X30) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,X30),s(bool,'T_0'),s(bool,X31))),s(bool,X33))) | s(bool,'T_0') = s(bool,X32)) ) | $spl12),
  inference(superposition,[],[f8080,f326])).
fof(f13719,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'const.bool.COND_3'(s(bool,X1),s(bool,'T_0'),s(bool,X2))))) | s(bool,'F_0') = s(bool,X1)) ) | $spl12),
  inference(condensation,[],[f13718])).
fof(f13718,plain,(
  ( ! [X592,X594,X593,X591] : (s(bool,'F_0') = s(bool,X591) | s(bool,'T_0') = s(bool,X593) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X594),s(bool,'const.bool.COND_3'(s(bool,X591),s(bool,'T_0'),s(bool,X592)))))) ) | $spl12),
  inference(subsumption_resolution,[],[f12653,f76])).
fof(f12653,plain,(
  ( ! [X592,X594,X593,X591] : (s(bool,'F_0') != s(bool,X593) | s(bool,'F_0') = s(bool,X591) | s(bool,'T_0') = s(bool,X593) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X594),s(bool,'const.bool.COND_3'(s(bool,X591),s(bool,'T_0'),s(bool,X592)))))) ) | $spl12),
  inference(superposition,[],[f8080,f1050])).
fof(f5267,plain,(
  ( ! [X35,X33,X36,X34,X32] : (s(bool,X33) = s(bool,X35) | s(bool,'F_0') = s(bool,X32) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,X32),s(bool,X33),s(bool,X34))),s(bool,X36))) | s(bool,'T_0') = s(bool,X35)) ) | $spl12),
  inference(superposition,[],[f3095,f326])).
fof(f5288,plain,(
  ( ! [X35,X36,X34] : (s(bool,'T_0') != s(bool,X35) | p(s(bool,'const.bool.COND_3'(s(bool,X34),s(bool,X35),s(bool,X36)))) | s(bool,'F_0') = s(bool,X34)) ) | $spl12),
  inference(superposition,[],[f79,f3095])).
fof(f5289,plain,(
  ( ! [X39,X37,X38] : (~p(s(bool,X38)) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X37),s(bool,X38),s(bool,X39))) | s(bool,'F_0') = s(bool,X37)) ) | $spl12),
  inference(superposition,[],[f80,f3095])).
fof(f5290,plain,(
  ( ! [X41,X42,X40] : (~p(s(bool,X41)) | s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X40),s(bool,X41),s(bool,X42))) | s(bool,'F_0') = s(bool,X40)) ) | $spl12),
  inference(superposition,[],[f83,f3095])).
fof(f5291,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,X44)) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X43),s(bool,X44),s(bool,X45))) | s(bool,'F_0') = s(bool,X43)) ) | $spl12),
  inference(superposition,[],[f84,f3095])).
fof(f5292,plain,(
  ( ! [X47,X48,X46,X49] : (p(s(bool,X47)) | p(s(bool,X49)) | s(bool,X49) = s(bool,'const.bool.COND_3'(s(bool,X46),s(bool,X47),s(bool,X48))) | s(bool,'F_0') = s(bool,X46)) ) | $spl12),
  inference(superposition,[],[f104,f3095])).
fof(f5294,plain,(
  ( ! [X57,X54,X56,X55] : (~p(s(bool,X55)) | ~p(s(bool,X57)) | s(bool,X57) = s(bool,'const.bool.COND_3'(s(bool,X54),s(bool,X55),s(bool,X56))) | s(bool,'F_0') = s(bool,X54)) ) | $spl12),
  inference(superposition,[],[f107,f3095])).
fof(f5299,plain,(
  ( ! [X76,X74,X77,X75] : (p(s(bool,X75)) | s(bool,X77) = s(bool,'const.bool.COND_3'(s(bool,X74),s(bool,X75),s(bool,X76))) | s(bool,'T_0') = s(bool,X77) | s(bool,'F_0') = s(bool,X74)) ) | $spl12),
  inference(superposition,[],[f162,f3095])).
fof(f5305,plain,(
  ( ! [X103,X105,X106,X104] : (s(bool,'T_0') != s(bool,X104) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,X103),s(bool,X104),s(bool,X105))),s(bool,X106))) | s(bool,'F_0') = s(bool,X103)) ) | $spl12),
  inference(superposition,[],[f376,f3095])).
fof(f5306,plain,(
  ( ! [X109,X107,X110,X108] : (~p(s(bool,X108)) | s(bool,X110) = s(bool,'const.bool.COND_3'(s(bool,X107),s(bool,X108),s(bool,X109))) | s(bool,'F_0') = s(bool,X110) | s(bool,'F_0') = s(bool,X107)) ) | $spl12),
  inference(superposition,[],[f453,f3095])).
fof(f5309,plain,(
  ( ! [X123,X121,X122,X120,X119] : (~p(s(bool,X120)) | s(bool,X122) = s(bool,'const.bool.COND_3'(s(bool,X119),s(bool,X120),s(bool,X121))) | s(bool,X122) = s(bool,X123) | s(bool,'T_0') = s(bool,X123) | s(bool,'F_0') = s(bool,X119)) ) | $spl12),
  inference(superposition,[],[f917,f3095])).
fof(f5310,plain,(
  ( ! [X127,X125,X128,X126,X124] : (p(s(bool,X125)) | s(bool,X127) != s(bool,'const.bool.COND_3'(s(bool,X124),s(bool,X125),s(bool,X126))) | s(bool,'T_0') = s(bool,X128) | s(bool,X127) = s(bool,X128) | s(bool,'F_0') = s(bool,X124)) ) | $spl12),
  inference(superposition,[],[f923,f3095])).
fof(f5321,plain,(
  ( ! [X177,X175,X178,X176] : (p(s(bool,X176)) | s(bool,'F_0') = s(bool,X178) | s(bool,X178) != s(bool,'const.bool.COND_3'(s(bool,X175),s(bool,X176),s(bool,X177))) | s(bool,'F_0') = s(bool,X175)) ) | ($spl10 | $spl12)),
  inference(superposition,[],[f3094,f3095])).
fof(f6058,plain,(
  ( ! [X144,X142,X145,X143,X141] : (p(s(bool,X142)) | s(bool,X144) = s(bool,'const.bool.COND_3'(s(bool,X141),s(bool,X142),s(bool,X143))) | s(bool,X144) = s(bool,X145) | s(bool,'F_0') = s(bool,X145) | s(bool,'F_0') = s(bool,X141)) ) | $spl12),
  inference(superposition,[],[f3148,f3095])).
fof(f6515,plain,(
  ( ! [X39,X41,X38,X42,X40] : (s(bool,X40) = s(bool,X41) | s(bool,'T_0') = s(bool,X38) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,X38),s(bool,X39),s(bool,X40))),s(bool,X42))) | s(bool,'T_0') = s(bool,X41)) )),
  inference(superposition,[],[f207,f326])).
fof(f6543,plain,(
  ( ! [X54,X56,X55] : (s(bool,'T_0') != s(bool,X56) | p(s(bool,'const.bool.COND_3'(s(bool,X54),s(bool,X55),s(bool,X56)))) | s(bool,'T_0') = s(bool,X54)) )),
  inference(superposition,[],[f79,f207])).
fof(f6544,plain,(
  ( ! [X59,X57,X58] : (~p(s(bool,X59)) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X57),s(bool,X58),s(bool,X59))) | s(bool,'T_0') = s(bool,X57)) )),
  inference(superposition,[],[f80,f207])).
fof(f6545,plain,(
  ( ! [X61,X62,X60] : (~p(s(bool,X62)) | s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X60),s(bool,X61),s(bool,X62))) | s(bool,'T_0') = s(bool,X60)) )),
  inference(superposition,[],[f83,f207])).
fof(f6546,plain,(
  ( ! [X64,X65,X63] : (p(s(bool,X65)) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X63),s(bool,X64),s(bool,X65))) | s(bool,'T_0') = s(bool,X63)) )),
  inference(superposition,[],[f84,f207])).
fof(f6547,plain,(
  ( ! [X68,X66,X69,X67] : (p(s(bool,X68)) | p(s(bool,X69)) | s(bool,X69) = s(bool,'const.bool.COND_3'(s(bool,X66),s(bool,X67),s(bool,X68))) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f104,f207])).
fof(f6549,plain,(
  ( ! [X76,X74,X77,X75] : (~p(s(bool,X76)) | ~p(s(bool,X77)) | s(bool,X77) = s(bool,'const.bool.COND_3'(s(bool,X74),s(bool,X75),s(bool,X76))) | s(bool,'T_0') = s(bool,X74)) )),
  inference(superposition,[],[f107,f207])).
fof(f6554,plain,(
  ( ! [X94,X97,X95,X96] : (p(s(bool,X96)) | s(bool,X97) = s(bool,'const.bool.COND_3'(s(bool,X94),s(bool,X95),s(bool,X96))) | s(bool,'T_0') = s(bool,X97) | s(bool,'T_0') = s(bool,X94)) )),
  inference(superposition,[],[f162,f207])).
fof(f6563,plain,(
  ( ! [X134,X136,X135,X137] : (s(bool,'T_0') != s(bool,X136) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,X134),s(bool,X135),s(bool,X136))),s(bool,X137))) | s(bool,'T_0') = s(bool,X134)) )),
  inference(superposition,[],[f376,f207])).
fof(f6564,plain,(
  ( ! [X140,X138,X141,X139] : (~p(s(bool,X140)) | s(bool,X141) = s(bool,'const.bool.COND_3'(s(bool,X138),s(bool,X139),s(bool,X140))) | s(bool,'F_0') = s(bool,X141) | s(bool,'T_0') = s(bool,X138)) )),
  inference(superposition,[],[f453,f207])).
fof(f6567,plain,(
  ( ! [X154,X152,X151,X153,X150] : (~p(s(bool,X152)) | s(bool,X153) = s(bool,'const.bool.COND_3'(s(bool,X150),s(bool,X151),s(bool,X152))) | s(bool,X153) = s(bool,X154) | s(bool,'T_0') = s(bool,X154) | s(bool,'T_0') = s(bool,X150)) )),
  inference(superposition,[],[f917,f207])).
fof(f6568,plain,(
  ( ! [X158,X156,X159,X157,X155] : (p(s(bool,X157)) | s(bool,X158) != s(bool,'const.bool.COND_3'(s(bool,X155),s(bool,X156),s(bool,X157))) | s(bool,'T_0') = s(bool,X159) | s(bool,X158) = s(bool,X159) | s(bool,'T_0') = s(bool,X155)) )),
  inference(superposition,[],[f923,f207])).
fof(f6579,plain,(
  ( ! [X208,X206,X209,X207] : (p(s(bool,X208)) | s(bool,'F_0') = s(bool,X209) | s(bool,X209) != s(bool,'const.bool.COND_3'(s(bool,X206),s(bool,X207),s(bool,X208))) | s(bool,'T_0') = s(bool,X206)) ) | $spl10),
  inference(superposition,[],[f3094,f207])).
fof(f6583,plain,(
  ( ! [X227,X225,X228,X226,X224] : (p(s(bool,X226)) | s(bool,X227) = s(bool,'const.bool.COND_3'(s(bool,X224),s(bool,X225),s(bool,X226))) | s(bool,X227) = s(bool,X228) | s(bool,'F_0') = s(bool,X228) | s(bool,'T_0') = s(bool,X224)) )),
  inference(superposition,[],[f3148,f207])).
fof(f7107,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | p(s(bool,'const.bool.COND_3'(s(bool,X2),s(bool,X1),s(bool,X3)))) | s(bool,'F_0') = s(bool,X2) | s(bool,X4) = s(bool,X5) | s(bool,'F_0') = s(bool,X5) | s(bool,X0) = s(bool,X4)) ) | $spl12),
  inference(superposition,[],[f5288,f1664])).
fof(f7108,plain,(
  ( ! [X6,X10,X8,X7,X9] : (s(bool,X6) != s(bool,X7) | p(s(bool,'const.bool.COND_3'(s(bool,X8),s(bool,X7),s(bool,X9)))) | s(bool,'F_0') = s(bool,X8) | s(bool,X6) = s(bool,X10) | s(bool,'T_0') = s(bool,X10)) ) | $spl12),
  inference(superposition,[],[f5288,f131])).
fof(f7162,plain,(
  ( ! [X158,X156,X154,X157,X155] : (s(bool,'T_0') != s(bool,X156) | p(s(bool,'const.bool.COND_3'(s(bool,X157),s(bool,'const.bool.COND_3'(s(bool,X154),s(bool,X155),s(bool,X156))),s(bool,X158)))) | s(bool,'F_0') = s(bool,X157) | s(bool,'T_0') = s(bool,X154)) ) | $spl12),
  inference(superposition,[],[f5288,f207])).
fof(f7163,plain,(
  ( ! [X163,X161,X159,X162,X160] : (s(bool,'T_0') != s(bool,X160) | p(s(bool,'const.bool.COND_3'(s(bool,X162),s(bool,'const.bool.COND_3'(s(bool,X159),s(bool,X160),s(bool,X161))),s(bool,X163)))) | s(bool,'F_0') = s(bool,X162) | s(bool,'F_0') = s(bool,X159)) ) | $spl12),
  inference(superposition,[],[f5288,f3095])).
fof(f7177,plain,(
  ( ! [X54,X52,X56,X55,X53] : (p(s(bool,'const.bool.COND_3'(s(bool,X53),s(bool,X52),s(bool,X54)))) | s(bool,'F_0') = s(bool,X53) | s(bool,X55) = s(bool,X56) | s(bool,'F_0') = s(bool,X56) | s(bool,X52) = s(bool,X55)) ) | $spl12),
  inference(trivial_inequality_removal,[],[f7137])).
fof(f7137,plain,(
  ( ! [X54,X52,X56,X55,X53] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,'const.bool.COND_3'(s(bool,X53),s(bool,X52),s(bool,X54)))) | s(bool,'F_0') = s(bool,X53) | s(bool,X55) = s(bool,X56) | s(bool,'F_0') = s(bool,X56) | s(bool,X52) = s(bool,X55)) ) | $spl12),
  inference(superposition,[],[f5288,f1664])).
fof(f7180,plain,(
  ( ! [X6,X4,X5,X3] : (p(s(bool,'const.bool.COND_3'(s(bool,X4),s(bool,X3),s(bool,X5)))) | s(bool,'F_0') = s(bool,X4) | s(bool,X3) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) ) | $spl12),
  inference(trivial_inequality_removal,[],[f7126])).
fof(f7126,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,'const.bool.COND_3'(s(bool,X4),s(bool,X3),s(bool,X5)))) | s(bool,'F_0') = s(bool,X4) | s(bool,X3) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) ) | $spl12),
  inference(superposition,[],[f5288,f131])).
fof(f7199,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X2)))) | s(bool,'F_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X1)) ) | $spl12),
  inference(condensation,[],[f7198])).
fof(f7198,plain,(
  ( ! [X74,X72,X71,X73] : (p(s(bool,'const.bool.COND_3'(s(bool,X73),s(bool,X71),s(bool,X74)))) | s(bool,'F_0') = s(bool,X73) | s(bool,'F_0') = s(bool,X71) | s(bool,'F_0') = s(bool,X72)) ) | $spl12),
  inference(subsumption_resolution,[],[f7141,f76])).
fof(f7141,plain,(
  ( ! [X74,X72,X71,X73] : (s(bool,'T_0') != s(bool,X72) | p(s(bool,'const.bool.COND_3'(s(bool,X73),s(bool,X71),s(bool,X74)))) | s(bool,'F_0') = s(bool,X73) | s(bool,'F_0') = s(bool,X71) | s(bool,'F_0') = s(bool,X72)) ) | $spl12),
  inference(superposition,[],[f5288,f3059])).
fof(f7987,plain,(
  ( ! [X158,X156,X154,X157,X155] : (~p(s(bool,X156)) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X157),s(bool,'const.bool.COND_3'(s(bool,X154),s(bool,X155),s(bool,X156))),s(bool,X158))) | s(bool,'F_0') = s(bool,X157) | s(bool,'T_0') = s(bool,X154)) ) | $spl12),
  inference(superposition,[],[f5289,f207])).
fof(f7988,plain,(
  ( ! [X163,X161,X159,X162,X160] : (~p(s(bool,X160)) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X162),s(bool,'const.bool.COND_3'(s(bool,X159),s(bool,X160),s(bool,X161))),s(bool,X163))) | s(bool,'F_0') = s(bool,X162) | s(bool,'F_0') = s(bool,X159)) ) | $spl12),
  inference(superposition,[],[f5289,f3095])).
fof(f8057,plain,(
  ( ! [X158,X156,X154,X157,X155] : (~p(s(bool,X156)) | s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X157),s(bool,'const.bool.COND_3'(s(bool,X154),s(bool,X155),s(bool,X156))),s(bool,X158))) | s(bool,'F_0') = s(bool,X157) | s(bool,'T_0') = s(bool,X154)) ) | $spl12),
  inference(superposition,[],[f5290,f207])).
fof(f8058,plain,(
  ( ! [X163,X161,X159,X162,X160] : (~p(s(bool,X160)) | s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X162),s(bool,'const.bool.COND_3'(s(bool,X159),s(bool,X160),s(bool,X161))),s(bool,X163))) | s(bool,'F_0') = s(bool,X162) | s(bool,'F_0') = s(bool,X159)) ) | $spl12),
  inference(superposition,[],[f5290,f3095])).
fof(f9080,plain,(
  ( ! [X158,X156,X159,X157,X160] : (p(s(bool,X158)) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X159),s(bool,'const.bool.COND_3'(s(bool,X156),s(bool,X157),s(bool,X158))),s(bool,X160))) | s(bool,'F_0') = s(bool,X159) | s(bool,'T_0') = s(bool,X156)) ) | $spl12),
  inference(superposition,[],[f5291,f207])).
fof(f9081,plain,(
  ( ! [X165,X163,X161,X164,X162] : (p(s(bool,X162)) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X164),s(bool,'const.bool.COND_3'(s(bool,X161),s(bool,X162),s(bool,X163))),s(bool,X165))) | s(bool,'F_0') = s(bool,X164) | s(bool,'F_0') = s(bool,X161)) ) | $spl12),
  inference(superposition,[],[f5291,f3095])).
fof(f9149,plain,(
  ( ! [X156,X154,X152,X155,X153] : (p(s(bool,'const.bool.COND_3'(s(bool,X154),s(bool,X155),s(bool,X156)))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X152),s(bool,X153),s(bool,X154))) | s(bool,'F_0') = s(bool,X155) | s(bool,'T_0') = s(bool,X152)) ) | $spl12),
  inference(superposition,[],[f7199,f207])).
fof(f9150,plain,(
  ( ! [X158,X161,X159,X157,X160] : (p(s(bool,'const.bool.COND_3'(s(bool,X158),s(bool,X160),s(bool,X161)))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X157),s(bool,X158),s(bool,X159))) | s(bool,'F_0') = s(bool,X160) | s(bool,'F_0') = s(bool,X157)) ) | $spl12),
  inference(superposition,[],[f7199,f3095])).
fof(f9189,plain,(
  ( ! [X156,X154,X152,X155,X153] : (p(s(bool,'const.bool.COND_3'(s(bool,X155),s(bool,X154),s(bool,X156)))) | s(bool,'F_0') = s(bool,X155) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X152),s(bool,X153),s(bool,X154))) | s(bool,'T_0') = s(bool,X152)) ) | $spl12),
  inference(superposition,[],[f7199,f207])).
fof(f9190,plain,(
  ( ! [X158,X161,X159,X157,X160] : (p(s(bool,'const.bool.COND_3'(s(bool,X160),s(bool,X158),s(bool,X161)))) | s(bool,'F_0') = s(bool,X160) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X157),s(bool,X158),s(bool,X159))) | s(bool,'F_0') = s(bool,X157)) ) | $spl12),
  inference(superposition,[],[f7199,f3095])).
fof(f9314,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,'const.bool.COND_3'(s(bool,X1),s(bool,X0),s(bool,X2)))) | s(bool,X1) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) ) | $spl12),
  inference(condensation,[],[f9313])).
fof(f9313,plain,(
  ( ! [X45,X43,X46,X44,X42] : (p(s(bool,'const.bool.COND_3'(s(bool,X43),s(bool,X44),s(bool,X45)))) | s(bool,'F_0') = s(bool,X42) | s(bool,'F_0') = s(bool,X44) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) ) | $spl12),
  inference(duplicate_literal_removal,[],[f9121])).
fof(f9121,plain,(
  ( ! [X45,X43,X46,X44,X42] : (p(s(bool,'const.bool.COND_3'(s(bool,X43),s(bool,X44),s(bool,X45)))) | s(bool,'F_0') = s(bool,X42) | s(bool,'F_0') = s(bool,X44) | s(bool,'F_0') = s(bool,X42) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) ) | $spl12),
  inference(superposition,[],[f7199,f1664])).
fof(f9332,plain,(
  ( ! [X90,X88,X87,X91,X89] : (p(s(bool,'const.bool.COND_3'(s(bool,X87),s(bool,X88),s(bool,X89)))) | s(bool,'F_0') = s(bool,X88) | s(bool,X90) = s(bool,X91) | s(bool,'F_0') = s(bool,X91) | s(bool,X87) = s(bool,X90)) ) | $spl12),
  inference(subsumption_resolution,[],[f9131,f110])).
fof(f9131,plain,(
  ( ! [X90,X88,X87,X91,X89] : (p(s(bool,'const.bool.COND_3'(s(bool,X87),s(bool,X88),s(bool,X89)))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X88) | s(bool,X90) = s(bool,X91) | s(bool,'F_0') = s(bool,X91) | s(bool,X87) = s(bool,X90)) ) | $spl12),
  inference(superposition,[],[f7199,f1664])).
fof(f10017,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | p(s(bool,'const.bool.COND_3'(s(bool,X2),s(bool,X3),s(bool,X1)))) | s(bool,'T_0') = s(bool,X2) | s(bool,X4) = s(bool,X5) | s(bool,'F_0') = s(bool,X5) | s(bool,X0) = s(bool,X4)) )),
  inference(superposition,[],[f6543,f1664])).
fof(f10018,plain,(
  ( ! [X6,X10,X8,X7,X9] : (s(bool,X6) != s(bool,X7) | p(s(bool,'const.bool.COND_3'(s(bool,X8),s(bool,X9),s(bool,X7)))) | s(bool,'T_0') = s(bool,X8) | s(bool,X6) = s(bool,X10) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f6543,f131])).
fof(f10073,plain,(
  ( ! [X158,X156,X159,X157,X160] : (s(bool,'T_0') != s(bool,X158) | p(s(bool,'const.bool.COND_3'(s(bool,X159),s(bool,X160),s(bool,'const.bool.COND_3'(s(bool,X156),s(bool,X157),s(bool,X158)))))) | s(bool,'T_0') = s(bool,X159) | s(bool,'T_0') = s(bool,X156)) )),
  inference(superposition,[],[f6543,f207])).
fof(f10074,plain,(
  ( ! [X165,X163,X161,X164,X162] : (s(bool,'T_0') != s(bool,X162) | p(s(bool,'const.bool.COND_3'(s(bool,X164),s(bool,X165),s(bool,'const.bool.COND_3'(s(bool,X161),s(bool,X162),s(bool,X163)))))) | s(bool,'T_0') = s(bool,X164) | s(bool,'F_0') = s(bool,X161)) ) | $spl12),
  inference(superposition,[],[f6543,f3095])).
fof(f10090,plain,(
  ( ! [X54,X52,X56,X55,X53] : (p(s(bool,'const.bool.COND_3'(s(bool,X53),s(bool,X54),s(bool,X52)))) | s(bool,'T_0') = s(bool,X53) | s(bool,X55) = s(bool,X56) | s(bool,'F_0') = s(bool,X56) | s(bool,X52) = s(bool,X55)) )),
  inference(trivial_inequality_removal,[],[f10047])).
fof(f10047,plain,(
  ( ! [X54,X52,X56,X55,X53] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,'const.bool.COND_3'(s(bool,X53),s(bool,X54),s(bool,X52)))) | s(bool,'T_0') = s(bool,X53) | s(bool,X55) = s(bool,X56) | s(bool,'F_0') = s(bool,X56) | s(bool,X52) = s(bool,X55)) )),
  inference(superposition,[],[f6543,f1664])).
fof(f10093,plain,(
  ( ! [X6,X4,X5,X3] : (p(s(bool,'const.bool.COND_3'(s(bool,X4),s(bool,X5),s(bool,X3)))) | s(bool,'T_0') = s(bool,X4) | s(bool,X3) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(trivial_inequality_removal,[],[f10036])).
fof(f10036,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,'const.bool.COND_3'(s(bool,X4),s(bool,X5),s(bool,X3)))) | s(bool,'T_0') = s(bool,X4) | s(bool,X3) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f6543,f131])).
fof(f10095,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) != s(bool,X2) | p(s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,X3),s(bool,X2)))) | s(bool,'F_0') = s(bool,X1)) )),
  inference(condensation,[],[f10094])).
fof(f10094,plain,(
  ( ! [X57,X54,X56,X55,X53] : (s(bool,X53) != s(bool,X54) | p(s(bool,'const.bool.COND_3'(s(bool,X55),s(bool,X56),s(bool,X54)))) | s(bool,'T_0') = s(bool,X55) | s(bool,'F_0') = s(bool,X53) | s(bool,'T_0') = s(bool,X57)) )),
  inference(duplicate_literal_removal,[],[f10028])).
fof(f10028,plain,(
  ( ! [X57,X54,X56,X55,X53] : (s(bool,X53) != s(bool,X54) | p(s(bool,'const.bool.COND_3'(s(bool,X55),s(bool,X56),s(bool,X54)))) | s(bool,'T_0') = s(bool,X55) | s(bool,'F_0') = s(bool,X53) | s(bool,'T_0') = s(bool,X57) | s(bool,'T_0') = s(bool,X57)) )),
  inference(superposition,[],[f6543,f1664])).
fof(f10120,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X2)))) | s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X2)) )),
  inference(duplicate_literal_removal,[],[f10119])).
fof(f10119,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X2)))) | s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) )),
  inference(condensation,[],[f10118])).
fof(f10118,plain,(
  ( ! [X45,X43,X46,X44,X42] : (p(s(bool,'const.bool.COND_3'(s(bool,X44),s(bool,X45),s(bool,X42)))) | s(bool,'T_0') = s(bool,X44) | s(bool,'F_0') = s(bool,X42) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f10045,f131])).
fof(f10045,plain,(
  ( ! [X45,X43,X46,X44,X42] : (s(bool,'T_0') != s(bool,X43) | p(s(bool,'const.bool.COND_3'(s(bool,X44),s(bool,X45),s(bool,X42)))) | s(bool,'T_0') = s(bool,X44) | s(bool,'F_0') = s(bool,X42) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f6543,f1664])).
fof(f10673,plain,(
  ( ! [X163,X161,X164,X162,X160] : (~p(s(bool,X162)) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X163),s(bool,X164),s(bool,'const.bool.COND_3'(s(bool,X160),s(bool,X161),s(bool,X162))))) | s(bool,'T_0') = s(bool,X163) | s(bool,'T_0') = s(bool,X160)) )),
  inference(superposition,[],[f6544,f207])).
fof(f10674,plain,(
  ( ! [X167,X165,X169,X166,X168] : (~p(s(bool,X166)) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X168),s(bool,X169),s(bool,'const.bool.COND_3'(s(bool,X165),s(bool,X166),s(bool,X167))))) | s(bool,'T_0') = s(bool,X168) | s(bool,'F_0') = s(bool,X165)) ) | $spl12),
  inference(superposition,[],[f6544,f3095])).
fof(f10927,plain,(
  ( ! [X158,X156,X154,X157,X155] : (p(s(bool,'const.bool.COND_3'(s(bool,X156),s(bool,X157),s(bool,X158)))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X154),s(bool,X155),s(bool,X156))) | s(bool,'F_0') = s(bool,X158) | s(bool,'T_0') = s(bool,X154)) )),
  inference(superposition,[],[f10120,f207])).
fof(f10928,plain,(
  ( ! [X163,X161,X159,X162,X160] : (p(s(bool,'const.bool.COND_3'(s(bool,X160),s(bool,X162),s(bool,X163)))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X159),s(bool,X160),s(bool,X161))) | s(bool,'F_0') = s(bool,X163) | s(bool,'F_0') = s(bool,X159)) ) | $spl12),
  inference(superposition,[],[f10120,f3095])).
fof(f11009,plain,(
  ( ! [X158,X156,X159,X157,X160] : (p(s(bool,'const.bool.COND_3'(s(bool,X159),s(bool,X160),s(bool,X158)))) | s(bool,'T_0') = s(bool,X159) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X156),s(bool,X157),s(bool,X158))) | s(bool,'T_0') = s(bool,X156)) )),
  inference(superposition,[],[f10120,f207])).
fof(f11010,plain,(
  ( ! [X165,X163,X161,X164,X162] : (p(s(bool,'const.bool.COND_3'(s(bool,X164),s(bool,X165),s(bool,X162)))) | s(bool,'T_0') = s(bool,X164) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X161),s(bool,X162),s(bool,X163))) | s(bool,'F_0') = s(bool,X161)) ) | $spl12),
  inference(superposition,[],[f10120,f3095])).
fof(f11118,plain,(
  ( ! [X101,X99,X102,X100] : (p(s(bool,'const.bool.COND_3'(s(bool,X99),s(bool,X100),s(bool,X101)))) | s(bool,'F_0') = s(bool,X101) | s(bool,'F_0') = s(bool,X102) | s(bool,X99) = s(bool,X102)) )),
  inference(subsumption_resolution,[],[f10912,f110])).
fof(f10912,plain,(
  ( ! [X101,X99,X102,X100] : (p(s(bool,'const.bool.COND_3'(s(bool,X99),s(bool,X100),s(bool,X101)))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X101) | s(bool,'F_0') = s(bool,X102) | s(bool,X99) = s(bool,X102)) )),
  inference(superposition,[],[f10120,f3059])).
fof(f11119,plain,(
  ( ! [X111,X109,X107,X110,X108] : (p(s(bool,'const.bool.COND_3'(s(bool,X107),s(bool,X108),s(bool,X109)))) | s(bool,'F_0') = s(bool,X109) | s(bool,X107) = s(bool,X110) | s(bool,X110) = s(bool,X111) | s(bool,'T_0') = s(bool,X111)) )),
  inference(subsumption_resolution,[],[f10914,f110])).
fof(f10914,plain,(
  ( ! [X111,X109,X107,X110,X108] : (p(s(bool,'const.bool.COND_3'(s(bool,X107),s(bool,X108),s(bool,X109)))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X109) | s(bool,X107) = s(bool,X110) | s(bool,X110) = s(bool,X111) | s(bool,'T_0') = s(bool,X111)) )),
  inference(superposition,[],[f10120,f1664])).
fof(f11187,plain,(
  ( ! [X163,X161,X164,X162,X160] : (~p(s(bool,X162)) | s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X163),s(bool,X164),s(bool,'const.bool.COND_3'(s(bool,X160),s(bool,X161),s(bool,X162))))) | s(bool,'T_0') = s(bool,X163) | s(bool,'T_0') = s(bool,X160)) )),
  inference(superposition,[],[f6545,f207])).
fof(f11188,plain,(
  ( ! [X167,X165,X169,X166,X168] : (~p(s(bool,X166)) | s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X168),s(bool,X169),s(bool,'const.bool.COND_3'(s(bool,X165),s(bool,X166),s(bool,X167))))) | s(bool,'T_0') = s(bool,X168) | s(bool,'F_0') = s(bool,X165)) ) | $spl12),
  inference(superposition,[],[f6545,f3095])).
fof(f11552,plain,(
  ( ! [X163,X161,X164,X162,X160] : (p(s(bool,X162)) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X163),s(bool,X164),s(bool,'const.bool.COND_3'(s(bool,X160),s(bool,X161),s(bool,X162))))) | s(bool,'T_0') = s(bool,X163) | s(bool,'T_0') = s(bool,X160)) )),
  inference(superposition,[],[f6546,f207])).
fof(f11553,plain,(
  ( ! [X167,X165,X169,X166,X168] : (p(s(bool,X166)) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X168),s(bool,X169),s(bool,'const.bool.COND_3'(s(bool,X165),s(bool,X166),s(bool,X167))))) | s(bool,'T_0') = s(bool,X168) | s(bool,'F_0') = s(bool,X165)) ) | $spl12),
  inference(superposition,[],[f6546,f3095])).
fof(f12274,plain,(
  ( ! [X191,X189,X192,X190,X188] : (s(bool,X190) = s(bool,X191) | s(bool,'T_0') = s(bool,X191) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X192),s(bool,'const.bool.COND_3'(s(bool,X188),s(bool,X189),s(bool,X190))))) | s(bool,'T_0') = s(bool,X188)) )),
  inference(superposition,[],[f1050,f207])).
fof(f12275,plain,(
  ( ! [X196,X194,X197,X195,X193] : (s(bool,X194) = s(bool,X196) | s(bool,'T_0') = s(bool,X196) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X197),s(bool,'const.bool.COND_3'(s(bool,X193),s(bool,X194),s(bool,X195))))) | s(bool,'F_0') = s(bool,X193)) ) | $spl12),
  inference(superposition,[],[f1050,f3095])).
fof(f11962,plain,(
  ( ! [X127,X128,X129] : (p(s(bool,'const.bool.COND_3'(s(bool,X128),s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X127))),s(bool,X129)))) | s(bool,'F_0') = s(bool,X128) | s('type.realax.real','const.real.abs_1'(s('type.realax.real',X127))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X127)))) ) | $spl12),
  inference(trivial_inequality_removal,[],[f11922])).
fof(f11922,plain,(
  ( ! [X127,X128,X129] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,'const.bool.COND_3'(s(bool,X128),s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X127))),s(bool,X129)))) | s(bool,'F_0') = s(bool,X128) | s('type.realax.real','const.real.abs_1'(s('type.realax.real',X127))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X127)))) ) | $spl12),
  inference(superposition,[],[f5288,f6529])).
fof(f10086,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,'T_0'))))) )),
  inference(duplicate_literal_removal,[],[f10085])).
fof(f10085,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(condensation,[],[f10084])).
fof(f10084,plain,(
  ( ! [X94,X92,X95,X93] : (p(s(bool,'const.bool.COND_3'(s(bool,X93),s(bool,X94),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,X93) | s(bool,X92) = s(bool,X95) | s(bool,'T_0') = s(bool,X95)) )),
  inference(trivial_inequality_removal,[],[f10056])).
fof(f10056,plain,(
  ( ! [X94,X92,X95,X93] : (s(bool,X92) != s(bool,X92) | p(s(bool,'const.bool.COND_3'(s(bool,X93),s(bool,X94),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,X93) | s(bool,X92) = s(bool,X95) | s(bool,'T_0') = s(bool,X95)) )),
  inference(superposition,[],[f6543,f131])).
fof(f11961,plain,(
  ( ! [X140,X141,X139] : (p(s(bool,'const.bool.COND_3'(s(bool,X140),s(bool,X141),s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X139)))))) | s(bool,'T_0') = s(bool,X140) | s('type.realax.real','const.real.abs_1'(s('type.realax.real',X139))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X139)))) )),
  inference(trivial_inequality_removal,[],[f11926])).
fof(f11926,plain,(
  ( ! [X140,X141,X139] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,'const.bool.COND_3'(s(bool,X140),s(bool,X141),s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X139)))))) | s(bool,'T_0') = s(bool,X140) | s('type.realax.real','const.real.abs_1'(s('type.realax.real',X139))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X139)))) )),
  inference(superposition,[],[f6543,f6529])).
fof(f468,plain,(
  ( ! [X4,X3] : (s('type.realax.real','const.real.abs_1'(s('type.realax.real',X3))) = s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real',X3),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X3))))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X3))),s(bool,X4)))) )),
  inference(forward_demodulation,[],[f457,f118])).
fof(f457,plain,(
  ( ! [X4,X3] : (s('type.realax.real','const.real.abs_1'(s('type.realax.real',X3))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))),s('type.realax.real',X3),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X3))))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X3))),s(bool,X4)))) )),
  inference(superposition,[],[f74,f263])).
fof(f469,plain,(
  ( ! [X5] : (s('type.realax.real','const.real.abs_1'(s('type.realax.real',X5))) = s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real',X5),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X5))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X5)))) )),
  inference(forward_demodulation,[],[f458,f118])).
fof(f458,plain,(
  ( ! [X5] : (s('type.realax.real','const.real.abs_1'(s('type.realax.real',X5))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))),s('type.realax.real',X5),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X5))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X5)))) )),
  inference(superposition,[],[f74,f249])).
fof(f2354,plain,(
  ( ! [X144,X145,X143] : (s('type.realax.real','const.real.abs_1'(s('type.realax.real',X143))) = s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real',X143),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X143))))) | s(bool,X144) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X143))) | s(bool,X144) = s(bool,X145) | s(bool,'T_0') = s(bool,X145)) )),
  inference(superposition,[],[f74,f1664])).
fof(f5146,plain,(
  ( ! [X37,X36] : (s('type.realax.real','const.real.abs_1'(s('type.realax.real',X36))) = s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real',X36),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X36))))) | s(bool,X37) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X36))) | s(bool,'F_0') = s(bool,X37)) )),
  inference(superposition,[],[f74,f3059])).
fof(f6770,plain,(
  ( ! [X3] : (s('type.realax.real','const.real.abs_1'(s('type.realax.real',X3))) = s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real',X3),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X3))))) | s('type.realax.real',X3) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X3)))) ) | $spl12),
  inference(superposition,[],[f74,f5279])).
fof(f642,plain,(
  ( ! [X101,X102,X100] : (s('type.realax.real','const.real.abs_1'(s('type.realax.real',X100))) = s('type.realax.real','const.bool.COND_3'(s(bool,X101),s('type.realax.real',X100),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X100))))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X101),s(bool,X102))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X100)))) )),
  inference(superposition,[],[f74,f326])).
fof(f681,plain,(
  ( ! [X101,X102,X100] : (s('type.realax.real','const.real.abs_1'(s('type.realax.real',X100))) = s('type.realax.real','const.bool.COND_3'(s(bool,X101),s('type.realax.real',X100),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X100))))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X100))),s(bool,X102))) | s(bool,'T_0') = s(bool,X101)) )),
  inference(superposition,[],[f74,f326])).
fof(f1157,plain,(
  ( ! [X109,X108] : (s('type.realax.real','const.real.abs_1'(s('type.realax.real',X108))) = s('type.realax.real','const.bool.COND_3'(s(bool,X109),s('type.realax.real',X108),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X108))))) | s(bool,'T_0') = s(bool,X109) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X108)))) )),
  inference(superposition,[],[f74,f131])).
fof(f2207,plain,(
  ( ! [X194,X195,X193] : (s('type.realax.real','const.real.abs_1'(s('type.realax.real',X193))) = s('type.realax.real','const.bool.COND_3'(s(bool,X194),s('type.realax.real',X193),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X193))))) | s(bool,'F_0') = s(bool,X194) | s(bool,X195) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X193))) | s(bool,'T_0') = s(bool,X195)) )),
  inference(superposition,[],[f74,f1664])).
fof(f2270,plain,(
  ( ! [X191,X189,X190] : (s('type.realax.real','const.real.abs_1'(s('type.realax.real',X189))) = s('type.realax.real','const.bool.COND_3'(s(bool,X190),s('type.realax.real',X189),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X189))))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X189))) | s(bool,X190) = s(bool,X191) | s(bool,'T_0') = s(bool,X191)) )),
  inference(superposition,[],[f74,f1664])).
fof(f2417,plain,(
  ( ! [X194,X195,X193] : (s('type.realax.real','const.real.abs_1'(s('type.realax.real',X193))) = s('type.realax.real','const.bool.COND_3'(s(bool,X194),s('type.realax.real',X193),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X193))))) | s(bool,X195) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X193))) | s(bool,'F_0') = s(bool,X195) | s(bool,'T_0') = s(bool,X194)) )),
  inference(superposition,[],[f74,f1664])).
fof(f2480,plain,(
  ( ! [X194,X195,X193] : (s('type.realax.real','const.real.abs_1'(s('type.realax.real',X193))) = s('type.realax.real','const.bool.COND_3'(s(bool,X194),s('type.realax.real',X193),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X193))))) | s(bool,X194) = s(bool,X195) | s(bool,'F_0') = s(bool,X195) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X193)))) )),
  inference(superposition,[],[f74,f1664])).
fof(f5147,plain,(
  ( ! [X39,X38] : (s('type.realax.real','const.real.abs_1'(s('type.realax.real',X38))) = s('type.realax.real','const.bool.COND_3'(s(bool,X39),s('type.realax.real',X38),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X38))))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X38))) | s(bool,'F_0') = s(bool,X39)) )),
  inference(superposition,[],[f74,f3059])).
fof(f12642,plain,(
  ( ! [X549,X548,X550] : (s('type.realax.real','const.real.abs_1'(s('type.realax.real',X548))) = s('type.realax.real','const.bool.COND_3'(s(bool,X549),s('type.realax.real',X548),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X548))))) | s(bool,'T_0') = s(bool,X549) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X550),s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X548)))))) )),
  inference(superposition,[],[f74,f1050])).
fof(f12802,plain,(
  ( ! [X528,X530,X529] : (s('type.realax.real','const.real.abs_1'(s('type.realax.real',X528))) = s('type.realax.real','const.bool.COND_3'(s(bool,X529),s('type.realax.real',X528),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X528))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X528))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X530),s(bool,X529)))) )),
  inference(superposition,[],[f74,f1050])).
fof(f74,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))),s('type.realax.real',X0),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))) )),
  inference(cnf_transformation,[],[f26])).
fof(f26,plain,(
  ! [X0] : s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))),s('type.realax.real',X0),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X6] : s('type.realax.real','const.real.abs_1'(s('type.realax.real',X6))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X6))),s('type.realax.real',X6),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X6)))))),
  file('thm.real.ABS_POS','thm.real.abs')).
fof(f5327,plain,(
  ( ! [X196,X197,X195] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X196)))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.bool.COND_3'(s(bool,X195),s('type.realax.real',X196),s('type.realax.real',X197)))))) | s(bool,'F_0') = s(bool,X195)) ) | $spl12),
  inference(superposition,[],[f77,f3095])).
fof(f6589,plain,(
  ( ! [X246,X247,X245] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X247)))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.bool.COND_3'(s(bool,X245),s('type.realax.real',X246),s('type.realax.real',X247)))))) | s(bool,'T_0') = s(bool,X245)) )),
  inference(superposition,[],[f77,f207])).
fof(f6747,plain,(
  ( ! [X6,X7,X5] : (s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X7))) | s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real',X6),s('type.realax.real',X7))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real',X6),s('type.realax.real',X7))))) | s(bool,'T_0') = s(bool,X5)) ) | $spl12),
  inference(superposition,[],[f5279,f207])).
fof(f6748,plain,(
  ( ! [X10,X8,X9] : (s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X9))) | s('type.realax.real','const.bool.COND_3'(s(bool,X8),s('type.realax.real',X9),s('type.realax.real',X10))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.bool.COND_3'(s(bool,X8),s('type.realax.real',X9),s('type.realax.real',X10))))) | s(bool,'F_0') = s(bool,X8)) ) | $spl12),
  inference(superposition,[],[f5279,f3095])).
fof(f8096,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X7)))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real',X6),s('type.realax.real',X7))))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f258,f207])).
fof(f8097,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X9)))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.bool.COND_3'(s(bool,X8),s('type.realax.real',X9),s('type.realax.real',X10))))))) | s(bool,'F_0') = s(bool,X8)) ) | $spl12),
  inference(superposition,[],[f258,f3095])).
fof(f11841,plain,(
  ( ! [X6,X7,X5] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X7))) | s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real',X6),s('type.realax.real',X7))))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.bool.COND_3'(s(bool,X5),s('type.realax.real',X6),s('type.realax.real',X7))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f6529,f207])).
fof(f11842,plain,(
  ( ! [X10,X8,X9] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X9))) | s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.bool.COND_3'(s(bool,X8),s('type.realax.real',X9),s('type.realax.real',X10))))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.bool.COND_3'(s(bool,X8),s('type.realax.real',X9),s('type.realax.real',X10))))) | s(bool,'F_0') = s(bool,X8)) ) | $spl12),
  inference(superposition,[],[f6529,f3095])).
fof(f12033,plain,(
  ( ! [X6,X8,X7] : (s('type.realax.real','const.real.abs_1'(s('type.realax.real',X8))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X8))) | s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real',X7),s('type.realax.real',X8))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.bool.COND_3'(s(bool,X6),s('type.realax.real',X7),s('type.realax.real',X8))))) | s(bool,'T_0') = s(bool,X6)) ) | $spl12),
  inference(superposition,[],[f11987,f207])).
fof(f12034,plain,(
  ( ! [X10,X11,X9] : (s('type.realax.real','const.real.abs_1'(s('type.realax.real',X10))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X10))) | s('type.realax.real','const.bool.COND_3'(s(bool,X9),s('type.realax.real',X10),s('type.realax.real',X11))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.bool.COND_3'(s(bool,X9),s('type.realax.real',X10),s('type.realax.real',X11))))) | s(bool,'F_0') = s(bool,X9)) ) | $spl12),
  inference(superposition,[],[f11987,f3095])).
fof(f12035,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))),s('type.realax.real',X0),s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0))))) | s('type.realax.real',X0) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0)))) ) | $spl12),
  inference(superposition,[],[f74,f11987])).
fof(f73,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f25])).
fof(f25,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.real.ABS_POS',aHL_FALSITY)).
fof(f83,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f34])).
fof(f34,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f33])).
fof(f33,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X9] : (s(bool,'F_0') = s(bool,X9) <=> ~p(s(bool,X9)))),
  file('thm.real.ABS_POS','thm.bool.EQ_CLAUSES|split|3')).
fof(f84,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f57])).
fof(f103,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f45])).
fof(f45,plain,(
  ! [X0,X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f44])).
fof(f44,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X10] : ! [X7,X8] : s(X10,X8) = s(X10,'const.bool.COND_3'(s(bool,'F_0'),s(X10,X7),s(X10,X8)))),
  file('thm.real.ABS_POS','thm.bool.COND_CLAUSES|split|1')).
fof(f118,plain,(
  s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0)))))),
  inference(unit_resulting_resolution,[],[f70,f84])).
fof(f187,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f102,f76])).
fof(f219,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,X1))) | s(bool,'F_0') = s(bool,X1)) )),
  inference(superposition,[],[f103,f76])).
fof(f264,plain,(
  ( ! [X6,X7] : (s(bool,X7) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X6),s(bool,X7)))) )),
  inference(superposition,[],[f249,f103])).
fof(f417,plain,(
  ( ! [X21,X22,X20] : (s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X20),s(bool,X21))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X21),s(bool,X22)))) )),
  inference(forward_demodulation,[],[f340,f118])).
fof(f340,plain,(
  ( ! [X21,X22,X20] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X21),s(bool,X22))) | s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X20),s(bool,X21))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0)))))) )),
  inference(superposition,[],[f263,f103])).
fof(f453,plain,(
  ( ! [X18,X16] : (~p(s(bool,X18)) | s(bool,X16) = s(bool,X18) | s(bool,'F_0') = s(bool,X16)) )),
  inference(forward_demodulation,[],[f452,f118])).
fof(f452,plain,(
  ( ! [X18,X16] : (s(bool,X16) = s(bool,X18) | ~p(s(bool,X18)) | s(bool,X16) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0)))))) )),
  inference(forward_demodulation,[],[f451,f102])).
fof(f451,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X18)) | s(bool,X18) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X16),s(bool,X17))) | s(bool,X16) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0)))))) )),
  inference(subsumption_resolution,[],[f397,f71])).
fof(f397,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X18)) | s(bool,X18) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X16),s(bool,X17))) | s(bool,X16) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0)))))) )),
  inference(superposition,[],[f107,f263])).
fof(f493,plain,(
  ( ! [X17,X15,X16] : (s(bool,X17) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X15),s(bool,X16))) | s(bool,'F_0') = s(bool,X17) | s(bool,X15) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0)))))) )),
  inference(subsumption_resolution,[],[f483,f71])).
fof(f483,plain,(
  ( ! [X17,X15,X16] : (~p(s(bool,'T_0')) | s(bool,X17) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X15),s(bool,X16))) | s(bool,'F_0') = s(bool,X17) | s(bool,X15) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0)))))) )),
  inference(superposition,[],[f453,f263])).
fof(f522,plain,(
  ( ! [X2,X0,X3,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,sK1(X0,X1,X2,bool)),s(bool,X3)))) )),
  inference(forward_demodulation,[],[f516,f118])).
fof(f516,plain,(
  ( ! [X2,X0,X3,X1] : (s(X2,f(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))))) != s(X2,g(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,sK1(X0,X1,X2,bool)),s(bool,X3)))) )),
  inference(superposition,[],[f108,f263])).
fof(f523,plain,(
  ( ! [X6,X4,X5] : (s(X6,f(s(bool,'F_0'))) != s(X6,g(s(bool,'F_0'))) | s(fun(bool,X6),X4) = s(fun(bool,X6),X5) | s(bool,'T_0') = s(bool,sK1(X4,X5,X6,bool))) )),
  inference(forward_demodulation,[],[f517,f118])).
fof(f517,plain,(
  ( ! [X6,X4,X5] : (s(X6,f(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))))) != s(X6,g(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))))) | s(fun(bool,X6),X4) = s(fun(bool,X6),X5) | s(bool,'T_0') = s(bool,sK1(X4,X5,X6,bool))) )),
  inference(superposition,[],[f108,f249])).
fof(f524,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK1(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,g(s(X0,sK1(X1,X2,bool,X0)))),s(bool,X3)))) )),
  inference(forward_demodulation,[],[f519,f118])).
fof(f519,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,f(s(X0,sK1(X1,X2,bool,X0)))) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,g(s(X0,sK1(X1,X2,bool,X0)))),s(bool,X3)))) )),
  inference(superposition,[],[f108,f263])).
fof(f525,plain,(
  ( ! [X6,X4,X5] : (s(bool,'F_0') != s(bool,f(s(X4,sK1(X5,X6,bool,X4)))) | s(fun(X4,bool),X5) = s(fun(X4,bool),X6) | s(bool,'T_0') = s(bool,g(s(X4,sK1(X5,X6,bool,X4))))) )),
  inference(forward_demodulation,[],[f520,f118])).
fof(f520,plain,(
  ( ! [X6,X4,X5] : (s(bool,f(s(X4,sK1(X5,X6,bool,X4)))) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))) | s(fun(X4,bool),X5) = s(fun(X4,bool),X6) | s(bool,'T_0') = s(bool,g(s(X4,sK1(X5,X6,bool,X4))))) )),
  inference(superposition,[],[f108,f249])).
fof(f542,plain,(
  ( ! [X54,X56,X55,X53] : (s(bool,X54) = s(bool,X55) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X55),s(bool,X56))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X53),s(bool,X54)))) )),
  inference(superposition,[],[f326,f103])).
fof(f576,plain,(
  ( ! [X57,X58,X56,X55] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X56),s(bool,X57))) | s(bool,X58) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X55),s(bool,X56))) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f326,f103])).
fof(f612,plain,(
  ( ! [X45,X43,X44,X42] : (s(bool,'T_0') = s(bool,X43) | s(bool,X44) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X42),s(bool,X43))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X44),s(bool,X45)))) )),
  inference(superposition,[],[f326,f103])).
fof(f921,plain,(
  ( ! [X54,X52,X55,X53] : (s(bool,X54) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X52),s(bool,X53))) | s(bool,'F_0') = s(bool,X54) | s(bool,X52) = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) )),
  inference(subsumption_resolution,[],[f705,f71])).
fof(f705,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,'T_0')) | s(bool,X54) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X52),s(bool,X53))) | s(bool,'F_0') = s(bool,X54) | s(bool,X52) = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) )),
  inference(superposition,[],[f453,f326])).
fof(f934,plain,(
  ( ! [X10,X11,X9] : (s(bool,'F_0') != s(bool,X9) | s(bool,X9) = s(bool,X10) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X10),s(bool,X11)))) )),
  inference(subsumption_resolution,[],[f735,f71])).
fof(f735,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X9) | s(bool,X9) = s(bool,X10) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X10),s(bool,X11)))) )),
  inference(superposition,[],[f83,f326])).
fof(f937,plain,(
  ( ! [X52,X50,X53,X51] : (s(bool,X50) = s(bool,X51) | s(bool,'F_0') = s(bool,X51) | s(bool,X50) = s(bool,X52) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X52),s(bool,X53)))) )),
  inference(subsumption_resolution,[],[f746,f71])).
fof(f746,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,'T_0')) | s(bool,X50) = s(bool,X51) | s(bool,'F_0') = s(bool,X51) | s(bool,X50) = s(bool,X52) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X52),s(bool,X53)))) )),
  inference(superposition,[],[f453,f326])).
fof(f957,plain,(
  ( ! [X39,X41,X40] : (s(bool,X40) != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X41))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X39),s(bool,X40)))) )),
  inference(superposition,[],[f792,f103])).
fof(f1050,plain,(
  ( ! [X57,X56,X55] : (s(bool,X56) = s(bool,X57) | s(bool,'T_0') = s(bool,X57) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X55),s(bool,X56)))) )),
  inference(superposition,[],[f131,f103])).
fof(f1094,plain,(
  ( ! [X43,X44,X42] : (s(bool,'T_0') = s(bool,X43) | s(bool,X44) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X42),s(bool,X43))) | s(bool,'T_0') = s(bool,X44)) )),
  inference(superposition,[],[f131,f103])).
fof(f1799,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X2))) | s(bool,'F_0') = s(bool,X1) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f187,f131])).
fof(f1801,plain,(
  ( ! [X12,X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.bool.COND_3'(s(bool,X8),s(bool,X9),s(bool,X10))) | s(bool,'F_0') = s(bool,X9) | s(bool,X8) = s(bool,X11) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X11),s(bool,X12)))) )),
  inference(superposition,[],[f187,f326])).
fof(f1951,plain,(
  ( ! [X70,X72,X71] : (s(bool,X72) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X70),s(bool,X71))) | s(bool,'F_0') = s(bool,X72) | s(bool,'F_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f1882,f71])).
fof(f1882,plain,(
  ( ! [X70,X72,X71] : (~p(s(bool,'T_0')) | s(bool,X72) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X70),s(bool,X71))) | s(bool,'F_0') = s(bool,X72) | s(bool,'F_0') = s(bool,X70)) )),
  inference(superposition,[],[f453,f187])).
fof(f1952,plain,(
  ( ! [X80,X81,X79,X82] : (s(bool,X81) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X79),s(bool,X80))) | s(bool,X81) = s(bool,X82) | s(bool,'T_0') = s(bool,X82) | s(bool,'F_0') = s(bool,X79)) )),
  inference(subsumption_resolution,[],[f1885,f71])).
fof(f1885,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,'T_0')) | s(bool,X81) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X79),s(bool,X80))) | s(bool,X81) = s(bool,X82) | s(bool,'T_0') = s(bool,X82) | s(bool,'F_0') = s(bool,X79)) )),
  inference(superposition,[],[f917,f187])).
fof(f1993,plain,(
  ( ! [X121,X122,X120,X119] : (s(bool,X119) = s(bool,X121) | s(bool,'F_0') = s(bool,X121) | s(bool,X122) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X119),s(bool,X120))) | s(bool,'T_0') = s(bool,X122)) )),
  inference(superposition,[],[f1664,f102])).
fof(f1994,plain,(
  ( ! [X125,X123,X126,X124] : (s(bool,X124) = s(bool,X125) | s(bool,'F_0') = s(bool,X125) | s(bool,X126) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X123),s(bool,X124))) | s(bool,'T_0') = s(bool,X126)) )),
  inference(superposition,[],[f1664,f103])).
fof(f2024,plain,(
  ( ! [X121,X122,X120,X119] : (s(bool,X120) = s(bool,X121) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X119),s(bool,X120))) | s(bool,X121) = s(bool,X122) | s(bool,'T_0') = s(bool,X122)) )),
  inference(superposition,[],[f1664,f103])).
fof(f2061,plain,(
  ( ! [X97,X95,X98,X96] : (s(bool,'F_0') = s(bool,X96) | s(bool,X97) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X95),s(bool,X96))) | s(bool,X97) = s(bool,X98) | s(bool,'T_0') = s(bool,X98)) )),
  inference(superposition,[],[f1664,f103])).
fof(f2091,plain,(
  ( ! [X125,X123,X126,X124] : (s(bool,X124) = s(bool,X125) | s(bool,X126) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X123),s(bool,X124))) | s(bool,'F_0') = s(bool,X126) | s(bool,'T_0') = s(bool,X125)) )),
  inference(superposition,[],[f1664,f103])).
fof(f2121,plain,(
  ( ! [X118,X120,X119,X117] : (s(bool,X118) = s(bool,X119) | s(bool,X119) = s(bool,X120) | s(bool,'F_0') = s(bool,X120) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X117),s(bool,X118)))) )),
  inference(superposition,[],[f1664,f103])).
fof(f2125,plain,(
  ( ! [X21,X19,X17,X20,X18,X16] : (s(bool,X16) = s(bool,X17) | s(bool,X18) = s(bool,X19) | s(bool,'F_0') = s(bool,X19) | s(bool,X17) = s(bool,X18) | s(bool,X16) = s(bool,X20) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X20),s(bool,X21)))) )),
  inference(superposition,[],[f1664,f326])).
fof(f2155,plain,(
  ( ! [X87,X85,X86,X84] : (s(bool,'T_0') = s(bool,X84) | s(bool,X86) = s(bool,X87) | s(bool,'F_0') = s(bool,X87) | s(bool,X86) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X84),s(bool,X85)))) )),
  inference(superposition,[],[f1664,f102])).
fof(f2156,plain,(
  ( ! [X90,X88,X91,X89] : (s(bool,'T_0') = s(bool,X89) | s(bool,X90) = s(bool,X91) | s(bool,'F_0') = s(bool,X91) | s(bool,X90) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X88),s(bool,X89)))) )),
  inference(superposition,[],[f1664,f103])).
fof(f2216,plain,(
  ( ! [X237,X235,X233,X236,X234] : (s(bool,X236) != s(bool,f(s(X233,sK1(X234,X235,bool,X233)))) | s(fun(X233,bool),X234) = s(fun(X233,bool),X235) | s(bool,'F_0') = s(bool,X236) | s(bool,X237) = s(bool,g(s(X233,sK1(X234,X235,bool,X233)))) | s(bool,'T_0') = s(bool,X237)) )),
  inference(superposition,[],[f108,f1664])).
fof(f2219,plain,(
  ( ! [X249,X246,X250,X248,X247] : (s(X248,f(s(bool,X249))) != s(X248,g(s(bool,X249))) | s(fun(bool,X248),X246) = s(fun(bool,X248),X247) | s(bool,'F_0') = s(bool,X249) | s(bool,X250) = s(bool,sK1(X246,X247,X248,bool)) | s(bool,'T_0') = s(bool,X250)) )),
  inference(superposition,[],[f108,f1664])).
fof(f2279,plain,(
  ( ! [X231,X229,X233,X230,X232] : (s(bool,X232) != s(bool,f(s(X229,sK1(X230,X231,bool,X229)))) | s(fun(X229,bool),X230) = s(fun(X229,bool),X231) | s(bool,'F_0') = s(bool,g(s(X229,sK1(X230,X231,bool,X229)))) | s(bool,X232) = s(bool,X233) | s(bool,'T_0') = s(bool,X233)) )),
  inference(superposition,[],[f108,f1664])).
fof(f2282,plain,(
  ( ! [X243,X246,X244,X242,X245] : (s(X244,f(s(bool,X245))) != s(X244,g(s(bool,X245))) | s(fun(bool,X244),X242) = s(fun(bool,X244),X243) | s(bool,'F_0') = s(bool,sK1(X242,X243,X244,bool)) | s(bool,X245) = s(bool,X246) | s(bool,'T_0') = s(bool,X246)) )),
  inference(superposition,[],[f108,f1664])).
fof(f2363,plain,(
  ( ! [X185,X182,X186,X184,X183] : (s(bool,'F_0') != s(bool,f(s(X182,sK1(X183,X184,bool,X182)))) | s(fun(X182,bool),X183) = s(fun(X182,bool),X184) | s(bool,X185) = s(bool,g(s(X182,sK1(X183,X184,bool,X182)))) | s(bool,X185) = s(bool,X186) | s(bool,'T_0') = s(bool,X186)) )),
  inference(superposition,[],[f108,f1664])).
fof(f2366,plain,(
  ( ! [X198,X196,X199,X197,X195] : (s(X197,f(s(bool,'F_0'))) != s(X197,g(s(bool,'F_0'))) | s(fun(bool,X197),X195) = s(fun(bool,X197),X196) | s(bool,X198) = s(bool,sK1(X195,X196,X197,bool)) | s(bool,X198) = s(bool,X199) | s(bool,'T_0') = s(bool,X199)) )),
  inference(superposition,[],[f108,f1664])).
fof(f2426,plain,(
  ( ! [X237,X235,X233,X236,X234] : (s(bool,X236) != s(bool,f(s(X233,sK1(X234,X235,bool,X233)))) | s(fun(X233,bool),X234) = s(fun(X233,bool),X235) | s(bool,X237) = s(bool,g(s(X233,sK1(X234,X235,bool,X233)))) | s(bool,'F_0') = s(bool,X237) | s(bool,'T_0') = s(bool,X236)) )),
  inference(superposition,[],[f108,f1664])).
fof(f2429,plain,(
  ( ! [X249,X246,X250,X248,X247] : (s(X248,f(s(bool,X249))) != s(X248,g(s(bool,X249))) | s(fun(bool,X248),X246) = s(fun(bool,X248),X247) | s(bool,X250) = s(bool,sK1(X246,X247,X248,bool)) | s(bool,'F_0') = s(bool,X250) | s(bool,'T_0') = s(bool,X249)) )),
  inference(superposition,[],[f108,f1664])).
fof(f2489,plain,(
  ( ! [X237,X235,X233,X236,X234] : (s(bool,X236) != s(bool,f(s(X233,sK1(X234,X235,bool,X233)))) | s(fun(X233,bool),X234) = s(fun(X233,bool),X235) | s(bool,X236) = s(bool,X237) | s(bool,'F_0') = s(bool,X237) | s(bool,'T_0') = s(bool,g(s(X233,sK1(X234,X235,bool,X233))))) )),
  inference(superposition,[],[f108,f1664])).
fof(f2492,plain,(
  ( ! [X249,X246,X250,X248,X247] : (s(X248,f(s(bool,X249))) != s(X248,g(s(bool,X249))) | s(fun(bool,X248),X246) = s(fun(bool,X248),X247) | s(bool,X249) = s(bool,X250) | s(bool,'F_0') = s(bool,X250) | s(bool,'T_0') = s(bool,sK1(X246,X247,X248,bool))) )),
  inference(superposition,[],[f108,f1664])).
fof(f2494,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,X3) != s(bool,X4) | p(s(bool,X4)) | s(bool,X5) = s(bool,X6) | s(bool,'F_0') = s(bool,X6) | s(bool,X3) = s(bool,X5)) )),
  inference(superposition,[],[f79,f1664])).
fof(f2495,plain,(
  ( ! [X12,X10,X8,X7,X11,X9] : (s(X8,X9) = s(X8,'const.bool.COND_3'(s(bool,X7),s(X8,X9),s(X8,X10))) | s(bool,X11) = s(bool,X12) | s(bool,'F_0') = s(bool,X12) | s(bool,X7) = s(bool,X11)) )),
  inference(superposition,[],[f102,f1664])).
fof(f2498,plain,(
  ( ! [X28,X26,X29,X27,X25] : (s(bool,X25) = s(bool,'const.bool.COND_3'(s(bool,X25),s(bool,X26),s(bool,X27))) | s(bool,'F_0') = s(bool,X26) | s(bool,X28) = s(bool,X29) | s(bool,'F_0') = s(bool,X29) | s(bool,X25) = s(bool,X28)) )),
  inference(superposition,[],[f187,f1664])).
fof(f2499,plain,(
  ( ! [X30,X33,X31,X34,X32] : (s(bool,X30) = s(bool,'const.bool.COND_3'(s(bool,X30),s(bool,X31),s(bool,X32))) | s(bool,X31) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))) | s(bool,X33) = s(bool,X34) | s(bool,'F_0') = s(bool,X34) | s(bool,X30) = s(bool,X33)) )),
  inference(superposition,[],[f263,f1664])).
fof(f2500,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (s(bool,X35) = s(bool,'const.bool.COND_3'(s(bool,X35),s(bool,X36),s(bool,X37))) | s(bool,X36) = s(bool,X38) | s(bool,'T_0') = s(bool,X38) | s(bool,X39) = s(bool,X40) | s(bool,'F_0') = s(bool,X40) | s(bool,X35) = s(bool,X39)) )),
  inference(superposition,[],[f326,f1664])).
fof(f2502,plain,(
  ( ! [X47,X50,X48,X51,X49] : (s(bool,X47) != s(bool,X48) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X48),s(bool,X49))) | s(bool,X50) = s(bool,X51) | s(bool,'F_0') = s(bool,X51) | s(bool,X47) = s(bool,X50)) )),
  inference(superposition,[],[f376,f1664])).
fof(f2503,plain,(
  ( ! [X54,X52,X56,X55,X53] : (s(bool,X53) != s(bool,'const.bool.COND_3'(s(bool,X52),s(bool,X52),s(bool,X54))) | s(bool,'T_0') = s(bool,X53) | s(bool,X55) = s(bool,X56) | s(bool,'F_0') = s(bool,X56) | s(bool,X52) = s(bool,X55)) )),
  inference(superposition,[],[f792,f1664])).
fof(f2585,plain,(
  ( ! [X216,X215,X213,X217,X214] : (s(bool,'T_0') != s(bool,f(s(X213,sK1(X214,X215,bool,X213)))) | s(fun(X213,bool),X214) = s(fun(X213,bool),X215) | s(bool,X216) = s(bool,X217) | s(bool,'F_0') = s(bool,X217) | s(bool,X216) = s(bool,g(s(X213,sK1(X214,X215,bool,X213))))) )),
  inference(superposition,[],[f108,f1664])).
fof(f2588,plain,(
  ( ! [X229,X227,X230,X228,X226] : (s(X228,f(s(bool,'T_0'))) != s(X228,g(s(bool,'T_0'))) | s(fun(bool,X228),X226) = s(fun(bool,X228),X227) | s(bool,X229) = s(bool,X230) | s(bool,'F_0') = s(bool,X230) | s(bool,X229) = s(bool,sK1(X226,X227,X228,bool))) )),
  inference(superposition,[],[f108,f1664])).
fof(f2829,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X0) != s(bool,'const.bool.COND_3'(s(bool,X1),s(bool,X1),s(bool,X2))) | s(bool,'F_0') = s(bool,X1)) )),
  inference(condensation,[],[f2828])).
fof(f2828,plain,(
  ( ! [X165,X163,X166,X164] : (s(bool,X164) != s(bool,'const.bool.COND_3'(s(bool,X163),s(bool,X163),s(bool,X165))) | s(bool,'T_0') = s(bool,X164) | s(bool,'F_0') = s(bool,X163) | s(bool,'T_0') = s(bool,X166)) )),
  inference(duplicate_literal_removal,[],[f2197])).
fof(f2197,plain,(
  ( ! [X165,X163,X166,X164] : (s(bool,X164) != s(bool,'const.bool.COND_3'(s(bool,X163),s(bool,X163),s(bool,X165))) | s(bool,'T_0') = s(bool,X164) | s(bool,'F_0') = s(bool,X163) | s(bool,'T_0') = s(bool,X166) | s(bool,'T_0') = s(bool,X166)) )),
  inference(superposition,[],[f792,f1664])).
fof(f2831,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1))) | s(bool,X0) != s(bool,X2) | s(bool,'F_0') = s(bool,X2)) )),
  inference(condensation,[],[f2830])).
fof(f2830,plain,(
  ( ! [X161,X159,X162,X160] : (s(bool,X159) != s(bool,X160) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X160),s(bool,X161))) | s(bool,'F_0') = s(bool,X159) | s(bool,'T_0') = s(bool,X162)) )),
  inference(duplicate_literal_removal,[],[f2196])).
fof(f2196,plain,(
  ( ! [X161,X159,X162,X160] : (s(bool,X159) != s(bool,X160) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X160),s(bool,X161))) | s(bool,'F_0') = s(bool,X159) | s(bool,'T_0') = s(bool,X162) | s(bool,'T_0') = s(bool,X162)) )),
  inference(superposition,[],[f376,f1664])).
fof(f2835,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,'const.bool.COND_3'(s(bool,X1),s(bool,X2),s(bool,X3))) | s(bool,X0) = s(bool,X2) | s(bool,'F_0') = s(bool,X1)) )),
  inference(condensation,[],[f2834])).
fof(f2834,plain,(
  ( ! [X152,X151,X149,X153,X150] : (s(bool,X149) = s(bool,'const.bool.COND_3'(s(bool,X149),s(bool,X150),s(bool,X151))) | s(bool,X150) = s(bool,X152) | s(bool,'T_0') = s(bool,X152) | s(bool,'F_0') = s(bool,X149) | s(bool,'T_0') = s(bool,X153)) )),
  inference(duplicate_literal_removal,[],[f2194])).
fof(f2194,plain,(
  ( ! [X152,X151,X149,X153,X150] : (s(bool,X149) = s(bool,'const.bool.COND_3'(s(bool,X149),s(bool,X150),s(bool,X151))) | s(bool,X150) = s(bool,X152) | s(bool,'T_0') = s(bool,X152) | s(bool,'F_0') = s(bool,X149) | s(bool,'T_0') = s(bool,X153) | s(bool,'T_0') = s(bool,X153)) )),
  inference(superposition,[],[f326,f1664])).
fof(f3076,plain,(
  ( ! [X76,X77] : (s(bool,X76) = s(bool,X77) | s(bool,'F_0') = s(bool,X77) | s(bool,X76) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0)))))) )),
  inference(subsumption_resolution,[],[f2154,f110])).
fof(f2154,plain,(
  ( ! [X76,X77] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X76) = s(bool,X77) | s(bool,'F_0') = s(bool,X77) | s(bool,X76) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0)))))) )),
  inference(superposition,[],[f1664,f118])).
fof(f3148,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2)) )),
  inference(condensation,[],[f3147])).
fof(f3147,plain,(
  ( ! [X17,X15,X18,X16] : (p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X15) = s(bool,X17) | s(bool,X15) = s(bool,X18) | s(bool,'F_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f2372,f79])).
fof(f2372,plain,(
  ( ! [X17,X15,X18,X16] : (p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X15) = s(bool,X17) | s(bool,X15) = s(bool,X18) | s(bool,'F_0') = s(bool,X18) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f104,f1664])).
fof(f3170,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) | s(bool,X1) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))) | s(bool,'F_0') = s(bool,X1)) )),
  inference(condensation,[],[f3169])).
fof(f3169,plain,(
  ( ! [X191,X192,X190] : (p(s(bool,X191)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X190)))) | s(bool,X192) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X190))))) | s(bool,'F_0') = s(bool,X192)) )),
  inference(subsumption_resolution,[],[f2416,f79])).
fof(f2416,plain,(
  ( ! [X191,X192,X190] : (p(s(bool,X191)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X190)))) | s(bool,X192) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X190))))) | s(bool,'F_0') = s(bool,X192) | s(bool,'T_0') = s(bool,X191)) )),
  inference(superposition,[],[f77,f1664])).
fof(f3201,plain,(
  ( ! [X26,X24,X23,X25] : (~p(s(bool,X24)) | s(bool,X23) = s(bool,X24) | s(bool,X25) = s(bool,X26) | s(bool,'F_0') = s(bool,X26) | s(bool,X23) = s(bool,X25)) )),
  inference(subsumption_resolution,[],[f2539,f71])).
fof(f2539,plain,(
  ( ! [X26,X24,X23,X25] : (~p(s(bool,'T_0')) | ~p(s(bool,X24)) | s(bool,X23) = s(bool,X24) | s(bool,X25) = s(bool,X26) | s(bool,'F_0') = s(bool,X26) | s(bool,X23) = s(bool,X25)) )),
  inference(superposition,[],[f107,f1664])).
fof(f3219,plain,(
  ( ! [X182,X183,X181] : (s('type.realax.real',X181) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X181))) | s(bool,X182) = s(bool,X183) | s(bool,'F_0') = s(bool,X183) | s(bool,X182) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X181)))) )),
  inference(forward_demodulation,[],[f2579,f102])).
fof(f2579,plain,(
  ( ! [X182,X183,X181] : (s('type.realax.real','const.real.abs_1'(s('type.realax.real',X181))) = s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real',X181),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X181))))) | s(bool,X182) = s(bool,X183) | s(bool,'F_0') = s(bool,X183) | s(bool,X182) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X181)))) )),
  inference(superposition,[],[f74,f1664])).
fof(f4099,plain,(
  ( ! [X36] : (s(bool,X36) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))) | s(bool,'F_0') = s(bool,X36)) ) | $spl10),
  inference(subsumption_resolution,[],[f4081,f73])).
fof(f4081,plain,(
  ( ! [X36] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,X36) | s(bool,X36) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0)))))) ) | $spl10),
  inference(superposition,[],[f3094,f118])).
fof(f4143,plain,(
  ( ! [X70,X71,X69] : (s(bool,'F_0') = s(bool,X70) | s(bool,X71) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X69),s(bool,X70))) | s(bool,'F_0') = s(bool,X71)) )),
  inference(superposition,[],[f3059,f103])).
fof(f4173,plain,(
  ( ! [X94,X92,X93] : (s(bool,X93) = s(bool,X94) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X92),s(bool,X93))) | s(bool,'F_0') = s(bool,X94)) )),
  inference(superposition,[],[f3059,f103])).
fof(f4245,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(X6,X7) = s(X6,'const.bool.COND_3'(s(bool,X5),s(X6,X8),s(X6,X7))) | s(bool,X5) = s(bool,X9) | s(bool,'F_0') = s(bool,X9)) )),
  inference(superposition,[],[f103,f3059])).
fof(f4321,plain,(
  ( ! [X134,X132,X135,X133] : (s(bool,'F_0') != s(bool,f(s(X132,sK1(X133,X134,bool,X132)))) | s(fun(X132,bool),X133) = s(fun(X132,bool),X134) | s(bool,X135) = s(bool,g(s(X132,sK1(X133,X134,bool,X132)))) | s(bool,'F_0') = s(bool,X135)) )),
  inference(superposition,[],[f108,f3059])).
fof(f4324,plain,(
  ( ! [X144,X142,X145,X143] : (s(X144,f(s(bool,'F_0'))) != s(X144,g(s(bool,'F_0'))) | s(fun(bool,X144),X142) = s(fun(bool,X144),X143) | s(bool,X145) = s(bool,sK1(X142,X143,X144,bool)) | s(bool,'F_0') = s(bool,X145)) )),
  inference(superposition,[],[f108,f3059])).
fof(f4379,plain,(
  ( ! [X173,X171,X174,X172] : (s(bool,X174) != s(bool,f(s(X171,sK1(X172,X173,bool,X171)))) | s(fun(X171,bool),X172) = s(fun(X171,bool),X173) | s(bool,'F_0') = s(bool,g(s(X171,sK1(X172,X173,bool,X171)))) | s(bool,'F_0') = s(bool,X174)) )),
  inference(superposition,[],[f108,f3059])).
fof(f4382,plain,(
  ( ! [X182,X184,X183,X181] : (s(X183,f(s(bool,X184))) != s(X183,g(s(bool,X184))) | s(fun(bool,X183),X181) = s(fun(bool,X183),X182) | s(bool,'F_0') = s(bool,sK1(X181,X182,X183,bool)) | s(bool,'F_0') = s(bool,X184)) )),
  inference(superposition,[],[f108,f3059])).
fof(f4966,plain,(
  ( ! [X92,X90,X91] : (s(bool,X91) != s(bool,X92) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X90),s(bool,X91))) | s(bool,'F_0') = s(bool,X92)) )),
  inference(superposition,[],[f3107,f103])).
fof(f4996,plain,(
  ( ! [X94,X92,X93] : (s(bool,X93) != s(bool,X94) | s(bool,'T_0') = s(bool,X94) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X92),s(bool,X93)))) )),
  inference(superposition,[],[f3107,f103])).
fof(f5047,plain,(
  ( ! [X14,X17,X15,X13,X16] : (s(bool,X13) != s(bool,X14) | s(bool,'F_0') = s(bool,X15) | s(bool,X14) = s(bool,X15) | s(bool,X13) = s(bool,X16) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X16),s(bool,X17)))) )),
  inference(superposition,[],[f3180,f326])).
fof(f5090,plain,(
  ( ! [X94,X95,X93] : (s(bool,'T_0') != s(bool,X93) | s(bool,'F_0') = s(bool,X95) | s(bool,X95) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X93),s(bool,X94)))) )),
  inference(superposition,[],[f3180,f102])).
fof(f5091,plain,(
  ( ! [X97,X98,X96] : (s(bool,'T_0') != s(bool,X97) | s(bool,'F_0') = s(bool,X98) | s(bool,X98) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X96),s(bool,X97)))) )),
  inference(superposition,[],[f3180,f103])).
fof(f5193,plain,(
  ( ! [X154,X152,X151,X153,X150] : (s(X152,X153) = s(X152,'const.bool.COND_3'(s(bool,X151),s(X152,X153),s(X152,X154))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X150),s(bool,X151)))) ) | $spl12),
  inference(superposition,[],[f3095,f103])).
fof(f5194,plain,(
  ( ! [X158,X156,X159,X157,X155,X160] : (s(X158,X159) = s(X158,'const.bool.COND_3'(s(bool,X156),s(X158,X159),s(X158,X160))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X155),s(bool,X156),s(bool,X157))) | s(bool,'F_0') = s(bool,X155)) ) | $spl12),
  inference(superposition,[],[f3095,f3095])).
fof(f5260,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') = s(bool,X4) | s(bool,'F_0') = s(bool,X3) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X3),s(bool,X4),s(bool,X5)))) ) | $spl12),
  inference(superposition,[],[f3095,f76])).
fof(f5261,plain,(
  ( ! [X6,X8,X7,X9] : (s(bool,'T_0') = s(bool,X7) | s(bool,'F_0') = s(bool,X6) | s(bool,X9) = s(bool,'const.bool.COND_3'(s(bool,X6),s(bool,X7),s(bool,X8))) | s(bool,'T_0') = s(bool,X9)) ) | $spl12),
  inference(superposition,[],[f3095,f131])).
fof(f5262,plain,(
  ( ! [X12,X10,X13,X11] : (s(bool,X11) = s(bool,X13) | s(bool,'F_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X10),s(bool,X11),s(bool,X12)))) ) | $spl12),
  inference(superposition,[],[f3095,f131])).
fof(f5265,plain,(
  ( ! [X26,X24,X23,X25,X22] : (s(bool,'T_0') = s(bool,X23) | s(bool,'F_0') = s(bool,X22) | s(bool,X25) = s(bool,'const.bool.COND_3'(s(bool,X22),s(bool,X23),s(bool,X24))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X25),s(bool,X26)))) ) | $spl12),
  inference(superposition,[],[f3095,f326])).
fof(f5266,plain,(
  ( ! [X30,X28,X31,X29,X27] : (s(bool,X28) = s(bool,X30) | s(bool,'F_0') = s(bool,X27) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X30),s(bool,X31))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X27),s(bool,X28),s(bool,X29)))) ) | $spl12),
  inference(superposition,[],[f3095,f326])).
fof(f5268,plain,(
  ( ! [X39,X37,X41,X38,X40] : (s(bool,X38) = s(bool,X40) | s(bool,'F_0') = s(bool,X37) | s(bool,'F_0') = s(bool,X40) | s(bool,X41) = s(bool,'const.bool.COND_3'(s(bool,X37),s(bool,X38),s(bool,X39))) | s(bool,'T_0') = s(bool,X41)) ) | $spl12),
  inference(superposition,[],[f3095,f1664])).
fof(f5269,plain,(
  ( ! [X45,X43,X46,X44,X42] : (s(bool,X43) = s(bool,X45) | s(bool,'F_0') = s(bool,X42) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X42),s(bool,X43),s(bool,X44))) | s(bool,X45) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) ) | $spl12),
  inference(superposition,[],[f3095,f1664])).
fof(f5270,plain,(
  ( ! [X47,X50,X48,X51,X49] : (s(bool,'F_0') = s(bool,X48) | s(bool,'F_0') = s(bool,X47) | s(bool,X50) = s(bool,'const.bool.COND_3'(s(bool,X47),s(bool,X48),s(bool,X49))) | s(bool,X50) = s(bool,X51) | s(bool,'T_0') = s(bool,X51)) ) | $spl12),
  inference(superposition,[],[f3095,f1664])).
fof(f5271,plain,(
  ( ! [X54,X52,X56,X55,X53] : (s(bool,'T_0') = s(bool,X53) | s(bool,'F_0') = s(bool,X52) | s(bool,X55) = s(bool,X56) | s(bool,'F_0') = s(bool,X56) | s(bool,X55) = s(bool,'const.bool.COND_3'(s(bool,X52),s(bool,X53),s(bool,X54)))) ) | $spl12),
  inference(superposition,[],[f3095,f1664])).
fof(f5272,plain,(
  ( ! [X61,X59,X57,X60,X58] : (s(bool,X58) = s(bool,X60) | s(bool,'F_0') = s(bool,X57) | s(bool,X61) = s(bool,'const.bool.COND_3'(s(bool,X57),s(bool,X58),s(bool,X59))) | s(bool,'F_0') = s(bool,X61) | s(bool,'T_0') = s(bool,X60)) ) | $spl12),
  inference(superposition,[],[f3095,f1664])).
fof(f5273,plain,(
  ( ! [X66,X64,X62,X65,X63] : (s(bool,X63) = s(bool,X65) | s(bool,'F_0') = s(bool,X62) | s(bool,X65) = s(bool,X66) | s(bool,'F_0') = s(bool,X66) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X62),s(bool,X63),s(bool,X64)))) ) | $spl12),
  inference(superposition,[],[f3095,f1664])).
fof(f5274,plain,(
  ( ! [X70,X68,X69,X67] : (s(bool,'F_0') = s(bool,X68) | s(bool,'F_0') = s(bool,X67) | s(bool,X70) = s(bool,'const.bool.COND_3'(s(bool,X67),s(bool,X68),s(bool,X69))) | s(bool,'F_0') = s(bool,X70)) ) | $spl12),
  inference(superposition,[],[f3095,f3059])).
fof(f5275,plain,(
  ( ! [X74,X72,X71,X73] : (s(bool,X72) = s(bool,X74) | s(bool,'F_0') = s(bool,X71) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X71),s(bool,X72),s(bool,X73))) | s(bool,'F_0') = s(bool,X74)) ) | $spl12),
  inference(superposition,[],[f3095,f3059])).
fof(f5279,plain,(
  ( ! [X93] : (s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X93))) | s('type.realax.real',X93) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X93)))) ) | $spl12),
  inference(superposition,[],[f3095,f74])).
fof(f5300,plain,(
  ( ! [X80,X78,X81,X79,X82] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X79),s(bool,X81))) | s(bool,X82) = s(bool,'const.bool.COND_3'(s(bool,X78),s(bool,X79),s(bool,X80))) | s(bool,'T_0') = s(bool,X82) | s(bool,'F_0') = s(bool,X78)) ) | $spl12),
  inference(superposition,[],[f326,f3095])).
fof(f5307,plain,(
  ( ! [X111,X114,X112,X113] : (s(bool,X112) != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X114))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X111),s(bool,X112),s(bool,X113))) | s(bool,'F_0') = s(bool,X111)) ) | $spl12),
  inference(superposition,[],[f792,f3095])).
fof(f5322,plain,(
  ( ! [X182,X180,X181,X179] : (s(bool,X180) != s(bool,X182) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X179),s(bool,X180),s(bool,X181))) | s(bool,'F_0') = s(bool,X182) | s(bool,'F_0') = s(bool,X179)) ) | $spl12),
  inference(superposition,[],[f3107,f3095])).
fof(f5323,plain,(
  ( ! [X185,X186,X184,X183] : (s(bool,X184) != s(bool,X186) | s(bool,'T_0') = s(bool,X186) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X183),s(bool,X184),s(bool,X185))) | s(bool,'F_0') = s(bool,X183)) ) | $spl12),
  inference(superposition,[],[f3107,f3095])).
fof(f5324,plain,(
  ( ! [X189,X187,X190,X188] : (s(bool,'T_0') != s(bool,X188) | s(bool,'F_0') = s(bool,X190) | s(bool,X190) = s(bool,'const.bool.COND_3'(s(bool,X187),s(bool,X188),s(bool,X189))) | s(bool,'F_0') = s(bool,X187)) ) | $spl12),
  inference(superposition,[],[f3180,f3095])).
fof(f5816,plain,(
  ( ! [X107,X105,X106,X104] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X105),s(bool,X107))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X104),s(bool,X105),s(bool,X106))) | s(bool,'F_0') = s(bool,X104)) ) | $spl12),
  inference(superposition,[],[f187,f3095])).
fof(f6124,plain,(
  ( ! [X68,X69,X67] : (s(bool,X68) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X67),s(bool,X68),s(bool,X69))) | s(bool,'F_0') = s(bool,X67)) ) | $spl12),
  inference(superposition,[],[f249,f3095])).
fof(f6425,plain,(
  ( ! [X175,X173,X171,X174,X172] : (s(X173,X174) = s(X173,'const.bool.COND_3'(s(bool,X172),s(X173,X175),s(X173,X174))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X171),s(bool,X172)))) )),
  inference(superposition,[],[f207,f103])).
fof(f6427,plain,(
  ( ! [X187,X185,X182,X186,X184,X183] : (s(X185,X186) = s(X185,'const.bool.COND_3'(s(bool,X183),s(X185,X187),s(X185,X186))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X182),s(bool,X183),s(bool,X184))) | s(bool,'F_0') = s(bool,X182)) ) | $spl12),
  inference(superposition,[],[f207,f3095])).
fof(f6507,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X6),s(bool,X7),s(bool,X8)))) )),
  inference(superposition,[],[f207,f76])).
fof(f6516,plain,(
  ( ! [X47,X45,X43,X46,X44] : (s(bool,X45) = s(bool,X46) | s(bool,'T_0') = s(bool,X43) | s(bool,'F_0') = s(bool,X46) | s(bool,X47) = s(bool,'const.bool.COND_3'(s(bool,X43),s(bool,X44),s(bool,X45))) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f207,f1664])).
fof(f6517,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s(bool,X50) = s(bool,X51) | s(bool,'T_0') = s(bool,X48) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X48),s(bool,X49),s(bool,X50))) | s(bool,X51) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(superposition,[],[f207,f1664])).
fof(f6518,plain,(
  ( ! [X57,X54,X56,X55,X53] : (s(bool,'F_0') = s(bool,X55) | s(bool,'T_0') = s(bool,X53) | s(bool,X56) = s(bool,'const.bool.COND_3'(s(bool,X53),s(bool,X54),s(bool,X55))) | s(bool,X56) = s(bool,X57) | s(bool,'T_0') = s(bool,X57)) )),
  inference(superposition,[],[f207,f1664])).
fof(f6519,plain,(
  ( ! [X61,X59,X62,X60,X58] : (s(bool,'T_0') = s(bool,X60) | s(bool,'T_0') = s(bool,X58) | s(bool,X61) = s(bool,X62) | s(bool,'F_0') = s(bool,X62) | s(bool,X61) = s(bool,'const.bool.COND_3'(s(bool,X58),s(bool,X59),s(bool,X60)))) )),
  inference(superposition,[],[f207,f1664])).
fof(f6520,plain,(
  ( ! [X66,X64,X67,X65,X63] : (s(bool,X65) = s(bool,X66) | s(bool,'T_0') = s(bool,X63) | s(bool,X67) = s(bool,'const.bool.COND_3'(s(bool,X63),s(bool,X64),s(bool,X65))) | s(bool,'F_0') = s(bool,X67) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f207,f1664])).
fof(f6521,plain,(
  ( ! [X70,X68,X72,X71,X69] : (s(bool,X70) = s(bool,X71) | s(bool,'T_0') = s(bool,X68) | s(bool,X71) = s(bool,X72) | s(bool,'F_0') = s(bool,X72) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X68),s(bool,X69),s(bool,X70)))) )),
  inference(superposition,[],[f207,f1664])).
fof(f6522,plain,(
  ( ! [X76,X74,X75,X73] : (s(bool,'F_0') = s(bool,X75) | s(bool,'T_0') = s(bool,X73) | s(bool,X76) = s(bool,'const.bool.COND_3'(s(bool,X73),s(bool,X74),s(bool,X75))) | s(bool,'F_0') = s(bool,X76)) )),
  inference(superposition,[],[f207,f3059])).
fof(f6523,plain,(
  ( ! [X80,X78,X79,X77] : (s(bool,X79) = s(bool,X80) | s(bool,'T_0') = s(bool,X77) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X77),s(bool,X78),s(bool,X79))) | s(bool,'F_0') = s(bool,X80)) )),
  inference(superposition,[],[f207,f3059])).
fof(f6555,plain,(
  ( ! [X101,X99,X100,X98] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X100),s(bool,X101))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X98),s(bool,X99),s(bool,X100))) | s(bool,'T_0') = s(bool,X98)) )),
  inference(superposition,[],[f187,f207])).
fof(f6580,plain,(
  ( ! [X212,X210,X215,X213,X211,X214] : (s(X213,X214) = s(X213,'const.bool.COND_3'(s(bool,X212),s(X213,X214),s(X213,X215))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X210),s(bool,X211),s(bool,X212))) | s(bool,'T_0') = s(bool,X210)) ) | $spl12),
  inference(superposition,[],[f3095,f207])).
fof(f6581,plain,(
  ( ! [X218,X216,X219,X217] : (s(bool,X218) != s(bool,X219) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X216),s(bool,X217),s(bool,X218))) | s(bool,'F_0') = s(bool,X219) | s(bool,'T_0') = s(bool,X216)) )),
  inference(superposition,[],[f3107,f207])).
fof(f6582,plain,(
  ( ! [X222,X220,X223,X221] : (s(bool,X222) != s(bool,X223) | s(bool,'T_0') = s(bool,X223) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X220),s(bool,X221),s(bool,X222))) | s(bool,'T_0') = s(bool,X220)) )),
  inference(superposition,[],[f3107,f207])).
fof(f6586,plain,(
  ( ! [X239,X237,X240,X238] : (s(bool,'T_0') != s(bool,X239) | s(bool,'F_0') = s(bool,X240) | s(bool,X240) = s(bool,'const.bool.COND_3'(s(bool,X237),s(bool,X238),s(bool,X239))) | s(bool,'T_0') = s(bool,X237)) )),
  inference(superposition,[],[f3180,f207])).
fof(f6691,plain,(
  ( ! [X66,X65] : (s(bool,X66) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X65),s(bool,X66)))) ) | $spl10),
  inference(superposition,[],[f4099,f103])).
fof(f6692,plain,(
  ( ! [X68,X69,X67] : (s(bool,X69) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X67),s(bool,X68),s(bool,X69))) | s(bool,'T_0') = s(bool,X67)) ) | $spl10),
  inference(superposition,[],[f4099,f207])).
fof(f6693,plain,(
  ( ! [X70,X72,X71] : (s(bool,X71) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X70),s(bool,X71),s(bool,X72))) | s(bool,'F_0') = s(bool,X70)) ) | ($spl10 | $spl12)),
  inference(superposition,[],[f4099,f3095])).
fof(f6845,plain,(
  ( ! [X85,X86] : (s(bool,'F_0') = s(bool,X86) | s(bool,X86) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X85))) | s('type.realax.real',X85) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X85)))) ) | ($spl10 | $spl12)),
  inference(subsumption_resolution,[],[f6805,f73])).
fof(f6805,plain,(
  ( ! [X85,X86] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,X86) | s(bool,X86) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X85))) | s('type.realax.real',X85) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X85)))) ) | ($spl10 | $spl12)),
  inference(superposition,[],[f3094,f5279])).
fof(f6848,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X2))) | s(bool,'F_0') = s(bool,X2) | s(bool,'F_0') = s(bool,X3) | s(bool,X0) = s(bool,X3)) )),
  inference(superposition,[],[f219,f3059])).
fof(f6850,plain,(
  ( ! [X12,X10,X8,X11,X9] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X8),s(bool,X9),s(bool,X10))) | s(bool,'F_0') = s(bool,X10) | s(bool,X8) = s(bool,X11) | s(bool,X11) = s(bool,X12) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f219,f1664])).
fof(f6890,plain,(
  ( ! [X70,X68,X69,X67] : (s(bool,X67) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X67),s(bool,X68))) | s(bool,'F_0') = s(bool,X68) | s(bool,X69) = s(bool,X70) | s(bool,'F_0') = s(bool,X70) | s(bool,X67) = s(bool,X69)) )),
  inference(superposition,[],[f219,f1664])).
fof(f6893,plain,(
  ( ! [X80,X78,X79,X77] : (s(bool,X77) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X77),s(bool,X78))) | s(bool,'F_0') = s(bool,X78) | s(bool,X77) = s(bool,X79) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X79),s(bool,X80)))) )),
  inference(superposition,[],[f219,f326])).
fof(f6945,plain,(
  ( ! [X111,X112,X110] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X112),s(bool,X111))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X110),s(bool,X111)))) )),
  inference(superposition,[],[f219,f103])).
fof(f6946,plain,(
  ( ! [X116,X114,X115,X113] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X116),s(bool,X115))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X113),s(bool,X114),s(bool,X115))) | s(bool,'T_0') = s(bool,X113)) )),
  inference(superposition,[],[f219,f207])).
fof(f6947,plain,(
  ( ! [X118,X120,X119,X117] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X120),s(bool,X118))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X117),s(bool,X118),s(bool,X119))) | s(bool,'F_0') = s(bool,X117)) ) | $spl12),
  inference(superposition,[],[f219,f3095])).
fof(f7090,plain,(
  ( ! [X37,X35,X36,X34] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X34),s(bool,X35),s(bool,X36))) | s(bool,'F_0') = s(bool,X36) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X34),s(bool,X37)))) )),
  inference(subsumption_resolution,[],[f6858,f110])).
fof(f6858,plain,(
  ( ! [X37,X35,X36,X34] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X34),s(bool,X35),s(bool,X36))) | s(bool,'F_0') = s(bool,X36) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X34),s(bool,X37))) | s(bool,'T_0') = s(bool,'F_0')) )),
  inference(superposition,[],[f219,f326])).
fof(f7097,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X6),s(bool,X7))) | s(bool,'F_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f6970,f71])).
fof(f6970,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X6),s(bool,X7))) | s(bool,'F_0') = s(bool,X7)) )),
  inference(superposition,[],[f83,f219])).
fof(f7406,plain,(
  ( ! [X45,X43,X44,X42] : (s(bool,'F_0') = s(bool,X42) | s(bool,X44) = s(bool,'const.bool.COND_3'(s(bool,X42),s(bool,'T_0'),s(bool,X43))) | s(bool,X44) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) ) | $spl12),
  inference(subsumption_resolution,[],[f7320,f73])).
fof(f7320,plain,(
  ( ! [X45,X43,X44,X42] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,X42) | s(bool,X44) = s(bool,'const.bool.COND_3'(s(bool,X42),s(bool,'T_0'),s(bool,X43))) | s(bool,X44) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) ) | $spl12),
  inference(superposition,[],[f7176,f1664])).
fof(f7411,plain,(
  ( ! [X59,X60,X58] : (s(bool,'F_0') = s(bool,X58) | s(bool,X60) = s(bool,'const.bool.COND_3'(s(bool,X58),s(bool,'T_0'),s(bool,X59))) | s(bool,'F_0') = s(bool,X60)) ) | $spl12),
  inference(subsumption_resolution,[],[f7324,f73])).
fof(f7324,plain,(
  ( ! [X59,X60,X58] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,X58) | s(bool,X60) = s(bool,'const.bool.COND_3'(s(bool,X58),s(bool,'T_0'),s(bool,X59))) | s(bool,'F_0') = s(bool,X60)) ) | $spl12),
  inference(superposition,[],[f7176,f3059])).
fof(f8060,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X4),s(bool,X3),s(bool,X5))) | s(bool,'F_0') = s(bool,X4) | s(bool,X3) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) ) | $spl12),
  inference(subsumption_resolution,[],[f8021,f71])).
fof(f8021,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X4),s(bool,X3),s(bool,X5))) | s(bool,'F_0') = s(bool,X4) | s(bool,X3) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) ) | $spl12),
  inference(superposition,[],[f5290,f131])).
fof(f8062,plain,(
  ( ! [X26,X24,X23,X25,X22] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X23),s(bool,X22),s(bool,X24))) | s(bool,'F_0') = s(bool,X23) | s(bool,X22) = s(bool,X25) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X25),s(bool,X26)))) ) | $spl12),
  inference(subsumption_resolution,[],[f8026,f71])).
fof(f8026,plain,(
  ( ! [X26,X24,X23,X25,X22] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X23),s(bool,X22),s(bool,X24))) | s(bool,'F_0') = s(bool,X23) | s(bool,X22) = s(bool,X25) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X25),s(bool,X26)))) ) | $spl12),
  inference(superposition,[],[f5290,f326])).
fof(f8068,plain,(
  ( ! [X54,X52,X56,X55,X53] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X53),s(bool,X52),s(bool,X54))) | s(bool,'F_0') = s(bool,X53) | s(bool,X55) = s(bool,X56) | s(bool,'F_0') = s(bool,X56) | s(bool,X52) = s(bool,X55)) ) | $spl12),
  inference(subsumption_resolution,[],[f8032,f71])).
fof(f8032,plain,(
  ( ! [X54,X52,X56,X55,X53] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X53),s(bool,X52),s(bool,X54))) | s(bool,'F_0') = s(bool,X53) | s(bool,X55) = s(bool,X56) | s(bool,'F_0') = s(bool,X56) | s(bool,X52) = s(bool,X55)) ) | $spl12),
  inference(superposition,[],[f5290,f1664])).
fof(f8135,plain,(
  ( ! [X39,X40] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X39))))) | s(bool,X40) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X39))) | s(bool,'F_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f8116,f73])).
fof(f8116,plain,(
  ( ! [X39,X40] : (p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X39))))) | s(bool,X40) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X39))) | s(bool,'F_0') = s(bool,X40)) )),
  inference(superposition,[],[f258,f3059])).
fof(f8341,plain,(
  ( ! [X22] : (s(bool,X22) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))))) | s(bool,'F_0') = s(bool,X22)) )),
  inference(backward_demodulation,[],[f8314,f8264])).
fof(f8568,plain,(
  ( ! [X103,X102,X104] : (s(bool,'F_0') != s(bool,X102) | s(bool,X104) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X102),s(bool,X103))) | s(bool,'T_0') = s(bool,X104)) )),
  inference(superposition,[],[f1663,f102])).
fof(f8569,plain,(
  ( ! [X107,X105,X106] : (s(bool,'F_0') != s(bool,X106) | s(bool,X107) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X105),s(bool,X106))) | s(bool,'T_0') = s(bool,X107)) )),
  inference(superposition,[],[f1663,f103])).
fof(f8570,plain,(
  ( ! [X111,X109,X110,X108] : (s(bool,'F_0') != s(bool,X110) | s(bool,X111) = s(bool,'const.bool.COND_3'(s(bool,X108),s(bool,X109),s(bool,X110))) | s(bool,'T_0') = s(bool,X111) | s(bool,'T_0') = s(bool,X108)) )),
  inference(superposition,[],[f1663,f207])).
fof(f8571,plain,(
  ( ! [X114,X112,X115,X113] : (s(bool,'F_0') != s(bool,X113) | s(bool,X115) = s(bool,'const.bool.COND_3'(s(bool,X112),s(bool,X113),s(bool,X114))) | s(bool,'T_0') = s(bool,X115) | s(bool,'F_0') = s(bool,X112)) ) | $spl12),
  inference(superposition,[],[f1663,f3095])).
fof(f8724,plain,(
  ( ! [X83,X81,X82] : (s(bool,X81) != s(bool,'const.bool.COND_3'(s(bool,X82),s(bool,'T_0'),s(bool,X81))) | s(bool,'F_0') = s(bool,X82) | s(bool,'F_0') = s(bool,X83) | s(bool,X81) = s(bool,X83)) ) | $spl12),
  inference(superposition,[],[f8080,f3059])).
fof(f8726,plain,(
  ( ! [X90,X88,X87,X89] : (s(bool,X87) != s(bool,'const.bool.COND_3'(s(bool,X88),s(bool,'T_0'),s(bool,X87))) | s(bool,'F_0') = s(bool,X88) | s(bool,X87) = s(bool,X89) | s(bool,X89) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) ) | $spl12),
  inference(superposition,[],[f8080,f1664])).
fof(f8727,plain,(
  ( ! [X92,X91] : (s(bool,X91) != s(bool,'const.bool.COND_3'(s(bool,X92),s(bool,'T_0'),s(bool,X91))) | s(bool,'F_0') = s(bool,X92) | s(bool,'T_0') = s(bool,X91)) ) | $spl12),
  inference(superposition,[],[f8080,f76])).
fof(f8841,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X2))) | s(bool,'F_0') = s(bool,X1) | s(bool,X3) = s(bool,X4) | s(bool,'F_0') = s(bool,X4) | s(bool,X0) = s(bool,X3)) )),
  inference(superposition,[],[f1944,f1664])).
fof(f8842,plain,(
  ( ! [X6,X8,X7,X5] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X5),s(bool,X6),s(bool,X7))) | s(bool,'F_0') = s(bool,X6) | s(bool,X5) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f1944,f131])).
fof(f8844,plain,(
  ( ! [X14,X17,X15,X13,X16] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X13),s(bool,X14),s(bool,X15))) | s(bool,'F_0') = s(bool,X14) | s(bool,X13) = s(bool,X16) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X16),s(bool,X17)))) )),
  inference(superposition,[],[f1944,f326])).
fof(f8923,plain,(
  ( ! [X83,X81,X82] : (s(bool,X81) != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X82),s(bool,X81))) | s(bool,'F_0') = s(bool,X82) | s(bool,'F_0') = s(bool,X83) | s(bool,X81) = s(bool,X83)) )),
  inference(superposition,[],[f1944,f3059])).
fof(f8925,plain,(
  ( ! [X90,X88,X87,X89] : (s(bool,X87) != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X88),s(bool,X87))) | s(bool,'F_0') = s(bool,X88) | s(bool,X87) = s(bool,X89) | s(bool,X89) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f1944,f1664])).
fof(f8926,plain,(
  ( ! [X92,X91] : (s(bool,X91) != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X92),s(bool,X91))) | s(bool,'F_0') = s(bool,X92) | s(bool,'T_0') = s(bool,X91)) )),
  inference(superposition,[],[f1944,f76])).
fof(f9082,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X1),s(bool,X0),s(bool,X2))) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl12),
  inference(subsumption_resolution,[],[f9042,f73])).
fof(f9042,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X1),s(bool,X0),s(bool,X2))) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl12),
  inference(superposition,[],[f5291,f76])).
fof(f9107,plain,(
  ( ! [X130,X128,X129] : (s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X129),s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X128))),s(bool,X130))) | s(bool,'F_0') = s(bool,X129) | s('type.realax.real',X128) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X128)))) ) | $spl12),
  inference(subsumption_resolution,[],[f9073,f73])).
fof(f9073,plain,(
  ( ! [X130,X128,X129] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X129),s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X128))),s(bool,X130))) | s(bool,'F_0') = s(bool,X129) | s('type.realax.real',X128) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X128)))) ) | $spl12),
  inference(superposition,[],[f5291,f5279])).
fof(f9375,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X2))) | s(bool,'F_0') = s(bool,X2) | s(bool,'F_0') = s(bool,X3) | s(bool,X0) = s(bool,X3)) )),
  inference(superposition,[],[f7097,f3059])).
fof(f9377,plain,(
  ( ! [X12,X10,X8,X11,X9] : (s(bool,X8) != s(bool,'const.bool.COND_3'(s(bool,X8),s(bool,X9),s(bool,X10))) | s(bool,'F_0') = s(bool,X10) | s(bool,X8) = s(bool,X11) | s(bool,X11) = s(bool,X12) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f7097,f1664])).
fof(f9378,plain,(
  ( ! [X14,X15,X13] : (s(bool,X13) != s(bool,'const.bool.COND_3'(s(bool,X13),s(bool,X14),s(bool,X15))) | s(bool,'F_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f7097,f76])).
fof(f9474,plain,(
  ( ! [X111,X112,X113] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X113),s(bool,X112))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X111),s(bool,X112)))) )),
  inference(superposition,[],[f7097,f103])).
fof(f9475,plain,(
  ( ! [X116,X114,X117,X115] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X117),s(bool,X116))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X114),s(bool,X115),s(bool,X116))) | s(bool,'T_0') = s(bool,X114)) )),
  inference(superposition,[],[f7097,f207])).
fof(f9476,plain,(
  ( ! [X121,X118,X120,X119] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X121),s(bool,X119))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X118),s(bool,X119),s(bool,X120))) | s(bool,'F_0') = s(bool,X118)) ) | $spl12),
  inference(superposition,[],[f7097,f3095])).
fof(f9546,plain,(
  ( ! [X37,X35,X36,X34] : (s(bool,X34) != s(bool,'const.bool.COND_3'(s(bool,X34),s(bool,X35),s(bool,X36))) | s(bool,'F_0') = s(bool,X36) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X34),s(bool,X37)))) )),
  inference(subsumption_resolution,[],[f9385,f110])).
fof(f9385,plain,(
  ( ! [X37,X35,X36,X34] : (s(bool,X34) != s(bool,'const.bool.COND_3'(s(bool,X34),s(bool,X35),s(bool,X36))) | s(bool,'F_0') = s(bool,X36) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X34),s(bool,X37))) | s(bool,'T_0') = s(bool,'F_0')) )),
  inference(superposition,[],[f7097,f326])).
fof(f9562,plain,(
  ( ! [X70,X68,X69,X67] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X68),s(bool,X67))) | s(bool,X69) = s(bool,X70) | s(bool,'F_0') = s(bool,X70) | s(bool,X67) = s(bool,X69)) )),
  inference(subsumption_resolution,[],[f9457,f110])).
fof(f9457,plain,(
  ( ! [X70,X68,X69,X67] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X68),s(bool,X67))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X69) = s(bool,X70) | s(bool,'F_0') = s(bool,X70) | s(bool,X67) = s(bool,X69)) )),
  inference(superposition,[],[f7097,f1664])).
fof(f9563,plain,(
  ( ! [X80,X78,X79,X77] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X78),s(bool,X77))) | s(bool,X77) = s(bool,X79) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X79),s(bool,X80)))) )),
  inference(subsumption_resolution,[],[f9460,f110])).
fof(f9460,plain,(
  ( ! [X80,X78,X79,X77] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X78),s(bool,X77))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X77) = s(bool,X79) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X79),s(bool,X80)))) )),
  inference(superposition,[],[f7097,f326])).
fof(f11192,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X4),s(bool,X5),s(bool,X3))) | s(bool,'T_0') = s(bool,X4) | s(bool,X3) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f11149,f71])).
fof(f11149,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X4),s(bool,X5),s(bool,X3))) | s(bool,'T_0') = s(bool,X4) | s(bool,X3) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f6545,f131])).
fof(f11194,plain,(
  ( ! [X26,X24,X23,X25,X22] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X23),s(bool,X24),s(bool,X22))) | s(bool,'T_0') = s(bool,X23) | s(bool,X22) = s(bool,X25) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X25),s(bool,X26)))) )),
  inference(subsumption_resolution,[],[f11154,f71])).
fof(f11154,plain,(
  ( ! [X26,X24,X23,X25,X22] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X23),s(bool,X24),s(bool,X22))) | s(bool,'T_0') = s(bool,X23) | s(bool,X22) = s(bool,X25) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X25),s(bool,X26)))) )),
  inference(superposition,[],[f6545,f326])).
fof(f11199,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X2))) | s(bool,'F_0') = s(bool,X2)) )),
  inference(duplicate_literal_removal,[],[f11198])).
fof(f11198,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X2))) | s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) )),
  inference(condensation,[],[f11197])).
fof(f11197,plain,(
  ( ! [X45,X43,X46,X44,X42] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X44),s(bool,X45),s(bool,X42))) | s(bool,'T_0') = s(bool,X44) | s(bool,'F_0') = s(bool,X42) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f11158,f162])).
fof(f11158,plain,(
  ( ! [X45,X43,X46,X44,X42] : (~p(s(bool,X43)) | s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X44),s(bool,X45),s(bool,X42))) | s(bool,'T_0') = s(bool,X44) | s(bool,'F_0') = s(bool,X42) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f6545,f1664])).
fof(f11200,plain,(
  ( ! [X54,X52,X56,X55,X53] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X53),s(bool,X54),s(bool,X52))) | s(bool,'T_0') = s(bool,X53) | s(bool,X55) = s(bool,X56) | s(bool,'F_0') = s(bool,X56) | s(bool,X52) = s(bool,X55)) )),
  inference(subsumption_resolution,[],[f11160,f71])).
fof(f11160,plain,(
  ( ! [X54,X52,X56,X55,X53] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X53),s(bool,X54),s(bool,X52))) | s(bool,'T_0') = s(bool,X53) | s(bool,X55) = s(bool,X56) | s(bool,'F_0') = s(bool,X56) | s(bool,X52) = s(bool,X55)) )),
  inference(superposition,[],[f6545,f1664])).
fof(f11209,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,'T_0')))) )),
  inference(duplicate_literal_removal,[],[f11208])).
fof(f11208,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(condensation,[],[f11207])).
fof(f11207,plain,(
  ( ! [X94,X92,X95,X93] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X93),s(bool,X94),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X93) | s(bool,X92) = s(bool,X95) | s(bool,'T_0') = s(bool,X95)) )),
  inference(subsumption_resolution,[],[f11169,f162])).
fof(f11169,plain,(
  ( ! [X94,X92,X95,X93] : (~p(s(bool,X92)) | s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X93),s(bool,X94),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X93) | s(bool,X92) = s(bool,X95) | s(bool,'T_0') = s(bool,X95)) )),
  inference(superposition,[],[f6545,f131])).
fof(f11361,plain,(
  ( ! [X152,X151,X153,X150] : (s(bool,X151) != s(bool,X152) | s(bool,'T_0') = s(bool,X153) | s(bool,X152) = s(bool,X153) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X150),s(bool,X151)))) )),
  inference(superposition,[],[f1783,f103])).
fof(f11363,plain,(
  ( ! [X163,X161,X159,X162,X160] : (s(bool,X160) != s(bool,X162) | s(bool,'T_0') = s(bool,X163) | s(bool,X162) = s(bool,X163) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X159),s(bool,X160),s(bool,X161))) | s(bool,'F_0') = s(bool,X159)) ) | $spl12),
  inference(superposition,[],[f1783,f3095])).
fof(f11402,plain,(
  ( ! [X158,X156,X159,X157] : (s(bool,X157) != s(bool,X158) | s(bool,'T_0') = s(bool,X159) | s(bool,X159) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X156),s(bool,X157))) | s(bool,'T_0') = s(bool,X158)) )),
  inference(superposition,[],[f1783,f103])).
fof(f11404,plain,(
  ( ! [X167,X165,X169,X166,X168] : (s(bool,X166) != s(bool,X168) | s(bool,'T_0') = s(bool,X169) | s(bool,X169) = s(bool,'const.bool.COND_3'(s(bool,X165),s(bool,X166),s(bool,X167))) | s(bool,'T_0') = s(bool,X168) | s(bool,'F_0') = s(bool,X165)) ) | $spl12),
  inference(superposition,[],[f1783,f3095])).
fof(f11555,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X1),s(bool,X2),s(bool,X0))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f11513,f73])).
fof(f11513,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X1),s(bool,X2),s(bool,X0))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f6546,f76])).
fof(f11581,plain,(
  ( ! [X130,X128,X129] : (s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X129),s(bool,X130),s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X128))))) | s(bool,'T_0') = s(bool,X129) | s('type.realax.real',X128) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X128)))) ) | $spl12),
  inference(subsumption_resolution,[],[f11544,f73])).
fof(f11544,plain,(
  ( ! [X130,X128,X129] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X129),s(bool,X130),s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X128))))) | s(bool,'T_0') = s(bool,X129) | s('type.realax.real',X128) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X128)))) ) | $spl12),
  inference(superposition,[],[f6546,f5279])).
fof(f11650,plain,(
  ( ! [X158,X156,X159,X157] : (s(bool,'T_0') != s(bool,X157) | s(bool,X158) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X156),s(bool,X157))) | s(bool,X158) = s(bool,X159) | s(bool,'T_0') = s(bool,X159)) )),
  inference(superposition,[],[f2285,f103])).
fof(f11652,plain,(
  ( ! [X167,X165,X169,X166,X168] : (s(bool,'T_0') != s(bool,X166) | s(bool,X168) = s(bool,'const.bool.COND_3'(s(bool,X165),s(bool,X166),s(bool,X167))) | s(bool,X168) = s(bool,X169) | s(bool,'T_0') = s(bool,X169) | s(bool,'F_0') = s(bool,X165)) ) | $spl12),
  inference(superposition,[],[f2285,f3095])).
fof(f11777,plain,(
  ( ! [X140,X138,X141,X139] : (s(bool,'F_0') != s(bool,X138) | s(bool,X140) = s(bool,X141) | s(bool,'F_0') = s(bool,X141) | s(bool,X140) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X138),s(bool,X139)))) )),
  inference(superposition,[],[f3200,f102])).
fof(f11778,plain,(
  ( ! [X144,X142,X145,X143] : (s(bool,'F_0') != s(bool,X143) | s(bool,X144) = s(bool,X145) | s(bool,'F_0') = s(bool,X145) | s(bool,X144) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X142),s(bool,X143)))) )),
  inference(superposition,[],[f3200,f103])).
fof(f11779,plain,(
  ( ! [X146,X149,X147,X150,X148] : (s(bool,'F_0') != s(bool,X148) | s(bool,X149) = s(bool,X150) | s(bool,'F_0') = s(bool,X150) | s(bool,X149) = s(bool,'const.bool.COND_3'(s(bool,X146),s(bool,X147),s(bool,X148))) | s(bool,'T_0') = s(bool,X146)) )),
  inference(superposition,[],[f3200,f207])).
fof(f11780,plain,(
  ( ! [X154,X152,X151,X155,X153] : (s(bool,'F_0') != s(bool,X152) | s(bool,X154) = s(bool,X155) | s(bool,'F_0') = s(bool,X155) | s(bool,X154) = s(bool,'const.bool.COND_3'(s(bool,X151),s(bool,X152),s(bool,X153))) | s(bool,'F_0') = s(bool,X151)) ) | $spl12),
  inference(superposition,[],[f3200,f3095])).
fof(f11995,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X9))) | s('type.realax.real','const.real.abs_1'(s('type.realax.real',X9))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X9)))) )),
  inference(subsumption_resolution,[],[f11873,f71])).
fof(f11873,plain,(
  ( ! [X9] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X9))) | s('type.realax.real','const.real.abs_1'(s('type.realax.real',X9))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X9)))) )),
  inference(superposition,[],[f83,f6529])).
fof(f11998,plain,(
  ( ! [X54,X53] : (s(bool,X54) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X53))) | s(bool,'F_0') = s(bool,X54) | s('type.realax.real','const.real.abs_1'(s('type.realax.real',X53))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X53)))) )),
  inference(subsumption_resolution,[],[f11892,f71])).
fof(f11892,plain,(
  ( ! [X54,X53] : (~p(s(bool,'T_0')) | s(bool,X54) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X53))) | s(bool,'F_0') = s(bool,X54) | s('type.realax.real','const.real.abs_1'(s('type.realax.real',X53))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X53)))) )),
  inference(superposition,[],[f453,f6529])).
fof(f12008,plain,(
  ( ! [X132,X130,X131] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X131),s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X130))),s(bool,X132))) | s(bool,'F_0') = s(bool,X131) | s('type.realax.real','const.real.abs_1'(s('type.realax.real',X130))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X130)))) ) | $spl12),
  inference(subsumption_resolution,[],[f11923,f71])).
fof(f11923,plain,(
  ( ! [X132,X130,X131] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X131),s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X130))),s(bool,X132))) | s(bool,'F_0') = s(bool,X131) | s('type.realax.real','const.real.abs_1'(s('type.realax.real',X130))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X130)))) ) | $spl12),
  inference(superposition,[],[f5289,f6529])).
fof(f12009,plain,(
  ( ! [X134,X135,X133] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X134),s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X133))),s(bool,X135))) | s(bool,'F_0') = s(bool,X134) | s('type.realax.real','const.real.abs_1'(s('type.realax.real',X133))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X133)))) ) | $spl12),
  inference(subsumption_resolution,[],[f11924,f71])).
fof(f11924,plain,(
  ( ! [X134,X135,X133] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X134),s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X133))),s(bool,X135))) | s(bool,'F_0') = s(bool,X134) | s('type.realax.real','const.real.abs_1'(s('type.realax.real',X133))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X133)))) ) | $spl12),
  inference(superposition,[],[f5290,f6529])).
fof(f12011,plain,(
  ( ! [X146,X147,X145] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X146),s(bool,X147),s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X145))))) | s(bool,'T_0') = s(bool,X146) | s('type.realax.real','const.real.abs_1'(s('type.realax.real',X145))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X145)))) )),
  inference(subsumption_resolution,[],[f11928,f71])).
fof(f11928,plain,(
  ( ! [X146,X147,X145] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X146),s(bool,X147),s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X145))))) | s(bool,'T_0') = s(bool,X146) | s('type.realax.real','const.real.abs_1'(s('type.realax.real',X145))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X145)))) )),
  inference(superposition,[],[f6545,f6529])).
fof(f12041,plain,(
  ( ! [X6] : (s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',X6))))) | s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X6))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X6))))) | s('type.realax.real',X6) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X6)))) ) | $spl12),
  inference(superposition,[],[f5279,f11987])).
fof(f12321,plain,(
  ( ! [X177,X175,X178,X176,X179] : (s(bool,X177) = s(bool,X178) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X175),s(bool,X176),s(bool,X177))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X179),s(bool,X178))) | s(bool,'T_0') = s(bool,X175)) )),
  inference(superposition,[],[f1050,f207])).
fof(f12322,plain,(
  ( ! [X182,X180,X184,X183,X181] : (s(bool,X181) = s(bool,X183) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X180),s(bool,X181),s(bool,X182))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X184),s(bool,X183))) | s(bool,'F_0') = s(bool,X180)) ) | $spl12),
  inference(superposition,[],[f1050,f3095])).
fof(f12323,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,X2))) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X5),s(bool,X4)))) )),
  inference(superposition,[],[f1050,f1050])).
fof(f12324,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (s(bool,X6) = s(bool,X7) | s(bool,X7) = s(bool,X8) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X9),s(bool,X8))) | s(bool,X10) = s(bool,X11) | s(bool,'F_0') = s(bool,X11) | s(bool,X6) = s(bool,X10)) )),
  inference(superposition,[],[f1050,f1664])).
fof(f12327,plain,(
  ( ! [X26,X24,X23,X27,X25,X22] : (s(bool,X22) = s(bool,X23) | s(bool,X23) = s(bool,X24) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X25),s(bool,X24))) | s(bool,X22) = s(bool,X26) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X26),s(bool,X27)))) )),
  inference(superposition,[],[f1050,f326])).
fof(f12371,plain,(
  ( ! [X130,X128,X131,X129] : (s(bool,'T_0') = s(bool,X128) | s(bool,X130) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X128),s(bool,X129))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X131),s(bool,X130)))) )),
  inference(superposition,[],[f1050,f102])).
fof(f12372,plain,(
  ( ! [X134,X132,X135,X133] : (s(bool,'T_0') = s(bool,X133) | s(bool,X134) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X132),s(bool,X133))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X135),s(bool,X134)))) )),
  inference(superposition,[],[f1050,f103])).
fof(f12373,plain,(
  ( ! [X140,X138,X136,X139,X137] : (s(bool,'T_0') = s(bool,X138) | s(bool,X139) = s(bool,'const.bool.COND_3'(s(bool,X136),s(bool,X137),s(bool,X138))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X140),s(bool,X139))) | s(bool,'T_0') = s(bool,X136)) )),
  inference(superposition,[],[f1050,f207])).
fof(f12374,plain,(
  ( ! [X144,X142,X145,X143,X141] : (s(bool,'T_0') = s(bool,X142) | s(bool,X144) = s(bool,'const.bool.COND_3'(s(bool,X141),s(bool,X142),s(bool,X143))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X145),s(bool,X144))) | s(bool,'F_0') = s(bool,X141)) ) | $spl12),
  inference(superposition,[],[f1050,f3095])).
fof(f12375,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X2))) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X3) | s(bool,'F_0') = s(bool,X4) | s(bool,X0) = s(bool,X4)) )),
  inference(superposition,[],[f1050,f3059])).
fof(f12423,plain,(
  ( ! [X103,X105,X102,X106,X104] : (s(bool,X102) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X102),s(bool,X103))) | s(bool,X103) = s(bool,X104) | s(bool,'T_0') = s(bool,X104) | s(bool,X102) = s(bool,X105) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X106),s(bool,X105)))) )),
  inference(superposition,[],[f1050,f1050])).
fof(f12424,plain,(
  ( ! [X111,X109,X107,X110,X108] : (s(bool,X107) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X107),s(bool,X108))) | s(bool,X108) = s(bool,X109) | s(bool,'T_0') = s(bool,X109) | s(bool,X110) = s(bool,X111) | s(bool,'F_0') = s(bool,X111) | s(bool,X107) = s(bool,X110)) )),
  inference(superposition,[],[f1050,f1664])).
fof(f12427,plain,(
  ( ! [X123,X121,X124,X122,X120] : (s(bool,X120) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X120),s(bool,X121))) | s(bool,X121) = s(bool,X122) | s(bool,'T_0') = s(bool,X122) | s(bool,X120) = s(bool,X123) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X123),s(bool,X124)))) )),
  inference(superposition,[],[f1050,f326])).
fof(f12490,plain,(
  ( ! [X182,X180,X183,X181] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X182),s(bool,X180))) | s(bool,X183) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X180),s(bool,X181))) | s(bool,'T_0') = s(bool,X183)) )),
  inference(superposition,[],[f1050,f102])).
fof(f12491,plain,(
  ( ! [X187,X185,X186,X184] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X186),s(bool,X185))) | s(bool,X187) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X184),s(bool,X185))) | s(bool,'T_0') = s(bool,X187)) )),
  inference(superposition,[],[f1050,f103])).
fof(f12492,plain,(
  ( ! [X191,X189,X192,X190,X188] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X191),s(bool,X190))) | s(bool,X192) = s(bool,'const.bool.COND_3'(s(bool,X188),s(bool,X189),s(bool,X190))) | s(bool,'T_0') = s(bool,X192) | s(bool,'T_0') = s(bool,X188)) )),
  inference(superposition,[],[f1050,f207])).
fof(f12493,plain,(
  ( ! [X196,X194,X197,X195,X193] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X196),s(bool,X194))) | s(bool,X197) = s(bool,'const.bool.COND_3'(s(bool,X193),s(bool,X194),s(bool,X195))) | s(bool,'T_0') = s(bool,X197) | s(bool,'F_0') = s(bool,X193)) ) | $spl12),
  inference(superposition,[],[f1050,f3095])).
fof(f12669,plain,(
  ( ! [X666,X668,X670,X667,X669] : (s(bool,X669) != s(bool,f(s(X666,sK1(X667,X668,bool,X666)))) | s(fun(X666,bool),X667) = s(fun(X666,bool),X668) | s(bool,'T_0') = s(bool,X669) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X670),s(bool,g(s(X666,sK1(X667,X668,bool,X666))))))) )),
  inference(superposition,[],[f108,f1050])).
fof(f12674,plain,(
  ( ! [X693,X689,X691,X692,X690] : (s(X691,f(s(bool,X692))) != s(X691,g(s(bool,X692))) | s(fun(bool,X691),X689) = s(fun(bool,X691),X690) | s(bool,'T_0') = s(bool,X692) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X693),s(bool,sK1(X689,X690,X691,bool))))) )),
  inference(superposition,[],[f108,f1050])).
fof(f12829,plain,(
  ( ! [X648,X650,X646,X649,X647] : (s(bool,X649) != s(bool,f(s(X646,sK1(X647,X648,bool,X646)))) | s(fun(X646,bool),X647) = s(fun(X646,bool),X648) | s(bool,'T_0') = s(bool,g(s(X646,sK1(X647,X648,bool,X646)))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X650),s(bool,X649)))) )),
  inference(superposition,[],[f108,f1050])).
fof(f12834,plain,(
  ( ! [X670,X673,X669,X671,X672] : (s(X671,f(s(bool,X672))) != s(X671,g(s(bool,X672))) | s(fun(bool,X671),X669) = s(fun(bool,X671),X670) | s(bool,'T_0') = s(bool,sK1(X669,X670,X671,bool)) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X673),s(bool,X672)))) )),
  inference(superposition,[],[f108,f1050])).
fof(f12837,plain,(
  ( ! [X12,X10,X8,X7,X11,X9] : (s(X8,X9) = s(X8,'const.bool.COND_3'(s(bool,X7),s(X8,X9),s(X8,X10))) | s(bool,X7) = s(bool,X11) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X12),s(bool,X11)))) )),
  inference(superposition,[],[f102,f1050])).
fof(f12840,plain,(
  ( ! [X28,X26,X29,X27,X25] : (s(bool,X25) = s(bool,'const.bool.COND_3'(s(bool,X25),s(bool,X26),s(bool,X27))) | s(bool,'F_0') = s(bool,X26) | s(bool,X25) = s(bool,X28) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X29),s(bool,X28)))) )),
  inference(superposition,[],[f187,f1050])).
fof(f12841,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (s(bool,X30) = s(bool,'const.bool.COND_3'(s(bool,X30),s(bool,X31),s(bool,X32))) | s(bool,X31) = s(bool,X33) | s(bool,'T_0') = s(bool,X33) | s(bool,X30) = s(bool,X34) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X35),s(bool,X34)))) )),
  inference(superposition,[],[f326,f1050])).
fof(f12843,plain,(
  ( ! [X45,X43,X46,X44,X42] : (s(bool,X42) != s(bool,X43) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X43),s(bool,X44))) | s(bool,X42) = s(bool,X45) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X46),s(bool,X45)))) )),
  inference(superposition,[],[f376,f1050])).
fof(f12844,plain,(
  ( ! [X47,X50,X48,X51,X49] : (s(bool,X48) != s(bool,'const.bool.COND_3'(s(bool,X47),s(bool,X47),s(bool,X49))) | s(bool,'T_0') = s(bool,X48) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X51),s(bool,X50)))) )),
  inference(superposition,[],[f792,f1050])).
fof(f12846,plain,(
  ( ! [X61,X59,X62,X60,X58] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X58),s(bool,X59),s(bool,X60))) | s(bool,'F_0') = s(bool,X59) | s(bool,X58) = s(bool,X61) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X62),s(bool,X61)))) )),
  inference(superposition,[],[f1944,f1050])).
fof(f12847,plain,(
  ( ! [X68,X66,X64,X67,X65,X63] : (s(bool,X63) != s(bool,X64) | s(bool,X64) = s(bool,X65) | s(bool,X65) = s(bool,X66) | s(bool,'T_0') = s(bool,X66) | s(bool,X63) = s(bool,X67) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X68),s(bool,X67)))) )),
  inference(superposition,[],[f2285,f1050])).
fof(f12849,plain,(
  ( ! [X78,X76,X74,X77,X75] : (s(bool,X74) != s(bool,X75) | s(bool,'F_0') = s(bool,X76) | s(bool,X75) = s(bool,X76) | s(bool,X74) = s(bool,X77) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X78),s(bool,X77)))) )),
  inference(superposition,[],[f3180,f1050])).
fof(f12854,plain,(
  ( ! [X103,X105,X102,X106,X104] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X103),s(bool,X102),s(bool,X104))) | s(bool,'F_0') = s(bool,X103) | s(bool,X102) = s(bool,X105) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X106),s(bool,X105)))) ) | $spl12),
  inference(superposition,[],[f8080,f1050])).
fof(f12868,plain,(
  ( ! [X154,X152,X155,X153] : (s(bool,X152) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X152),s(bool,X153))) | s(bool,'F_0') = s(bool,X153) | s(bool,X152) = s(bool,X154) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X155),s(bool,X154)))) )),
  inference(superposition,[],[f219,f1050])).
fof(f13084,plain,(
  ( ! [X633,X635,X636,X634,X637] : (s(bool,'T_0') != s(bool,f(s(X633,sK1(X634,X635,bool,X633)))) | s(fun(X633,bool),X634) = s(fun(X633,bool),X635) | s(bool,X636) = s(bool,g(s(X633,sK1(X634,X635,bool,X633)))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X637),s(bool,X636)))) )),
  inference(superposition,[],[f108,f1050])).
fof(f13089,plain,(
  ( ! [X660,X656,X658,X657,X659] : (s(X658,f(s(bool,'T_0'))) != s(X658,g(s(bool,'T_0'))) | s(fun(bool,X658),X656) = s(fun(bool,X658),X657) | s(bool,X659) = s(bool,sK1(X656,X657,X658,bool)) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X660),s(bool,X659)))) )),
  inference(superposition,[],[f108,f1050])).
fof(f13203,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,X2))) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f1050])).
fof(f13215,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(duplicate_literal_removal,[],[f13197])).
fof(f13197,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f1050])).
fof(f13313,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,X2)))) )),
  inference(condensation,[],[f13312])).
fof(f13312,plain,(
  ( ! [X231,X229,X230,X228,X232] : (s(bool,X228) != s(bool,X229) | s(bool,'T_0') = s(bool,X230) | s(bool,'T_0') = s(bool,X229) | s(bool,X228) = s(bool,X231) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X232),s(bool,X231)))) )),
  inference(duplicate_literal_removal,[],[f12885])).
fof(f12885,plain,(
  ( ! [X231,X229,X230,X228,X232] : (s(bool,X228) != s(bool,X229) | s(bool,'T_0') = s(bool,X230) | s(bool,'T_0') = s(bool,X230) | s(bool,'T_0') = s(bool,X229) | s(bool,X228) = s(bool,X231) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X232),s(bool,X231)))) )),
  inference(superposition,[],[f1783,f1050])).
fof(f13348,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,X1)))) )),
  inference(condensation,[],[f13347])).
fof(f13347,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') != s(bool,X1) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X1)))) )),
  inference(condensation,[],[f13346])).
fof(f13346,plain,(
  ( ! [X4,X2,X3,X1] : (s(bool,X1) = s(bool,X2) | s(bool,'T_0') != s(bool,X3) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X4),s(bool,X3)))) )),
  inference(duplicate_literal_removal,[],[f13345])).
fof(f13345,plain,(
  ( ! [X4,X2,X3,X1] : (s(bool,X1) = s(bool,X2) | s(bool,'T_0') != s(bool,X3) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X4),s(bool,X3)))) )),
  inference(condensation,[],[f12716])).
fof(f12716,plain,(
  ( ! [X177,X175,X178,X176,X179] : (s(bool,'T_0') != s(bool,X176) | s(bool,X175) = s(bool,X177) | s(bool,X177) = s(bool,X178) | s(bool,'T_0') = s(bool,X178) | s(bool,'T_0') = s(bool,X175) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X179),s(bool,X176)))) )),
  inference(superposition,[],[f2285,f1050])).
fof(f13532,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1))))) | s('type.realax.real','const.real.abs_1'(s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1)))) )),
  inference(condensation,[],[f13531])).
fof(f13531,plain,(
  ( ! [X149,X150,X148] : (s(bool,'T_0') = s(bool,X149) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X150),s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X148))))) | s('type.realax.real','const.real.abs_1'(s('type.realax.real',X148))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X148)))) )),
  inference(duplicate_literal_removal,[],[f12264])).
fof(f12264,plain,(
  ( ! [X149,X150,X148] : (s(bool,'T_0') = s(bool,X149) | s(bool,'T_0') = s(bool,X149) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X150),s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X148))))) | s('type.realax.real','const.real.abs_1'(s('type.realax.real',X148))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X148)))) )),
  inference(superposition,[],[f1050,f6529])).
fof(f13550,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,X1))) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,X2)))) )),
  inference(condensation,[],[f13549])).
fof(f13549,plain,(
  ( ! [X47,X50,X48,X51,X49] : (s(bool,'T_0') = s(bool,X48) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X49),s(bool,X47))) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X51),s(bool,X50)))) )),
  inference(duplicate_literal_removal,[],[f12240])).
fof(f12240,plain,(
  ( ! [X47,X50,X48,X51,X49] : (s(bool,'T_0') = s(bool,X48) | s(bool,'T_0') = s(bool,X48) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X49),s(bool,X47))) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X51),s(bool,X50)))) )),
  inference(superposition,[],[f1050,f1050])).
fof(f13591,plain,(
  ( ! [X123,X121,X124,X122,X120] : (s(bool,X120) = s(bool,X121) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X122),s(bool,X121))) | s(bool,X120) = s(bool,X123) | s(bool,X123) = s(bool,X124) | s(bool,'T_0') = s(bool,X124)) )),
  inference(subsumption_resolution,[],[f12306,f110])).
fof(f12306,plain,(
  ( ! [X123,X121,X124,X122,X120] : (s(bool,X120) = s(bool,X121) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X122),s(bool,X121))) | s(bool,X120) = s(bool,X123) | s(bool,X123) = s(bool,X124) | s(bool,'T_0') = s(bool,X124)) )),
  inference(superposition,[],[f1050,f1664])).
fof(f13601,plain,(
  ( ! [X111,X112,X113] : (s(bool,X112) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X111))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X113),s(bool,X112))) | s('type.realax.real',X111) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X111)))) ) | $spl12),
  inference(subsumption_resolution,[],[f12369,f110])).
fof(f12369,plain,(
  ( ! [X111,X112,X113] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X112) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X111))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X113),s(bool,X112))) | s('type.realax.real',X111) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X111)))) ) | $spl12),
  inference(superposition,[],[f1050,f5279])).
fof(f13603,plain,(
  ( ! [X61,X59,X62,X60,X63] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X59),s(bool,X60),s(bool,X61))) | s(bool,X61) = s(bool,X62) | s(bool,'T_0') = s(bool,X62) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X63),s(bool,X59)))) )),
  inference(subsumption_resolution,[],[f12388,f110])).
fof(f12388,plain,(
  ( ! [X61,X59,X62,X60,X63] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X59),s(bool,X60),s(bool,X61))) | s(bool,X61) = s(bool,X62) | s(bool,'T_0') = s(bool,X62) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X63),s(bool,X59)))) )),
  inference(superposition,[],[f1050,f1050])).
fof(f13624,plain,(
  ( ! [X61,X59,X57,X60,X58] : (s(bool,X58) = s(bool,X59) | s(bool,'T_0') = s(bool,X59) | s(bool,X60) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X57),s(bool,X58))) | s(bool,X60) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) )),
  inference(subsumption_resolution,[],[f12504,f110])).
fof(f12504,plain,(
  ( ! [X61,X59,X57,X60,X58] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X58) = s(bool,X59) | s(bool,'T_0') = s(bool,X59) | s(bool,X60) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X57),s(bool,X58))) | s(bool,X60) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) )),
  inference(superposition,[],[f1050,f1664])).
fof(f13775,plain,(
  ( ! [X499,X496,X498,X495,X497] : (s(bool,X495) != s(bool,X496) | s(bool,X497) = s(bool,X498) | s(bool,'F_0') = s(bool,X498) | s(bool,X496) = s(bool,X497) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X499),s(bool,X495)))) )),
  inference(subsumption_resolution,[],[f12793,f110])).
fof(f12793,plain,(
  ( ! [X499,X496,X498,X495,X497] : (s(bool,X495) != s(bool,X496) | s(bool,X497) = s(bool,X498) | s(bool,'F_0') = s(bool,X498) | s(bool,X496) = s(bool,X497) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X499),s(bool,X495)))) )),
  inference(superposition,[],[f3200,f1050])).
fof(f13776,plain,(
  ( ! [X500,X502,X501,X503] : (s(bool,X500) != s(bool,'const.bool.COND_3'(s(bool,X500),s(bool,X501),s(bool,X502))) | s(bool,'F_0') = s(bool,X502) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X503),s(bool,X500)))) )),
  inference(subsumption_resolution,[],[f12794,f110])).
fof(f12794,plain,(
  ( ! [X500,X502,X501,X503] : (s(bool,X500) != s(bool,'const.bool.COND_3'(s(bool,X500),s(bool,X501),s(bool,X502))) | s(bool,'F_0') = s(bool,X502) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X503),s(bool,X500)))) )),
  inference(superposition,[],[f7097,f1050])).
fof(f13808,plain,(
  ( ! [X344,X341,X343,X342] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X342),s(bool,X341))) | s(bool,X341) = s(bool,X343) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X344),s(bool,X343)))) )),
  inference(subsumption_resolution,[],[f12909,f110])).
fof(f12909,plain,(
  ( ! [X344,X341,X343,X342] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X342),s(bool,X341))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X341) = s(bool,X343) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X344),s(bool,X343)))) )),
  inference(superposition,[],[f7097,f1050])).
fof(f13824,plain,(
  ( ! [X10,X11,X9] : (s(bool,'F_0') != s(bool,X9) | s(bool,X9) = s(bool,X10) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X11),s(bool,X10)))) )),
  inference(subsumption_resolution,[],[f12943,f71])).
fof(f12943,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X9) | s(bool,X9) = s(bool,X10) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X11),s(bool,X10)))) )),
  inference(superposition,[],[f83,f1050])).
fof(f13840,plain,(
  ( ! [X269,X265,X267,X268,X266] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X266),s(bool,X267),s(bool,X265))) | s(bool,'T_0') = s(bool,X266) | s(bool,X265) = s(bool,X268) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X269),s(bool,X268)))) )),
  inference(subsumption_resolution,[],[f13000,f71])).
fof(f13000,plain,(
  ( ! [X269,X265,X267,X268,X266] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X266),s(bool,X267),s(bool,X265))) | s(bool,'T_0') = s(bool,X266) | s(bool,X265) = s(bool,X268) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X269),s(bool,X268)))) )),
  inference(superposition,[],[f6545,f1050])).
fof(f648,plain,(
  ( ! [X127,X130,X128,X131,X129] : (s(bool,X130) != s(bool,f(s(X127,sK1(X128,X129,bool,X127)))) | s(fun(X127,bool),X128) = s(fun(X127,bool),X129) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X130),s(bool,X131))) | s(bool,'T_0') = s(bool,g(s(X127,sK1(X128,X129,bool,X127))))) )),
  inference(superposition,[],[f108,f326])).
fof(f687,plain,(
  ( ! [X127,X130,X128,X131,X129] : (s(bool,X130) != s(bool,f(s(X127,sK1(X128,X129,bool,X127)))) | s(fun(X127,bool),X128) = s(fun(X127,bool),X129) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,g(s(X127,sK1(X128,X129,bool,X127)))),s(bool,X131))) | s(bool,'T_0') = s(bool,X130)) )),
  inference(superposition,[],[f108,f326])).
fof(f766,plain,(
  ( ! [X125,X123,X126,X124,X122] : (s(bool,'T_0') != s(bool,f(s(X122,sK1(X123,X124,bool,X122)))) | s(fun(X122,bool),X123) = s(fun(X122,bool),X124) | s(bool,X125) = s(bool,g(s(X122,sK1(X123,X124,bool,X122)))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X125),s(bool,X126)))) )),
  inference(superposition,[],[f108,f326])).
fof(f1165,plain,(
  ( ! [X138,X136,X139,X137] : (s(bool,X139) != s(bool,f(s(X136,sK1(X137,X138,bool,X136)))) | s(fun(X136,bool),X137) = s(fun(X136,bool),X138) | s(bool,'T_0') = s(bool,X139) | s(bool,'T_0') = s(bool,g(s(X136,sK1(X137,X138,bool,X136))))) )),
  inference(superposition,[],[f108,f131])).
fof(f1294,plain,(
  ( ! [X127,X130,X128,X129] : (s(bool,'T_0') != s(bool,f(s(X127,sK1(X128,X129,bool,X127)))) | s(fun(X127,bool),X128) = s(fun(X127,bool),X129) | s(bool,X130) = s(bool,g(s(X127,sK1(X128,X129,bool,X127)))) | s(bool,'T_0') = s(bool,X130)) )),
  inference(superposition,[],[f108,f131])).
fof(f77,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))))) )),
  inference(cnf_transformation,[],[f29])).
fof(f29,plain,(
  ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X6] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X6)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X6)))))))),
  file('thm.real.ABS_POS','thm.real.REAL_LE_NEGTOTAL')).
fof(f258,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))) )),
  inference(subsumption_resolution,[],[f255,f73])).
fof(f255,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))) )),
  inference(superposition,[],[f77,f76])).
fof(f3141,plain,(
  ( ! [X142,X140,X141] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X140)))) | s(bool,X141) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X140))))) | s(bool,X141) = s(bool,X142) | s(bool,'T_0') = s(bool,X142)) )),
  inference(subsumption_resolution,[],[f2353,f73])).
fof(f2353,plain,(
  ( ! [X142,X140,X141] : (p(s(bool,'F_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X140)))) | s(bool,X141) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X140))))) | s(bool,X141) = s(bool,X142) | s(bool,'T_0') = s(bool,X142)) )),
  inference(superposition,[],[f77,f1664])).
fof(f6838,plain,(
  ( ! [X2] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2)))) | s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2)))))) ) | $spl12),
  inference(subsumption_resolution,[],[f6769,f73])).
fof(f6769,plain,(
  ( ! [X2] : (p(s(bool,'F_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2)))) | s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2)))))) ) | $spl12),
  inference(superposition,[],[f77,f5279])).
fof(f6839,plain,(
  ( ! [X10,X9] : (p(s(bool,X10)) | s(bool,X10) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X9))) | s('type.realax.real',X9) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X9)))) ) | $spl12),
  inference(subsumption_resolution,[],[f6775,f73])).
fof(f6775,plain,(
  ( ! [X10,X9] : (p(s(bool,'F_0')) | p(s(bool,X10)) | s(bool,X10) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X9))) | s('type.realax.real',X9) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X9)))) ) | $spl12),
  inference(superposition,[],[f104,f5279])).
fof(f8127,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2)))) )),
  inference(subsumption_resolution,[],[f8101,f73])).
fof(f8101,plain,(
  ( ! [X2] : (p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2)))) )),
  inference(superposition,[],[f258,f76])).
fof(f8132,plain,(
  ( ! [X28,X29,X27] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X27))))) | s(bool,X28) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X27))) | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29)) )),
  inference(subsumption_resolution,[],[f8112,f73])).
fof(f8112,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X27))))) | s(bool,X28) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X27))) | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29)) )),
  inference(superposition,[],[f258,f1664])).
fof(f11973,plain,(
  ( ! [X7] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X7)))) | s('type.realax.real','const.real.abs_1'(s('type.realax.real',X7))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X7)))) )),
  inference(trivial_inequality_removal,[],[f11871])).
fof(f11871,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X7)))) | s('type.realax.real','const.real.abs_1'(s('type.realax.real',X7))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X7)))) )),
  inference(superposition,[],[f79,f6529])).
fof(f11978,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))),s(bool,X1))) | s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))) )),
  inference(condensation,[],[f11977])).
fof(f11977,plain,(
  ( ! [X17,X15,X16] : (s(bool,'T_0') = s(bool,X16) | s('type.realax.real','const.real.abs_1'(s('type.realax.real',X15))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X15))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X15))),s(bool,X17)))) )),
  inference(duplicate_literal_removal,[],[f11851])).
fof(f11851,plain,(
  ( ! [X17,X15,X16] : (s(bool,'T_0') = s(bool,X16) | s('type.realax.real','const.real.abs_1'(s('type.realax.real',X15))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X15))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X15))),s(bool,X17))) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f6529,f326])).
fof(f11996,plain,(
  ( ! [X15,X16] : (~p(s(bool,X16)) | s(bool,X16) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X15))) | s('type.realax.real','const.real.abs_1'(s('type.realax.real',X15))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X15)))) )),
  inference(subsumption_resolution,[],[f11877,f71])).
fof(f11877,plain,(
  ( ! [X15,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X16)) | s(bool,X16) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X15))) | s('type.realax.real','const.real.abs_1'(s('type.realax.real',X15))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X15)))) )),
  inference(superposition,[],[f107,f6529])).
fof(f12010,plain,(
  ( ! [X144,X142,X143] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X143),s(bool,X144),s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X142))))) | s(bool,'T_0') = s(bool,X143) | s('type.realax.real','const.real.abs_1'(s('type.realax.real',X142))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X142)))) )),
  inference(subsumption_resolution,[],[f11927,f71])).
fof(f11927,plain,(
  ( ! [X144,X142,X143] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X143),s(bool,X144),s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X142))))) | s(bool,'T_0') = s(bool,X143) | s('type.realax.real','const.real.abs_1'(s('type.realax.real',X142))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X142)))) )),
  inference(superposition,[],[f6544,f6529])).
fof(f12036,plain,(
  ( ! [X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',X1)))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) | s('type.realax.real',X1) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X1)))) ) | $spl12),
  inference(superposition,[],[f77,f11987])).
fof(f12038,plain,(
  ( ! [X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',X3)))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X3))))))) | s('type.realax.real',X3) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X3)))) ) | $spl12),
  inference(superposition,[],[f258,f11987])).
fof(f12040,plain,(
  ( ! [X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',X5)))))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X5)))))) | s('type.realax.real',X5) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X5)))) ) | $spl12),
  inference(superposition,[],[f77,f11987])).
fof(f690,plain,(
  ( ! [X144,X142,X140,X143,X141] : (s(X142,f(s(bool,X143))) != s(X142,g(s(bool,X143))) | s(fun(bool,X142),X140) = s(fun(bool,X142),X141) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,sK1(X140,X141,X142,bool)),s(bool,X144))) | s(bool,'T_0') = s(bool,X143)) )),
  inference(superposition,[],[f108,f326])).
fof(f12425,plain,(
  ( ! [X114,X112,X115,X113] : (s(bool,X112) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X112),s(bool,X113))) | s(bool,X113) = s(bool,X114) | s(bool,'T_0') = s(bool,X114) | s(bool,X112) = s(bool,X115) | s(bool,'T_0') = s(bool,X115)) )),
  inference(superposition,[],[f1050,f131])).
fof(f13602,plain,(
  ( ! [X47,X45,X48,X46,X44] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X44),s(bool,X45),s(bool,X46))) | s(bool,X46) = s(bool,X47) | s(bool,'T_0') = s(bool,X47) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X44),s(bool,X48)))) )),
  inference(subsumption_resolution,[],[f12385,f110])).
fof(f12385,plain,(
  ( ! [X47,X45,X48,X46,X44] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X44),s(bool,X45),s(bool,X46))) | s(bool,X46) = s(bool,X47) | s(bool,'T_0') = s(bool,X47) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X44),s(bool,X48))) | s(bool,'T_0') = s(bool,'F_0')) )),
  inference(superposition,[],[f1050,f326])).
fof(f12377,plain,(
  ( ! [X14,X12,X10,X15,X13,X11] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X10),s(bool,X11),s(bool,X12))) | s(bool,X12) = s(bool,X13) | s(bool,'T_0') = s(bool,X13) | s(bool,X10) = s(bool,X14) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f1050,f1664])).
fof(f12037,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',X2))))) | s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))) | s('type.realax.real',X2) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X2)))) ) | $spl12),
  inference(superposition,[],[f6529,f11987])).
fof(f12029,plain,(
  ( ! [X1] : (s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',X1))))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',X1))))) | s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))))) | s('type.realax.real',X1) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X1)))) ) | $spl12),
  inference(superposition,[],[f11987,f11987])).
fof(f11989,plain,(
  ( ! [X26,X24,X25] : (s('type.realax.real','const.real.abs_1'(s('type.realax.real',X24))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X24))) | s(bool,X25) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X24))) | s(bool,X25) = s(bool,X26) | s(bool,'T_0') = s(bool,X26)) )),
  inference(subsumption_resolution,[],[f11854,f110])).
fof(f11854,plain,(
  ( ! [X26,X24,X25] : (s(bool,'T_0') = s(bool,'F_0') | s('type.realax.real','const.real.abs_1'(s('type.realax.real',X24))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X24))) | s(bool,X25) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X24))) | s(bool,X25) = s(bool,X26) | s(bool,'T_0') = s(bool,X26)) )),
  inference(superposition,[],[f6529,f1664])).
fof(f11987,plain,(
  ( ! [X2] : (s('type.realax.real','const.real.abs_1'(s('type.realax.real',X2))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))) | s('type.realax.real',X2) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X2)))) ) | $spl12),
  inference(subsumption_resolution,[],[f11845,f110])).
fof(f11845,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,'F_0') | s('type.realax.real','const.real.abs_1'(s('type.realax.real',X2))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))) | s('type.realax.real',X2) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X2)))) ) | $spl12),
  inference(superposition,[],[f6529,f5279])).
fof(f11651,plain,(
  ( ! [X163,X161,X164,X162,X160] : (s(bool,'T_0') != s(bool,X162) | s(bool,X163) = s(bool,'const.bool.COND_3'(s(bool,X160),s(bool,X161),s(bool,X162))) | s(bool,X163) = s(bool,X164) | s(bool,'T_0') = s(bool,X164) | s(bool,'T_0') = s(bool,X160)) )),
  inference(superposition,[],[f2285,f207])).
fof(f11649,plain,(
  ( ! [X154,X152,X155,X153] : (s(bool,'T_0') != s(bool,X152) | s(bool,X154) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X152),s(bool,X153))) | s(bool,X154) = s(bool,X155) | s(bool,'T_0') = s(bool,X155)) )),
  inference(superposition,[],[f2285,f102])).
fof(f11597,plain,(
  ( ! [X21,X19,X17,X20,X18,X16] : (s(bool,X16) != s(bool,X17) | s(bool,X17) = s(bool,X18) | s(bool,X18) = s(bool,X19) | s(bool,'T_0') = s(bool,X19) | s(bool,X16) = s(bool,X20) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X20),s(bool,X21)))) )),
  inference(superposition,[],[f2285,f326])).
fof(f11594,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X3) | s(bool,X4) = s(bool,X5) | s(bool,'F_0') = s(bool,X5) | s(bool,X0) = s(bool,X4)) )),
  inference(superposition,[],[f2285,f1664])).
fof(f11403,plain,(
  ( ! [X163,X161,X164,X162,X160] : (s(bool,X162) != s(bool,X163) | s(bool,'T_0') = s(bool,X164) | s(bool,X164) = s(bool,'const.bool.COND_3'(s(bool,X160),s(bool,X161),s(bool,X162))) | s(bool,'T_0') = s(bool,X163) | s(bool,'T_0') = s(bool,X160)) )),
  inference(superposition,[],[f1783,f207])).
fof(f11401,plain,(
  ( ! [X154,X152,X155,X153] : (s(bool,X152) != s(bool,X154) | s(bool,'T_0') = s(bool,X155) | s(bool,X155) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X152),s(bool,X153))) | s(bool,'T_0') = s(bool,X154)) )),
  inference(superposition,[],[f1783,f102])).
fof(f11362,plain,(
  ( ! [X158,X156,X154,X157,X155] : (s(bool,X156) != s(bool,X157) | s(bool,'T_0') = s(bool,X158) | s(bool,X157) = s(bool,X158) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X154),s(bool,X155),s(bool,X156))) | s(bool,'T_0') = s(bool,X154)) )),
  inference(superposition,[],[f1783,f207])).
fof(f10008,plain,(
  ( ! [X61,X59,X57,X60,X58] : (s(bool,'T_0') = s(bool,X59) | s(bool,X57) = s(bool,X59) | s(bool,X60) = s(bool,'const.bool.COND_3'(s(bool,X57),s(bool,X57),s(bool,X58))) | s(bool,X60) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) ) | $spl12),
  inference(subsumption_resolution,[],[f9907,f73])).
fof(f9907,plain,(
  ( ! [X61,X59,X57,X60,X58] : (p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,X59) | s(bool,X57) = s(bool,X59) | s(bool,X60) = s(bool,'const.bool.COND_3'(s(bool,X57),s(bool,X57),s(bool,X58))) | s(bool,X60) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) ) | $spl12),
  inference(superposition,[],[f7389,f1664])).
fof(f10000,plain,(
  ( ! [X12,X10,X11] : (s(bool,'T_0') = s(bool,X12) | s(bool,X10) = s(bool,X12) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X10),s(bool,X10),s(bool,X11)))) ) | $spl12),
  inference(subsumption_resolution,[],[f9896,f73])).
fof(f9896,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,X12) | s(bool,X10) = s(bool,X12) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X10),s(bool,X10),s(bool,X11)))) ) | $spl12),
  inference(superposition,[],[f7389,f76])).
fof(f8340,plain,(
  ( ! [X10,X11] : (s(bool,X10) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))))) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11)) )),
  inference(backward_demodulation,[],[f8314,f8251])).
fof(f8339,plain,(
  s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))))) | $spl12),
  inference(backward_demodulation,[],[f8314,f8249])).
fof(f8338,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0)))))))),
  inference(backward_demodulation,[],[f8314,f8125])).
fof(f8314,plain,(
  s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0)))))),
  inference(superposition,[],[f456,f103])).
fof(f8294,plain,(
  ( ! [X6,X4,X5] : (s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))) = s('type.realax.real','const.bool.COND_3'(s(bool,X4),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))))) | s(bool,X4) = s(bool,X5) | s(bool,X5) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f456,f1664])).
fof(f8292,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))))) | s(bool,'F_0') = s(bool,X1) | s(bool,X0) = s(bool,X1)) )),
  inference(superposition,[],[f456,f3059])).
fof(f8264,plain,(
  ( ! [X22] : (s(bool,X22) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))))) | s(bool,'F_0') = s(bool,X22)) )),
  inference(subsumption_resolution,[],[f8188,f71])).
fof(f8188,plain,(
  ( ! [X22] : (~p(s(bool,'T_0')) | s(bool,X22) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))))) | s(bool,'F_0') = s(bool,X22)) )),
  inference(superposition,[],[f453,f8125])).
fof(f8251,plain,(
  ( ! [X10,X11] : (s(bool,X10) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))))) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11)) )),
  inference(subsumption_resolution,[],[f8154,f110])).
fof(f8154,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X10) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))))) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f8125,f1664])).
fof(f8249,plain,(
  s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))))) | $spl12),
  inference(subsumption_resolution,[],[f8145,f110])).
fof(f8145,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))))) | $spl12),
  inference(superposition,[],[f8125,f5279])).
fof(f8126,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))))) | s('type.realax.real',X1) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X1)))) ) | $spl12),
  inference(subsumption_resolution,[],[f8100,f73])).
fof(f8100,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))))) | s('type.realax.real',X1) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X1)))) ) | $spl12),
  inference(superposition,[],[f258,f5279])).
fof(f8125,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0)))))))),
  inference(condensation,[],[f8124])).
fof(f8124,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f8098,f80])).
fof(f8098,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f258,f249])).
fof(f6841,plain,(
  ( ! [X61,X62,X60] : (s(bool,X61) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X60))) | s(bool,'T_0') = s(bool,X62) | s(bool,X61) = s(bool,X62) | s('type.realax.real',X60) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X60)))) ) | $spl12),
  inference(subsumption_resolution,[],[f6797,f73])).
fof(f6797,plain,(
  ( ! [X61,X62,X60] : (p(s(bool,'F_0')) | s(bool,X61) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X60))) | s(bool,'T_0') = s(bool,X62) | s(bool,X61) = s(bool,X62) | s('type.realax.real',X60) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X60)))) ) | $spl12),
  inference(superposition,[],[f923,f5279])).
fof(f6565,plain,(
  ( ! [X144,X142,X145,X143] : (s(bool,X144) != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X145))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X142),s(bool,X143),s(bool,X144))) | s(bool,'T_0') = s(bool,X142)) )),
  inference(superposition,[],[f792,f207])).
fof(f6558,plain,(
  ( ! [X111,X109,X112,X110,X113] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X111),s(bool,X112))) | s(bool,X113) = s(bool,'const.bool.COND_3'(s(bool,X109),s(bool,X110),s(bool,X111))) | s(bool,'T_0') = s(bool,X113) | s(bool,'T_0') = s(bool,X109)) )),
  inference(superposition,[],[f326,f207])).
fof(f6529,plain,(
  ( ! [X107] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X107))) | s('type.realax.real','const.real.abs_1'(s('type.realax.real',X107))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X107)))) )),
  inference(superposition,[],[f207,f74])).
fof(f6514,plain,(
  ( ! [X37,X35,X33,X36,X34] : (s(bool,X35) = s(bool,X36) | s(bool,'T_0') = s(bool,X33) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X36),s(bool,X37))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X33),s(bool,X34),s(bool,X35)))) )),
  inference(superposition,[],[f207,f326])).
fof(f6513,plain,(
  ( ! [X30,X28,X31,X29,X32] : (s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,X28) | s(bool,X31) = s(bool,'const.bool.COND_3'(s(bool,X28),s(bool,X29),s(bool,X30))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X31),s(bool,X32)))) )),
  inference(superposition,[],[f207,f326])).
fof(f6512,plain,(
  ( ! [X26,X27,X25] : (s(bool,X27) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))) | s(bool,'T_0') = s(bool,X25) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X25),s(bool,X26),s(bool,X27)))) )),
  inference(superposition,[],[f207,f249])).
fof(f6509,plain,(
  ( ! [X14,X15,X13,X16] : (s(bool,X15) = s(bool,X16) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X16) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X13),s(bool,X14),s(bool,X15)))) )),
  inference(superposition,[],[f207,f131])).
fof(f6508,plain,(
  ( ! [X12,X10,X11,X9] : (s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X9) | s(bool,X12) = s(bool,'const.bool.COND_3'(s(bool,X9),s(bool,X10),s(bool,X11))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f207,f131])).
fof(f6426,plain,(
  ( ! [X177,X180,X178,X176,X181,X179] : (s(X179,X180) = s(X179,'const.bool.COND_3'(s(bool,X178),s(X179,X181),s(X179,X180))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X176),s(bool,X177),s(bool,X178))) | s(bool,'T_0') = s(bool,X176)) )),
  inference(superposition,[],[f207,f207])).
fof(f5044,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X2) | s(bool,X1) = s(bool,X2) | s(bool,X3) = s(bool,X4) | s(bool,'F_0') = s(bool,X4) | s(bool,X0) = s(bool,X3)) )),
  inference(superposition,[],[f3180,f1664])).
fof(f5039,plain,(
  ( ! [X74,X72,X75,X73] : (s(bool,X72) != s(bool,X73) | s(bool,'T_0') = s(bool,X73) | s(bool,X72) = s(bool,X74) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X74),s(bool,X75)))) )),
  inference(subsumption_resolution,[],[f4988,f110])).
fof(f4988,plain,(
  ( ! [X74,X72,X75,X73] : (s(bool,X72) != s(bool,X73) | s(bool,'T_0') = s(bool,X73) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X72) = s(bool,X74) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X74),s(bool,X75)))) )),
  inference(superposition,[],[f3107,f326])).
fof(f5038,plain,(
  ( ! [X64,X62,X65,X63] : (s(bool,X62) != s(bool,X63) | s(bool,'T_0') = s(bool,X63) | s(bool,X64) = s(bool,X65) | s(bool,'F_0') = s(bool,X65) | s(bool,X62) = s(bool,X64)) )),
  inference(subsumption_resolution,[],[f4985,f110])).
fof(f4985,plain,(
  ( ! [X64,X62,X65,X63] : (s(bool,X62) != s(bool,X63) | s(bool,'T_0') = s(bool,X63) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X64) = s(bool,X65) | s(bool,'F_0') = s(bool,X65) | s(bool,X62) = s(bool,X64)) )),
  inference(superposition,[],[f3107,f1664])).
fof(f3771,plain,(
  ( ! [X37,X35,X36,X34] : (s(bool,'F_0') = s(bool,X35) | s(bool,X34) != s(bool,X35) | s(bool,X34) = s(bool,X36) | s(bool,X36) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) ) | $spl10),
  inference(subsumption_resolution,[],[f3753,f73])).
fof(f3753,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,X35) | s(bool,X34) != s(bool,X35) | s(bool,X34) = s(bool,X36) | s(bool,X36) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) ) | $spl10),
  inference(superposition,[],[f3094,f1664])).
fof(f454,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))) = s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)))) )),
  inference(superposition,[],[f74,f263])).
fof(f455,plain,(
  ( ! [X2] : (s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))) = s('type.realax.real','const.bool.COND_3'(s(bool,X2),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f74,f249])).
fof(f111,plain,(
  s(bool,'T_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0)))))),
  inference(unit_resulting_resolution,[],[f70,f79])).
fof(f249,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f244,f73])).
fof(f244,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | s(bool,X0) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f162,f118])).
fof(f263,plain,(
  ( ! [X4,X5] : (s(bool,X4) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X4),s(bool,X5)))) )),
  inference(superposition,[],[f249,f102])).
fof(f714,plain,(
  ( ! [X19,X17,X15,X18,X16] : (s(bool,X15) = s(bool,'const.bool.COND_3'(s(bool,X15),s(bool,X16),s(bool,X17))) | s(bool,X16) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))) | s(bool,X15) = s(bool,X18) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X18),s(bool,X19)))) )),
  inference(superposition,[],[f263,f326])).
fof(f1223,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.bool.COND_3'(s(bool,X12),s(bool,X13),s(bool,X14))) | s(bool,X13) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))) | s(bool,X12) = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f263,f131])).
fof(f1740,plain,(
  ( ! [X74,X72,X71,X73] : (s(bool,X73) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X71),s(bool,X72))) | s(bool,X73) = s(bool,X74) | s(bool,'T_0') = s(bool,X74) | s(bool,X71) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0)))))) )),
  inference(subsumption_resolution,[],[f1704,f71])).
fof(f1704,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,'T_0')) | s(bool,X73) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X71),s(bool,X72))) | s(bool,X73) = s(bool,X74) | s(bool,'T_0') = s(bool,X74) | s(bool,X71) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0)))))) )),
  inference(superposition,[],[f917,f263])).
fof(f1795,plain,(
  ( ! [X64,X65] : (s(bool,X64) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))) | s(bool,'T_0') = s(bool,X65) | s(bool,X64) = s(bool,X65)) )),
  inference(subsumption_resolution,[],[f1772,f73])).
fof(f1772,plain,(
  ( ! [X64,X65] : (p(s(bool,'F_0')) | s(bool,X64) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))) | s(bool,'T_0') = s(bool,X65) | s(bool,X64) = s(bool,X65)) )),
  inference(superposition,[],[f923,f118])).
fof(f938,plain,(
  ( ! [X99,X100,X98] : (s('type.realax.real',X98) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X98))) | s(bool,X99) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X98))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X99),s(bool,X100)))) )),
  inference(forward_demodulation,[],[f761,f102])).
fof(f761,plain,(
  ( ! [X99,X100,X98] : (s('type.realax.real','const.real.abs_1'(s('type.realax.real',X98))) = s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real',X98),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X98))))) | s(bool,X99) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X98))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X99),s(bool,X100)))) )),
  inference(superposition,[],[f74,f326])).
fof(f1665,plain,(
  ( ! [X103,X104] : (s('type.realax.real',X103) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X103))) | s(bool,X104) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X103))) | s(bool,'T_0') = s(bool,X104)) )),
  inference(forward_demodulation,[],[f1288,f102])).
fof(f1288,plain,(
  ( ! [X103,X104] : (s('type.realax.real','const.real.abs_1'(s('type.realax.real',X103))) = s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real',X103),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X103))))) | s(bool,X104) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X103))) | s(bool,'T_0') = s(bool,X104)) )),
  inference(superposition,[],[f74,f131])).
fof(f456,plain,(
  s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))) = s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0)))))))),
  inference(superposition,[],[f74,f118])).
fof(f3202,plain,(
  ( ! [X92,X90,X91,X89] : (s(bool,X89) = s(bool,X90) | s(bool,'F_0') = s(bool,X90) | s(bool,X91) = s(bool,X92) | s(bool,'F_0') = s(bool,X92) | s(bool,X89) = s(bool,X91)) )),
  inference(subsumption_resolution,[],[f2554,f71])).
fof(f2554,plain,(
  ( ! [X92,X90,X91,X89] : (~p(s(bool,'T_0')) | s(bool,X89) = s(bool,X90) | s(bool,'F_0') = s(bool,X90) | s(bool,X91) = s(bool,X92) | s(bool,'F_0') = s(bool,X92) | s(bool,X89) = s(bool,X91)) )),
  inference(superposition,[],[f453,f1664])).
fof(f3200,plain,(
  ( ! [X10,X11,X9] : (s(bool,'F_0') != s(bool,X9) | s(bool,X10) = s(bool,X11) | s(bool,'F_0') = s(bool,X11) | s(bool,X9) = s(bool,X10)) )),
  inference(subsumption_resolution,[],[f2535,f71])).
fof(f2535,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X9) | s(bool,X10) = s(bool,X11) | s(bool,'F_0') = s(bool,X11) | s(bool,X9) = s(bool,X10)) )),
  inference(superposition,[],[f83,f1664])).
fof(f3180,plain,(
  ( ! [X98,X96] : (s(bool,'T_0') != s(bool,X96) | s(bool,'F_0') = s(bool,X98) | s(bool,X96) = s(bool,X98)) )),
  inference(global_subsumption,[],[f70,f71,f73,f74,f76,f77,f80,f79,f84,f83,f102,f103,f107,f104,f108,f111,f110,f118,f131,f162,f187,f207,f219,f249,f258,f263,f264,f325,f326,f416,f417,f376,f450,f453,f454,f455,f456,f468,f469,f493,f522,f523,f524,f525,f542,f558,f856,f575,f576,f598,f611,f612,f891,f642,f648,f651,f681,f687,f690,f915,f917,f921,f923,f713,f714,f715,f934,f937,f938,f766,f769,f792,f780,f996,f957,f958,f1050,f1072,f1074,f1093,f1094,f1157,f1165,f1168,f1222,f1223,f1224,f1226,f1227,f1663,f1664,f1665,f1294,f1297,f1739,f1740,f1783,f1795,f1799,f1801,f1944,f1951,f1952,f2996,f2989,f1993,f1994,f2966,f3031,f2023,f2024,f2025,f2948,f2928,f2061,f2919,f2916,f2091,f2900,f2895,f3048,f2889,f2886,f3059,f2121,f2122,f2125,f2881,f2871,f3075,f2866,f3076,f2155,f2156,f2843,f2842,f2837,f2836,f2835,f2831,f2829,f2207,f3107,f2216,f2219,f2270,f2279,f2282,f2283,f2285,f3133,f2333,f3137,f3141,f2354,f3143,f2363,f2366,f3148,f3163,f3170,f2417,f3172,f2426,f2429,f3174,f3179])).
fof(f3179,plain,(
  ( ! [X97,X95,X98,X96] : (s(bool,'T_0') != s(bool,X96) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X95),s(bool,X97))) | s(bool,X96) = s(bool,X98) | s(bool,'F_0') = s(bool,X98)) )),
  inference(subsumption_resolution,[],[f2454,f376])).
fof(f2454,plain,(
  ( ! [X97,X95,X98,X96] : (s(bool,'T_0') != s(bool,X96) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X95),s(bool,X97))) | s(bool,X96) = s(bool,X98) | s(bool,'F_0') = s(bool,X98) | s(bool,'T_0') = s(bool,X95)) )),
  inference(superposition,[],[f376,f1664])).
fof(f3143,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X2))) | s(bool,X0) = s(bool,X1)) )),
  inference(condensation,[],[f3142])).
fof(f3142,plain,(
  ( ! [X161,X159,X162,X160] : (s(bool,'T_0') = s(bool,X160) | s(bool,X161) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X159))) | s(bool,X161) = s(bool,X162) | s(bool,'T_0') = s(bool,X162)) )),
  inference(subsumption_resolution,[],[f2358,f76])).
fof(f2358,plain,(
  ( ! [X161,X159,X162,X160] : (s(bool,'F_0') != s(bool,X160) | s(bool,'T_0') = s(bool,X160) | s(bool,X161) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X159))) | s(bool,X161) = s(bool,X162) | s(bool,'T_0') = s(bool,X162)) )),
  inference(superposition,[],[f792,f1664])).
fof(f3137,plain,(
  ( ! [X118,X116,X114,X117,X115] : (s(bool,X114) != s(bool,X115) | s(bool,'T_0') = s(bool,X116) | s(bool,X115) = s(bool,X116) | s(bool,X114) = s(bool,X117) | s(bool,X117) = s(bool,X118) | s(bool,'T_0') = s(bool,X118)) )),
  inference(subsumption_resolution,[],[f2345,f73])).
fof(f2345,plain,(
  ( ! [X118,X116,X114,X117,X115] : (p(s(bool,'F_0')) | s(bool,X114) != s(bool,X115) | s(bool,'T_0') = s(bool,X116) | s(bool,X115) = s(bool,X116) | s(bool,X114) = s(bool,X117) | s(bool,X117) = s(bool,X118) | s(bool,'T_0') = s(bool,X118)) )),
  inference(superposition,[],[f923,f1664])).
fof(f3133,plain,(
  ( ! [X14,X15,X13,X16] : (p(s(bool,X14)) | s(bool,X13) = s(bool,X14) | s(bool,X13) = s(bool,X15) | s(bool,X15) = s(bool,X16) | s(bool,'T_0') = s(bool,X16)) )),
  inference(subsumption_resolution,[],[f2321,f73])).
fof(f2321,plain,(
  ( ! [X14,X15,X13,X16] : (p(s(bool,'F_0')) | p(s(bool,X14)) | s(bool,X13) = s(bool,X14) | s(bool,X13) = s(bool,X15) | s(bool,X15) = s(bool,X16) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f104,f1664])).
fof(f2285,plain,(
  ( ! [X12,X10,X11] : (s(bool,'T_0') != s(bool,X10) | s(bool,X10) = s(bool,X11) | s(bool,X11) = s(bool,X12) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f110,f1664])).
fof(f2283,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X1,X2) = s(X1,'const.bool.COND_3'(s(bool,X0),s(X1,X3),s(X1,X2))) | s(bool,X0) = s(bool,X4) | s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f103,f1664])).
fof(f3107,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X1)) )),
  inference(condensation,[],[f3106])).
fof(f3106,plain,(
  ( ! [X212,X210,X211] : (s(bool,X210) != s(bool,X211) | s(bool,'T_0') = s(bool,X211) | s(bool,'F_0') = s(bool,X210) | s(bool,'T_0') = s(bool,X212)) )),
  inference(subsumption_resolution,[],[f2211,f792])).
fof(f2211,plain,(
  ( ! [X212,X210,X211,X209] : (s(bool,X210) != s(bool,X211) | s(bool,'T_0') = s(bool,X211) | s(bool,'F_0') = s(bool,X210) | s(bool,X212) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X209))) | s(bool,'T_0') = s(bool,X212)) )),
  inference(superposition,[],[f792,f1664])).
fof(f2122,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X2) = s(bool,X3) | s(bool,'F_0') = s(bool,X3) | s(bool,X1) = s(bool,X2) | s(bool,X4) = s(bool,X5) | s(bool,'F_0') = s(bool,X5) | s(bool,X0) = s(bool,X4)) )),
  inference(superposition,[],[f1664,f1664])).
fof(f3059,plain,(
  ( ! [X99,X98] : (s(bool,'F_0') = s(bool,X98) | s(bool,X98) = s(bool,X99) | s(bool,'F_0') = s(bool,X99)) )),
  inference(subsumption_resolution,[],[f2116,f111])).
fof(f2116,plain,(
  ( ! [X99,X98] : (s(bool,'F_0') = s(bool,X98) | s(bool,X98) = s(bool,X99) | s(bool,'F_0') = s(bool,X99) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0)))))) )),
  inference(superposition,[],[f1664,f118])).
fof(f2025,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X3) | s(bool,X0) = s(bool,X4) | s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1664,f1664])).
fof(f3031,plain,(
  ( ! [X74,X72,X71,X75,X73] : (s(bool,X71) = s(bool,X72) | s(bool,X72) = s(bool,X73) | s(bool,'T_0') = s(bool,X73) | s(bool,X74) = s(bool,X75) | s(bool,'F_0') = s(bool,X75) | s(bool,X71) = s(bool,X74)) )),
  inference(subsumption_resolution,[],[f2011,f110])).
fof(f2011,plain,(
  ( ! [X74,X72,X71,X75,X73] : (s(bool,X71) = s(bool,X72) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X72) = s(bool,X73) | s(bool,'T_0') = s(bool,X73) | s(bool,X74) = s(bool,X75) | s(bool,'F_0') = s(bool,X75) | s(bool,X71) = s(bool,X74)) )),
  inference(superposition,[],[f1664,f1664])).
fof(f1783,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f1756,f73])).
fof(f1756,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f923,f76])).
fof(f1739,plain,(
  ( ! [X70,X68,X66,X69,X67] : (s(bool,X68) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X66),s(bool,X67))) | s(bool,X68) = s(bool,X69) | s(bool,'T_0') = s(bool,X69) | s(bool,X66) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f1703,f71])).
fof(f1703,plain,(
  ( ! [X70,X68,X66,X69,X67] : (~p(s(bool,'T_0')) | s(bool,X68) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X66),s(bool,X67))) | s(bool,X68) = s(bool,X69) | s(bool,'T_0') = s(bool,X69) | s(bool,X66) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f917,f326])).
fof(f1297,plain,(
  ( ! [X140,X138,X139,X137] : (s(X139,f(s(bool,'T_0'))) != s(X139,g(s(bool,'T_0'))) | s(fun(bool,X139),X137) = s(fun(bool,X139),X138) | s(bool,X140) = s(bool,sK1(X137,X138,X139,bool)) | s(bool,'T_0') = s(bool,X140)) )),
  inference(superposition,[],[f108,f131])).
fof(f1664,plain,(
  ( ! [X54,X55,X53] : (s(bool,X53) = s(bool,X54) | s(bool,'F_0') = s(bool,X54) | s(bool,X53) = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) )),
  inference(subsumption_resolution,[],[f1268,f71])).
fof(f1268,plain,(
  ( ! [X54,X55,X53] : (~p(s(bool,'T_0')) | s(bool,X53) = s(bool,X54) | s(bool,'F_0') = s(bool,X54) | s(bool,X53) = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) )),
  inference(superposition,[],[f453,f131])).
fof(f1663,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f1253,f71])).
fof(f1253,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f83,f131])).
fof(f1227,plain,(
  ( ! [X30,X33,X31,X32] : (s(bool,X31) != s(bool,'const.bool.COND_3'(s(bool,X30),s(bool,X30),s(bool,X32))) | s(bool,'T_0') = s(bool,X31) | s(bool,X30) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f792,f131])).
fof(f1226,plain,(
  ( ! [X28,X26,X29,X27] : (s(bool,X26) != s(bool,X27) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X27),s(bool,X28))) | s(bool,X26) = s(bool,X29) | s(bool,'T_0') = s(bool,X29)) )),
  inference(superposition,[],[f376,f131])).
fof(f1224,plain,(
  ( ! [X19,X17,X20,X18,X16] : (s(bool,X16) = s(bool,'const.bool.COND_3'(s(bool,X16),s(bool,X17),s(bool,X18))) | s(bool,X17) = s(bool,X19) | s(bool,'T_0') = s(bool,X19) | s(bool,X16) = s(bool,X20) | s(bool,'T_0') = s(bool,X20)) )),
  inference(superposition,[],[f326,f131])).
fof(f1222,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(X6,X7) = s(X6,'const.bool.COND_3'(s(bool,X5),s(X6,X7),s(X6,X8))) | s(bool,X5) = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f102,f131])).
fof(f1168,plain,(
  ( ! [X146,X149,X147,X148] : (s(X148,f(s(bool,X149))) != s(X148,g(s(bool,X149))) | s(fun(bool,X148),X146) = s(fun(bool,X148),X147) | s(bool,'T_0') = s(bool,X149) | s(bool,'T_0') = s(bool,sK1(X146,X147,X148,bool))) )),
  inference(superposition,[],[f108,f131])).
fof(f1093,plain,(
  ( ! [X39,X41,X40] : (s(bool,'T_0') = s(bool,X39) | s(bool,X41) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X39),s(bool,X40))) | s(bool,'T_0') = s(bool,X41)) )),
  inference(superposition,[],[f131,f102])).
fof(f1074,plain,(
  ( ! [X12,X10,X8,X11,X9] : (s(bool,X8) = s(bool,X9) | s(bool,X9) = s(bool,X10) | s(bool,'T_0') = s(bool,X10) | s(bool,X8) = s(bool,X11) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X11),s(bool,X12)))) )),
  inference(superposition,[],[f131,f326])).
fof(f1072,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f131,f131])).
fof(f958,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X1) != s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,X0),s(bool,X2))) | s(bool,'T_0') = s(bool,X1) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X3),s(bool,X4)))) )),
  inference(superposition,[],[f792,f326])).
fof(f996,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1))) | s(bool,X0) != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X2)))) )),
  inference(condensation,[],[f995])).
fof(f995,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X9) != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X10))) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X9),s(bool,X11)))) )),
  inference(duplicate_literal_removal,[],[f946])).
fof(f946,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X9) != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X10))) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X9),s(bool,X11))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f792,f326])).
fof(f780,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f326])).
fof(f792,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X1))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(duplicate_literal_removal,[],[f773])).
fof(f773,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X1))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f326])).
fof(f769,plain,(
  ( ! [X138,X136,X135,X139,X137] : (s(X137,f(s(bool,'T_0'))) != s(X137,g(s(bool,'T_0'))) | s(fun(bool,X137),X135) = s(fun(bool,X137),X136) | s(bool,X138) = s(bool,sK1(X135,X136,X137,bool)) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X138),s(bool,X139)))) )),
  inference(superposition,[],[f108,f326])).
fof(f715,plain,(
  ( ! [X24,X23,X21,X22,X20] : (s(bool,X20) != s(bool,X21) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X21),s(bool,X22))) | s(bool,X20) = s(bool,X23) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X23),s(bool,X24)))) )),
  inference(superposition,[],[f376,f326])).
fof(f713,plain,(
  ( ! [X12,X10,X8,X7,X11,X9] : (s(X8,X9) = s(X8,'const.bool.COND_3'(s(bool,X7),s(X8,X9),s(X8,X10))) | s(bool,X7) = s(bool,X11) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X11),s(bool,X12)))) )),
  inference(superposition,[],[f102,f326])).
fof(f923,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X4)) | s(bool,X3) != s(bool,X4) | s(bool,'T_0') = s(bool,X5) | s(bool,X3) = s(bool,X5)) )),
  inference(forward_demodulation,[],[f712,f102])).
fof(f712,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,X3) != s(bool,X4) | p(s(bool,X4)) | s(bool,X3) = s(bool,X5) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X5),s(bool,X6)))) )),
  inference(superposition,[],[f79,f326])).
fof(f917,plain,(
  ( ! [X28,X27,X25] : (~p(s(bool,X27)) | s(bool,X25) = s(bool,X27) | s(bool,X25) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(forward_demodulation,[],[f916,f102])).
fof(f916,plain,(
  ( ! [X28,X26,X27,X25] : (~p(s(bool,X27)) | s(bool,X27) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X25),s(bool,X26))) | s(bool,X25) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f698,f71])).
fof(f698,plain,(
  ( ! [X28,X26,X27,X25] : (~p(s(bool,'T_0')) | ~p(s(bool,X27)) | s(bool,X27) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X25),s(bool,X26))) | s(bool,X25) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f107,f326])).
fof(f651,plain,(
  ( ! [X144,X142,X140,X143,X141] : (s(X142,f(s(bool,X143))) != s(X142,g(s(bool,X143))) | s(fun(bool,X142),X140) = s(fun(bool,X142),X141) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X143),s(bool,X144))) | s(bool,'T_0') = s(bool,sK1(X140,X141,X142,bool))) )),
  inference(superposition,[],[f108,f326])).
fof(f891,plain,(
  ( ! [X78,X76,X74,X77,X75] : (s(X75,X76) = s(X75,'const.bool.COND_3'(s(bool,X74),s(X75,X77),s(X75,X76))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X74),s(bool,X78)))) )),
  inference(subsumption_resolution,[],[f633,f110])).
fof(f633,plain,(
  ( ! [X78,X76,X74,X77,X75] : (s(X75,X76) = s(X75,'const.bool.COND_3'(s(bool,X74),s(X75,X77),s(X75,X76))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X74),s(bool,X78))) | s(bool,'T_0') = s(bool,'F_0')) )),
  inference(superposition,[],[f103,f326])).
fof(f611,plain,(
  ( ! [X39,X41,X38,X40] : (s(bool,'T_0') = s(bool,X38) | s(bool,X40) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X38),s(bool,X39))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X40),s(bool,X41)))) )),
  inference(superposition,[],[f326,f102])).
fof(f598,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X2),s(bool,X3))) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X4),s(bool,X5)))) )),
  inference(superposition,[],[f326,f326])).
fof(f575,plain,(
  ( ! [X54,X52,X53,X51] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X51),s(bool,X53))) | s(bool,X54) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X51),s(bool,X52))) | s(bool,'T_0') = s(bool,X54)) )),
  inference(superposition,[],[f326,f102])).
fof(f856,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1))) | s(bool,X2) = s(bool,'const.bool.COND_3'(s(bool,X2),s(bool,X2),s(bool,X3))) | s(bool,X0) = s(bool,X2)) )),
  inference(condensation,[],[f855])).
fof(f855,plain,(
  ( ! [X28,X26,X29,X27,X25] : (s(bool,X25) = s(bool,'const.bool.COND_3'(s(bool,X25),s(bool,X25),s(bool,X26))) | s(bool,'T_0') = s(bool,X27) | s(bool,X25) = s(bool,X28) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X28),s(bool,X29)))) )),
  inference(duplicate_literal_removal,[],[f568])).
fof(f568,plain,(
  ( ! [X28,X26,X29,X27,X25] : (s(bool,X25) = s(bool,'const.bool.COND_3'(s(bool,X25),s(bool,X25),s(bool,X26))) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,X27) | s(bool,X25) = s(bool,X28) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X28),s(bool,X29)))) )),
  inference(superposition,[],[f326,f326])).
fof(f558,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,X0) = s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X2))) | s(bool,X1) = s(bool,X3) | s(bool,'T_0') = s(bool,X3) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X4),s(bool,X5)))) )),
  inference(superposition,[],[f326,f326])).
fof(f376,plain,(
  ( ! [X6,X5] : (s(bool,'T_0') != s(bool,X5) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X5),s(bool,X6)))) )),
  inference(superposition,[],[f111,f263])).
fof(f326,plain,(
  ( ! [X6,X4,X5] : (s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X5),s(bool,X6))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f263,f249])).
fof(f325,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X3)))) )),
  inference(superposition,[],[f263,f263])).
fof(f207,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.bool.COND_3'(s(bool,X0),s(X1,X3),s(X1,X2))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f103,f76])).
fof(f162,plain,(
  ( ! [X4,X3] : (p(s(bool,X4)) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f156,f73])).
fof(f156,plain,(
  ( ! [X4,X3] : (p(s(bool,'F_0')) | p(s(bool,X4)) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f104,f76])).
fof(f131,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f76,f76])).
fof(f110,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f73,f79])).
fof(f108,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK1(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK1(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f69])).
fof(f69,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK1(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK1(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1])],[f53])).
fof(f53,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f48])).
fof(f48,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f47])).
fof(f47,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.real.ABS_POS',aHL_EXT)).
fof(f104,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f51])).
fof(f51,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f46])).
fof(f46,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X7,X8] : ((p(s(bool,X7)) => p(s(bool,X8))) => ((p(s(bool,X8)) => p(s(bool,X7))) => s(bool,X7) = s(bool,X8)))),
  file('thm.real.ABS_POS','thm.bool.IMP_ANTISYM_AX')).
fof(f107,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f52])).
fof(f102,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f42])).
fof(f42,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X10] : ! [X7,X8] : s(X10,X7) = s(X10,'const.bool.COND_3'(s(bool,'T_0'),s(X10,X7),s(X10,X8)))),
  file('thm.real.ABS_POS','thm.bool.COND_CLAUSES|split|0')).
fof(f79,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X9] : (s(bool,'T_0') = s(bool,X9) <=> p(s(bool,X9)))),
  file('thm.real.ABS_POS','thm.bool.EQ_CLAUSES|split|0')).
fof(f80,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f55])).
fof(f76,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.real.ABS_POS',aHL_BOOL_CASES)).
fof(f71,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'T_0'))),
  file('thm.real.ABS_POS','thm.bool.TRUTH')).
fof(f70,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))))),
  inference(cnf_transformation,[],[f54])).
fof(f54,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f49])).
fof(f49,plain,(
  ? [X0] : ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0))))))),
  inference(ennf_transformation,[],[f24])).
fof(f24,plain,(
  ~! [X0] : p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',X6))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',X6))))))),
  file('thm.real.ABS_POS','thm.real.ABS_POS')).
fof(f21693,plain,(
  $false | ($spl10 | $spl12)),
  inference(global_subsumption,[],[f70,f71,f76,f80,f79,f102,f107,f104,f108,f110,f131,f162,f207,f325,f326,f376,f558,f856,f575,f598,f611,f891,f651,f917,f923,f713,f715,f769,f792,f780,f996,f958,f1072,f1074,f1093,f1168,f1222,f1224,f1226,f1227,f1663,f1664,f1297,f1739,f1783,f2996,f2989,f2966,f3031,f2025,f2948,f2928,f2919,f2916,f2900,f2895,f3048,f2889,f2886,f3059,f2122,f2881,f2871,f3075,f2866,f2843,f2842,f2837,f2836,f3107,f2283,f2285,f3133,f2333,f3137,f3143,f3163,f3172,f3174,f3179,f3180,f2700,f3200,f3202,f456,f1665,f938,f1795,f1740,f1223,f714,f263,f249,f111,f455,f454,f3771,f3772,f5038,f5039,f5044,f6426,f6508,f6509,f6512,f6513,f6514,f6529,f6558,f6565,f6841,f7181,f7195,f7372,f7407,f8066,f8072,f8074,f8125,f8126,f8133,f8249,f8251,f8235,f8264,f8292,f8294,f8314,f8338,f8339,f8340,f8810,f9513,f9961,f9983,f9985,f10000,f10008,f10089,f11088,f11205,f11362,f11401,f11403,f11594,f11597,f11649,f11651,f11987,f11989,f11976,f12029,f12037,f12377,f13602,f12425,f13740,f690,f12040,f12038,f12036,f12010,f11996,f11978,f11973,f8132,f8127,f6839,f6838,f3141,f258,f77,f1294,f1165,f766,f687,f648,f13840,f13824,f13808,f13776,f13775,f13624,f13603,f13601,f13591,f13550,f13532,f13348,f13313,f13215,f13203,f13089,f13084,f12868,f12854,f12849,f12847,f12846,f12844,f12843,f12841,f12840,f12837,f12834,f12829,f12674,f12669,f12493,f12492,f12491,f12490,f12427,f12424,f12423,f12375,f12374,f12373,f12372,f12371,f12327,f12324,f12323,f12322,f12321,f12041,f12011,f12009,f12008,f11998,f11995,f11780,f11779,f11778,f11777,f11652,f11650,f11581,f11555,f11404,f11402,f11363,f11361,f11209,f11200,f11199,f11194,f11192,f9563,f9562,f9546,f9476,f9475,f9474,f9378,f9377,f9375,f9107,f9082,f8926,f8925,f8923,f8844,f8842,f8841,f8727,f8726,f8724,f8571,f8570,f8569,f8568,f8341,f8135,f8068,f8062,f8060,f7411,f7406,f7097,f7090,f6947,f6946,f6945,f6893,f6890,f6850,f6848,f6845,f6693,f6692,f6691,f6586,f6582,f6581,f6580,f6555,f6523,f6522,f6521,f6520,f6519,f6518,f6517,f6516,f6507,f6427,f6425,f6124,f5816,f5324,f5323,f5322,f5307,f5300,f5279,f5275,f5274,f5273,f5272,f5271,f5270,f5269,f5268,f5266,f5265,f5262,f5261,f5260,f5194,f5193,f5091,f5090,f5047,f4996,f4966,f4382,f4379,f4324,f4321,f4245,f4173,f4143,f4099,f3219,f3201,f3170,f3148,f3095,f3094,f3076,f2835,f2831,f2829,f2588,f2585,f2503,f2502,f2500,f2499,f2498,f2495,f2494,f2492,f2489,f2429,f2426,f2366,f2363,f2282,f2279,f2219,f2216,f2156,f2155,f2125,f2121,f2091,f2061,f2024,f1994,f1993,f1952,f1951,f1801,f1799,f1094,f1050,f957,f937,f934,f921,f612,f576,f542,f525,f524,f523,f522,f493,f453,f417,f264,f219,f187,f118,f103,f84,f83,f73,f12035,f12034,f12033,f11842,f11841,f8097,f8096,f6748,f6747,f6589,f5327,f74,f12802,f12642,f5147,f2480,f2417,f2270,f2207,f1157,f681,f642,f6770,f5146,f2354,f469,f468,f11961,f10086,f11962,f12275,f12274,f11553,f11552,f11188,f11187,f11119,f11118,f11010,f11009,f10928,f10927,f10674,f10673,f10120,f10095,f10093,f10090,f10074,f10073,f10018,f10017,f9332,f9314,f9190,f9189,f9150,f9149,f9081,f9080,f8058,f8057,f7988,f7987,f7199,f7180,f7177,f7163,f7162,f7108,f7107,f6583,f6579,f6568,f6567,f6564,f6563,f6554,f6549,f6547,f6546,f6545,f6544,f6543,f6515,f6058,f5321,f5310,f5309,f5306,f5305,f5299,f5294,f5292,f5291,f5290,f5289,f5288,f5267,f13719,f8832,f8681,f8680,f8679,f8080,f7247,f7246,f7176,f9851,f9850,f7389,f7387,f9473,f8898,f8897,f8896,f8895,f8678,f6944,f6690,f5192,f4995,f4172,f2023,f1944,f915,f450,f416,f16317,f16368,f16369,f20133,f20134,f20135,f20136,f20527,f20723,f20686,f20687,f20688,f20850,f20854,f20855,f20892,f20893,f20898,f20899,f20903,f20907,f20913,f20914,f20932,f20918,f20919,f20920,f20922,f20923,f21054,f21053,f21025,f21026,f21061,f21064,f21065,f21066,f21067,f21068,f21076,f21078,f21081,f21114,f21115,f21116,f21117,f21118,f21124,f21602,f21128,f21607,f21623,f21625,f21633,f21638,f21648,f21652,f21664,f21669,f21672,f21674,f21692])).
fof(f21675,plain,(
  $false | ($spl10 | $spl12)),
  inference(global_subsumption,[],[f70,f71,f76,f80,f79,f102,f107,f104,f108,f110,f131,f162,f207,f325,f326,f376,f558,f856,f575,f598,f611,f891,f651,f917,f923,f713,f715,f769,f792,f780,f996,f958,f1072,f1074,f1093,f1168,f1222,f1224,f1226,f1227,f1663,f1664,f1297,f1739,f1783,f2996,f2989,f2966,f3031,f2025,f2948,f2928,f2919,f2916,f2900,f2895,f3048,f2889,f2886,f3059,f2122,f2881,f2871,f3075,f2866,f2843,f2842,f2837,f2836,f3107,f2283,f2285,f3133,f2333,f3137,f3143,f3163,f3172,f3174,f3179,f3180,f2700,f3200,f3202,f456,f1665,f938,f1795,f1740,f1223,f714,f263,f249,f111,f455,f454,f3771,f3772,f5038,f5039,f5044,f6426,f6508,f6509,f6512,f6513,f6514,f6529,f6558,f6565,f6841,f7181,f7195,f7372,f7407,f8066,f8072,f8074,f8125,f8126,f8133,f8249,f8251,f8235,f8264,f8292,f8294,f8314,f8338,f8339,f8340,f8810,f9513,f9961,f9983,f9985,f10000,f10008,f10089,f11088,f11205,f11362,f11401,f11403,f11594,f11597,f11649,f11651,f11987,f11989,f11976,f12029,f12037,f12377,f13602,f12425,f13740,f690,f12040,f12038,f12036,f12010,f11996,f11978,f11973,f8132,f8127,f6839,f6838,f3141,f258,f77,f1294,f1165,f766,f687,f648,f13840,f13824,f13808,f13776,f13775,f13624,f13603,f13601,f13591,f13550,f13532,f13348,f13313,f13215,f13203,f13089,f13084,f12868,f12854,f12849,f12847,f12846,f12844,f12843,f12841,f12840,f12837,f12834,f12829,f12674,f12669,f12493,f12492,f12491,f12490,f12427,f12424,f12423,f12375,f12374,f12373,f12372,f12371,f12327,f12324,f12323,f12322,f12321,f12041,f12011,f12009,f12008,f11998,f11995,f11780,f11779,f11778,f11777,f11652,f11650,f11581,f11555,f11404,f11402,f11363,f11361,f11209,f11200,f11199,f11194,f11192,f9563,f9562,f9546,f9476,f9475,f9474,f9378,f9377,f9375,f9107,f9082,f8926,f8925,f8923,f8844,f8842,f8841,f8727,f8726,f8724,f8571,f8570,f8569,f8568,f8341,f8135,f8068,f8062,f8060,f7411,f7406,f7097,f7090,f6947,f6946,f6945,f6893,f6890,f6850,f6848,f6845,f6693,f6692,f6691,f6586,f6582,f6581,f6580,f6555,f6523,f6522,f6521,f6520,f6519,f6518,f6517,f6516,f6507,f6427,f6425,f6124,f5816,f5324,f5323,f5322,f5307,f5300,f5279,f5275,f5274,f5273,f5272,f5271,f5270,f5269,f5268,f5266,f5265,f5262,f5261,f5260,f5194,f5193,f5091,f5090,f5047,f4996,f4966,f4382,f4379,f4324,f4321,f4245,f4173,f4143,f4099,f3219,f3201,f3170,f3148,f3095,f3094,f3076,f2835,f2831,f2829,f2588,f2585,f2503,f2502,f2500,f2499,f2498,f2495,f2494,f2492,f2489,f2429,f2426,f2366,f2363,f2282,f2279,f2219,f2216,f2156,f2155,f2125,f2121,f2091,f2061,f2024,f1994,f1993,f1952,f1951,f1801,f1799,f1094,f1050,f957,f937,f934,f921,f612,f576,f542,f525,f524,f523,f522,f493,f453,f417,f264,f219,f187,f118,f103,f84,f83,f73,f12035,f12034,f12033,f11842,f11841,f8097,f8096,f6748,f6747,f6589,f5327,f74,f12802,f12642,f5147,f2480,f2417,f2270,f2207,f1157,f681,f642,f6770,f5146,f2354,f469,f468,f11961,f10086,f11962,f12275,f12274,f11553,f11552,f11188,f11187,f11119,f11118,f11010,f11009,f10928,f10927,f10674,f10673,f10120,f10095,f10093,f10090,f10074,f10073,f10018,f10017,f9332,f9314,f9190,f9189,f9150,f9149,f9081,f9080,f8058,f8057,f7988,f7987,f7199,f7180,f7177,f7163,f7162,f7108,f7107,f6583,f6579,f6568,f6567,f6564,f6563,f6554,f6549,f6547,f6546,f6545,f6544,f6543,f6515,f6058,f5321,f5310,f5309,f5306,f5305,f5299,f5294,f5292,f5291,f5290,f5289,f5288,f5267,f13719,f8832,f8681,f8680,f8679,f8080,f7247,f7246,f7176,f9851,f9850,f7389,f7387,f9473,f8898,f8897,f8896,f8895,f8678,f6944,f6690,f5192,f4995,f4172,f2023,f1944,f915,f450,f416,f16317,f16368,f16369,f20133,f20134,f20135,f20136,f20527,f20723,f20686,f20687,f20688,f20850,f20854,f20855,f20892,f20893,f20898,f20899,f20903,f20907,f20913,f20914,f20932,f20918,f20919,f20920,f20922,f20923,f21054,f21053,f21025,f21026,f21061,f21064,f21065,f21066,f21067,f21068,f21076,f21078,f21081,f21114,f21115,f21116,f21117,f21118,f21124,f21602,f21128,f21607,f21623,f21625,f21633,f21638,f21648,f21652,f21664,f21669,f21672,f21674])).
fof(f21673,plain,(
  $false | ($spl10 | $spl12)),
  inference(global_subsumption,[],[f70,f71,f76,f80,f79,f102,f107,f104,f108,f110,f131,f162,f207,f325,f326,f376,f558,f856,f575,f598,f611,f891,f651,f917,f923,f713,f715,f769,f792,f780,f996,f958,f1072,f1074,f1093,f1168,f1222,f1224,f1226,f1227,f1663,f1664,f1297,f1739,f1783,f2996,f2989,f2966,f3031,f2025,f2948,f2928,f2919,f2916,f2900,f2895,f3048,f2889,f2886,f3059,f2122,f2881,f2871,f3075,f2866,f2843,f2842,f2837,f2836,f3107,f2283,f2285,f3133,f2333,f3137,f3143,f3163,f3172,f3174,f3179,f3180,f2700,f3200,f3202,f456,f1665,f938,f1795,f1740,f1223,f714,f263,f249,f111,f455,f454,f3771,f3772,f5038,f5039,f5044,f6426,f6508,f6509,f6512,f6513,f6514,f6529,f6558,f6565,f6841,f7181,f7195,f7372,f7407,f8066,f8072,f8074,f8125,f8126,f8133,f8249,f8251,f8235,f8264,f8292,f8294,f8314,f8338,f8339,f8340,f8810,f9513,f9961,f9983,f9985,f10000,f10008,f10089,f11088,f11205,f11362,f11401,f11403,f11594,f11597,f11649,f11651,f11987,f11989,f11976,f12029,f12037,f12377,f13602,f12425,f13740,f690,f12040,f12038,f12036,f12010,f11996,f11978,f11973,f8132,f8127,f6839,f6838,f3141,f258,f77,f1294,f1165,f766,f687,f648,f13840,f13824,f13808,f13776,f13775,f13624,f13603,f13601,f13591,f13550,f13532,f13348,f13313,f13215,f13203,f13089,f13084,f12868,f12854,f12849,f12847,f12846,f12844,f12843,f12841,f12840,f12837,f12834,f12829,f12674,f12669,f12493,f12492,f12491,f12490,f12427,f12424,f12423,f12375,f12374,f12373,f12372,f12371,f12327,f12324,f12323,f12322,f12321,f12041,f12011,f12009,f12008,f11998,f11995,f11780,f11779,f11778,f11777,f11652,f11650,f11581,f11555,f11404,f11402,f11363,f11361,f11209,f11200,f11199,f11194,f11192,f9563,f9562,f9546,f9476,f9475,f9474,f9378,f9377,f9375,f9107,f9082,f8926,f8925,f8923,f8844,f8842,f8841,f8727,f8726,f8724,f8571,f8570,f8569,f8568,f8341,f8135,f8068,f8062,f8060,f7411,f7406,f7097,f7090,f6947,f6946,f6945,f6893,f6890,f6850,f6848,f6845,f6693,f6692,f6691,f6586,f6582,f6581,f6580,f6555,f6523,f6522,f6521,f6520,f6519,f6518,f6517,f6516,f6507,f6427,f6425,f6124,f5816,f5324,f5323,f5322,f5307,f5300,f5279,f5275,f5274,f5273,f5272,f5271,f5270,f5269,f5268,f5266,f5265,f5262,f5261,f5260,f5194,f5193,f5091,f5090,f5047,f4996,f4966,f4382,f4379,f4324,f4321,f4245,f4173,f4143,f4099,f3219,f3201,f3170,f3148,f3095,f3094,f3076,f2835,f2831,f2829,f2588,f2585,f2503,f2502,f2500,f2499,f2498,f2495,f2494,f2492,f2489,f2429,f2426,f2366,f2363,f2282,f2279,f2219,f2216,f2156,f2155,f2125,f2121,f2091,f2061,f2024,f1994,f1993,f1952,f1951,f1801,f1799,f1094,f1050,f957,f937,f934,f921,f612,f576,f542,f525,f524,f523,f522,f493,f453,f417,f264,f219,f187,f118,f103,f84,f83,f73,f12035,f12034,f12033,f11842,f11841,f8097,f8096,f6748,f6747,f6589,f5327,f74,f12802,f12642,f5147,f2480,f2417,f2270,f2207,f1157,f681,f642,f6770,f5146,f2354,f469,f468,f11961,f10086,f11962,f12275,f12274,f11553,f11552,f11188,f11187,f11119,f11118,f11010,f11009,f10928,f10927,f10674,f10673,f10120,f10095,f10093,f10090,f10074,f10073,f10018,f10017,f9332,f9314,f9190,f9189,f9150,f9149,f9081,f9080,f8058,f8057,f7988,f7987,f7199,f7180,f7177,f7163,f7162,f7108,f7107,f6583,f6579,f6568,f6567,f6564,f6563,f6554,f6549,f6547,f6546,f6545,f6544,f6543,f6515,f6058,f5321,f5310,f5309,f5306,f5305,f5299,f5294,f5292,f5291,f5290,f5289,f5288,f5267,f13719,f8832,f8681,f8680,f8679,f8080,f7247,f7246,f7176,f9851,f9850,f7389,f7387,f9473,f8898,f8897,f8896,f8895,f8678,f6944,f6690,f5192,f4995,f4172,f2023,f1944,f915,f450,f416,f16317,f16368,f16369,f20133,f20134,f20135,f20136,f20527,f20723,f20686,f20687,f20688,f20850,f20854,f20855,f20892,f20893,f20898,f20899,f20903,f20907,f20913,f20914,f20932,f20918,f20919,f20920,f20922,f20923,f21054,f21053,f21025,f21026,f21061,f21064,f21065,f21066,f21067,f21068,f21076,f21078,f21081,f21114,f21115,f21116,f21117,f21118,f21124,f21602,f21128,f21607,f21623,f21625,f21633,f21638,f21648,f21652,f21664,f21669,f21672])).
fof(f21670,plain,(
  $false | ($spl10 | $spl12)),
  inference(global_subsumption,[],[f70,f71,f76,f80,f79,f102,f107,f104,f108,f110,f131,f162,f207,f325,f326,f376,f558,f856,f575,f598,f611,f891,f651,f917,f923,f713,f715,f769,f792,f780,f996,f958,f1072,f1074,f1093,f1168,f1222,f1224,f1226,f1227,f1663,f1664,f1297,f1739,f1783,f2996,f2989,f2966,f3031,f2025,f2948,f2928,f2919,f2916,f2900,f2895,f3048,f2889,f2886,f3059,f2122,f2881,f2871,f3075,f2866,f2843,f2842,f2837,f2836,f3107,f2283,f2285,f3133,f2333,f3137,f3143,f3163,f3172,f3174,f3179,f3180,f2700,f3200,f3202,f456,f1665,f938,f1795,f1740,f1223,f714,f263,f249,f111,f455,f454,f3771,f3772,f5038,f5039,f5044,f6426,f6508,f6509,f6512,f6513,f6514,f6529,f6558,f6565,f6841,f7181,f7195,f7372,f7407,f8066,f8072,f8074,f8125,f8126,f8133,f8249,f8251,f8235,f8264,f8292,f8294,f8314,f8338,f8339,f8340,f8810,f9513,f9961,f9983,f9985,f10000,f10008,f10089,f11088,f11205,f11362,f11401,f11403,f11594,f11597,f11649,f11651,f11987,f11989,f11976,f12029,f12037,f12377,f13602,f12425,f13740,f690,f12040,f12038,f12036,f12010,f11996,f11978,f11973,f8132,f8127,f6839,f6838,f3141,f258,f77,f1294,f1165,f766,f687,f648,f13840,f13824,f13808,f13776,f13775,f13624,f13603,f13601,f13591,f13550,f13532,f13348,f13313,f13215,f13203,f13089,f13084,f12868,f12854,f12849,f12847,f12846,f12844,f12843,f12841,f12840,f12837,f12834,f12829,f12674,f12669,f12493,f12492,f12491,f12490,f12427,f12424,f12423,f12375,f12374,f12373,f12372,f12371,f12327,f12324,f12323,f12322,f12321,f12041,f12011,f12009,f12008,f11998,f11995,f11780,f11779,f11778,f11777,f11652,f11650,f11581,f11555,f11404,f11402,f11363,f11361,f11209,f11200,f11199,f11194,f11192,f9563,f9562,f9546,f9476,f9475,f9474,f9378,f9377,f9375,f9107,f9082,f8926,f8925,f8923,f8844,f8842,f8841,f8727,f8726,f8724,f8571,f8570,f8569,f8568,f8341,f8135,f8068,f8062,f8060,f7411,f7406,f7097,f7090,f6947,f6946,f6945,f6893,f6890,f6850,f6848,f6845,f6693,f6692,f6691,f6586,f6582,f6581,f6580,f6555,f6523,f6522,f6521,f6520,f6519,f6518,f6517,f6516,f6507,f6427,f6425,f6124,f5816,f5324,f5323,f5322,f5307,f5300,f5279,f5275,f5274,f5273,f5272,f5271,f5270,f5269,f5268,f5266,f5265,f5262,f5261,f5260,f5194,f5193,f5091,f5090,f5047,f4996,f4966,f4382,f4379,f4324,f4321,f4245,f4173,f4143,f4099,f3219,f3201,f3170,f3148,f3095,f3094,f3076,f2835,f2831,f2829,f2588,f2585,f2503,f2502,f2500,f2499,f2498,f2495,f2494,f2492,f2489,f2429,f2426,f2366,f2363,f2282,f2279,f2219,f2216,f2156,f2155,f2125,f2121,f2091,f2061,f2024,f1994,f1993,f1952,f1951,f1801,f1799,f1094,f1050,f957,f937,f934,f921,f612,f576,f542,f525,f524,f523,f522,f493,f453,f417,f264,f219,f187,f118,f103,f84,f83,f73,f12035,f12034,f12033,f11842,f11841,f8097,f8096,f6748,f6747,f6589,f5327,f74,f12802,f12642,f5147,f2480,f2417,f2270,f2207,f1157,f681,f642,f6770,f5146,f2354,f469,f468,f11961,f10086,f11962,f12275,f12274,f11553,f11552,f11188,f11187,f11119,f11118,f11010,f11009,f10928,f10927,f10674,f10673,f10120,f10095,f10093,f10090,f10074,f10073,f10018,f10017,f9332,f9314,f9190,f9189,f9150,f9149,f9081,f9080,f8058,f8057,f7988,f7987,f7199,f7180,f7177,f7163,f7162,f7108,f7107,f6583,f6579,f6568,f6567,f6564,f6563,f6554,f6549,f6547,f6546,f6545,f6544,f6543,f6515,f6058,f5321,f5310,f5309,f5306,f5305,f5299,f5294,f5292,f5291,f5290,f5289,f5288,f5267,f13719,f8832,f8681,f8680,f8679,f8080,f7247,f7246,f7176,f9851,f9850,f7389,f7387,f9473,f8898,f8897,f8896,f8895,f8678,f6944,f6690,f5192,f4995,f4172,f2023,f1944,f915,f450,f416,f16317,f16368,f16369,f20133,f20134,f20135,f20136,f20527,f20723,f20686,f20687,f20688,f20850,f20854,f20855,f20892,f20893,f20898,f20899,f20903,f20907,f20913,f20914,f20932,f20918,f20919,f20920,f20922,f20923,f21054,f21053,f21025,f21026,f21061,f21064,f21065,f21066,f21067,f21068,f21076,f21078,f21081,f21114,f21115,f21116,f21117,f21118,f21124,f21602,f21128,f21607,f21623,f21625,f21633,f21638,f21648,f21652,f21664,f21669])).
fof(f21665,plain,(
  $false | ($spl10 | $spl12)),
  inference(global_subsumption,[],[f70,f71,f76,f80,f79,f102,f107,f104,f108,f110,f131,f162,f207,f325,f326,f376,f558,f856,f575,f598,f611,f891,f651,f917,f923,f713,f715,f769,f792,f780,f996,f958,f1072,f1074,f1093,f1168,f1222,f1224,f1226,f1227,f1663,f1664,f1297,f1739,f1783,f2996,f2989,f2966,f3031,f2025,f2948,f2928,f2919,f2916,f2900,f2895,f3048,f2889,f2886,f3059,f2122,f2881,f2871,f3075,f2866,f2843,f2842,f2837,f2836,f3107,f2283,f2285,f3133,f2333,f3137,f3143,f3163,f3172,f3174,f3179,f3180,f2700,f3200,f3202,f456,f1665,f938,f1795,f1740,f1223,f714,f263,f249,f111,f455,f454,f3771,f3772,f5038,f5039,f5044,f6426,f6508,f6509,f6512,f6513,f6514,f6529,f6558,f6565,f6841,f7181,f7195,f7372,f7407,f8066,f8072,f8074,f8125,f8126,f8133,f8249,f8251,f8235,f8264,f8292,f8294,f8314,f8338,f8339,f8340,f8810,f9513,f9961,f9983,f9985,f10000,f10008,f10089,f11088,f11205,f11362,f11401,f11403,f11594,f11597,f11649,f11651,f11987,f11989,f11976,f12029,f12037,f12377,f13602,f12425,f13740,f690,f12040,f12038,f12036,f12010,f11996,f11978,f11973,f8132,f8127,f6839,f6838,f3141,f258,f77,f1294,f1165,f766,f687,f648,f13840,f13824,f13808,f13776,f13775,f13624,f13603,f13601,f13591,f13550,f13532,f13348,f13313,f13215,f13203,f13089,f13084,f12868,f12854,f12849,f12847,f12846,f12844,f12843,f12841,f12840,f12837,f12834,f12829,f12674,f12669,f12493,f12492,f12491,f12490,f12427,f12424,f12423,f12375,f12374,f12373,f12372,f12371,f12327,f12324,f12323,f12322,f12321,f12041,f12011,f12009,f12008,f11998,f11995,f11780,f11779,f11778,f11777,f11652,f11650,f11581,f11555,f11404,f11402,f11363,f11361,f11209,f11200,f11199,f11194,f11192,f9563,f9562,f9546,f9476,f9475,f9474,f9378,f9377,f9375,f9107,f9082,f8926,f8925,f8923,f8844,f8842,f8841,f8727,f8726,f8724,f8571,f8570,f8569,f8568,f8341,f8135,f8068,f8062,f8060,f7411,f7406,f7097,f7090,f6947,f6946,f6945,f6893,f6890,f6850,f6848,f6845,f6693,f6692,f6691,f6586,f6582,f6581,f6580,f6555,f6523,f6522,f6521,f6520,f6519,f6518,f6517,f6516,f6507,f6427,f6425,f6124,f5816,f5324,f5323,f5322,f5307,f5300,f5279,f5275,f5274,f5273,f5272,f5271,f5270,f5269,f5268,f5266,f5265,f5262,f5261,f5260,f5194,f5193,f5091,f5090,f5047,f4996,f4966,f4382,f4379,f4324,f4321,f4245,f4173,f4143,f4099,f3219,f3201,f3170,f3148,f3095,f3094,f3076,f2835,f2831,f2829,f2588,f2585,f2503,f2502,f2500,f2499,f2498,f2495,f2494,f2492,f2489,f2429,f2426,f2366,f2363,f2282,f2279,f2219,f2216,f2156,f2155,f2125,f2121,f2091,f2061,f2024,f1994,f1993,f1952,f1951,f1801,f1799,f1094,f1050,f957,f937,f934,f921,f612,f576,f542,f525,f524,f523,f522,f493,f453,f417,f264,f219,f187,f118,f103,f84,f83,f73,f12035,f12034,f12033,f11842,f11841,f8097,f8096,f6748,f6747,f6589,f5327,f74,f12802,f12642,f5147,f2480,f2417,f2270,f2207,f1157,f681,f642,f6770,f5146,f2354,f469,f468,f11961,f10086,f11962,f12275,f12274,f11553,f11552,f11188,f11187,f11119,f11118,f11010,f11009,f10928,f10927,f10674,f10673,f10120,f10095,f10093,f10090,f10074,f10073,f10018,f10017,f9332,f9314,f9190,f9189,f9150,f9149,f9081,f9080,f8058,f8057,f7988,f7987,f7199,f7180,f7177,f7163,f7162,f7108,f7107,f6583,f6579,f6568,f6567,f6564,f6563,f6554,f6549,f6547,f6546,f6545,f6544,f6543,f6515,f6058,f5321,f5310,f5309,f5306,f5305,f5299,f5294,f5292,f5291,f5290,f5289,f5288,f5267,f13719,f8832,f8681,f8680,f8679,f8080,f7247,f7246,f7176,f9851,f9850,f7389,f7387,f9473,f8898,f8897,f8896,f8895,f8678,f6944,f6690,f5192,f4995,f4172,f2023,f1944,f915,f450,f416,f16317,f16368,f16369,f20133,f20134,f20135,f20136,f20527,f20723,f20686,f20687,f20688,f20850,f20854,f20855,f20892,f20893,f20898,f20899,f20903,f20907,f20913,f20914,f20932,f20918,f20919,f20920,f20922,f20923,f21054,f21053,f21025,f21026,f21061,f21064,f21065,f21066,f21067,f21068,f21076,f21078,f21081,f21114,f21115,f21116,f21117,f21118,f21124,f21602,f21128,f21607,f21623,f21625,f21633,f21638,f21648,f21652,f21664])).
fof(f21658,plain,(
  $false | ($spl10 | $spl12 | $spl68)),
  inference(global_subsumption,[],[f21657,f70,f71,f76,f80,f79,f102,f107,f104,f108,f110,f131,f162,f207,f325,f326,f376,f558,f856,f575,f598,f611,f891,f651,f917,f923,f713,f715,f769,f792,f780,f996,f958,f1072,f1074,f1093,f1168,f1222,f1224,f1226,f1227,f1663,f1664,f1297,f1739,f1783,f2996,f2989,f2966,f3031,f2025,f2948,f2928,f2919,f2916,f2900,f2895,f3048,f2889,f2886,f3059,f2122,f2881,f2871,f3075,f2866,f2843,f2842,f2837,f2836,f3107,f2283,f2285,f3133,f2333,f3137,f3143,f3163,f3172,f3174,f3179,f3180,f2700,f3200,f3202,f456,f1665,f938,f1795,f1740,f1223,f714,f263,f249,f111,f455,f454,f3771,f3772,f5038,f5039,f5044,f6426,f6508,f6509,f6512,f6513,f6514,f6529,f6558,f6565,f6841,f7181,f7195,f7372,f7407,f8066,f8072,f8074,f8125,f8126,f8133,f8249,f8251,f8235,f8264,f8292,f8294,f8314,f8338,f8339,f8340,f8810,f9513,f9961,f9983,f9985,f10000,f10008,f10089,f11088,f11205,f11362,f11401,f11403,f11594,f11597,f11649,f11651,f11987,f11989,f11976,f12029,f12037,f12377,f13602,f12425,f13740,f690,f12040,f12038,f12036,f12010,f11996,f11978,f11973,f8132,f8127,f6839,f6838,f3141,f258,f77,f1294,f1165,f766,f687,f648,f13840,f13824,f13808,f13776,f13775,f13624,f13603,f13601,f13591,f13550,f13532,f13348,f13313,f13215,f13203,f13089,f13084,f12868,f12854,f12849,f12847,f12846,f12844,f12843,f12841,f12840,f12837,f12834,f12829,f12674,f12669,f12493,f12492,f12491,f12490,f12427,f12424,f12423,f12375,f12374,f12373,f12372,f12371,f12327,f12324,f12323,f12322,f12321,f12041,f12011,f12009,f12008,f11998,f11995,f11780,f11779,f11778,f11777,f11652,f11650,f11581,f11555,f11404,f11402,f11363,f11361,f11209,f11200,f11199,f11194,f11192,f9563,f9562,f9546,f9476,f9475,f9474,f9378,f9377,f9375,f9107,f9082,f8926,f8925,f8923,f8844,f8842,f8841,f8727,f8726,f8724,f8571,f8570,f8569,f8568,f8341,f8135,f8068,f8062,f8060,f7411,f7406,f7097,f7090,f6947,f6946,f6945,f6893,f6890,f6850,f6848,f6845,f6693,f6692,f6691,f6586,f6582,f6581,f6580,f6555,f6523,f6522,f6521,f6520,f6519,f6518,f6517,f6516,f6507,f6427,f6425,f6124,f5816,f5324,f5323,f5322,f5307,f5300,f5279,f5275,f5274,f5273,f5272,f5271,f5270,f5269,f5268,f5266,f5265,f5262,f5261,f5260,f5194,f5193,f5091,f5090,f5047,f4996,f4966,f4382,f4379,f4324,f4321,f4245,f4173,f4143,f4099,f3219,f3201,f3170,f3148,f3095,f3094,f3076,f2835,f2831,f2829,f2588,f2585,f2503,f2502,f2500,f2499,f2498,f2495,f2494,f2492,f2489,f2429,f2426,f2366,f2363,f2282,f2279,f2219,f2216,f2156,f2155,f2125,f2121,f2091,f2061,f2024,f1994,f1993,f1952,f1951,f1801,f1799,f1094,f1050,f957,f937,f934,f921,f612,f576,f542,f525,f524,f523,f522,f493,f453,f417,f264,f219,f187,f118,f103,f84,f83,f73,f12035,f12034,f12033,f11842,f11841,f8097,f8096,f6748,f6747,f6589,f5327,f74,f12802,f12642,f5147,f2480,f2417,f2270,f2207,f1157,f681,f642,f6770,f5146,f2354,f469,f468,f11961,f10086,f11962,f12275,f12274,f11553,f11552,f11188,f11187,f11119,f11118,f11010,f11009,f10928,f10927,f10674,f10673,f10120,f10095,f10093,f10090,f10074,f10073,f10018,f10017,f9332,f9314,f9190,f9189,f9150,f9149,f9081,f9080,f8058,f8057,f7988,f7987,f7199,f7180,f7177,f7163,f7162,f7108,f7107,f6583,f6579,f6568,f6567,f6564,f6563,f6554,f6549,f6547,f6546,f6545,f6544,f6543,f6515,f6058,f5321,f5310,f5309,f5306,f5305,f5299,f5294,f5292,f5291,f5290,f5289,f5288,f5267,f13719,f8832,f8681,f8680,f8679,f8080,f7247,f7246,f7176,f9851,f9850,f7389,f7387,f9473,f8898,f8897,f8896,f8895,f8678,f6944,f6690,f5192,f4995,f4172,f2023,f1944,f915,f450,f416,f16317,f16368,f16369,f20133,f20134,f20135,f20136,f20527,f20723,f20686,f20687,f20688,f20850,f20854,f20855,f20892,f20893,f20898,f20899,f20903,f20907,f20913,f20914,f20932,f20918,f20919,f20920,f20922,f20923,f21054,f21053,f21025,f21026,f21061,f21064,f21065,f21066,f21067,f21068,f21076,f21078,f21081,f21114,f21115,f21116,f21117,f21118,f21124,f21602,f21128,f21607,f21623,f21625,f21633,f21638,f21648,f21652])).
fof(f21657,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | ($spl10 | $spl12 | $spl68)),
  inference(condensation,[],[f21656])).
fof(f21656,plain,(
  ( ! [X191,X189] : (s(bool,'F_0') = s(bool,X191) | s(bool,'F_0') = s(bool,X189)) ) | ($spl10 | $spl12 | $spl68)),
  inference(subsumption_resolution,[],[f21392,f21124])).
fof(f21392,plain,(
  ( ! [X191,X189,X192,X190] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,X191),s(bool,'const.bool.COND_3'(s(bool,X189),s(bool,X189),s(bool,X190))),s(bool,X192))) | s(bool,'F_0') = s(bool,X191) | s(bool,'F_0') = s(bool,X189)) ) | ($spl10 | $spl12 | $spl68)),
  inference(backward_demodulation,[],[f21124,f10616])).
fof(f10616,plain,(
  ( ! [X191,X189,X192,X190] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X191),s(bool,'const.bool.COND_3'(s(bool,X189),s(bool,X189),s(bool,X190))),s(bool,X192))) | s(bool,'F_0') = s(bool,X191) | s(bool,'F_0') = s(bool,X189)) ) | ($spl12 | $spl68)),
  inference(subsumption_resolution,[],[f10471,f71])).
fof(f10471,plain,(
  ( ! [X191,X189,X192,X190] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X191),s(bool,'const.bool.COND_3'(s(bool,X189),s(bool,X189),s(bool,X190))),s(bool,X192))) | s(bool,'F_0') = s(bool,X191) | s(bool,'F_0') = s(bool,X189)) ) | ($spl12 | $spl68)),
  inference(superposition,[],[f5290,f10298])).
fof(f10298,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X6),s(bool,X6),s(bool,X7))) | s(bool,'F_0') = s(bool,X6)) ) | $spl68),
  inference(subsumption_resolution,[],[f10220,f73])).
fof(f10220,plain,(
  ( ! [X6,X7] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,X6) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X6),s(bool,X6),s(bool,X7)))) ) | $spl68),
  inference(superposition,[],[f9977,f76])).
fof(f21655,plain,(
  $false | ($spl10 | $spl12 | $spl68)),
  inference(global_subsumption,[],[f21654,f70,f71,f76,f80,f79,f102,f107,f104,f108,f110,f131,f162,f207,f325,f326,f376,f558,f856,f575,f598,f611,f891,f651,f917,f923,f713,f715,f769,f792,f780,f996,f958,f1072,f1074,f1093,f1168,f1222,f1224,f1226,f1227,f1663,f1664,f1297,f1739,f1783,f2996,f2989,f2966,f3031,f2025,f2948,f2928,f2919,f2916,f2900,f2895,f3048,f2889,f2886,f3059,f2122,f2881,f2871,f3075,f2866,f2843,f2842,f2837,f2836,f3107,f2283,f2285,f3133,f2333,f3137,f3143,f3163,f3172,f3174,f3179,f3180,f2700,f3200,f3202,f456,f1665,f938,f1795,f1740,f1223,f714,f263,f249,f111,f455,f454,f3771,f3772,f5038,f5039,f5044,f6426,f6508,f6509,f6512,f6513,f6514,f6529,f6558,f6565,f6841,f7181,f7195,f7372,f7407,f8066,f8072,f8074,f8125,f8126,f8133,f8249,f8251,f8235,f8264,f8292,f8294,f8314,f8338,f8339,f8340,f8810,f9513,f9961,f9983,f9985,f10000,f10008,f10089,f11088,f11205,f11362,f11401,f11403,f11594,f11597,f11649,f11651,f11987,f11989,f11976,f12029,f12037,f12377,f13602,f12425,f13740,f690,f12040,f12038,f12036,f12010,f11996,f11978,f11973,f8132,f8127,f6839,f6838,f3141,f258,f77,f1294,f1165,f766,f687,f648,f13840,f13824,f13808,f13776,f13775,f13624,f13603,f13601,f13591,f13550,f13532,f13348,f13313,f13215,f13203,f13089,f13084,f12868,f12854,f12849,f12847,f12846,f12844,f12843,f12841,f12840,f12837,f12834,f12829,f12674,f12669,f12493,f12492,f12491,f12490,f12427,f12424,f12423,f12375,f12374,f12373,f12372,f12371,f12327,f12324,f12323,f12322,f12321,f12041,f12011,f12009,f12008,f11998,f11995,f11780,f11779,f11778,f11777,f11652,f11650,f11581,f11555,f11404,f11402,f11363,f11361,f11209,f11200,f11199,f11194,f11192,f9563,f9562,f9546,f9476,f9475,f9474,f9378,f9377,f9375,f9107,f9082,f8926,f8925,f8923,f8844,f8842,f8841,f8727,f8726,f8724,f8571,f8570,f8569,f8568,f8341,f8135,f8068,f8062,f8060,f7411,f7406,f7097,f7090,f6947,f6946,f6945,f6893,f6890,f6850,f6848,f6845,f6693,f6692,f6691,f6586,f6582,f6581,f6580,f6555,f6523,f6522,f6521,f6520,f6519,f6518,f6517,f6516,f6507,f6427,f6425,f6124,f5816,f5324,f5323,f5322,f5307,f5300,f5279,f5275,f5274,f5273,f5272,f5271,f5270,f5269,f5268,f5266,f5265,f5262,f5261,f5260,f5194,f5193,f5091,f5090,f5047,f4996,f4966,f4382,f4379,f4324,f4321,f4245,f4173,f4143,f4099,f3219,f3201,f3170,f3148,f3095,f3094,f3076,f2835,f2831,f2829,f2588,f2585,f2503,f2502,f2500,f2499,f2498,f2495,f2494,f2492,f2489,f2429,f2426,f2366,f2363,f2282,f2279,f2219,f2216,f2156,f2155,f2125,f2121,f2091,f2061,f2024,f1994,f1993,f1952,f1951,f1801,f1799,f1094,f1050,f957,f937,f934,f921,f612,f576,f542,f525,f524,f523,f522,f493,f453,f417,f264,f219,f187,f118,f103,f84,f83,f73,f12035,f12034,f12033,f11842,f11841,f8097,f8096,f6748,f6747,f6589,f5327,f74,f12802,f12642,f5147,f2480,f2417,f2270,f2207,f1157,f681,f642,f6770,f5146,f2354,f469,f468,f11961,f10086,f11962,f12275,f12274,f11553,f11552,f11188,f11187,f11119,f11118,f11010,f11009,f10928,f10927,f10674,f10673,f10120,f10095,f10093,f10090,f10074,f10073,f10018,f10017,f9332,f9314,f9190,f9189,f9150,f9149,f9081,f9080,f8058,f8057,f7988,f7987,f7199,f7180,f7177,f7163,f7162,f7108,f7107,f6583,f6579,f6568,f6567,f6564,f6563,f6554,f6549,f6547,f6546,f6545,f6544,f6543,f6515,f6058,f5321,f5310,f5309,f5306,f5305,f5299,f5294,f5292,f5291,f5290,f5289,f5288,f5267,f13719,f8832,f8681,f8680,f8679,f8080,f7247,f7246,f7176,f9851,f9850,f7389,f7387,f9473,f8898,f8897,f8896,f8895,f8678,f6944,f6690,f5192,f4995,f4172,f2023,f1944,f915,f450,f416,f16317,f16368,f16369,f20133,f20134,f20135,f20136,f20527,f20723,f20686,f20687,f20688,f20850,f20854,f20855,f20892,f20893,f20898,f20899,f20903,f20907,f20913,f20914,f20932,f20918,f20919,f20920,f20922,f20923,f21054,f21053,f21025,f21026,f21061,f21064,f21065,f21066,f21067,f21068,f21076,f21078,f21081,f21114,f21115,f21116,f21117,f21118,f21124,f21602,f21128,f21607,f21623,f21625,f21633,f21638,f21648,f21652])).
fof(f21654,plain,(
  ( ! [X26] : (s(bool,'F_0') = s(bool,X26)) ) | ($spl10 | $spl12 | $spl68)),
  inference(subsumption_resolution,[],[f21389,f21124])).
fof(f21389,plain,(
  ( ! [X26,X27] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,X26),s(bool,X26),s(bool,X27))) | s(bool,'F_0') = s(bool,X26)) ) | ($spl10 | $spl12 | $spl68)),
  inference(backward_demodulation,[],[f21124,f10603])).
fof(f10603,plain,(
  ( ! [X26,X27] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X26),s(bool,X26),s(bool,X27))) | s(bool,'F_0') = s(bool,X26)) ) | $spl68),
  inference(subsumption_resolution,[],[f10424,f71])).
fof(f10424,plain,(
  ( ! [X26,X27] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X26),s(bool,X26),s(bool,X27))) | s(bool,'F_0') = s(bool,X26)) ) | $spl68),
  inference(superposition,[],[f83,f10298])).
fof(f21653,plain,(
  $false | ($spl10 | $spl12)),
  inference(global_subsumption,[],[f70,f71,f76,f80,f79,f102,f107,f104,f108,f110,f131,f162,f207,f325,f326,f376,f558,f856,f575,f598,f611,f891,f651,f917,f923,f713,f715,f769,f792,f780,f996,f958,f1072,f1074,f1093,f1168,f1222,f1224,f1226,f1227,f1663,f1664,f1297,f1739,f1783,f2996,f2989,f2966,f3031,f2025,f2948,f2928,f2919,f2916,f2900,f2895,f3048,f2889,f2886,f3059,f2122,f2881,f2871,f3075,f2866,f2843,f2842,f2837,f2836,f3107,f2283,f2285,f3133,f2333,f3137,f3143,f3163,f3172,f3174,f3179,f3180,f2700,f3200,f3202,f456,f1665,f938,f1795,f1740,f1223,f714,f263,f249,f111,f455,f454,f3771,f3772,f5038,f5039,f5044,f6426,f6508,f6509,f6512,f6513,f6514,f6529,f6558,f6565,f6841,f7181,f7195,f7372,f7407,f8066,f8072,f8074,f8125,f8126,f8133,f8249,f8251,f8235,f8264,f8292,f8294,f8314,f8338,f8339,f8340,f8810,f9513,f9961,f9983,f9985,f10000,f10008,f10089,f11088,f11205,f11362,f11401,f11403,f11594,f11597,f11649,f11651,f11987,f11989,f11976,f12029,f12037,f12377,f13602,f12425,f13740,f690,f12040,f12038,f12036,f12010,f11996,f11978,f11973,f8132,f8127,f6839,f6838,f3141,f258,f77,f1294,f1165,f766,f687,f648,f13840,f13824,f13808,f13776,f13775,f13624,f13603,f13601,f13591,f13550,f13532,f13348,f13313,f13215,f13203,f13089,f13084,f12868,f12854,f12849,f12847,f12846,f12844,f12843,f12841,f12840,f12837,f12834,f12829,f12674,f12669,f12493,f12492,f12491,f12490,f12427,f12424,f12423,f12375,f12374,f12373,f12372,f12371,f12327,f12324,f12323,f12322,f12321,f12041,f12011,f12009,f12008,f11998,f11995,f11780,f11779,f11778,f11777,f11652,f11650,f11581,f11555,f11404,f11402,f11363,f11361,f11209,f11200,f11199,f11194,f11192,f9563,f9562,f9546,f9476,f9475,f9474,f9378,f9377,f9375,f9107,f9082,f8926,f8925,f8923,f8844,f8842,f8841,f8727,f8726,f8724,f8571,f8570,f8569,f8568,f8341,f8135,f8068,f8062,f8060,f7411,f7406,f7097,f7090,f6947,f6946,f6945,f6893,f6890,f6850,f6848,f6845,f6693,f6692,f6691,f6586,f6582,f6581,f6580,f6555,f6523,f6522,f6521,f6520,f6519,f6518,f6517,f6516,f6507,f6427,f6425,f6124,f5816,f5324,f5323,f5322,f5307,f5300,f5279,f5275,f5274,f5273,f5272,f5271,f5270,f5269,f5268,f5266,f5265,f5262,f5261,f5260,f5194,f5193,f5091,f5090,f5047,f4996,f4966,f4382,f4379,f4324,f4321,f4245,f4173,f4143,f4099,f3219,f3201,f3170,f3148,f3095,f3094,f3076,f2835,f2831,f2829,f2588,f2585,f2503,f2502,f2500,f2499,f2498,f2495,f2494,f2492,f2489,f2429,f2426,f2366,f2363,f2282,f2279,f2219,f2216,f2156,f2155,f2125,f2121,f2091,f2061,f2024,f1994,f1993,f1952,f1951,f1801,f1799,f1094,f1050,f957,f937,f934,f921,f612,f576,f542,f525,f524,f523,f522,f493,f453,f417,f264,f219,f187,f118,f103,f84,f83,f73,f12035,f12034,f12033,f11842,f11841,f8097,f8096,f6748,f6747,f6589,f5327,f74,f12802,f12642,f5147,f2480,f2417,f2270,f2207,f1157,f681,f642,f6770,f5146,f2354,f469,f468,f11961,f10086,f11962,f12275,f12274,f11553,f11552,f11188,f11187,f11119,f11118,f11010,f11009,f10928,f10927,f10674,f10673,f10120,f10095,f10093,f10090,f10074,f10073,f10018,f10017,f9332,f9314,f9190,f9189,f9150,f9149,f9081,f9080,f8058,f8057,f7988,f7987,f7199,f7180,f7177,f7163,f7162,f7108,f7107,f6583,f6579,f6568,f6567,f6564,f6563,f6554,f6549,f6547,f6546,f6545,f6544,f6543,f6515,f6058,f5321,f5310,f5309,f5306,f5305,f5299,f5294,f5292,f5291,f5290,f5289,f5288,f5267,f13719,f8832,f8681,f8680,f8679,f8080,f7247,f7246,f7176,f9851,f9850,f7389,f7387,f9473,f8898,f8897,f8896,f8895,f8678,f6944,f6690,f5192,f4995,f4172,f2023,f1944,f915,f450,f416,f16317,f16368,f16369,f20133,f20134,f20135,f20136,f20527,f20723,f20686,f20687,f20688,f20850,f20854,f20855,f20892,f20893,f20898,f20899,f20903,f20907,f20913,f20914,f20932,f20918,f20919,f20920,f20922,f20923,f21054,f21053,f21025,f21026,f21061,f21064,f21065,f21066,f21067,f21068,f21076,f21078,f21081,f21114,f21115,f21116,f21117,f21118,f21124,f21602,f21128,f21607,f21623,f21625,f21633,f21638,f21648,f21652])).
fof(f21649,plain,(
  $false | ($spl10 | $spl12)),
  inference(global_subsumption,[],[f70,f71,f76,f80,f79,f102,f107,f104,f108,f110,f131,f162,f207,f325,f326,f376,f558,f856,f575,f598,f611,f891,f651,f917,f923,f713,f715,f769,f792,f780,f996,f958,f1072,f1074,f1093,f1168,f1222,f1224,f1226,f1227,f1663,f1664,f1297,f1739,f1783,f2996,f2989,f2966,f3031,f2025,f2948,f2928,f2919,f2916,f2900,f2895,f3048,f2889,f2886,f3059,f2122,f2881,f2871,f3075,f2866,f2843,f2842,f2837,f2836,f3107,f2283,f2285,f3133,f2333,f3137,f3143,f3163,f3172,f3174,f3179,f3180,f2700,f3200,f3202,f456,f1665,f938,f1795,f1740,f1223,f714,f263,f249,f111,f455,f454,f3771,f3772,f5038,f5039,f5044,f6426,f6508,f6509,f6512,f6513,f6514,f6529,f6558,f6565,f6841,f7181,f7195,f7372,f7407,f8066,f8072,f8074,f8125,f8126,f8133,f8249,f8251,f8235,f8264,f8292,f8294,f8314,f8338,f8339,f8340,f8810,f9513,f9961,f9983,f9985,f10000,f10008,f10089,f11088,f11205,f11362,f11401,f11403,f11594,f11597,f11649,f11651,f11987,f11989,f11976,f12029,f12037,f12377,f13602,f12425,f13740,f690,f12040,f12038,f12036,f12010,f11996,f11978,f11973,f8132,f8127,f6839,f6838,f3141,f258,f77,f1294,f1165,f766,f687,f648,f13840,f13824,f13808,f13776,f13775,f13624,f13603,f13601,f13591,f13550,f13532,f13348,f13313,f13215,f13203,f13089,f13084,f12868,f12854,f12849,f12847,f12846,f12844,f12843,f12841,f12840,f12837,f12834,f12829,f12674,f12669,f12493,f12492,f12491,f12490,f12427,f12424,f12423,f12375,f12374,f12373,f12372,f12371,f12327,f12324,f12323,f12322,f12321,f12041,f12011,f12009,f12008,f11998,f11995,f11780,f11779,f11778,f11777,f11652,f11650,f11581,f11555,f11404,f11402,f11363,f11361,f11209,f11200,f11199,f11194,f11192,f9563,f9562,f9546,f9476,f9475,f9474,f9378,f9377,f9375,f9107,f9082,f8926,f8925,f8923,f8844,f8842,f8841,f8727,f8726,f8724,f8571,f8570,f8569,f8568,f8341,f8135,f8068,f8062,f8060,f7411,f7406,f7097,f7090,f6947,f6946,f6945,f6893,f6890,f6850,f6848,f6845,f6693,f6692,f6691,f6586,f6582,f6581,f6580,f6555,f6523,f6522,f6521,f6520,f6519,f6518,f6517,f6516,f6507,f6427,f6425,f6124,f5816,f5324,f5323,f5322,f5307,f5300,f5279,f5275,f5274,f5273,f5272,f5271,f5270,f5269,f5268,f5266,f5265,f5262,f5261,f5260,f5194,f5193,f5091,f5090,f5047,f4996,f4966,f4382,f4379,f4324,f4321,f4245,f4173,f4143,f4099,f3219,f3201,f3170,f3148,f3095,f3094,f3076,f2835,f2831,f2829,f2588,f2585,f2503,f2502,f2500,f2499,f2498,f2495,f2494,f2492,f2489,f2429,f2426,f2366,f2363,f2282,f2279,f2219,f2216,f2156,f2155,f2125,f2121,f2091,f2061,f2024,f1994,f1993,f1952,f1951,f1801,f1799,f1094,f1050,f957,f937,f934,f921,f612,f576,f542,f525,f524,f523,f522,f493,f453,f417,f264,f219,f187,f118,f103,f84,f83,f73,f12035,f12034,f12033,f11842,f11841,f8097,f8096,f6748,f6747,f6589,f5327,f74,f12802,f12642,f5147,f2480,f2417,f2270,f2207,f1157,f681,f642,f6770,f5146,f2354,f469,f468,f11961,f10086,f11962,f12275,f12274,f11553,f11552,f11188,f11187,f11119,f11118,f11010,f11009,f10928,f10927,f10674,f10673,f10120,f10095,f10093,f10090,f10074,f10073,f10018,f10017,f9332,f9314,f9190,f9189,f9150,f9149,f9081,f9080,f8058,f8057,f7988,f7987,f7199,f7180,f7177,f7163,f7162,f7108,f7107,f6583,f6579,f6568,f6567,f6564,f6563,f6554,f6549,f6547,f6546,f6545,f6544,f6543,f6515,f6058,f5321,f5310,f5309,f5306,f5305,f5299,f5294,f5292,f5291,f5290,f5289,f5288,f5267,f13719,f8832,f8681,f8680,f8679,f8080,f7247,f7246,f7176,f9851,f9850,f7389,f7387,f9473,f8898,f8897,f8896,f8895,f8678,f6944,f6690,f5192,f4995,f4172,f2023,f1944,f915,f450,f416,f16317,f16368,f16369,f20133,f20134,f20135,f20136,f20527,f20723,f20686,f20687,f20688,f20850,f20854,f20855,f20892,f20893,f20898,f20899,f20903,f20907,f20913,f20914,f20932,f20918,f20919,f20920,f20922,f20923,f21054,f21053,f21025,f21026,f21061,f21064,f21065,f21066,f21067,f21068,f21076,f21078,f21081,f21114,f21115,f21116,f21117,f21118,f21124,f21602,f21128,f21607,f21623,f21625,f21633,f21638,f21648])).
fof(f21639,plain,(
  $false | ($spl10 | $spl12)),
  inference(global_subsumption,[],[f70,f71,f76,f80,f79,f102,f107,f104,f108,f110,f131,f162,f207,f325,f326,f376,f558,f856,f575,f598,f611,f891,f651,f917,f923,f713,f715,f769,f792,f780,f996,f958,f1072,f1074,f1093,f1168,f1222,f1224,f1226,f1227,f1663,f1664,f1297,f1739,f1783,f2996,f2989,f2966,f3031,f2025,f2948,f2928,f2919,f2916,f2900,f2895,f3048,f2889,f2886,f3059,f2122,f2881,f2871,f3075,f2866,f2843,f2842,f2837,f2836,f3107,f2283,f2285,f3133,f2333,f3137,f3143,f3163,f3172,f3174,f3179,f3180,f2700,f3200,f3202,f456,f1665,f938,f1795,f1740,f1223,f714,f263,f249,f111,f455,f454,f3771,f3772,f5038,f5039,f5044,f6426,f6508,f6509,f6512,f6513,f6514,f6529,f6558,f6565,f6841,f7181,f7195,f7372,f7407,f8066,f8072,f8074,f8125,f8126,f8133,f8249,f8251,f8235,f8264,f8292,f8294,f8314,f8338,f8339,f8340,f8810,f9513,f9961,f9983,f9985,f10000,f10008,f10089,f11088,f11205,f11362,f11401,f11403,f11594,f11597,f11649,f11651,f11987,f11989,f11976,f12029,f12037,f12377,f13602,f12425,f13740,f690,f12040,f12038,f12036,f12010,f11996,f11978,f11973,f8132,f8127,f6839,f6838,f3141,f258,f77,f1294,f1165,f766,f687,f648,f13840,f13824,f13808,f13776,f13775,f13624,f13603,f13601,f13591,f13550,f13532,f13348,f13313,f13215,f13203,f13089,f13084,f12868,f12854,f12849,f12847,f12846,f12844,f12843,f12841,f12840,f12837,f12834,f12829,f12674,f12669,f12493,f12492,f12491,f12490,f12427,f12424,f12423,f12375,f12374,f12373,f12372,f12371,f12327,f12324,f12323,f12322,f12321,f12041,f12011,f12009,f12008,f11998,f11995,f11780,f11779,f11778,f11777,f11652,f11650,f11581,f11555,f11404,f11402,f11363,f11361,f11209,f11200,f11199,f11194,f11192,f9563,f9562,f9546,f9476,f9475,f9474,f9378,f9377,f9375,f9107,f9082,f8926,f8925,f8923,f8844,f8842,f8841,f8727,f8726,f8724,f8571,f8570,f8569,f8568,f8341,f8135,f8068,f8062,f8060,f7411,f7406,f7097,f7090,f6947,f6946,f6945,f6893,f6890,f6850,f6848,f6845,f6693,f6692,f6691,f6586,f6582,f6581,f6580,f6555,f6523,f6522,f6521,f6520,f6519,f6518,f6517,f6516,f6507,f6427,f6425,f6124,f5816,f5324,f5323,f5322,f5307,f5300,f5279,f5275,f5274,f5273,f5272,f5271,f5270,f5269,f5268,f5266,f5265,f5262,f5261,f5260,f5194,f5193,f5091,f5090,f5047,f4996,f4966,f4382,f4379,f4324,f4321,f4245,f4173,f4143,f4099,f3219,f3201,f3170,f3148,f3095,f3094,f3076,f2835,f2831,f2829,f2588,f2585,f2503,f2502,f2500,f2499,f2498,f2495,f2494,f2492,f2489,f2429,f2426,f2366,f2363,f2282,f2279,f2219,f2216,f2156,f2155,f2125,f2121,f2091,f2061,f2024,f1994,f1993,f1952,f1951,f1801,f1799,f1094,f1050,f957,f937,f934,f921,f612,f576,f542,f525,f524,f523,f522,f493,f453,f417,f264,f219,f187,f118,f103,f84,f83,f73,f12035,f12034,f12033,f11842,f11841,f8097,f8096,f6748,f6747,f6589,f5327,f74,f12802,f12642,f5147,f2480,f2417,f2270,f2207,f1157,f681,f642,f6770,f5146,f2354,f469,f468,f11961,f10086,f11962,f12275,f12274,f11553,f11552,f11188,f11187,f11119,f11118,f11010,f11009,f10928,f10927,f10674,f10673,f10120,f10095,f10093,f10090,f10074,f10073,f10018,f10017,f9332,f9314,f9190,f9189,f9150,f9149,f9081,f9080,f8058,f8057,f7988,f7987,f7199,f7180,f7177,f7163,f7162,f7108,f7107,f6583,f6579,f6568,f6567,f6564,f6563,f6554,f6549,f6547,f6546,f6545,f6544,f6543,f6515,f6058,f5321,f5310,f5309,f5306,f5305,f5299,f5294,f5292,f5291,f5290,f5289,f5288,f5267,f13719,f8832,f8681,f8680,f8679,f8080,f7247,f7246,f7176,f9851,f9850,f7389,f7387,f9473,f8898,f8897,f8896,f8895,f8678,f6944,f6690,f5192,f4995,f4172,f2023,f1944,f915,f450,f416,f16317,f16368,f16369,f20133,f20134,f20135,f20136,f20527,f20723,f20686,f20687,f20688,f20850,f20854,f20855,f20892,f20893,f20898,f20899,f20903,f20907,f20913,f20914,f20932,f20918,f20919,f20920,f20922,f20923,f21054,f21053,f21025,f21026,f21061,f21064,f21065,f21066,f21067,f21068,f21076,f21078,f21081,f21114,f21115,f21116,f21117,f21118,f21124,f21602,f21128,f21607,f21623,f21625,f21633,f21638])).
fof(f21637,plain,(
  $false | ($spl10 | $spl12 | $spl26)),
  inference(global_subsumption,[],[f21636,f70,f71,f76,f80,f79,f102,f107,f104,f108,f110,f131,f162,f207,f325,f326,f376,f558,f856,f575,f598,f611,f891,f651,f917,f923,f713,f715,f769,f792,f780,f996,f958,f1072,f1074,f1093,f1168,f1222,f1224,f1226,f1227,f1663,f1664,f1297,f1739,f1783,f2996,f2989,f2966,f3031,f2025,f2948,f2928,f2919,f2916,f2900,f2895,f3048,f2889,f2886,f3059,f2122,f2881,f2871,f3075,f2866,f2843,f2842,f2837,f2836,f3107,f2283,f2285,f3133,f2333,f3137,f3143,f3163,f3172,f3174,f3179,f3180,f2700,f3200,f3202,f456,f1665,f938,f1795,f1740,f1223,f714,f263,f249,f111,f455,f454,f3771,f3772,f5038,f5039,f5044,f6426,f6508,f6509,f6512,f6513,f6514,f6529,f6558,f6565,f6841,f7181,f7195,f7372,f7407,f8066,f8072,f8074,f8125,f8126,f8133,f8249,f8251,f8235,f8264,f8292,f8294,f8314,f8338,f8339,f8340,f8810,f9513,f9961,f9983,f9985,f10000,f10008,f10089,f11088,f11205,f11362,f11401,f11403,f11594,f11597,f11649,f11651,f11987,f11989,f11976,f12029,f12037,f12377,f13602,f12425,f13740,f690,f12040,f12038,f12036,f12010,f11996,f11978,f11973,f8132,f8127,f6839,f6838,f3141,f258,f77,f1294,f1165,f766,f687,f648,f13840,f13824,f13808,f13776,f13775,f13624,f13603,f13601,f13591,f13550,f13532,f13348,f13313,f13215,f13203,f13089,f13084,f12868,f12854,f12849,f12847,f12846,f12844,f12843,f12841,f12840,f12837,f12834,f12829,f12674,f12669,f12493,f12492,f12491,f12490,f12427,f12424,f12423,f12375,f12374,f12373,f12372,f12371,f12327,f12324,f12323,f12322,f12321,f12041,f12011,f12009,f12008,f11998,f11995,f11780,f11779,f11778,f11777,f11652,f11650,f11581,f11555,f11404,f11402,f11363,f11361,f11209,f11200,f11199,f11194,f11192,f9563,f9562,f9546,f9476,f9475,f9474,f9378,f9377,f9375,f9107,f9082,f8926,f8925,f8923,f8844,f8842,f8841,f8727,f8726,f8724,f8571,f8570,f8569,f8568,f8341,f8135,f8068,f8062,f8060,f7411,f7406,f7097,f7090,f6947,f6946,f6945,f6893,f6890,f6850,f6848,f6845,f6693,f6692,f6691,f6586,f6582,f6581,f6580,f6555,f6523,f6522,f6521,f6520,f6519,f6518,f6517,f6516,f6507,f6427,f6425,f6124,f5816,f5324,f5323,f5322,f5307,f5300,f5279,f5275,f5274,f5273,f5272,f5271,f5270,f5269,f5268,f5266,f5265,f5262,f5261,f5260,f5194,f5193,f5091,f5090,f5047,f4996,f4966,f4382,f4379,f4324,f4321,f4245,f4173,f4143,f4099,f3219,f3201,f3170,f3148,f3095,f3094,f3076,f2835,f2831,f2829,f2588,f2585,f2503,f2502,f2500,f2499,f2498,f2495,f2494,f2492,f2489,f2429,f2426,f2366,f2363,f2282,f2279,f2219,f2216,f2156,f2155,f2125,f2121,f2091,f2061,f2024,f1994,f1993,f1952,f1951,f1801,f1799,f1094,f1050,f957,f937,f934,f921,f612,f576,f542,f525,f524,f523,f522,f493,f453,f417,f264,f219,f187,f118,f103,f84,f83,f73,f12035,f12034,f12033,f11842,f11841,f8097,f8096,f6748,f6747,f6589,f5327,f74,f12802,f12642,f5147,f2480,f2417,f2270,f2207,f1157,f681,f642,f6770,f5146,f2354,f469,f468,f11961,f10086,f11962,f12275,f12274,f11553,f11552,f11188,f11187,f11119,f11118,f11010,f11009,f10928,f10927,f10674,f10673,f10120,f10095,f10093,f10090,f10074,f10073,f10018,f10017,f9332,f9314,f9190,f9189,f9150,f9149,f9081,f9080,f8058,f8057,f7988,f7987,f7199,f7180,f7177,f7163,f7162,f7108,f7107,f6583,f6579,f6568,f6567,f6564,f6563,f6554,f6549,f6547,f6546,f6545,f6544,f6543,f6515,f6058,f5321,f5310,f5309,f5306,f5305,f5299,f5294,f5292,f5291,f5290,f5289,f5288,f5267,f13719,f8832,f8681,f8680,f8679,f8080,f7247,f7246,f7176,f9851,f9850,f7389,f7387,f9473,f8898,f8897,f8896,f8895,f8678,f6944,f6690,f5192,f4995,f4172,f2023,f1944,f915,f450,f416,f16317,f16368,f16369,f20133,f20134,f20135,f20136,f20527,f20723,f20686,f20687,f20688,f20850,f20854,f20855,f20892,f20893,f20898,f20899,f20903,f20907,f20913,f20914,f20932,f20918,f20919,f20920,f20922,f20923,f21054,f21053,f21025,f21026,f21061,f21064,f21065,f21066,f21067,f21068,f21076,f21078,f21081,f21114,f21115,f21116,f21117,f21118,f21124,f21602,f21128,f21607,f21623,f21625,f21633])).
fof(f21636,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | ($spl10 | $spl12 | $spl26)),
  inference(condensation,[],[f21635])).
fof(f21635,plain,(
  ( ! [X85,X83] : (s(bool,'F_0') = s(bool,X85) | s(bool,'F_0') = s(bool,X83)) ) | ($spl10 | $spl12 | $spl26)),
  inference(subsumption_resolution,[],[f21320,f21124])).
fof(f21320,plain,(
  ( ! [X85,X83,X86] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,X85),s(bool,X83),s(bool,X86))) | s(bool,'F_0') = s(bool,X85) | s(bool,'F_0') = s(bool,X83)) ) | ($spl10 | $spl12 | $spl26)),
  inference(backward_demodulation,[],[f21124,f8076])).
fof(f8076,plain,(
  ( ! [X85,X83,X86] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X85),s(bool,X83),s(bool,X86))) | s(bool,'F_0') = s(bool,X85) | s(bool,'F_0') = s(bool,X83)) ) | ($spl12 | $spl26)),
  inference(subsumption_resolution,[],[f8075,f71])).
fof(f8075,plain,(
  ( ! [X85,X83,X86] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X85),s(bool,X83),s(bool,X86))) | s(bool,'F_0') = s(bool,X85) | s(bool,'F_0') = s(bool,X83)) ) | ($spl12 | $spl26)),
  inference(forward_demodulation,[],[f8039,f102])).
fof(f8039,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X84)))) | s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X85),s(bool,X83),s(bool,X86))) | s(bool,'F_0') = s(bool,X85) | s(bool,'F_0') = s(bool,X83)) ) | ($spl12 | $spl26)),
  inference(superposition,[],[f5290,f3173])).
fof(f21634,plain,(
  $false | ($spl10 | $spl12)),
  inference(global_subsumption,[],[f70,f71,f76,f80,f79,f102,f107,f104,f108,f110,f131,f162,f207,f325,f326,f376,f558,f856,f575,f598,f611,f891,f651,f917,f923,f713,f715,f769,f792,f780,f996,f958,f1072,f1074,f1093,f1168,f1222,f1224,f1226,f1227,f1663,f1664,f1297,f1739,f1783,f2996,f2989,f2966,f3031,f2025,f2948,f2928,f2919,f2916,f2900,f2895,f3048,f2889,f2886,f3059,f2122,f2881,f2871,f3075,f2866,f2843,f2842,f2837,f2836,f3107,f2283,f2285,f3133,f2333,f3137,f3143,f3163,f3172,f3174,f3179,f3180,f2700,f3200,f3202,f456,f1665,f938,f1795,f1740,f1223,f714,f263,f249,f111,f455,f454,f3771,f3772,f5038,f5039,f5044,f6426,f6508,f6509,f6512,f6513,f6514,f6529,f6558,f6565,f6841,f7181,f7195,f7372,f7407,f8066,f8072,f8074,f8125,f8126,f8133,f8249,f8251,f8235,f8264,f8292,f8294,f8314,f8338,f8339,f8340,f8810,f9513,f9961,f9983,f9985,f10000,f10008,f10089,f11088,f11205,f11362,f11401,f11403,f11594,f11597,f11649,f11651,f11987,f11989,f11976,f12029,f12037,f12377,f13602,f12425,f13740,f690,f12040,f12038,f12036,f12010,f11996,f11978,f11973,f8132,f8127,f6839,f6838,f3141,f258,f77,f1294,f1165,f766,f687,f648,f13840,f13824,f13808,f13776,f13775,f13624,f13603,f13601,f13591,f13550,f13532,f13348,f13313,f13215,f13203,f13089,f13084,f12868,f12854,f12849,f12847,f12846,f12844,f12843,f12841,f12840,f12837,f12834,f12829,f12674,f12669,f12493,f12492,f12491,f12490,f12427,f12424,f12423,f12375,f12374,f12373,f12372,f12371,f12327,f12324,f12323,f12322,f12321,f12041,f12011,f12009,f12008,f11998,f11995,f11780,f11779,f11778,f11777,f11652,f11650,f11581,f11555,f11404,f11402,f11363,f11361,f11209,f11200,f11199,f11194,f11192,f9563,f9562,f9546,f9476,f9475,f9474,f9378,f9377,f9375,f9107,f9082,f8926,f8925,f8923,f8844,f8842,f8841,f8727,f8726,f8724,f8571,f8570,f8569,f8568,f8341,f8135,f8068,f8062,f8060,f7411,f7406,f7097,f7090,f6947,f6946,f6945,f6893,f6890,f6850,f6848,f6845,f6693,f6692,f6691,f6586,f6582,f6581,f6580,f6555,f6523,f6522,f6521,f6520,f6519,f6518,f6517,f6516,f6507,f6427,f6425,f6124,f5816,f5324,f5323,f5322,f5307,f5300,f5279,f5275,f5274,f5273,f5272,f5271,f5270,f5269,f5268,f5266,f5265,f5262,f5261,f5260,f5194,f5193,f5091,f5090,f5047,f4996,f4966,f4382,f4379,f4324,f4321,f4245,f4173,f4143,f4099,f3219,f3201,f3170,f3148,f3095,f3094,f3076,f2835,f2831,f2829,f2588,f2585,f2503,f2502,f2500,f2499,f2498,f2495,f2494,f2492,f2489,f2429,f2426,f2366,f2363,f2282,f2279,f2219,f2216,f2156,f2155,f2125,f2121,f2091,f2061,f2024,f1994,f1993,f1952,f1951,f1801,f1799,f1094,f1050,f957,f937,f934,f921,f612,f576,f542,f525,f524,f523,f522,f493,f453,f417,f264,f219,f187,f118,f103,f84,f83,f73,f12035,f12034,f12033,f11842,f11841,f8097,f8096,f6748,f6747,f6589,f5327,f74,f12802,f12642,f5147,f2480,f2417,f2270,f2207,f1157,f681,f642,f6770,f5146,f2354,f469,f468,f11961,f10086,f11962,f12275,f12274,f11553,f11552,f11188,f11187,f11119,f11118,f11010,f11009,f10928,f10927,f10674,f10673,f10120,f10095,f10093,f10090,f10074,f10073,f10018,f10017,f9332,f9314,f9190,f9189,f9150,f9149,f9081,f9080,f8058,f8057,f7988,f7987,f7199,f7180,f7177,f7163,f7162,f7108,f7107,f6583,f6579,f6568,f6567,f6564,f6563,f6554,f6549,f6547,f6546,f6545,f6544,f6543,f6515,f6058,f5321,f5310,f5309,f5306,f5305,f5299,f5294,f5292,f5291,f5290,f5289,f5288,f5267,f13719,f8832,f8681,f8680,f8679,f8080,f7247,f7246,f7176,f9851,f9850,f7389,f7387,f9473,f8898,f8897,f8896,f8895,f8678,f6944,f6690,f5192,f4995,f4172,f2023,f1944,f915,f450,f416,f16317,f16368,f16369,f20133,f20134,f20135,f20136,f20527,f20723,f20686,f20687,f20688,f20850,f20854,f20855,f20892,f20893,f20898,f20899,f20903,f20907,f20913,f20914,f20932,f20918,f20919,f20920,f20922,f20923,f21054,f21053,f21025,f21026,f21061,f21064,f21065,f21066,f21067,f21068,f21076,f21078,f21081,f21114,f21115,f21116,f21117,f21118,f21124,f21602,f21128,f21607,f21623,f21625,f21633])).
fof(f21626,plain,(
  $false | ($spl10 | $spl12)),
  inference(global_subsumption,[],[f70,f71,f76,f80,f79,f102,f107,f104,f108,f110,f131,f162,f207,f325,f326,f376,f558,f856,f575,f598,f611,f891,f651,f917,f923,f713,f715,f769,f792,f780,f996,f958,f1072,f1074,f1093,f1168,f1222,f1224,f1226,f1227,f1663,f1664,f1297,f1739,f1783,f2996,f2989,f2966,f3031,f2025,f2948,f2928,f2919,f2916,f2900,f2895,f3048,f2889,f2886,f3059,f2122,f2881,f2871,f3075,f2866,f2843,f2842,f2837,f2836,f3107,f2283,f2285,f3133,f2333,f3137,f3143,f3163,f3172,f3174,f3179,f3180,f2700,f3200,f3202,f456,f1665,f938,f1795,f1740,f1223,f714,f263,f249,f111,f455,f454,f3771,f3772,f5038,f5039,f5044,f6426,f6508,f6509,f6512,f6513,f6514,f6529,f6558,f6565,f6841,f7181,f7195,f7372,f7407,f8066,f8072,f8074,f8125,f8126,f8133,f8249,f8251,f8235,f8264,f8292,f8294,f8314,f8338,f8339,f8340,f8810,f9513,f9961,f9983,f9985,f10000,f10008,f10089,f11088,f11205,f11362,f11401,f11403,f11594,f11597,f11649,f11651,f11987,f11989,f11976,f12029,f12037,f12377,f13602,f12425,f13740,f690,f12040,f12038,f12036,f12010,f11996,f11978,f11973,f8132,f8127,f6839,f6838,f3141,f258,f77,f1294,f1165,f766,f687,f648,f13840,f13824,f13808,f13776,f13775,f13624,f13603,f13601,f13591,f13550,f13532,f13348,f13313,f13215,f13203,f13089,f13084,f12868,f12854,f12849,f12847,f12846,f12844,f12843,f12841,f12840,f12837,f12834,f12829,f12674,f12669,f12493,f12492,f12491,f12490,f12427,f12424,f12423,f12375,f12374,f12373,f12372,f12371,f12327,f12324,f12323,f12322,f12321,f12041,f12011,f12009,f12008,f11998,f11995,f11780,f11779,f11778,f11777,f11652,f11650,f11581,f11555,f11404,f11402,f11363,f11361,f11209,f11200,f11199,f11194,f11192,f9563,f9562,f9546,f9476,f9475,f9474,f9378,f9377,f9375,f9107,f9082,f8926,f8925,f8923,f8844,f8842,f8841,f8727,f8726,f8724,f8571,f8570,f8569,f8568,f8341,f8135,f8068,f8062,f8060,f7411,f7406,f7097,f7090,f6947,f6946,f6945,f6893,f6890,f6850,f6848,f6845,f6693,f6692,f6691,f6586,f6582,f6581,f6580,f6555,f6523,f6522,f6521,f6520,f6519,f6518,f6517,f6516,f6507,f6427,f6425,f6124,f5816,f5324,f5323,f5322,f5307,f5300,f5279,f5275,f5274,f5273,f5272,f5271,f5270,f5269,f5268,f5266,f5265,f5262,f5261,f5260,f5194,f5193,f5091,f5090,f5047,f4996,f4966,f4382,f4379,f4324,f4321,f4245,f4173,f4143,f4099,f3219,f3201,f3170,f3148,f3095,f3094,f3076,f2835,f2831,f2829,f2588,f2585,f2503,f2502,f2500,f2499,f2498,f2495,f2494,f2492,f2489,f2429,f2426,f2366,f2363,f2282,f2279,f2219,f2216,f2156,f2155,f2125,f2121,f2091,f2061,f2024,f1994,f1993,f1952,f1951,f1801,f1799,f1094,f1050,f957,f937,f934,f921,f612,f576,f542,f525,f524,f523,f522,f493,f453,f417,f264,f219,f187,f118,f103,f84,f83,f73,f12035,f12034,f12033,f11842,f11841,f8097,f8096,f6748,f6747,f6589,f5327,f74,f12802,f12642,f5147,f2480,f2417,f2270,f2207,f1157,f681,f642,f6770,f5146,f2354,f469,f468,f11961,f10086,f11962,f12275,f12274,f11553,f11552,f11188,f11187,f11119,f11118,f11010,f11009,f10928,f10927,f10674,f10673,f10120,f10095,f10093,f10090,f10074,f10073,f10018,f10017,f9332,f9314,f9190,f9189,f9150,f9149,f9081,f9080,f8058,f8057,f7988,f7987,f7199,f7180,f7177,f7163,f7162,f7108,f7107,f6583,f6579,f6568,f6567,f6564,f6563,f6554,f6549,f6547,f6546,f6545,f6544,f6543,f6515,f6058,f5321,f5310,f5309,f5306,f5305,f5299,f5294,f5292,f5291,f5290,f5289,f5288,f5267,f13719,f8832,f8681,f8680,f8679,f8080,f7247,f7246,f7176,f9851,f9850,f7389,f7387,f9473,f8898,f8897,f8896,f8895,f8678,f6944,f6690,f5192,f4995,f4172,f2023,f1944,f915,f450,f416,f16317,f16368,f16369,f20133,f20134,f20135,f20136,f20527,f20723,f20686,f20687,f20688,f20850,f20854,f20855,f20892,f20893,f20898,f20899,f20903,f20907,f20913,f20914,f20932,f20918,f20919,f20920,f20922,f20923,f21054,f21053,f21025,f21026,f21061,f21064,f21065,f21066,f21067,f21068,f21076,f21078,f21081,f21114,f21115,f21116,f21117,f21118,f21124,f21602,f21128,f21607,f21623,f21625])).
fof(f21624,plain,(
  $false | ($spl10 | $spl12)),
  inference(global_subsumption,[],[f70,f71,f76,f80,f79,f102,f107,f104,f108,f110,f131,f162,f207,f325,f326,f376,f558,f856,f575,f598,f611,f891,f651,f917,f923,f713,f715,f769,f792,f780,f996,f958,f1072,f1074,f1093,f1168,f1222,f1224,f1226,f1227,f1663,f1664,f1297,f1739,f1783,f2996,f2989,f2966,f3031,f2025,f2948,f2928,f2919,f2916,f2900,f2895,f3048,f2889,f2886,f3059,f2122,f2881,f2871,f3075,f2866,f2843,f2842,f2837,f2836,f3107,f2283,f2285,f3133,f2333,f3137,f3143,f3163,f3172,f3174,f3179,f3180,f2700,f3200,f3202,f456,f1665,f938,f1795,f1740,f1223,f714,f263,f249,f111,f455,f454,f3771,f3772,f5038,f5039,f5044,f6426,f6508,f6509,f6512,f6513,f6514,f6529,f6558,f6565,f6841,f7181,f7195,f7372,f7407,f8066,f8072,f8074,f8125,f8126,f8133,f8249,f8251,f8235,f8264,f8292,f8294,f8314,f8338,f8339,f8340,f8810,f9513,f9961,f9983,f9985,f10000,f10008,f10089,f11088,f11205,f11362,f11401,f11403,f11594,f11597,f11649,f11651,f11987,f11989,f11976,f12029,f12037,f12377,f13602,f12425,f13740,f690,f12040,f12038,f12036,f12010,f11996,f11978,f11973,f8132,f8127,f6839,f6838,f3141,f258,f77,f1294,f1165,f766,f687,f648,f13840,f13824,f13808,f13776,f13775,f13624,f13603,f13601,f13591,f13550,f13532,f13348,f13313,f13215,f13203,f13089,f13084,f12868,f12854,f12849,f12847,f12846,f12844,f12843,f12841,f12840,f12837,f12834,f12829,f12674,f12669,f12493,f12492,f12491,f12490,f12427,f12424,f12423,f12375,f12374,f12373,f12372,f12371,f12327,f12324,f12323,f12322,f12321,f12041,f12011,f12009,f12008,f11998,f11995,f11780,f11779,f11778,f11777,f11652,f11650,f11581,f11555,f11404,f11402,f11363,f11361,f11209,f11200,f11199,f11194,f11192,f9563,f9562,f9546,f9476,f9475,f9474,f9378,f9377,f9375,f9107,f9082,f8926,f8925,f8923,f8844,f8842,f8841,f8727,f8726,f8724,f8571,f8570,f8569,f8568,f8341,f8135,f8068,f8062,f8060,f7411,f7406,f7097,f7090,f6947,f6946,f6945,f6893,f6890,f6850,f6848,f6845,f6693,f6692,f6691,f6586,f6582,f6581,f6580,f6555,f6523,f6522,f6521,f6520,f6519,f6518,f6517,f6516,f6507,f6427,f6425,f6124,f5816,f5324,f5323,f5322,f5307,f5300,f5279,f5275,f5274,f5273,f5272,f5271,f5270,f5269,f5268,f5266,f5265,f5262,f5261,f5260,f5194,f5193,f5091,f5090,f5047,f4996,f4966,f4382,f4379,f4324,f4321,f4245,f4173,f4143,f4099,f3219,f3201,f3170,f3148,f3095,f3094,f3076,f2835,f2831,f2829,f2588,f2585,f2503,f2502,f2500,f2499,f2498,f2495,f2494,f2492,f2489,f2429,f2426,f2366,f2363,f2282,f2279,f2219,f2216,f2156,f2155,f2125,f2121,f2091,f2061,f2024,f1994,f1993,f1952,f1951,f1801,f1799,f1094,f1050,f957,f937,f934,f921,f612,f576,f542,f525,f524,f523,f522,f493,f453,f417,f264,f219,f187,f118,f103,f84,f83,f73,f12035,f12034,f12033,f11842,f11841,f8097,f8096,f6748,f6747,f6589,f5327,f74,f12802,f12642,f5147,f2480,f2417,f2270,f2207,f1157,f681,f642,f6770,f5146,f2354,f469,f468,f11961,f10086,f11962,f12275,f12274,f11553,f11552,f11188,f11187,f11119,f11118,f11010,f11009,f10928,f10927,f10674,f10673,f10120,f10095,f10093,f10090,f10074,f10073,f10018,f10017,f9332,f9314,f9190,f9189,f9150,f9149,f9081,f9080,f8058,f8057,f7988,f7987,f7199,f7180,f7177,f7163,f7162,f7108,f7107,f6583,f6579,f6568,f6567,f6564,f6563,f6554,f6549,f6547,f6546,f6545,f6544,f6543,f6515,f6058,f5321,f5310,f5309,f5306,f5305,f5299,f5294,f5292,f5291,f5290,f5289,f5288,f5267,f13719,f8832,f8681,f8680,f8679,f8080,f7247,f7246,f7176,f9851,f9850,f7389,f7387,f9473,f8898,f8897,f8896,f8895,f8678,f6944,f6690,f5192,f4995,f4172,f2023,f1944,f915,f450,f416,f16317,f16368,f16369,f20133,f20134,f20135,f20136,f20527,f20723,f20686,f20687,f20688,f20850,f20854,f20855,f20892,f20893,f20898,f20899,f20903,f20907,f20913,f20914,f20932,f20918,f20919,f20920,f20922,f20923,f21054,f21053,f21025,f21026,f21061,f21064,f21065,f21066,f21067,f21068,f21076,f21078,f21081,f21114,f21115,f21116,f21117,f21118,f21124,f21602,f21128,f21607,f21623])).
fof(f21608,plain,(
  $false | ($spl10 | $spl12)),
  inference(global_subsumption,[],[f70,f71,f76,f80,f79,f102,f107,f104,f108,f110,f131,f162,f207,f325,f326,f376,f558,f856,f575,f598,f611,f891,f651,f917,f923,f713,f715,f769,f792,f780,f996,f958,f1072,f1074,f1093,f1168,f1222,f1224,f1226,f1227,f1663,f1664,f1297,f1739,f1783,f2996,f2989,f2966,f3031,f2025,f2948,f2928,f2919,f2916,f2900,f2895,f3048,f2889,f2886,f3059,f2122,f2881,f2871,f3075,f2866,f2843,f2842,f2837,f2836,f3107,f2283,f2285,f3133,f2333,f3137,f3143,f3163,f3172,f3174,f3179,f3180,f2700,f3200,f3202,f456,f1665,f938,f1795,f1740,f1223,f714,f263,f249,f111,f455,f454,f3771,f3772,f5038,f5039,f5044,f6426,f6508,f6509,f6512,f6513,f6514,f6529,f6558,f6565,f6841,f7181,f7195,f7372,f7407,f8066,f8072,f8074,f8125,f8126,f8133,f8249,f8251,f8235,f8264,f8292,f8294,f8314,f8338,f8339,f8340,f8810,f9513,f9961,f9983,f9985,f10000,f10008,f10089,f11088,f11205,f11362,f11401,f11403,f11594,f11597,f11649,f11651,f11987,f11989,f11976,f12029,f12037,f12377,f13602,f12425,f13740,f690,f12040,f12038,f12036,f12010,f11996,f11978,f11973,f8132,f8127,f6839,f6838,f3141,f258,f77,f1294,f1165,f766,f687,f648,f13840,f13824,f13808,f13776,f13775,f13624,f13603,f13601,f13591,f13550,f13532,f13348,f13313,f13215,f13203,f13089,f13084,f12868,f12854,f12849,f12847,f12846,f12844,f12843,f12841,f12840,f12837,f12834,f12829,f12674,f12669,f12493,f12492,f12491,f12490,f12427,f12424,f12423,f12375,f12374,f12373,f12372,f12371,f12327,f12324,f12323,f12322,f12321,f12041,f12011,f12009,f12008,f11998,f11995,f11780,f11779,f11778,f11777,f11652,f11650,f11581,f11555,f11404,f11402,f11363,f11361,f11209,f11200,f11199,f11194,f11192,f9563,f9562,f9546,f9476,f9475,f9474,f9378,f9377,f9375,f9107,f9082,f8926,f8925,f8923,f8844,f8842,f8841,f8727,f8726,f8724,f8571,f8570,f8569,f8568,f8341,f8135,f8068,f8062,f8060,f7411,f7406,f7097,f7090,f6947,f6946,f6945,f6893,f6890,f6850,f6848,f6845,f6693,f6692,f6691,f6586,f6582,f6581,f6580,f6555,f6523,f6522,f6521,f6520,f6519,f6518,f6517,f6516,f6507,f6427,f6425,f6124,f5816,f5324,f5323,f5322,f5307,f5300,f5279,f5275,f5274,f5273,f5272,f5271,f5270,f5269,f5268,f5266,f5265,f5262,f5261,f5260,f5194,f5193,f5091,f5090,f5047,f4996,f4966,f4382,f4379,f4324,f4321,f4245,f4173,f4143,f4099,f3219,f3201,f3170,f3148,f3095,f3094,f3076,f2835,f2831,f2829,f2588,f2585,f2503,f2502,f2500,f2499,f2498,f2495,f2494,f2492,f2489,f2429,f2426,f2366,f2363,f2282,f2279,f2219,f2216,f2156,f2155,f2125,f2121,f2091,f2061,f2024,f1994,f1993,f1952,f1951,f1801,f1799,f1094,f1050,f957,f937,f934,f921,f612,f576,f542,f525,f524,f523,f522,f493,f453,f417,f264,f219,f187,f118,f103,f84,f83,f73,f12035,f12034,f12033,f11842,f11841,f8097,f8096,f6748,f6747,f6589,f5327,f74,f12802,f12642,f5147,f2480,f2417,f2270,f2207,f1157,f681,f642,f6770,f5146,f2354,f469,f468,f11961,f10086,f11962,f12275,f12274,f11553,f11552,f11188,f11187,f11119,f11118,f11010,f11009,f10928,f10927,f10674,f10673,f10120,f10095,f10093,f10090,f10074,f10073,f10018,f10017,f9332,f9314,f9190,f9189,f9150,f9149,f9081,f9080,f8058,f8057,f7988,f7987,f7199,f7180,f7177,f7163,f7162,f7108,f7107,f6583,f6579,f6568,f6567,f6564,f6563,f6554,f6549,f6547,f6546,f6545,f6544,f6543,f6515,f6058,f5321,f5310,f5309,f5306,f5305,f5299,f5294,f5292,f5291,f5290,f5289,f5288,f5267,f13719,f8832,f8681,f8680,f8679,f8080,f7247,f7246,f7176,f9851,f9850,f7389,f7387,f9473,f8898,f8897,f8896,f8895,f8678,f6944,f6690,f5192,f4995,f4172,f2023,f1944,f915,f450,f416,f16317,f16368,f16369,f20133,f20134,f20135,f20136,f20527,f20723,f20686,f20687,f20688,f20850,f20854,f20855,f20892,f20893,f20898,f20899,f20903,f20907,f20913,f20914,f20932,f20918,f20919,f20920,f20922,f20923,f21054,f21053,f21025,f21026,f21061,f21064,f21065,f21066,f21067,f21068,f21076,f21078,f21081,f21114,f21115,f21116,f21117,f21118,f21124,f21602,f21128,f21607])).
fof(f21604,plain,(
  $false | ($spl10 | $spl12)),
  inference(global_subsumption,[],[f70,f71,f76,f80,f79,f102,f107,f104,f108,f110,f131,f162,f207,f325,f326,f376,f558,f856,f575,f598,f611,f891,f651,f917,f923,f713,f715,f769,f792,f780,f996,f958,f1072,f1074,f1093,f1168,f1222,f1224,f1226,f1227,f1663,f1664,f1297,f1739,f1783,f2996,f2989,f2966,f3031,f2025,f2948,f2928,f2919,f2916,f2900,f2895,f3048,f2889,f2886,f3059,f2122,f2881,f2871,f3075,f2866,f2843,f2842,f2837,f2836,f3107,f2283,f2285,f3133,f2333,f3137,f3143,f3163,f3172,f3174,f3179,f3180,f2700,f3200,f3202,f456,f1665,f938,f1795,f1740,f1223,f714,f263,f249,f111,f455,f454,f3771,f3772,f5038,f5039,f5044,f6426,f6508,f6509,f6512,f6513,f6514,f6529,f6558,f6565,f6841,f7181,f7195,f7372,f7407,f8066,f8072,f8074,f8125,f8126,f8133,f8249,f8251,f8235,f8264,f8292,f8294,f8314,f8338,f8339,f8340,f8810,f9513,f9961,f9983,f9985,f10000,f10008,f10089,f11088,f11205,f11362,f11401,f11403,f11594,f11597,f11649,f11651,f11987,f11989,f11976,f12029,f12037,f12377,f13602,f12425,f13740,f690,f12040,f12038,f12036,f12010,f11996,f11978,f11973,f8132,f8127,f6839,f6838,f3141,f258,f77,f1294,f1165,f766,f687,f648,f13840,f13824,f13808,f13776,f13775,f13624,f13603,f13601,f13591,f13550,f13532,f13348,f13313,f13215,f13203,f13089,f13084,f12868,f12854,f12849,f12847,f12846,f12844,f12843,f12841,f12840,f12837,f12834,f12829,f12674,f12669,f12493,f12492,f12491,f12490,f12427,f12424,f12423,f12375,f12374,f12373,f12372,f12371,f12327,f12324,f12323,f12322,f12321,f12041,f12011,f12009,f12008,f11998,f11995,f11780,f11779,f11778,f11777,f11652,f11650,f11581,f11555,f11404,f11402,f11363,f11361,f11209,f11200,f11199,f11194,f11192,f9563,f9562,f9546,f9476,f9475,f9474,f9378,f9377,f9375,f9107,f9082,f8926,f8925,f8923,f8844,f8842,f8841,f8727,f8726,f8724,f8571,f8570,f8569,f8568,f8341,f8135,f8068,f8062,f8060,f7411,f7406,f7097,f7090,f6947,f6946,f6945,f6893,f6890,f6850,f6848,f6845,f6693,f6692,f6691,f6586,f6582,f6581,f6580,f6555,f6523,f6522,f6521,f6520,f6519,f6518,f6517,f6516,f6507,f6427,f6425,f6124,f5816,f5324,f5323,f5322,f5307,f5300,f5279,f5275,f5274,f5273,f5272,f5271,f5270,f5269,f5268,f5266,f5265,f5262,f5261,f5260,f5194,f5193,f5091,f5090,f5047,f4996,f4966,f4382,f4379,f4324,f4321,f4245,f4173,f4143,f4099,f3219,f3201,f3170,f3148,f3095,f3094,f3076,f2835,f2831,f2829,f2588,f2585,f2503,f2502,f2500,f2499,f2498,f2495,f2494,f2492,f2489,f2429,f2426,f2366,f2363,f2282,f2279,f2219,f2216,f2156,f2155,f2125,f2121,f2091,f2061,f2024,f1994,f1993,f1952,f1951,f1801,f1799,f1094,f1050,f957,f937,f934,f921,f612,f576,f542,f525,f524,f523,f522,f493,f453,f417,f264,f219,f187,f118,f103,f84,f83,f73,f12035,f12034,f12033,f11842,f11841,f8097,f8096,f6748,f6747,f6589,f5327,f74,f12802,f12642,f5147,f2480,f2417,f2270,f2207,f1157,f681,f642,f6770,f5146,f2354,f469,f468,f11961,f10086,f11962,f12275,f12274,f11553,f11552,f11188,f11187,f11119,f11118,f11010,f11009,f10928,f10927,f10674,f10673,f10120,f10095,f10093,f10090,f10074,f10073,f10018,f10017,f9332,f9314,f9190,f9189,f9150,f9149,f9081,f9080,f8058,f8057,f7988,f7987,f7199,f7180,f7177,f7163,f7162,f7108,f7107,f6583,f6579,f6568,f6567,f6564,f6563,f6554,f6549,f6547,f6546,f6545,f6544,f6543,f6515,f6058,f5321,f5310,f5309,f5306,f5305,f5299,f5294,f5292,f5291,f5290,f5289,f5288,f5267,f13719,f8832,f8681,f8680,f8679,f8080,f7247,f7246,f7176,f9851,f9850,f7389,f7387,f9473,f8898,f8897,f8896,f8895,f8678,f6944,f6690,f5192,f4995,f4172,f2023,f1944,f915,f450,f416,f16317,f16368,f16369,f20133,f20134,f20135,f20136,f20527,f20723,f20686,f20687,f20688,f20850,f20854,f20855,f20892,f20893,f20898,f20899,f20903,f20907,f20913,f20914,f20932,f20918,f20919,f20920,f20922,f20923,f21054,f21053,f21025,f21026,f21061,f21064,f21065,f21066,f21067,f21068,f21076,f21078,f21081,f21114,f21115,f21116,f21117,f21118,f21124,f21602,f21128])).
fof(f21603,plain,(
  $false | ($spl10 | $spl12)),
  inference(global_subsumption,[],[f70,f71,f76,f80,f79,f102,f107,f104,f108,f110,f131,f162,f207,f325,f326,f376,f558,f856,f575,f598,f611,f891,f651,f917,f923,f713,f715,f769,f792,f780,f996,f958,f1072,f1074,f1093,f1168,f1222,f1224,f1226,f1227,f1663,f1664,f1297,f1739,f1783,f2996,f2989,f2966,f3031,f2025,f2948,f2928,f2919,f2916,f2900,f2895,f3048,f2889,f2886,f3059,f2122,f2881,f2871,f3075,f2866,f2843,f2842,f2837,f2836,f3107,f2283,f2285,f3133,f2333,f3137,f3143,f3163,f3172,f3174,f3179,f3180,f2700,f3200,f3202,f456,f1665,f938,f1795,f1740,f1223,f714,f263,f249,f111,f455,f454,f3771,f3772,f5038,f5039,f5044,f6426,f6508,f6509,f6512,f6513,f6514,f6529,f6558,f6565,f6841,f7181,f7195,f7372,f7407,f8066,f8072,f8074,f8125,f8126,f8133,f8249,f8251,f8235,f8264,f8292,f8294,f8314,f8338,f8339,f8340,f8810,f9513,f9961,f9983,f9985,f10000,f10008,f10089,f11088,f11205,f11362,f11401,f11403,f11594,f11597,f11649,f11651,f11987,f11989,f11976,f12029,f12037,f12377,f13602,f12425,f13740,f690,f12040,f12038,f12036,f12010,f11996,f11978,f11973,f8132,f8127,f6839,f6838,f3141,f258,f77,f1294,f1165,f766,f687,f648,f13840,f13824,f13808,f13776,f13775,f13624,f13603,f13601,f13591,f13550,f13532,f13348,f13313,f13215,f13203,f13089,f13084,f12868,f12854,f12849,f12847,f12846,f12844,f12843,f12841,f12840,f12837,f12834,f12829,f12674,f12669,f12493,f12492,f12491,f12490,f12427,f12424,f12423,f12375,f12374,f12373,f12372,f12371,f12327,f12324,f12323,f12322,f12321,f12041,f12011,f12009,f12008,f11998,f11995,f11780,f11779,f11778,f11777,f11652,f11650,f11581,f11555,f11404,f11402,f11363,f11361,f11209,f11200,f11199,f11194,f11192,f9563,f9562,f9546,f9476,f9475,f9474,f9378,f9377,f9375,f9107,f9082,f8926,f8925,f8923,f8844,f8842,f8841,f8727,f8726,f8724,f8571,f8570,f8569,f8568,f8341,f8135,f8068,f8062,f8060,f7411,f7406,f7097,f7090,f6947,f6946,f6945,f6893,f6890,f6850,f6848,f6845,f6693,f6692,f6691,f6586,f6582,f6581,f6580,f6555,f6523,f6522,f6521,f6520,f6519,f6518,f6517,f6516,f6507,f6427,f6425,f6124,f5816,f5324,f5323,f5322,f5307,f5300,f5279,f5275,f5274,f5273,f5272,f5271,f5270,f5269,f5268,f5266,f5265,f5262,f5261,f5260,f5194,f5193,f5091,f5090,f5047,f4996,f4966,f4382,f4379,f4324,f4321,f4245,f4173,f4143,f4099,f3219,f3201,f3170,f3148,f3095,f3094,f3076,f2835,f2831,f2829,f2588,f2585,f2503,f2502,f2500,f2499,f2498,f2495,f2494,f2492,f2489,f2429,f2426,f2366,f2363,f2282,f2279,f2219,f2216,f2156,f2155,f2125,f2121,f2091,f2061,f2024,f1994,f1993,f1952,f1951,f1801,f1799,f1094,f1050,f957,f937,f934,f921,f612,f576,f542,f525,f524,f523,f522,f493,f453,f417,f264,f219,f187,f118,f103,f84,f83,f73,f12035,f12034,f12033,f11842,f11841,f8097,f8096,f6748,f6747,f6589,f5327,f74,f12802,f12642,f5147,f2480,f2417,f2270,f2207,f1157,f681,f642,f6770,f5146,f2354,f469,f468,f11961,f10086,f11962,f12275,f12274,f11553,f11552,f11188,f11187,f11119,f11118,f11010,f11009,f10928,f10927,f10674,f10673,f10120,f10095,f10093,f10090,f10074,f10073,f10018,f10017,f9332,f9314,f9190,f9189,f9150,f9149,f9081,f9080,f8058,f8057,f7988,f7987,f7199,f7180,f7177,f7163,f7162,f7108,f7107,f6583,f6579,f6568,f6567,f6564,f6563,f6554,f6549,f6547,f6546,f6545,f6544,f6543,f6515,f6058,f5321,f5310,f5309,f5306,f5305,f5299,f5294,f5292,f5291,f5290,f5289,f5288,f5267,f13719,f8832,f8681,f8680,f8679,f8080,f7247,f7246,f7176,f9851,f9850,f7389,f7387,f9473,f8898,f8897,f8896,f8895,f8678,f6944,f6690,f5192,f4995,f4172,f2023,f1944,f915,f450,f416,f16317,f16368,f16369,f20133,f20134,f20135,f20136,f20527,f20723,f20686,f20687,f20688,f20850,f20854,f20855,f20892,f20893,f20898,f20899,f20903,f20907,f20913,f20914,f20932,f20918,f20919,f20920,f20922,f20923,f21054,f21053,f21025,f21026,f21061,f21064,f21065,f21066,f21067,f21068,f21076,f21078,f21081,f21114,f21115,f21116,f21117,f21118,f21124,f21602])).
fof(f21601,plain,(
  $false | ($spl10 | $spl12)),
  inference(subsumption_resolution,[],[f21125,f71])).
fof(f21125,plain,(
  ~p(s(bool,'T_0')) | ($spl10 | $spl12)),
  inference(backward_demodulation,[],[f21124,f73])).
fof(f21120,plain,(
  $false | ($spl10 | $spl12)),
  inference(subsumption_resolution,[],[f21119,f21066])).
fof(f21119,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | ($spl10 | $spl12)),
  inference(forward_demodulation,[],[f21089,f21065])).
fof(f21089,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))) | ($spl10 | $spl12)),
  inference(backward_demodulation,[],[f21065,f8338])).
fof(f21063,plain,(
  s('type.realax.real',sK0) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))) | $spl84),
  inference(cnf_transformation,[],[f21063_D])).
fof(f21063_D,plain,(
  s('type.realax.real',sK0) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f3008,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl2),
  inference(cnf_transformation,[],[f3008_D])).
fof(f3008_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f21061,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s('type.realax.real',sK0) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0)))) ) | $spl12),
  inference(subsumption_resolution,[],[f21030,f249])).
fof(f21030,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))) | s(bool,'T_0') = s(bool,X0) | s('type.realax.real',sK0) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0)))) ) | $spl12),
  inference(superposition,[],[f20527,f20932])).
fof(f20900,plain,(
  ( ! [X37,X35] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X35))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X35))))) | s(bool,'F_0') = s(bool,X37) | s(bool,X37) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X35)))) ) | $spl82),
  inference(cnf_transformation,[],[f20900_D])).
fof(f20900_D,plain,(
  ( ! [X37,X35] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X35))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X35))))) | s(bool,'F_0') = s(bool,X37) | s(bool,X37) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X35)))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f3165,plain,(
  ( ! [X117] : (p(s(bool,X117))) ) | $spl24),
  inference(cnf_transformation,[],[f3165_D])).
fof(f3165_D,plain,(
  ( ! [X117] : (p(s(bool,X117))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f20899,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,X36)) | s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X35))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X35))))) | s(bool,X37) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X35))) | s(bool,'F_0') = s(bool,X37)) ) | $spl12),
  inference(subsumption_resolution,[],[f20875,f79])).
fof(f20875,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,X36)) | s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X35))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X35))))) | s(bool,X37) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X35))) | s(bool,'F_0') = s(bool,X37) | s(bool,'T_0') = s(bool,X36)) ) | $spl12),
  inference(superposition,[],[f6838,f1664])).
fof(f14575,plain,(
  $false | ($spl8 | $spl12)),
  inference(global_subsumption,[],[f70,f71,f76,f80,f79,f102,f107,f104,f108,f110,f131,f162,f207,f325,f326,f376,f558,f856,f575,f598,f611,f891,f651,f917,f923,f713,f715,f769,f792,f780,f996,f958,f1072,f1074,f1093,f1168,f1222,f1224,f1226,f1227,f1663,f1664,f1297,f1739,f1783,f2996,f2989,f2966,f3031,f2025,f2948,f2928,f2919,f2916,f2900,f2895,f3048,f2889,f2886,f3059,f2122,f2881,f2871,f3075,f2866,f2843,f2842,f2837,f2836,f3107,f2283,f2285,f3133,f2333,f3137,f3143,f3163,f3172,f3174,f3179,f3180,f2700,f3200,f3202,f103,f118,f187,f219,f264,f417,f456,f493,f523,f542,f576,f612,f921,f934,f937,f957,f1050,f1094,f1799,f1801,f1951,f1952,f1993,f1994,f2024,f2061,f2091,f2121,f2125,f2155,f2156,f2219,f2282,f2366,f2429,f2492,f2495,f2498,f2499,f2500,f2502,f2503,f2588,f2829,f2831,f2835,f3076,f3219,f1665,f938,f1795,f1740,f1223,f714,f263,f249,f111,f455,f454,f73,f83,f84,f453,f2494,f3148,f3201,f3095,f2023,f1944,f915,f450,f416,f4143,f4172,f4173,f4245,f4324,f4382,f4966,f5038,f5039,f4995,f4996,f5044,f5047,f5090,f5091,f5192,f5193,f5194,f5260,f5261,f5262,f5265,f5266,f5268,f5269,f5270,f5271,f5272,f5273,f5274,f5275,f5279,f5300,f5307,f5322,f5323,f5324,f5816,f6124,f6425,f6426,f6427,f6507,f6508,f6509,f6512,f6513,f6514,f6516,f6517,f6518,f6519,f6520,f6521,f6522,f6523,f6529,f6555,f6558,f6565,f6580,f6581,f6582,f6586,f6770,f6841,f6848,f6850,f7090,f6890,f6893,f6944,f6945,f6946,f6947,f7097,f7181,f7195,f7372,f7406,f7407,f7411,f690,f522,f7199,f7180,f7177,f7163,f7162,f7108,f7107,f6583,f6568,f6567,f6564,f6563,f6554,f6549,f6547,f6546,f6545,f6544,f6543,f6515,f6058,f5310,f5309,f5306,f5305,f5299,f5294,f5292,f5291,f5290,f5289,f5288,f5267,f7176,f7246,f7247,f7389,f7387,f5327,f6589,f6839,f6838,f6748,f6747,f5147,f5146,f3170,f3141,f2480,f2417,f2354,f2270,f2207,f1157,f681,f642,f469,f468,f258,f77,f74,f4379,f4321,f2585,f2489,f2426,f2363,f2279,f2216,f1294,f1165,f766,f687,f648,f525,f524,f7987,f7988,f8060,f8062,f8066,f8068,f8072,f8074,f8080,f8057,f8058,f8096,f8097,f8125,f8126,f8127,f8132,f8133,f8135,f8249,f8251,f8235,f8264,f8292,f8294,f8314,f8338,f8339,f8340,f8341,f8568,f8569,f8570,f8571,f8810,f8678,f8679,f8680,f8681,f8724,f8726,f8727,f8832,f8841,f8842,f8844,f8895,f8896,f8897,f8898,f8923,f8925,f8926,f9082,f9107,f9080,f9081,f9314,f9332,f9149,f9150,f9189,f9190,f9375,f9377,f9378,f9546,f9513,f9562,f9563,f9473,f9474,f9475,f9476,f9961,f9983,f9985,f9850,f9851,f10000,f10008,f10017,f10018,f10095,f10093,f10120,f10090,f10089,f10086,f10073,f10074,f10673,f10674,f11088,f11118,f11119,f10927,f10928,f11009,f11010,f11192,f11194,f11199,f11200,f11205,f11209,f11187,f11188,f11361,f11362,f11363,f11401,f11402,f11403,f11404,f11555,f11581,f11552,f11553,f11594,f11597,f11649,f11650,f11651,f11652,f11777,f11778,f11779,f11780,f11841,f11842,f11987,f11978,f11989,f11976,f11973,f11995,f11996,f11998,f11962,f12008,f12009,f11961,f12010,f12011,f12029,f12033,f12034,f12035,f12036,f12037,f12038,f12040,f12041,f13550,f13532,f12274,f12275,f13591,f12321,f12322,f12323,f12324,f12327,f13601,f12371,f12372,f12373,f12374,f12375,f12377,f13602,f13603,f12423,f12424,f12425,f12427,f12490,f12491,f12492,f12493,f13624,f12642,f13719,f12669,f12674,f13740,f13348,f13775,f13776,f12802,f12829,f12834,f12837,f12840,f12841,f12843,f12844,f12846,f12847,f12849,f12854,f12868,f13313,f13808,f13824,f13840,f13084,f13089,f13215,f13203,f3049,f13962,f13963,f14173,f14174,f14175,f14191,f14010,f14193,f14162,f14194,f14497,f14198,f14512,f14518,f14529,f14535,f14547,f14108,f14085,f14097])).
fof(f14097,plain,(
  ( ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0)))) ) | ($spl8 | $spl12)),
  inference(duplicate_literal_removal,[],[f14096])).
fof(f14096,plain,(
  ( ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0))) | s('type.realax.real',X0) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0)))) ) | ($spl8 | $spl12)),
  inference(backward_demodulation,[],[f13963,f12035])).
fof(f14085,plain,(
  ( ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0)))) ) | ($spl8 | $spl12)),
  inference(backward_demodulation,[],[f13963,f74])).
fof(f14108,plain,(
  ( ! [X3] : (s('type.realax.real',X3) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X3)))) ) | ($spl8 | $spl12)),
  inference(duplicate_literal_removal,[],[f14074])).
fof(f14074,plain,(
  ( ! [X3] : (s('type.realax.real',X3) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X3))) | s('type.realax.real',X3) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X3)))) ) | ($spl8 | $spl12)),
  inference(backward_demodulation,[],[f13963,f6770])).
fof(f14547,plain,(
  ( ! [X9] : (s('type.realax.real','const.real.abs_1'(s('type.realax.real',X9))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X9)))) ) | ($spl8 | $spl12)),
  inference(subsumption_resolution,[],[f14408,f14194])).
fof(f14408,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X9))) | s('type.realax.real','const.real.abs_1'(s('type.realax.real',X9))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X9)))) ) | ($spl8 | $spl12)),
  inference(backward_demodulation,[],[f14194,f11995])).
fof(f14535,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | ($spl8 | $spl12)),
  inference(condensation,[],[f14534])).
fof(f14534,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X1)) ) | ($spl8 | $spl12)),
  inference(condensation,[],[f14533])).
fof(f14533,plain,(
  ( ! [X70,X69,X67] : (s(bool,X69) = s(bool,X70) | s(bool,'F_0') = s(bool,X70) | s(bool,X67) = s(bool,X69)) ) | ($spl8 | $spl12)),
  inference(subsumption_resolution,[],[f14370,f14194])).
fof(f14370,plain,(
  ( ! [X70,X68,X69,X67] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X68),s(bool,X67))) | s(bool,X69) = s(bool,X70) | s(bool,'F_0') = s(bool,X70) | s(bool,X67) = s(bool,X69)) ) | ($spl8 | $spl12)),
  inference(backward_demodulation,[],[f14194,f9562])).
fof(f14529,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | ($spl8 | $spl12)),
  inference(condensation,[],[f14528])).
fof(f14528,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X1)) ) | ($spl8 | $spl12)),
  inference(duplicate_literal_removal,[],[f14527])).
fof(f14527,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X1)) ) | ($spl8 | $spl12)),
  inference(condensation,[],[f14526])).
fof(f14526,plain,(
  ( ! [X4,X0,X3,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X3) = s(bool,X4) | s(bool,'F_0') = s(bool,X4) | s(bool,X0) = s(bool,X3)) ) | ($spl8 | $spl12)),
  inference(subsumption_resolution,[],[f14354,f14194])).
fof(f14354,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X2))) | s(bool,'F_0') = s(bool,X1) | s(bool,X3) = s(bool,X4) | s(bool,'F_0') = s(bool,X4) | s(bool,X0) = s(bool,X3)) ) | ($spl8 | $spl12)),
  inference(backward_demodulation,[],[f14194,f8841])).
fof(f14518,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | ($spl8 | $spl12)),
  inference(condensation,[],[f14517])).
fof(f14517,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X1)) ) | ($spl8 | $spl12)),
  inference(duplicate_literal_removal,[],[f14516])).
fof(f14516,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X1)) ) | ($spl8 | $spl12)),
  inference(condensation,[],[f14515])).
fof(f14515,plain,(
  ( ! [X52,X56,X55,X53] : (s(bool,'F_0') = s(bool,X53) | s(bool,X55) = s(bool,X56) | s(bool,'F_0') = s(bool,X56) | s(bool,X52) = s(bool,X55)) ) | ($spl8 | $spl12)),
  inference(subsumption_resolution,[],[f14343,f14194])).
fof(f14343,plain,(
  ( ! [X54,X52,X56,X55,X53] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,X53),s(bool,X52),s(bool,X54))) | s(bool,'F_0') = s(bool,X53) | s(bool,X55) = s(bool,X56) | s(bool,'F_0') = s(bool,X56) | s(bool,X52) = s(bool,X55)) ) | ($spl8 | $spl12)),
  inference(backward_demodulation,[],[f14194,f8068])).
fof(f14512,plain,(
  ( ! [X7] : (s(bool,'F_0') = s(bool,X7)) ) | ($spl8 | $spl12)),
  inference(subsumption_resolution,[],[f14338,f14194])).
fof(f14338,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X6),s(bool,X7))) | s(bool,'F_0') = s(bool,X7)) ) | ($spl8 | $spl12)),
  inference(backward_demodulation,[],[f14194,f7097])).
fof(f14198,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) ) | ($spl8 | $spl12)),
  inference(backward_demodulation,[],[f14194,f103])).
fof(f14497,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl8 | $spl12)),
  inference(subsumption_resolution,[],[f14196,f14194])).
fof(f14196,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | ($spl8 | $spl12)),
  inference(backward_demodulation,[],[f14194,f83])).
fof(f14194,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl8 | $spl12)),
  inference(global_subsumption,[],[f70,f71,f76,f80,f79,f102,f107,f104,f108,f110,f131,f162,f207,f325,f326,f376,f558,f856,f575,f598,f611,f891,f651,f917,f923,f713,f715,f769,f792,f780,f996,f958,f1072,f1074,f1093,f1168,f1222,f1224,f1226,f1227,f1663,f1664,f1297,f1739,f1783,f2996,f2989,f2966,f3031,f2025,f2948,f2928,f2919,f2916,f2900,f2895,f3048,f2889,f2886,f3059,f2122,f2881,f2871,f3075,f2866,f2843,f2842,f2837,f2836,f3107,f2283,f2285,f3133,f2333,f3137,f3143,f3163,f3172,f3174,f3179,f3180,f2700,f3200,f3202,f103,f118,f187,f219,f264,f417,f456,f493,f523,f542,f576,f612,f921,f934,f937,f957,f1050,f1094,f1799,f1801,f1951,f1952,f1993,f1994,f2024,f2061,f2091,f2121,f2125,f2155,f2156,f2219,f2282,f2366,f2429,f2492,f2495,f2498,f2499,f2500,f2502,f2503,f2588,f2829,f2831,f2835,f3076,f3219,f1665,f938,f1795,f1740,f1223,f714,f263,f249,f111,f455,f454,f73,f83,f84,f453,f2494,f3148,f3201,f3095,f2023,f1944,f915,f450,f416,f4143,f4172,f4173,f4245,f4324,f4382,f4966,f5038,f5039,f4995,f4996,f5044,f5047,f5090,f5091,f5192,f5193,f5194,f5260,f5261,f5262,f5265,f5266,f5268,f5269,f5270,f5271,f5272,f5273,f5274,f5275,f5279,f5300,f5307,f5322,f5323,f5324,f5816,f6124,f6425,f6426,f6427,f6507,f6508,f6509,f6512,f6513,f6514,f6516,f6517,f6518,f6519,f6520,f6521,f6522,f6523,f6529,f6555,f6558,f6565,f6580,f6581,f6582,f6586,f6770,f6841,f6848,f6850,f7090,f6890,f6893,f6944,f6945,f6946,f6947,f7097,f7181,f7195,f7372,f7406,f7407,f7411,f690,f522,f7199,f7180,f7177,f7163,f7162,f7108,f7107,f6583,f6568,f6567,f6564,f6563,f6554,f6549,f6547,f6546,f6545,f6544,f6543,f6515,f6058,f5310,f5309,f5306,f5305,f5299,f5294,f5292,f5291,f5290,f5289,f5288,f5267,f7176,f7246,f7247,f7389,f7387,f5327,f6589,f6839,f6838,f6748,f6747,f5147,f5146,f3170,f3141,f2480,f2417,f2354,f2270,f2207,f1157,f681,f642,f469,f468,f258,f77,f74,f4379,f4321,f2585,f2489,f2426,f2363,f2279,f2216,f1294,f1165,f766,f687,f648,f525,f524,f7987,f7988,f8060,f8062,f8066,f8068,f8072,f8074,f8080,f8057,f8058,f8096,f8097,f8125,f8126,f8127,f8132,f8133,f8135,f8249,f8251,f8235,f8264,f8292,f8294,f8314,f8338,f8339,f8340,f8341,f8568,f8569,f8570,f8571,f8810,f8678,f8679,f8680,f8681,f8724,f8726,f8727,f8832,f8841,f8842,f8844,f8895,f8896,f8897,f8898,f8923,f8925,f8926,f9082,f9107,f9080,f9081,f9314,f9332,f9149,f9150,f9189,f9190,f9375,f9377,f9378,f9546,f9513,f9562,f9563,f9473,f9474,f9475,f9476,f9961,f9983,f9985,f9850,f9851,f10000,f10008,f10017,f10018,f10095,f10093,f10120,f10090,f10089,f10086,f10073,f10074,f10673,f10674,f11088,f11118,f11119,f10927,f10928,f11009,f11010,f11192,f11194,f11199,f11200,f11205,f11209,f11187,f11188,f11361,f11362,f11363,f11401,f11402,f11403,f11404,f11555,f11581,f11552,f11553,f11594,f11597,f11649,f11650,f11651,f11652,f11777,f11778,f11779,f11780,f11841,f11842,f11987,f11978,f11989,f11976,f11973,f11995,f11996,f11998,f11962,f12008,f12009,f11961,f12010,f12011,f12029,f12033,f12034,f12035,f12036,f12037,f12038,f12040,f12041,f13550,f13532,f12274,f12275,f13591,f12321,f12322,f12323,f12324,f12327,f13601,f12371,f12372,f12373,f12374,f12375,f12377,f13602,f13603,f12423,f12424,f12425,f12427,f12490,f12491,f12492,f12493,f13624,f12642,f13719,f12669,f12674,f13740,f13348,f13775,f13776,f12802,f12829,f12834,f12837,f12840,f12841,f12843,f12844,f12846,f12847,f12849,f12854,f12868,f13313,f13808,f13824,f13840,f13084,f13089,f13215,f13203,f3049,f13962,f13963,f14173,f14174,f14175,f14191,f14010,f14193,f14162])).
fof(f14162,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1))) ) | ($spl8 | $spl12)),
  inference(condensation,[],[f14011])).
fof(f14011,plain,(
  ( ! [X59,X57,X58] : (s(bool,'T_0') = s(bool,X58) | ~p(s(bool,X59)) | s(bool,'T_0') = s(bool,X57)) ) | ($spl8 | $spl12)),
  inference(backward_demodulation,[],[f13963,f6544])).
fof(f14010,plain,(
  ( ! [X54,X56,X55] : (p(s(bool,X55)) | s(bool,'T_0') != s(bool,X56) | s(bool,'T_0') = s(bool,X54)) ) | ($spl8 | $spl12)),
  inference(backward_demodulation,[],[f13963,f6543])).
fof(f14191,plain,(
  ( ! [X127,X125,X128,X124] : (s(bool,X125) != s(bool,X127) | p(s(bool,X125)) | s(bool,'T_0') = s(bool,X128) | s(bool,'F_0') = s(bool,X124)) ) | ($spl8 | $spl12)),
  inference(subsumption_resolution,[],[f14006,f13962])).
fof(f14006,plain,(
  ( ! [X127,X125,X128,X124] : (s(bool,X125) != s(bool,X127) | p(s(bool,X125)) | s(bool,'T_0') = s(bool,X128) | s(bool,X127) = s(bool,X128) | s(bool,'F_0') = s(bool,X124)) ) | ($spl8 | $spl12)),
  inference(backward_demodulation,[],[f13963,f5310])).
fof(f14174,plain,(
  ( ! [X165,X161,X162] : (p(s(bool,X162)) | s(bool,'T_0') = s(bool,X165) | s(bool,'F_0') = s(bool,X161)) ) | ($spl8 | $spl12)),
  inference(subsumption_resolution,[],[f13967,f13962])).
fof(f13967,plain,(
  ( ! [X165,X163,X161,X162] : (p(s(bool,X162)) | s(bool,'T_0') = s(bool,X165) | s(bool,X165) = s(bool,'const.bool.COND_3'(s(bool,X161),s(bool,X162),s(bool,X163))) | s(bool,'F_0') = s(bool,X161)) ) | ($spl8 | $spl12)),
  inference(backward_demodulation,[],[f13963,f9851])).
fof(f13963,plain,(
  ( ! [X149,X147,X145,X148] : (s(X147,X148) = s(X147,'const.bool.COND_3'(s(bool,X145),s(X147,X148),s(X147,X149)))) ) | ($spl8 | $spl12)),
  inference(subsumption_resolution,[],[f13936,f110])).
fof(f13936,plain,(
  ( ! [X149,X147,X145,X148] : (s(bool,'T_0') = s(bool,'F_0') | s(X147,X148) = s(X147,'const.bool.COND_3'(s(bool,X145),s(X147,X148),s(X147,X149)))) ) | ($spl8 | $spl12)),
  inference(backward_demodulation,[],[f3049,f5192])).
fof(f13962,plain,(
  ( ! [X91,X89] : (s(bool,X89) != s(bool,X91) | s(bool,'T_0') = s(bool,X91)) ) | $spl8),
  inference(subsumption_resolution,[],[f13935,f110])).
fof(f13935,plain,(
  ( ! [X91,X89] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X89) != s(bool,X91) | s(bool,'T_0') = s(bool,X91)) ) | $spl8),
  inference(backward_demodulation,[],[f3049,f4995])).
fof(f14572,plain,(
  $false | ($spl8 | $spl12)),
  inference(global_subsumption,[],[f70,f71,f76,f80,f79,f102,f107,f104,f108,f110,f131,f162,f207,f325,f326,f376,f558,f856,f575,f598,f611,f891,f651,f917,f923,f713,f715,f769,f792,f780,f996,f958,f1072,f1074,f1093,f1168,f1222,f1224,f1226,f1227,f1663,f1664,f1297,f1739,f1783,f2996,f2989,f2966,f3031,f2025,f2948,f2928,f2919,f2916,f2900,f2895,f3048,f2889,f2886,f3059,f2122,f2881,f2871,f3075,f2866,f2843,f2842,f2837,f2836,f3107,f2283,f2285,f3133,f2333,f3137,f3143,f3163,f3172,f3174,f3179,f3180,f2700,f3200,f3202,f103,f118,f187,f219,f264,f417,f456,f493,f523,f542,f576,f612,f921,f934,f937,f957,f1050,f1094,f1799,f1801,f1951,f1952,f1993,f1994,f2024,f2061,f2091,f2121,f2125,f2155,f2156,f2219,f2282,f2366,f2429,f2492,f2495,f2498,f2499,f2500,f2502,f2503,f2588,f2829,f2831,f2835,f3076,f3219,f1665,f938,f1795,f1740,f1223,f714,f263,f249,f111,f455,f454,f73,f83,f84,f453,f2494,f3148,f3201,f3095,f2023,f1944,f915,f450,f416,f4143,f4172,f4173,f4245,f4324,f4382,f4966,f5038,f5039,f4995,f4996,f5044,f5047,f5090,f5091,f5192,f5193,f5194,f5260,f5261,f5262,f5265,f5266,f5268,f5269,f5270,f5271,f5272,f5273,f5274,f5275,f5279,f5300,f5307,f5322,f5323,f5324,f5816,f6124,f6425,f6426,f6427,f6507,f6508,f6509,f6512,f6513,f6514,f6516,f6517,f6518,f6519,f6520,f6521,f6522,f6523,f6529,f6555,f6558,f6565,f6580,f6581,f6582,f6586,f6770,f6841,f6848,f6850,f7090,f6890,f6893,f6944,f6945,f6946,f6947,f7097,f7181,f7195,f7372,f7406,f7407,f7411,f690,f522,f7199,f7180,f7177,f7163,f7162,f7108,f7107,f6583,f6568,f6567,f6564,f6563,f6554,f6549,f6547,f6546,f6545,f6544,f6543,f6515,f6058,f5310,f5309,f5306,f5305,f5299,f5294,f5292,f5291,f5290,f5289,f5288,f5267,f7176,f7246,f7247,f7389,f7387,f5327,f6589,f6839,f6838,f6748,f6747,f5147,f5146,f3170,f3141,f2480,f2417,f2354,f2270,f2207,f1157,f681,f642,f469,f468,f258,f77,f74,f4379,f4321,f2585,f2489,f2426,f2363,f2279,f2216,f1294,f1165,f766,f687,f648,f525,f524,f7987,f7988,f8060,f8062,f8066,f8068,f8072,f8074,f8080,f8057,f8058,f8096,f8097,f8125,f8126,f8127,f8132,f8133,f8135,f8249,f8251,f8235,f8264,f8292,f8294,f8314,f8338,f8339,f8340,f8341,f8568,f8569,f8570,f8571,f8810,f8678,f8679,f8680,f8681,f8724,f8726,f8727,f8832,f8841,f8842,f8844,f8895,f8896,f8897,f8898,f8923,f8925,f8926,f9082,f9107,f9080,f9081,f9314,f9332,f9149,f9150,f9189,f9190,f9375,f9377,f9378,f9546,f9513,f9562,f9563,f9473,f9474,f9475,f9476,f9961,f9983,f9985,f9850,f9851,f10000,f10008,f10017,f10018,f10095,f10093,f10120,f10090,f10089,f10086,f10073,f10074,f10673,f10674,f11088,f11118,f11119,f10927,f10928,f11009,f11010,f11192,f11194,f11199,f11200,f11205,f11209,f11187,f11188,f11361,f11362,f11363,f11401,f11402,f11403,f11404,f11555,f11581,f11552,f11553,f11594,f11597,f11649,f11650,f11651,f11652,f11777,f11778,f11779,f11780,f11841,f11842,f11987,f11978,f11989,f11976,f11973,f11995,f11996,f11998,f11962,f12008,f12009,f11961,f12010,f12011,f12029,f12033,f12034,f12035,f12036,f12037,f12038,f12040,f12041,f13550,f13532,f12274,f12275,f13591,f12321,f12322,f12323,f12324,f12327,f13601,f12371,f12372,f12373,f12374,f12375,f12377,f13602,f13603,f12423,f12424,f12425,f12427,f12490,f12491,f12492,f12493,f13624,f12642,f13719,f12669,f12674,f13740,f13348,f13775,f13776,f12802,f12829,f12834,f12837,f12840,f12841,f12843,f12844,f12846,f12847,f12849,f12854,f12868,f13313,f13808,f13824,f13840,f13084,f13089,f13215,f13203,f3049,f13962,f13963,f14173,f14174,f14175,f14191,f14010,f14193,f14162,f14194,f14497,f14198,f14512,f14518,f14529,f14535,f14547,f14108,f14085])).
fof(f14570,plain,(
  $false | ($spl8 | $spl12)),
  inference(global_subsumption,[],[f70,f71,f76,f80,f79,f102,f107,f104,f108,f110,f131,f162,f207,f325,f326,f376,f558,f856,f575,f598,f611,f891,f651,f917,f923,f713,f715,f769,f792,f780,f996,f958,f1072,f1074,f1093,f1168,f1222,f1224,f1226,f1227,f1663,f1664,f1297,f1739,f1783,f2996,f2989,f2966,f3031,f2025,f2948,f2928,f2919,f2916,f2900,f2895,f3048,f2889,f2886,f3059,f2122,f2881,f2871,f3075,f2866,f2843,f2842,f2837,f2836,f3107,f2283,f2285,f3133,f2333,f3137,f3143,f3163,f3172,f3174,f3179,f3180,f2700,f3200,f3202,f103,f118,f187,f219,f264,f417,f456,f493,f523,f542,f576,f612,f921,f934,f937,f957,f1050,f1094,f1799,f1801,f1951,f1952,f1993,f1994,f2024,f2061,f2091,f2121,f2125,f2155,f2156,f2219,f2282,f2366,f2429,f2492,f2495,f2498,f2499,f2500,f2502,f2503,f2588,f2829,f2831,f2835,f3076,f3219,f1665,f938,f1795,f1740,f1223,f714,f263,f249,f111,f455,f454,f73,f83,f84,f453,f2494,f3148,f3201,f3095,f2023,f1944,f915,f450,f416,f4143,f4172,f4173,f4245,f4324,f4382,f4966,f5038,f5039,f4995,f4996,f5044,f5047,f5090,f5091,f5192,f5193,f5194,f5260,f5261,f5262,f5265,f5266,f5268,f5269,f5270,f5271,f5272,f5273,f5274,f5275,f5279,f5300,f5307,f5322,f5323,f5324,f5816,f6124,f6425,f6426,f6427,f6507,f6508,f6509,f6512,f6513,f6514,f6516,f6517,f6518,f6519,f6520,f6521,f6522,f6523,f6529,f6555,f6558,f6565,f6580,f6581,f6582,f6586,f6770,f6841,f6848,f6850,f7090,f6890,f6893,f6944,f6945,f6946,f6947,f7097,f7181,f7195,f7372,f7406,f7407,f7411,f690,f522,f7199,f7180,f7177,f7163,f7162,f7108,f7107,f6583,f6568,f6567,f6564,f6563,f6554,f6549,f6547,f6546,f6545,f6544,f6543,f6515,f6058,f5310,f5309,f5306,f5305,f5299,f5294,f5292,f5291,f5290,f5289,f5288,f5267,f7176,f7246,f7247,f7389,f7387,f5327,f6589,f6839,f6838,f6748,f6747,f5147,f5146,f3170,f3141,f2480,f2417,f2354,f2270,f2207,f1157,f681,f642,f469,f468,f258,f77,f74,f4379,f4321,f2585,f2489,f2426,f2363,f2279,f2216,f1294,f1165,f766,f687,f648,f525,f524,f7987,f7988,f8060,f8062,f8066,f8068,f8072,f8074,f8080,f8057,f8058,f8096,f8097,f8125,f8126,f8127,f8132,f8133,f8135,f8249,f8251,f8235,f8264,f8292,f8294,f8314,f8338,f8339,f8340,f8341,f8568,f8569,f8570,f8571,f8810,f8678,f8679,f8680,f8681,f8724,f8726,f8727,f8832,f8841,f8842,f8844,f8895,f8896,f8897,f8898,f8923,f8925,f8926,f9082,f9107,f9080,f9081,f9314,f9332,f9149,f9150,f9189,f9190,f9375,f9377,f9378,f9546,f9513,f9562,f9563,f9473,f9474,f9475,f9476,f9961,f9983,f9985,f9850,f9851,f10000,f10008,f10017,f10018,f10095,f10093,f10120,f10090,f10089,f10086,f10073,f10074,f10673,f10674,f11088,f11118,f11119,f10927,f10928,f11009,f11010,f11192,f11194,f11199,f11200,f11205,f11209,f11187,f11188,f11361,f11362,f11363,f11401,f11402,f11403,f11404,f11555,f11581,f11552,f11553,f11594,f11597,f11649,f11650,f11651,f11652,f11777,f11778,f11779,f11780,f11841,f11842,f11987,f11978,f11989,f11976,f11973,f11995,f11996,f11998,f11962,f12008,f12009,f11961,f12010,f12011,f12029,f12033,f12034,f12035,f12036,f12037,f12038,f12040,f12041,f13550,f13532,f12274,f12275,f13591,f12321,f12322,f12323,f12324,f12327,f13601,f12371,f12372,f12373,f12374,f12375,f12377,f13602,f13603,f12423,f12424,f12425,f12427,f12490,f12491,f12492,f12493,f13624,f12642,f13719,f12669,f12674,f13740,f13348,f13775,f13776,f12802,f12829,f12834,f12837,f12840,f12841,f12843,f12844,f12846,f12847,f12849,f12854,f12868,f13313,f13808,f13824,f13840,f13084,f13089,f13215,f13203,f3049,f13962,f13963,f14173,f14174,f14175,f14191,f14010,f14193,f14162,f14194,f14497,f14198,f14512,f14518,f14529,f14535,f14547,f14108])).
fof(f14548,plain,(
  $false | ($spl8 | $spl12)),
  inference(global_subsumption,[],[f70,f71,f76,f80,f79,f102,f107,f104,f108,f110,f131,f162,f207,f325,f326,f376,f558,f856,f575,f598,f611,f891,f651,f917,f923,f713,f715,f769,f792,f780,f996,f958,f1072,f1074,f1093,f1168,f1222,f1224,f1226,f1227,f1663,f1664,f1297,f1739,f1783,f2996,f2989,f2966,f3031,f2025,f2948,f2928,f2919,f2916,f2900,f2895,f3048,f2889,f2886,f3059,f2122,f2881,f2871,f3075,f2866,f2843,f2842,f2837,f2836,f3107,f2283,f2285,f3133,f2333,f3137,f3143,f3163,f3172,f3174,f3179,f3180,f2700,f3200,f3202,f103,f118,f187,f219,f264,f417,f456,f493,f523,f542,f576,f612,f921,f934,f937,f957,f1050,f1094,f1799,f1801,f1951,f1952,f1993,f1994,f2024,f2061,f2091,f2121,f2125,f2155,f2156,f2219,f2282,f2366,f2429,f2492,f2495,f2498,f2499,f2500,f2502,f2503,f2588,f2829,f2831,f2835,f3076,f3219,f1665,f938,f1795,f1740,f1223,f714,f263,f249,f111,f455,f454,f73,f83,f84,f453,f2494,f3148,f3201,f3095,f2023,f1944,f915,f450,f416,f4143,f4172,f4173,f4245,f4324,f4382,f4966,f5038,f5039,f4995,f4996,f5044,f5047,f5090,f5091,f5192,f5193,f5194,f5260,f5261,f5262,f5265,f5266,f5268,f5269,f5270,f5271,f5272,f5273,f5274,f5275,f5279,f5300,f5307,f5322,f5323,f5324,f5816,f6124,f6425,f6426,f6427,f6507,f6508,f6509,f6512,f6513,f6514,f6516,f6517,f6518,f6519,f6520,f6521,f6522,f6523,f6529,f6555,f6558,f6565,f6580,f6581,f6582,f6586,f6770,f6841,f6848,f6850,f7090,f6890,f6893,f6944,f6945,f6946,f6947,f7097,f7181,f7195,f7372,f7406,f7407,f7411,f690,f522,f7199,f7180,f7177,f7163,f7162,f7108,f7107,f6583,f6568,f6567,f6564,f6563,f6554,f6549,f6547,f6546,f6545,f6544,f6543,f6515,f6058,f5310,f5309,f5306,f5305,f5299,f5294,f5292,f5291,f5290,f5289,f5288,f5267,f7176,f7246,f7247,f7389,f7387,f5327,f6589,f6839,f6838,f6748,f6747,f5147,f5146,f3170,f3141,f2480,f2417,f2354,f2270,f2207,f1157,f681,f642,f469,f468,f258,f77,f74,f4379,f4321,f2585,f2489,f2426,f2363,f2279,f2216,f1294,f1165,f766,f687,f648,f525,f524,f7987,f7988,f8060,f8062,f8066,f8068,f8072,f8074,f8080,f8057,f8058,f8096,f8097,f8125,f8126,f8127,f8132,f8133,f8135,f8249,f8251,f8235,f8264,f8292,f8294,f8314,f8338,f8339,f8340,f8341,f8568,f8569,f8570,f8571,f8810,f8678,f8679,f8680,f8681,f8724,f8726,f8727,f8832,f8841,f8842,f8844,f8895,f8896,f8897,f8898,f8923,f8925,f8926,f9082,f9107,f9080,f9081,f9314,f9332,f9149,f9150,f9189,f9190,f9375,f9377,f9378,f9546,f9513,f9562,f9563,f9473,f9474,f9475,f9476,f9961,f9983,f9985,f9850,f9851,f10000,f10008,f10017,f10018,f10095,f10093,f10120,f10090,f10089,f10086,f10073,f10074,f10673,f10674,f11088,f11118,f11119,f10927,f10928,f11009,f11010,f11192,f11194,f11199,f11200,f11205,f11209,f11187,f11188,f11361,f11362,f11363,f11401,f11402,f11403,f11404,f11555,f11581,f11552,f11553,f11594,f11597,f11649,f11650,f11651,f11652,f11777,f11778,f11779,f11780,f11841,f11842,f11987,f11978,f11989,f11976,f11973,f11995,f11996,f11998,f11962,f12008,f12009,f11961,f12010,f12011,f12029,f12033,f12034,f12035,f12036,f12037,f12038,f12040,f12041,f13550,f13532,f12274,f12275,f13591,f12321,f12322,f12323,f12324,f12327,f13601,f12371,f12372,f12373,f12374,f12375,f12377,f13602,f13603,f12423,f12424,f12425,f12427,f12490,f12491,f12492,f12493,f13624,f12642,f13719,f12669,f12674,f13740,f13348,f13775,f13776,f12802,f12829,f12834,f12837,f12840,f12841,f12843,f12844,f12846,f12847,f12849,f12854,f12868,f13313,f13808,f13824,f13840,f13084,f13089,f13215,f13203,f3049,f13962,f13963,f14173,f14174,f14175,f14191,f14010,f14193,f14162,f14194,f14497,f14198,f14512,f14518,f14529,f14535,f14547])).
fof(f14538,plain,(
  $false | ($spl8 | $spl12 | $spl68)),
  inference(global_subsumption,[],[f14537,f70,f71,f76,f80,f79,f102,f107,f104,f108,f110,f131,f162,f207,f325,f326,f376,f558,f856,f575,f598,f611,f891,f651,f917,f923,f713,f715,f769,f792,f780,f996,f958,f1072,f1074,f1093,f1168,f1222,f1224,f1226,f1227,f1663,f1664,f1297,f1739,f1783,f2996,f2989,f2966,f3031,f2025,f2948,f2928,f2919,f2916,f2900,f2895,f3048,f2889,f2886,f3059,f2122,f2881,f2871,f3075,f2866,f2843,f2842,f2837,f2836,f3107,f2283,f2285,f3133,f2333,f3137,f3143,f3163,f3172,f3174,f3179,f3180,f2700,f3200,f3202,f103,f118,f187,f219,f264,f417,f456,f493,f523,f542,f576,f612,f921,f934,f937,f957,f1050,f1094,f1799,f1801,f1951,f1952,f1993,f1994,f2024,f2061,f2091,f2121,f2125,f2155,f2156,f2219,f2282,f2366,f2429,f2492,f2495,f2498,f2499,f2500,f2502,f2503,f2588,f2829,f2831,f2835,f3076,f3219,f1665,f938,f1795,f1740,f1223,f714,f263,f249,f111,f455,f454,f73,f83,f84,f453,f2494,f3148,f3201,f3095,f2023,f1944,f915,f450,f416,f4143,f4172,f4173,f4245,f4324,f4382,f4966,f5038,f5039,f4995,f4996,f5044,f5047,f5090,f5091,f5192,f5193,f5194,f5260,f5261,f5262,f5265,f5266,f5268,f5269,f5270,f5271,f5272,f5273,f5274,f5275,f5279,f5300,f5307,f5322,f5323,f5324,f5816,f6124,f6425,f6426,f6427,f6507,f6508,f6509,f6512,f6513,f6514,f6516,f6517,f6518,f6519,f6520,f6521,f6522,f6523,f6529,f6555,f6558,f6565,f6580,f6581,f6582,f6586,f6770,f6841,f6848,f6850,f7090,f6890,f6893,f6944,f6945,f6946,f6947,f7097,f7181,f7195,f7372,f7406,f7407,f7411,f690,f522,f7199,f7180,f7177,f7163,f7162,f7108,f7107,f6583,f6568,f6567,f6564,f6563,f6554,f6549,f6547,f6546,f6545,f6544,f6543,f6515,f6058,f5310,f5309,f5306,f5305,f5299,f5294,f5292,f5291,f5290,f5289,f5288,f5267,f7176,f7246,f7247,f7389,f7387,f5327,f6589,f6839,f6838,f6748,f6747,f5147,f5146,f3170,f3141,f2480,f2417,f2354,f2270,f2207,f1157,f681,f642,f469,f468,f258,f77,f74,f4379,f4321,f2585,f2489,f2426,f2363,f2279,f2216,f1294,f1165,f766,f687,f648,f525,f524,f7987,f7988,f8060,f8062,f8066,f8068,f8072,f8074,f8080,f8057,f8058,f8096,f8097,f8125,f8126,f8127,f8132,f8133,f8135,f8249,f8251,f8235,f8264,f8292,f8294,f8314,f8338,f8339,f8340,f8341,f8568,f8569,f8570,f8571,f8810,f8678,f8679,f8680,f8681,f8724,f8726,f8727,f8832,f8841,f8842,f8844,f8895,f8896,f8897,f8898,f8923,f8925,f8926,f9082,f9107,f9080,f9081,f9314,f9332,f9149,f9150,f9189,f9190,f9375,f9377,f9378,f9546,f9513,f9562,f9563,f9473,f9474,f9475,f9476,f9961,f9983,f9985,f9850,f9851,f10000,f10008,f10017,f10018,f10095,f10093,f10120,f10090,f10089,f10086,f10073,f10074,f10673,f10674,f11088,f11118,f11119,f10927,f10928,f11009,f11010,f11192,f11194,f11199,f11200,f11205,f11209,f11187,f11188,f11361,f11362,f11363,f11401,f11402,f11403,f11404,f11555,f11581,f11552,f11553,f11594,f11597,f11649,f11650,f11651,f11652,f11777,f11778,f11779,f11780,f11841,f11842,f11987,f11978,f11989,f11976,f11973,f11995,f11996,f11998,f11962,f12008,f12009,f11961,f12010,f12011,f12029,f12033,f12034,f12035,f12036,f12037,f12038,f12040,f12041,f13550,f13532,f12274,f12275,f13591,f12321,f12322,f12323,f12324,f12327,f13601,f12371,f12372,f12373,f12374,f12375,f12377,f13602,f13603,f12423,f12424,f12425,f12427,f12490,f12491,f12492,f12493,f13624,f12642,f13719,f12669,f12674,f13740,f13348,f13775,f13776,f12802,f12829,f12834,f12837,f12840,f12841,f12843,f12844,f12846,f12847,f12849,f12854,f12868,f13313,f13808,f13824,f13840,f13084,f13089,f13215,f13203,f3049,f13962,f13963,f14173,f14174,f14175,f14191,f14010,f14193,f14162,f14194,f14497,f14198,f14512,f14518,f14529,f14535])).
fof(f14537,plain,(
  ( ! [X26] : (s(bool,'F_0') = s(bool,X26)) ) | ($spl8 | $spl12 | $spl68)),
  inference(subsumption_resolution,[],[f14380,f14194])).
fof(f14380,plain,(
  ( ! [X26,X27] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,X26),s(bool,X26),s(bool,X27))) | s(bool,'F_0') = s(bool,X26)) ) | ($spl8 | $spl12 | $spl68)),
  inference(backward_demodulation,[],[f14194,f10603])).
fof(f14536,plain,(
  $false | ($spl8 | $spl12)),
  inference(global_subsumption,[],[f70,f71,f76,f80,f79,f102,f107,f104,f108,f110,f131,f162,f207,f325,f326,f376,f558,f856,f575,f598,f611,f891,f651,f917,f923,f713,f715,f769,f792,f780,f996,f958,f1072,f1074,f1093,f1168,f1222,f1224,f1226,f1227,f1663,f1664,f1297,f1739,f1783,f2996,f2989,f2966,f3031,f2025,f2948,f2928,f2919,f2916,f2900,f2895,f3048,f2889,f2886,f3059,f2122,f2881,f2871,f3075,f2866,f2843,f2842,f2837,f2836,f3107,f2283,f2285,f3133,f2333,f3137,f3143,f3163,f3172,f3174,f3179,f3180,f2700,f3200,f3202,f103,f118,f187,f219,f264,f417,f456,f493,f523,f542,f576,f612,f921,f934,f937,f957,f1050,f1094,f1799,f1801,f1951,f1952,f1993,f1994,f2024,f2061,f2091,f2121,f2125,f2155,f2156,f2219,f2282,f2366,f2429,f2492,f2495,f2498,f2499,f2500,f2502,f2503,f2588,f2829,f2831,f2835,f3076,f3219,f1665,f938,f1795,f1740,f1223,f714,f263,f249,f111,f455,f454,f73,f83,f84,f453,f2494,f3148,f3201,f3095,f2023,f1944,f915,f450,f416,f4143,f4172,f4173,f4245,f4324,f4382,f4966,f5038,f5039,f4995,f4996,f5044,f5047,f5090,f5091,f5192,f5193,f5194,f5260,f5261,f5262,f5265,f5266,f5268,f5269,f5270,f5271,f5272,f5273,f5274,f5275,f5279,f5300,f5307,f5322,f5323,f5324,f5816,f6124,f6425,f6426,f6427,f6507,f6508,f6509,f6512,f6513,f6514,f6516,f6517,f6518,f6519,f6520,f6521,f6522,f6523,f6529,f6555,f6558,f6565,f6580,f6581,f6582,f6586,f6770,f6841,f6848,f6850,f7090,f6890,f6893,f6944,f6945,f6946,f6947,f7097,f7181,f7195,f7372,f7406,f7407,f7411,f690,f522,f7199,f7180,f7177,f7163,f7162,f7108,f7107,f6583,f6568,f6567,f6564,f6563,f6554,f6549,f6547,f6546,f6545,f6544,f6543,f6515,f6058,f5310,f5309,f5306,f5305,f5299,f5294,f5292,f5291,f5290,f5289,f5288,f5267,f7176,f7246,f7247,f7389,f7387,f5327,f6589,f6839,f6838,f6748,f6747,f5147,f5146,f3170,f3141,f2480,f2417,f2354,f2270,f2207,f1157,f681,f642,f469,f468,f258,f77,f74,f4379,f4321,f2585,f2489,f2426,f2363,f2279,f2216,f1294,f1165,f766,f687,f648,f525,f524,f7987,f7988,f8060,f8062,f8066,f8068,f8072,f8074,f8080,f8057,f8058,f8096,f8097,f8125,f8126,f8127,f8132,f8133,f8135,f8249,f8251,f8235,f8264,f8292,f8294,f8314,f8338,f8339,f8340,f8341,f8568,f8569,f8570,f8571,f8810,f8678,f8679,f8680,f8681,f8724,f8726,f8727,f8832,f8841,f8842,f8844,f8895,f8896,f8897,f8898,f8923,f8925,f8926,f9082,f9107,f9080,f9081,f9314,f9332,f9149,f9150,f9189,f9190,f9375,f9377,f9378,f9546,f9513,f9562,f9563,f9473,f9474,f9475,f9476,f9961,f9983,f9985,f9850,f9851,f10000,f10008,f10017,f10018,f10095,f10093,f10120,f10090,f10089,f10086,f10073,f10074,f10673,f10674,f11088,f11118,f11119,f10927,f10928,f11009,f11010,f11192,f11194,f11199,f11200,f11205,f11209,f11187,f11188,f11361,f11362,f11363,f11401,f11402,f11403,f11404,f11555,f11581,f11552,f11553,f11594,f11597,f11649,f11650,f11651,f11652,f11777,f11778,f11779,f11780,f11841,f11842,f11987,f11978,f11989,f11976,f11973,f11995,f11996,f11998,f11962,f12008,f12009,f11961,f12010,f12011,f12029,f12033,f12034,f12035,f12036,f12037,f12038,f12040,f12041,f13550,f13532,f12274,f12275,f13591,f12321,f12322,f12323,f12324,f12327,f13601,f12371,f12372,f12373,f12374,f12375,f12377,f13602,f13603,f12423,f12424,f12425,f12427,f12490,f12491,f12492,f12493,f13624,f12642,f13719,f12669,f12674,f13740,f13348,f13775,f13776,f12802,f12829,f12834,f12837,f12840,f12841,f12843,f12844,f12846,f12847,f12849,f12854,f12868,f13313,f13808,f13824,f13840,f13084,f13089,f13215,f13203,f3049,f13962,f13963,f14173,f14174,f14175,f14191,f14010,f14193,f14162,f14194,f14497,f14198,f14512,f14518,f14529,f14535])).
fof(f14530,plain,(
  $false | ($spl8 | $spl12)),
  inference(global_subsumption,[],[f70,f71,f76,f80,f79,f102,f107,f104,f108,f110,f131,f162,f207,f325,f326,f376,f558,f856,f575,f598,f611,f891,f651,f917,f923,f713,f715,f769,f792,f780,f996,f958,f1072,f1074,f1093,f1168,f1222,f1224,f1226,f1227,f1663,f1664,f1297,f1739,f1783,f2996,f2989,f2966,f3031,f2025,f2948,f2928,f2919,f2916,f2900,f2895,f3048,f2889,f2886,f3059,f2122,f2881,f2871,f3075,f2866,f2843,f2842,f2837,f2836,f3107,f2283,f2285,f3133,f2333,f3137,f3143,f3163,f3172,f3174,f3179,f3180,f2700,f3200,f3202,f103,f118,f187,f219,f264,f417,f456,f493,f523,f542,f576,f612,f921,f934,f937,f957,f1050,f1094,f1799,f1801,f1951,f1952,f1993,f1994,f2024,f2061,f2091,f2121,f2125,f2155,f2156,f2219,f2282,f2366,f2429,f2492,f2495,f2498,f2499,f2500,f2502,f2503,f2588,f2829,f2831,f2835,f3076,f3219,f1665,f938,f1795,f1740,f1223,f714,f263,f249,f111,f455,f454,f73,f83,f84,f453,f2494,f3148,f3201,f3095,f2023,f1944,f915,f450,f416,f4143,f4172,f4173,f4245,f4324,f4382,f4966,f5038,f5039,f4995,f4996,f5044,f5047,f5090,f5091,f5192,f5193,f5194,f5260,f5261,f5262,f5265,f5266,f5268,f5269,f5270,f5271,f5272,f5273,f5274,f5275,f5279,f5300,f5307,f5322,f5323,f5324,f5816,f6124,f6425,f6426,f6427,f6507,f6508,f6509,f6512,f6513,f6514,f6516,f6517,f6518,f6519,f6520,f6521,f6522,f6523,f6529,f6555,f6558,f6565,f6580,f6581,f6582,f6586,f6770,f6841,f6848,f6850,f7090,f6890,f6893,f6944,f6945,f6946,f6947,f7097,f7181,f7195,f7372,f7406,f7407,f7411,f690,f522,f7199,f7180,f7177,f7163,f7162,f7108,f7107,f6583,f6568,f6567,f6564,f6563,f6554,f6549,f6547,f6546,f6545,f6544,f6543,f6515,f6058,f5310,f5309,f5306,f5305,f5299,f5294,f5292,f5291,f5290,f5289,f5288,f5267,f7176,f7246,f7247,f7389,f7387,f5327,f6589,f6839,f6838,f6748,f6747,f5147,f5146,f3170,f3141,f2480,f2417,f2354,f2270,f2207,f1157,f681,f642,f469,f468,f258,f77,f74,f4379,f4321,f2585,f2489,f2426,f2363,f2279,f2216,f1294,f1165,f766,f687,f648,f525,f524,f7987,f7988,f8060,f8062,f8066,f8068,f8072,f8074,f8080,f8057,f8058,f8096,f8097,f8125,f8126,f8127,f8132,f8133,f8135,f8249,f8251,f8235,f8264,f8292,f8294,f8314,f8338,f8339,f8340,f8341,f8568,f8569,f8570,f8571,f8810,f8678,f8679,f8680,f8681,f8724,f8726,f8727,f8832,f8841,f8842,f8844,f8895,f8896,f8897,f8898,f8923,f8925,f8926,f9082,f9107,f9080,f9081,f9314,f9332,f9149,f9150,f9189,f9190,f9375,f9377,f9378,f9546,f9513,f9562,f9563,f9473,f9474,f9475,f9476,f9961,f9983,f9985,f9850,f9851,f10000,f10008,f10017,f10018,f10095,f10093,f10120,f10090,f10089,f10086,f10073,f10074,f10673,f10674,f11088,f11118,f11119,f10927,f10928,f11009,f11010,f11192,f11194,f11199,f11200,f11205,f11209,f11187,f11188,f11361,f11362,f11363,f11401,f11402,f11403,f11404,f11555,f11581,f11552,f11553,f11594,f11597,f11649,f11650,f11651,f11652,f11777,f11778,f11779,f11780,f11841,f11842,f11987,f11978,f11989,f11976,f11973,f11995,f11996,f11998,f11962,f12008,f12009,f11961,f12010,f12011,f12029,f12033,f12034,f12035,f12036,f12037,f12038,f12040,f12041,f13550,f13532,f12274,f12275,f13591,f12321,f12322,f12323,f12324,f12327,f13601,f12371,f12372,f12373,f12374,f12375,f12377,f13602,f13603,f12423,f12424,f12425,f12427,f12490,f12491,f12492,f12493,f13624,f12642,f13719,f12669,f12674,f13740,f13348,f13775,f13776,f12802,f12829,f12834,f12837,f12840,f12841,f12843,f12844,f12846,f12847,f12849,f12854,f12868,f13313,f13808,f13824,f13840,f13084,f13089,f13215,f13203,f3049,f13962,f13963,f14173,f14174,f14175,f14191,f14010,f14193,f14162,f14194,f14497,f14198,f14512,f14518,f14529])).
fof(f14522,plain,(
  $false | ($spl8 | $spl12 | $spl26)),
  inference(global_subsumption,[],[f14521,f70,f71,f76,f80,f79,f102,f107,f104,f108,f110,f131,f162,f207,f325,f326,f376,f558,f856,f575,f598,f611,f891,f651,f917,f923,f713,f715,f769,f792,f780,f996,f958,f1072,f1074,f1093,f1168,f1222,f1224,f1226,f1227,f1663,f1664,f1297,f1739,f1783,f2996,f2989,f2966,f3031,f2025,f2948,f2928,f2919,f2916,f2900,f2895,f3048,f2889,f2886,f3059,f2122,f2881,f2871,f3075,f2866,f2843,f2842,f2837,f2836,f3107,f2283,f2285,f3133,f2333,f3137,f3143,f3163,f3172,f3174,f3179,f3180,f2700,f3200,f3202,f103,f118,f187,f219,f264,f417,f456,f493,f523,f542,f576,f612,f921,f934,f937,f957,f1050,f1094,f1799,f1801,f1951,f1952,f1993,f1994,f2024,f2061,f2091,f2121,f2125,f2155,f2156,f2219,f2282,f2366,f2429,f2492,f2495,f2498,f2499,f2500,f2502,f2503,f2588,f2829,f2831,f2835,f3076,f3219,f1665,f938,f1795,f1740,f1223,f714,f263,f249,f111,f455,f454,f73,f83,f84,f453,f2494,f3148,f3201,f3095,f2023,f1944,f915,f450,f416,f4143,f4172,f4173,f4245,f4324,f4382,f4966,f5038,f5039,f4995,f4996,f5044,f5047,f5090,f5091,f5192,f5193,f5194,f5260,f5261,f5262,f5265,f5266,f5268,f5269,f5270,f5271,f5272,f5273,f5274,f5275,f5279,f5300,f5307,f5322,f5323,f5324,f5816,f6124,f6425,f6426,f6427,f6507,f6508,f6509,f6512,f6513,f6514,f6516,f6517,f6518,f6519,f6520,f6521,f6522,f6523,f6529,f6555,f6558,f6565,f6580,f6581,f6582,f6586,f6770,f6841,f6848,f6850,f7090,f6890,f6893,f6944,f6945,f6946,f6947,f7097,f7181,f7195,f7372,f7406,f7407,f7411,f690,f522,f7199,f7180,f7177,f7163,f7162,f7108,f7107,f6583,f6568,f6567,f6564,f6563,f6554,f6549,f6547,f6546,f6545,f6544,f6543,f6515,f6058,f5310,f5309,f5306,f5305,f5299,f5294,f5292,f5291,f5290,f5289,f5288,f5267,f7176,f7246,f7247,f7389,f7387,f5327,f6589,f6839,f6838,f6748,f6747,f5147,f5146,f3170,f3141,f2480,f2417,f2354,f2270,f2207,f1157,f681,f642,f469,f468,f258,f77,f74,f4379,f4321,f2585,f2489,f2426,f2363,f2279,f2216,f1294,f1165,f766,f687,f648,f525,f524,f7987,f7988,f8060,f8062,f8066,f8068,f8072,f8074,f8080,f8057,f8058,f8096,f8097,f8125,f8126,f8127,f8132,f8133,f8135,f8249,f8251,f8235,f8264,f8292,f8294,f8314,f8338,f8339,f8340,f8341,f8568,f8569,f8570,f8571,f8810,f8678,f8679,f8680,f8681,f8724,f8726,f8727,f8832,f8841,f8842,f8844,f8895,f8896,f8897,f8898,f8923,f8925,f8926,f9082,f9107,f9080,f9081,f9314,f9332,f9149,f9150,f9189,f9190,f9375,f9377,f9378,f9546,f9513,f9562,f9563,f9473,f9474,f9475,f9476,f9961,f9983,f9985,f9850,f9851,f10000,f10008,f10017,f10018,f10095,f10093,f10120,f10090,f10089,f10086,f10073,f10074,f10673,f10674,f11088,f11118,f11119,f10927,f10928,f11009,f11010,f11192,f11194,f11199,f11200,f11205,f11209,f11187,f11188,f11361,f11362,f11363,f11401,f11402,f11403,f11404,f11555,f11581,f11552,f11553,f11594,f11597,f11649,f11650,f11651,f11652,f11777,f11778,f11779,f11780,f11841,f11842,f11987,f11978,f11989,f11976,f11973,f11995,f11996,f11998,f11962,f12008,f12009,f11961,f12010,f12011,f12029,f12033,f12034,f12035,f12036,f12037,f12038,f12040,f12041,f13550,f13532,f12274,f12275,f13591,f12321,f12322,f12323,f12324,f12327,f13601,f12371,f12372,f12373,f12374,f12375,f12377,f13602,f13603,f12423,f12424,f12425,f12427,f12490,f12491,f12492,f12493,f13624,f12642,f13719,f12669,f12674,f13740,f13348,f13775,f13776,f12802,f12829,f12834,f12837,f12840,f12841,f12843,f12844,f12846,f12847,f12849,f12854,f12868,f13313,f13808,f13824,f13840,f13084,f13089,f13215,f13203,f3049,f13962,f13963,f14173,f14174,f14175,f14191,f14010,f14193,f14162,f14194,f14497,f14198,f14512,f14518])).
fof(f14521,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | ($spl8 | $spl12 | $spl26)),
  inference(condensation,[],[f14520])).
fof(f14520,plain,(
  ( ! [X85,X83] : (s(bool,'F_0') = s(bool,X85) | s(bool,'F_0') = s(bool,X83)) ) | ($spl8 | $spl12 | $spl26)),
  inference(subsumption_resolution,[],[f14344,f14194])).
fof(f14344,plain,(
  ( ! [X85,X83,X86] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,X85),s(bool,X83),s(bool,X86))) | s(bool,'F_0') = s(bool,X85) | s(bool,'F_0') = s(bool,X83)) ) | ($spl8 | $spl12 | $spl26)),
  inference(backward_demodulation,[],[f14194,f8076])).
fof(f14519,plain,(
  $false | ($spl8 | $spl12)),
  inference(global_subsumption,[],[f70,f71,f76,f80,f79,f102,f107,f104,f108,f110,f131,f162,f207,f325,f326,f376,f558,f856,f575,f598,f611,f891,f651,f917,f923,f713,f715,f769,f792,f780,f996,f958,f1072,f1074,f1093,f1168,f1222,f1224,f1226,f1227,f1663,f1664,f1297,f1739,f1783,f2996,f2989,f2966,f3031,f2025,f2948,f2928,f2919,f2916,f2900,f2895,f3048,f2889,f2886,f3059,f2122,f2881,f2871,f3075,f2866,f2843,f2842,f2837,f2836,f3107,f2283,f2285,f3133,f2333,f3137,f3143,f3163,f3172,f3174,f3179,f3180,f2700,f3200,f3202,f103,f118,f187,f219,f264,f417,f456,f493,f523,f542,f576,f612,f921,f934,f937,f957,f1050,f1094,f1799,f1801,f1951,f1952,f1993,f1994,f2024,f2061,f2091,f2121,f2125,f2155,f2156,f2219,f2282,f2366,f2429,f2492,f2495,f2498,f2499,f2500,f2502,f2503,f2588,f2829,f2831,f2835,f3076,f3219,f1665,f938,f1795,f1740,f1223,f714,f263,f249,f111,f455,f454,f73,f83,f84,f453,f2494,f3148,f3201,f3095,f2023,f1944,f915,f450,f416,f4143,f4172,f4173,f4245,f4324,f4382,f4966,f5038,f5039,f4995,f4996,f5044,f5047,f5090,f5091,f5192,f5193,f5194,f5260,f5261,f5262,f5265,f5266,f5268,f5269,f5270,f5271,f5272,f5273,f5274,f5275,f5279,f5300,f5307,f5322,f5323,f5324,f5816,f6124,f6425,f6426,f6427,f6507,f6508,f6509,f6512,f6513,f6514,f6516,f6517,f6518,f6519,f6520,f6521,f6522,f6523,f6529,f6555,f6558,f6565,f6580,f6581,f6582,f6586,f6770,f6841,f6848,f6850,f7090,f6890,f6893,f6944,f6945,f6946,f6947,f7097,f7181,f7195,f7372,f7406,f7407,f7411,f690,f522,f7199,f7180,f7177,f7163,f7162,f7108,f7107,f6583,f6568,f6567,f6564,f6563,f6554,f6549,f6547,f6546,f6545,f6544,f6543,f6515,f6058,f5310,f5309,f5306,f5305,f5299,f5294,f5292,f5291,f5290,f5289,f5288,f5267,f7176,f7246,f7247,f7389,f7387,f5327,f6589,f6839,f6838,f6748,f6747,f5147,f5146,f3170,f3141,f2480,f2417,f2354,f2270,f2207,f1157,f681,f642,f469,f468,f258,f77,f74,f4379,f4321,f2585,f2489,f2426,f2363,f2279,f2216,f1294,f1165,f766,f687,f648,f525,f524,f7987,f7988,f8060,f8062,f8066,f8068,f8072,f8074,f8080,f8057,f8058,f8096,f8097,f8125,f8126,f8127,f8132,f8133,f8135,f8249,f8251,f8235,f8264,f8292,f8294,f8314,f8338,f8339,f8340,f8341,f8568,f8569,f8570,f8571,f8810,f8678,f8679,f8680,f8681,f8724,f8726,f8727,f8832,f8841,f8842,f8844,f8895,f8896,f8897,f8898,f8923,f8925,f8926,f9082,f9107,f9080,f9081,f9314,f9332,f9149,f9150,f9189,f9190,f9375,f9377,f9378,f9546,f9513,f9562,f9563,f9473,f9474,f9475,f9476,f9961,f9983,f9985,f9850,f9851,f10000,f10008,f10017,f10018,f10095,f10093,f10120,f10090,f10089,f10086,f10073,f10074,f10673,f10674,f11088,f11118,f11119,f10927,f10928,f11009,f11010,f11192,f11194,f11199,f11200,f11205,f11209,f11187,f11188,f11361,f11362,f11363,f11401,f11402,f11403,f11404,f11555,f11581,f11552,f11553,f11594,f11597,f11649,f11650,f11651,f11652,f11777,f11778,f11779,f11780,f11841,f11842,f11987,f11978,f11989,f11976,f11973,f11995,f11996,f11998,f11962,f12008,f12009,f11961,f12010,f12011,f12029,f12033,f12034,f12035,f12036,f12037,f12038,f12040,f12041,f13550,f13532,f12274,f12275,f13591,f12321,f12322,f12323,f12324,f12327,f13601,f12371,f12372,f12373,f12374,f12375,f12377,f13602,f13603,f12423,f12424,f12425,f12427,f12490,f12491,f12492,f12493,f13624,f12642,f13719,f12669,f12674,f13740,f13348,f13775,f13776,f12802,f12829,f12834,f12837,f12840,f12841,f12843,f12844,f12846,f12847,f12849,f12854,f12868,f13313,f13808,f13824,f13840,f13084,f13089,f13215,f13203,f3049,f13962,f13963,f14173,f14174,f14175,f14191,f14010,f14193,f14162,f14194,f14497,f14198,f14512,f14518])).
fof(f14513,plain,(
  $false | ($spl8 | $spl12)),
  inference(global_subsumption,[],[f70,f71,f76,f80,f79,f102,f107,f104,f108,f110,f131,f162,f207,f325,f326,f376,f558,f856,f575,f598,f611,f891,f651,f917,f923,f713,f715,f769,f792,f780,f996,f958,f1072,f1074,f1093,f1168,f1222,f1224,f1226,f1227,f1663,f1664,f1297,f1739,f1783,f2996,f2989,f2966,f3031,f2025,f2948,f2928,f2919,f2916,f2900,f2895,f3048,f2889,f2886,f3059,f2122,f2881,f2871,f3075,f2866,f2843,f2842,f2837,f2836,f3107,f2283,f2285,f3133,f2333,f3137,f3143,f3163,f3172,f3174,f3179,f3180,f2700,f3200,f3202,f103,f118,f187,f219,f264,f417,f456,f493,f523,f542,f576,f612,f921,f934,f937,f957,f1050,f1094,f1799,f1801,f1951,f1952,f1993,f1994,f2024,f2061,f2091,f2121,f2125,f2155,f2156,f2219,f2282,f2366,f2429,f2492,f2495,f2498,f2499,f2500,f2502,f2503,f2588,f2829,f2831,f2835,f3076,f3219,f1665,f938,f1795,f1740,f1223,f714,f263,f249,f111,f455,f454,f73,f83,f84,f453,f2494,f3148,f3201,f3095,f2023,f1944,f915,f450,f416,f4143,f4172,f4173,f4245,f4324,f4382,f4966,f5038,f5039,f4995,f4996,f5044,f5047,f5090,f5091,f5192,f5193,f5194,f5260,f5261,f5262,f5265,f5266,f5268,f5269,f5270,f5271,f5272,f5273,f5274,f5275,f5279,f5300,f5307,f5322,f5323,f5324,f5816,f6124,f6425,f6426,f6427,f6507,f6508,f6509,f6512,f6513,f6514,f6516,f6517,f6518,f6519,f6520,f6521,f6522,f6523,f6529,f6555,f6558,f6565,f6580,f6581,f6582,f6586,f6770,f6841,f6848,f6850,f7090,f6890,f6893,f6944,f6945,f6946,f6947,f7097,f7181,f7195,f7372,f7406,f7407,f7411,f690,f522,f7199,f7180,f7177,f7163,f7162,f7108,f7107,f6583,f6568,f6567,f6564,f6563,f6554,f6549,f6547,f6546,f6545,f6544,f6543,f6515,f6058,f5310,f5309,f5306,f5305,f5299,f5294,f5292,f5291,f5290,f5289,f5288,f5267,f7176,f7246,f7247,f7389,f7387,f5327,f6589,f6839,f6838,f6748,f6747,f5147,f5146,f3170,f3141,f2480,f2417,f2354,f2270,f2207,f1157,f681,f642,f469,f468,f258,f77,f74,f4379,f4321,f2585,f2489,f2426,f2363,f2279,f2216,f1294,f1165,f766,f687,f648,f525,f524,f7987,f7988,f8060,f8062,f8066,f8068,f8072,f8074,f8080,f8057,f8058,f8096,f8097,f8125,f8126,f8127,f8132,f8133,f8135,f8249,f8251,f8235,f8264,f8292,f8294,f8314,f8338,f8339,f8340,f8341,f8568,f8569,f8570,f8571,f8810,f8678,f8679,f8680,f8681,f8724,f8726,f8727,f8832,f8841,f8842,f8844,f8895,f8896,f8897,f8898,f8923,f8925,f8926,f9082,f9107,f9080,f9081,f9314,f9332,f9149,f9150,f9189,f9190,f9375,f9377,f9378,f9546,f9513,f9562,f9563,f9473,f9474,f9475,f9476,f9961,f9983,f9985,f9850,f9851,f10000,f10008,f10017,f10018,f10095,f10093,f10120,f10090,f10089,f10086,f10073,f10074,f10673,f10674,f11088,f11118,f11119,f10927,f10928,f11009,f11010,f11192,f11194,f11199,f11200,f11205,f11209,f11187,f11188,f11361,f11362,f11363,f11401,f11402,f11403,f11404,f11555,f11581,f11552,f11553,f11594,f11597,f11649,f11650,f11651,f11652,f11777,f11778,f11779,f11780,f11841,f11842,f11987,f11978,f11989,f11976,f11973,f11995,f11996,f11998,f11962,f12008,f12009,f11961,f12010,f12011,f12029,f12033,f12034,f12035,f12036,f12037,f12038,f12040,f12041,f13550,f13532,f12274,f12275,f13591,f12321,f12322,f12323,f12324,f12327,f13601,f12371,f12372,f12373,f12374,f12375,f12377,f13602,f13603,f12423,f12424,f12425,f12427,f12490,f12491,f12492,f12493,f13624,f12642,f13719,f12669,f12674,f13740,f13348,f13775,f13776,f12802,f12829,f12834,f12837,f12840,f12841,f12843,f12844,f12846,f12847,f12849,f12854,f12868,f13313,f13808,f13824,f13840,f13084,f13089,f13215,f13203,f3049,f13962,f13963,f14173,f14174,f14175,f14191,f14010,f14193,f14162,f14194,f14497,f14198,f14512])).
fof(f14499,plain,(
  $false | ($spl8 | $spl12)),
  inference(global_subsumption,[],[f70,f71,f76,f80,f79,f102,f107,f104,f108,f110,f131,f162,f207,f325,f326,f376,f558,f856,f575,f598,f611,f891,f651,f917,f923,f713,f715,f769,f792,f780,f996,f958,f1072,f1074,f1093,f1168,f1222,f1224,f1226,f1227,f1663,f1664,f1297,f1739,f1783,f2996,f2989,f2966,f3031,f2025,f2948,f2928,f2919,f2916,f2900,f2895,f3048,f2889,f2886,f3059,f2122,f2881,f2871,f3075,f2866,f2843,f2842,f2837,f2836,f3107,f2283,f2285,f3133,f2333,f3137,f3143,f3163,f3172,f3174,f3179,f3180,f2700,f3200,f3202,f103,f118,f187,f219,f264,f417,f456,f493,f523,f542,f576,f612,f921,f934,f937,f957,f1050,f1094,f1799,f1801,f1951,f1952,f1993,f1994,f2024,f2061,f2091,f2121,f2125,f2155,f2156,f2219,f2282,f2366,f2429,f2492,f2495,f2498,f2499,f2500,f2502,f2503,f2588,f2829,f2831,f2835,f3076,f3219,f1665,f938,f1795,f1740,f1223,f714,f263,f249,f111,f455,f454,f73,f83,f84,f453,f2494,f3148,f3201,f3095,f2023,f1944,f915,f450,f416,f4143,f4172,f4173,f4245,f4324,f4382,f4966,f5038,f5039,f4995,f4996,f5044,f5047,f5090,f5091,f5192,f5193,f5194,f5260,f5261,f5262,f5265,f5266,f5268,f5269,f5270,f5271,f5272,f5273,f5274,f5275,f5279,f5300,f5307,f5322,f5323,f5324,f5816,f6124,f6425,f6426,f6427,f6507,f6508,f6509,f6512,f6513,f6514,f6516,f6517,f6518,f6519,f6520,f6521,f6522,f6523,f6529,f6555,f6558,f6565,f6580,f6581,f6582,f6586,f6770,f6841,f6848,f6850,f7090,f6890,f6893,f6944,f6945,f6946,f6947,f7097,f7181,f7195,f7372,f7406,f7407,f7411,f690,f522,f7199,f7180,f7177,f7163,f7162,f7108,f7107,f6583,f6568,f6567,f6564,f6563,f6554,f6549,f6547,f6546,f6545,f6544,f6543,f6515,f6058,f5310,f5309,f5306,f5305,f5299,f5294,f5292,f5291,f5290,f5289,f5288,f5267,f7176,f7246,f7247,f7389,f7387,f5327,f6589,f6839,f6838,f6748,f6747,f5147,f5146,f3170,f3141,f2480,f2417,f2354,f2270,f2207,f1157,f681,f642,f469,f468,f258,f77,f74,f4379,f4321,f2585,f2489,f2426,f2363,f2279,f2216,f1294,f1165,f766,f687,f648,f525,f524,f7987,f7988,f8060,f8062,f8066,f8068,f8072,f8074,f8080,f8057,f8058,f8096,f8097,f8125,f8126,f8127,f8132,f8133,f8135,f8249,f8251,f8235,f8264,f8292,f8294,f8314,f8338,f8339,f8340,f8341,f8568,f8569,f8570,f8571,f8810,f8678,f8679,f8680,f8681,f8724,f8726,f8727,f8832,f8841,f8842,f8844,f8895,f8896,f8897,f8898,f8923,f8925,f8926,f9082,f9107,f9080,f9081,f9314,f9332,f9149,f9150,f9189,f9190,f9375,f9377,f9378,f9546,f9513,f9562,f9563,f9473,f9474,f9475,f9476,f9961,f9983,f9985,f9850,f9851,f10000,f10008,f10017,f10018,f10095,f10093,f10120,f10090,f10089,f10086,f10073,f10074,f10673,f10674,f11088,f11118,f11119,f10927,f10928,f11009,f11010,f11192,f11194,f11199,f11200,f11205,f11209,f11187,f11188,f11361,f11362,f11363,f11401,f11402,f11403,f11404,f11555,f11581,f11552,f11553,f11594,f11597,f11649,f11650,f11651,f11652,f11777,f11778,f11779,f11780,f11841,f11842,f11987,f11978,f11989,f11976,f11973,f11995,f11996,f11998,f11962,f12008,f12009,f11961,f12010,f12011,f12029,f12033,f12034,f12035,f12036,f12037,f12038,f12040,f12041,f13550,f13532,f12274,f12275,f13591,f12321,f12322,f12323,f12324,f12327,f13601,f12371,f12372,f12373,f12374,f12375,f12377,f13602,f13603,f12423,f12424,f12425,f12427,f12490,f12491,f12492,f12493,f13624,f12642,f13719,f12669,f12674,f13740,f13348,f13775,f13776,f12802,f12829,f12834,f12837,f12840,f12841,f12843,f12844,f12846,f12847,f12849,f12854,f12868,f13313,f13808,f13824,f13840,f13084,f13089,f13215,f13203,f3049,f13962,f13963,f14173,f14174,f14175,f14191,f14010,f14193,f14162,f14194,f14497,f14198])).
fof(f14498,plain,(
  $false | ($spl8 | $spl12)),
  inference(global_subsumption,[],[f70,f71,f76,f80,f79,f102,f107,f104,f108,f110,f131,f162,f207,f325,f326,f376,f558,f856,f575,f598,f611,f891,f651,f917,f923,f713,f715,f769,f792,f780,f996,f958,f1072,f1074,f1093,f1168,f1222,f1224,f1226,f1227,f1663,f1664,f1297,f1739,f1783,f2996,f2989,f2966,f3031,f2025,f2948,f2928,f2919,f2916,f2900,f2895,f3048,f2889,f2886,f3059,f2122,f2881,f2871,f3075,f2866,f2843,f2842,f2837,f2836,f3107,f2283,f2285,f3133,f2333,f3137,f3143,f3163,f3172,f3174,f3179,f3180,f2700,f3200,f3202,f103,f118,f187,f219,f264,f417,f456,f493,f523,f542,f576,f612,f921,f934,f937,f957,f1050,f1094,f1799,f1801,f1951,f1952,f1993,f1994,f2024,f2061,f2091,f2121,f2125,f2155,f2156,f2219,f2282,f2366,f2429,f2492,f2495,f2498,f2499,f2500,f2502,f2503,f2588,f2829,f2831,f2835,f3076,f3219,f1665,f938,f1795,f1740,f1223,f714,f263,f249,f111,f455,f454,f73,f83,f84,f453,f2494,f3148,f3201,f3095,f2023,f1944,f915,f450,f416,f4143,f4172,f4173,f4245,f4324,f4382,f4966,f5038,f5039,f4995,f4996,f5044,f5047,f5090,f5091,f5192,f5193,f5194,f5260,f5261,f5262,f5265,f5266,f5268,f5269,f5270,f5271,f5272,f5273,f5274,f5275,f5279,f5300,f5307,f5322,f5323,f5324,f5816,f6124,f6425,f6426,f6427,f6507,f6508,f6509,f6512,f6513,f6514,f6516,f6517,f6518,f6519,f6520,f6521,f6522,f6523,f6529,f6555,f6558,f6565,f6580,f6581,f6582,f6586,f6770,f6841,f6848,f6850,f7090,f6890,f6893,f6944,f6945,f6946,f6947,f7097,f7181,f7195,f7372,f7406,f7407,f7411,f690,f522,f7199,f7180,f7177,f7163,f7162,f7108,f7107,f6583,f6568,f6567,f6564,f6563,f6554,f6549,f6547,f6546,f6545,f6544,f6543,f6515,f6058,f5310,f5309,f5306,f5305,f5299,f5294,f5292,f5291,f5290,f5289,f5288,f5267,f7176,f7246,f7247,f7389,f7387,f5327,f6589,f6839,f6838,f6748,f6747,f5147,f5146,f3170,f3141,f2480,f2417,f2354,f2270,f2207,f1157,f681,f642,f469,f468,f258,f77,f74,f4379,f4321,f2585,f2489,f2426,f2363,f2279,f2216,f1294,f1165,f766,f687,f648,f525,f524,f7987,f7988,f8060,f8062,f8066,f8068,f8072,f8074,f8080,f8057,f8058,f8096,f8097,f8125,f8126,f8127,f8132,f8133,f8135,f8249,f8251,f8235,f8264,f8292,f8294,f8314,f8338,f8339,f8340,f8341,f8568,f8569,f8570,f8571,f8810,f8678,f8679,f8680,f8681,f8724,f8726,f8727,f8832,f8841,f8842,f8844,f8895,f8896,f8897,f8898,f8923,f8925,f8926,f9082,f9107,f9080,f9081,f9314,f9332,f9149,f9150,f9189,f9190,f9375,f9377,f9378,f9546,f9513,f9562,f9563,f9473,f9474,f9475,f9476,f9961,f9983,f9985,f9850,f9851,f10000,f10008,f10017,f10018,f10095,f10093,f10120,f10090,f10089,f10086,f10073,f10074,f10673,f10674,f11088,f11118,f11119,f10927,f10928,f11009,f11010,f11192,f11194,f11199,f11200,f11205,f11209,f11187,f11188,f11361,f11362,f11363,f11401,f11402,f11403,f11404,f11555,f11581,f11552,f11553,f11594,f11597,f11649,f11650,f11651,f11652,f11777,f11778,f11779,f11780,f11841,f11842,f11987,f11978,f11989,f11976,f11973,f11995,f11996,f11998,f11962,f12008,f12009,f11961,f12010,f12011,f12029,f12033,f12034,f12035,f12036,f12037,f12038,f12040,f12041,f13550,f13532,f12274,f12275,f13591,f12321,f12322,f12323,f12324,f12327,f13601,f12371,f12372,f12373,f12374,f12375,f12377,f13602,f13603,f12423,f12424,f12425,f12427,f12490,f12491,f12492,f12493,f13624,f12642,f13719,f12669,f12674,f13740,f13348,f13775,f13776,f12802,f12829,f12834,f12837,f12840,f12841,f12843,f12844,f12846,f12847,f12849,f12854,f12868,f13313,f13808,f13824,f13840,f13084,f13089,f13215,f13203,f3049,f13962,f13963,f14173,f14174,f14175,f14191,f14010,f14193,f14162,f14194,f14497])).
fof(f14496,plain,(
  $false | ($spl8 | $spl12)),
  inference(subsumption_resolution,[],[f14195,f71])).
fof(f14195,plain,(
  ~p(s(bool,'T_0')) | ($spl8 | $spl12)),
  inference(backward_demodulation,[],[f14194,f73])).
fof(f14193,plain,(
  ( ! [X54,X55] : (s(bool,'T_0') = s(bool,X54) | p(s(bool,X55))) ) | ($spl8 | $spl12)),
  inference(global_subsumption,[],[f70,f71,f76,f80,f79,f102,f107,f104,f108,f110,f131,f162,f207,f325,f326,f376,f558,f856,f575,f598,f611,f891,f651,f917,f923,f713,f715,f769,f792,f780,f996,f958,f1072,f1074,f1093,f1168,f1222,f1224,f1226,f1227,f1663,f1664,f1297,f1739,f1783,f2996,f2989,f2966,f3031,f2025,f2948,f2928,f2919,f2916,f2900,f2895,f3048,f2889,f2886,f3059,f2122,f2881,f2871,f3075,f2866,f2843,f2842,f2837,f2836,f3107,f2283,f2285,f3133,f2333,f3137,f3143,f3163,f3172,f3174,f3179,f3180,f2700,f3200,f3202,f103,f118,f187,f219,f264,f417,f456,f493,f523,f542,f576,f612,f921,f934,f937,f957,f1050,f1094,f1799,f1801,f1951,f1952,f1993,f1994,f2024,f2061,f2091,f2121,f2125,f2155,f2156,f2219,f2282,f2366,f2429,f2492,f2495,f2498,f2499,f2500,f2502,f2503,f2588,f2829,f2831,f2835,f3076,f3219,f1665,f938,f1795,f1740,f1223,f714,f263,f249,f111,f455,f454,f73,f83,f84,f453,f2494,f3148,f3201,f3095,f2023,f1944,f915,f450,f416,f4143,f4172,f4173,f4245,f4324,f4382,f4966,f5038,f5039,f4995,f4996,f5044,f5047,f5090,f5091,f5192,f5193,f5194,f5260,f5261,f5262,f5265,f5266,f5268,f5269,f5270,f5271,f5272,f5273,f5274,f5275,f5279,f5300,f5307,f5322,f5323,f5324,f5816,f6124,f6425,f6426,f6427,f6507,f6508,f6509,f6512,f6513,f6514,f6516,f6517,f6518,f6519,f6520,f6521,f6522,f6523,f6529,f6555,f6558,f6565,f6580,f6581,f6582,f6586,f6770,f6841,f6848,f6850,f7090,f6890,f6893,f6944,f6945,f6946,f6947,f7097,f7181,f7195,f7372,f7406,f7407,f7411,f690,f522,f7199,f7180,f7177,f7163,f7162,f7108,f7107,f6583,f6568,f6567,f6564,f6563,f6554,f6549,f6547,f6546,f6545,f6544,f6543,f6515,f6058,f5310,f5309,f5306,f5305,f5299,f5294,f5292,f5291,f5290,f5289,f5288,f5267,f7176,f7246,f7247,f7389,f7387,f5327,f6589,f6839,f6838,f6748,f6747,f5147,f5146,f3170,f3141,f2480,f2417,f2354,f2270,f2207,f1157,f681,f642,f469,f468,f258,f77,f74,f4379,f4321,f2585,f2489,f2426,f2363,f2279,f2216,f1294,f1165,f766,f687,f648,f525,f524,f7987,f7988,f8060,f8062,f8066,f8068,f8072,f8074,f8080,f8057,f8058,f8096,f8097,f8125,f8126,f8127,f8132,f8133,f8135,f8249,f8251,f8235,f8264,f8292,f8294,f8314,f8338,f8339,f8340,f8341,f8568,f8569,f8570,f8571,f8810,f8678,f8679,f8680,f8681,f8724,f8726,f8727,f8832,f8841,f8842,f8844,f8895,f8896,f8897,f8898,f8923,f8925,f8926,f9082,f9107,f9080,f9081,f9314,f9332,f9149,f9150,f9189,f9190,f9375,f9377,f9378,f9546,f9513,f9562,f9563,f9473,f9474,f9475,f9476,f9961,f9983,f9985,f9850,f9851,f10000,f10008,f10017,f10018,f10095,f10093,f10120,f10090,f10089,f10086,f10073,f10074,f10673,f10674,f11088,f11118,f11119,f10927,f10928,f11009,f11010,f11192,f11194,f11199,f11200,f11205,f11209,f11187,f11188,f11361,f11362,f11363,f11401,f11402,f11403,f11404,f11555,f11581,f11552,f11553,f11594,f11597,f11649,f11650,f11651,f11652,f11777,f11778,f11779,f11780,f11841,f11842,f11987,f11978,f11989,f11976,f11973,f11995,f11996,f11998,f11962,f12008,f12009,f11961,f12010,f12011,f12029,f12033,f12034,f12035,f12036,f12037,f12038,f12040,f12041,f13550,f13532,f12274,f12275,f13591,f12321,f12322,f12323,f12324,f12327,f13601,f12371,f12372,f12373,f12374,f12375,f12377,f13602,f13603,f12423,f12424,f12425,f12427,f12490,f12491,f12492,f12493,f13624,f12642,f13719,f12669,f12674,f13740,f13348,f13775,f13776,f12802,f12829,f12834,f12837,f12840,f12841,f12843,f12844,f12846,f12847,f12849,f12854,f12868,f13313,f13808,f13824,f13840,f13084,f13089,f13215,f13203,f3049,f13962,f13963,f14173,f14174,f14175,f14191,f14010])).
fof(f7382,plain,(
  ( ! [X32] : (s(bool,'F_0') = s(bool,X32)) ) | $spl44),
  inference(cnf_transformation,[],[f7382_D])).
fof(f7382_D,plain,(
  ( ! [X32] : (s(bool,'F_0') = s(bool,X32)) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f14192,plain,(
  ( ! [X128,X124] : (s(bool,'F_0') = s(bool,X124) | s(bool,'T_0') = s(bool,X128)) ) | ($spl8 | $spl12 | $spl68)),
  inference(global_subsumption,[],[f70,f71,f76,f80,f79,f102,f107,f104,f108,f110,f131,f162,f207,f325,f326,f376,f558,f856,f575,f598,f611,f891,f651,f917,f923,f713,f715,f769,f792,f780,f996,f958,f1072,f1074,f1093,f1168,f1222,f1224,f1226,f1227,f1663,f1664,f1297,f1739,f1783,f2996,f2989,f2966,f3031,f2025,f2948,f2928,f2919,f2916,f2900,f2895,f3048,f2889,f2886,f3059,f2122,f2881,f2871,f3075,f2866,f2843,f2842,f2837,f2836,f3107,f2283,f2285,f3133,f2333,f3137,f3143,f3163,f3172,f3174,f3179,f3180,f2700,f3200,f3202,f103,f118,f187,f219,f264,f417,f456,f493,f523,f542,f576,f612,f921,f934,f937,f957,f1050,f1094,f1799,f1801,f1951,f1952,f1993,f1994,f2024,f2061,f2091,f2121,f2125,f2155,f2156,f2219,f2282,f2366,f2429,f2492,f2495,f2498,f2499,f2500,f2502,f2503,f2588,f2829,f2831,f2835,f3076,f3219,f1665,f938,f1795,f1740,f1223,f714,f263,f249,f111,f455,f454,f73,f83,f84,f453,f2494,f3148,f3201,f3095,f2023,f1944,f915,f450,f416,f4143,f4172,f4173,f4245,f4324,f4382,f4966,f5038,f5039,f4995,f4996,f5044,f5047,f5090,f5091,f5192,f5193,f5194,f5260,f5261,f5262,f5265,f5266,f5268,f5269,f5270,f5271,f5272,f5273,f5274,f5275,f5279,f5300,f5307,f5322,f5323,f5324,f5816,f6124,f6425,f6426,f6427,f6507,f6508,f6509,f6512,f6513,f6514,f6516,f6517,f6518,f6519,f6520,f6521,f6522,f6523,f6529,f6555,f6558,f6565,f6580,f6581,f6582,f6586,f6770,f6841,f6848,f6850,f7090,f6890,f6893,f6944,f6945,f6946,f6947,f7097,f7181,f7195,f7372,f7406,f7407,f7411,f690,f522,f7199,f7180,f7177,f7163,f7162,f7108,f7107,f6583,f6568,f6567,f6564,f6563,f6554,f6549,f6547,f6546,f6545,f6544,f6543,f6515,f6058,f5310,f5309,f5306,f5305,f5299,f5294,f5292,f5291,f5290,f5289,f5288,f5267,f7176,f7246,f7247,f7389,f7387,f5327,f6589,f6839,f6838,f6748,f6747,f5147,f5146,f3170,f3141,f2480,f2417,f2354,f2270,f2207,f1157,f681,f642,f469,f468,f258,f77,f74,f4379,f4321,f2585,f2489,f2426,f2363,f2279,f2216,f1294,f1165,f766,f687,f648,f525,f524,f7987,f7988,f8060,f8062,f8066,f8068,f8072,f8074,f8080,f8057,f8058,f8096,f8097,f8125,f8126,f8127,f8132,f8133,f8135,f8249,f8251,f8235,f8264,f8292,f8294,f8314,f8338,f8339,f8340,f8341,f8568,f8569,f8570,f8571,f8810,f8678,f8679,f8680,f8681,f8724,f8726,f8727,f8832,f8841,f8842,f8844,f8895,f8896,f8897,f8898,f8923,f8925,f8926,f9082,f9107,f9080,f9081,f9314,f9332,f9149,f9150,f9189,f9190,f9375,f9377,f9378,f9546,f9513,f9562,f9563,f9473,f9474,f9475,f9476,f9961,f9983,f9985,f9850,f9851,f10000,f10008,f9977,f10017,f10018,f10095,f10093,f10120,f10090,f10089,f10086,f10073,f10074,f10174,f10175,f10298,f10303,f10346,f10347,f10348,f10349,f10370,f10373,f10536,f10603,f10604,f10510,f10615,f10616,f10509,f10700,f10673,f10674,f10857,f10742,f10743,f10744,f10745,f10770,f10772,f11088,f11118,f11119,f10927,f10928,f11009,f11010,f11192,f11194,f11199,f11200,f11205,f11209,f11222,f11187,f11188,f11269,f11274,f11291,f11292,f11295,f11296,f11265,f11266,f11361,f11362,f11363,f11401,f11402,f11403,f11404,f11555,f11581,f11552,f11553,f11594,f11597,f11649,f11650,f11651,f11652,f11777,f11778,f11779,f11780,f11841,f11842,f11987,f11978,f11989,f11976,f11973,f11995,f11996,f11998,f11962,f12008,f12009,f11961,f12010,f12011,f12026,f12029,f12033,f12034,f12035,f12036,f12037,f12038,f12040,f12041,f13550,f13532,f13522,f12274,f12275,f13591,f12321,f12322,f12323,f12324,f12327,f13601,f12371,f12372,f12373,f12374,f12375,f12377,f13602,f13603,f12423,f12424,f12425,f12427,f12490,f12491,f12492,f12493,f13624,f12642,f13719,f12669,f12674,f13740,f13348,f13775,f13776,f12802,f12829,f12834,f12837,f12840,f12841,f12843,f12844,f12846,f12847,f12849,f12854,f12868,f13313,f13808,f12927,f13824,f13840,f13863,f13084,f13089,f13928,f13215,f13203,f3049,f13962,f13963,f14173,f14174,f14175,f14176,f14177,f14180,f14181,f14185,f14186,f14191])).
fof(f14185,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X0) | ~p(s(bool,X1)) | s(bool,'T_0') = s(bool,X2)) ) | ($spl8 | $spl12 | $spl68)),
  inference(condensation,[],[f14184])).
fof(f14184,plain,(
  ( ! [X163,X161,X162,X160] : (s(bool,X161) = s(bool,X163) | ~p(s(bool,X162)) | s(bool,'F_0') = s(bool,X163) | s(bool,'T_0') = s(bool,X160)) ) | ($spl8 | $spl12 | $spl68)),
  inference(forward_demodulation,[],[f13979,f13963])).
fof(f13979,plain,(
  ( ! [X163,X161,X162,X160] : (s(bool,X163) = s(bool,'const.bool.COND_3'(s(bool,X160),s(bool,X161),s(bool,X162))) | ~p(s(bool,X162)) | s(bool,'F_0') = s(bool,X163) | s(bool,'T_0') = s(bool,X160)) ) | ($spl8 | $spl12 | $spl68)),
  inference(backward_demodulation,[],[f13963,f11265])).
fof(f14180,plain,(
  ( ! [X200,X199,X197] : (p(s(bool,X200)) | s(bool,'T_0') = s(bool,X199) | s(bool,'F_0') = s(bool,X197)) ) | ($spl8 | $spl12 | $spl68)),
  inference(forward_demodulation,[],[f13971,f13963])).
fof(f13971,plain,(
  ( ! [X200,X199,X197] : (p(s(bool,'const.bool.COND_3'(s(bool,X199),s(bool,X200),s(bool,X197)))) | s(bool,'T_0') = s(bool,X199) | s(bool,'F_0') = s(bool,X197)) ) | ($spl8 | $spl12 | $spl68)),
  inference(backward_demodulation,[],[f13963,f10509])).
fof(f14176,plain,(
  ( ! [X116,X114,X115] : (s(bool,'F_0') = s(bool,X115) | p(s(bool,X116)) | s(bool,'T_0') = s(bool,X114)) ) | ($spl8 | $spl12 | $spl68)),
  inference(forward_demodulation,[],[f13969,f13963])).
fof(f13969,plain,(
  ( ! [X116,X114,X115] : (p(s(bool,X116)) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X114),s(bool,X115),s(bool,X116))) | s(bool,'T_0') = s(bool,X114)) ) | ($spl8 | $spl12 | $spl68)),
  inference(backward_demodulation,[],[f13963,f10174])).
fof(f13928,plain,(
  ( ! [X399,X401,X398,X400,X402] : (s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X398),s(bool,X399))) = s(bool,'const.bool.COND_3'(s(bool,X400),s(bool,X400),s(bool,X401))) | s(bool,'F_0') = s(bool,X400) | s(bool,X399) = s(bool,X402) | s(bool,'T_0') = s(bool,X402)) ) | $spl68),
  inference(subsumption_resolution,[],[f13183,f71])).
fof(f13183,plain,(
  ( ! [X399,X401,X398,X400,X402] : (~p(s(bool,'T_0')) | s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X398),s(bool,X399))) = s(bool,'const.bool.COND_3'(s(bool,X400),s(bool,X400),s(bool,X401))) | s(bool,'F_0') = s(bool,X400) | s(bool,X399) = s(bool,X402) | s(bool,'T_0') = s(bool,X402)) ) | $spl68),
  inference(superposition,[],[f10604,f1050])).
fof(f13863,plain,(
  ( ! [X377,X374,X376,X373,X375] : (s(bool,X373) = s(bool,'const.bool.COND_3'(s(bool,X374),s(bool,X374),s(bool,X375))) | s(bool,'F_0') = s(bool,X374) | s(bool,X373) = s(bool,X376) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X377),s(bool,X376)))) ) | $spl68),
  inference(subsumption_resolution,[],[f13024,f71])).
fof(f13024,plain,(
  ( ! [X377,X374,X376,X373,X375] : (~p(s(bool,'T_0')) | s(bool,X373) = s(bool,'const.bool.COND_3'(s(bool,X374),s(bool,X374),s(bool,X375))) | s(bool,'F_0') = s(bool,X374) | s(bool,X373) = s(bool,X376) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X377),s(bool,X376)))) ) | $spl68),
  inference(superposition,[],[f10604,f1050])).
fof(f12927,plain,(
  ( ! [X422,X424,X421,X423] : (s(bool,X421) = s(bool,'const.bool.COND_3'(s(bool,X422),s(bool,X422),s(bool,X421))) | s(bool,'F_0') = s(bool,X422) | s(bool,X421) = s(bool,X423) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X424),s(bool,X423)))) ) | $spl68),
  inference(superposition,[],[f10298,f1050])).
fof(f13522,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'const.bool.COND_3'(s(bool,X1),s(bool,X1),s(bool,X2))))) | s(bool,'F_0') = s(bool,X1)) ) | $spl68),
  inference(condensation,[],[f13521])).
fof(f13521,plain,(
  ( ! [X177,X178,X176,X179] : (s(bool,'T_0') = s(bool,X178) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X179),s(bool,'const.bool.COND_3'(s(bool,X176),s(bool,X176),s(bool,X177))))) | s(bool,'F_0') = s(bool,X176)) ) | $spl68),
  inference(duplicate_literal_removal,[],[f12271])).
fof(f12271,plain,(
  ( ! [X177,X178,X176,X179] : (s(bool,'T_0') = s(bool,X178) | s(bool,'T_0') = s(bool,X178) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X179),s(bool,'const.bool.COND_3'(s(bool,X176),s(bool,X176),s(bool,X177))))) | s(bool,'F_0') = s(bool,X176)) ) | $spl68),
  inference(superposition,[],[f1050,f10298])).
fof(f12026,plain,(
  ( ! [X206,X207,X205] : (s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X205))) = s(bool,'const.bool.COND_3'(s(bool,X206),s(bool,X206),s(bool,X207))) | s(bool,'F_0') = s(bool,X206) | s('type.realax.real','const.real.abs_1'(s('type.realax.real',X205))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X205)))) ) | $spl68),
  inference(subsumption_resolution,[],[f11952,f71])).
fof(f11952,plain,(
  ( ! [X206,X207,X205] : (~p(s(bool,'T_0')) | s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X205))) = s(bool,'const.bool.COND_3'(s(bool,X206),s(bool,X206),s(bool,X207))) | s(bool,'F_0') = s(bool,X206) | s('type.realax.real','const.real.abs_1'(s('type.realax.real',X205))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X205)))) ) | $spl68),
  inference(superposition,[],[f10604,f6529])).
fof(f11266,plain,(
  ( ! [X167,X165,X169,X166,X168] : (~p(s(bool,X166)) | s(bool,'const.bool.COND_3'(s(bool,X165),s(bool,X166),s(bool,X167))) = s(bool,'const.bool.COND_3'(s(bool,X168),s(bool,X168),s(bool,X169))) | s(bool,'F_0') = s(bool,X168) | s(bool,'F_0') = s(bool,X165)) ) | ($spl12 | $spl68)),
  inference(superposition,[],[f10604,f3095])).
fof(f11265,plain,(
  ( ! [X163,X161,X164,X162,X160] : (~p(s(bool,X162)) | s(bool,'const.bool.COND_3'(s(bool,X160),s(bool,X161),s(bool,X162))) = s(bool,'const.bool.COND_3'(s(bool,X163),s(bool,X163),s(bool,X164))) | s(bool,'F_0') = s(bool,X163) | s(bool,'T_0') = s(bool,X160)) ) | $spl68),
  inference(superposition,[],[f10604,f207])).
fof(f11296,plain,(
  ( ! [X151,X149,X150,X148] : (s(bool,'const.bool.COND_3'(s(bool,X148),s(bool,X148),s(bool,X149))) = s(bool,'const.bool.COND_3'(s(bool,X150),s(bool,X150),s(bool,X151))) | s(bool,'F_0') = s(bool,X150) | s(bool,'F_0') = s(bool,X148)) ) | $spl68),
  inference(subsumption_resolution,[],[f11262,f71])).
fof(f11262,plain,(
  ( ! [X151,X149,X150,X148] : (~p(s(bool,'T_0')) | s(bool,'const.bool.COND_3'(s(bool,X148),s(bool,X148),s(bool,X149))) = s(bool,'const.bool.COND_3'(s(bool,X150),s(bool,X150),s(bool,X151))) | s(bool,'F_0') = s(bool,X150) | s(bool,'F_0') = s(bool,X148)) ) | $spl68),
  inference(superposition,[],[f10604,f10298])).
fof(f11295,plain,(
  ( ! [X146,X144,X147,X145] : (s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X144),s(bool,X145))) = s(bool,'const.bool.COND_3'(s(bool,X146),s(bool,X146),s(bool,X147))) | s(bool,'F_0') = s(bool,X146) | s(bool,'F_0') = s(bool,X145)) ) | $spl68),
  inference(subsumption_resolution,[],[f11261,f71])).
fof(f11261,plain,(
  ( ! [X146,X144,X147,X145] : (~p(s(bool,'T_0')) | s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X144),s(bool,X145))) = s(bool,'const.bool.COND_3'(s(bool,X146),s(bool,X146),s(bool,X147))) | s(bool,'F_0') = s(bool,X146) | s(bool,'F_0') = s(bool,X145)) ) | $spl68),
  inference(superposition,[],[f10604,f219])).
fof(f11292,plain,(
  ( ! [X138,X136,X135,X139,X137] : (s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X135),s(bool,X136))) = s(bool,'const.bool.COND_3'(s(bool,X137),s(bool,X137),s(bool,X138))) | s(bool,'F_0') = s(bool,X137) | s(bool,X135) = s(bool,X139) | s(bool,'T_0') = s(bool,X139)) ) | $spl68),
  inference(subsumption_resolution,[],[f11259,f71])).
fof(f11259,plain,(
  ( ! [X138,X136,X135,X139,X137] : (~p(s(bool,'T_0')) | s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X135),s(bool,X136))) = s(bool,'const.bool.COND_3'(s(bool,X137),s(bool,X137),s(bool,X138))) | s(bool,'F_0') = s(bool,X137) | s(bool,X135) = s(bool,X139) | s(bool,'T_0') = s(bool,X139)) ) | $spl68),
  inference(superposition,[],[f10604,f326])).
fof(f11291,plain,(
  ( ! [X134,X132,X133,X131] : (s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X131),s(bool,X132))) = s(bool,'const.bool.COND_3'(s(bool,X133),s(bool,X133),s(bool,X134))) | s(bool,'F_0') = s(bool,X133) | s(bool,'F_0') = s(bool,X131)) ) | $spl68),
  inference(subsumption_resolution,[],[f11258,f71])).
fof(f11258,plain,(
  ( ! [X134,X132,X133,X131] : (~p(s(bool,'T_0')) | s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X131),s(bool,X132))) = s(bool,'const.bool.COND_3'(s(bool,X133),s(bool,X133),s(bool,X134))) | s(bool,'F_0') = s(bool,X133) | s(bool,'F_0') = s(bool,X131)) ) | $spl68),
  inference(superposition,[],[f10604,f187])).
fof(f11274,plain,(
  ( ! [X54,X52,X56,X55,X53] : (s(bool,X52) = s(bool,'const.bool.COND_3'(s(bool,X53),s(bool,X53),s(bool,X54))) | s(bool,'F_0') = s(bool,X53) | s(bool,X55) = s(bool,X56) | s(bool,'F_0') = s(bool,X56) | s(bool,X52) = s(bool,X55)) ) | $spl68),
  inference(subsumption_resolution,[],[f11238,f71])).
fof(f11238,plain,(
  ( ! [X54,X52,X56,X55,X53] : (~p(s(bool,'T_0')) | s(bool,X52) = s(bool,'const.bool.COND_3'(s(bool,X53),s(bool,X53),s(bool,X54))) | s(bool,'F_0') = s(bool,X53) | s(bool,X55) = s(bool,X56) | s(bool,'F_0') = s(bool,X56) | s(bool,X52) = s(bool,X55)) ) | $spl68),
  inference(superposition,[],[f10604,f1664])).
fof(f11269,plain,(
  ( ! [X26,X24,X23,X25,X22] : (s(bool,X22) = s(bool,'const.bool.COND_3'(s(bool,X23),s(bool,X23),s(bool,X24))) | s(bool,'F_0') = s(bool,X23) | s(bool,X22) = s(bool,X25) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X25),s(bool,X26)))) ) | $spl68),
  inference(subsumption_resolution,[],[f11232,f71])).
fof(f11232,plain,(
  ( ! [X26,X24,X23,X25,X22] : (~p(s(bool,'T_0')) | s(bool,X22) = s(bool,'const.bool.COND_3'(s(bool,X23),s(bool,X23),s(bool,X24))) | s(bool,'F_0') = s(bool,X23) | s(bool,X22) = s(bool,X25) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X25),s(bool,X26)))) ) | $spl68),
  inference(superposition,[],[f10604,f326])).
fof(f11222,plain,(
  ( ! [X151,X149,X150,X148] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X150),s(bool,X151),s(bool,'const.bool.COND_3'(s(bool,X148),s(bool,X148),s(bool,X149))))) | s(bool,'T_0') = s(bool,X150) | s(bool,'F_0') = s(bool,X148)) ) | $spl68),
  inference(subsumption_resolution,[],[f11184,f71])).
fof(f11184,plain,(
  ( ! [X151,X149,X150,X148] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X150),s(bool,X151),s(bool,'const.bool.COND_3'(s(bool,X148),s(bool,X148),s(bool,X149))))) | s(bool,'T_0') = s(bool,X150) | s(bool,'F_0') = s(bool,X148)) ) | $spl68),
  inference(superposition,[],[f6545,f10298])).
fof(f10772,plain,(
  ( ! [X90,X88,X87,X89] : (s(bool,X87) != s(bool,'const.bool.COND_3'(s(bool,X88),s(bool,X88),s(bool,X87))) | s(bool,'F_0') = s(bool,X88) | s(bool,X87) = s(bool,X89) | s(bool,X89) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) ) | $spl68),
  inference(superposition,[],[f10603,f1664])).
fof(f10770,plain,(
  ( ! [X83,X81,X82] : (s(bool,X81) != s(bool,'const.bool.COND_3'(s(bool,X82),s(bool,X82),s(bool,X81))) | s(bool,'F_0') = s(bool,X82) | s(bool,'F_0') = s(bool,X83) | s(bool,X81) = s(bool,X83)) ) | $spl68),
  inference(superposition,[],[f10603,f3059])).
fof(f10745,plain,(
  ( ! [X123,X121,X124,X122] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X122),s(bool,X122),s(bool,X124))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X121),s(bool,X122),s(bool,X123))) | s(bool,'F_0') = s(bool,X121)) ) | ($spl12 | $spl68)),
  inference(superposition,[],[f10603,f3095])).
fof(f10744,plain,(
  ( ! [X118,X120,X119,X117] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X119),s(bool,X119),s(bool,X120))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X117),s(bool,X118),s(bool,X119))) | s(bool,'T_0') = s(bool,X117)) ) | $spl68),
  inference(superposition,[],[f10603,f207])).
fof(f10743,plain,(
  ( ! [X116,X114,X115] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X115),s(bool,X115),s(bool,X116))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X114),s(bool,X115)))) ) | $spl68),
  inference(superposition,[],[f10603,f103])).
fof(f10742,plain,(
  ( ! [X111,X112,X113] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X111),s(bool,X111),s(bool,X113))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X111),s(bool,X112)))) ) | $spl68),
  inference(superposition,[],[f10603,f102])).
fof(f10700,plain,(
  ( ! [X151,X149,X150,X148] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X150),s(bool,X151),s(bool,'const.bool.COND_3'(s(bool,X148),s(bool,X148),s(bool,X149))))) | s(bool,'T_0') = s(bool,X150) | s(bool,'F_0') = s(bool,X148)) ) | $spl68),
  inference(subsumption_resolution,[],[f10670,f71])).
fof(f10670,plain,(
  ( ! [X151,X149,X150,X148] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X150),s(bool,X151),s(bool,'const.bool.COND_3'(s(bool,X148),s(bool,X148),s(bool,X149))))) | s(bool,'T_0') = s(bool,X150) | s(bool,'F_0') = s(bool,X148)) ) | $spl68),
  inference(superposition,[],[f6544,f10298])).
fof(f10509,plain,(
  ( ! [X198,X200,X199,X197] : (p(s(bool,'const.bool.COND_3'(s(bool,X199),s(bool,X200),s(bool,'const.bool.COND_3'(s(bool,X197),s(bool,X197),s(bool,X198)))))) | s(bool,'T_0') = s(bool,X199) | s(bool,'F_0') = s(bool,X197)) ) | $spl68),
  inference(trivial_inequality_removal,[],[f10473])).
fof(f10473,plain,(
  ( ! [X198,X200,X199,X197] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,'const.bool.COND_3'(s(bool,X199),s(bool,X200),s(bool,'const.bool.COND_3'(s(bool,X197),s(bool,X197),s(bool,X198)))))) | s(bool,'T_0') = s(bool,X199) | s(bool,'F_0') = s(bool,X197)) ) | $spl68),
  inference(superposition,[],[f6543,f10298])).
fof(f10615,plain,(
  ( ! [X187,X185,X188,X186] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X187),s(bool,'const.bool.COND_3'(s(bool,X185),s(bool,X185),s(bool,X186))),s(bool,X188))) | s(bool,'F_0') = s(bool,X187) | s(bool,'F_0') = s(bool,X185)) ) | ($spl12 | $spl68)),
  inference(subsumption_resolution,[],[f10470,f71])).
fof(f10470,plain,(
  ( ! [X187,X185,X188,X186] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X187),s(bool,'const.bool.COND_3'(s(bool,X185),s(bool,X185),s(bool,X186))),s(bool,X188))) | s(bool,'F_0') = s(bool,X187) | s(bool,'F_0') = s(bool,X185)) ) | ($spl12 | $spl68)),
  inference(superposition,[],[f5289,f10298])).
fof(f10510,plain,(
  ( ! [X182,X184,X183,X181] : (p(s(bool,'const.bool.COND_3'(s(bool,X183),s(bool,'const.bool.COND_3'(s(bool,X181),s(bool,X181),s(bool,X182))),s(bool,X184)))) | s(bool,'F_0') = s(bool,X183) | s(bool,'F_0') = s(bool,X181)) ) | ($spl12 | $spl68)),
  inference(trivial_inequality_removal,[],[f10469])).
fof(f10469,plain,(
  ( ! [X182,X184,X183,X181] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,'const.bool.COND_3'(s(bool,X183),s(bool,'const.bool.COND_3'(s(bool,X181),s(bool,X181),s(bool,X182))),s(bool,X184)))) | s(bool,'F_0') = s(bool,X183) | s(bool,'F_0') = s(bool,X181)) ) | ($spl12 | $spl68)),
  inference(superposition,[],[f5288,f10298])).
fof(f10604,plain,(
  ( ! [X37,X38,X36] : (~p(s(bool,X38)) | s(bool,X38) = s(bool,'const.bool.COND_3'(s(bool,X36),s(bool,X36),s(bool,X37))) | s(bool,'F_0') = s(bool,X36)) ) | $spl68),
  inference(subsumption_resolution,[],[f10428,f71])).
fof(f10428,plain,(
  ( ! [X37,X38,X36] : (~p(s(bool,'T_0')) | ~p(s(bool,X38)) | s(bool,X38) = s(bool,'const.bool.COND_3'(s(bool,X36),s(bool,X36),s(bool,X37))) | s(bool,'F_0') = s(bool,X36)) ) | $spl68),
  inference(superposition,[],[f107,f10298])).
fof(f10536,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,X0),s(bool,X1))),s(bool,X2))) | s(bool,'F_0') = s(bool,X0)) ) | $spl68),
  inference(condensation,[],[f10535])).
fof(f10535,plain,(
  ( ! [X30,X28,X29,X27] : (s(bool,'T_0') = s(bool,X29) | s(bool,'F_0') = s(bool,X27) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,X27),s(bool,X27),s(bool,X28))),s(bool,X30)))) ) | $spl68),
  inference(duplicate_literal_removal,[],[f10397])).
fof(f10397,plain,(
  ( ! [X30,X28,X29,X27] : (s(bool,'T_0') = s(bool,X29) | s(bool,'F_0') = s(bool,X27) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,X27),s(bool,X27),s(bool,X28))),s(bool,X30))) | s(bool,'T_0') = s(bool,X29)) ) | $spl68),
  inference(superposition,[],[f10298,f326])).
fof(f10373,plain,(
  ( ! [X80,X78,X79,X77] : (s(bool,X77) = s(bool,'const.bool.COND_3'(s(bool,X78),s(bool,X78),s(bool,X77))) | s(bool,'F_0') = s(bool,X78) | s(bool,X77) = s(bool,X79) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X79),s(bool,X80)))) ) | $spl68),
  inference(superposition,[],[f10298,f326])).
fof(f10370,plain,(
  ( ! [X70,X68,X69,X67] : (s(bool,X67) = s(bool,'const.bool.COND_3'(s(bool,X68),s(bool,X68),s(bool,X67))) | s(bool,'F_0') = s(bool,X68) | s(bool,X69) = s(bool,X70) | s(bool,'F_0') = s(bool,X70) | s(bool,X67) = s(bool,X69)) ) | $spl68),
  inference(superposition,[],[f10298,f1664])).
fof(f10349,plain,(
  ( ! [X123,X121,X124,X122] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X122),s(bool,X122),s(bool,X124))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X121),s(bool,X122),s(bool,X123))) | s(bool,'F_0') = s(bool,X121)) ) | ($spl12 | $spl68)),
  inference(superposition,[],[f10298,f3095])).
fof(f10348,plain,(
  ( ! [X118,X120,X119,X117] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X119),s(bool,X119),s(bool,X120))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X117),s(bool,X118),s(bool,X119))) | s(bool,'T_0') = s(bool,X117)) ) | $spl68),
  inference(superposition,[],[f10298,f207])).
fof(f10347,plain,(
  ( ! [X116,X114,X115] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X115),s(bool,X115),s(bool,X116))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X114),s(bool,X115)))) ) | $spl68),
  inference(superposition,[],[f10298,f103])).
fof(f10346,plain,(
  ( ! [X111,X112,X113] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X111),s(bool,X111),s(bool,X113))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X111),s(bool,X112)))) ) | $spl68),
  inference(superposition,[],[f10298,f102])).
fof(f10303,plain,(
  ( ! [X45,X43,X44,X42] : (s(bool,'F_0') = s(bool,X42) | s(bool,X44) = s(bool,'const.bool.COND_3'(s(bool,X42),s(bool,X42),s(bool,X43))) | s(bool,X44) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) ) | $spl68),
  inference(subsumption_resolution,[],[f10231,f73])).
fof(f10231,plain,(
  ( ! [X45,X43,X44,X42] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,X42) | s(bool,X44) = s(bool,'const.bool.COND_3'(s(bool,X42),s(bool,X42),s(bool,X43))) | s(bool,X44) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) ) | $spl68),
  inference(superposition,[],[f9977,f1664])).
fof(f10175,plain,(
  ( ! [X121,X118,X120,X119] : (p(s(bool,'const.bool.COND_3'(s(bool,X119),s(bool,X119),s(bool,X121)))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X118),s(bool,X119),s(bool,X120))) | s(bool,'F_0') = s(bool,X118)) ) | ($spl12 | $spl68)),
  inference(superposition,[],[f9977,f3095])).
fof(f10174,plain,(
  ( ! [X116,X114,X117,X115] : (p(s(bool,'const.bool.COND_3'(s(bool,X116),s(bool,X116),s(bool,X117)))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X114),s(bool,X115),s(bool,X116))) | s(bool,'T_0') = s(bool,X114)) ) | $spl68),
  inference(superposition,[],[f9977,f207])).
fof(f14186,plain,(
  ( ! [X2,X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X2)) ) | ($spl8 | $spl12 | $spl68)),
  inference(global_subsumption,[],[f70,f71,f76,f80,f79,f102,f107,f104,f108,f110,f131,f162,f207,f325,f326,f376,f558,f856,f575,f598,f611,f891,f651,f917,f923,f713,f715,f769,f792,f780,f996,f958,f1072,f1074,f1093,f1168,f1222,f1224,f1226,f1227,f1663,f1664,f1297,f1739,f1783,f2996,f2989,f2966,f3031,f2025,f2948,f2928,f2919,f2916,f2900,f2895,f3048,f2889,f2886,f3059,f2122,f2881,f2871,f3075,f2866,f2843,f2842,f2837,f2836,f3107,f2283,f2285,f3133,f2333,f3137,f3143,f3163,f3172,f3174,f3179,f3180,f2700,f3200,f3202,f103,f118,f187,f219,f264,f417,f456,f493,f523,f542,f576,f612,f921,f934,f937,f957,f1050,f1094,f1799,f1801,f1951,f1952,f1993,f1994,f2024,f2061,f2091,f2121,f2125,f2155,f2156,f2219,f2282,f2366,f2429,f2492,f2495,f2498,f2499,f2500,f2502,f2503,f2588,f2829,f2831,f2835,f3076,f3219,f1665,f938,f1795,f1740,f1223,f714,f263,f249,f111,f455,f454,f73,f83,f84,f453,f2494,f3148,f3201,f3095,f2023,f1944,f915,f450,f416,f4143,f4172,f4173,f4245,f4324,f4382,f4966,f5038,f5039,f4995,f4996,f5044,f5047,f5090,f5091,f5192,f5193,f5194,f5260,f5261,f5262,f5265,f5266,f5268,f5269,f5270,f5271,f5272,f5273,f5274,f5275,f5279,f5300,f5307,f5322,f5323,f5324,f5816,f6124,f6425,f6426,f6427,f6507,f6508,f6509,f6512,f6513,f6514,f6516,f6517,f6518,f6519,f6520,f6521,f6522,f6523,f6529,f6555,f6558,f6565,f6580,f6581,f6582,f6586,f6770,f6841,f6848,f6850,f7090,f6890,f6893,f6944,f6945,f6946,f6947,f7097,f7181,f7195,f7372,f7406,f7407,f7411,f690,f522,f7199,f7180,f7177,f7163,f7162,f7108,f7107,f6583,f6568,f6567,f6564,f6563,f6554,f6549,f6547,f6546,f6545,f6544,f6543,f6515,f6058,f5310,f5309,f5306,f5305,f5299,f5294,f5292,f5291,f5290,f5289,f5288,f5267,f7176,f7246,f7247,f7389,f7387,f5327,f6589,f6839,f6838,f6748,f6747,f5147,f5146,f3170,f3141,f2480,f2417,f2354,f2270,f2207,f1157,f681,f642,f469,f468,f258,f77,f74,f4379,f4321,f2585,f2489,f2426,f2363,f2279,f2216,f1294,f1165,f766,f687,f648,f525,f524,f7987,f7988,f8060,f8062,f8066,f8068,f8072,f8074,f8080,f8057,f8058,f8096,f8097,f8125,f8126,f8127,f8132,f8133,f8135,f8249,f8251,f8235,f8264,f8292,f8294,f8314,f8338,f8339,f8340,f8341,f8568,f8569,f8570,f8571,f8810,f8678,f8679,f8680,f8681,f8724,f8726,f8727,f8832,f8841,f8842,f8844,f8895,f8896,f8897,f8898,f8923,f8925,f8926,f9082,f9107,f9080,f9081,f9314,f9332,f9149,f9150,f9189,f9190,f9375,f9377,f9378,f9546,f9513,f9562,f9563,f9473,f9474,f9475,f9476,f9961,f9983,f9985,f9850,f9851,f10000,f10008,f9977,f10017,f10018,f10095,f10093,f10120,f10090,f10089,f10086,f10073,f10074,f10174,f10175,f10298,f10303,f10346,f10347,f10348,f10349,f10370,f10373,f10536,f10603,f10604,f10510,f10615,f10616,f10509,f10700,f10673,f10674,f10857,f10742,f10743,f10744,f10745,f10770,f10772,f11088,f11118,f11119,f10927,f10928,f11009,f11010,f11192,f11194,f11199,f11200,f11205,f11209,f11222,f11187,f11188,f11269,f11274,f11291,f11292,f11295,f11296,f11265,f11266,f11361,f11362,f11363,f11401,f11402,f11403,f11404,f11555,f11581,f11552,f11553,f11594,f11597,f11649,f11650,f11651,f11652,f11777,f11778,f11779,f11780,f11841,f11842,f11987,f11978,f11989,f11976,f11973,f11995,f11996,f11998,f11962,f12008,f12009,f11961,f12010,f12011,f12026,f12029,f12033,f12034,f12035,f12036,f12037,f12038,f12040,f12041,f13550,f13532,f13522,f12274,f12275,f13591,f12321,f12322,f12323,f12324,f12327,f13601,f12371,f12372,f12373,f12374,f12375,f12377,f13602,f13603,f12423,f12424,f12425,f12427,f12490,f12491,f12492,f12493,f13624,f12642,f13719,f12669,f12674,f13740,f13348,f13775,f13776,f12802,f12829,f12834,f12837,f12840,f12841,f12843,f12844,f12846,f12847,f12849,f12854,f12868,f13313,f13808,f12927,f13824,f13840,f13863,f13084,f13089,f13928,f13215,f13203,f3049,f13962,f13963,f14173,f14174,f14175,f14176,f14177,f14180,f14181,f14185])).
fof(f14181,plain,(
  ( ! [X200,X199] : (s(bool,'T_0') = s(bool,X199) | p(s(bool,X200))) ) | ($spl8 | $spl12 | $spl68)),
  inference(global_subsumption,[],[f14180,f70,f71,f76,f80,f79,f102,f107,f104,f108,f110,f131,f162,f207,f325,f326,f376,f558,f856,f575,f598,f611,f891,f651,f917,f923,f713,f715,f769,f792,f780,f996,f958,f1072,f1074,f1093,f1168,f1222,f1224,f1226,f1227,f1663,f1664,f1297,f1739,f1783,f2996,f2989,f2966,f3031,f2025,f2948,f2928,f2919,f2916,f2900,f2895,f3048,f2889,f2886,f3059,f2122,f2881,f2871,f3075,f2866,f2843,f2842,f2837,f2836,f3107,f2283,f2285,f3133,f2333,f3137,f3143,f3163,f3172,f3174,f3179,f3180,f2700,f3200,f3202,f103,f118,f187,f219,f264,f417,f456,f493,f523,f542,f576,f612,f921,f934,f937,f957,f1050,f1094,f1799,f1801,f1951,f1952,f1993,f1994,f2024,f2061,f2091,f2121,f2125,f2155,f2156,f2219,f2282,f2366,f2429,f2492,f2495,f2498,f2499,f2500,f2502,f2503,f2588,f2829,f2831,f2835,f3076,f3219,f1665,f938,f1795,f1740,f1223,f714,f263,f249,f111,f455,f454,f73,f83,f84,f453,f2494,f3148,f3201,f2023,f1944,f915,f450,f416,f4143,f4172,f4173,f4245,f4324,f4382,f4966,f5038,f5039,f4995,f4996,f5044,f5047,f5090,f5091,f6425,f6426,f6507,f6508,f6509,f6512,f6513,f6514,f6516,f6517,f6518,f6519,f6520,f6521,f6522,f6523,f6529,f6555,f6558,f6565,f6581,f6582,f6586,f6848,f6850,f7090,f6890,f6893,f6944,f6945,f6946,f7097,f690,f522,f6583,f6568,f6567,f6564,f6563,f6554,f6549,f6547,f6546,f6545,f6544,f6543,f6515,f6589,f5147,f5146,f3170,f3141,f2480,f2417,f2354,f2270,f2207,f1157,f681,f642,f469,f468,f258,f77,f74,f4379,f4321,f2585,f2489,f2426,f2363,f2279,f2216,f1294,f1165,f766,f687,f648,f525,f524,f8096,f8125,f8127,f8132,f8133,f8135,f8251,f8235,f8264,f8292,f8294,f8314,f8338,f8340,f8341,f8568,f8569,f8570,f8841,f8842,f8844,f8895,f8896,f8897,f8923,f8925,f8926,f9375,f9377,f9378,f9546,f9513,f9562,f9563,f9473,f9474,f9475,f10017,f10018,f10095,f10093,f10120,f10090,f10089,f10086,f10073,f10673,f11088,f11118,f11119,f10927,f11009,f11192,f11194,f11199,f11200,f11205,f11209,f11187,f11361,f11362,f11401,f11402,f11403,f11555,f11552,f11594,f11597,f11649,f11650,f11651,f11777,f11778,f11779,f11841,f11978,f11989,f11976,f11973,f11995,f11996,f11998,f11961,f12010,f12011,f13550,f13532,f12274,f13591,f12321,f12323,f12324,f12327,f12371,f12372,f12373,f12375,f12377,f13602,f13603,f12423,f12424,f12425,f12427,f12490,f12491,f12492,f13624,f12642,f12669,f12674,f13740,f13348,f13775,f13776,f12802,f12829,f12834,f12837,f12840,f12841,f12843,f12844,f12846,f12847,f12849,f12868,f13313,f13808,f13824,f13840,f13084,f13089,f13215,f13203,f3049,f13962])).
fof(f14177,plain,(
  ( ! [X116,X114] : (p(s(bool,X116)) | s(bool,'T_0') = s(bool,X114)) ) | ($spl8 | $spl12 | $spl68)),
  inference(global_subsumption,[],[f14176,f70,f71,f76,f80,f79,f102,f107,f104,f108,f110,f131,f162,f207,f325,f326,f376,f558,f856,f575,f598,f611,f891,f651,f917,f923,f713,f715,f769,f792,f780,f996,f958,f1072,f1074,f1093,f1168,f1222,f1224,f1226,f1227,f1663,f1664,f1297,f1739,f1783,f2996,f2989,f2966,f3031,f2025,f2948,f2928,f2919,f2916,f2900,f2895,f3048,f2889,f2886,f3059,f2122,f2881,f2871,f3075,f2866,f2843,f2842,f2837,f2836,f3107,f2283,f2285,f3133,f2333,f3137,f3143,f3163,f3172,f3174,f3179,f3180,f2700,f3200,f3202,f103,f118,f187,f219,f264,f417,f456,f493,f523,f542,f576,f612,f921,f934,f937,f957,f1050,f1094,f1799,f1801,f1951,f1952,f1993,f1994,f2024,f2061,f2091,f2121,f2125,f2155,f2156,f2219,f2282,f2366,f2429,f2492,f2495,f2498,f2499,f2500,f2502,f2503,f2588,f2829,f2831,f2835,f3076,f3219,f1665,f938,f1795,f1740,f1223,f714,f263,f249,f111,f455,f454,f73,f83,f84,f453,f2494,f3148,f3201,f2023,f1944,f915,f450,f416,f4143,f4172,f4173,f4245,f4324,f4382,f4966,f5038,f5039,f4995,f4996,f5044,f5047,f5090,f5091,f6425,f6426,f6507,f6508,f6509,f6512,f6513,f6514,f6516,f6517,f6518,f6519,f6520,f6521,f6522,f6523,f6529,f6555,f6558,f6565,f6581,f6582,f6586,f6848,f6850,f7090,f6890,f6893,f6944,f6945,f6946,f7097,f690,f522,f6583,f6568,f6567,f6564,f6563,f6554,f6549,f6547,f6546,f6545,f6544,f6543,f6515,f6589,f5147,f5146,f3170,f3141,f2480,f2417,f2354,f2270,f2207,f1157,f681,f642,f469,f468,f258,f77,f74,f4379,f4321,f2585,f2489,f2426,f2363,f2279,f2216,f1294,f1165,f766,f687,f648,f525,f524,f8096,f8125,f8127,f8132,f8133,f8135,f8251,f8235,f8264,f8292,f8294,f8314,f8338,f8340,f8341,f8568,f8569,f8570,f8841,f8842,f8844,f8895,f8896,f8897,f8923,f8925,f8926,f9375,f9377,f9378,f9546,f9513,f9562,f9563,f9473,f9474,f9475,f10017,f10018,f10095,f10093,f10120,f10090,f10089,f10086,f10073,f10673,f11088,f11118,f11119,f10927,f11009,f11192,f11194,f11199,f11200,f11205,f11209,f11187,f11361,f11362,f11401,f11402,f11403,f11555,f11552,f11594,f11597,f11649,f11650,f11651,f11777,f11778,f11779,f11841,f11978,f11989,f11976,f11973,f11995,f11996,f11998,f11961,f12010,f12011,f13550,f13532,f12274,f13591,f12321,f12323,f12324,f12327,f12371,f12372,f12373,f12375,f12377,f13602,f13603,f12423,f12424,f12425,f12427,f12490,f12491,f12492,f13624,f12642,f12669,f12674,f13740,f13348,f13775,f13776,f12802,f12829,f12834,f12837,f12840,f12841,f12843,f12844,f12846,f12847,f12849,f12868,f13313,f13808,f13824,f13840,f13084,f13089,f13215,f13203,f3049,f13962])).
fof(f14175,plain,(
  ( ! [X165,X162] : (p(s(bool,X162)) | s(bool,'T_0') = s(bool,X165)) ) | ($spl8 | $spl12)),
  inference(global_subsumption,[],[f14174,f70,f71,f76,f80,f79,f102,f107,f104,f108,f110,f131,f162,f207,f325,f326,f376,f558,f856,f575,f598,f611,f891,f651,f917,f923,f713,f715,f769,f792,f780,f996,f958,f1072,f1074,f1093,f1168,f1222,f1224,f1226,f1227,f1663,f1664,f1297,f1739,f1783,f2996,f2989,f2966,f3031,f2025,f2948,f2928,f2919,f2916,f2900,f2895,f3048,f2889,f2886,f3059,f2122,f2881,f2871,f3075,f2866,f2843,f2842,f2837,f2836,f3107,f2283,f2285,f3133,f2333,f3137,f3143,f3163,f3172,f3174,f3179,f3180,f2700,f3200,f3202,f103,f118,f187,f219,f264,f417,f456,f493,f523,f542,f576,f612,f921,f934,f937,f957,f1050,f1094,f1799,f1801,f1951,f1952,f1993,f1994,f2024,f2061,f2091,f2121,f2125,f2155,f2156,f2219,f2282,f2366,f2429,f2492,f2495,f2498,f2499,f2500,f2502,f2503,f2588,f2829,f2831,f2835,f3076,f3219,f1665,f938,f1795,f1740,f1223,f714,f263,f249,f111,f455,f454,f73,f83,f84,f453,f2494,f3148,f3201,f2023,f1944,f915,f450,f416,f4143,f4172,f4173,f4245,f4324,f4382,f4966,f5038,f5039,f4995,f4996,f5044,f5047,f5090,f5091,f6425,f6426,f6507,f6508,f6509,f6512,f6513,f6514,f6516,f6517,f6518,f6519,f6520,f6521,f6522,f6523,f6529,f6555,f6558,f6565,f6581,f6582,f6586,f6848,f6850,f7090,f6890,f6893,f6944,f6945,f6946,f7097,f690,f522,f6583,f6568,f6567,f6564,f6563,f6554,f6549,f6547,f6546,f6545,f6544,f6543,f6515,f6589,f5147,f5146,f3170,f3141,f2480,f2417,f2354,f2270,f2207,f1157,f681,f642,f469,f468,f258,f77,f74,f4379,f4321,f2585,f2489,f2426,f2363,f2279,f2216,f1294,f1165,f766,f687,f648,f525,f524,f8096,f8125,f8127,f8132,f8133,f8135,f8251,f8235,f8264,f8292,f8294,f8314,f8338,f8340,f8341,f8568,f8569,f8570,f8841,f8842,f8844,f8895,f8896,f8897,f8923,f8925,f8926,f9375,f9377,f9378,f9546,f9513,f9562,f9563,f9473,f9474,f9475,f10017,f10018,f10095,f10093,f10120,f10090,f10089,f10086,f10073,f10673,f11088,f11118,f11119,f10927,f11009,f11192,f11194,f11199,f11200,f11205,f11209,f11187,f11361,f11362,f11401,f11402,f11403,f11555,f11552,f11594,f11597,f11649,f11650,f11651,f11777,f11778,f11779,f11841,f11978,f11989,f11976,f11973,f11995,f11996,f11998,f11961,f12010,f12011,f13550,f13532,f12274,f13591,f12321,f12323,f12324,f12327,f12371,f12372,f12373,f12375,f12377,f13602,f13603,f12423,f12424,f12425,f12427,f12490,f12491,f12492,f13624,f12642,f12669,f12674,f13740,f13348,f13775,f13776,f12802,f12829,f12834,f12837,f12840,f12841,f12843,f12844,f12846,f12847,f12849,f12868,f13313,f13808,f13824,f13840,f13084,f13089,f13215,f13203,f3049,f13962])).
fof(f14173,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1))) ) | ($spl8 | $spl12)),
  inference(condensation,[],[f14172])).
fof(f14172,plain,(
  ( ! [X158,X156,X160] : (p(s(bool,X158)) | s(bool,'T_0') = s(bool,X160) | s(bool,'T_0') = s(bool,X156)) ) | ($spl8 | $spl12)),
  inference(subsumption_resolution,[],[f13966,f13962])).
fof(f13966,plain,(
  ( ! [X158,X156,X157,X160] : (p(s(bool,X158)) | s(bool,'T_0') = s(bool,X160) | s(bool,X160) = s(bool,'const.bool.COND_3'(s(bool,X156),s(bool,X157),s(bool,X158))) | s(bool,'T_0') = s(bool,X156)) ) | ($spl8 | $spl12)),
  inference(backward_demodulation,[],[f13963,f9850])).
fof(f13961,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f13955,f110])).
fof(f13955,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl8),
  inference(duplicate_literal_removal,[],[f13954])).
fof(f13954,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,'F_0') | $spl8),
  inference(condensation,[],[f13934])).
fof(f13934,plain,(
  ( ! [X91,X89] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X89) = s(bool,X91) | s(bool,'F_0') = s(bool,X91)) ) | $spl8),
  inference(backward_demodulation,[],[f3049,f4172])).
fof(f13960,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f13958,f110])).
fof(f13958,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl8),
  inference(condensation,[],[f13957])).
fof(f13957,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X1)) ) | $spl8),
  inference(duplicate_literal_removal,[],[f13956])).
fof(f13956,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | $spl8),
  inference(condensation,[],[f13933])).
fof(f13933,plain,(
  ( ! [X118,X117,X115] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X115) = s(bool,X117) | s(bool,X117) = s(bool,X118) | s(bool,'T_0') = s(bool,X118)) ) | $spl8),
  inference(backward_demodulation,[],[f3049,f2023])).
fof(f3049,plain,(
  ( ! [X25,X22] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X22),s(bool,X25)))) ) | $spl8),
  inference(cnf_transformation,[],[f3049_D])).
fof(f3049_D,plain,(
  ( ! [X25,X22] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X22),s(bool,X25)))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f13741,plain,(
  ( ! [X100,X98] : (s(bool,'T_0') != s(bool,X98) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X100),s(bool,X98)))) ) | $spl80),
  inference(cnf_transformation,[],[f13741_D])).
fof(f13741_D,plain,(
  ( ! [X100,X98] : (s(bool,'T_0') != s(bool,X98) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X100),s(bool,X98)))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f13740,plain,(
  ( ! [X99,X97,X100,X98] : (s(bool,'T_0') != s(bool,X98) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X97),s(bool,X99))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X100),s(bool,X98)))) )),
  inference(subsumption_resolution,[],[f12699,f376])).
fof(f12699,plain,(
  ( ! [X99,X97,X100,X98] : (s(bool,'T_0') != s(bool,X98) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X97),s(bool,X99))) | s(bool,'T_0') = s(bool,X97) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X100),s(bool,X98)))) )),
  inference(superposition,[],[f376,f1050])).
fof(f11990,plain,(
  ( ! [X30,X28] : (s('type.realax.real','const.real.abs_1'(s('type.realax.real',X28))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X28))) | s(bool,'F_0') = s(bool,X30) | s(bool,X30) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X28)))) ) | $spl78),
  inference(cnf_transformation,[],[f11990_D])).
fof(f11990_D,plain,(
  ( ! [X30,X28] : (s('type.realax.real','const.real.abs_1'(s('type.realax.real',X28))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X28))) | s(bool,'F_0') = s(bool,X30) | s(bool,X30) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X28)))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f11976,plain,(
  ( ! [X30,X28,X29] : (s(bool,'T_0') = s(bool,X29) | s('type.realax.real','const.real.abs_1'(s('type.realax.real',X28))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X28))) | s(bool,X30) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X28))) | s(bool,'F_0') = s(bool,X30)) )),
  inference(duplicate_literal_removal,[],[f11855])).
fof(f11855,plain,(
  ( ! [X30,X28,X29] : (s(bool,'T_0') = s(bool,X29) | s('type.realax.real','const.real.abs_1'(s('type.realax.real',X28))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X28))) | s(bool,X30) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X28))) | s(bool,'F_0') = s(bool,X30) | s(bool,'T_0') = s(bool,X29)) )),
  inference(superposition,[],[f6529,f1664])).
fof(f11206,plain,(
  ( ! [X2,X1] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X1),s(bool,X2),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X1)) ) | $spl76),
  inference(cnf_transformation,[],[f11206_D])).
fof(f11206_D,plain,(
  ( ! [X2,X1] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X1),s(bool,X2),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X1)) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f11205,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X1),s(bool,X2),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(condensation,[],[f11204])).
fof(f11204,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X1),s(bool,X2),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X1) | s(bool,'F_0') = s(bool,X3)) )),
  inference(condensation,[],[f11203])).
fof(f11203,plain,(
  ( ! [X90,X88,X87,X91,X89] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X88),s(bool,X89),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X88) | s(bool,X90) = s(bool,X91) | s(bool,'F_0') = s(bool,X91) | s(bool,X87) = s(bool,X90)) )),
  inference(subsumption_resolution,[],[f11168,f3148])).
fof(f11168,plain,(
  ( ! [X90,X88,X87,X91,X89] : (~p(s(bool,X87)) | s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X88),s(bool,X89),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X88) | s(bool,X90) = s(bool,X91) | s(bool,'F_0') = s(bool,X91) | s(bool,X87) = s(bool,X90)) )),
  inference(superposition,[],[f6545,f1664])).
fof(f11111,plain,(
  ( ! [X66,X64,X65,X63] : (p(s(bool,'const.bool.COND_3'(s(bool,X63),s(bool,X64),s(bool,X65)))) | s(bool,'F_0') = s(bool,X66) | s(bool,X63) = s(bool,X66) | s(bool,'F_0') = s(bool,X65)) ) | $spl74),
  inference(cnf_transformation,[],[f11111_D])).
fof(f11111_D,plain,(
  ( ! [X66,X64,X65,X63] : (p(s(bool,'const.bool.COND_3'(s(bool,X63),s(bool,X64),s(bool,X65)))) | s(bool,'F_0') = s(bool,X66) | s(bool,X63) = s(bool,X66) | s(bool,'F_0') = s(bool,X65)) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f11088,plain,(
  ( ! [X66,X64,X62,X65,X63] : (p(s(bool,'const.bool.COND_3'(s(bool,X63),s(bool,X64),s(bool,X65)))) | s(bool,'T_0') = s(bool,X62) | s(bool,'F_0') = s(bool,X65) | s(bool,X63) = s(bool,X66) | s(bool,'F_0') = s(bool,X66)) )),
  inference(duplicate_literal_removal,[],[f10902])).
fof(f10902,plain,(
  ( ! [X66,X64,X62,X65,X63] : (p(s(bool,'const.bool.COND_3'(s(bool,X63),s(bool,X64),s(bool,X65)))) | s(bool,'T_0') = s(bool,X62) | s(bool,'F_0') = s(bool,X65) | s(bool,X63) = s(bool,X66) | s(bool,'F_0') = s(bool,X66) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f10120,f1664])).
fof(f10867,plain,(
  ( ! [X35,X33,X34] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X33),s(bool,X33),s(bool,X34))) | s(bool,'T_0') = s(bool,X35) | s(bool,X33) = s(bool,X35)) ) | $spl72),
  inference(cnf_transformation,[],[f10867_D])).
fof(f10867_D,plain,(
  ( ! [X35,X33,X34] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X33),s(bool,X33),s(bool,X34))) | s(bool,'T_0') = s(bool,X35) | s(bool,X33) = s(bool,X35)) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f10857,plain,(
  ( ! [X35,X33,X34,X32] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X33),s(bool,X33),s(bool,X34))) | s(bool,'F_0') = s(bool,X32) | s(bool,X33) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) ) | $spl68),
  inference(duplicate_literal_removal,[],[f10715])).
fof(f10715,plain,(
  ( ! [X35,X33,X34,X32] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X33),s(bool,X33),s(bool,X34))) | s(bool,'F_0') = s(bool,X32) | s(bool,'F_0') = s(bool,X32) | s(bool,X33) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) ) | $spl68),
  inference(superposition,[],[f10603,f1664])).
fof(f10126,plain,(
  ( ! [X2,X1] : (p(s(bool,'const.bool.COND_3'(s(bool,X1),s(bool,X2),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,X1)) ) | $spl70),
  inference(cnf_transformation,[],[f10126_D])).
fof(f10126_D,plain,(
  ( ! [X2,X1] : (p(s(bool,'const.bool.COND_3'(s(bool,X1),s(bool,X2),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,X1)) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f10089,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,'const.bool.COND_3'(s(bool,X1),s(bool,X2),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(condensation,[],[f10088])).
fof(f10088,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,'const.bool.COND_3'(s(bool,X1),s(bool,X2),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,X1) | s(bool,'F_0') = s(bool,X3)) )),
  inference(condensation,[],[f10087])).
fof(f10087,plain,(
  ( ! [X90,X88,X87,X91,X89] : (p(s(bool,'const.bool.COND_3'(s(bool,X88),s(bool,X89),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,X88) | s(bool,X90) = s(bool,X91) | s(bool,'F_0') = s(bool,X91) | s(bool,X87) = s(bool,X90)) )),
  inference(trivial_inequality_removal,[],[f10055])).
fof(f10055,plain,(
  ( ! [X90,X88,X87,X91,X89] : (s(bool,X87) != s(bool,X87) | p(s(bool,'const.bool.COND_3'(s(bool,X88),s(bool,X89),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,X88) | s(bool,X90) = s(bool,X91) | s(bool,'F_0') = s(bool,X91) | s(bool,X87) = s(bool,X90)) )),
  inference(superposition,[],[f6543,f1664])).
fof(f9977,plain,(
  ( ! [X2,X1] : (p(s(bool,'const.bool.COND_3'(s(bool,X1),s(bool,X1),s(bool,X2)))) | s(bool,'F_0') = s(bool,X1)) ) | $spl68),
  inference(cnf_transformation,[],[f9977_D])).
fof(f9977_D,plain,(
  ( ! [X2,X1] : (p(s(bool,'const.bool.COND_3'(s(bool,X1),s(bool,X1),s(bool,X2)))) | s(bool,'F_0') = s(bool,X1)) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f9995,plain,(
  ( ! [X142,X143,X141] : (p(s(bool,'const.bool.COND_3'(s(bool,X141),s(bool,X141),s(bool,X142)))) | s(bool,'T_0') = s(bool,X143) | s(bool,'F_0') = s(bool,X141)) ) | ($spl12 | $spl26)),
  inference(subsumption_resolution,[],[f9846,f792])).
fof(f9846,plain,(
  ( ! [X142,X140,X143,X141] : (p(s(bool,'const.bool.COND_3'(s(bool,X141),s(bool,X141),s(bool,X142)))) | s(bool,'T_0') = s(bool,X143) | s(bool,X143) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X140))) | s(bool,'F_0') = s(bool,X141)) ) | ($spl12 | $spl26)),
  inference(superposition,[],[f7389,f3173])).
fof(f9985,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,'const.bool.COND_3'(s(bool,X1),s(bool,X1),s(bool,X2)))) | s(bool,'F_0') = s(bool,X1)) ) | $spl12),
  inference(condensation,[],[f9984])).
fof(f9984,plain,(
  ( ! [X78,X76,X77,X75] : (p(s(bool,'const.bool.COND_3'(s(bool,X76),s(bool,X76),s(bool,X77)))) | s(bool,'T_0') = s(bool,X78) | s(bool,X75) = s(bool,X78) | s(bool,'F_0') = s(bool,X76)) ) | $spl12),
  inference(subsumption_resolution,[],[f9829,f1663])).
fof(f9829,plain,(
  ( ! [X78,X76,X77,X75] : (p(s(bool,'const.bool.COND_3'(s(bool,X76),s(bool,X76),s(bool,X77)))) | s(bool,'T_0') = s(bool,X78) | s(bool,X75) = s(bool,X78) | s(bool,'F_0') = s(bool,X76) | s(bool,'F_0') = s(bool,X75)) ) | $spl12),
  inference(superposition,[],[f7389,f3059])).
fof(f9983,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,'const.bool.COND_3'(s(bool,X1),s(bool,X1),s(bool,X2)))) | s(bool,'F_0') = s(bool,X1)) ) | $spl12),
  inference(condensation,[],[f9982])).
fof(f9982,plain,(
  ( ! [X74,X72,X71,X73] : (p(s(bool,'const.bool.COND_3'(s(bool,X72),s(bool,X72),s(bool,X73)))) | s(bool,'T_0') = s(bool,X74) | s(bool,X71) = s(bool,X74) | s(bool,'F_0') = s(bool,X72)) ) | $spl12),
  inference(subsumption_resolution,[],[f9828,f1663])).
fof(f9828,plain,(
  ( ! [X74,X72,X71,X73] : (p(s(bool,'const.bool.COND_3'(s(bool,X72),s(bool,X72),s(bool,X73)))) | s(bool,'T_0') = s(bool,X74) | s(bool,X71) = s(bool,X74) | s(bool,'F_0') = s(bool,X71) | s(bool,'F_0') = s(bool,X72)) ) | $spl12),
  inference(superposition,[],[f7389,f3059])).
fof(f9961,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,'const.bool.COND_3'(s(bool,X1),s(bool,X1),s(bool,X2)))) | s(bool,'F_0') = s(bool,X1)) ) | $spl12),
  inference(condensation,[],[f9960])).
fof(f9960,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,'const.bool.COND_3'(s(bool,X1),s(bool,X1),s(bool,X2)))) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X3)) ) | $spl12),
  inference(duplicate_literal_removal,[],[f9959])).
fof(f9959,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,'const.bool.COND_3'(s(bool,X1),s(bool,X1),s(bool,X2)))) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) ) | $spl12),
  inference(condensation,[],[f9821])).
fof(f9821,plain,(
  ( ! [X39,X37,X41,X38,X40] : (p(s(bool,'const.bool.COND_3'(s(bool,X38),s(bool,X38),s(bool,X39)))) | s(bool,'T_0') = s(bool,X40) | s(bool,X37) = s(bool,X40) | s(bool,'F_0') = s(bool,X38) | s(bool,X37) = s(bool,X41) | s(bool,'T_0') = s(bool,X41)) ) | $spl12),
  inference(superposition,[],[f7389,f1664])).
fof(f9801,plain,(
  ( ! [X70,X68,X71,X67] : (s(bool,'T_0') = s(bool,X70) | s(bool,'F_0') = s(bool,X71) | s(bool,X71) = s(bool,'const.bool.COND_3'(s(bool,X67),s(bool,'T_0'),s(bool,X68))) | s(bool,X67) = s(bool,X70)) ) | $spl66),
  inference(cnf_transformation,[],[f9801_D])).
fof(f9801_D,plain,(
  ( ! [X70,X68,X71,X67] : (s(bool,'T_0') = s(bool,X70) | s(bool,'F_0') = s(bool,X71) | s(bool,X71) = s(bool,'const.bool.COND_3'(s(bool,X67),s(bool,'T_0'),s(bool,X68))) | s(bool,X67) = s(bool,X70)) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f9800,plain,(
  ( ! [X70,X68,X71,X69,X67] : (p(s(bool,X69)) | s(bool,'T_0') = s(bool,X70) | s(bool,X67) = s(bool,X70) | s(bool,X71) = s(bool,'const.bool.COND_3'(s(bool,X67),s(bool,'T_0'),s(bool,X68))) | s(bool,'F_0') = s(bool,X71)) ) | $spl46),
  inference(subsumption_resolution,[],[f9697,f79])).
fof(f9697,plain,(
  ( ! [X70,X68,X71,X69,X67] : (p(s(bool,X69)) | s(bool,'T_0') = s(bool,X70) | s(bool,X67) = s(bool,X70) | s(bool,X71) = s(bool,'const.bool.COND_3'(s(bool,X67),s(bool,'T_0'),s(bool,X68))) | s(bool,'F_0') = s(bool,X71) | s(bool,'T_0') = s(bool,X69)) ) | $spl46),
  inference(superposition,[],[f7383,f1664])).
fof(f9555,plain,(
  ( ! [X35,X33,X34] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X34),s(bool,X33))) | s(bool,'T_0') = s(bool,X35) | s(bool,X33) = s(bool,X35)) ) | $spl64),
  inference(cnf_transformation,[],[f9555_D])).
fof(f9555_D,plain,(
  ( ! [X35,X33,X34] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X34),s(bool,X33))) | s(bool,'T_0') = s(bool,X35) | s(bool,X33) = s(bool,X35)) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f9513,plain,(
  ( ! [X35,X33,X34,X32] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X34),s(bool,X33))) | s(bool,'F_0') = s(bool,X32) | s(bool,X33) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) )),
  inference(duplicate_literal_removal,[],[f9447])).
fof(f9447,plain,(
  ( ! [X35,X33,X34,X32] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X34),s(bool,X33))) | s(bool,'F_0') = s(bool,X32) | s(bool,'F_0') = s(bool,X32) | s(bool,X33) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) )),
  inference(superposition,[],[f7097,f1664])).
fof(f8820,plain,(
  ( ! [X35,X33,X34] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X33),s(bool,'T_0'),s(bool,X34))) | s(bool,'T_0') = s(bool,X35) | s(bool,X33) = s(bool,X35)) ) | $spl62),
  inference(cnf_transformation,[],[f8820_D])).
fof(f8820_D,plain,(
  ( ! [X35,X33,X34] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X33),s(bool,'T_0'),s(bool,X34))) | s(bool,'T_0') = s(bool,X35) | s(bool,X33) = s(bool,X35)) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f8810,plain,(
  ( ! [X35,X33,X34,X32] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X33),s(bool,'T_0'),s(bool,X34))) | s(bool,'F_0') = s(bool,X32) | s(bool,X33) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) ) | $spl12),
  inference(duplicate_literal_removal,[],[f8652])).
fof(f8652,plain,(
  ( ! [X35,X33,X34,X32] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X33),s(bool,'T_0'),s(bool,X34))) | s(bool,'F_0') = s(bool,X32) | s(bool,'F_0') = s(bool,X32) | s(bool,X33) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) ) | $spl12),
  inference(superposition,[],[f8080,f1664])).
fof(f8252,plain,(
  ( ! [X13] : (s(bool,X13) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))))) | s(bool,'F_0') = s(bool,X13)) ) | $spl60),
  inference(cnf_transformation,[],[f8252_D])).
fof(f8252_D,plain,(
  ( ! [X13] : (s(bool,X13) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))))) | s(bool,'F_0') = s(bool,X13)) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f8235,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') = s(bool,X12) | s(bool,X13) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))))) | s(bool,'F_0') = s(bool,X13)) )),
  inference(duplicate_literal_removal,[],[f8155])).
fof(f8155,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') = s(bool,X12) | s(bool,X13) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))))) | s(bool,'F_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f8125,f1664])).
fof(f8134,plain,(
  ( ! [X35,X33] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X33))))) | s(bool,'F_0') = s(bool,X35) | s(bool,X35) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X33)))) ) | $spl58),
  inference(cnf_transformation,[],[f8134_D])).
fof(f8134_D,plain,(
  ( ! [X35,X33] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X33))))) | s(bool,'F_0') = s(bool,X35) | s(bool,X35) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X33)))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f8133,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X34)) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X33))))) | s(bool,X35) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X33))) | s(bool,'F_0') = s(bool,X35)) )),
  inference(subsumption_resolution,[],[f8114,f79])).
fof(f8114,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X34)) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X33))))) | s(bool,X35) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X33))) | s(bool,'F_0') = s(bool,X35) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f258,f1664])).
fof(f8067,plain,(
  ( ! [X2,X3,X1] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X1),s(bool,X2),s(bool,X3))) | s(bool,'F_0') = s(bool,X2) | s(bool,'F_0') = s(bool,X1)) ) | $spl54),
  inference(cnf_transformation,[],[f8067_D])).
fof(f8067_D,plain,(
  ( ! [X2,X3,X1] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X1),s(bool,X2),s(bool,X3))) | s(bool,'F_0') = s(bool,X2) | s(bool,'F_0') = s(bool,X1)) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f8073,plain,(
  ( ! [X72] : (~p(s(bool,X72))) ) | $spl56),
  inference(cnf_transformation,[],[f8073_D])).
fof(f8073_D,plain,(
  ( ! [X72] : (~p(s(bool,X72))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f8074,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,X76)) | s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X77),s(bool,X75),s(bool,X78))) | s(bool,'F_0') = s(bool,X77) | s(bool,'F_0') = s(bool,X75)) ) | $spl12),
  inference(subsumption_resolution,[],[f8037,f83])).
fof(f8037,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,X76)) | s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X77),s(bool,X75),s(bool,X78))) | s(bool,'F_0') = s(bool,X77) | s(bool,'F_0') = s(bool,X76) | s(bool,'F_0') = s(bool,X75)) ) | $spl12),
  inference(superposition,[],[f5290,f3059])).
fof(f8072,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,X72)) | s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X73),s(bool,X71),s(bool,X74))) | s(bool,'F_0') = s(bool,X73) | s(bool,'F_0') = s(bool,X71)) ) | $spl12),
  inference(subsumption_resolution,[],[f8036,f83])).
fof(f8036,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,X72)) | s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X73),s(bool,X71),s(bool,X74))) | s(bool,'F_0') = s(bool,X73) | s(bool,'F_0') = s(bool,X71) | s(bool,'F_0') = s(bool,X72)) ) | $spl12),
  inference(superposition,[],[f5290,f3059])).
fof(f8066,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X1),s(bool,X2),s(bool,X3))) | s(bool,'F_0') = s(bool,X1) | s(bool,'F_0') = s(bool,X2)) ) | $spl12),
  inference(condensation,[],[f8065])).
fof(f8065,plain,(
  ( ! [X45,X43,X46,X44,X42] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X44),s(bool,X42),s(bool,X45))) | s(bool,'F_0') = s(bool,X44) | s(bool,'F_0') = s(bool,X42) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) ) | $spl12),
  inference(subsumption_resolution,[],[f8030,f162])).
fof(f8030,plain,(
  ( ! [X45,X43,X46,X44,X42] : (~p(s(bool,X43)) | s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X44),s(bool,X42),s(bool,X45))) | s(bool,'F_0') = s(bool,X44) | s(bool,'F_0') = s(bool,X42) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) ) | $spl12),
  inference(superposition,[],[f5290,f1664])).
fof(f7567,plain,(
  $false | ($spl12 | $spl44)),
  inference(global_subsumption,[],[f70,f71,f76,f80,f79,f102,f107,f104,f108,f110,f131,f162,f207,f325,f326,f376,f558,f856,f575,f598,f611,f891,f651,f917,f923,f713,f715,f769,f792,f780,f996,f958,f1072,f1074,f1093,f1168,f1222,f1224,f1226,f1227,f1663,f1664,f1297,f1739,f1783,f2996,f2989,f2966,f3031,f2025,f2948,f2928,f2919,f2916,f2900,f2895,f3048,f2889,f2886,f3059,f2122,f2881,f2871,f3075,f2866,f2843,f2842,f2837,f2836,f3107,f2283,f2285,f3133,f2333,f3137,f3143,f3163,f3172,f3174,f3179,f3180,f2700,f3200,f3202,f103,f118,f187,f219,f264,f417,f456,f493,f523,f542,f576,f612,f921,f934,f937,f957,f1050,f1094,f1799,f1801,f1951,f1952,f1993,f1994,f2024,f2061,f2091,f2121,f2125,f2155,f2156,f2219,f2282,f2366,f2429,f2492,f2495,f2498,f2499,f2500,f2502,f2503,f2588,f2829,f2831,f2835,f3076,f3219,f1665,f938,f1795,f1740,f1223,f714,f263,f249,f111,f455,f454,f73,f83,f84,f453,f2494,f3148,f3201,f3095,f690,f522,f3170,f3141,f2480,f2417,f2354,f2270,f2207,f1157,f681,f642,f469,f468,f258,f77,f74,f2585,f2489,f2426,f2363,f2279,f2216,f1294,f1165,f766,f687,f648,f525,f524,f2023,f1944,f915,f450,f416,f4143,f4172,f4173,f4245,f4321,f4324,f4379,f4382,f4966,f5038,f5039,f4995,f4996,f5044,f5047,f5090,f5091,f5146,f5147,f5192,f5193,f5194,f5260,f5261,f5262,f5265,f5266,f5267,f5268,f5269,f5270,f5271,f5272,f5273,f5274,f5275,f5279,f5288,f5289,f5290,f5291,f5292,f5294,f5299,f5300,f5305,f5306,f5307,f5309,f5310,f5322,f5323,f5324,f5327,f5816,f6058,f6124,f6425,f6426,f6427,f6507,f6508,f6509,f6512,f6513,f6514,f6515,f6516,f6517,f6518,f6519,f6520,f6521,f6522,f6523,f6529,f6543,f6544,f6545,f6546,f6547,f6549,f6554,f6555,f6558,f6563,f6564,f6565,f6567,f6568,f6580,f6581,f6582,f6583,f6586,f6589,f6747,f6748,f6838,f6770,f6839,f6841,f6848,f6850,f7090,f6890,f6893,f6944,f6945,f6946,f6947,f7097,f7107,f7108,f7181,f7180,f7195,f7177,f7199,f7176,f7162,f7163,f7372,f7387,f7389,f7246,f7247,f7406,f7407,f7411,f7382,f7530,f7528,f7535,f7527,f7539,f7428,f7543,f7434,f7544,f7546,f7548,f7549,f7552,f7554,f7507,f7557,f7503,f7560,f7499,f7564,f7486,f7487])).
fof(f7487,plain,(
  ( ! [X144,X142,X140,X143,X141] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X144))) | s(X142,f(s(bool,X143))) != s(X142,g(s(bool,X143))) | s(fun(bool,X142),X140) = s(fun(bool,X142),X141) | s(bool,'T_0') = s(bool,X143)) ) | $spl44),
  inference(backward_demodulation,[],[f7382,f690])).
fof(f7486,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X3))) | s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1)) ) | $spl44),
  inference(backward_demodulation,[],[f7382,f522])).
fof(f7564,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1))) ) | $spl44),
  inference(condensation,[],[f7563])).
fof(f7563,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) ) | $spl44),
  inference(subsumption_resolution,[],[f7497,f1663])).
fof(f7497,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X1) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) ) | $spl44),
  inference(condensation,[],[f7474])).
fof(f7474,plain,(
  ( ! [X154,X152,X153,X150] : (s(bool,'F_0') = s(bool,X153) | ~p(s(bool,X152)) | s(bool,X153) = s(bool,X154) | s(bool,'T_0') = s(bool,X154) | s(bool,'T_0') = s(bool,X150)) ) | $spl44),
  inference(backward_demodulation,[],[f7382,f6567])).
fof(f7499,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X0))) | s(bool,'T_0') != s(bool,X1)) ) | $spl44),
  inference(condensation,[],[f7472])).
fof(f7472,plain,(
  ( ! [X134,X136,X137] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X137))) | s(bool,'T_0') != s(bool,X136) | s(bool,'T_0') = s(bool,X134)) ) | $spl44),
  inference(backward_demodulation,[],[f7382,f6563])).
fof(f7560,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X1)) ) | $spl44),
  inference(subsumption_resolution,[],[f7501,f83])).
fof(f7501,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1)) ) | $spl44),
  inference(condensation,[],[f7470])).
fof(f7470,plain,(
  ( ! [X76,X74,X77] : (s(bool,'F_0') = s(bool,X77) | ~p(s(bool,X76)) | ~p(s(bool,X77)) | s(bool,'T_0') = s(bool,X74)) ) | $spl44),
  inference(backward_demodulation,[],[f7382,f6549])).
fof(f7503,plain,(
  ( ! [X62,X60] : (~p(s(bool,X62)) | s(bool,'T_0') = s(bool,X60)) ) | $spl44),
  inference(trivial_inequality_removal,[],[f7468])).
fof(f7468,plain,(
  ( ! [X62,X60] : (s(bool,'F_0') != s(bool,'F_0') | ~p(s(bool,X62)) | s(bool,'T_0') = s(bool,X60)) ) | $spl44),
  inference(backward_demodulation,[],[f7382,f6545])).
fof(f7557,plain,(
  ( ! [X54,X56] : (s(bool,'T_0') != s(bool,X56) | s(bool,'T_0') = s(bool,X54)) ) | $spl44),
  inference(subsumption_resolution,[],[f7466,f73])).
fof(f7466,plain,(
  ( ! [X54,X56] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,X56) | s(bool,'T_0') = s(bool,X54)) ) | $spl44),
  inference(backward_demodulation,[],[f7382,f6543])).
fof(f7507,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X0)))) ) | $spl44),
  inference(condensation,[],[f7506])).
fof(f7506,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X0))) | s(bool,'T_0') = s(bool,X1)) ) | $spl44),
  inference(duplicate_literal_removal,[],[f7505])).
fof(f7505,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X0))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X0)))) ) | $spl44),
  inference(condensation,[],[f7465])).
fof(f7465,plain,(
  ( ! [X41,X38,X42,X40] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X42))) | s(bool,X40) = s(bool,X41) | s(bool,'T_0') = s(bool,X38) | s(bool,'T_0') = s(bool,X41)) ) | $spl44),
  inference(backward_demodulation,[],[f7382,f6515])).
fof(f7554,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl12 | $spl44)),
  inference(subsumption_resolution,[],[f7514,f83])).
fof(f7514,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) | ($spl12 | $spl44)),
  inference(condensation,[],[f7513])).
fof(f7513,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X0) | ~p(s(bool,X1)) | ~p(s(bool,X0))) ) | ($spl12 | $spl44)),
  inference(condensation,[],[f7457])).
fof(f7457,plain,(
  ( ! [X57,X54,X55] : (s(bool,'F_0') = s(bool,X57) | ~p(s(bool,X55)) | ~p(s(bool,X57)) | s(bool,'F_0') = s(bool,X54)) ) | ($spl12 | $spl44)),
  inference(backward_demodulation,[],[f7382,f5294])).
fof(f7552,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl12 | $spl44)),
  inference(subsumption_resolution,[],[f7521,f73])).
fof(f7521,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,'F_0'))) ) | ($spl12 | $spl44)),
  inference(condensation,[],[f7448])).
fof(f7448,plain,(
  ( ! [X79,X77] : (p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,X79) | s(bool,X77) = s(bool,X79)) ) | ($spl12 | $spl44)),
  inference(backward_demodulation,[],[f7382,f7389])).
fof(f7549,plain,(
  ( ! [X2] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2)))))) ) | ($spl12 | $spl44)),
  inference(subsumption_resolution,[],[f7443,f73])).
fof(f7443,plain,(
  ( ! [X2] : (p(s(bool,'F_0')) | s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2)))))) ) | ($spl12 | $spl44)),
  inference(backward_demodulation,[],[f7382,f6838])).
fof(f7548,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1)))))) ) | $spl44),
  inference(condensation,[],[f7547])).
fof(f7547,plain,(
  ( ! [X142,X140,X141] : (s(bool,X141) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X140))))) | s(bool,X141) = s(bool,X142)) ) | $spl44),
  inference(global_subsumption,[],[f70,f71,f76,f80,f79,f102,f107,f104,f108,f110,f131,f162,f207,f325,f326,f376,f558,f856,f575,f598,f611,f891,f651,f917,f923,f713,f715,f769,f792,f780,f996,f958,f1072,f1074,f1093,f1168,f1222,f1224,f1226,f1227,f1663,f1664,f1297,f1739,f1783,f2996,f2989,f2966,f3031,f2025,f2948,f2928,f2919,f2916,f2900,f2895,f3048,f2889,f2886,f3059,f2122,f2881,f2871,f3075,f2866,f2843,f2842,f2837,f2836,f3107,f2283,f2285,f3133,f2333,f3137,f3143,f3163,f3172,f3174,f3179,f3180,f2700,f3200,f3202,f103,f118,f187,f219,f264,f417,f456,f493,f523,f542,f576,f612,f921,f934,f937,f957,f1050,f1094,f1799,f1801,f1951,f1952,f1993,f1994,f2024,f2061,f2091,f2121,f2125,f2155,f2156,f2219,f2282,f2366,f2429,f2492,f2495,f2498,f2499,f2500,f2502,f2503,f2588,f2829,f2831,f2835,f3076,f3219,f1665,f938,f1795,f1740,f1223,f714,f263,f249,f111,f455,f454,f73,f83,f84,f453,f2494,f3148,f3201,f690,f522,f3170,f3141,f2480,f2417,f2354,f2270,f2207,f1157,f681,f642,f469,f468,f258,f77,f74,f2585,f2489,f2426,f2363,f2279,f2216,f1294,f1165,f766,f687,f648,f525,f524,f2023,f1944,f915,f450,f416,f4143,f4172,f4173,f4245,f4321,f4324,f4379,f4382,f4966,f5038,f5039,f4995,f4996,f5044,f5047,f5090,f5091,f5146,f5147,f6425,f6426,f6507,f6508,f6509,f6512,f6513,f6514,f6515,f6516,f6517,f6518,f6519,f6520,f6521,f6522,f6523,f6529,f6543,f6544,f6545,f6546,f6547,f6549,f6554,f6555,f6558,f6563,f6564,f6565,f6567,f6568,f6581,f6582,f6583,f6586,f6589,f6848,f6850,f7090,f6890,f6893,f6944,f6945,f6946,f7097,f7382,f7530,f7528,f7535,f7527,f7539,f7428,f7543,f7434,f7544,f7546])).
fof(f7546,plain,(
  ( ! [X142,X140,X141] : (s(bool,X141) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X140))))) | s(bool,X141) = s(bool,X142) | s(bool,'T_0') = s(bool,X142)) ) | $spl44),
  inference(subsumption_resolution,[],[f7440,f73])).
fof(f7440,plain,(
  ( ! [X142,X140,X141] : (p(s(bool,'F_0')) | s(bool,X141) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X140))))) | s(bool,X141) = s(bool,X142) | s(bool,'T_0') = s(bool,X142)) ) | $spl44),
  inference(backward_demodulation,[],[f7382,f3141])).
fof(f7434,plain,(
  ( ! [X101,X102,X100] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X102))) | s('type.realax.real','const.real.abs_1'(s('type.realax.real',X100))) = s('type.realax.real','const.bool.COND_3'(s(bool,X101),s('type.realax.real',X100),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X100))))) | s(bool,'T_0') = s(bool,X101)) ) | $spl44),
  inference(backward_demodulation,[],[f7382,f681])).
fof(f7543,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))) ) | $spl44),
  inference(subsumption_resolution,[],[f7430,f73])).
fof(f7430,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))) ) | $spl44),
  inference(backward_demodulation,[],[f7382,f258])).
fof(f7428,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0))) = s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real',X0),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))) ) | $spl44),
  inference(backward_demodulation,[],[f7382,f74])).
fof(f7539,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,g(s(X1,sK1(X2,X3,bool,X1)))) | s(fun(X1,bool),X2) = s(fun(X1,bool),X3)) ) | $spl44),
  inference(condensation,[],[f7538])).
fof(f7538,plain,(
  ( ! [X185,X182,X186,X184,X183] : (s(bool,X185) = s(bool,g(s(X182,sK1(X183,X184,bool,X182)))) | s(fun(X182,bool),X183) = s(fun(X182,bool),X184) | s(bool,X185) = s(bool,X186)) ) | $spl44),
  inference(global_subsumption,[],[f70,f71,f76,f80,f79,f102,f107,f104,f108,f110,f131,f162,f207,f325,f326,f376,f558,f856,f575,f598,f611,f891,f651,f917,f923,f713,f715,f769,f792,f780,f996,f958,f1072,f1074,f1093,f1168,f1222,f1224,f1226,f1227,f1663,f1664,f1297,f1739,f1783,f2996,f2989,f2966,f3031,f2025,f2948,f2928,f2919,f2916,f2900,f2895,f3048,f2889,f2886,f3059,f2122,f2881,f2871,f3075,f2866,f2843,f2842,f2837,f2836,f3107,f2283,f2285,f3133,f2333,f3137,f3143,f3163,f3172,f3174,f3179,f3180,f2700,f3200,f3202,f103,f118,f187,f219,f264,f417,f456,f493,f523,f542,f576,f612,f921,f934,f937,f957,f1050,f1094,f1799,f1801,f1951,f1952,f1993,f1994,f2024,f2061,f2091,f2121,f2125,f2155,f2156,f2219,f2282,f2366,f2429,f2492,f2495,f2498,f2499,f2500,f2502,f2503,f2588,f2829,f2831,f2835,f3076,f3219,f1665,f938,f1795,f1740,f1223,f714,f263,f249,f111,f455,f454,f73,f83,f84,f453,f2494,f3148,f3201,f690,f522,f3170,f3141,f2480,f2417,f2354,f2270,f2207,f1157,f681,f642,f469,f468,f258,f77,f74,f2585,f2489,f2426,f2363,f2279,f2216,f1294,f1165,f766,f687,f648,f525,f524,f2023,f1944,f915,f450,f416,f4143,f4172,f4173,f4245,f4321,f4324,f4379,f4382,f4966,f5038,f5039,f4995,f4996,f5044,f5047,f5090,f5091,f5146,f5147,f6425,f6426,f6507,f6508,f6509,f6512,f6513,f6514,f6515,f6516,f6517,f6518,f6519,f6520,f6521,f6522,f6523,f6529,f6543,f6544,f6545,f6546,f6547,f6549,f6554,f6555,f6558,f6563,f6564,f6565,f6567,f6568,f6581,f6582,f6583,f6586,f6589,f6848,f6850,f7090,f6890,f6893,f6944,f6945,f6946,f7097,f7382,f7530,f7528,f7535,f7527])).
fof(f7527,plain,(
  ( ! [X185,X182,X186,X184,X183] : (s(fun(X182,bool),X183) = s(fun(X182,bool),X184) | s(bool,X185) = s(bool,g(s(X182,sK1(X183,X184,bool,X182)))) | s(bool,X185) = s(bool,X186) | s(bool,'T_0') = s(bool,X186)) ) | $spl44),
  inference(trivial_inequality_removal,[],[f7421])).
fof(f7421,plain,(
  ( ! [X185,X182,X186,X184,X183] : (s(bool,'F_0') != s(bool,'F_0') | s(fun(X182,bool),X183) = s(fun(X182,bool),X184) | s(bool,X185) = s(bool,g(s(X182,sK1(X183,X184,bool,X182)))) | s(bool,X185) = s(bool,X186) | s(bool,'T_0') = s(bool,X186)) ) | $spl44),
  inference(backward_demodulation,[],[f7382,f2363])).
fof(f7528,plain,(
  ( ! [X6,X4,X5] : (s(fun(X4,bool),X5) = s(fun(X4,bool),X6) | s(bool,'T_0') = s(bool,g(s(X4,sK1(X5,X6,bool,X4))))) ) | $spl44),
  inference(trivial_inequality_removal,[],[f7413])).
fof(f7413,plain,(
  ( ! [X6,X4,X5] : (s(bool,'F_0') != s(bool,'F_0') | s(fun(X4,bool),X5) = s(fun(X4,bool),X6) | s(bool,'T_0') = s(bool,g(s(X4,sK1(X5,X6,bool,X4))))) ) | $spl44),
  inference(backward_demodulation,[],[f7382,f525])).
fof(f7566,plain,(
  $false | ($spl12 | $spl44)),
  inference(global_subsumption,[],[f70,f71,f76,f80,f79,f102,f107,f104,f108,f110,f131,f162,f207,f325,f326,f376,f558,f856,f575,f598,f611,f891,f651,f917,f923,f713,f715,f769,f792,f780,f996,f958,f1072,f1074,f1093,f1168,f1222,f1224,f1226,f1227,f1663,f1664,f1297,f1739,f1783,f2996,f2989,f2966,f3031,f2025,f2948,f2928,f2919,f2916,f2900,f2895,f3048,f2889,f2886,f3059,f2122,f2881,f2871,f3075,f2866,f2843,f2842,f2837,f2836,f3107,f2283,f2285,f3133,f2333,f3137,f3143,f3163,f3172,f3174,f3179,f3180,f2700,f3200,f3202,f103,f118,f187,f219,f264,f417,f456,f493,f523,f542,f576,f612,f921,f934,f937,f957,f1050,f1094,f1799,f1801,f1951,f1952,f1993,f1994,f2024,f2061,f2091,f2121,f2125,f2155,f2156,f2219,f2282,f2366,f2429,f2492,f2495,f2498,f2499,f2500,f2502,f2503,f2588,f2829,f2831,f2835,f3076,f3219,f1665,f938,f1795,f1740,f1223,f714,f263,f249,f111,f455,f454,f73,f83,f84,f453,f2494,f3148,f3201,f3095,f690,f522,f3170,f3141,f2480,f2417,f2354,f2270,f2207,f1157,f681,f642,f469,f468,f258,f77,f74,f2585,f2489,f2426,f2363,f2279,f2216,f1294,f1165,f766,f687,f648,f525,f524,f2023,f1944,f915,f450,f416,f4143,f4172,f4173,f4245,f4321,f4324,f4379,f4382,f4966,f5038,f5039,f4995,f4996,f5044,f5047,f5090,f5091,f5146,f5147,f5192,f5193,f5194,f5260,f5261,f5262,f5265,f5266,f5267,f5268,f5269,f5270,f5271,f5272,f5273,f5274,f5275,f5279,f5288,f5289,f5290,f5291,f5292,f5294,f5299,f5300,f5305,f5306,f5307,f5309,f5310,f5322,f5323,f5324,f5327,f5816,f6058,f6124,f6425,f6426,f6427,f6507,f6508,f6509,f6512,f6513,f6514,f6515,f6516,f6517,f6518,f6519,f6520,f6521,f6522,f6523,f6529,f6543,f6544,f6545,f6546,f6547,f6549,f6554,f6555,f6558,f6563,f6564,f6565,f6567,f6568,f6580,f6581,f6582,f6583,f6586,f6589,f6747,f6748,f6838,f6770,f6839,f6841,f6848,f6850,f7090,f6890,f6893,f6944,f6945,f6946,f6947,f7097,f7107,f7108,f7181,f7180,f7195,f7177,f7199,f7176,f7162,f7163,f7372,f7387,f7389,f7246,f7247,f7406,f7407,f7411,f7382,f7530,f7528,f7535,f7527,f7539,f7428,f7543,f7434,f7544,f7546,f7548,f7549,f7552,f7554,f7507,f7557,f7503,f7560,f7499,f7564,f7486])).
fof(f7565,plain,(
  $false | ($spl12 | $spl44)),
  inference(global_subsumption,[],[f70,f71,f76,f80,f79,f102,f107,f104,f108,f110,f131,f162,f207,f325,f326,f376,f558,f856,f575,f598,f611,f891,f651,f917,f923,f713,f715,f769,f792,f780,f996,f958,f1072,f1074,f1093,f1168,f1222,f1224,f1226,f1227,f1663,f1664,f1297,f1739,f1783,f2996,f2989,f2966,f3031,f2025,f2948,f2928,f2919,f2916,f2900,f2895,f3048,f2889,f2886,f3059,f2122,f2881,f2871,f3075,f2866,f2843,f2842,f2837,f2836,f3107,f2283,f2285,f3133,f2333,f3137,f3143,f3163,f3172,f3174,f3179,f3180,f2700,f3200,f3202,f103,f118,f187,f219,f264,f417,f456,f493,f523,f542,f576,f612,f921,f934,f937,f957,f1050,f1094,f1799,f1801,f1951,f1952,f1993,f1994,f2024,f2061,f2091,f2121,f2125,f2155,f2156,f2219,f2282,f2366,f2429,f2492,f2495,f2498,f2499,f2500,f2502,f2503,f2588,f2829,f2831,f2835,f3076,f3219,f1665,f938,f1795,f1740,f1223,f714,f263,f249,f111,f455,f454,f73,f83,f84,f453,f2494,f3148,f3201,f3095,f690,f522,f3170,f3141,f2480,f2417,f2354,f2270,f2207,f1157,f681,f642,f469,f468,f258,f77,f74,f2585,f2489,f2426,f2363,f2279,f2216,f1294,f1165,f766,f687,f648,f525,f524,f2023,f1944,f915,f450,f416,f4143,f4172,f4173,f4245,f4321,f4324,f4379,f4382,f4966,f5038,f5039,f4995,f4996,f5044,f5047,f5090,f5091,f5146,f5147,f5192,f5193,f5194,f5260,f5261,f5262,f5265,f5266,f5267,f5268,f5269,f5270,f5271,f5272,f5273,f5274,f5275,f5279,f5288,f5289,f5290,f5291,f5292,f5294,f5299,f5300,f5305,f5306,f5307,f5309,f5310,f5322,f5323,f5324,f5327,f5816,f6058,f6124,f6425,f6426,f6427,f6507,f6508,f6509,f6512,f6513,f6514,f6515,f6516,f6517,f6518,f6519,f6520,f6521,f6522,f6523,f6529,f6543,f6544,f6545,f6546,f6547,f6549,f6554,f6555,f6558,f6563,f6564,f6565,f6567,f6568,f6580,f6581,f6582,f6583,f6586,f6589,f6747,f6748,f6838,f6770,f6839,f6841,f6848,f6850,f7090,f6890,f6893,f6944,f6945,f6946,f6947,f7097,f7107,f7108,f7181,f7180,f7195,f7177,f7199,f7176,f7162,f7163,f7372,f7387,f7389,f7246,f7247,f7406,f7407,f7411,f7382,f7530,f7528,f7535,f7527,f7539,f7428,f7543,f7434,f7544,f7546,f7548,f7549,f7552,f7554,f7507,f7557,f7503,f7560,f7499,f7564])).
fof(f7562,plain,(
  $false | ($spl12 | $spl44)),
  inference(global_subsumption,[],[f70,f71,f76,f80,f79,f102,f107,f104,f108,f110,f131,f162,f207,f325,f326,f376,f558,f856,f575,f598,f611,f891,f651,f917,f923,f713,f715,f769,f792,f780,f996,f958,f1072,f1074,f1093,f1168,f1222,f1224,f1226,f1227,f1663,f1664,f1297,f1739,f1783,f2996,f2989,f2966,f3031,f2025,f2948,f2928,f2919,f2916,f2900,f2895,f3048,f2889,f2886,f3059,f2122,f2881,f2871,f3075,f2866,f2843,f2842,f2837,f2836,f3107,f2283,f2285,f3133,f2333,f3137,f3143,f3163,f3172,f3174,f3179,f3180,f2700,f3200,f3202,f103,f118,f187,f219,f264,f417,f456,f493,f523,f542,f576,f612,f921,f934,f937,f957,f1050,f1094,f1799,f1801,f1951,f1952,f1993,f1994,f2024,f2061,f2091,f2121,f2125,f2155,f2156,f2219,f2282,f2366,f2429,f2492,f2495,f2498,f2499,f2500,f2502,f2503,f2588,f2829,f2831,f2835,f3076,f3219,f1665,f938,f1795,f1740,f1223,f714,f263,f249,f111,f455,f454,f73,f83,f84,f453,f2494,f3148,f3201,f3095,f690,f522,f3170,f3141,f2480,f2417,f2354,f2270,f2207,f1157,f681,f642,f469,f468,f258,f77,f74,f2585,f2489,f2426,f2363,f2279,f2216,f1294,f1165,f766,f687,f648,f525,f524,f2023,f1944,f915,f450,f416,f4143,f4172,f4173,f4245,f4321,f4324,f4379,f4382,f4966,f5038,f5039,f4995,f4996,f5044,f5047,f5090,f5091,f5146,f5147,f5192,f5193,f5194,f5260,f5261,f5262,f5265,f5266,f5267,f5268,f5269,f5270,f5271,f5272,f5273,f5274,f5275,f5279,f5288,f5289,f5290,f5291,f5292,f5294,f5299,f5300,f5305,f5306,f5307,f5309,f5310,f5322,f5323,f5324,f5327,f5816,f6058,f6124,f6425,f6426,f6427,f6507,f6508,f6509,f6512,f6513,f6514,f6515,f6516,f6517,f6518,f6519,f6520,f6521,f6522,f6523,f6529,f6543,f6544,f6545,f6546,f6547,f6549,f6554,f6555,f6558,f6563,f6564,f6565,f6567,f6568,f6580,f6581,f6582,f6583,f6586,f6589,f6747,f6748,f6838,f6770,f6839,f6841,f6848,f6850,f7090,f6890,f6893,f6944,f6945,f6946,f6947,f7097,f7107,f7108,f7181,f7180,f7195,f7177,f7199,f7176,f7162,f7163,f7372,f7387,f7389,f7246,f7247,f7406,f7407,f7411,f7382,f7530,f7528,f7535,f7527,f7539,f7428,f7543,f7434,f7544,f7546,f7548,f7549,f7552,f7554,f7507,f7557,f7503,f7560,f7499])).
fof(f7561,plain,(
  $false | ($spl12 | $spl44)),
  inference(global_subsumption,[],[f70,f71,f76,f80,f79,f102,f107,f104,f108,f110,f131,f162,f207,f325,f326,f376,f558,f856,f575,f598,f611,f891,f651,f917,f923,f713,f715,f769,f792,f780,f996,f958,f1072,f1074,f1093,f1168,f1222,f1224,f1226,f1227,f1663,f1664,f1297,f1739,f1783,f2996,f2989,f2966,f3031,f2025,f2948,f2928,f2919,f2916,f2900,f2895,f3048,f2889,f2886,f3059,f2122,f2881,f2871,f3075,f2866,f2843,f2842,f2837,f2836,f3107,f2283,f2285,f3133,f2333,f3137,f3143,f3163,f3172,f3174,f3179,f3180,f2700,f3200,f3202,f103,f118,f187,f219,f264,f417,f456,f493,f523,f542,f576,f612,f921,f934,f937,f957,f1050,f1094,f1799,f1801,f1951,f1952,f1993,f1994,f2024,f2061,f2091,f2121,f2125,f2155,f2156,f2219,f2282,f2366,f2429,f2492,f2495,f2498,f2499,f2500,f2502,f2503,f2588,f2829,f2831,f2835,f3076,f3219,f1665,f938,f1795,f1740,f1223,f714,f263,f249,f111,f455,f454,f73,f83,f84,f453,f2494,f3148,f3201,f3095,f690,f522,f3170,f3141,f2480,f2417,f2354,f2270,f2207,f1157,f681,f642,f469,f468,f258,f77,f74,f2585,f2489,f2426,f2363,f2279,f2216,f1294,f1165,f766,f687,f648,f525,f524,f2023,f1944,f915,f450,f416,f4143,f4172,f4173,f4245,f4321,f4324,f4379,f4382,f4966,f5038,f5039,f4995,f4996,f5044,f5047,f5090,f5091,f5146,f5147,f5192,f5193,f5194,f5260,f5261,f5262,f5265,f5266,f5267,f5268,f5269,f5270,f5271,f5272,f5273,f5274,f5275,f5279,f5288,f5289,f5290,f5291,f5292,f5294,f5299,f5300,f5305,f5306,f5307,f5309,f5310,f5322,f5323,f5324,f5327,f5816,f6058,f6124,f6425,f6426,f6427,f6507,f6508,f6509,f6512,f6513,f6514,f6515,f6516,f6517,f6518,f6519,f6520,f6521,f6522,f6523,f6529,f6543,f6544,f6545,f6546,f6547,f6549,f6554,f6555,f6558,f6563,f6564,f6565,f6567,f6568,f6580,f6581,f6582,f6583,f6586,f6589,f6747,f6748,f6838,f6770,f6839,f6841,f6848,f6850,f7090,f6890,f6893,f6944,f6945,f6946,f6947,f7097,f7107,f7108,f7181,f7180,f7195,f7177,f7199,f7176,f7162,f7163,f7372,f7387,f7389,f7246,f7247,f7406,f7407,f7411,f7382,f7530,f7528,f7535,f7527,f7539,f7428,f7543,f7434,f7544,f7546,f7548,f7549,f7552,f7554,f7507,f7557,f7503,f7560])).
fof(f7559,plain,(
  $false | ($spl12 | $spl44)),
  inference(global_subsumption,[],[f70,f71,f76,f80,f79,f102,f107,f104,f108,f110,f131,f162,f207,f325,f326,f376,f558,f856,f575,f598,f611,f891,f651,f917,f923,f713,f715,f769,f792,f780,f996,f958,f1072,f1074,f1093,f1168,f1222,f1224,f1226,f1227,f1663,f1664,f1297,f1739,f1783,f2996,f2989,f2966,f3031,f2025,f2948,f2928,f2919,f2916,f2900,f2895,f3048,f2889,f2886,f3059,f2122,f2881,f2871,f3075,f2866,f2843,f2842,f2837,f2836,f3107,f2283,f2285,f3133,f2333,f3137,f3143,f3163,f3172,f3174,f3179,f3180,f2700,f3200,f3202,f103,f118,f187,f219,f264,f417,f456,f493,f523,f542,f576,f612,f921,f934,f937,f957,f1050,f1094,f1799,f1801,f1951,f1952,f1993,f1994,f2024,f2061,f2091,f2121,f2125,f2155,f2156,f2219,f2282,f2366,f2429,f2492,f2495,f2498,f2499,f2500,f2502,f2503,f2588,f2829,f2831,f2835,f3076,f3219,f1665,f938,f1795,f1740,f1223,f714,f263,f249,f111,f455,f454,f73,f83,f84,f453,f2494,f3148,f3201,f3095,f690,f522,f3170,f3141,f2480,f2417,f2354,f2270,f2207,f1157,f681,f642,f469,f468,f258,f77,f74,f2585,f2489,f2426,f2363,f2279,f2216,f1294,f1165,f766,f687,f648,f525,f524,f2023,f1944,f915,f450,f416,f4143,f4172,f4173,f4245,f4321,f4324,f4379,f4382,f4966,f5038,f5039,f4995,f4996,f5044,f5047,f5090,f5091,f5146,f5147,f5192,f5193,f5194,f5260,f5261,f5262,f5265,f5266,f5267,f5268,f5269,f5270,f5271,f5272,f5273,f5274,f5275,f5279,f5288,f5289,f5290,f5291,f5292,f5294,f5299,f5300,f5305,f5306,f5307,f5309,f5310,f5322,f5323,f5324,f5327,f5816,f6058,f6124,f6425,f6426,f6427,f6507,f6508,f6509,f6512,f6513,f6514,f6515,f6516,f6517,f6518,f6519,f6520,f6521,f6522,f6523,f6529,f6543,f6544,f6545,f6546,f6547,f6549,f6554,f6555,f6558,f6563,f6564,f6565,f6567,f6568,f6580,f6581,f6582,f6583,f6586,f6589,f6747,f6748,f6838,f6770,f6839,f6841,f6848,f6850,f7090,f6890,f6893,f6944,f6945,f6946,f6947,f7097,f7107,f7108,f7181,f7180,f7195,f7177,f7199,f7176,f7162,f7163,f7372,f7387,f7389,f7246,f7247,f7406,f7407,f7411,f7382,f7530,f7528,f7535,f7527,f7539,f7428,f7543,f7434,f7544,f7546,f7548,f7549,f7552,f7554,f7507,f7557,f7503])).
fof(f7558,plain,(
  $false | ($spl12 | $spl44)),
  inference(global_subsumption,[],[f70,f71,f76,f80,f79,f102,f107,f104,f108,f110,f131,f162,f207,f325,f326,f376,f558,f856,f575,f598,f611,f891,f651,f917,f923,f713,f715,f769,f792,f780,f996,f958,f1072,f1074,f1093,f1168,f1222,f1224,f1226,f1227,f1663,f1664,f1297,f1739,f1783,f2996,f2989,f2966,f3031,f2025,f2948,f2928,f2919,f2916,f2900,f2895,f3048,f2889,f2886,f3059,f2122,f2881,f2871,f3075,f2866,f2843,f2842,f2837,f2836,f3107,f2283,f2285,f3133,f2333,f3137,f3143,f3163,f3172,f3174,f3179,f3180,f2700,f3200,f3202,f103,f118,f187,f219,f264,f417,f456,f493,f523,f542,f576,f612,f921,f934,f937,f957,f1050,f1094,f1799,f1801,f1951,f1952,f1993,f1994,f2024,f2061,f2091,f2121,f2125,f2155,f2156,f2219,f2282,f2366,f2429,f2492,f2495,f2498,f2499,f2500,f2502,f2503,f2588,f2829,f2831,f2835,f3076,f3219,f1665,f938,f1795,f1740,f1223,f714,f263,f249,f111,f455,f454,f73,f83,f84,f453,f2494,f3148,f3201,f3095,f690,f522,f3170,f3141,f2480,f2417,f2354,f2270,f2207,f1157,f681,f642,f469,f468,f258,f77,f74,f2585,f2489,f2426,f2363,f2279,f2216,f1294,f1165,f766,f687,f648,f525,f524,f2023,f1944,f915,f450,f416,f4143,f4172,f4173,f4245,f4321,f4324,f4379,f4382,f4966,f5038,f5039,f4995,f4996,f5044,f5047,f5090,f5091,f5146,f5147,f5192,f5193,f5194,f5260,f5261,f5262,f5265,f5266,f5267,f5268,f5269,f5270,f5271,f5272,f5273,f5274,f5275,f5279,f5288,f5289,f5290,f5291,f5292,f5294,f5299,f5300,f5305,f5306,f5307,f5309,f5310,f5322,f5323,f5324,f5327,f5816,f6058,f6124,f6425,f6426,f6427,f6507,f6508,f6509,f6512,f6513,f6514,f6515,f6516,f6517,f6518,f6519,f6520,f6521,f6522,f6523,f6529,f6543,f6544,f6545,f6546,f6547,f6549,f6554,f6555,f6558,f6563,f6564,f6565,f6567,f6568,f6580,f6581,f6582,f6583,f6586,f6589,f6747,f6748,f6838,f6770,f6839,f6841,f6848,f6850,f7090,f6890,f6893,f6944,f6945,f6946,f6947,f7097,f7107,f7108,f7181,f7180,f7195,f7177,f7199,f7176,f7162,f7163,f7372,f7387,f7389,f7246,f7247,f7406,f7407,f7411,f7382,f7530,f7528,f7535,f7527,f7539,f7428,f7543,f7434,f7544,f7546,f7548,f7549,f7552,f7554,f7507,f7557])).
fof(f7556,plain,(
  $false | ($spl12 | $spl44)),
  inference(global_subsumption,[],[f70,f71,f76,f80,f79,f102,f107,f104,f108,f110,f131,f162,f207,f325,f326,f376,f558,f856,f575,f598,f611,f891,f651,f917,f923,f713,f715,f769,f792,f780,f996,f958,f1072,f1074,f1093,f1168,f1222,f1224,f1226,f1227,f1663,f1664,f1297,f1739,f1783,f2996,f2989,f2966,f3031,f2025,f2948,f2928,f2919,f2916,f2900,f2895,f3048,f2889,f2886,f3059,f2122,f2881,f2871,f3075,f2866,f2843,f2842,f2837,f2836,f3107,f2283,f2285,f3133,f2333,f3137,f3143,f3163,f3172,f3174,f3179,f3180,f2700,f3200,f3202,f103,f118,f187,f219,f264,f417,f456,f493,f523,f542,f576,f612,f921,f934,f937,f957,f1050,f1094,f1799,f1801,f1951,f1952,f1993,f1994,f2024,f2061,f2091,f2121,f2125,f2155,f2156,f2219,f2282,f2366,f2429,f2492,f2495,f2498,f2499,f2500,f2502,f2503,f2588,f2829,f2831,f2835,f3076,f3219,f1665,f938,f1795,f1740,f1223,f714,f263,f249,f111,f455,f454,f73,f83,f84,f453,f2494,f3148,f3201,f3095,f690,f522,f3170,f3141,f2480,f2417,f2354,f2270,f2207,f1157,f681,f642,f469,f468,f258,f77,f74,f2585,f2489,f2426,f2363,f2279,f2216,f1294,f1165,f766,f687,f648,f525,f524,f2023,f1944,f915,f450,f416,f4143,f4172,f4173,f4245,f4321,f4324,f4379,f4382,f4966,f5038,f5039,f4995,f4996,f5044,f5047,f5090,f5091,f5146,f5147,f5192,f5193,f5194,f5260,f5261,f5262,f5265,f5266,f5267,f5268,f5269,f5270,f5271,f5272,f5273,f5274,f5275,f5279,f5288,f5289,f5290,f5291,f5292,f5294,f5299,f5300,f5305,f5306,f5307,f5309,f5310,f5322,f5323,f5324,f5327,f5816,f6058,f6124,f6425,f6426,f6427,f6507,f6508,f6509,f6512,f6513,f6514,f6515,f6516,f6517,f6518,f6519,f6520,f6521,f6522,f6523,f6529,f6543,f6544,f6545,f6546,f6547,f6549,f6554,f6555,f6558,f6563,f6564,f6565,f6567,f6568,f6580,f6581,f6582,f6583,f6586,f6589,f6747,f6748,f6838,f6770,f6839,f6841,f6848,f6850,f7090,f6890,f6893,f6944,f6945,f6946,f6947,f7097,f7107,f7108,f7181,f7180,f7195,f7177,f7199,f7176,f7162,f7163,f7372,f7387,f7389,f7246,f7247,f7406,f7407,f7411,f7382,f7530,f7528,f7535,f7527,f7539,f7428,f7543,f7434,f7544,f7546,f7548,f7549,f7552,f7554,f7507])).
fof(f7555,plain,(
  $false | ($spl12 | $spl44)),
  inference(global_subsumption,[],[f70,f71,f76,f80,f79,f102,f107,f104,f108,f110,f131,f162,f207,f325,f326,f376,f558,f856,f575,f598,f611,f891,f651,f917,f923,f713,f715,f769,f792,f780,f996,f958,f1072,f1074,f1093,f1168,f1222,f1224,f1226,f1227,f1663,f1664,f1297,f1739,f1783,f2996,f2989,f2966,f3031,f2025,f2948,f2928,f2919,f2916,f2900,f2895,f3048,f2889,f2886,f3059,f2122,f2881,f2871,f3075,f2866,f2843,f2842,f2837,f2836,f3107,f2283,f2285,f3133,f2333,f3137,f3143,f3163,f3172,f3174,f3179,f3180,f2700,f3200,f3202,f103,f118,f187,f219,f264,f417,f456,f493,f523,f542,f576,f612,f921,f934,f937,f957,f1050,f1094,f1799,f1801,f1951,f1952,f1993,f1994,f2024,f2061,f2091,f2121,f2125,f2155,f2156,f2219,f2282,f2366,f2429,f2492,f2495,f2498,f2499,f2500,f2502,f2503,f2588,f2829,f2831,f2835,f3076,f3219,f1665,f938,f1795,f1740,f1223,f714,f263,f249,f111,f455,f454,f73,f83,f84,f453,f2494,f3148,f3201,f3095,f690,f522,f3170,f3141,f2480,f2417,f2354,f2270,f2207,f1157,f681,f642,f469,f468,f258,f77,f74,f2585,f2489,f2426,f2363,f2279,f2216,f1294,f1165,f766,f687,f648,f525,f524,f2023,f1944,f915,f450,f416,f4143,f4172,f4173,f4245,f4321,f4324,f4379,f4382,f4966,f5038,f5039,f4995,f4996,f5044,f5047,f5090,f5091,f5146,f5147,f5192,f5193,f5194,f5260,f5261,f5262,f5265,f5266,f5267,f5268,f5269,f5270,f5271,f5272,f5273,f5274,f5275,f5279,f5288,f5289,f5290,f5291,f5292,f5294,f5299,f5300,f5305,f5306,f5307,f5309,f5310,f5322,f5323,f5324,f5327,f5816,f6058,f6124,f6425,f6426,f6427,f6507,f6508,f6509,f6512,f6513,f6514,f6515,f6516,f6517,f6518,f6519,f6520,f6521,f6522,f6523,f6529,f6543,f6544,f6545,f6546,f6547,f6549,f6554,f6555,f6558,f6563,f6564,f6565,f6567,f6568,f6580,f6581,f6582,f6583,f6586,f6589,f6747,f6748,f6838,f6770,f6839,f6841,f6848,f6850,f7090,f6890,f6893,f6944,f6945,f6946,f6947,f7097,f7107,f7108,f7181,f7180,f7195,f7177,f7199,f7176,f7162,f7163,f7372,f7387,f7389,f7246,f7247,f7406,f7407,f7411,f7382,f7530,f7528,f7535,f7527,f7539,f7428,f7543,f7434,f7544,f7546,f7548,f7549,f7552,f7554])).
fof(f7553,plain,(
  $false | ($spl12 | $spl44)),
  inference(global_subsumption,[],[f70,f71,f76,f80,f79,f102,f107,f104,f108,f110,f131,f162,f207,f325,f326,f376,f558,f856,f575,f598,f611,f891,f651,f917,f923,f713,f715,f769,f792,f780,f996,f958,f1072,f1074,f1093,f1168,f1222,f1224,f1226,f1227,f1663,f1664,f1297,f1739,f1783,f2996,f2989,f2966,f3031,f2025,f2948,f2928,f2919,f2916,f2900,f2895,f3048,f2889,f2886,f3059,f2122,f2881,f2871,f3075,f2866,f2843,f2842,f2837,f2836,f3107,f2283,f2285,f3133,f2333,f3137,f3143,f3163,f3172,f3174,f3179,f3180,f2700,f3200,f3202,f103,f118,f187,f219,f264,f417,f456,f493,f523,f542,f576,f612,f921,f934,f937,f957,f1050,f1094,f1799,f1801,f1951,f1952,f1993,f1994,f2024,f2061,f2091,f2121,f2125,f2155,f2156,f2219,f2282,f2366,f2429,f2492,f2495,f2498,f2499,f2500,f2502,f2503,f2588,f2829,f2831,f2835,f3076,f3219,f1665,f938,f1795,f1740,f1223,f714,f263,f249,f111,f455,f454,f73,f83,f84,f453,f2494,f3148,f3201,f3095,f690,f522,f3170,f3141,f2480,f2417,f2354,f2270,f2207,f1157,f681,f642,f469,f468,f258,f77,f74,f2585,f2489,f2426,f2363,f2279,f2216,f1294,f1165,f766,f687,f648,f525,f524,f2023,f1944,f915,f450,f416,f4143,f4172,f4173,f4245,f4321,f4324,f4379,f4382,f4966,f5038,f5039,f4995,f4996,f5044,f5047,f5090,f5091,f5146,f5147,f5192,f5193,f5194,f5260,f5261,f5262,f5265,f5266,f5267,f5268,f5269,f5270,f5271,f5272,f5273,f5274,f5275,f5279,f5288,f5289,f5290,f5291,f5292,f5294,f5299,f5300,f5305,f5306,f5307,f5309,f5310,f5322,f5323,f5324,f5327,f5816,f6058,f6124,f6425,f6426,f6427,f6507,f6508,f6509,f6512,f6513,f6514,f6515,f6516,f6517,f6518,f6519,f6520,f6521,f6522,f6523,f6529,f6543,f6544,f6545,f6546,f6547,f6549,f6554,f6555,f6558,f6563,f6564,f6565,f6567,f6568,f6580,f6581,f6582,f6583,f6586,f6589,f6747,f6748,f6838,f6770,f6839,f6841,f6848,f6850,f7090,f6890,f6893,f6944,f6945,f6946,f6947,f7097,f7107,f7108,f7181,f7180,f7195,f7177,f7199,f7176,f7162,f7163,f7372,f7387,f7389,f7246,f7247,f7406,f7407,f7411,f7382,f7530,f7528,f7535,f7527,f7539,f7428,f7543,f7434,f7544,f7546,f7548,f7549,f7552])).
fof(f3135,plain,(
  ( ! [X61] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X61)))) ) | $spl20),
  inference(cnf_transformation,[],[f3135_D])).
fof(f3135_D,plain,(
  ( ! [X61] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X61)))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f7545,plain,(
  ( ! [X101,X100] : (s('type.realax.real','const.real.abs_1'(s('type.realax.real',X100))) = s('type.realax.real','const.bool.COND_3'(s(bool,X101),s('type.realax.real',X100),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X100)))))) ) | $spl52),
  inference(cnf_transformation,[],[f7545_D])).
fof(f7545_D,plain,(
  ( ! [X101,X100] : (s('type.realax.real','const.real.abs_1'(s('type.realax.real',X100))) = s('type.realax.real','const.bool.COND_3'(s(bool,X101),s('type.realax.real',X100),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X100)))))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f7544,plain,(
  ( ! [X101,X102,X100] : (s('type.realax.real','const.real.abs_1'(s('type.realax.real',X100))) = s('type.realax.real','const.bool.COND_3'(s(bool,X101),s('type.realax.real',X100),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X100))))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X102)))) ) | $spl44),
  inference(global_subsumption,[],[f70,f71,f76,f80,f79,f102,f107,f104,f108,f110,f131,f162,f207,f325,f326,f376,f558,f856,f575,f598,f611,f891,f651,f917,f923,f713,f715,f769,f792,f780,f996,f958,f1072,f1074,f1093,f1168,f1222,f1224,f1226,f1227,f1663,f1664,f1297,f1739,f1783,f2996,f2989,f2966,f3031,f2025,f2948,f2928,f2919,f2916,f2900,f2895,f3048,f2889,f2886,f3059,f2122,f2881,f2871,f3075,f2866,f2843,f2842,f2837,f2836,f3107,f2283,f2285,f3133,f2333,f3137,f3143,f3163,f3172,f3174,f3179,f3180,f2700,f3200,f3202,f103,f118,f187,f219,f264,f417,f456,f493,f523,f542,f576,f612,f921,f934,f937,f957,f1050,f1094,f1799,f1801,f1951,f1952,f1993,f1994,f2024,f2061,f2091,f2121,f2125,f2155,f2156,f2219,f2282,f2366,f2429,f2492,f2495,f2498,f2499,f2500,f2502,f2503,f2588,f2829,f2831,f2835,f3076,f3219,f1665,f938,f1795,f1740,f1223,f714,f263,f249,f111,f455,f454,f73,f83,f84,f453,f2494,f3148,f3201,f690,f522,f3170,f3141,f2480,f2417,f2354,f2270,f2207,f1157,f681,f642,f469,f468,f258,f77,f74,f2585,f2489,f2426,f2363,f2279,f2216,f1294,f1165,f766,f687,f648,f525,f524,f2023,f1944,f915,f450,f416,f4143,f4172,f4173,f4245,f4321,f4324,f4379,f4382,f4966,f5038,f5039,f4995,f4996,f5044,f5047,f5090,f5091,f5146,f5147,f6425,f6426,f6507,f6508,f6509,f6512,f6513,f6514,f6515,f6516,f6517,f6518,f6519,f6520,f6521,f6522,f6523,f6529,f6543,f6544,f6545,f6546,f6547,f6549,f6554,f6555,f6558,f6563,f6564,f6565,f6567,f6568,f6581,f6582,f6583,f6586,f6589,f6848,f6850,f7090,f6890,f6893,f6944,f6945,f6946,f7097,f7382,f7530,f7528,f7535,f7527,f7539,f7428,f7543,f7434])).
fof(f7542,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f7541,f73])).
fof(f7541,plain,(
  p(s(bool,'F_0')) | $spl44),
  inference(forward_demodulation,[],[f7540,f7382])).
fof(f7540,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))))) ) | $spl44),
  inference(subsumption_resolution,[],[f7429,f73])).
fof(f7429,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))))) ) | $spl44),
  inference(backward_demodulation,[],[f7382,f77])).
fof(f7531,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | $spl50),
  inference(cnf_transformation,[],[f7531_D])).
fof(f7531_D,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f7535,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X3))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | $spl44),
  inference(forward_demodulation,[],[f7534,f7382])).
fof(f7534,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,g(s(X0,sK1(X1,X2,bool,X0)))),s(bool,X3))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | $spl44),
  inference(condensation,[],[f7533])).
fof(f7533,plain,(
  ( ! [X127,X130,X128,X131,X129] : (s(fun(X127,bool),X128) = s(fun(X127,bool),X129) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,g(s(X127,sK1(X128,X129,bool,X127)))),s(bool,X131))) | s(bool,'T_0') = s(bool,X130)) ) | $spl44),
  inference(subsumption_resolution,[],[f7415,f76])).
fof(f7415,plain,(
  ( ! [X127,X130,X128,X131,X129] : (s(bool,'F_0') != s(bool,X130) | s(fun(X127,bool),X128) = s(fun(X127,bool),X129) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,g(s(X127,sK1(X128,X129,bool,X127)))),s(bool,X131))) | s(bool,'T_0') = s(bool,X130)) ) | $spl44),
  inference(backward_demodulation,[],[f7382,f687])).
fof(f7530,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X3))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | $spl44),
  inference(forward_demodulation,[],[f7529,f7382])).
fof(f7529,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,g(s(X0,sK1(X1,X2,bool,X0)))),s(bool,X3)))) ) | $spl44),
  inference(trivial_inequality_removal,[],[f7412])).
fof(f7412,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') != s(bool,'F_0') | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,g(s(X0,sK1(X1,X2,bool,X0)))),s(bool,X3)))) ) | $spl44),
  inference(backward_demodulation,[],[f7382,f524])).
fof(f7408,plain,(
  ( ! [X50,X53,X51] : (s(bool,'F_0') = s(bool,X50) | s(bool,'F_0') = s(bool,X53) | s(bool,X53) = s(bool,'const.bool.COND_3'(s(bool,X50),s(bool,'T_0'),s(bool,X51)))) ) | $spl48),
  inference(cnf_transformation,[],[f7408_D])).
fof(f7408_D,plain,(
  ( ! [X50,X53,X51] : (s(bool,'F_0') = s(bool,X50) | s(bool,'F_0') = s(bool,X53) | s(bool,X53) = s(bool,'const.bool.COND_3'(s(bool,X50),s(bool,'T_0'),s(bool,X51)))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f7407,plain,(
  ( ! [X52,X50,X53,X51] : (p(s(bool,X52)) | s(bool,'F_0') = s(bool,X50) | s(bool,X53) = s(bool,'const.bool.COND_3'(s(bool,X50),s(bool,'T_0'),s(bool,X51))) | s(bool,'F_0') = s(bool,X53)) ) | $spl12),
  inference(subsumption_resolution,[],[f7322,f79])).
fof(f7322,plain,(
  ( ! [X52,X50,X53,X51] : (p(s(bool,X52)) | s(bool,'F_0') = s(bool,X50) | s(bool,X53) = s(bool,'const.bool.COND_3'(s(bool,X50),s(bool,'T_0'),s(bool,X51))) | s(bool,'F_0') = s(bool,X53) | s(bool,'T_0') = s(bool,X52)) ) | $spl12),
  inference(superposition,[],[f7176,f1664])).
fof(f7383,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,'const.bool.COND_3'(s(bool,X33),s(bool,'T_0'),s(bool,X34)))) | s(bool,'T_0') = s(bool,X35) | s(bool,X33) = s(bool,X35)) ) | $spl46),
  inference(cnf_transformation,[],[f7383_D])).
fof(f7383_D,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,'const.bool.COND_3'(s(bool,X33),s(bool,'T_0'),s(bool,X34)))) | s(bool,'T_0') = s(bool,X35) | s(bool,X33) = s(bool,X35)) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f7372,plain,(
  ( ! [X35,X33,X34,X32] : (p(s(bool,'const.bool.COND_3'(s(bool,X33),s(bool,'T_0'),s(bool,X34)))) | s(bool,'F_0') = s(bool,X32) | s(bool,X33) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) ) | $spl12),
  inference(duplicate_literal_removal,[],[f7219])).
fof(f7219,plain,(
  ( ! [X35,X33,X34,X32] : (p(s(bool,'const.bool.COND_3'(s(bool,X33),s(bool,'T_0'),s(bool,X34)))) | s(bool,'F_0') = s(bool,X32) | s(bool,'F_0') = s(bool,X32) | s(bool,X33) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) ) | $spl12),
  inference(superposition,[],[f7176,f1664])).
fof(f7196,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,'const.bool.COND_3'(s(bool,X1),s(bool,X2),s(bool,X3)))) | s(bool,'F_0') = s(bool,X2) | s(bool,'F_0') = s(bool,X1)) ) | $spl42),
  inference(cnf_transformation,[],[f7196_D])).
fof(f7196_D,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,'const.bool.COND_3'(s(bool,X1),s(bool,X2),s(bool,X3)))) | s(bool,'F_0') = s(bool,X2) | s(bool,'F_0') = s(bool,X1)) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f7195,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,'const.bool.COND_3'(s(bool,X1),s(bool,X2),s(bool,X3)))) | s(bool,'F_0') = s(bool,X1) | s(bool,'F_0') = s(bool,X2)) ) | $spl12),
  inference(condensation,[],[f7194])).
fof(f7194,plain,(
  ( ! [X45,X43,X46,X44,X42] : (p(s(bool,'const.bool.COND_3'(s(bool,X44),s(bool,X42),s(bool,X45)))) | s(bool,'F_0') = s(bool,X44) | s(bool,'F_0') = s(bool,X42) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) ) | $spl12),
  inference(subsumption_resolution,[],[f7135,f131])).
fof(f7135,plain,(
  ( ! [X45,X43,X46,X44,X42] : (s(bool,'T_0') != s(bool,X43) | p(s(bool,'const.bool.COND_3'(s(bool,X44),s(bool,X42),s(bool,X45)))) | s(bool,'F_0') = s(bool,X44) | s(bool,'F_0') = s(bool,X42) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) ) | $spl12),
  inference(superposition,[],[f5288,f1664])).
fof(f7189,plain,(
  ( ! [X54,X56,X55,X53] : (s(bool,X53) != s(bool,X54) | s(bool,'F_0') = s(bool,X53) | s(bool,'F_0') = s(bool,X55) | p(s(bool,'const.bool.COND_3'(s(bool,X55),s(bool,X54),s(bool,X56))))) ) | $spl40),
  inference(cnf_transformation,[],[f7189_D])).
fof(f7189_D,plain,(
  ( ! [X54,X56,X55,X53] : (s(bool,X53) != s(bool,X54) | s(bool,'F_0') = s(bool,X53) | s(bool,'F_0') = s(bool,X55) | p(s(bool,'const.bool.COND_3'(s(bool,X55),s(bool,X54),s(bool,X56))))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f7181,plain,(
  ( ! [X57,X54,X56,X55,X53] : (s(bool,X53) != s(bool,X54) | p(s(bool,'const.bool.COND_3'(s(bool,X55),s(bool,X54),s(bool,X56)))) | s(bool,'F_0') = s(bool,X55) | s(bool,'F_0') = s(bool,X53) | s(bool,'T_0') = s(bool,X57)) ) | $spl12),
  inference(duplicate_literal_removal,[],[f7118])).
fof(f7118,plain,(
  ( ! [X57,X54,X56,X55,X53] : (s(bool,X53) != s(bool,X54) | p(s(bool,'const.bool.COND_3'(s(bool,X55),s(bool,X54),s(bool,X56)))) | s(bool,'F_0') = s(bool,X55) | s(bool,'F_0') = s(bool,X53) | s(bool,'T_0') = s(bool,X57) | s(bool,'T_0') = s(bool,X57)) ) | $spl12),
  inference(superposition,[],[f5288,f1664])).
fof(f5683,plain,(
  ( ! [X43,X44,X42] : (s(bool,X43) = s(bool,'const.bool.COND_3'(s(bool,X42),s(bool,X42),s(bool,X44))) | s(bool,'F_0') = s(bool,X42) | s(bool,'F_0') = s(bool,X43)) ) | $spl38),
  inference(cnf_transformation,[],[f5683_D])).
fof(f5683_D,plain,(
  ( ! [X43,X44,X42] : (s(bool,X43) = s(bool,'const.bool.COND_3'(s(bool,X42),s(bool,X42),s(bool,X44))) | s(bool,'F_0') = s(bool,X42) | s(bool,'F_0') = s(bool,X43)) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f5651,plain,(
  ( ! [X45,X43,X44,X42] : (s(bool,X43) = s(bool,'const.bool.COND_3'(s(bool,X42),s(bool,X42),s(bool,X44))) | s(bool,'F_0') = s(bool,X43) | s(bool,'F_0') = s(bool,X42) | s(bool,'T_0') = s(bool,X45)) ) | $spl26),
  inference(duplicate_literal_removal,[],[f5423])).
fof(f5423,plain,(
  ( ! [X45,X43,X44,X42] : (s(bool,X43) = s(bool,'const.bool.COND_3'(s(bool,X42),s(bool,X42),s(bool,X44))) | s(bool,'F_0') = s(bool,X43) | s(bool,'F_0') = s(bool,X42) | s(bool,'T_0') = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) ) | $spl26),
  inference(superposition,[],[f3173,f1664])).
fof(f3773,plain,(
  ( ! [X45,X44,X42] : (s(bool,X42) != s(bool,X44) | s(bool,'F_0') = s(bool,X45) | s(bool,X42) = s(bool,X45) | s(bool,'F_0') = s(bool,X44)) ) | $spl36),
  inference(cnf_transformation,[],[f3773_D])).
fof(f3773_D,plain,(
  ( ! [X45,X44,X42] : (s(bool,X42) != s(bool,X44) | s(bool,'F_0') = s(bool,X45) | s(bool,X42) = s(bool,X45) | s(bool,'F_0') = s(bool,X44)) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f3772,plain,(
  ( ! [X45,X43,X44,X42] : (p(s(bool,X43)) | s(bool,'F_0') = s(bool,X44) | s(bool,X42) != s(bool,X44) | s(bool,X42) = s(bool,X45) | s(bool,'F_0') = s(bool,X45)) ) | $spl10),
  inference(subsumption_resolution,[],[f3755,f79])).
fof(f3755,plain,(
  ( ! [X45,X43,X44,X42] : (p(s(bool,X43)) | s(bool,'F_0') = s(bool,X44) | s(bool,X42) != s(bool,X44) | s(bool,X42) = s(bool,X45) | s(bool,'F_0') = s(bool,X45) | s(bool,'T_0') = s(bool,X43)) ) | $spl10),
  inference(superposition,[],[f3094,f1664])).
fof(f3655,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f3654,f73])).
fof(f3654,plain,(
  p(s(bool,'F_0')) | $spl20),
  inference(forward_demodulation,[],[f3653,f3613])).
fof(f3613,plain,(
  ( ! [X10] : (s(bool,'F_0') = s(bool,X10)) ) | $spl20),
  inference(subsumption_resolution,[],[f3612,f84])).
fof(f3612,plain,(
  ( ! [X10] : (s(bool,'F_0') = s(bool,X10) | ~p(s(bool,X10))) ) | $spl20),
  inference(forward_demodulation,[],[f3611,f102])).
fof(f3611,plain,(
  ( ! [X10,X9] : (~p(s(bool,X10)) | s(bool,X10) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X9)))) ) | $spl20),
  inference(subsumption_resolution,[],[f3509,f71])).
fof(f3509,plain,(
  ( ! [X10,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,X10)) | s(bool,X10) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X9)))) ) | $spl20),
  inference(superposition,[],[f107,f3135])).
fof(f3653,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))))) ) | $spl20),
  inference(subsumption_resolution,[],[f3628,f73])).
fof(f3628,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))))) ) | $spl20),
  inference(backward_demodulation,[],[f3613,f77])).
fof(f3610,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f3609,f102])).
fof(f3609,plain,(
  ( ! [X4] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X4)))) ) | $spl20),
  inference(subsumption_resolution,[],[f3506,f71])).
fof(f3506,plain,(
  ( ! [X4] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X4)))) ) | $spl20),
  inference(superposition,[],[f83,f3135])).
fof(f3608,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f3607,f73])).
fof(f3607,plain,(
  p(s(bool,'F_0')) | $spl20),
  inference(forward_demodulation,[],[f3538,f102])).
fof(f3538,plain,(
  ( ! [X2] : (p(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X2))))) ) | $spl20),
  inference(trivial_inequality_removal,[],[f3504])).
fof(f3504,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X2))))) ) | $spl20),
  inference(superposition,[],[f79,f3135])).
fof(f3606,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f3502,f110])).
fof(f3502,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl20),
  inference(superposition,[],[f3135,f102])).
fof(f3602,plain,(
  ( ! [X28,X26] : (s(bool,X28) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X26))) | s(bool,'F_0') = s(bool,X28)) ) | $spl34),
  inference(cnf_transformation,[],[f3602_D])).
fof(f3602_D,plain,(
  ( ! [X28,X26] : (s(bool,X28) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X26))) | s(bool,'F_0') = s(bool,X28)) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f3539,plain,(
  ( ! [X28,X26,X27] : (s(bool,'T_0') = s(bool,X27) | s(bool,X28) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X26))) | s(bool,'F_0') = s(bool,X28)) ) | $spl20),
  inference(duplicate_literal_removal,[],[f3500])).
fof(f3500,plain,(
  ( ! [X28,X26,X27] : (s(bool,'T_0') = s(bool,X27) | s(bool,X28) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X26))) | s(bool,'F_0') = s(bool,X28) | s(bool,'T_0') = s(bool,X27)) ) | $spl20),
  inference(superposition,[],[f3135,f1664])).
fof(f3594,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f3592,f110])).
fof(f3592,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl20),
  inference(condensation,[],[f3591])).
fof(f3591,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X1)) ) | $spl20),
  inference(duplicate_literal_removal,[],[f3590])).
fof(f3590,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | $spl20),
  inference(condensation,[],[f3589])).
fof(f3589,plain,(
  ( ! [X118,X117,X115] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X115) = s(bool,X117) | s(bool,X117) = s(bool,X118) | s(bool,'T_0') = s(bool,X118)) ) | $spl20),
  inference(backward_demodulation,[],[f3584,f2023])).
fof(f3584,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)))) ) | $spl20),
  inference(duplicate_literal_removal,[],[f3583])).
fof(f3583,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)))) ) | $spl20),
  inference(condensation,[],[f3582])).
fof(f3582,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1))) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) ) | $spl20),
  inference(subsumption_resolution,[],[f3453,f1226])).
fof(f3453,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1))) | s(bool,X0) = s(bool,X2) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) ) | $spl20),
  inference(superposition,[],[f3135,f1664])).
fof(f3578,plain,(
  ( ! [X31,X32] : (s(bool,X31) = s(bool,'const.bool.COND_3'(s(bool,X31),s(bool,'F_0'),s(bool,X32))) | s(bool,'F_0') = s(bool,X31)) ) | $spl32),
  inference(cnf_transformation,[],[f3578_D])).
fof(f3578_D,plain,(
  ( ! [X31,X32] : (s(bool,X31) = s(bool,'const.bool.COND_3'(s(bool,X31),s(bool,'F_0'),s(bool,X32))) | s(bool,'F_0') = s(bool,X31)) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f3568,plain,(
  ( ! [X33,X31,X32] : (s(bool,X31) = s(bool,'const.bool.COND_3'(s(bool,X31),s(bool,'F_0'),s(bool,X32))) | s(bool,'F_0') = s(bool,X31) | s(bool,'T_0') = s(bool,X33)) ) | $spl20),
  inference(duplicate_literal_removal,[],[f3448])).
fof(f3448,plain,(
  ( ! [X33,X31,X32] : (s(bool,X31) = s(bool,'const.bool.COND_3'(s(bool,X31),s(bool,'F_0'),s(bool,X32))) | s(bool,'F_0') = s(bool,X31) | s(bool,'T_0') = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) ) | $spl20),
  inference(superposition,[],[f3135,f1664])).
fof(f3576,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f3575,f73])).
fof(f3575,plain,(
  p(s(bool,'F_0')) | $spl20),
  inference(forward_demodulation,[],[f3435,f102])).
fof(f3435,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X0))))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f3135,f79])).
fof(f3574,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f3573,f73])).
fof(f3573,plain,(
  p(s(bool,'F_0')) | $spl20),
  inference(forward_demodulation,[],[f3436,f102])).
fof(f3436,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X0))))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f110,f110,f3135,f923])).
fof(f3370,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f73,f3165])).
fof(f3369,plain,(
  $false | $spl24),
  inference(global_subsumption,[],[f70,f71,f73,f74,f76,f77,f80,f79,f84,f83,f102,f103,f107,f104,f108,f111,f110,f118,f131,f162,f187,f207,f219,f249,f258,f263,f264,f325,f326,f416,f417,f376,f450,f453,f454,f455,f456,f468,f469,f493,f522,f523,f524,f525,f542,f558,f856,f575,f576,f598,f611,f612,f891,f642,f648,f651,f681,f687,f690,f915,f917,f921,f923,f713,f714,f715,f934,f937,f938,f766,f769,f792,f780,f996,f957,f958,f1050,f1072,f1074,f1093,f1094,f1157,f1165,f1168,f1222,f1223,f1224,f1226,f1227,f1663,f1664,f1665,f1294,f1297,f1739,f1740,f1783,f1795,f1799,f1801,f1944,f1951,f1952,f2996,f2989,f1993,f1994,f2966,f3031,f2023,f2024,f2025,f2948,f2928,f2061,f2919,f2916,f2091,f2900,f2895,f3048,f2889,f2886,f3059,f2121,f2122,f2125,f2881,f2871,f3075,f2866,f3076,f2155,f2156,f2843,f2842,f2837,f2836,f2835,f2831,f2829,f2207,f3107,f2216,f2219,f2270,f2279,f2282,f2283,f2285,f3133,f2333,f3137,f3141,f2354,f3143,f2363,f2366,f3148,f3163,f3170,f2417,f3172,f2426,f2429,f3174,f3179,f3180,f2700,f2480,f2489,f2492,f2494,f2495,f2498,f2499,f2500,f2502,f2503,f3200,f3201,f3202,f3219,f2585,f2588,f3165,f3345,f3362,f3368])).
fof(f3368,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | $spl24),
  inference(subsumption_resolution,[],[f83,f3165])).
fof(f3362,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl24),
  inference(condensation,[],[f3361])).
fof(f3361,plain,(
  ( ! [X18,X16] : (s(bool,X16) = s(bool,X18) | s(bool,'F_0') = s(bool,X16)) ) | $spl24),
  inference(subsumption_resolution,[],[f453,f3165])).
fof(f3345,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl24),
  inference(condensation,[],[f3344])).
fof(f3344,plain,(
  ( ! [X2,X3] : (s(bool,X2) = s(bool,X3) | s(bool,'F_0') = s(bool,X2)) ) | $spl24),
  inference(duplicate_literal_removal,[],[f3343])).
fof(f3343,plain,(
  ( ! [X2,X3] : (s(bool,X2) = s(bool,X3) | s(bool,'F_0') = s(bool,X2) | s(bool,X2) = s(bool,X3)) ) | $spl24),
  inference(condensation,[],[f3342])).
fof(f3342,plain,(
  ( ! [X26,X24,X23,X25] : (s(bool,X23) = s(bool,X24) | s(bool,X25) = s(bool,X26) | s(bool,'F_0') = s(bool,X26) | s(bool,X23) = s(bool,X25)) ) | $spl24),
  inference(subsumption_resolution,[],[f3201,f3165])).
fof(f3367,plain,(
  $false | ($spl2 | $spl24)),
  inference(global_subsumption,[],[f3366,f70,f71,f73,f74,f76,f77,f80,f79,f84,f83,f102,f103,f107,f104,f108,f111,f110,f118,f131,f162,f187,f207,f219,f249,f258,f263,f264,f325,f326,f416,f417,f376,f450,f453,f454,f455,f456,f468,f469,f493,f522,f523,f524,f525,f542,f558,f856,f575,f576,f598,f611,f612,f891,f642,f648,f651,f681,f687,f690,f915,f917,f921,f923,f713,f714,f715,f934,f937,f938,f766,f769,f792,f780,f996,f957,f958,f1050,f1072,f1074,f1093,f1094,f1157,f1165,f1168,f1222,f1223,f1224,f1226,f1227,f1663,f1664,f1665,f1294,f1297,f1739,f1740,f1783,f1795,f1799,f1801,f1944,f1951,f1952,f2996,f2989,f1993,f1994,f2966,f3031,f2023,f2024,f2025,f2948,f2928,f2061,f2919,f2916,f2091,f2900,f2895,f3048,f2889,f2886,f3059,f2121,f2122,f2125,f2881,f2871,f3075,f2866,f3076,f2155,f2156,f2843,f2842,f2837,f2836,f2835,f2831,f2829,f2207,f3107,f2216,f2219,f2270,f2279,f2282,f2283,f2285,f3133,f2333,f3137,f3141,f2354,f3143,f2363,f2366,f3148,f3163,f3170,f2417,f3172,f2426,f2429,f3174,f3179,f3180,f2700,f2480,f2489,f2492,f2494,f2495,f2498,f2499,f2500,f2502,f2503,f3200,f3201,f3202,f3219,f2585,f2588,f3165,f3345,f3362])).
fof(f3366,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) ) | $spl2),
  inference(forward_demodulation,[],[f103,f3008])).
fof(f3363,plain,(
  $false | $spl24),
  inference(global_subsumption,[],[f70,f71,f73,f74,f76,f77,f80,f79,f84,f83,f102,f103,f107,f104,f108,f111,f110,f118,f131,f162,f187,f207,f219,f249,f258,f263,f264,f325,f326,f416,f417,f376,f450,f453,f454,f455,f456,f468,f469,f493,f522,f523,f524,f525,f542,f558,f856,f575,f576,f598,f611,f612,f891,f642,f648,f651,f681,f687,f690,f915,f917,f921,f923,f713,f714,f715,f934,f937,f938,f766,f769,f792,f780,f996,f957,f958,f1050,f1072,f1074,f1093,f1094,f1157,f1165,f1168,f1222,f1223,f1224,f1226,f1227,f1663,f1664,f1665,f1294,f1297,f1739,f1740,f1783,f1795,f1799,f1801,f1944,f1951,f1952,f2996,f2989,f1993,f1994,f2966,f3031,f2023,f2024,f2025,f2948,f2928,f2061,f2919,f2916,f2091,f2900,f2895,f3048,f2889,f2886,f3059,f2121,f2122,f2125,f2881,f2871,f3075,f2866,f3076,f2155,f2156,f2843,f2842,f2837,f2836,f2835,f2831,f2829,f2207,f3107,f2216,f2219,f2270,f2279,f2282,f2283,f2285,f3133,f2333,f3137,f3141,f2354,f3143,f2363,f2366,f3148,f3163,f3170,f2417,f3172,f2426,f2429,f3174,f3179,f3180,f2700,f2480,f2489,f2492,f2494,f2495,f2498,f2499,f2500,f2502,f2503,f3200,f3201,f3202,f3219,f2585,f2588,f3165,f3345,f3362])).
fof(f3346,plain,(
  $false | $spl24),
  inference(global_subsumption,[],[f70,f71,f73,f74,f76,f77,f80,f79,f84,f83,f102,f103,f107,f104,f108,f111,f110,f118,f131,f162,f187,f207,f219,f249,f258,f263,f264,f325,f326,f416,f417,f376,f450,f453,f454,f455,f456,f468,f469,f493,f522,f523,f524,f525,f542,f558,f856,f575,f576,f598,f611,f612,f891,f642,f648,f651,f681,f687,f690,f915,f917,f921,f923,f713,f714,f715,f934,f937,f938,f766,f769,f792,f780,f996,f957,f958,f1050,f1072,f1074,f1093,f1094,f1157,f1165,f1168,f1222,f1223,f1224,f1226,f1227,f1663,f1664,f1665,f1294,f1297,f1739,f1740,f1783,f1795,f1799,f1801,f1944,f1951,f1952,f2996,f2989,f1993,f1994,f2966,f3031,f2023,f2024,f2025,f2948,f2928,f2061,f2919,f2916,f2091,f2900,f2895,f3048,f2889,f2886,f3059,f2121,f2122,f2125,f2881,f2871,f3075,f2866,f3076,f2155,f2156,f2843,f2842,f2837,f2836,f2835,f2831,f2829,f2207,f3107,f2216,f2219,f2270,f2279,f2282,f2283,f2285,f3133,f2333,f3137,f3141,f2354,f3143,f2363,f2366,f3148,f3163,f3170,f2417,f3172,f2426,f2429,f3174,f3179,f3180,f2700,f2480,f2489,f2492,f2494,f2495,f2498,f2499,f2500,f2502,f2503,f3200,f3201,f3202,f3219,f2585,f2588,f3165,f3345])).
fof(f3341,plain,(
  $false | $spl2),
  inference(subsumption_resolution,[],[f3334,f3008])).
fof(f3334,plain,(
  s(bool,'T_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | $spl2),
  inference(backward_demodulation,[],[f3333,f111])).
fof(f3333,plain,(
  ( ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0)))) ) | $spl2),
  inference(forward_demodulation,[],[f3332,f102])).
fof(f3332,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0))) = s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real',X0),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))) ) | $spl2),
  inference(forward_demodulation,[],[f74,f3008])).
fof(f3330,plain,(
  $false | ($spl18 | $spl24)),
  inference(global_subsumption,[],[f70,f71,f73,f74,f76,f77,f80,f79,f84,f83,f102,f103,f107,f104,f108,f111,f110,f118,f131,f162,f187,f207,f219,f249,f258,f263,f264,f325,f326,f416,f417,f376,f450,f453,f454,f455,f456,f468,f469,f493,f522,f523,f524,f525,f542,f558,f856,f575,f576,f598,f611,f612,f891,f642,f648,f651,f681,f687,f690,f915,f917,f921,f923,f713,f714,f715,f934,f937,f938,f766,f769,f792,f780,f996,f957,f958,f1050,f1072,f1074,f1093,f1094,f1157,f1165,f1168,f1222,f1223,f1224,f1226,f1227,f1663,f1664,f1665,f1294,f1297,f1739,f1740,f1783,f1795,f1799,f1801,f1944,f1951,f1952,f2996,f2989,f1993,f1994,f2966,f3031,f2023,f2024,f2025,f2948,f2928,f2061,f2919,f2916,f2091,f2900,f2895,f3048,f2889,f2886,f3059,f2121,f2122,f2125,f2881,f2871,f3075,f2866,f3076,f2155,f2156,f2843,f2842,f2837,f2836,f2835,f2831,f2829,f2207,f3107,f2216,f2219,f2270,f2279,f2282,f2283,f2285,f3133,f2333,f3137,f3141,f2354,f3143,f2363,f2366,f3148,f3163,f3170,f2417,f3172,f2426,f2429,f3174,f3179,f3180,f2700,f2480,f2489,f2492,f2494,f2495,f2498,f2499,f2500,f2502,f2503,f3200,f3201,f3202,f3219,f2585,f2588,f3165,f3134,f3226,f3230,f3321,f3324,f3327,f3329])).
fof(f3329,plain,(
  ( ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0)))) ) | ($spl18 | $spl24)),
  inference(forward_demodulation,[],[f3309,f102])).
fof(f3309,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0))) = s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real',X0),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))) ) | ($spl18 | $spl24)),
  inference(backward_demodulation,[],[f3226,f74])).
fof(f3327,plain,(
  ( ! [X12] : (s(bool,'F_0') = s(bool,X12)) ) | ($spl18 | $spl24)),
  inference(subsumption_resolution,[],[f3259,f3226])).
fof(f3259,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X12),s(bool,X13))) | s(bool,'F_0') = s(bool,X12)) ) | ($spl18 | $spl24)),
  inference(backward_demodulation,[],[f3226,f1944])).
fof(f3324,plain,(
  s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))) | ($spl18 | $spl24)),
  inference(forward_demodulation,[],[f3239,f102])).
fof(f3239,plain,(
  s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))) = s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))))) | ($spl18 | $spl24)),
  inference(backward_demodulation,[],[f3226,f456])).
fof(f3321,plain,(
  ( ! [X6] : (s(bool,X6) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0)))))) ) | ($spl18 | $spl24)),
  inference(subsumption_resolution,[],[f3237,f3226])).
fof(f3237,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X6),s(bool,X7))) | s(bool,X6) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0)))))) ) | ($spl18 | $spl24)),
  inference(backward_demodulation,[],[f3226,f450])).
fof(f3230,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) ) | ($spl18 | $spl24)),
  inference(backward_demodulation,[],[f3226,f103])).
fof(f3226,plain,(
  ( ! [X63] : (s(bool,'T_0') = s(bool,X63)) ) | ($spl18 | $spl24)),
  inference(global_subsumption,[],[f3134,f70,f71,f73,f74,f76,f77,f80,f79,f84,f83,f102,f103,f107,f104,f108,f111,f110,f118,f131,f162,f187,f207,f219,f249,f258,f263,f264,f325,f326,f416,f417,f376,f450,f453,f454,f455,f456,f468,f469,f493,f522,f523,f524,f525,f542,f558,f856,f575,f576,f598,f611,f612,f891,f642,f648,f651,f681,f687,f690,f915,f917,f921,f923,f713,f714,f715,f934,f937,f938,f766,f769,f792,f780,f996,f957,f958,f1050,f1072,f1074,f1093,f1094,f1157,f1165,f1168,f1222,f1223,f1224,f1226,f1227,f1663,f1664,f1665,f1294,f1297,f1739,f1740,f1783,f1795,f1799,f1801,f1944,f1951,f1952,f2996,f2989,f1993,f1994,f2966,f3031,f2023,f2024,f2025,f2948,f2928,f2061,f2919,f2916,f2091,f2900,f2895,f3048,f2889,f2886,f3059,f2121,f2122,f2125,f2881,f2871,f3075,f2866,f3076,f2155,f2156,f2843,f2842,f2837,f2836,f2835,f2831,f2829,f2207,f3107,f2216,f2219,f2270,f2279,f2282,f2283,f2285,f3133,f2333,f3137,f3141,f2354,f3143,f2363,f2366,f3148,f3163,f3170,f2417,f3172,f2426,f2429,f3174,f3179,f3180,f2700,f2480,f2489,f2492,f2494,f2495,f2498,f2499,f2500,f2502,f2503,f3200,f3201,f3202,f3219,f2585,f2588,f3165])).
fof(f3328,plain,(
  $false | ($spl18 | $spl24)),
  inference(global_subsumption,[],[f70,f71,f73,f74,f76,f77,f80,f79,f84,f83,f102,f103,f107,f104,f108,f111,f110,f118,f131,f162,f187,f207,f219,f249,f258,f263,f264,f325,f326,f416,f417,f376,f450,f453,f454,f455,f456,f468,f469,f493,f522,f523,f524,f525,f542,f558,f856,f575,f576,f598,f611,f612,f891,f642,f648,f651,f681,f687,f690,f915,f917,f921,f923,f713,f714,f715,f934,f937,f938,f766,f769,f792,f780,f996,f957,f958,f1050,f1072,f1074,f1093,f1094,f1157,f1165,f1168,f1222,f1223,f1224,f1226,f1227,f1663,f1664,f1665,f1294,f1297,f1739,f1740,f1783,f1795,f1799,f1801,f1944,f1951,f1952,f2996,f2989,f1993,f1994,f2966,f3031,f2023,f2024,f2025,f2948,f2928,f2061,f2919,f2916,f2091,f2900,f2895,f3048,f2889,f2886,f3059,f2121,f2122,f2125,f2881,f2871,f3075,f2866,f3076,f2155,f2156,f2843,f2842,f2837,f2836,f2835,f2831,f2829,f2207,f3107,f2216,f2219,f2270,f2279,f2282,f2283,f2285,f3133,f2333,f3137,f3141,f2354,f3143,f2363,f2366,f3148,f3163,f3170,f2417,f3172,f2426,f2429,f3174,f3179,f3180,f2700,f2480,f2489,f2492,f2494,f2495,f2498,f2499,f2500,f2502,f2503,f3200,f3201,f3202,f3219,f2585,f2588,f3165,f3134,f3226,f3230,f3321,f3324,f3327])).
fof(f3320,plain,(
  $false | ($spl18 | $spl24)),
  inference(subsumption_resolution,[],[f3319,f3165])).
fof(f3319,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl18 | $spl24)),
  inference(subsumption_resolution,[],[f3228,f3226])).
fof(f3228,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | ($spl18 | $spl24)),
  inference(backward_demodulation,[],[f3226,f83])).
fof(f3318,plain,(
  $false | ($spl18 | $spl24)),
  inference(subsumption_resolution,[],[f3227,f71])).
fof(f3227,plain,(
  ~p(s(bool,'T_0')) | ($spl18 | $spl24)),
  inference(backward_demodulation,[],[f3226,f73])).
fof(f3181,plain,(
  ( ! [X105,X106,X104] : (s(bool,X104) != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X105))) | s(bool,'F_0') = s(bool,X106) | s(bool,X104) = s(bool,X106)) ) | $spl30),
  inference(cnf_transformation,[],[f3181_D])).
fof(f3181_D,plain,(
  ( ! [X105,X106,X104] : (s(bool,X104) != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X105))) | s(bool,'F_0') = s(bool,X106) | s(bool,X104) = s(bool,X106)) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f2700,plain,(
  ( ! [X103,X105,X106,X104] : (s(bool,X104) != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X105))) | s(bool,'T_0') = s(bool,X103) | s(bool,X104) = s(bool,X106) | s(bool,'F_0') = s(bool,X106)) )),
  inference(duplicate_literal_removal,[],[f2456])).
fof(f2456,plain,(
  ( ! [X103,X105,X106,X104] : (s(bool,X104) != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X105))) | s(bool,'T_0') = s(bool,X103) | s(bool,X104) = s(bool,X106) | s(bool,'F_0') = s(bool,X106) | s(bool,'T_0') = s(bool,X103)) )),
  inference(superposition,[],[f792,f1664])).
fof(f3175,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | s(bool,'F_0') = s(bool,X5) | s(bool,X4) = s(bool,X5)) ) | $spl28),
  inference(cnf_transformation,[],[f3175_D])).
fof(f3175_D,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | s(bool,'F_0') = s(bool,X5) | s(bool,X4) = s(bool,X5)) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f3174,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,X4) | p(s(bool,X3)) | s(bool,X4) = s(bool,X5) | s(bool,'F_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f2431,f79])).
fof(f2431,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,X4) | p(s(bool,X3)) | s(bool,X4) = s(bool,X5) | s(bool,'F_0') = s(bool,X5) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f79,f1664])).
fof(f3173,plain,(
  ( ! [X2,X1] : (s(bool,X1) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X2))) | s(bool,'F_0') = s(bool,X1)) ) | $spl26),
  inference(cnf_transformation,[],[f3173_D])).
fof(f3173_D,plain,(
  ( ! [X2,X1] : (s(bool,X1) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X2))) | s(bool,'F_0') = s(bool,X1)) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f3172,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X2))) | s(bool,'F_0') = s(bool,X1)) )),
  inference(condensation,[],[f3171])).
fof(f3171,plain,(
  ( ! [X212,X210,X211,X209] : (s(bool,'T_0') = s(bool,X211) | s(bool,X212) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X209))) | s(bool,'F_0') = s(bool,X212) | s(bool,'T_0') = s(bool,X210)) )),
  inference(subsumption_resolution,[],[f2421,f131])).
fof(f2421,plain,(
  ( ! [X212,X210,X211,X209] : (s(bool,X210) != s(bool,X211) | s(bool,'T_0') = s(bool,X211) | s(bool,X212) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X209))) | s(bool,'F_0') = s(bool,X212) | s(bool,'T_0') = s(bool,X210)) )),
  inference(superposition,[],[f792,f1664])).
fof(f3164,plain,(
  ( ! [X118,X116,X120,X119] : (s(bool,X116) != s(bool,X118) | s(bool,'F_0') = s(bool,X120) | s(bool,X116) = s(bool,X120) | s(bool,X118) = s(bool,X119) | s(bool,'T_0') = s(bool,X119)) ) | $spl22),
  inference(cnf_transformation,[],[f3164_D])).
fof(f3164_D,plain,(
  ( ! [X118,X116,X120,X119] : (s(bool,X116) != s(bool,X118) | s(bool,'F_0') = s(bool,X120) | s(bool,X116) = s(bool,X120) | s(bool,X118) = s(bool,X119) | s(bool,'T_0') = s(bool,X119)) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f3163,plain,(
  ( ! [X118,X116,X120,X119,X117] : (p(s(bool,X117)) | s(bool,X116) != s(bool,X118) | s(bool,'T_0') = s(bool,X119) | s(bool,X118) = s(bool,X119) | s(bool,X116) = s(bool,X120) | s(bool,'F_0') = s(bool,X120)) )),
  inference(subsumption_resolution,[],[f2396,f79])).
fof(f2396,plain,(
  ( ! [X118,X116,X120,X119,X117] : (p(s(bool,X117)) | s(bool,X116) != s(bool,X118) | s(bool,'T_0') = s(bool,X119) | s(bool,X118) = s(bool,X119) | s(bool,X116) = s(bool,X120) | s(bool,'F_0') = s(bool,X120) | s(bool,'T_0') = s(bool,X117)) )),
  inference(superposition,[],[f923,f1664])).
fof(f3134,plain,(
  ( ! [X62,X60,X63] : (s(bool,X60) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))) | s(bool,'T_0') = s(bool,X63) | s(bool,X62) = s(bool,X63) | s(bool,X60) = s(bool,X62)) ) | $spl18),
  inference(cnf_transformation,[],[f3134_D])).
fof(f3134_D,plain,(
  ( ! [X62,X60,X63] : (s(bool,X60) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))) | s(bool,'T_0') = s(bool,X63) | s(bool,X62) = s(bool,X63) | s(bool,X60) = s(bool,X62)) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f2333,plain,(
  ( ! [X61,X62,X60,X63] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X61))) | s(bool,X60) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))) | s(bool,X60) = s(bool,X62) | s(bool,X62) = s(bool,X63) | s(bool,'T_0') = s(bool,X63)) )),
  inference(superposition,[],[f263,f1664])).
fof(f3097,plain,(
  ( ! [X146,X147,X145] : (s(bool,X145) = s(bool,'const.bool.COND_3'(s(bool,X145),s(bool,X146),s(bool,X147))) | s(bool,'F_0') = s(bool,X145) | s(bool,X146) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0)))))) ) | $spl16),
  inference(cnf_transformation,[],[f3097_D])).
fof(f3097_D,plain,(
  ( ! [X146,X147,X145] : (s(bool,X145) = s(bool,'const.bool.COND_3'(s(bool,X145),s(bool,X146),s(bool,X147))) | s(bool,'F_0') = s(bool,X145) | s(bool,X146) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0)))))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f2836,plain,(
  ( ! [X146,X147,X145,X148] : (s(bool,X145) = s(bool,'const.bool.COND_3'(s(bool,X145),s(bool,X146),s(bool,X147))) | s(bool,X146) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))) | s(bool,'F_0') = s(bool,X145) | s(bool,'T_0') = s(bool,X148)) )),
  inference(duplicate_literal_removal,[],[f2193])).
fof(f2193,plain,(
  ( ! [X146,X147,X145,X148] : (s(bool,X145) = s(bool,'const.bool.COND_3'(s(bool,X145),s(bool,X146),s(bool,X147))) | s(bool,X146) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))) | s(bool,'F_0') = s(bool,X145) | s(bool,'T_0') = s(bool,X148) | s(bool,'T_0') = s(bool,X148)) )),
  inference(superposition,[],[f263,f1664])).
fof(f3096,plain,(
  ( ! [X142,X143,X141] : (s(bool,X141) = s(bool,'const.bool.COND_3'(s(bool,X141),s(bool,X142),s(bool,X143))) | s(bool,'F_0') = s(bool,X141) | s(bool,'F_0') = s(bool,X142)) ) | $spl14),
  inference(cnf_transformation,[],[f3096_D])).
fof(f3096_D,plain,(
  ( ! [X142,X143,X141] : (s(bool,X141) = s(bool,'const.bool.COND_3'(s(bool,X141),s(bool,X142),s(bool,X143))) | s(bool,'F_0') = s(bool,X141) | s(bool,'F_0') = s(bool,X142)) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f2837,plain,(
  ( ! [X144,X142,X143,X141] : (s(bool,X141) = s(bool,'const.bool.COND_3'(s(bool,X141),s(bool,X142),s(bool,X143))) | s(bool,'F_0') = s(bool,X142) | s(bool,'F_0') = s(bool,X141) | s(bool,'T_0') = s(bool,X144)) )),
  inference(duplicate_literal_removal,[],[f2192])).
fof(f2192,plain,(
  ( ! [X144,X142,X143,X141] : (s(bool,X141) = s(bool,'const.bool.COND_3'(s(bool,X141),s(bool,X142),s(bool,X143))) | s(bool,'F_0') = s(bool,X142) | s(bool,'F_0') = s(bool,X141) | s(bool,'T_0') = s(bool,X144) | s(bool,'T_0') = s(bool,X144)) )),
  inference(superposition,[],[f187,f1664])).
fof(f3095,plain,(
  ( ! [X127,X128,X126,X129] : (s(X127,X128) = s(X127,'const.bool.COND_3'(s(bool,X126),s(X127,X128),s(X127,X129))) | s(bool,'F_0') = s(bool,X126)) ) | $spl12),
  inference(cnf_transformation,[],[f3095_D])).
fof(f3095_D,plain,(
  ( ! [X127,X128,X126,X129] : (s(X127,X128) = s(X127,'const.bool.COND_3'(s(bool,X126),s(X127,X128),s(X127,X129))) | s(bool,'F_0') = s(bool,X126)) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f2842,plain,(
  ( ! [X127,X130,X128,X126,X129] : (s(X127,X128) = s(X127,'const.bool.COND_3'(s(bool,X126),s(X127,X128),s(X127,X129))) | s(bool,'F_0') = s(bool,X126) | s(bool,'T_0') = s(bool,X130)) )),
  inference(duplicate_literal_removal,[],[f2189])).
fof(f2189,plain,(
  ( ! [X127,X130,X128,X126,X129] : (s(X127,X128) = s(X127,'const.bool.COND_3'(s(bool,X126),s(X127,X128),s(X127,X129))) | s(bool,'F_0') = s(bool,X126) | s(bool,'T_0') = s(bool,X130) | s(bool,'T_0') = s(bool,X130)) )),
  inference(superposition,[],[f102,f1664])).
fof(f3094,plain,(
  ( ! [X123,X124] : (p(s(bool,X124)) | s(bool,'F_0') = s(bool,X123) | s(bool,X123) != s(bool,X124)) ) | $spl10),
  inference(cnf_transformation,[],[f3094_D])).
fof(f3094_D,plain,(
  ( ! [X123,X124] : (p(s(bool,X124)) | s(bool,'F_0') = s(bool,X123) | s(bool,X123) != s(bool,X124)) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f2843,plain,(
  ( ! [X125,X123,X124] : (s(bool,X123) != s(bool,X124) | p(s(bool,X124)) | s(bool,'F_0') = s(bool,X123) | s(bool,'T_0') = s(bool,X125)) )),
  inference(duplicate_literal_removal,[],[f2188])).
fof(f2188,plain,(
  ( ! [X125,X123,X124] : (s(bool,X123) != s(bool,X124) | p(s(bool,X124)) | s(bool,'F_0') = s(bool,X123) | s(bool,'T_0') = s(bool,X125) | s(bool,'T_0') = s(bool,X125)) )),
  inference(superposition,[],[f79,f1664])).
fof(f3047,plain,(
  ( ! [X21,X20] : (s(bool,X20) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))) | s(bool,'F_0') = s(bool,X21) | s(bool,X20) = s(bool,X21)) ) | $spl6),
  inference(cnf_transformation,[],[f3047_D])).
fof(f3047_D,plain,(
  ( ! [X21,X20] : (s(bool,X20) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))) | s(bool,'F_0') = s(bool,X21) | s(bool,X20) = s(bool,X21)) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f2866,plain,(
  ( ! [X74,X75,X73] : (s(bool,'T_0') = s(bool,X73) | s(bool,X74) = s(bool,X75) | s(bool,'F_0') = s(bool,X75) | s(bool,X74) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0)))))) )),
  inference(duplicate_literal_removal,[],[f2153])).
fof(f2153,plain,(
  ( ! [X74,X75,X73] : (s(bool,'T_0') = s(bool,X73) | s(bool,X74) = s(bool,X75) | s(bool,'F_0') = s(bool,X75) | s(bool,X74) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))) | s(bool,'T_0') = s(bool,X73)) )),
  inference(superposition,[],[f1664,f249])).
fof(f3075,plain,(
  ( ! [X70,X72,X71,X69] : (s(bool,X70) = s(bool,X71) | s(bool,'F_0') = s(bool,X71) | s(bool,X70) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X69),s(bool,X72)))) )),
  inference(subsumption_resolution,[],[f2152,f376])).
fof(f2152,plain,(
  ( ! [X70,X72,X71,X69] : (s(bool,'T_0') = s(bool,X69) | s(bool,X70) = s(bool,X71) | s(bool,'F_0') = s(bool,X71) | s(bool,X70) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X69),s(bool,X72)))) )),
  inference(superposition,[],[f1664,f263])).
fof(f3015,plain,(
  ( ! [X74,X72,X71,X75] : (s(bool,X71) = s(bool,X72) | s(bool,X71) = s(bool,X74) | s(bool,X74) = s(bool,X75) | s(bool,'F_0') = s(bool,X75) | s(bool,'F_0') = s(bool,X72)) ) | $spl4),
  inference(cnf_transformation,[],[f3015_D])).
fof(f3015_D,plain,(
  ( ! [X74,X72,X71,X75] : (s(bool,X71) = s(bool,X72) | s(bool,X71) = s(bool,X74) | s(bool,X74) = s(bool,X75) | s(bool,'F_0') = s(bool,X75) | s(bool,'F_0') = s(bool,X72)) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f2871,plain,(
  ( ! [X47,X50,X48,X51,X49] : (s(bool,'T_0') = s(bool,X48) | s(bool,X49) = s(bool,X50) | s(bool,'F_0') = s(bool,X50) | s(bool,X47) = s(bool,X49) | s(bool,X47) = s(bool,X51) | s(bool,'F_0') = s(bool,X51)) )),
  inference(duplicate_literal_removal,[],[f2148])).
fof(f2148,plain,(
  ( ! [X47,X50,X48,X51,X49] : (s(bool,'T_0') = s(bool,X48) | s(bool,X49) = s(bool,X50) | s(bool,'F_0') = s(bool,X50) | s(bool,X47) = s(bool,X49) | s(bool,X47) = s(bool,X51) | s(bool,'F_0') = s(bool,X51) | s(bool,'T_0') = s(bool,X48)) )),
  inference(superposition,[],[f1664,f1664])).
fof(f2881,plain,(
  ( ! [X57,X54,X56,X55,X53] : (s(bool,X53) = s(bool,X54) | s(bool,X55) = s(bool,X56) | s(bool,'F_0') = s(bool,X56) | s(bool,X54) = s(bool,X55) | s(bool,'F_0') = s(bool,X53) | s(bool,'T_0') = s(bool,X57)) )),
  inference(duplicate_literal_removal,[],[f2133])).
fof(f2133,plain,(
  ( ! [X57,X54,X56,X55,X53] : (s(bool,X53) = s(bool,X54) | s(bool,X55) = s(bool,X56) | s(bool,'F_0') = s(bool,X56) | s(bool,X54) = s(bool,X55) | s(bool,'F_0') = s(bool,X53) | s(bool,'T_0') = s(bool,X57) | s(bool,'T_0') = s(bool,X57)) )),
  inference(superposition,[],[f1664,f1664])).
fof(f2886,plain,(
  ( ! [X70,X68,X66,X69,X67] : (s(bool,X67) = s(bool,X68) | s(bool,X68) = s(bool,X69) | s(bool,'F_0') = s(bool,X69) | s(bool,'T_0') = s(bool,X66) | s(bool,X67) = s(bool,X70) | s(bool,'F_0') = s(bool,X70)) )),
  inference(duplicate_literal_removal,[],[f2107])).
fof(f2107,plain,(
  ( ! [X70,X68,X66,X69,X67] : (s(bool,X67) = s(bool,X68) | s(bool,X68) = s(bool,X69) | s(bool,'F_0') = s(bool,X69) | s(bool,'T_0') = s(bool,X66) | s(bool,X67) = s(bool,X70) | s(bool,'F_0') = s(bool,X70) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f1664,f1664])).
fof(f3007,plain,(
  ( ! [X2,X1] : (s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X1) | s(bool,'F_0') = s(bool,X2)) ) | $spl0),
  inference(cnf_transformation,[],[f3007_D])).
fof(f3007_D,plain,(
  ( ! [X2,X1] : (s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X1) | s(bool,'F_0') = s(bool,X2)) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f2889,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2)) )),
  inference(condensation,[],[f2888])).
fof(f2888,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X3)) )),
  inference(duplicate_literal_removal,[],[f2887])).
fof(f2887,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(condensation,[],[f2104])).
fof(f2104,plain,(
  ( ! [X54,X52,X55,X53,X51] : (s(bool,'F_0') = s(bool,X52) | s(bool,X52) = s(bool,X53) | s(bool,'F_0') = s(bool,X53) | s(bool,'T_0') = s(bool,X51) | s(bool,X51) = s(bool,X54) | s(bool,X54) = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) )),
  inference(superposition,[],[f1664,f1664])).
fof(f3048,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,X23) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))) | s(bool,X23) = s(bool,X24) | s(bool,'F_0') = s(bool,X24) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X22),s(bool,X25)))) )),
  inference(subsumption_resolution,[],[f2098,f376])).
fof(f2098,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,X23) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))) | s(bool,X23) = s(bool,X24) | s(bool,'F_0') = s(bool,X24) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X22),s(bool,X25)))) )),
  inference(superposition,[],[f1664,f263])).
fof(f2895,plain,(
  ( ! [X21,X19,X20] : (s(bool,X20) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))) | s(bool,X20) = s(bool,X21) | s(bool,'F_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X19)) )),
  inference(duplicate_literal_removal,[],[f2097])).
fof(f2097,plain,(
  ( ! [X21,X19,X20] : (s(bool,X20) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK0))))) | s(bool,X20) = s(bool,X21) | s(bool,'F_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f1664,f249])).
fof(f2900,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) )),
  inference(duplicate_literal_removal,[],[f2092])).
fof(f2092,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1664,f76])).
fof(f2916,plain,(
  ( ! [X59,X57,X60,X58,X56] : (s(bool,'T_0') = s(bool,X57) | s(bool,X56) = s(bool,X58) | s(bool,'F_0') = s(bool,X58) | s(bool,X59) = s(bool,X60) | s(bool,'F_0') = s(bool,X60) | s(bool,X56) = s(bool,X59)) )),
  inference(duplicate_literal_removal,[],[f2075])).
fof(f2075,plain,(
  ( ! [X59,X57,X60,X58,X56] : (s(bool,'T_0') = s(bool,X57) | s(bool,X56) = s(bool,X58) | s(bool,'F_0') = s(bool,X58) | s(bool,'T_0') = s(bool,X57) | s(bool,X59) = s(bool,X60) | s(bool,'F_0') = s(bool,X60) | s(bool,X56) = s(bool,X59)) )),
  inference(superposition,[],[f1664,f1664])).
fof(f2919,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,'F_0') = s(bool,X1)) )),
  inference(condensation,[],[f2918])).
fof(f2918,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X3)) )),
  inference(duplicate_literal_removal,[],[f2917])).
fof(f2917,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(condensation,[],[f2073])).
fof(f2073,plain,(
  ( ! [X47,X50,X48,X46,X49] : (s(bool,X47) = s(bool,X48) | s(bool,X46) = s(bool,X49) | s(bool,'F_0') = s(bool,X49) | s(bool,'T_0') = s(bool,X48) | s(bool,'F_0') = s(bool,X46) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f1664,f1664])).
fof(f2928,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2)) )),
  inference(condensation,[],[f2927])).
fof(f2927,plain,(
  ( ! [X4,X2,X3,X1] : (s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X1) | s(bool,X3) = s(bool,X4) | s(bool,'F_0') = s(bool,X4)) )),
  inference(duplicate_literal_removal,[],[f2926])).
fof(f2926,plain,(
  ( ! [X4,X2,X3,X1] : (s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X1) | s(bool,X3) = s(bool,X4) | s(bool,'F_0') = s(bool,X4) | s(bool,'T_0') = s(bool,X1)) )),
  inference(condensation,[],[f2054])).
fof(f2054,plain,(
  ( ! [X61,X64,X62,X60,X63] : (s(bool,'F_0') = s(bool,X61) | s(bool,X60) = s(bool,X62) | s(bool,X62) = s(bool,X63) | s(bool,'T_0') = s(bool,X63) | s(bool,X61) = s(bool,X64) | s(bool,'F_0') = s(bool,X64) | s(bool,'T_0') = s(bool,X60)) )),
  inference(superposition,[],[f1664,f1664])).
fof(f2948,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,'F_0') = s(bool,X1)) )),
  inference(condensation,[],[f2947])).
fof(f2947,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1) | s(bool,X2) = s(bool,X3) | s(bool,'F_0') = s(bool,X3) | s(bool,'F_0') = s(bool,X2)) )),
  inference(duplicate_literal_removal,[],[f2946])).
fof(f2946,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,X2) = s(bool,X3) | s(bool,'F_0') = s(bool,X3) | s(bool,'F_0') = s(bool,X2)) )),
  inference(condensation,[],[f2038])).
fof(f2038,plain,(
  ( ! [X61,X59,X62,X60,X58] : (s(bool,'T_0') = s(bool,X58) | s(bool,X58) = s(bool,X59) | s(bool,X59) = s(bool,X60) | s(bool,'T_0') = s(bool,X60) | s(bool,X61) = s(bool,X62) | s(bool,'F_0') = s(bool,X62) | s(bool,'F_0') = s(bool,X61)) )),
  inference(superposition,[],[f1664,f1664])).
fof(f2966,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2)) )),
  inference(condensation,[],[f2965])).
fof(f2965,plain,(
  ( ! [X4,X2,X3,X1] : (s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X1) | s(bool,X3) = s(bool,X4) | s(bool,'F_0') = s(bool,X4)) )),
  inference(duplicate_literal_removal,[],[f2964])).
fof(f2964,plain,(
  ( ! [X4,X2,X3,X1] : (s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X1) | s(bool,X3) = s(bool,X4) | s(bool,'F_0') = s(bool,X4) | s(bool,'T_0') = s(bool,X1)) )),
  inference(condensation,[],[f2009])).
fof(f2009,plain,(
  ( ! [X61,X64,X62,X65,X63] : (s(bool,X62) = s(bool,X63) | s(bool,'F_0') = s(bool,X61) | s(bool,X63) = s(bool,X64) | s(bool,'T_0') = s(bool,X64) | s(bool,X61) = s(bool,X65) | s(bool,'F_0') = s(bool,X65) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f1664,f1664])).
fof(f2989,plain,(
  ( ! [X74,X72,X71,X75,X73] : (s(bool,X71) = s(bool,X72) | s(bool,'F_0') = s(bool,X72) | s(bool,'T_0') = s(bool,X73) | s(bool,X74) = s(bool,X75) | s(bool,'F_0') = s(bool,X75) | s(bool,X71) = s(bool,X74)) )),
  inference(duplicate_literal_removal,[],[f1981])).
fof(f1981,plain,(
  ( ! [X74,X72,X71,X75,X73] : (s(bool,X71) = s(bool,X72) | s(bool,'F_0') = s(bool,X72) | s(bool,'T_0') = s(bool,X73) | s(bool,'T_0') = s(bool,X73) | s(bool,X74) = s(bool,X75) | s(bool,'F_0') = s(bool,X75) | s(bool,X71) = s(bool,X74)) )),
  inference(superposition,[],[f1664,f1664])).
fof(f2996,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,'F_0') = s(bool,X1)) )),
  inference(condensation,[],[f2995])).
fof(f2995,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X3)) )),
  inference(duplicate_literal_removal,[],[f2994])).
fof(f2994,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(condensation,[],[f1975])).
fof(f1975,plain,(
  ( ! [X45,X43,X41,X44,X42] : (s(bool,X42) = s(bool,X43) | s(bool,'F_0') = s(bool,X43) | s(bool,X41) = s(bool,X44) | s(bool,'T_0') = s(bool,X44) | s(bool,'F_0') = s(bool,X42) | s(bool,X41) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) )),
  inference(superposition,[],[f1664,f1664])).
% SZS output end Proof for thm.real
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 135
% Passive clauses: 1315
% Generated clauses: 21904
% Final active clauses: 30
% Final passive clauses: 79
% Input formulas: 22
% Initial clauses: 39
% 
% Duplicate literals: 1666
% Trivial inequalities: 296
% Fw subsumption resolutions: 1577
% Fw demodulations: 1401
% Bw demodulations: 1159
% Condensations: 3342
% Global subsumptions: 61
% 
% Simple tautologies: 318
% Equational tautologies: 166
% Forward subsumptions: 11704
% Bw demodulations to eq. taut.: 17
% 
% Unit resulting resolution: 492
% Factoring: 1
% Forward superposition: 7061
% Backward superposition: 3700
% Self superposition: 219
% Equality factoring: 132
% Equality resolution: 19
% 
% Split clauses: 63
% Split components: 126
% SAT solver clauses: 1567
% SAT solver unit clauses: 40
% SAT solver binary clauses: 210
% 
% Sat splits: 63
% Sat splitting refutations: 63
% 
% Memory used [KB]: 8827
% Time elapsed: 1.328 s
% ------------------------------
% ------------------------------
% Success in time 2.645 s
